2020-05-31 21:21:43.939683: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_69
2020-05-31 21:21:44.060047: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-05-31 21:21:44.324989: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-05-31 21:21:45.078547: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-05-31 21:21:47.346965: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 3407292 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_69' loaded (query=128; context=256; weigths=real; len_mult=0.000000).
# Preprocessing time       : 4.560 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(free_g1_pre_topc, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>![X3, X4]:(g1_pre_topc(X1,X2)=g1_pre_topc(X3,X4)=>(X1=X3&X2=X4))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', free_g1_pre_topc)).
fof(d5_pcomps_1, axiom, ![X1]:(l1_metric_1(X1)=>k3_pcomps_1(X1)=g1_pre_topc(u1_struct_0(X1),k2_pcomps_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', d5_pcomps_1)).
fof(dt_k2_pcomps_1, axiom, ![X1]:(l1_metric_1(X1)=>m1_subset_1(k2_pcomps_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', dt_k2_pcomps_1)).
fof(abstractness_v1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v1_pre_topc(X1)=>X1=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', abstractness_v1_pre_topc)).
fof(dt_k3_pcomps_1, axiom, ![X1]:(l1_metric_1(X1)=>l1_pre_topc(k3_pcomps_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', dt_k3_pcomps_1)).
fof(fc1_pcomps_1, axiom, ![X1]:(l1_metric_1(X1)=>(v1_pre_topc(k3_pcomps_1(X1))&v2_pre_topc(k3_pcomps_1(X1)))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', fc1_pcomps_1)).
fof(free_g1_metric_1, axiom, ![X1, X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k2_zfmisc_1(X1,X1),k1_numbers))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers))))=>![X3, X4]:(g1_metric_1(X1,X2)=g1_metric_1(X3,X4)=>(X1=X3&X2=X4))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', free_g1_metric_1)).
fof(dt_k7_metric_1, axiom, ((v1_funct_1(k7_metric_1)&v1_funct_2(k7_metric_1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))&m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', dt_k7_metric_1)).
fof(d13_metric_1, axiom, k8_metric_1=g1_metric_1(k1_numbers,k7_metric_1), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', d13_metric_1)).
fof(abstractness_v1_metric_1, axiom, ![X1]:(l1_metric_1(X1)=>(v1_metric_1(X1)=>X1=g1_metric_1(u1_struct_0(X1),u1_metric_1(X1)))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', abstractness_v1_metric_1)).
fof(t24_tops_4, conjecture, ![X1]:(v7_ordinal1(X1)=>![X2]:(((v1_funct_1(X2)&v1_funct_2(X2,u1_struct_0(k15_euclid(X1)),u1_struct_0(k3_topmetr)))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k3_topmetr)))))=>(v5_pre_topc(X2,k15_euclid(X1),k3_topmetr)<=>![X3]:(m1_subset_1(X3,u1_struct_0(k15_euclid(X1)))=>![X4]:((v1_xreal_0(X4)&v2_xxreal_0(X4))=>?[X5]:((v1_xreal_0(X5)&v2_xxreal_0(X5))&r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k3_topmetr),X2,k1_topreal9(X1,X3,X5)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(X1)),u1_struct_0(k3_topmetr),X2,X3),X4),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(X1)),u1_struct_0(k3_topmetr),X2,X3),X4))))))))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', t24_tops_4)).
fof(t7_frechet, axiom, ![X1]:(m1_subset_1(X1,u1_struct_0(k8_metric_1))=>![X2]:(m1_subset_1(X2,k1_numbers)=>k9_metric_1(k8_metric_1,X1,X2)=k2_rcomp_1(k5_real_1(X1,X2),k3_real_1(X1,X2)))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', t7_frechet)).
fof(dt_k8_metric_1, axiom, (v1_metric_1(k8_metric_1)&l1_metric_1(k8_metric_1)), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', dt_k8_metric_1)).
fof(d6_topmetr, axiom, k3_topmetr=k3_pcomps_1(k8_metric_1), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', d6_topmetr)).
fof(fc5_metric_1, axiom, (~(v2_struct_0(k8_metric_1))&v1_metric_1(k8_metric_1)), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', fc5_metric_1)).
fof(redefinition_k3_funct_2, axiom, ![X1, X2, X3, X4]:(((((~(v1_xboole_0(X1))&v1_funct_1(X3))&v1_funct_2(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&m1_subset_1(X4,X1))=>k3_funct_2(X1,X2,X3,X4)=k1_funct_1(X3,X4)), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', redefinition_k3_funct_2)).
fof(redefinition_k5_real_1, axiom, ![X1, X2]:((v1_xreal_0(X1)&m1_subset_1(X2,k1_numbers))=>k5_real_1(X1,X2)=k6_xcmplx_0(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', redefinition_k5_real_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', t1_subset)).
fof(d1_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)<=>r2_hidden(X1,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', d1_xreal_0)).
fof(t17_tops_4, axiom, ![X1]:((((((~(v2_struct_0(X1))&v6_metric_1(X1))&v7_metric_1(X1))&v8_metric_1(X1))&v9_metric_1(X1))&l1_metric_1(X1))=>![X2]:((((((~(v2_struct_0(X2))&v6_metric_1(X2))&v7_metric_1(X2))&v8_metric_1(X2))&v9_metric_1(X2))&l1_metric_1(X2))=>![X3]:(((v1_funct_1(X3)&v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2))))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2))))))=>(v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>![X6]:((v1_xreal_0(X6)&v2_xxreal_0(X6))=>~((X5=k1_funct_1(X3,X4)&![X7]:((v1_xreal_0(X7)&v2_xxreal_0(X7))=>~(r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)),X3,k9_metric_1(X1,X4,X7)),k9_metric_1(X2,X5,X6))))))))))))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', t17_tops_4)).
fof(redefinition_k7_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k7_relset_1(X1,X2,X3,X4)=k7_relat_1(X3,X4)), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', redefinition_k7_relset_1)).
fof(t13_topreal9, axiom, ![X1]:(m1_subset_1(X1,k5_numbers)=>![X2]:(v1_xreal_0(X2)=>![X3]:(m1_subset_1(X3,u1_struct_0(k15_euclid(X1)))=>![X4]:(m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))=>(X3=X4=>k9_metric_1(k14_euclid(X1),X4,X2)=k1_topreal9(X1,X3,X2)))))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', t13_topreal9)).
fof(redefinition_k3_real_1, axiom, ![X1, X2]:((v1_xreal_0(X1)&m1_subset_1(X2,k1_numbers))=>k3_real_1(X1,X2)=k2_xcmplx_0(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', redefinition_k3_real_1)).
fof(d12_ordinal1, axiom, ![X1]:(v7_ordinal1(X1)<=>r2_hidden(X1,k4_ordinal1)), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', d12_ordinal1)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', redefinition_k5_numbers)).
fof(fc6_metric_1, axiom, ((((v1_metric_1(k8_metric_1)&v6_metric_1(k8_metric_1))&v7_metric_1(k8_metric_1))&v8_metric_1(k8_metric_1))&v9_metric_1(k8_metric_1)), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', fc6_metric_1)).
fof(t67_euclid, axiom, ![X1]:(v7_ordinal1(X1)=>u1_struct_0(k14_euclid(X1))=u1_struct_0(k15_euclid(X1))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', t67_euclid)).
fof(dt_k14_euclid, axiom, ![X1]:(v7_ordinal1(X1)=>(((((v1_metric_1(k14_euclid(X1))&v6_metric_1(k14_euclid(X1)))&v7_metric_1(k14_euclid(X1)))&v8_metric_1(k14_euclid(X1)))&v9_metric_1(k14_euclid(X1)))&l1_metric_1(k14_euclid(X1)))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', dt_k14_euclid)).
fof(fc4_euclid, axiom, ![X1]:(v7_ordinal1(X1)=>(((((~(v2_struct_0(k14_euclid(X1)))&v1_metric_1(k14_euclid(X1)))&v6_metric_1(k14_euclid(X1)))&v7_metric_1(k14_euclid(X1)))&v8_metric_1(k14_euclid(X1)))&v9_metric_1(k14_euclid(X1)))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', fc4_euclid)).
fof(cc18_valued_0, axiom, ![X1, X2]:(v3_membered(X2)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v3_valued_0(X3))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', cc18_valued_0)).
fof(fc63_valued_0, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v3_valued_0(X1))=>v1_xreal_0(k1_funct_1(X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', fc63_valued_0)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', cc2_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', fc6_relat_1)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', fc3_membered)).
fof(t36_yellow12, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:((v2_pre_topc(X2)&l1_pre_topc(X2))=>![X3]:((v2_pre_topc(X3)&l1_pre_topc(X3))=>![X4]:((v2_pre_topc(X4)&l1_pre_topc(X4))=>![X5]:(((v1_funct_1(X5)&v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X3)))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3)))))=>![X6]:(((v1_funct_1(X6)&v1_funct_2(X6,u1_struct_0(X2),u1_struct_0(X4)))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X4)))))=>((((g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))&g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))=g1_pre_topc(u1_struct_0(X4),u1_pre_topc(X4)))&X5=X6)&v5_pre_topc(X5,X1,X3))=>v5_pre_topc(X6,X2,X4)))))))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', t36_yellow12)).
fof(fc2_topmetr, axiom, ((~(v2_struct_0(k3_topmetr))&v1_pre_topc(k3_topmetr))&v2_pre_topc(k3_topmetr)), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', fc2_topmetr)).
fof(dt_k3_topmetr, axiom, (v2_pre_topc(k3_topmetr)&l1_pre_topc(k3_topmetr)), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', dt_k3_topmetr)).
fof(d8_euclid, axiom, ![X1]:(v7_ordinal1(X1)=>![X2]:((v5_rltopsp1(X2)&l1_rltopsp1(X2))=>(X2=k15_euclid(X1)<=>(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))=k3_pcomps_1(k14_euclid(X1))&g1_rlvect_1(u1_struct_0(X2),u2_struct_0(X2),u1_algstr_0(X2),u1_rlvect_1(X2))=k10_funcsdom(k2_finseq_1(X1)))))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', d8_euclid)).
fof(dt_k15_euclid, axiom, ![X1]:(v7_ordinal1(X1)=>(v5_rltopsp1(k15_euclid(X1))&l1_rltopsp1(k15_euclid(X1)))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', dt_k15_euclid)).
fof(dt_l1_rltopsp1, axiom, ![X1]:(l1_rltopsp1(X1)=>(l1_rlvect_1(X1)&l1_pre_topc(X1))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', dt_l1_rltopsp1)).
fof(commutativity_k3_real_1, axiom, ![X1, X2]:((v1_xreal_0(X1)&m1_subset_1(X2,k1_numbers))=>k3_real_1(X1,X2)=k3_real_1(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', commutativity_k3_real_1)).
fof(fc6_euclid, axiom, ![X1]:(v7_ordinal1(X1)=>(((((((((v2_pre_topc(k15_euclid(X1))&v13_algstr_0(k15_euclid(X1)))&v2_rlvect_1(k15_euclid(X1)))&v3_rlvect_1(k15_euclid(X1)))&v4_rlvect_1(k15_euclid(X1)))&v5_rlvect_1(k15_euclid(X1)))&v6_rlvect_1(k15_euclid(X1)))&v7_rlvect_1(k15_euclid(X1)))&v8_rlvect_1(k15_euclid(X1)))&v5_rltopsp1(k15_euclid(X1)))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', fc6_euclid)).
fof(commutativity_k2_xcmplx_0, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', commutativity_k2_xcmplx_0)).
fof(cc3_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xcmplx_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', cc3_xreal_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', fc2_struct_0)).
fof(dt_l1_metric_1, axiom, ![X1]:(l1_metric_1(X1)=>l1_struct_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', dt_l1_metric_1)).
fof(cc1_xreal_0, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>v1_xreal_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', cc1_xreal_0)).
fof(fc7_pre_topc, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>(~(v2_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))&v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', fc7_pre_topc)).
fof(dt_u1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', dt_u1_pre_topc)).
fof(fc5_euclid, axiom, ![X1]:(v7_ordinal1(X1)=>(~(v2_struct_0(k15_euclid(X1)))&v5_rltopsp1(k15_euclid(X1)))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', fc5_euclid)).
fof(fc6_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>(v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))&v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', fc6_pre_topc)).
fof(dt_g1_pre_topc, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>(v1_pre_topc(g1_pre_topc(X1,X2))&l1_pre_topc(g1_pre_topc(X1,X2)))), file('/scratch/mptp/allpred/preds__0.005/t24_tops_4', dt_g1_pre_topc)).
fof(c_0_52, plain, ![X31, X32, X33, X34]:((X31=X33|g1_pre_topc(X31,X32)!=g1_pre_topc(X33,X34)|~m1_subset_1(X32,k1_zfmisc_1(k1_zfmisc_1(X31))))&(X32=X34|g1_pre_topc(X31,X32)!=g1_pre_topc(X33,X34)|~m1_subset_1(X32,k1_zfmisc_1(k1_zfmisc_1(X31))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[free_g1_pre_topc])])])])).
fof(c_0_53, plain, ![X62]:(~l1_metric_1(X62)|k3_pcomps_1(X62)=g1_pre_topc(u1_struct_0(X62),k2_pcomps_1(X62))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_pcomps_1])])).
cnf(c_0_54, plain, (X1=X2|g1_pre_topc(X1,X3)!=g1_pre_topc(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_55, plain, (k3_pcomps_1(X1)=g1_pre_topc(u1_struct_0(X1),k2_pcomps_1(X1))|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
fof(c_0_56, plain, ![X84]:(~l1_metric_1(X84)|m1_subset_1(k2_pcomps_1(X84),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X84))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_pcomps_1])])).
cnf(c_0_57, plain, (X1=X2|g1_pre_topc(X1,X3)!=g1_pre_topc(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(evalgc,[status(thm)],[c_0_54])).
cnf(c_0_58, plain, (g1_pre_topc(u1_struct_0(X1),k2_pcomps_1(X1))=k3_pcomps_1(X1)|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[c_0_55])).
cnf(c_0_59, plain, (m1_subset_1(k2_pcomps_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
fof(c_0_60, plain, ![X54]:(~l1_pre_topc(X54)|(~v1_pre_topc(X54)|X54=g1_pre_topc(u1_struct_0(X54),u1_pre_topc(X54)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[abstractness_v1_pre_topc])])).
cnf(c_0_61, plain, (u1_struct_0(X1)=X2|k3_pcomps_1(X1)!=g1_pre_topc(X2,X3)|~m1_subset_1(k2_pcomps_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58])])).
cnf(c_0_62, plain, (m1_subset_1(k2_pcomps_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[c_0_59])).
cnf(c_0_63, plain, (X1=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X1)|~v1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_64, plain, (u1_struct_0(X1)=X2|k3_pcomps_1(X1)!=g1_pre_topc(X2,X3)|~l1_metric_1(X1)), inference(csr,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_65, plain, (g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=X1|~l1_pre_topc(X1)|~v1_pre_topc(X1)), inference(evalgc,[status(thm)],[c_0_63])).
fof(c_0_66, plain, ![X52]:(~l1_metric_1(X52)|l1_pre_topc(k3_pcomps_1(X52))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_pcomps_1])])).
fof(c_0_67, plain, ![X58]:((v1_pre_topc(k3_pcomps_1(X58))|~l1_metric_1(X58))&(v2_pre_topc(k3_pcomps_1(X58))|~l1_metric_1(X58))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_pcomps_1])])])).
fof(c_0_68, plain, ![X79, X80, X81, X82]:((X79=X81|g1_metric_1(X79,X80)!=g1_metric_1(X81,X82)|(~v1_funct_1(X80)|~v1_funct_2(X80,k2_zfmisc_1(X79,X79),k1_numbers)|~m1_subset_1(X80,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X79,X79),k1_numbers)))))&(X80=X82|g1_metric_1(X79,X80)!=g1_metric_1(X81,X82)|(~v1_funct_1(X80)|~v1_funct_2(X80,k2_zfmisc_1(X79,X79),k1_numbers)|~m1_subset_1(X80,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X79,X79),k1_numbers)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[free_g1_metric_1])])])])).
cnf(c_0_69, plain, (u1_struct_0(X1)=u1_struct_0(X2)|k3_pcomps_1(X1)!=X2|~l1_metric_1(X1)|~v1_pre_topc(X2)|~l1_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65])])).
cnf(c_0_70, plain, (l1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_71, plain, (v1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_72, plain, (X1=X2|g1_metric_1(X1,X3)!=g1_metric_1(X2,X4)|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_73, plain, (v1_funct_2(k7_metric_1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)), inference(split_conjunct,[status(thm)],[dt_k7_metric_1])).
cnf(c_0_74, plain, (k8_metric_1=g1_metric_1(k1_numbers,k7_metric_1)), inference(split_conjunct,[status(thm)],[d13_metric_1])).
cnf(c_0_75, plain, (m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))), inference(split_conjunct,[status(thm)],[dt_k7_metric_1])).
cnf(c_0_76, plain, (v1_funct_1(k7_metric_1)), inference(split_conjunct,[status(thm)],[dt_k7_metric_1])).
fof(c_0_77, plain, ![X83]:(~l1_metric_1(X83)|(~v1_metric_1(X83)|X83=g1_metric_1(u1_struct_0(X83),u1_metric_1(X83)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[abstractness_v1_metric_1])])).
fof(c_0_78, negated_conjecture, ~(![X1]:(v7_ordinal1(X1)=>![X2]:(((v1_funct_1(X2)&v1_funct_2(X2,u1_struct_0(k15_euclid(X1)),u1_struct_0(k3_topmetr)))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k3_topmetr)))))=>(v5_pre_topc(X2,k15_euclid(X1),k3_topmetr)<=>![X3]:(m1_subset_1(X3,u1_struct_0(k15_euclid(X1)))=>![X4]:((v1_xreal_0(X4)&v2_xxreal_0(X4))=>?[X5]:((v1_xreal_0(X5)&v2_xxreal_0(X5))&r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k3_topmetr),X2,k1_topreal9(X1,X3,X5)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(X1)),u1_struct_0(k3_topmetr),X2,X3),X4),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(X1)),u1_struct_0(k3_topmetr),X2,X3),X4)))))))))), inference(assume_negation,[status(cth)],[t24_tops_4])).
fof(c_0_79, plain, ![X64, X65]:(~m1_subset_1(X64,u1_struct_0(k8_metric_1))|(~m1_subset_1(X65,k1_numbers)|k9_metric_1(k8_metric_1,X64,X65)=k2_rcomp_1(k5_real_1(X64,X65),k3_real_1(X64,X65)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_frechet])])])).
cnf(c_0_80, plain, (u1_struct_0(k3_pcomps_1(X1))=u1_struct_0(X1)|~v1_pre_topc(k3_pcomps_1(X1))|~l1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_69])])).
cnf(c_0_81, plain, (l1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[c_0_70])).
cnf(c_0_82, plain, (v1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[c_0_71])).
cnf(c_0_83, plain, (l1_metric_1(k8_metric_1)), inference(split_conjunct,[status(thm)],[dt_k8_metric_1])).
cnf(c_0_84, plain, (k3_topmetr=k3_pcomps_1(k8_metric_1)), inference(split_conjunct,[status(thm)],[d6_topmetr])).
cnf(c_0_85, plain, (X1=X2|g1_metric_1(X1,X3)!=g1_metric_1(X2,X4)|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)))), inference(evalgc,[status(thm)],[c_0_72])).
cnf(c_0_86, plain, (v1_funct_2(k7_metric_1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)), inference(evalgc,[status(thm)],[c_0_73])).
cnf(c_0_87, plain, (g1_metric_1(k1_numbers,k7_metric_1)=k8_metric_1), inference(evalgc,[status(thm)],[c_0_74])).
cnf(c_0_88, plain, (m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))), inference(evalgc,[status(thm)],[c_0_75])).
cnf(c_0_89, plain, (v1_funct_1(k7_metric_1)), inference(evalgc,[status(thm)],[c_0_76])).
cnf(c_0_90, plain, (X1=g1_metric_1(u1_struct_0(X1),u1_metric_1(X1))|~l1_metric_1(X1)|~v1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_77])).
fof(c_0_91, plain, (~v2_struct_0(k8_metric_1)&v1_metric_1(k8_metric_1)), inference(fof_simplification,[status(thm)],[fc5_metric_1])).
fof(c_0_92, negated_conjecture, ![X12, X13, X14]:(v7_ordinal1(esk1_0)&(((v1_funct_1(esk2_0)&v1_funct_2(esk2_0,u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr)))&m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr)))))&(((m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0)))|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr))&(((v1_xreal_0(esk4_0)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr))&(v2_xxreal_0(esk4_0)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)))&(~v1_xreal_0(X12)|~v2_xxreal_0(X12)|~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,k1_topreal9(esk1_0,esk3_0,X12)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0),esk4_0)))|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr))))&(((v1_xreal_0(esk5_2(X13,X14))|(~v1_xreal_0(X14)|~v2_xxreal_0(X14))|~m1_subset_1(X13,u1_struct_0(k15_euclid(esk1_0)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr))&(v2_xxreal_0(esk5_2(X13,X14))|(~v1_xreal_0(X14)|~v2_xxreal_0(X14))|~m1_subset_1(X13,u1_struct_0(k15_euclid(esk1_0)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)))&(r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,k1_topreal9(esk1_0,X13,esk5_2(X13,X14))),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X13),X14),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X13),X14)))|(~v1_xreal_0(X14)|~v2_xxreal_0(X14))|~m1_subset_1(X13,u1_struct_0(k15_euclid(esk1_0)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_78])])])])])).
fof(c_0_93, plain, ![X35, X36, X37, X38]:(v1_xboole_0(X35)|~v1_funct_1(X37)|~v1_funct_2(X37,X35,X36)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(X35,X36)))|~m1_subset_1(X38,X35)|k3_funct_2(X35,X36,X37,X38)=k1_funct_1(X37,X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_funct_2])])])).
cnf(c_0_94, plain, (k9_metric_1(k8_metric_1,X1,X2)=k2_rcomp_1(k5_real_1(X1,X2),k3_real_1(X1,X2))|~m1_subset_1(X1,u1_struct_0(k8_metric_1))|~m1_subset_1(X2,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_95, plain, (u1_struct_0(k3_pcomps_1(X1))=u1_struct_0(X1)|~l1_metric_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_80, c_0_81]), c_0_82])).
cnf(c_0_96, plain, (l1_metric_1(k8_metric_1)), inference(evalgc,[status(thm)],[c_0_83])).
cnf(c_0_97, plain, (k3_pcomps_1(k8_metric_1)=k3_topmetr), inference(evalgc,[status(thm)],[c_0_84])).
cnf(c_0_98, plain, (k1_numbers=X1|g1_metric_1(X1,X2)!=k8_metric_1), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_87]), c_0_88]), c_0_89])])])).
cnf(c_0_99, plain, (g1_metric_1(u1_struct_0(X1),u1_metric_1(X1))=X1|~l1_metric_1(X1)|~v1_metric_1(X1)), inference(evalgc,[status(thm)],[c_0_90])).
cnf(c_0_100, plain, (v1_metric_1(k8_metric_1)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_101, negated_conjecture, (r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X1),X2),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X1),X2)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v1_xreal_0(X2)|~v2_xxreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_102, plain, (v1_xboole_0(X1)|k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|~v1_funct_1(X2)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~m1_subset_1(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_103, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr))))), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_104, negated_conjecture, (v1_funct_2(esk2_0,u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr))), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_105, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_106, plain, (k2_rcomp_1(k5_real_1(X1,X2),k3_real_1(X1,X2))=k9_metric_1(k8_metric_1,X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,u1_struct_0(k8_metric_1))), inference(evalgc,[status(thm)],[c_0_94])).
cnf(c_0_107, plain, (u1_struct_0(k8_metric_1)=u1_struct_0(k3_topmetr)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_97])])).
cnf(c_0_108, plain, (u1_struct_0(X1)=k1_numbers|X1!=k8_metric_1|~v1_metric_1(X1)|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99])])).
cnf(c_0_109, plain, (v1_metric_1(k8_metric_1)), inference(evalgc,[status(thm)],[c_0_100])).
fof(c_0_110, plain, ![X49, X50]:(~v1_xreal_0(X49)|~m1_subset_1(X50,k1_numbers)|k5_real_1(X49,X50)=k6_xcmplx_0(X49,X50)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k5_real_1])])).
fof(c_0_111, plain, ![X77, X78]:(~r2_hidden(X77,X78)|m1_subset_1(X77,X78)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
fof(c_0_112, plain, ![X53]:((~v1_xreal_0(X53)|r2_hidden(X53,k1_numbers))&(~r2_hidden(X53,k1_numbers)|v1_xreal_0(X53))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_xreal_0])])).
cnf(c_0_113, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X1),X2),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X1),X2)))|~v1_xreal_0(X2)|~v2_xxreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))), inference(evalgc,[status(thm)],[c_0_101])).
cnf(c_0_114, plain, (k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(evalgc,[status(thm)],[c_0_102])).
cnf(c_0_115, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr))))), inference(evalgc,[status(thm)],[c_0_103])).
cnf(c_0_116, negated_conjecture, (v1_funct_2(esk2_0,u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr))), inference(evalgc,[status(thm)],[c_0_104])).
cnf(c_0_117, negated_conjecture, (v1_funct_1(esk2_0)), inference(evalgc,[status(thm)],[c_0_105])).
fof(c_0_118, plain, ![X16, X17, X18, X19, X20, X21, X26]:((((v1_xreal_0(esk6_6(X16,X17,X18,X19,X20,X21))|X20!=k1_funct_1(X18,X19)|(~v1_xreal_0(X21)|~v2_xxreal_0(X21))|~m1_subset_1(X20,u1_struct_0(X17))|~m1_subset_1(X19,u1_struct_0(X16))|~v5_pre_topc(X18,k3_pcomps_1(X16),k3_pcomps_1(X17))|(~v1_funct_1(X18)|~v1_funct_2(X18,u1_struct_0(k3_pcomps_1(X16)),u1_struct_0(k3_pcomps_1(X17)))|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X16)),u1_struct_0(k3_pcomps_1(X17))))))|(v2_struct_0(X17)|~v6_metric_1(X17)|~v7_metric_1(X17)|~v8_metric_1(X17)|~v9_metric_1(X17)|~l1_metric_1(X17))|(v2_struct_0(X16)|~v6_metric_1(X16)|~v7_metric_1(X16)|~v8_metric_1(X16)|~v9_metric_1(X16)|~l1_metric_1(X16)))&(v2_xxreal_0(esk6_6(X16,X17,X18,X19,X20,X21))|X20!=k1_funct_1(X18,X19)|(~v1_xreal_0(X21)|~v2_xxreal_0(X21))|~m1_subset_1(X20,u1_struct_0(X17))|~m1_subset_1(X19,u1_struct_0(X16))|~v5_pre_topc(X18,k3_pcomps_1(X16),k3_pcomps_1(X17))|(~v1_funct_1(X18)|~v1_funct_2(X18,u1_struct_0(k3_pcomps_1(X16)),u1_struct_0(k3_pcomps_1(X17)))|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X16)),u1_struct_0(k3_pcomps_1(X17))))))|(v2_struct_0(X17)|~v6_metric_1(X17)|~v7_metric_1(X17)|~v8_metric_1(X17)|~v9_metric_1(X17)|~l1_metric_1(X17))|(v2_struct_0(X16)|~v6_metric_1(X16)|~v7_metric_1(X16)|~v8_metric_1(X16)|~v9_metric_1(X16)|~l1_metric_1(X16))))&(r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(X16)),u1_struct_0(k3_pcomps_1(X17)),X18,k9_metric_1(X16,X19,esk6_6(X16,X17,X18,X19,X20,X21))),k9_metric_1(X17,X20,X21))|X20!=k1_funct_1(X18,X19)|(~v1_xreal_0(X21)|~v2_xxreal_0(X21))|~m1_subset_1(X20,u1_struct_0(X17))|~m1_subset_1(X19,u1_struct_0(X16))|~v5_pre_topc(X18,k3_pcomps_1(X16),k3_pcomps_1(X17))|(~v1_funct_1(X18)|~v1_funct_2(X18,u1_struct_0(k3_pcomps_1(X16)),u1_struct_0(k3_pcomps_1(X17)))|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X16)),u1_struct_0(k3_pcomps_1(X17))))))|(v2_struct_0(X17)|~v6_metric_1(X17)|~v7_metric_1(X17)|~v8_metric_1(X17)|~v9_metric_1(X17)|~l1_metric_1(X17))|(v2_struct_0(X16)|~v6_metric_1(X16)|~v7_metric_1(X16)|~v8_metric_1(X16)|~v9_metric_1(X16)|~l1_metric_1(X16))))&((m1_subset_1(esk7_3(X16,X17,X18),u1_struct_0(X16))|v5_pre_topc(X18,k3_pcomps_1(X16),k3_pcomps_1(X17))|(~v1_funct_1(X18)|~v1_funct_2(X18,u1_struct_0(k3_pcomps_1(X16)),u1_struct_0(k3_pcomps_1(X17)))|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X16)),u1_struct_0(k3_pcomps_1(X17))))))|(v2_struct_0(X17)|~v6_metric_1(X17)|~v7_metric_1(X17)|~v8_metric_1(X17)|~v9_metric_1(X17)|~l1_metric_1(X17))|(v2_struct_0(X16)|~v6_metric_1(X16)|~v7_metric_1(X16)|~v8_metric_1(X16)|~v9_metric_1(X16)|~l1_metric_1(X16)))&((m1_subset_1(esk8_3(X16,X17,X18),u1_struct_0(X17))|v5_pre_topc(X18,k3_pcomps_1(X16),k3_pcomps_1(X17))|(~v1_funct_1(X18)|~v1_funct_2(X18,u1_struct_0(k3_pcomps_1(X16)),u1_struct_0(k3_pcomps_1(X17)))|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X16)),u1_struct_0(k3_pcomps_1(X17))))))|(v2_struct_0(X17)|~v6_metric_1(X17)|~v7_metric_1(X17)|~v8_metric_1(X17)|~v9_metric_1(X17)|~l1_metric_1(X17))|(v2_struct_0(X16)|~v6_metric_1(X16)|~v7_metric_1(X16)|~v8_metric_1(X16)|~v9_metric_1(X16)|~l1_metric_1(X16)))&(((v1_xreal_0(esk9_3(X16,X17,X18))|v5_pre_topc(X18,k3_pcomps_1(X16),k3_pcomps_1(X17))|(~v1_funct_1(X18)|~v1_funct_2(X18,u1_struct_0(k3_pcomps_1(X16)),u1_struct_0(k3_pcomps_1(X17)))|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X16)),u1_struct_0(k3_pcomps_1(X17))))))|(v2_struct_0(X17)|~v6_metric_1(X17)|~v7_metric_1(X17)|~v8_metric_1(X17)|~v9_metric_1(X17)|~l1_metric_1(X17))|(v2_struct_0(X16)|~v6_metric_1(X16)|~v7_metric_1(X16)|~v8_metric_1(X16)|~v9_metric_1(X16)|~l1_metric_1(X16)))&(v2_xxreal_0(esk9_3(X16,X17,X18))|v5_pre_topc(X18,k3_pcomps_1(X16),k3_pcomps_1(X17))|(~v1_funct_1(X18)|~v1_funct_2(X18,u1_struct_0(k3_pcomps_1(X16)),u1_struct_0(k3_pcomps_1(X17)))|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X16)),u1_struct_0(k3_pcomps_1(X17))))))|(v2_struct_0(X17)|~v6_metric_1(X17)|~v7_metric_1(X17)|~v8_metric_1(X17)|~v9_metric_1(X17)|~l1_metric_1(X17))|(v2_struct_0(X16)|~v6_metric_1(X16)|~v7_metric_1(X16)|~v8_metric_1(X16)|~v9_metric_1(X16)|~l1_metric_1(X16))))&((esk8_3(X16,X17,X18)=k1_funct_1(X18,esk7_3(X16,X17,X18))|v5_pre_topc(X18,k3_pcomps_1(X16),k3_pcomps_1(X17))|(~v1_funct_1(X18)|~v1_funct_2(X18,u1_struct_0(k3_pcomps_1(X16)),u1_struct_0(k3_pcomps_1(X17)))|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X16)),u1_struct_0(k3_pcomps_1(X17))))))|(v2_struct_0(X17)|~v6_metric_1(X17)|~v7_metric_1(X17)|~v8_metric_1(X17)|~v9_metric_1(X17)|~l1_metric_1(X17))|(v2_struct_0(X16)|~v6_metric_1(X16)|~v7_metric_1(X16)|~v8_metric_1(X16)|~v9_metric_1(X16)|~l1_metric_1(X16)))&(~v1_xreal_0(X26)|~v2_xxreal_0(X26)|~r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(X16)),u1_struct_0(k3_pcomps_1(X17)),X18,k9_metric_1(X16,esk7_3(X16,X17,X18),X26)),k9_metric_1(X17,esk8_3(X16,X17,X18),esk9_3(X16,X17,X18)))|v5_pre_topc(X18,k3_pcomps_1(X16),k3_pcomps_1(X17))|(~v1_funct_1(X18)|~v1_funct_2(X18,u1_struct_0(k3_pcomps_1(X16)),u1_struct_0(k3_pcomps_1(X17)))|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X16)),u1_struct_0(k3_pcomps_1(X17))))))|(v2_struct_0(X17)|~v6_metric_1(X17)|~v7_metric_1(X17)|~v8_metric_1(X17)|~v9_metric_1(X17)|~l1_metric_1(X17))|(v2_struct_0(X16)|~v6_metric_1(X16)|~v7_metric_1(X16)|~v8_metric_1(X16)|~v9_metric_1(X16)|~l1_metric_1(X16)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t17_tops_4])])])])])])).
cnf(c_0_119, plain, (k2_rcomp_1(k5_real_1(X1,X2),k3_real_1(X1,X2))=k9_metric_1(k8_metric_1,X1,X2)|~m1_subset_1(X1,u1_struct_0(k3_topmetr))|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_106, c_0_107])])).
cnf(c_0_120, plain, (u1_struct_0(k3_topmetr)=k1_numbers), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_107]), c_0_96])])])).
cnf(c_0_121, plain, (k5_real_1(X1,X2)=k6_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_122, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_123, plain, (r2_hidden(X1,k1_numbers)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_112])).
fof(c_0_124, plain, ![X71, X72, X73, X74]:(~m1_subset_1(X73,k1_zfmisc_1(k2_zfmisc_1(X71,X72)))|k7_relset_1(X71,X72,X73,X74)=k7_relat_1(X73,X74)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_relset_1])])).
fof(c_0_125, plain, ![X27, X28, X29, X30]:(~m1_subset_1(X27,k5_numbers)|(~v1_xreal_0(X28)|(~m1_subset_1(X29,u1_struct_0(k15_euclid(X27)))|(~m1_subset_1(X30,u1_struct_0(k14_euclid(X27)))|(X29!=X30|k9_metric_1(k14_euclid(X27),X30,X28)=k1_topreal9(X27,X29,X28)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_topreal9])])])).
cnf(c_0_126, negated_conjecture, (r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,X1),X2),k2_xcmplx_0(k1_funct_1(esk2_0,X1),X2)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114]), c_0_115]), c_0_116]), c_0_117])])])).
cnf(c_0_127, plain, (esk8_3(X1,X2,X3)=k1_funct_1(X3,esk7_3(X1,X2,X3))|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_128, plain, (k2_rcomp_1(k5_real_1(X1,X2),k3_real_1(X1,X2))=k9_metric_1(k8_metric_1,X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(rw,[status(thm)],[c_0_119, c_0_120])).
cnf(c_0_129, plain, (k5_real_1(X1,X2)=k6_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[c_0_121])).
cnf(c_0_130, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_122])).
cnf(c_0_131, plain, (r2_hidden(X1,k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_123])).
fof(c_0_132, plain, ![X45, X46]:(~v1_xreal_0(X45)|~m1_subset_1(X46,k1_numbers)|k3_real_1(X45,X46)=k2_xcmplx_0(X45,X46)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k3_real_1])])).
fof(c_0_133, plain, ![X60]:((~v7_ordinal1(X60)|r2_hidden(X60,k4_ordinal1))&(~r2_hidden(X60,k4_ordinal1)|v7_ordinal1(X60))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_ordinal1])])).
cnf(c_0_134, plain, (v5_pre_topc(X4,k3_pcomps_1(X2),k3_pcomps_1(X3))|v2_struct_0(X3)|v2_struct_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1)|~r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X3)),X4,k9_metric_1(X2,esk7_3(X2,X3,X4),X1)),k9_metric_1(X3,esk8_3(X2,X3,X4),esk9_3(X2,X3,X4)))|~v1_funct_1(X4)|~v1_funct_2(X4,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X3)))))|~v6_metric_1(X3)|~v7_metric_1(X3)|~v8_metric_1(X3)|~v9_metric_1(X3)|~l1_metric_1(X3)|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_135, plain, (k7_relset_1(X2,X3,X1,X4)=k7_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_124])).
cnf(c_0_136, plain, (k9_metric_1(k14_euclid(X1),X4,X2)=k1_topreal9(X1,X3,X2)|~m1_subset_1(X1,k5_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X3,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|X3!=X4), inference(split_conjunct,[status(thm)],[c_0_125])).
cnf(c_0_137, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,X1),X2),k2_xcmplx_0(k1_funct_1(esk2_0,X1),X2)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_126, c_0_120])])).
cnf(c_0_138, plain, (k1_funct_1(X3,esk7_3(X1,X2,X3))=esk8_3(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v1_funct_1(X3)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~l1_metric_1(X2)|~l1_metric_1(X1)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))), inference(evalgc,[status(thm)],[c_0_127])).
cnf(c_0_139, plain, (k2_rcomp_1(k6_xcmplx_0(X1,X2),k3_real_1(X1,X2))=k9_metric_1(k8_metric_1,X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_129])])).
cnf(c_0_140, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_131])])).
cnf(c_0_141, plain, (k3_real_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_132])).
cnf(c_0_142, plain, (r2_hidden(X1,k4_ordinal1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_133])).
cnf(c_0_143, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
cnf(c_0_144, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v5_pre_topc(X4,k3_pcomps_1(X2),k3_pcomps_1(X3))|~v1_funct_1(X4)|~v1_xreal_0(X1)|~v2_xxreal_0(X1)|~v6_metric_1(X3)|~v6_metric_1(X2)|~v7_metric_1(X3)|~v7_metric_1(X2)|~v8_metric_1(X3)|~v8_metric_1(X2)|~v9_metric_1(X3)|~v9_metric_1(X2)|~l1_metric_1(X3)|~l1_metric_1(X2)|~v1_funct_2(X4,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X3)))))|~r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X3)),X4,k9_metric_1(X2,esk7_3(X2,X3,X4),X1)),k9_metric_1(X3,esk8_3(X2,X3,X4),esk9_3(X2,X3,X4)))), inference(evalgc,[status(thm)],[c_0_134])).
cnf(c_0_145, plain, (k7_relset_1(X2,X3,X1,X4)=k7_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_135])).
cnf(c_0_146, plain, (k9_metric_1(k14_euclid(X1),X4,X2)=k1_topreal9(X1,X3,X2)|X3!=X4|~v1_xreal_0(X2)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(X1)))), inference(evalgc,[status(thm)],[c_0_136])).
cnf(c_0_147, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk7_3(X1,X2,esk2_0),esk5_2(esk7_3(X1,X2,esk2_0),X3))),k2_rcomp_1(k6_xcmplx_0(esk8_3(X1,X2,esk2_0),X3),k2_xcmplx_0(esk8_3(X1,X2,esk2_0),X3)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v2_struct_0(X1)|v2_struct_0(X2)|v5_pre_topc(esk2_0,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(esk7_3(X1,X2,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_138]), c_0_117])])])).
cnf(c_0_148, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers)))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_115, c_0_120])])).
cnf(c_0_149, plain, (k2_rcomp_1(k6_xcmplx_0(X1,X2),k3_real_1(X1,X2))=k9_metric_1(k8_metric_1,X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_139, c_0_140])).
cnf(c_0_150, plain, (k3_real_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[c_0_141])).
cnf(c_0_151, plain, (r2_hidden(X1,k4_ordinal1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_142])).
cnf(c_0_152, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_143])).
cnf(c_0_153, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~r1_tarski(k7_relat_1(X3,k9_metric_1(X1,esk7_3(X1,X2,X3),X4)),k9_metric_1(X2,esk8_3(X1,X2,X3),esk9_3(X1,X2,X3)))|~v2_xxreal_0(X4)|~v1_xreal_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_145])])).
cnf(c_0_154, plain, (k9_metric_1(k14_euclid(X1),X2,X3)=k1_topreal9(X1,X2,X3)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_146])])).
cnf(c_0_155, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_156, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(X1)|v2_struct_0(X2)|r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk7_3(X2,X1,esk2_0),esk5_2(esk7_3(X2,X1,esk2_0),X3))),k2_rcomp_1(k6_xcmplx_0(esk8_3(X2,X1,esk2_0),X3),k2_xcmplx_0(esk8_3(X2,X1,esk2_0),X3)))|v5_pre_topc(esk2_0,k3_pcomps_1(X2),k3_pcomps_1(X1))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~l1_metric_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X1)|~v9_metric_1(X2)|~v8_metric_1(X1)|~v8_metric_1(X2)|~v7_metric_1(X1)|~v7_metric_1(X2)|~v6_metric_1(X1)|~v6_metric_1(X2)|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X1)))))|~m1_subset_1(esk7_3(X2,X1,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X1)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_145]), c_0_148])])])).
cnf(c_0_157, plain, (k2_rcomp_1(k6_xcmplx_0(X1,X2),k2_xcmplx_0(X1,X2))=k9_metric_1(k8_metric_1,X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_150])])).
cnf(c_0_158, plain, (r2_hidden(X1,k5_numbers)|~v7_ordinal1(X1)), inference(rw,[status(thm)],[c_0_151, c_0_152])).
cnf(c_0_159, negated_conjecture, (v7_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_160, plain, (v2_struct_0(X1)|v2_struct_0(k14_euclid(X2))|v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X2)),k3_pcomps_1(X1))|~l1_metric_1(X1)|~l1_metric_1(k14_euclid(X2))|~v9_metric_1(X1)|~v9_metric_1(k14_euclid(X2))|~v8_metric_1(X1)|~v8_metric_1(k14_euclid(X2))|~v7_metric_1(X1)|~v7_metric_1(k14_euclid(X2))|~v6_metric_1(X1)|~v6_metric_1(k14_euclid(X2))|~r1_tarski(k7_relat_1(X3,k1_topreal9(X2,esk7_3(k14_euclid(X2),X1,X3),X4)),k9_metric_1(X1,esk8_3(k14_euclid(X2),X1,X3),esk9_3(k14_euclid(X2),X1,X3)))|~v2_xxreal_0(X4)|~v1_xreal_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X2))),u1_struct_0(k3_pcomps_1(X1)))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X2))),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X3)|~m1_subset_1(esk7_3(k14_euclid(X2),X1,X3),u1_struct_0(k14_euclid(X2)))|~m1_subset_1(esk7_3(k14_euclid(X2),X1,X3),u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_154])])).
cnf(c_0_161, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|~v1_funct_1(X3)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~l1_metric_1(X2)|~l1_metric_1(X1)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))), inference(evalgc,[status(thm)],[c_0_155])).
cnf(c_0_162, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(X1)|v2_struct_0(X2)|r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk7_3(X1,X2,esk2_0),esk5_2(esk7_3(X1,X2,esk2_0),X3))),k9_metric_1(k8_metric_1,esk8_3(X1,X2,esk2_0),X3))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k3_pcomps_1(X1),k3_pcomps_1(X2))|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(esk7_3(X1,X2,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~v1_xreal_0(esk8_3(X1,X2,esk2_0))|~m1_subset_1(X3,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_157])])).
cnf(c_0_163, plain, (v9_metric_1(k8_metric_1)), inference(split_conjunct,[status(thm)],[fc6_metric_1])).
cnf(c_0_164, plain, (v8_metric_1(k8_metric_1)), inference(split_conjunct,[status(thm)],[fc6_metric_1])).
cnf(c_0_165, plain, (v7_metric_1(k8_metric_1)), inference(split_conjunct,[status(thm)],[fc6_metric_1])).
cnf(c_0_166, plain, (v6_metric_1(k8_metric_1)), inference(split_conjunct,[status(thm)],[fc6_metric_1])).
cnf(c_0_167, plain, (m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_158])])).
cnf(c_0_168, negated_conjecture, (v7_ordinal1(esk1_0)), inference(evalgc,[status(thm)],[c_0_159])).
cnf(c_0_169, plain, (~v2_struct_0(k8_metric_1)), inference(split_conjunct,[status(thm)],[c_0_91])).
fof(c_0_170, plain, ![X55]:(~v7_ordinal1(X55)|u1_struct_0(k14_euclid(X55))=u1_struct_0(k15_euclid(X55))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t67_euclid])])).
cnf(c_0_171, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(X2)|v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(X2))|~l1_metric_1(k14_euclid(X1))|~l1_metric_1(X2)|~v9_metric_1(k14_euclid(X1))|~v9_metric_1(X2)|~v8_metric_1(k14_euclid(X1))|~v8_metric_1(X2)|~v7_metric_1(k14_euclid(X1))|~v7_metric_1(X2)|~v6_metric_1(k14_euclid(X1))|~v6_metric_1(X2)|~r1_tarski(k7_relat_1(X3,k1_topreal9(X1,esk7_3(k14_euclid(X1),X2,X3),X4)),k9_metric_1(X2,esk8_3(k14_euclid(X1),X2,X3),esk9_3(k14_euclid(X1),X2,X3)))|~v2_xxreal_0(X4)|~v1_xreal_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(esk7_3(k14_euclid(X1),X2,X3),u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)), inference(csr,[status(thm)],[c_0_160, c_0_161])).
cnf(c_0_172, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(X1)|v2_struct_0(X2)|r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk7_3(X2,X1,esk2_0),esk5_2(esk7_3(X2,X1,esk2_0),X3))),k9_metric_1(k8_metric_1,esk8_3(X2,X1,esk2_0),X3))|v5_pre_topc(esk2_0,k3_pcomps_1(X2),k3_pcomps_1(X1))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~l1_metric_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X1)|~v9_metric_1(X2)|~v8_metric_1(X1)|~v8_metric_1(X2)|~v7_metric_1(X1)|~v7_metric_1(X2)|~v6_metric_1(X1)|~v6_metric_1(X2)|~v2_xxreal_0(X3)|~v1_xreal_0(esk8_3(X2,X1,esk2_0))|~v1_xreal_0(X3)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X1)))))|~m1_subset_1(esk7_3(X2,X1,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X1)))), inference(csr,[status(thm)],[c_0_162, c_0_140])).
cnf(c_0_173, plain, (v9_metric_1(k8_metric_1)), inference(evalgc,[status(thm)],[c_0_163])).
cnf(c_0_174, plain, (v8_metric_1(k8_metric_1)), inference(evalgc,[status(thm)],[c_0_164])).
cnf(c_0_175, plain, (v7_metric_1(k8_metric_1)), inference(evalgc,[status(thm)],[c_0_165])).
cnf(c_0_176, plain, (v6_metric_1(k8_metric_1)), inference(evalgc,[status(thm)],[c_0_166])).
cnf(c_0_177, negated_conjecture, (m1_subset_1(esk1_0,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_168])])).
cnf(c_0_178, plain, (~v2_struct_0(k8_metric_1)), inference(evalgc,[status(thm)],[c_0_169])).
cnf(c_0_179, negated_conjecture, (v1_xreal_0(esk5_2(X1,X2))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v1_xreal_0(X2)|~v2_xxreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_180, negated_conjecture, (v2_xxreal_0(esk5_2(X1,X2))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v1_xreal_0(X2)|~v2_xxreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_181, plain, (u1_struct_0(k14_euclid(X1))=u1_struct_0(k15_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_170])).
fof(c_0_182, plain, ![X51]:((((((v1_metric_1(k14_euclid(X51))|~v7_ordinal1(X51))&(v6_metric_1(k14_euclid(X51))|~v7_ordinal1(X51)))&(v7_metric_1(k14_euclid(X51))|~v7_ordinal1(X51)))&(v8_metric_1(k14_euclid(X51))|~v7_ordinal1(X51)))&(v9_metric_1(k14_euclid(X51))|~v7_ordinal1(X51)))&(l1_metric_1(k14_euclid(X51))|~v7_ordinal1(X51))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k14_euclid])])])).
fof(c_0_183, plain, ![X76]:((((((~v2_struct_0(k14_euclid(X76))|~v7_ordinal1(X76))&(v1_metric_1(k14_euclid(X76))|~v7_ordinal1(X76)))&(v6_metric_1(k14_euclid(X76))|~v7_ordinal1(X76)))&(v7_metric_1(k14_euclid(X76))|~v7_ordinal1(X76)))&(v8_metric_1(k14_euclid(X76))|~v7_ordinal1(X76)))&(v9_metric_1(k14_euclid(X76))|~v7_ordinal1(X76))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc4_euclid])])])])).
fof(c_0_184, plain, ![X93, X94, X95]:(~v3_membered(X94)|(~m1_subset_1(X95,k1_zfmisc_1(k2_zfmisc_1(X93,X94)))|v3_valued_0(X95))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc18_valued_0])])])).
cnf(c_0_185, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v2_xxreal_0(esk5_2(esk7_3(k14_euclid(esk1_0),k8_metric_1,esk2_0),esk9_3(k14_euclid(esk1_0),k8_metric_1,esk2_0)))|~v1_xreal_0(esk5_2(esk7_3(k14_euclid(esk1_0),k8_metric_1,esk2_0),esk9_3(k14_euclid(esk1_0),k8_metric_1,esk2_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk7_3(k14_euclid(esk1_0),k8_metric_1,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v2_xxreal_0(esk9_3(k14_euclid(esk1_0),k8_metric_1,esk2_0))|~v1_xreal_0(esk8_3(k14_euclid(esk1_0),k8_metric_1,esk2_0))|~v1_xreal_0(esk9_3(k14_euclid(esk1_0),k8_metric_1,esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_172]), c_0_97]), c_0_96]), c_0_173]), c_0_174]), c_0_175]), c_0_176]), c_0_97]), c_0_120]), c_0_177]), c_0_97]), c_0_120]), c_0_117])]), c_0_178])])).
cnf(c_0_186, negated_conjecture, (v1_xreal_0(esk5_2(X1,X2))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v1_xreal_0(X2)|~v2_xxreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))), inference(evalgc,[status(thm)],[c_0_179])).
cnf(c_0_187, negated_conjecture, (v2_xxreal_0(esk5_2(X1,X2))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v1_xreal_0(X2)|~v2_xxreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))), inference(evalgc,[status(thm)],[c_0_180])).
cnf(c_0_188, plain, (v2_xxreal_0(esk9_3(X1,X2,X3))|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_189, plain, (u1_struct_0(k14_euclid(X1))=u1_struct_0(k15_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_181])).
cnf(c_0_190, plain, (v6_metric_1(k14_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_182])).
cnf(c_0_191, plain, (v7_metric_1(k14_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_182])).
cnf(c_0_192, plain, (v8_metric_1(k14_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_182])).
cnf(c_0_193, plain, (v9_metric_1(k14_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_182])).
cnf(c_0_194, plain, (l1_metric_1(k14_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_182])).
cnf(c_0_195, plain, (~v2_struct_0(k14_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_183])).
fof(c_0_196, plain, ![X89, X90]:(~v1_relat_1(X89)|~v1_funct_1(X89)|~v3_valued_0(X89)|v1_xreal_0(k1_funct_1(X89,X90))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc63_valued_0])])).
cnf(c_0_197, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_184])).
fof(c_0_198, plain, ![X103, X104]:(~v1_relat_1(X103)|(~m1_subset_1(X104,k1_zfmisc_1(X103))|v1_relat_1(X104))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_199, plain, ![X100, X101]:v1_relat_1(k2_zfmisc_1(X100,X101)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
cnf(c_0_200, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v2_xxreal_0(esk9_3(k14_euclid(esk1_0),k8_metric_1,esk2_0))|~v1_xreal_0(esk8_3(k14_euclid(esk1_0),k8_metric_1,esk2_0))|~v1_xreal_0(esk9_3(k14_euclid(esk1_0),k8_metric_1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk7_3(k14_euclid(esk1_0),k8_metric_1,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_185, c_0_186]), c_0_187])).
cnf(c_0_201, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v2_xxreal_0(esk9_3(X1,X2,X3))|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v1_funct_1(X3)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~l1_metric_1(X2)|~l1_metric_1(X1)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))), inference(evalgc,[status(thm)],[c_0_188])).
cnf(c_0_202, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(X2)|v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(X2))|m1_subset_1(esk7_3(k14_euclid(X1),X2,X3),u1_struct_0(k15_euclid(X1)))|~l1_metric_1(X2)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(X2)|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(X2)|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(X2)|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(X2)|~v6_metric_1(k14_euclid(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_189])])).
cnf(c_0_203, plain, (v6_metric_1(k14_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_190])).
cnf(c_0_204, plain, (v7_metric_1(k14_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_191])).
cnf(c_0_205, plain, (v8_metric_1(k14_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_192])).
cnf(c_0_206, plain, (v9_metric_1(k14_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_193])).
cnf(c_0_207, plain, (l1_metric_1(k14_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_194])).
cnf(c_0_208, plain, (~v7_ordinal1(X1)|~v2_struct_0(k14_euclid(X1))), inference(evalgc,[status(thm)],[c_0_195])).
cnf(c_0_209, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_196])).
cnf(c_0_210, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_197])).
cnf(c_0_211, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
cnf(c_0_212, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_198])).
cnf(c_0_213, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_199])).
fof(c_0_214, plain, ![X39, X40, X41, X42, X43, X44]:(~v2_pre_topc(X39)|~l1_pre_topc(X39)|(~v2_pre_topc(X40)|~l1_pre_topc(X40)|(~v2_pre_topc(X41)|~l1_pre_topc(X41)|(~v2_pre_topc(X42)|~l1_pre_topc(X42)|(~v1_funct_1(X43)|~v1_funct_2(X43,u1_struct_0(X39),u1_struct_0(X41))|~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X39),u1_struct_0(X41))))|(~v1_funct_1(X44)|~v1_funct_2(X44,u1_struct_0(X40),u1_struct_0(X42))|~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X40),u1_struct_0(X42))))|(g1_pre_topc(u1_struct_0(X39),u1_pre_topc(X39))!=g1_pre_topc(u1_struct_0(X40),u1_pre_topc(X40))|g1_pre_topc(u1_struct_0(X41),u1_pre_topc(X41))!=g1_pre_topc(u1_struct_0(X42),u1_pre_topc(X42))|X43!=X44|~v5_pre_topc(X43,X39,X41)|v5_pre_topc(X44,X40,X42)))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t36_yellow12])])])).
cnf(c_0_215, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(esk8_3(k14_euclid(esk1_0),k8_metric_1,esk2_0))|~v1_xreal_0(esk9_3(k14_euclid(esk1_0),k8_metric_1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk7_3(k14_euclid(esk1_0),k8_metric_1,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_201]), c_0_97]), c_0_96]), c_0_173]), c_0_174]), c_0_175]), c_0_176]), c_0_97]), c_0_120]), c_0_97]), c_0_120]), c_0_117])]), c_0_178])])).
cnf(c_0_216, plain, (v2_struct_0(X1)|v5_pre_topc(X2,k3_pcomps_1(k14_euclid(X3)),k3_pcomps_1(X1))|m1_subset_1(esk7_3(k14_euclid(X3),X1,X2),u1_struct_0(k15_euclid(X3)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X3))),u1_struct_0(k3_pcomps_1(X1)))))|~v1_funct_2(X2,u1_struct_0(k3_pcomps_1(k14_euclid(X3))),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X2)|~v7_ordinal1(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_202, c_0_203]), c_0_204]), c_0_205]), c_0_206]), c_0_207]), c_0_208])).
cnf(c_0_217, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_209])).
cnf(c_0_218, negated_conjecture, (v3_valued_0(esk2_0)|~v3_membered(u1_struct_0(k3_topmetr))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_115])])).
cnf(c_0_219, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_211])).
cnf(c_0_220, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_212])).
cnf(c_0_221, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(evalgc,[status(thm)],[c_0_213])).
fof(c_0_222, plain, ((~v2_struct_0(k3_topmetr)&v1_pre_topc(k3_topmetr))&v2_pre_topc(k3_topmetr)), inference(fof_simplification,[status(thm)],[fc2_topmetr])).
cnf(c_0_223, plain, (v5_pre_topc(X6,X2,X4)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~l1_pre_topc(X3)|~v2_pre_topc(X4)|~l1_pre_topc(X4)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X3))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))|~v1_funct_1(X6)|~v1_funct_2(X6,u1_struct_0(X2),u1_struct_0(X4))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X4))))|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))|g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))!=g1_pre_topc(u1_struct_0(X4),u1_pre_topc(X4))|X5!=X6|~v5_pre_topc(X5,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_214])).
cnf(c_0_224, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(esk8_3(k14_euclid(esk1_0),k8_metric_1,esk2_0))|~v1_xreal_0(esk9_3(k14_euclid(esk1_0),k8_metric_1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_216]), c_0_97]), c_0_96]), c_0_173]), c_0_174]), c_0_175]), c_0_176]), c_0_97]), c_0_120]), c_0_97]), c_0_120]), c_0_117]), c_0_168])]), c_0_178])])).
cnf(c_0_225, plain, (v1_xreal_0(esk8_3(X2,X3,X1))|v2_struct_0(X2)|v2_struct_0(X3)|v5_pre_topc(X1,k3_pcomps_1(X2),k3_pcomps_1(X3))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~l1_metric_1(X3)|~l1_metric_1(X2)|~v9_metric_1(X3)|~v9_metric_1(X2)|~v8_metric_1(X3)|~v8_metric_1(X2)|~v7_metric_1(X3)|~v7_metric_1(X2)|~v6_metric_1(X3)|~v6_metric_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X3)))))|~v1_funct_2(X1,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X3)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_138])])).
cnf(c_0_226, negated_conjecture, (v3_valued_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_218, c_0_120]), c_0_219])])])).
cnf(c_0_227, negated_conjecture, (v1_relat_1(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_115]), c_0_221])])])).
cnf(c_0_228, plain, (v1_xreal_0(esk9_3(X1,X2,X3))|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_229, plain, (v1_pre_topc(k3_topmetr)), inference(split_conjunct,[status(thm)],[c_0_222])).
cnf(c_0_230, plain, (l1_pre_topc(k3_topmetr)), inference(split_conjunct,[status(thm)],[dt_k3_topmetr])).
cnf(c_0_231, plain, (v5_pre_topc(X6,X2,X4)|X5!=X6|g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))!=g1_pre_topc(u1_struct_0(X4),u1_pre_topc(X4))|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))|~v1_funct_1(X6)|~v1_funct_1(X5)|~v2_pre_topc(X4)|~v2_pre_topc(X3)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X4)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v5_pre_topc(X5,X1,X3)|~v1_funct_2(X6,u1_struct_0(X2),u1_struct_0(X4))|~v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X3))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X4))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))), inference(evalgc,[status(thm)],[c_0_223])).
cnf(c_0_232, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(esk9_3(k14_euclid(esk1_0),k8_metric_1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_224, c_0_225]), c_0_97]), c_0_226]), c_0_227]), c_0_96]), c_0_173]), c_0_174]), c_0_175]), c_0_176]), c_0_97]), c_0_120]), c_0_97]), c_0_120]), c_0_117])]), c_0_178])])).
cnf(c_0_233, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v1_xreal_0(esk9_3(X1,X2,X3))|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v1_funct_1(X3)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~l1_metric_1(X2)|~l1_metric_1(X1)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))), inference(evalgc,[status(thm)],[c_0_228])).
cnf(c_0_234, plain, (v1_pre_topc(k3_topmetr)), inference(evalgc,[status(thm)],[c_0_229])).
cnf(c_0_235, plain, (l1_pre_topc(k3_topmetr)), inference(evalgc,[status(thm)],[c_0_230])).
cnf(c_0_236, plain, (v2_pre_topc(k3_topmetr)), inference(split_conjunct,[status(thm)],[dt_k3_topmetr])).
cnf(c_0_237, plain, (v5_pre_topc(X1,X2,X3)|g1_pre_topc(u1_struct_0(X4),u1_pre_topc(X4))!=g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))|g1_pre_topc(u1_struct_0(X5),u1_pre_topc(X5))!=g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))|~l1_pre_topc(X3)|~l1_pre_topc(X4)|~l1_pre_topc(X2)|~l1_pre_topc(X5)|~v2_pre_topc(X3)|~v2_pre_topc(X4)|~v2_pre_topc(X2)|~v2_pre_topc(X5)|~v5_pre_topc(X1,X5,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X4))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_2(X1,u1_struct_0(X5),u1_struct_0(X4))|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_231])])).
cnf(c_0_238, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_233]), c_0_97]), c_0_96]), c_0_173]), c_0_174]), c_0_175]), c_0_176]), c_0_97]), c_0_120]), c_0_97]), c_0_120]), c_0_117])]), c_0_178])])).
cnf(c_0_239, plain, (g1_pre_topc(k1_numbers,u1_pre_topc(k3_topmetr))=k3_topmetr), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_120]), c_0_234]), c_0_235])])])).
cnf(c_0_240, plain, (v2_pre_topc(k3_topmetr)), inference(evalgc,[status(thm)],[c_0_236])).
cnf(c_0_241, plain, (v2_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_242, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|g1_pre_topc(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~l1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0)))|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(k3_pcomps_1(k14_euclid(esk1_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_238]), c_0_120]), c_0_239]), c_0_235]), c_0_240]), c_0_120]), c_0_120]), c_0_117])])])).
cnf(c_0_243, plain, (v2_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[c_0_241])).
fof(c_0_244, plain, ![X47, X48]:(((g1_pre_topc(u1_struct_0(X48),u1_pre_topc(X48))=k3_pcomps_1(k14_euclid(X47))|X48!=k15_euclid(X47)|(~v5_rltopsp1(X48)|~l1_rltopsp1(X48))|~v7_ordinal1(X47))&(g1_rlvect_1(u1_struct_0(X48),u2_struct_0(X48),u1_algstr_0(X48),u1_rlvect_1(X48))=k10_funcsdom(k2_finseq_1(X47))|X48!=k15_euclid(X47)|(~v5_rltopsp1(X48)|~l1_rltopsp1(X48))|~v7_ordinal1(X47)))&(g1_pre_topc(u1_struct_0(X48),u1_pre_topc(X48))!=k3_pcomps_1(k14_euclid(X47))|g1_rlvect_1(u1_struct_0(X48),u2_struct_0(X48),u1_algstr_0(X48),u1_rlvect_1(X48))!=k10_funcsdom(k2_finseq_1(X47))|X48=k15_euclid(X47)|(~v5_rltopsp1(X48)|~l1_rltopsp1(X48))|~v7_ordinal1(X47))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_euclid])])])])).
cnf(c_0_245, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_242, c_0_243]), c_0_81])).
cnf(c_0_246, plain, (g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=k3_pcomps_1(k14_euclid(X2))|X1!=k15_euclid(X2)|~v5_rltopsp1(X1)|~l1_rltopsp1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_244])).
fof(c_0_247, plain, ![X63]:((v5_rltopsp1(k15_euclid(X63))|~v7_ordinal1(X63))&(l1_rltopsp1(k15_euclid(X63))|~v7_ordinal1(X63))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k15_euclid])])])).
cnf(c_0_248, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0)))|~l1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_65])])).
cnf(c_0_249, plain, (g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=k3_pcomps_1(k14_euclid(X2))|X1!=k15_euclid(X2)|~v7_ordinal1(X2)|~v5_rltopsp1(X1)|~l1_rltopsp1(X1)), inference(evalgc,[status(thm)],[c_0_246])).
cnf(c_0_250, plain, (v5_rltopsp1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_247])).
cnf(c_0_251, plain, (l1_rltopsp1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_247])).
fof(c_0_252, plain, ![X61]:((l1_rlvect_1(X61)|~l1_rltopsp1(X61))&(l1_pre_topc(X61)|~l1_rltopsp1(X61))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_rltopsp1])])])).
fof(c_0_253, plain, ![X56, X57]:(~v1_xreal_0(X56)|~m1_subset_1(X57,k1_numbers)|k3_real_1(X56,X57)=k3_real_1(X57,X56)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k3_real_1])])).
cnf(c_0_254, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_248, c_0_81]), c_0_82])).
cnf(c_0_255, plain, (g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k15_euclid(X1)))=k3_pcomps_1(k14_euclid(X1))|~l1_rltopsp1(k15_euclid(X1))|~v5_rltopsp1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_249])])).
cnf(c_0_256, plain, (v5_rltopsp1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_250])).
cnf(c_0_257, plain, (l1_rltopsp1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_251])).
fof(c_0_258, plain, ![X91]:((((((((((v2_pre_topc(k15_euclid(X91))|~v7_ordinal1(X91))&(v13_algstr_0(k15_euclid(X91))|~v7_ordinal1(X91)))&(v2_rlvect_1(k15_euclid(X91))|~v7_ordinal1(X91)))&(v3_rlvect_1(k15_euclid(X91))|~v7_ordinal1(X91)))&(v4_rlvect_1(k15_euclid(X91))|~v7_ordinal1(X91)))&(v5_rlvect_1(k15_euclid(X91))|~v7_ordinal1(X91)))&(v6_rlvect_1(k15_euclid(X91))|~v7_ordinal1(X91)))&(v7_rlvect_1(k15_euclid(X91))|~v7_ordinal1(X91)))&(v8_rlvect_1(k15_euclid(X91))|~v7_ordinal1(X91)))&(v5_rltopsp1(k15_euclid(X91))|~v7_ordinal1(X91))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc6_euclid])])])).
cnf(c_0_259, plain, (l1_pre_topc(X1)|~l1_rltopsp1(X1)), inference(split_conjunct,[status(thm)],[c_0_252])).
cnf(c_0_260, plain, (k3_real_1(X1,X2)=k3_real_1(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_253])).
cnf(c_0_261, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,k3_topmetr)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_120]), c_0_239]), c_0_235]), c_0_240])])])).
cnf(c_0_262, plain, (g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k15_euclid(X1)))=k3_pcomps_1(k14_euclid(X1))|~v7_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_255, c_0_256]), c_0_257])).
cnf(c_0_263, plain, (v2_pre_topc(k15_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_258])).
cnf(c_0_264, plain, (l1_pre_topc(X1)|~l1_rltopsp1(X1)), inference(evalgc,[status(thm)],[c_0_259])).
cnf(c_0_265, plain, (k3_real_1(X1,X2)=k3_real_1(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[c_0_260])).
cnf(c_0_266, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(k15_euclid(X1))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_261, c_0_262])])).
cnf(c_0_267, plain, (v2_pre_topc(k15_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_263])).
cnf(c_0_268, plain, (l1_pre_topc(k15_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_257])])).
cnf(c_0_269, plain, (k2_rcomp_1(k5_real_1(X1,X2),k3_real_1(X2,X1))=k9_metric_1(k8_metric_1,X1,X2)|~m1_subset_1(X1,u1_struct_0(k8_metric_1))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_265])])).
cnf(c_0_270, plain, (k3_real_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_265])])).
cnf(c_0_271, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_266, c_0_267]), c_0_268])).
cnf(c_0_272, plain, (u1_struct_0(k3_pcomps_1(k14_euclid(X1)))=u1_struct_0(k14_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_207])])).
cnf(c_0_273, plain, (k2_rcomp_1(k5_real_1(X1,X2),k3_real_1(X2,X1))=k9_metric_1(k8_metric_1,X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X1,u1_struct_0(k3_topmetr))|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_269, c_0_107])])).
cnf(c_0_274, plain, (k3_real_1(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_270, c_0_140]), c_0_140])).
fof(c_0_275, plain, ![X87, X88]:(~v1_xcmplx_0(X87)|~v1_xcmplx_0(X88)|k2_xcmplx_0(X87,X88)=k2_xcmplx_0(X88,X87)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k2_xcmplx_0])])).
fof(c_0_276, plain, ![X86]:(~v1_xreal_0(X86)|v1_xcmplx_0(X86)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_xreal_0])])).
cnf(c_0_277, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_272]), c_0_168])])])).
cnf(c_0_278, negated_conjecture, (~v1_xreal_0(X1)|~v2_xxreal_0(X1)|~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0),esk4_0)))|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_279, plain, (k2_rcomp_1(k5_real_1(X1,X2),k3_real_1(X2,X1))=k9_metric_1(k8_metric_1,X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[c_0_273, c_0_120]), c_0_140])).
cnf(c_0_280, plain, (k3_real_1(X2,X1)=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_274])])).
cnf(c_0_281, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_275])).
cnf(c_0_282, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_276])).
cnf(c_0_283, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_277, c_0_206]), c_0_168])])])).
cnf(c_0_284, negated_conjecture, (~v1_xreal_0(X1)|~v2_xxreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0),esk4_0)))), inference(evalgc,[status(thm)],[c_0_278])).
cnf(c_0_285, plain, (k2_rcomp_1(k6_xcmplx_0(X1,X2),k3_real_1(X2,X1))=k9_metric_1(k8_metric_1,X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_279, c_0_129])])).
cnf(c_0_286, plain, (k3_real_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_280, c_0_140])).
cnf(c_0_287, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_281])).
cnf(c_0_288, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_282])).
fof(c_0_289, plain, ![X66]:(v2_struct_0(X66)|~l1_struct_0(X66)|~v1_xboole_0(u1_struct_0(X66))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_290, plain, ![X75]:(~l1_metric_1(X75)|l1_struct_0(X75)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_metric_1])])).
cnf(c_0_291, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_205]), c_0_168])])])).
cnf(c_0_292, negated_conjecture, (~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_284, c_0_120]), c_0_120]), c_0_120])])).
cnf(c_0_293, negated_conjecture, (v1_funct_2(esk2_0,u1_struct_0(k15_euclid(esk1_0)),k1_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_116, c_0_120])])).
cnf(c_0_294, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0)))|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_295, plain, (k2_rcomp_1(k6_xcmplx_0(X1,X2),k2_xcmplx_0(X2,X1))=k9_metric_1(k8_metric_1,X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_285, c_0_286])])).
cnf(c_0_296, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_287, c_0_288])])).
cnf(c_0_297, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_289])).
cnf(c_0_298, plain, (l1_struct_0(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_290])).
cnf(c_0_299, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_204]), c_0_168])])])).
cnf(c_0_300, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_114]), c_0_293]), c_0_117])])])).
cnf(c_0_301, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0)))|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)), inference(evalgc,[status(thm)],[c_0_294])).
cnf(c_0_302, plain, (k2_rcomp_1(k6_xcmplx_0(X1,X2),k2_xcmplx_0(X2,X1))=k9_metric_1(k8_metric_1,X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_295, c_0_140])).
cnf(c_0_303, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_296, c_0_288])])).
cnf(c_0_304, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_297])).
cnf(c_0_305, plain, (l1_struct_0(X1)|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[c_0_298])).
cnf(c_0_306, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_299, c_0_203]), c_0_168])])])).
cnf(c_0_307, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_300, c_0_148])]), c_0_301])).
cnf(c_0_308, plain, (k2_rcomp_1(k6_xcmplx_0(X1,X2),k2_xcmplx_0(X1,X2))=k9_metric_1(k8_metric_1,X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_303])])).
cnf(c_0_309, negated_conjecture, (v1_xreal_0(esk4_0)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_310, plain, (r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)),X3,k9_metric_1(X1,X4,esk6_6(X1,X2,X3,X4,X5,X6))),k9_metric_1(X2,X5,X6))|v2_struct_0(X2)|v2_struct_0(X1)|X5!=k1_funct_1(X3,X4)|~v1_xreal_0(X6)|~v2_xxreal_0(X6)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X1))|~v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_311, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_304, c_0_305])])).
cnf(c_0_312, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_306, c_0_189]), c_0_148]), c_0_293]), c_0_168])])])).
cnf(c_0_313, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_307, c_0_308])])).
cnf(c_0_314, negated_conjecture, (v1_xreal_0(esk4_0)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)), inference(evalgc,[status(thm)],[c_0_309])).
cnf(c_0_315, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)),X3,k9_metric_1(X1,X4,esk6_6(X1,X2,X3,X4,X5,X6))),k9_metric_1(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~v1_funct_1(X3)|~v1_xreal_0(X6)|~v2_xxreal_0(X6)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~l1_metric_1(X2)|~l1_metric_1(X1)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X1))|~v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))), inference(evalgc,[status(thm)],[c_0_310])).
cnf(c_0_316, plain, (v1_xreal_0(esk6_6(X1,X2,X3,X4,X5,X6))|v2_struct_0(X2)|v2_struct_0(X1)|X5!=k1_funct_1(X3,X4)|~v1_xreal_0(X6)|~v2_xxreal_0(X6)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X1))|~v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_317, plain, (v2_struct_0(k14_euclid(X1))|~v1_xboole_0(u1_struct_0(k15_euclid(X1)))|~l1_metric_1(k14_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_189])])).
cnf(c_0_318, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_312, c_0_207]), c_0_168])])])).
cnf(c_0_319, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)), inference(csr,[status(thm)],[c_0_313, c_0_314])).
cnf(c_0_320, plain, (k7_relset_1(X1,X2,X3,X4)=k7_relset_1(X5,X6,X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_145])])).
cnf(c_0_321, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(X2)|r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)),X3,k1_topreal9(X1,X4,esk6_6(k14_euclid(X1),X2,X3,X4,X5,X6))),k9_metric_1(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(X2)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(X2)|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(X2)|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(X2)|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(X2)|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~v1_xreal_0(esk6_6(k14_euclid(X1),X2,X3,X4,X5,X6))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_154])])).
cnf(c_0_322, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v1_xreal_0(esk6_6(X1,X2,X3,X4,X5,X6))|X5!=k1_funct_1(X3,X4)|~v1_funct_1(X3)|~v1_xreal_0(X6)|~v2_xxreal_0(X6)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~l1_metric_1(X2)|~l1_metric_1(X1)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X1))|~v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))), inference(evalgc,[status(thm)],[c_0_316])).
cnf(c_0_323, plain, (~v1_xboole_0(u1_struct_0(k15_euclid(X1)))|~v7_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_317, c_0_207]), c_0_208])).
cnf(c_0_324, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_148]), c_0_293]), c_0_168])])])).
cnf(c_0_325, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relset_1(X2,X3,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_319, c_0_320]), c_0_148])])])).
cnf(c_0_326, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(X2)|r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)),X3,k1_topreal9(X1,X4,esk6_6(k14_euclid(X1),X2,X3,X4,X5,X6))),k9_metric_1(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X1))|~l1_metric_1(X2)|~v9_metric_1(k14_euclid(X1))|~v9_metric_1(X2)|~v8_metric_1(k14_euclid(X1))|~v8_metric_1(X2)|~v7_metric_1(k14_euclid(X1))|~v7_metric_1(X2)|~v6_metric_1(k14_euclid(X1))|~v6_metric_1(X2)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)), inference(csr,[status(thm)],[c_0_321, c_0_322])).
cnf(c_0_327, negated_conjecture, (v2_xxreal_0(esk4_0)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_328, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_324]), c_0_168])])])).
cnf(c_0_329, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~v2_xxreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v2_xxreal_0(esk4_0)|~v1_xreal_0(esk4_0)|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0)))|~m1_subset_1(k1_funct_1(esk2_0,esk3_0),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_326]), c_0_97]), c_0_120]), c_0_96]), c_0_173]), c_0_174]), c_0_175]), c_0_176]), c_0_97]), c_0_107]), c_0_120]), c_0_177]), c_0_97]), c_0_120]), c_0_117])]), c_0_178])])).
cnf(c_0_330, negated_conjecture, (v2_xxreal_0(esk4_0)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)), inference(evalgc,[status(thm)],[c_0_327])).
cnf(c_0_331, plain, (v2_xxreal_0(esk6_6(X1,X2,X3,X4,X5,X6))|v2_struct_0(X2)|v2_struct_0(X1)|X5!=k1_funct_1(X3,X4)|~v1_xreal_0(X6)|~v2_xxreal_0(X6)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X1))|~v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_332, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_328]), c_0_168])])])).
cnf(c_0_333, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v2_xxreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v1_xreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_329, c_0_140]), c_0_301]), c_0_314]), c_0_330])).
cnf(c_0_334, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v2_xxreal_0(esk6_6(X1,X2,X3,X4,X5,X6))|X5!=k1_funct_1(X3,X4)|~v1_funct_1(X3)|~v1_xreal_0(X6)|~v2_xxreal_0(X6)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~l1_metric_1(X2)|~l1_metric_1(X1)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X1))|~v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))), inference(evalgc,[status(thm)],[c_0_331])).
fof(c_0_335, plain, ![X85]:(~m1_subset_1(X85,k1_numbers)|v1_xreal_0(X85)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_xreal_0])])).
cnf(c_0_336, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|g1_pre_topc(u1_struct_0(k15_euclid(esk1_0)),u1_pre_topc(k15_euclid(esk1_0)))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~l1_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_332]), c_0_120]), c_0_239]), c_0_235]), c_0_240]), c_0_120]), c_0_148]), c_0_120]), c_0_293]), c_0_117])])])).
fof(c_0_337, plain, ![X146]:((~v2_struct_0(g1_pre_topc(u1_struct_0(X146),u1_pre_topc(X146)))|(v2_struct_0(X146)|~l1_pre_topc(X146)))&(v1_pre_topc(g1_pre_topc(u1_struct_0(X146),u1_pre_topc(X146)))|(v2_struct_0(X146)|~l1_pre_topc(X146)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc7_pre_topc])])])])).
fof(c_0_338, plain, ![X59]:(~l1_pre_topc(X59)|m1_subset_1(u1_pre_topc(X59),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X59))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_pre_topc])])).
cnf(c_0_339, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v2_xxreal_0(esk4_0)|~v1_xreal_0(esk4_0)|~m1_subset_1(k1_funct_1(esk2_0,esk3_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_333, c_0_334]), c_0_96]), c_0_173]), c_0_174]), c_0_175]), c_0_176]), c_0_97]), c_0_97]), c_0_120]), c_0_107]), c_0_120]), c_0_97]), c_0_120]), c_0_117])]), c_0_178])])).
cnf(c_0_340, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_335])).
cnf(c_0_341, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_336, c_0_262]), c_0_168])])])).
cnf(c_0_342, plain, (v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|v2_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_337])).
fof(c_0_343, plain, ![X107]:((~v2_struct_0(k15_euclid(X107))|~v7_ordinal1(X107))&(v5_rltopsp1(k15_euclid(X107))|~v7_ordinal1(X107))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc5_euclid])])])])).
fof(c_0_344, plain, ![X145]:((v1_pre_topc(g1_pre_topc(u1_struct_0(X145),u1_pre_topc(X145)))|(~v2_pre_topc(X145)|~l1_pre_topc(X145)))&(v2_pre_topc(g1_pre_topc(u1_struct_0(X145),u1_pre_topc(X145)))|(~v2_pre_topc(X145)|~l1_pre_topc(X145)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc6_pre_topc])])])).
fof(c_0_345, plain, ![X153, X154]:((v1_pre_topc(g1_pre_topc(X153,X154))|~m1_subset_1(X154,k1_zfmisc_1(k1_zfmisc_1(X153))))&(l1_pre_topc(g1_pre_topc(X153,X154))|~m1_subset_1(X154,k1_zfmisc_1(k1_zfmisc_1(X153))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_g1_pre_topc])])])).
cnf(c_0_346, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_338])).
cnf(c_0_347, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_339, c_0_140]), c_0_314]), c_0_330])).
cnf(c_0_348, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_340])).
cnf(c_0_349, negated_conjecture, (v5_pre_topc(esk2_0,X1,k3_topmetr)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_341, c_0_120]), c_0_239]), c_0_235]), c_0_240])])])).
cnf(c_0_350, plain, (v2_struct_0(X1)|v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[c_0_342])).
cnf(c_0_351, plain, (~v2_struct_0(k15_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_343])).
cnf(c_0_352, plain, (v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_344])).
cnf(c_0_353, plain, (l1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_345])).
cnf(c_0_354, plain, (g1_pre_topc(u1_struct_0(k15_euclid(X1)),k2_pcomps_1(k14_euclid(X1)))=k3_pcomps_1(k14_euclid(X1))|~l1_metric_1(k14_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_189])])).
cnf(c_0_355, plain, (u1_struct_0(k15_euclid(X1))=X2|k3_pcomps_1(k14_euclid(X1))!=g1_pre_topc(X2,X3)|~m1_subset_1(u1_pre_topc(k15_euclid(X1)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_262])])).
cnf(c_0_356, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[c_0_346])).
cnf(c_0_357, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~m1_subset_1(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k1_funct_1(esk2_0,esk3_0),esk4_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_347, c_0_348])])).
cnf(c_0_358, plain, (m1_subset_1(esk6_6(X1,X2,X3,X4,X5,X6),k1_numbers)|v2_struct_0(X1)|v2_struct_0(X2)|X5!=k1_funct_1(X3,X4)|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_322])])).
cnf(c_0_359, negated_conjecture, (v5_pre_topc(esk2_0,X1,k3_topmetr)|X1!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers)|~v1_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_349, c_0_65])])).
cnf(c_0_360, plain, (v1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|v2_struct_0(k15_euclid(X1))|~l1_pre_topc(k15_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_350, c_0_262])])).
cnf(c_0_361, plain, (~v7_ordinal1(X1)|~v2_struct_0(k15_euclid(X1))), inference(evalgc,[status(thm)],[c_0_351])).
cnf(c_0_362, plain, (v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[c_0_352])).
cnf(c_0_363, plain, (l1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(evalgc,[status(thm)],[c_0_353])).
cnf(c_0_364, plain, (g1_pre_topc(u1_struct_0(k15_euclid(X1)),k2_pcomps_1(k14_euclid(X1)))=k3_pcomps_1(k14_euclid(X1))|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_354, c_0_207])).
cnf(c_0_365, plain, (m1_subset_1(k2_pcomps_1(k14_euclid(X1)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))|~l1_metric_1(k14_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_189])])).
cnf(c_0_366, plain, (u1_struct_0(k15_euclid(X1))=X2|k3_pcomps_1(k14_euclid(X1))!=g1_pre_topc(X2,X3)|~v7_ordinal1(X1)|~l1_pre_topc(k15_euclid(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_355, c_0_356])])).
cnf(c_0_367, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v2_xxreal_0(esk4_0)|~v1_xreal_0(esk4_0)|~m1_subset_1(k1_funct_1(esk2_0,esk3_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_357, c_0_358]), c_0_96]), c_0_173]), c_0_174]), c_0_175]), c_0_176]), c_0_97]), c_0_97]), c_0_120]), c_0_107]), c_0_120]), c_0_97]), c_0_120]), c_0_117])]), c_0_178])])).
cnf(c_0_368, negated_conjecture, (v5_pre_topc(esk2_0,X1,k3_topmetr)|X1!=k3_pcomps_1(k14_euclid(esk1_0))|~v1_pre_topc(X1)|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_359, c_0_267]), c_0_168])])])).
cnf(c_0_369, plain, (v1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~v7_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_360, c_0_268]), c_0_361])).
cnf(c_0_370, plain, (v2_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(k15_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_362, c_0_262])])).
cnf(c_0_371, plain, (l1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(k2_pcomps_1(k14_euclid(X1)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_363, c_0_364])])).
cnf(c_0_372, plain, (m1_subset_1(k2_pcomps_1(k14_euclid(X1)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_365, c_0_207])).
cnf(c_0_373, plain, (u1_struct_0(k15_euclid(X1))=X2|k3_pcomps_1(k14_euclid(X1))!=g1_pre_topc(X2,X3)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_366, c_0_268])).
cnf(c_0_374, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_367, c_0_140]), c_0_314]), c_0_330])).
cnf(c_0_375, negated_conjecture, (v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~v2_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),k1_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_368, c_0_369])])).
cnf(c_0_376, plain, (v2_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~v7_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_370, c_0_267]), c_0_268])).
cnf(c_0_377, plain, (l1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_371, c_0_372])).
cnf(c_0_378, plain, (u1_struct_0(k15_euclid(X1))=u1_struct_0(X2)|k3_pcomps_1(k14_euclid(X1))!=X2|~v7_ordinal1(X1)|~v1_pre_topc(X2)|~l1_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_65])])).
cnf(c_0_379, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_374, c_0_332])])])).
cnf(c_0_380, negated_conjecture, (v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),k1_numbers)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_375, c_0_376]), c_0_377])).
cnf(c_0_381, plain, (u1_struct_0(k3_pcomps_1(k14_euclid(X1)))=u1_struct_0(k15_euclid(X1))|~v1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~l1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_378])])).
cnf(c_0_382, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~l1_pre_topc(k15_euclid(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_379, c_0_380]), c_0_168])])])).
cnf(c_0_383, plain, (u1_struct_0(k3_pcomps_1(k14_euclid(X1)))=u1_struct_0(k15_euclid(X1))|~v7_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_381, c_0_377]), c_0_369])).
cnf(c_0_384, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_382, c_0_383]), c_0_148]), c_0_293]), c_0_168])])])).
cnf(c_0_385, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_384, c_0_217]), c_0_226]), c_0_227]), c_0_117])])])).
cnf(c_0_386, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_385, c_0_206]), c_0_168])])])).
cnf(c_0_387, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_386, c_0_205]), c_0_168])])])).
cnf(c_0_388, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_204]), c_0_168])])])).
cnf(c_0_389, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_388, c_0_203]), c_0_168])])])).
cnf(c_0_390, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_389]), c_0_168])])])).
cnf(c_0_391, negated_conjecture, (~l1_pre_topc(k15_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_390]), c_0_168])])])).
cnf(c_0_392, negated_conjecture, (~l1_pre_topc(k15_euclid(esk1_0))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_207]), c_0_168])])])).
cnf(c_0_393, negated_conjecture, (~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_392, c_0_268]), c_0_168])])])).
cnf(c_0_394, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_301, c_0_332])])])).
cnf(c_0_395, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_189]), c_0_394]), c_0_168])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 396
# Proof object clause steps            : 299
# Proof object formula steps           : 97
# Proof object conjectures             : 94
# Proof object clause conjectures      : 91
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 82
# Proof object initial formulas used   : 52
# Proof object generating inferences   : 92
# Proof object simplifying inferences  : 299
# Training examples: 184 positive, 1294 negative
# Training: Positive examples begin
cnf(c_0_152, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_97, plain, (k3_pcomps_1(k8_metric_1)=k3_topmetr)).# trainpos
cnf(c_0_168, negated_conjecture, (v7_ordinal1(esk1_0))).# trainpos
cnf(c_0_89, plain, (v1_funct_1(k7_metric_1))).# trainpos
cnf(c_0_117, negated_conjecture, (v1_funct_1(esk2_0))).# trainpos
cnf(c_0_176, plain, (v6_metric_1(k8_metric_1))).# trainpos
cnf(c_0_175, plain, (v7_metric_1(k8_metric_1))).# trainpos
cnf(c_0_174, plain, (v8_metric_1(k8_metric_1))).# trainpos
cnf(c_0_173, plain, (v9_metric_1(k8_metric_1))).# trainpos
cnf(c_0_96, plain, (l1_metric_1(k8_metric_1))).# trainpos
cnf(c_0_240, plain, (v2_pre_topc(k3_topmetr))).# trainpos
cnf(c_0_116, negated_conjecture, (v1_funct_2(esk2_0,u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr)))).# trainpos
cnf(c_0_284, negated_conjecture, (~v1_xreal_0(X1)|~v2_xxreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0),esk4_0))))).# trainpos
cnf(c_0_178, plain, (~v2_struct_0(k8_metric_1))).# trainpos
cnf(c_0_235, plain, (l1_pre_topc(k3_topmetr))).# trainpos
cnf(c_0_234, plain, (v1_pre_topc(k3_topmetr))).# trainpos
cnf(c_0_109, plain, (v1_metric_1(k8_metric_1))).# trainpos
cnf(c_0_219, plain, (v3_membered(k1_numbers))).# trainpos
cnf(c_0_87, plain, (g1_metric_1(k1_numbers,k7_metric_1)=k8_metric_1)).# trainpos
cnf(c_0_221, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).# trainpos
cnf(c_0_264, plain, (l1_pre_topc(X1)|~l1_rltopsp1(X1))).# trainpos
cnf(c_0_314, negated_conjecture, (v1_xreal_0(esk4_0)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr))).# trainpos
cnf(c_0_115, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr)))))).# trainpos
cnf(c_0_86, plain, (v1_funct_2(k7_metric_1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))).# trainpos
cnf(c_0_305, plain, (l1_struct_0(X1)|~l1_metric_1(X1))).# trainpos
cnf(c_0_330, negated_conjecture, (v2_xxreal_0(esk4_0)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr))).# trainpos
cnf(c_0_301, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0)))|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr))).# trainpos
cnf(c_0_288, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_203, plain, (v6_metric_1(k14_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_204, plain, (v7_metric_1(k14_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_205, plain, (v8_metric_1(k14_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_88, plain, (m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))))).# trainpos
cnf(c_0_57, plain, (X1=X2|g1_pre_topc(X1,X3)!=g1_pre_topc(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1))))).# trainpos
cnf(c_0_361, plain, (~v7_ordinal1(X1)|~v2_struct_0(k15_euclid(X1)))).# trainpos
cnf(c_0_208, plain, (~v7_ordinal1(X1)|~v2_struct_0(k14_euclid(X1)))).# trainpos
cnf(c_0_267, plain, (v2_pre_topc(k15_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_256, plain, (v5_rltopsp1(k15_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_206, plain, (v9_metric_1(k14_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_207, plain, (l1_metric_1(k14_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_243, plain, (v2_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1))).# trainpos
cnf(c_0_81, plain, (l1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1))).# trainpos
cnf(c_0_257, plain, (l1_rltopsp1(k15_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_82, plain, (v1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1))).# trainpos
cnf(c_0_348, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_130, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_131, plain, (r2_hidden(X1,k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_151, plain, (r2_hidden(X1,k4_ordinal1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_189, plain, (u1_struct_0(k14_euclid(X1))=u1_struct_0(k15_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_304, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).# trainpos
cnf(c_0_58, plain, (g1_pre_topc(u1_struct_0(X1),k2_pcomps_1(X1))=k3_pcomps_1(X1)|~l1_metric_1(X1))).# trainpos
cnf(c_0_186, negated_conjecture, (v1_xreal_0(esk5_2(X1,X2))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v1_xreal_0(X2)|~v2_xxreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0))))).# trainpos
cnf(c_0_356, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).# trainpos
cnf(c_0_150, plain, (k3_real_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_220, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_65, plain, (g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=X1|~l1_pre_topc(X1)|~v1_pre_topc(X1))).# trainpos
cnf(c_0_99, plain, (g1_metric_1(u1_struct_0(X1),u1_metric_1(X1))=X1|~l1_metric_1(X1)|~v1_metric_1(X1))).# trainpos
cnf(c_0_62, plain, (m1_subset_1(k2_pcomps_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_metric_1(X1))).# trainpos
cnf(c_0_363, plain, (l1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).# trainpos
cnf(c_0_287, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_265, plain, (k3_real_1(X1,X2)=k3_real_1(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_129, plain, (k5_real_1(X1,X2)=k6_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_187, negated_conjecture, (v2_xxreal_0(esk5_2(X1,X2))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v1_xreal_0(X2)|~v2_xxreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0))))).# trainpos
cnf(c_0_210, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).# trainpos
cnf(c_0_350, plain, (v2_struct_0(X1)|v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1))).# trainpos
cnf(c_0_85, plain, (X1=X2|g1_metric_1(X1,X3)!=g1_metric_1(X2,X4)|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers))))).# trainpos
cnf(c_0_217, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).# trainpos
cnf(c_0_249, plain, (g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=k3_pcomps_1(k14_euclid(X2))|X1!=k15_euclid(X2)|~v7_ordinal1(X2)|~v5_rltopsp1(X1)|~l1_rltopsp1(X1))).# trainpos
cnf(c_0_362, plain, (v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_145, plain, (k7_relset_1(X2,X3,X1,X4)=k7_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_106, plain, (k2_rcomp_1(k5_real_1(X1,X2),k3_real_1(X1,X2))=k9_metric_1(k8_metric_1,X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,u1_struct_0(k8_metric_1)))).# trainpos
cnf(c_0_146, plain, (k9_metric_1(k14_euclid(X1),X4,X2)=k1_topreal9(X1,X3,X2)|X3!=X4|~v1_xreal_0(X2)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(X1))))).# trainpos
cnf(c_0_114, plain, (k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).# trainpos
cnf(c_0_144, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v5_pre_topc(X4,k3_pcomps_1(X2),k3_pcomps_1(X3))|~v1_funct_1(X4)|~v1_xreal_0(X1)|~v2_xxreal_0(X1)|~v6_metric_1(X3)|~v6_metric_1(X2)|~v7_metric_1(X3)|~v7_metric_1(X2)|~v8_metric_1(X3)|~v8_metric_1(X2)|~v9_metric_1(X3)|~v9_metric_1(X2)|~l1_metric_1(X3)|~l1_metric_1(X2)|~v1_funct_2(X4,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X3)))))|~r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X3)),X4,k9_metric_1(X2,esk7_3(X2,X3,X4),X1)),k9_metric_1(X3,esk8_3(X2,X3,X4),esk9_3(X2,X3,X4))))).# trainpos
cnf(c_0_113, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X1),X2),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X1),X2)))|~v1_xreal_0(X2)|~v2_xxreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0))))).# trainpos
cnf(c_0_233, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v1_xreal_0(esk9_3(X1,X2,X3))|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v1_funct_1(X3)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~l1_metric_1(X2)|~l1_metric_1(X1)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2))))))).# trainpos
cnf(c_0_201, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v2_xxreal_0(esk9_3(X1,X2,X3))|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v1_funct_1(X3)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~l1_metric_1(X2)|~l1_metric_1(X1)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2))))))).# trainpos
cnf(c_0_161, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|~v1_funct_1(X3)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~l1_metric_1(X2)|~l1_metric_1(X1)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2))))))).# trainpos
cnf(c_0_138, plain, (k1_funct_1(X3,esk7_3(X1,X2,X3))=esk8_3(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v1_funct_1(X3)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~l1_metric_1(X2)|~l1_metric_1(X1)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2))))))).# trainpos
cnf(c_0_231, plain, (v5_pre_topc(X6,X2,X4)|X5!=X6|g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))!=g1_pre_topc(u1_struct_0(X4),u1_pre_topc(X4))|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))|~v1_funct_1(X6)|~v1_funct_1(X5)|~v2_pre_topc(X4)|~v2_pre_topc(X3)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X4)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v5_pre_topc(X5,X1,X3)|~v1_funct_2(X6,u1_struct_0(X2),u1_struct_0(X4))|~v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X3))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X4))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3)))))).# trainpos
cnf(c_0_322, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v1_xreal_0(esk6_6(X1,X2,X3,X4,X5,X6))|X5!=k1_funct_1(X3,X4)|~v1_funct_1(X3)|~v1_xreal_0(X6)|~v2_xxreal_0(X6)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~l1_metric_1(X2)|~l1_metric_1(X1)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X1))|~v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2))))))).# trainpos
cnf(c_0_334, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v2_xxreal_0(esk6_6(X1,X2,X3,X4,X5,X6))|X5!=k1_funct_1(X3,X4)|~v1_funct_1(X3)|~v1_xreal_0(X6)|~v2_xxreal_0(X6)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~l1_metric_1(X2)|~l1_metric_1(X1)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X1))|~v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2))))))).# trainpos
cnf(c_0_315, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)),X3,k9_metric_1(X1,X4,esk6_6(X1,X2,X3,X4,X5,X6))),k9_metric_1(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~v1_funct_1(X3)|~v1_xreal_0(X6)|~v2_xxreal_0(X6)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~l1_metric_1(X2)|~l1_metric_1(X1)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X1))|~v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2))))))).# trainpos
cnf(c_0_255, plain, (g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k15_euclid(X1)))=k3_pcomps_1(k14_euclid(X1))|~l1_rltopsp1(k15_euclid(X1))|~v5_rltopsp1(k15_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_61, plain, (u1_struct_0(X1)=X2|k3_pcomps_1(X1)!=g1_pre_topc(X2,X3)|~m1_subset_1(k2_pcomps_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_metric_1(X1))).# trainpos
cnf(c_0_296, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_268, plain, (l1_pre_topc(k15_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_227, negated_conjecture, (v1_relat_1(esk2_0))).# trainpos
cnf(c_0_218, negated_conjecture, (v3_valued_0(esk2_0)|~v3_membered(u1_struct_0(k3_topmetr)))).# trainpos
cnf(c_0_140, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_167, plain, (m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_354, plain, (g1_pre_topc(u1_struct_0(k15_euclid(X1)),k2_pcomps_1(k14_euclid(X1)))=k3_pcomps_1(k14_euclid(X1))|~l1_metric_1(k14_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_98, plain, (k1_numbers=X1|g1_metric_1(X1,X2)!=k8_metric_1)).# trainpos
cnf(c_0_69, plain, (u1_struct_0(X1)=u1_struct_0(X2)|k3_pcomps_1(X1)!=X2|~l1_metric_1(X1)|~v1_pre_topc(X2)|~l1_pre_topc(X2))).# trainpos
cnf(c_0_154, plain, (k9_metric_1(k14_euclid(X1),X2,X3)=k1_topreal9(X1,X2,X3)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_237, plain, (v5_pre_topc(X1,X2,X3)|g1_pre_topc(u1_struct_0(X4),u1_pre_topc(X4))!=g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))|g1_pre_topc(u1_struct_0(X5),u1_pre_topc(X5))!=g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))|~l1_pre_topc(X3)|~l1_pre_topc(X4)|~l1_pre_topc(X2)|~l1_pre_topc(X5)|~v2_pre_topc(X3)|~v2_pre_topc(X4)|~v2_pre_topc(X2)|~v2_pre_topc(X5)|~v5_pre_topc(X1,X5,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X4))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_2(X1,u1_struct_0(X5),u1_struct_0(X4))|~v1_funct_1(X1))).# trainpos
cnf(c_0_177, negated_conjecture, (m1_subset_1(esk1_0,k5_numbers))).# trainpos
cnf(c_0_126, negated_conjecture, (r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,X1),X2),k2_xcmplx_0(k1_funct_1(esk2_0,X1),X2)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0))))).# trainpos
cnf(c_0_303, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_320, plain, (k7_relset_1(X1,X2,X3,X4)=k7_relset_1(X5,X6,X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).# trainpos
cnf(c_0_355, plain, (u1_struct_0(k15_euclid(X1))=X2|k3_pcomps_1(k14_euclid(X1))!=g1_pre_topc(X2,X3)|~m1_subset_1(u1_pre_topc(k15_euclid(X1)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_311, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_metric_1(X1))).# trainpos
cnf(c_0_153, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~r1_tarski(k7_relat_1(X3,k9_metric_1(X1,esk7_3(X1,X2,X3),X4)),k9_metric_1(X2,esk8_3(X1,X2,X3),esk9_3(X1,X2,X3)))|~v2_xxreal_0(X4)|~v1_xreal_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3))).# trainpos
cnf(c_0_269, plain, (k2_rcomp_1(k5_real_1(X1,X2),k3_real_1(X2,X1))=k9_metric_1(k8_metric_1,X1,X2)|~m1_subset_1(X1,u1_struct_0(k8_metric_1))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_80, plain, (u1_struct_0(k3_pcomps_1(X1))=u1_struct_0(X1)|~v1_pre_topc(k3_pcomps_1(X1))|~l1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1))).# trainpos
cnf(c_0_317, plain, (v2_struct_0(k14_euclid(X1))|~v1_xboole_0(u1_struct_0(k15_euclid(X1)))|~l1_metric_1(k14_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_107, plain, (u1_struct_0(k8_metric_1)=u1_struct_0(k3_topmetr))).# trainpos
cnf(c_0_108, plain, (u1_struct_0(X1)=k1_numbers|X1!=k8_metric_1|~v1_metric_1(X1)|~l1_metric_1(X1))).# trainpos
cnf(c_0_160, plain, (v2_struct_0(X1)|v2_struct_0(k14_euclid(X2))|v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X2)),k3_pcomps_1(X1))|~l1_metric_1(X1)|~l1_metric_1(k14_euclid(X2))|~v9_metric_1(X1)|~v9_metric_1(k14_euclid(X2))|~v8_metric_1(X1)|~v8_metric_1(k14_euclid(X2))|~v7_metric_1(X1)|~v7_metric_1(k14_euclid(X2))|~v6_metric_1(X1)|~v6_metric_1(k14_euclid(X2))|~r1_tarski(k7_relat_1(X3,k1_topreal9(X2,esk7_3(k14_euclid(X2),X1,X3),X4)),k9_metric_1(X1,esk8_3(k14_euclid(X2),X1,X3),esk9_3(k14_euclid(X2),X1,X3)))|~v2_xxreal_0(X4)|~v1_xreal_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X2))),u1_struct_0(k3_pcomps_1(X1)))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X2))),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X3)|~m1_subset_1(esk7_3(k14_euclid(X2),X1,X3),u1_struct_0(k14_euclid(X2)))|~m1_subset_1(esk7_3(k14_euclid(X2),X1,X3),u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_366, plain, (u1_struct_0(k15_euclid(X1))=X2|k3_pcomps_1(k14_euclid(X1))!=g1_pre_topc(X2,X3)|~v7_ordinal1(X1)|~l1_pre_topc(k15_euclid(X1)))).# trainpos
cnf(c_0_120, plain, (u1_struct_0(k3_topmetr)=k1_numbers)).# trainpos
cnf(c_0_273, plain, (k2_rcomp_1(k5_real_1(X1,X2),k3_real_1(X2,X1))=k9_metric_1(k8_metric_1,X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X1,u1_struct_0(k3_topmetr))|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_378, plain, (u1_struct_0(k15_euclid(X1))=u1_struct_0(X2)|k3_pcomps_1(k14_euclid(X1))!=X2|~v7_ordinal1(X1)|~v1_pre_topc(X2)|~l1_pre_topc(X2))).# trainpos
cnf(c_0_119, plain, (k2_rcomp_1(k5_real_1(X1,X2),k3_real_1(X1,X2))=k9_metric_1(k8_metric_1,X1,X2)|~m1_subset_1(X1,u1_struct_0(k3_topmetr))|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_226, negated_conjecture, (v3_valued_0(esk2_0))).# trainpos
cnf(c_0_293, negated_conjecture, (v1_funct_2(esk2_0,u1_struct_0(k15_euclid(esk1_0)),k1_numbers))).# trainpos
cnf(c_0_292, negated_conjecture, (~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr))).# trainpos
cnf(c_0_148, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers))))).# trainpos
cnf(c_0_239, plain, (g1_pre_topc(k1_numbers,u1_pre_topc(k3_topmetr))=k3_topmetr)).# trainpos
cnf(c_0_137, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,X1),X2),k2_xcmplx_0(k1_funct_1(esk2_0,X1),X2)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0))))).# trainpos
cnf(c_0_365, plain, (m1_subset_1(k2_pcomps_1(k14_euclid(X1)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))|~l1_metric_1(k14_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_147, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk7_3(X1,X2,esk2_0),esk5_2(esk7_3(X1,X2,esk2_0),X3))),k2_rcomp_1(k6_xcmplx_0(esk8_3(X1,X2,esk2_0),X3),k2_xcmplx_0(esk8_3(X1,X2,esk2_0),X3)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v2_struct_0(X1)|v2_struct_0(X2)|v5_pre_topc(esk2_0,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(esk7_3(X1,X2,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2))))).# trainpos
cnf(c_0_270, plain, (k3_real_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_300, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0))))).# trainpos
cnf(c_0_139, plain, (k2_rcomp_1(k6_xcmplx_0(X1,X2),k3_real_1(X1,X2))=k9_metric_1(k8_metric_1,X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_157, plain, (k2_rcomp_1(k6_xcmplx_0(X1,X2),k2_xcmplx_0(X1,X2))=k9_metric_1(k8_metric_1,X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_156, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(X1)|v2_struct_0(X2)|r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk7_3(X2,X1,esk2_0),esk5_2(esk7_3(X2,X1,esk2_0),X3))),k2_rcomp_1(k6_xcmplx_0(esk8_3(X2,X1,esk2_0),X3),k2_xcmplx_0(esk8_3(X2,X1,esk2_0),X3)))|v5_pre_topc(esk2_0,k3_pcomps_1(X2),k3_pcomps_1(X1))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~l1_metric_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X1)|~v9_metric_1(X2)|~v8_metric_1(X1)|~v8_metric_1(X2)|~v7_metric_1(X1)|~v7_metric_1(X2)|~v6_metric_1(X1)|~v6_metric_1(X2)|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X1)))))|~m1_subset_1(esk7_3(X2,X1,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X1))))).# trainpos
cnf(c_0_280, plain, (k3_real_1(X2,X1)=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_162, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(X1)|v2_struct_0(X2)|r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk7_3(X1,X2,esk2_0),esk5_2(esk7_3(X1,X2,esk2_0),X3))),k9_metric_1(k8_metric_1,esk8_3(X1,X2,esk2_0),X3))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k3_pcomps_1(X1),k3_pcomps_1(X2))|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(esk7_3(X1,X2,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~v1_xreal_0(esk8_3(X1,X2,esk2_0))|~m1_subset_1(X3,k1_numbers))).# trainpos
cnf(c_0_185, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v2_xxreal_0(esk5_2(esk7_3(k14_euclid(esk1_0),k8_metric_1,esk2_0),esk9_3(k14_euclid(esk1_0),k8_metric_1,esk2_0)))|~v1_xreal_0(esk5_2(esk7_3(k14_euclid(esk1_0),k8_metric_1,esk2_0),esk9_3(k14_euclid(esk1_0),k8_metric_1,esk2_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk7_3(k14_euclid(esk1_0),k8_metric_1,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v2_xxreal_0(esk9_3(k14_euclid(esk1_0),k8_metric_1,esk2_0))|~v1_xreal_0(esk8_3(k14_euclid(esk1_0),k8_metric_1,esk2_0))|~v1_xreal_0(esk9_3(k14_euclid(esk1_0),k8_metric_1,esk2_0)))).# trainpos
cnf(c_0_285, plain, (k2_rcomp_1(k6_xcmplx_0(X1,X2),k3_real_1(X2,X1))=k9_metric_1(k8_metric_1,X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_215, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(esk8_3(k14_euclid(esk1_0),k8_metric_1,esk2_0))|~v1_xreal_0(esk9_3(k14_euclid(esk1_0),k8_metric_1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk7_3(k14_euclid(esk1_0),k8_metric_1,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers))).# trainpos
cnf(c_0_295, plain, (k2_rcomp_1(k6_xcmplx_0(X1,X2),k2_xcmplx_0(X2,X1))=k9_metric_1(k8_metric_1,X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_308, plain, (k2_rcomp_1(k6_xcmplx_0(X1,X2),k2_xcmplx_0(X1,X2))=k9_metric_1(k8_metric_1,X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_202, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(X2)|v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(X2))|m1_subset_1(esk7_3(k14_euclid(X1),X2,X3),u1_struct_0(k15_euclid(X1)))|~l1_metric_1(X2)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(X2)|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(X2)|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(X2)|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(X2)|~v6_metric_1(k14_euclid(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_225, plain, (v1_xreal_0(esk8_3(X2,X3,X1))|v2_struct_0(X2)|v2_struct_0(X3)|v5_pre_topc(X1,k3_pcomps_1(X2),k3_pcomps_1(X3))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~l1_metric_1(X3)|~l1_metric_1(X2)|~v9_metric_1(X3)|~v9_metric_1(X2)|~v8_metric_1(X3)|~v8_metric_1(X2)|~v7_metric_1(X3)|~v7_metric_1(X2)|~v6_metric_1(X3)|~v6_metric_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X3)))))|~v1_funct_2(X1,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X3))))).# trainpos
cnf(c_0_313, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(esk4_0))).# trainpos
cnf(c_0_224, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(esk8_3(k14_euclid(esk1_0),k8_metric_1,esk2_0))|~v1_xreal_0(esk9_3(k14_euclid(esk1_0),k8_metric_1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers))).# trainpos
cnf(c_0_325, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relset_1(X2,X3,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_360, plain, (v1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|v2_struct_0(k15_euclid(X1))|~l1_pre_topc(k15_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_370, plain, (v2_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(k15_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_232, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(esk9_3(k14_euclid(esk1_0),k8_metric_1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers))).# trainpos
cnf(c_0_358, plain, (m1_subset_1(esk6_6(X1,X2,X3,X4,X5,X6),k1_numbers)|v2_struct_0(X1)|v2_struct_0(X2)|X5!=k1_funct_1(X3,X4)|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3))).# trainpos
cnf(c_0_371, plain, (l1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(k2_pcomps_1(k14_euclid(X1)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_238, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers))).# trainpos
cnf(c_0_321, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(X2)|r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)),X3,k1_topreal9(X1,X4,esk6_6(k14_euclid(X1),X2,X3,X4,X5,X6))),k9_metric_1(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(X2)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(X2)|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(X2)|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(X2)|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(X2)|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~v1_xreal_0(esk6_6(k14_euclid(X1),X2,X3,X4,X5,X6))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_242, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|g1_pre_topc(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~l1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0)))|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(k3_pcomps_1(k14_euclid(esk1_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0)))).# trainpos
cnf(c_0_329, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~v2_xxreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v2_xxreal_0(esk4_0)|~v1_xreal_0(esk4_0)|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0)))|~m1_subset_1(k1_funct_1(esk2_0,esk3_0),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers))).# trainpos
cnf(c_0_272, plain, (u1_struct_0(k3_pcomps_1(k14_euclid(X1)))=u1_struct_0(k14_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_339, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v2_xxreal_0(esk4_0)|~v1_xreal_0(esk4_0)|~m1_subset_1(k1_funct_1(esk2_0,esk3_0),k1_numbers))).# trainpos
cnf(c_0_248, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0)))|~l1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))).# trainpos
cnf(c_0_261, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,k3_topmetr)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers))).# trainpos
cnf(c_0_357, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~m1_subset_1(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k1_funct_1(esk2_0,esk3_0),esk4_0),k1_numbers))).# trainpos
cnf(c_0_367, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v2_xxreal_0(esk4_0)|~v1_xreal_0(esk4_0)|~m1_subset_1(k1_funct_1(esk2_0,esk3_0),k1_numbers))).# trainpos
cnf(c_0_266, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(k15_euclid(X1))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_277, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_283, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_291, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_299, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_306, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_312, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_318, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_324, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr))).# trainpos
cnf(c_0_328, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr))).# trainpos
cnf(c_0_332, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr))).# trainpos
cnf(c_0_394, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0))))).# trainpos
cnf(c_0_336, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|g1_pre_topc(u1_struct_0(k15_euclid(esk1_0)),u1_pre_topc(k15_euclid(esk1_0)))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~l1_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2)))).# trainpos
cnf(c_0_341, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2)))).# trainpos
cnf(c_0_349, negated_conjecture, (v5_pre_topc(esk2_0,X1,k3_topmetr)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers))).# trainpos
cnf(c_0_381, plain, (u1_struct_0(k3_pcomps_1(k14_euclid(X1)))=u1_struct_0(k15_euclid(X1))|~v1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~l1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_359, negated_conjecture, (v5_pre_topc(esk2_0,X1,k3_topmetr)|X1!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers)|~v1_pre_topc(X1))).# trainpos
cnf(c_0_368, negated_conjecture, (v5_pre_topc(esk2_0,X1,k3_topmetr)|X1!=k3_pcomps_1(k14_euclid(esk1_0))|~v1_pre_topc(X1)|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers))).# trainpos
cnf(c_0_379, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers))).# trainpos
cnf(c_0_375, negated_conjecture, (v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~v2_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),k1_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_382, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~l1_pre_topc(k15_euclid(esk1_0)))).# trainpos
cnf(c_0_384, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0))))).# trainpos
cnf(c_0_385, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0))))).# trainpos
cnf(c_0_386, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0))))).# trainpos
cnf(c_0_387, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0))))).# trainpos
cnf(c_0_388, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0))))).# trainpos
cnf(c_0_389, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0))))).# trainpos
cnf(c_0_390, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0))))).# trainpos
cnf(c_0_391, negated_conjecture, (~l1_pre_topc(k15_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0))))).# trainpos
cnf(c_0_392, negated_conjecture, (~l1_pre_topc(k15_euclid(esk1_0))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0))))).# trainpos
cnf(c_0_393, negated_conjecture, (~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0))))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_168, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_128, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_107, plain, (~v2_struct_0(k3_topmetr))).#trainneg
cnf(i_0_124, plain, (v3_topmetr(k3_topmetr))).#trainneg
cnf(i_0_155, plain, (v3_topmetr(k8_metric_1))).#trainneg
cnf(i_0_153, plain, (v6_membered(k4_ordinal1))).#trainneg
cnf(i_0_171, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_145, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_170, plain, (m1_subset_1(esk10_1(X1),X1))).#trainneg
cnf(i_0_50, plain, (l1_rlvect_1(X1)|~l1_rltopsp1(X1))).#trainneg
cnf(i_0_22, plain, (X1=X2|g1_pre_topc(X3,X1)!=g1_pre_topc(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X3))))).#trainneg
cnf(i_0_46, plain, (v7_ordinal1(X1)|~r2_hidden(X1,k4_ordinal1))).#trainneg
cnf(i_0_138, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_91, plain, (l1_struct_0(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_119, plain, (l1_struct_0(X1)|~l1_algstr_0(X1))).#trainneg
cnf(i_0_123, plain, (l1_struct_0(X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_148, plain, (l1_algstr_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_120, plain, (v1_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_149, plain, (l2_struct_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_130, plain, (k1_finseq_1(X1)=k2_finseq_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_152, plain, (v1_funct_1(u1_rlvect_1(X1))|~l1_rlvect_1(X1))).#trainneg
cnf(i_0_110, plain, (v1_funct_1(u1_metric_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_113, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_112, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_38, plain, (v1_xreal_0(X1)|~r2_hidden(X1,k1_numbers))).#trainneg
cnf(i_0_36, plain, (v1_metric_1(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_89, plain, (v13_algstr_0(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_88, plain, (v2_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_87, plain, (v3_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_27, plain, (X1=k15_euclid(X2)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(X2))|g1_rlvect_1(u1_struct_0(X1),u2_struct_0(X1),u1_algstr_0(X1),u1_rlvect_1(X1))!=k10_funcsdom(k2_finseq_1(X2))|~v7_ordinal1(X2)|~v5_rltopsp1(X1)|~l1_rltopsp1(X1))).#trainneg
cnf(i_0_86, plain, (v4_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_131, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_85, plain, (v5_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_84, plain, (v6_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_83, plain, (v7_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_82, plain, (v8_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_103, plain, (v2_monoid_0(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_140, plain, (m1_subset_1(k2_finseq_1(X1),k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_160, plain, (m1_subset_1(u2_struct_0(X1),u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_133, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_96, plain, (v3_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~v3_topmetr(X1))).#trainneg
cnf(i_0_165, plain, (v1_finset_1(u1_struct_0(X1))|~l1_struct_0(X1)|~v8_struct_0(X1))).#trainneg
cnf(i_0_135, plain, (v2_struct_0(X1)|~l1_metric_1(X1)|~v2_struct_0(k3_pcomps_1(X1)))).#trainneg
cnf(i_0_154, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_95, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_139, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_72, plain, (X1=X2|g1_metric_1(X3,X1)!=g1_metric_1(X4,X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k2_zfmisc_1(X3,X3),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X3,X3),k1_numbers))))).#trainneg
cnf(i_0_166, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_169, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_129, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_173, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_161, plain, (v1_xboole_0(k7_relat_1(X1,X2))|~v1_xboole_0(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_159, plain, (v1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_127, plain, (k4_xxreal_1(X1,X2)=k2_rcomp_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_132, plain, (v3_valued_0(X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1))))).#trainneg
cnf(i_0_157, plain, (v1_finseq_1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1))))).#trainneg
cnf(i_0_167, plain, (m1_subset_1(k2_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_162, plain, (m1_subset_1(k3_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_137, plain, (v3_card_1(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1))))).#trainneg
cnf(i_0_147, plain, (v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))).#trainneg
cnf(i_0_126, plain, (v1_valued_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_136, plain, (m1_subset_1(k7_relset_1(X2,X3,X1,X4),k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_163, plain, (m1_subset_1(k1_topreal9(X1,X2,X3),k1_zfmisc_1(u1_struct_0(k15_euclid(X1))))|~v7_ordinal1(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1))))).#trainneg
cnf(i_0_144, plain, (v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_151, plain, (v1_funct_2(u1_rlvect_1(X1),k2_zfmisc_1(k1_numbers,u1_struct_0(X1)),u1_struct_0(X1))|~l1_rlvect_1(X1))).#trainneg
cnf(i_0_109, plain, (v1_funct_2(u1_metric_1(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)|~l1_metric_1(X1))).#trainneg
cnf(i_0_150, plain, (m1_subset_1(u1_rlvect_1(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(X1)),u1_struct_0(X1))))|~l1_rlvect_1(X1))).#trainneg
cnf(i_0_108, plain, (m1_subset_1(u1_metric_1(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)))|~l1_metric_1(X1))).#trainneg
cnf(i_0_134, plain, (g1_rlvect_1(u1_struct_0(X1),u2_struct_0(X1),u1_algstr_0(X1),u1_rlvect_1(X1))=X1|~l1_rlvect_1(X1)|~v1_rlvect_1(X1))).#trainneg
cnf(i_0_164, plain, (g1_rltopsp1(u1_struct_0(X1),u2_struct_0(X1),u1_algstr_0(X1),u1_rlvect_1(X1),u1_pre_topc(X1))=X1|~v5_rltopsp1(X1)|~l1_rltopsp1(X1))).#trainneg
cnf(i_0_172, plain, (v9_rltopsp1(k1_topreal9(X1,X2,X3),k15_euclid(X1))|~v1_xreal_0(X3)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1))))).#trainneg
cnf(i_0_28, plain, (g1_rlvect_1(u1_struct_0(X1),u2_struct_0(X1),u1_algstr_0(X1),u1_rlvect_1(X1))=k10_funcsdom(k2_finseq_1(X2))|X1!=k15_euclid(X2)|~v7_ordinal1(X2)|~v5_rltopsp1(X1)|~l1_rltopsp1(X1))).#trainneg
cnf(i_0_121, plain, (l1_metric_1(g1_metric_1(X1,X2))|~v1_funct_1(X2)|~v1_funct_2(X2,k2_zfmisc_1(X1,X1),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers))))).#trainneg
cnf(i_0_122, plain, (v1_metric_1(g1_metric_1(X1,X2))|~v1_funct_1(X2)|~v1_funct_2(X2,k2_zfmisc_1(X1,X1),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers))))).#trainneg
cnf(i_0_58, plain, (v1_xboole_0(X1)|m1_subset_1(k3_funct_2(X1,X3,X2,X4),X3)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_141, plain, (l1_rlvect_1(g1_rlvect_1(X1,X2,X3,X4))|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X2,X1)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~v1_funct_2(X4,k2_zfmisc_1(k1_numbers,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,X1),X1))))).#trainneg
cnf(i_0_142, plain, (v1_rlvect_1(g1_rlvect_1(X1,X2,X3,X4))|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X2,X1)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~v1_funct_2(X4,k2_zfmisc_1(k1_numbers,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,X1),X1))))).#trainneg
cnf(i_0_16, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X2))|~v1_funct_1(X3)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~l1_metric_1(X2)|~l1_metric_1(X1)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2))))))).#trainneg
cnf(i_0_313, plain, (u1_struct_0(X1)=X2|g1_metric_1(u1_struct_0(X1),u1_metric_1(X1))!=g1_metric_1(X2,X3)|~m1_subset_1(u1_metric_1(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)))|~v1_funct_1(u1_metric_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_175, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0),esk4_0)))))).#trainneg
cnf(i_0_250, plain, (g1_metric_1(u1_struct_0(k15_euclid(X1)),u1_metric_1(k14_euclid(X1)))=k14_euclid(X1)|~v1_metric_1(k14_euclid(X1))|~l1_metric_1(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_177, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_219, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk10_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_238, plain, (v3_membered(k5_numbers))).#trainneg
cnf(i_0_229, plain, (k7_metric_1=X1|g1_metric_1(X2,X1)!=k8_metric_1)).#trainneg
cnf(i_0_241, plain, (v1_relat_1(k7_metric_1))).#trainneg
cnf(i_0_289, negated_conjecture, (v1_valued_0(esk2_0)|~v1_membered(u1_struct_0(k3_topmetr)))).#trainneg
cnf(i_0_312, plain, (u1_metric_1(X1)=X2|g1_metric_1(u1_struct_0(X1),u1_metric_1(X1))!=g1_metric_1(X3,X2)|~m1_subset_1(u1_metric_1(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)))|~v1_funct_1(u1_metric_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_215, plain, (X1=k2_pcomps_1(X2)|g1_pre_topc(X3,X1)!=k3_pcomps_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X3)))|~l1_metric_1(X2))).#trainneg
cnf(i_0_249, plain, (g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))=k14_euclid(X1)|~v1_pre_topc(k14_euclid(X1))|~l1_pre_topc(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_376, negated_conjecture, (r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X1),X2),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X1),X2)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_338, negated_conjecture, (~r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0),esk4_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr))).#trainneg
cnf(i_0_406, plain, (k3_pcomps_1(k14_euclid(X1))=k15_euclid(X1)|~v1_pre_topc(k15_euclid(X1))|~l1_pre_topc(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_411, plain, (u1_struct_0(X1)=u1_struct_0(X2)|k3_pcomps_1(X1)!=k3_pcomps_1(X2)|~l1_metric_1(X1)|~l1_metric_1(X2))).#trainneg
cnf(i_0_419, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_420, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_282, plain, (v3_valued_0(k7_metric_1))).#trainneg
cnf(i_0_191, plain, (k14_euclid(X1)=k15_euclid(X2)|g1_rlvect_1(u1_struct_0(k15_euclid(X1)),u2_struct_0(k14_euclid(X1)),u1_algstr_0(k14_euclid(X1)),u1_rlvect_1(k14_euclid(X1)))!=k10_funcsdom(k2_finseq_1(X2))|g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))!=k3_pcomps_1(k14_euclid(X2))|~l1_rltopsp1(k14_euclid(X1))|~v5_rltopsp1(k14_euclid(X1))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_244, plain, (v1_relat_1(k5_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_290, plain, (v1_valued_0(k7_metric_1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_197, plain, (v3_membered(u1_struct_0(k3_topmetr))|~l1_struct_0(k3_topmetr))).#trainneg
cnf(i_0_194, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_metric_1(X1))).#trainneg
cnf(i_0_193, plain, (m1_subset_1(u2_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X1)))|~l2_struct_0(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_195, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_412, plain, (u1_struct_0(X1)=u1_struct_0(k15_euclid(X2))|k3_pcomps_1(X1)!=k3_pcomps_1(k14_euclid(X2))|~l1_metric_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_339, plain, (m1_subset_1(k7_relat_1(X3,X4),k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_416, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0),esk4_0)))))).#trainneg
cnf(i_0_356, plain, (v2_struct_0(X1)|v5_pre_topc(X2,k3_topmetr,k3_pcomps_1(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~r1_tarski(k7_relset_1(u1_struct_0(k3_topmetr),u1_struct_0(k3_pcomps_1(X1)),X2,k9_metric_1(k8_metric_1,esk7_3(k8_metric_1,X1,X2),X3)),k9_metric_1(X1,esk8_3(k8_metric_1,X1,X2),esk9_3(k8_metric_1,X1,X2)))|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_topmetr),u1_struct_0(k3_pcomps_1(X1)))))|~v1_funct_2(X2,u1_struct_0(k3_topmetr),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_355, plain, (v2_struct_0(X1)|v5_pre_topc(X2,k3_pcomps_1(X1),k3_topmetr)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),X2,k9_metric_1(X1,esk7_3(X1,k8_metric_1,X2),X3)),k9_metric_1(k8_metric_1,esk8_3(X1,k8_metric_1,X2),esk9_3(X1,k8_metric_1,X2)))|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr))))|~v1_funct_2(X2,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr))|~v1_funct_1(X2))).#trainneg
cnf(i_0_432, negated_conjecture, (m1_subset_1(esk5_2(X1,X2),k1_numbers)|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_436, plain, (m1_subset_1(X1,k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_202, plain, (v7_ordinal1(esk10_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_196, plain, (v3_membered(u1_struct_0(k8_metric_1))|~l1_struct_0(k8_metric_1))).#trainneg
cnf(i_0_455, plain, (u1_struct_0(k15_euclid(X1))=k1_numbers|k14_euclid(X1)!=k8_metric_1|~v7_ordinal1(X1))).#trainneg
cnf(i_0_198, plain, (v1_finset_1(u1_struct_0(k15_euclid(X1)))|~v8_struct_0(k14_euclid(X1))|~l1_struct_0(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_223, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_207, plain, (v7_ordinal1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_213, plain, (v1_xreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_421, plain, (v3_valued_0(k2_zfmisc_1(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_422, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~v1_membered(X2))).#trainneg
cnf(i_0_201, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_372, negated_conjecture, (m1_subset_1(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X1),X2),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X1),X2))))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_442, plain, (u1_struct_0(k15_euclid(X1))=X2|k3_pcomps_1(k14_euclid(X1))!=g1_pre_topc(X2,X3)|~m1_subset_1(k2_pcomps_1(k14_euclid(X1)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_354, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_464, plain, (u1_struct_0(X1)=u1_struct_0(k15_euclid(X2))|k3_pcomps_1(X1)!=k14_euclid(X2)|~l1_metric_1(X1)|~v1_pre_topc(k14_euclid(X2))|~l1_pre_topc(k14_euclid(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_348, plain, (g1_rlvect_1(u1_struct_0(k15_euclid(X1)),u2_struct_0(k15_euclid(X1)),u1_algstr_0(k15_euclid(X1)),u1_rlvect_1(k15_euclid(X1)))=k10_funcsdom(k2_finseq_1(X1))|~l1_rltopsp1(k15_euclid(X1))|~v5_rltopsp1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_454, plain, (u1_pre_topc(X1)=k2_pcomps_1(X2)|X1!=k3_pcomps_1(X2)|~l1_metric_1(X2)|~m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_483, negated_conjecture, (v1_xreal_0(esk1_0))).#trainneg
cnf(i_0_499, negated_conjecture, (r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X1),X2),k2_xcmplx_0(X2,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X1))))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X1)))).#trainneg
cnf(i_0_557, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers))).#trainneg
cnf(i_0_235, plain, (v3_membered(k7_metric_1)|~v3_membered(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))).#trainneg
cnf(i_0_340, plain, (X1=k15_euclid(X2)|X1!=k10_funcsdom(k2_finseq_1(X2))|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(X2))|~l1_rltopsp1(X1)|~v5_rltopsp1(X1)|~v7_ordinal1(X2)|~v1_rlvect_1(X1)|~l1_rlvect_1(X1))).#trainneg
cnf(i_0_203, negated_conjecture, (v7_ordinal1(esk2_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr)))))).#trainneg
cnf(i_0_204, plain, (v7_ordinal1(k7_metric_1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))))).#trainneg
cnf(i_0_210, plain, (v1_xreal_0(k7_metric_1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))))).#trainneg
cnf(i_0_205, plain, (v7_ordinal1(k2_finseq_1(X1))|~v6_membered(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_206, plain, (v7_ordinal1(u2_struct_0(X1))|~v6_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_498, negated_conjecture, (r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X1),X2),k2_xcmplx_0(X2,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X1))))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X1)))).#trainneg
cnf(i_0_466, negated_conjecture, (r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,X1),X2),k2_xcmplx_0(k1_funct_1(esk2_0,X1),X2)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_468, negated_conjecture, (r1_tarski(k7_relset_1(X3,X4,esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X1),X2),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X1),X2)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_565, plain, (v2_struct_0(k14_euclid(X1))|~l1_pre_topc(k14_euclid(X1))|~v1_xboole_0(u1_struct_0(k15_euclid(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_480, plain, (u1_struct_0(k8_metric_1)=u1_struct_0(X1)|k3_pcomps_1(X1)!=k3_topmetr|~l1_metric_1(X1))).#trainneg
cnf(i_0_208, plain, (v1_xreal_0(esk10_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_234, negated_conjecture, (v3_membered(esk2_0)|~v3_membered(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_481, plain, (u1_struct_0(k14_euclid(X1))=u1_struct_0(X2)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(X2)|~l1_metric_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_209, negated_conjecture, (v1_xreal_0(esk2_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr)))))).#trainneg
cnf(i_0_560, plain, (m1_subset_1(X1,k1_numbers)|v1_xboole_0(k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_188, plain, (~r2_hidden(X1,k5_numbers)|~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_187, plain, (~r2_hidden(X1,k7_metric_1)|~v1_xboole_0(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))).#trainneg
cnf(i_0_272, plain, (v3_membered(k2_rcomp_1(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_220, negated_conjecture, (m1_subset_1(X1,k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr)))|~r2_hidden(X1,esk2_0))).#trainneg
cnf(i_0_211, plain, (v1_xreal_0(k2_finseq_1(X1))|~v3_membered(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_212, plain, (v1_xreal_0(u2_struct_0(X1))|~v3_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_494, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,X1),X2),k2_xcmplx_0(X2,k1_funct_1(esk2_0,X1))))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))|~v1_xreal_0(k1_funct_1(esk2_0,X1)))).#trainneg
cnf(i_0_545, plain, (v2_struct_0(X1)|v5_pre_topc(X2,k3_topmetr,k3_pcomps_1(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~r1_tarski(k7_relat_1(X2,k9_metric_1(k8_metric_1,esk7_3(k8_metric_1,X1,X2),X3)),k9_metric_1(X1,esk8_3(k8_metric_1,X1,X2),esk9_3(k8_metric_1,X1,X2)))|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_topmetr),u1_struct_0(k3_pcomps_1(X1)))))|~v1_funct_2(X2,u1_struct_0(k3_topmetr),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_477, negated_conjecture, (~r1_tarski(k7_relset_1(X2,X3,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0),esk4_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_496, negated_conjecture, (~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0),esk4_0),k2_xcmplx_0(esk4_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0))))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0))|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_516, plain, (m1_subset_1(k7_relset_1(X5,X6,X3,X4),k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_512, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|r1_tarski(k7_relset_1(X3,X4,esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,X1),X2),k2_xcmplx_0(k1_funct_1(esk2_0,X1),X2)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_470, plain, (u1_struct_0(k8_metric_1)=u1_struct_0(X1)|k3_topmetr!=X1|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_532, plain, (m1_subset_1(k7_relat_1(k7_metric_1,X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_186, negated_conjecture, (~r2_hidden(X1,esk2_0)|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_399, plain, (u1_struct_0(X1)=X2|X1!=g1_metric_1(X2,X3)|~l1_metric_1(X1)|~v1_metric_1(X1))).#trainneg
cnf(i_0_214, plain, (k2_pcomps_1(X1)=X2|k3_pcomps_1(X1)!=g1_pre_topc(X3,X2)|~m1_subset_1(k2_pcomps_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_metric_1(X1))).#trainneg
cnf(i_0_423, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_185, plain, (~r2_hidden(X1,esk10_1(k1_zfmisc_1(X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_296, plain, (v3_membered(k7_relset_1(X1,X2,X3,X4))|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_297, plain, (v1_relat_1(k7_relset_1(X1,X2,X3,X4))|~v1_relat_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_221, plain, (m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))|~r2_hidden(X1,k7_metric_1))).#trainneg
cnf(i_0_222, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_497, negated_conjecture, (~r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0),esk4_0),k2_xcmplx_0(esk4_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0))))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0))|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_491, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,k1_topreal9(esk1_0,esk7_3(X1,X2,esk2_0),esk5_2(esk7_3(X1,X2,esk2_0),X3))),k2_rcomp_1(k6_xcmplx_0(esk8_3(X1,X2,esk2_0),X3),k2_xcmplx_0(esk8_3(X1,X2,esk2_0),X3)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v2_struct_0(X1)|v2_struct_0(X2)|v5_pre_topc(esk2_0,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(esk7_3(X1,X2,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2))))).#trainneg
cnf(i_0_713, plain, (u1_struct_0(X1)=u1_struct_0(k15_euclid(X2))|X1!=k14_euclid(X2)|~v1_metric_1(X1)|~l1_metric_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_653, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,k1_topreal9(esk1_0,esk7_3(X1,X2,esk2_0),esk5_2(esk7_3(X1,X2,esk2_0),X3))),k2_rcomp_1(k6_xcmplx_0(esk8_3(X1,X2,esk2_0),X3),k2_xcmplx_0(X3,esk8_3(X1,X2,esk2_0))))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v2_struct_0(X1)|v2_struct_0(X2)|v5_pre_topc(esk2_0,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v2_xxreal_0(X3)|~v1_xreal_0(esk8_3(X1,X2,esk2_0))|~v1_xreal_0(X3)|~m1_subset_1(esk7_3(X1,X2,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2))))).#trainneg
cnf(i_0_707, plain, (v3_membered(k7_relat_1(k7_metric_1,X1)))).#trainneg
cnf(i_0_642, plain, (v1_xboole_0(k5_numbers)|m1_subset_1(esk10_1(k5_numbers),k1_numbers))).#trainneg
cnf(i_0_224, plain, (v7_ordinal1(u1_pre_topc(X1))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_231, plain, (v7_ordinal1(X1)|v1_xboole_0(k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_646, plain, (~v1_xboole_0(k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_721, plain, (~v1_xboole_0(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_649, negated_conjecture, (m1_subset_1(X1,k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr)))|v1_xboole_0(esk2_0)|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_291, plain, (v7_ordinal1(k7_relset_1(X1,X2,X3,X4))|~v6_membered(k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_225, plain, (v1_xreal_0(u1_pre_topc(X1))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_227, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,u1_pre_topc(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_610, negated_conjecture, (r1_tarski(k7_relset_1(X3,X4,esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X1),X2),k2_xcmplx_0(X2,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X1))))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_709, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k7_relat_1(k7_metric_1,X2)))).#trainneg
cnf(i_0_263, plain, (v3_valued_0(esk10_1(u1_struct_0(k15_euclid(X1))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_722, plain, (~v1_xboole_0(k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_717, plain, (k2_pcomps_1(X1)=k2_pcomps_1(X2)|k3_pcomps_1(X1)!=k3_pcomps_1(X2)|~l1_metric_1(X1)|~l1_metric_1(X2))).#trainneg
cnf(i_0_228, plain, (m1_subset_1(u1_pre_topc(k14_euclid(X1)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))|~l1_pre_topc(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_645, plain, (v1_xboole_0(k5_numbers)|~v1_xboole_0(k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_706, plain, (~r2_hidden(X1,k7_relat_1(k7_metric_1,X2))|~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_292, plain, (v1_xreal_0(k7_relset_1(X1,X2,X3,X4))|~v3_membered(k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_336, plain, (v1_xboole_0(k7_relset_1(X3,X4,X1,X2))|~v1_relat_1(X1)|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_233, plain, (v3_membered(esk10_1(k1_zfmisc_1(X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_236, plain, (v3_membered(k2_finseq_1(X1))|~v3_membered(k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_475, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_530, negated_conjecture, (m1_subset_1(k7_relat_1(esk2_0,X1),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_298, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relset_1(X3,X2,X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_265, plain, (v1_finseq_1(esk10_1(u1_struct_0(k15_euclid(X1))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_708, plain, (v1_relat_1(k7_relat_1(k7_metric_1,X1))|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_613, plain, (k10_funcsdom(k2_finseq_1(X1))=k15_euclid(X1)|g1_pre_topc(u1_struct_0(k10_funcsdom(k2_finseq_1(X1))),u1_pre_topc(k10_funcsdom(k2_finseq_1(X1))))!=k3_pcomps_1(k14_euclid(X1))|~v1_rlvect_1(k10_funcsdom(k2_finseq_1(X1)))|~l1_rltopsp1(k10_funcsdom(k2_finseq_1(X1)))|~v5_rltopsp1(k10_funcsdom(k2_finseq_1(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_237, plain, (v3_membered(u1_pre_topc(X1))|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_302, plain, (v3_membered(k1_topreal9(X1,X2,X3))|~v3_membered(u1_struct_0(k15_euclid(X1)))|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_303, plain, (v1_relat_1(k1_topreal9(X1,X2,X3))|~v1_relat_1(u1_struct_0(k15_euclid(X1)))|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_239, plain, (v1_relat_1(esk10_1(k1_zfmisc_1(X1)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_242, plain, (v1_relat_1(k2_finseq_1(X1))|~v1_relat_1(k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_472, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(X2)|v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(X2))|~l1_metric_1(X2)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(X2)|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(X2)|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(X2)|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(X2)|~v6_metric_1(k14_euclid(X1))|~r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)),X3,k1_topreal9(X1,esk7_3(k14_euclid(X1),X2,X3),X4)),k9_metric_1(X2,esk8_3(k14_euclid(X1),X2,X3),esk9_3(k14_euclid(X1),X2,X3)))|~v2_xxreal_0(X4)|~v1_xreal_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~m1_subset_1(esk7_3(k14_euclid(X1),X2,X3),u1_struct_0(k14_euclid(X1)))|~m1_subset_1(esk7_3(k14_euclid(X1),X2,X3),u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_606, negated_conjecture, (r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,X1),X2),k2_xcmplx_0(X2,k1_funct_1(esk2_0,X1))))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(X2)|~v1_xreal_0(k1_funct_1(esk2_0,X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_513, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(k7_relset_1(X2,X3,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k3_topmetr),esk2_0,esk3_0),esk4_0))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_480, plain, (u1_struct_0(k3_topmetr)=u1_struct_0(X1)|k3_pcomps_1(X1)!=k3_topmetr|~l1_metric_1(X1))).#trainneg
cnf(i_0_470, plain, (u1_struct_0(k3_topmetr)=u1_struct_0(X1)|k3_topmetr!=X1|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_817, plain, (v1_xboole_0(k5_numbers)|~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_243, plain, (v1_relat_1(u1_pre_topc(X1))|~v1_relat_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_779, plain, (v1_xreal_0(esk10_1(k5_numbers))|v1_xboole_0(k5_numbers))).#trainneg
cnf(i_0_782, plain, (v1_xboole_0(k5_numbers)|v7_ordinal1(esk10_1(k5_numbers)))).#trainneg
cnf(i_0_300, plain, (v1_xreal_0(k1_topreal9(X1,X2,X3))|~v3_membered(k1_zfmisc_1(u1_struct_0(k15_euclid(X1))))|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_293, plain, (v3_valued_0(k7_relset_1(X1,k2_zfmisc_1(X2,X3),X4,X5))|~v3_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_245, plain, (u1_pre_topc(X1)=X2|X1!=g1_pre_topc(X3,X2)|~m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_862, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk7_3(X1,X2,esk2_0),esk5_2(esk7_3(X1,X2,esk2_0),X3))),k2_rcomp_1(k6_xcmplx_0(esk8_3(X1,X2,esk2_0),X3),k2_xcmplx_0(X3,esk8_3(X1,X2,esk2_0))))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v2_struct_0(X1)|v2_struct_0(X2)|v5_pre_topc(esk2_0,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v2_xxreal_0(X3)|~v1_xreal_0(esk8_3(X1,X2,esk2_0))|~v1_xreal_0(X3)|~m1_subset_1(esk7_3(X1,X2,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2))))).#trainneg
cnf(i_0_866, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|r1_tarski(k7_relset_1(X3,X4,esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,X1),X2),k2_xcmplx_0(X2,k1_funct_1(esk2_0,X1))))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(k1_funct_1(esk2_0,X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_849, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relat_1(X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_434, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_247, plain, (X1=u1_struct_0(X2)|g1_pre_topc(X1,X3)!=X2|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v1_pre_topc(X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_731, plain, (m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))|v1_xboole_0(k7_metric_1)|~m1_subset_1(X1,k7_metric_1))).#trainneg
cnf(i_0_840, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0)))))).#trainneg
cnf(i_0_294, plain, (v1_valued_0(k7_relset_1(X1,k2_zfmisc_1(X2,X3),X4,X5))|~v1_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_248, plain, (u1_struct_0(X1)=X2|X1!=g1_pre_topc(X2,X3)|~m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_251, plain, (v7_ordinal1(k2_pcomps_1(X1))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_metric_1(X1))).#trainneg
cnf(i_0_1, negated_conjecture, (r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1),X2),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1),X2)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_603, plain, (u1_struct_0(k8_metric_1)=k1_numbers)).#trainneg
cnf(i_0_376, negated_conjecture, (r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1),X2),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1),X2)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_897, plain, (~v1_xboole_0(k1_numbers)|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_898, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_498, negated_conjecture, (r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1),X2),k2_xcmplx_0(X2,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1))))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_338, negated_conjecture, (~r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr))).#trainneg
cnf(i_0_957, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relat_1(k2_zfmisc_1(X3,X2),X4)))).#trainneg
cnf(i_0_951, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relat_1(esk10_1(k1_zfmisc_1(k2_zfmisc_1(X3,X2))),X4)))).#trainneg
cnf(i_0_289, negated_conjecture, (v1_valued_0(esk2_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_497, negated_conjecture, (~r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(esk4_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))))|~v2_xxreal_0(X1)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr))).#trainneg
cnf(i_0_252, plain, (v1_xreal_0(k2_pcomps_1(X1))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_metric_1(X1))).#trainneg
cnf(i_0_876, plain, (v1_finset_1(k1_numbers)|~v8_struct_0(k3_topmetr)|~l1_struct_0(k3_topmetr))).#trainneg
cnf(i_0_689, plain, (m1_subset_1(k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5),k1_zfmisc_1(X4))|~m1_subset_1(k2_zfmisc_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_254, plain, (v3_membered(k2_pcomps_1(X1))|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_metric_1(X1))).#trainneg
cnf(i_0_255, plain, (v1_relat_1(k2_pcomps_1(X1))|~v1_relat_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_metric_1(X1))).#trainneg
cnf(i_0_499, negated_conjecture, (r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1),X2),k2_xcmplx_0(X2,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1))))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_494, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,X1),X2),k2_xcmplx_0(X2,k1_funct_1(esk2_0,X1))))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(k1_funct_1(esk2_0,X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_1058, plain, (X1=u1_pre_topc(k3_topmetr)|g1_pre_topc(X2,X1)!=k3_topmetr|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_1106, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk7_3(X1,X2,esk2_0),esk5_2(esk7_3(X1,X2,esk2_0),X3))),k2_rcomp_1(k6_xcmplx_0(esk8_3(X1,X2,esk2_0),X3),k2_xcmplx_0(X3,esk8_3(X1,X2,esk2_0))))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v2_struct_0(X1)|v2_struct_0(X2)|v5_pre_topc(esk2_0,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v2_xxreal_0(X3)|~v1_xreal_0(esk8_3(X1,X2,esk2_0))|~v1_xreal_0(X3)|~m1_subset_1(esk7_3(X1,X2,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2))))).#trainneg
cnf(i_0_954, plain, (m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,k7_relat_1(u1_rlvect_1(X2),X3))|~l1_rlvect_1(X2))).#trainneg
cnf(i_0_996, plain, (g1_pre_topc(k1_numbers,k2_pcomps_1(k8_metric_1))=k3_topmetr)).#trainneg
cnf(i_0_416, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0)))))).#trainneg
cnf(i_0_256, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,k2_pcomps_1(X2))|~l1_metric_1(X2))).#trainneg
cnf(i_0_1004, plain, (~l1_pre_topc(k8_metric_1)|~v2_struct_0(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1))))).#trainneg
cnf(i_0_1092, plain, (v3_membered(k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5))|~v3_membered(X4)|~m1_subset_1(k2_zfmisc_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1093, plain, (v1_relat_1(k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5))|~v1_relat_1(X4)|~m1_subset_1(k2_zfmisc_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1067, plain, (u1_struct_0(k15_euclid(X1))=k1_numbers|k3_pcomps_1(k14_euclid(X1))!=k3_topmetr|~v7_ordinal1(X1))).#trainneg
cnf(i_0_933, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(X1)|v2_struct_0(X2)|r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk7_3(X2,X1,esk2_0),esk5_2(esk7_3(X2,X1,esk2_0),X3))),k2_rcomp_1(k6_xcmplx_0(esk8_3(X2,X1,esk2_0),X3),k2_xcmplx_0(esk8_3(X2,X1,esk2_0),X3)))|v5_pre_topc(esk2_0,k3_pcomps_1(X2),k3_pcomps_1(X1))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~l1_metric_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X1)|~v9_metric_1(X2)|~v8_metric_1(X1)|~v8_metric_1(X2)|~v7_metric_1(X1)|~v7_metric_1(X2)|~v6_metric_1(X1)|~v6_metric_1(X2)|~v2_xxreal_0(X3)|~v1_xreal_0(esk8_3(X2,X1,esk2_0))|~v1_xreal_0(X3)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X1)))))|~m1_subset_1(esk7_3(X2,X1,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X1))))).#trainneg
cnf(i_0_477, negated_conjecture, (~r1_tarski(k7_relset_1(X1,X2,esk2_0,k1_topreal9(esk1_0,esk3_0,X3)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0)))|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_496, negated_conjecture, (~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(esk4_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))))|~v2_xxreal_0(X1)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr))).#trainneg
cnf(i_0_994, plain, (v1_finset_1(k1_numbers)|~v8_struct_0(k8_metric_1)|~l1_struct_0(k8_metric_1))).#trainneg
cnf(i_0_234, negated_conjecture, (v3_membered(esk2_0)|~v3_membered(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers)))).#trainneg
cnf(i_0_1084, plain, (v7_ordinal1(k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5))|~v6_membered(k1_zfmisc_1(X4))|~m1_subset_1(k2_zfmisc_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1085, plain, (v1_xreal_0(k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5))|~v3_membered(k1_zfmisc_1(X4))|~m1_subset_1(k2_zfmisc_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_264, plain, (v3_valued_0(u2_struct_0(k15_euclid(X1)))|~v7_ordinal1(X1)|~l2_struct_0(k15_euclid(X1)))).#trainneg
cnf(i_0_266, plain, (v1_finseq_1(u2_struct_0(k15_euclid(X1)))|~v7_ordinal1(X1)|~l2_struct_0(k15_euclid(X1)))).#trainneg
cnf(i_0_1059, plain, (X1=k1_numbers|g1_pre_topc(X1,X2)!=k3_topmetr|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_1046, negated_conjecture, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k7_relat_1(esk2_0,X2)))).#trainneg
cnf(i_0_175, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0)))))).#trainneg
cnf(i_0_877, plain, (m1_subset_1(u2_struct_0(k3_topmetr),k1_numbers)|~l2_struct_0(k3_topmetr))).#trainneg
cnf(i_0_209, negated_conjecture, (v1_xreal_0(esk2_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers))))).#trainneg
cnf(i_0_203, negated_conjecture, (v7_ordinal1(esk2_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers))))).#trainneg
cnf(i_0_367, plain, (v7_ordinal1(k3_funct_2(X1,X2,X3,X4))|v1_xboole_0(X1)|~v6_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_368, plain, (v1_xreal_0(k3_funct_2(X1,X2,X3,X4))|v1_xboole_0(X1)|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_268, plain, (k3_real_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_269, plain, (v7_ordinal1(k2_rcomp_1(X1,X2))|~v6_membered(k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_610, negated_conjecture, (r1_tarski(k7_relset_1(X1,X2,esk2_0,k1_topreal9(esk1_0,X3,esk5_2(X3,X4))),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X3),X4),k2_xcmplx_0(X4,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X3))))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X4)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X3))|~v1_xreal_0(X4)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_468, negated_conjecture, (r1_tarski(k7_relset_1(X1,X2,esk2_0,k1_topreal9(esk1_0,X3,esk5_2(X3,X4))),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X3),X4),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X3),X4)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X4)|~v1_xreal_0(X4)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_1220, plain, (v1_xreal_0(u2_struct_0(k3_topmetr))|~l2_struct_0(k3_topmetr))).#trainneg
cnf(i_0_530, negated_conjecture, (m1_subset_1(k7_relat_1(esk2_0,X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1219, plain, (v7_ordinal1(u2_struct_0(k3_topmetr))|~v6_membered(k1_numbers)|~l2_struct_0(k3_topmetr))).#trainneg
cnf(i_0_513, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(k7_relset_1(X2,X3,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_270, plain, (v1_xreal_0(k2_rcomp_1(X1,X2))|~v3_membered(k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_995, plain, (m1_subset_1(u2_struct_0(k8_metric_1),k1_numbers)|~l2_struct_0(k8_metric_1))).#trainneg
cnf(i_0_299, plain, (v7_ordinal1(k1_topreal9(X1,X2,X3))|~v6_membered(k1_zfmisc_1(u1_struct_0(k15_euclid(X1))))|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_726, plain, (v3_membered(k7_relset_1(X5,X6,X3,X4))|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_273, plain, (v1_relat_1(k2_rcomp_1(X1,X2))|~v1_relat_1(k1_numbers)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_274, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k2_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_1141, plain, (u1_pre_topc(X1)=k2_pcomps_1(k8_metric_1)|X1!=k3_topmetr|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_977, plain, (u1_struct_0(g1_pre_topc(X1,X2))=X1|~v1_pre_topc(g1_pre_topc(X1,X2))|~l1_pre_topc(g1_pre_topc(X1,X2)))).#trainneg
cnf(i_0_1135, plain, (X1=k2_pcomps_1(k8_metric_1)|g1_pre_topc(X2,X1)!=k3_topmetr|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_832, plain, (u1_struct_0(k15_euclid(X1))=u1_struct_0(k15_euclid(X2))|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_884, plain, (m1_subset_1(u1_pre_topc(k3_topmetr),k1_zfmisc_1(k1_zfmisc_1(k1_numbers))))).#trainneg
cnf(i_0_186, negated_conjecture, (~r2_hidden(X1,esk2_0)|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers)))).#trainneg
cnf(i_0_275, plain, (v7_ordinal1(k3_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_998, plain, (g1_metric_1(k1_numbers,u1_metric_1(k8_metric_1))=k8_metric_1)).#trainneg
cnf(i_0_899, plain, (v7_ordinal1(u1_pre_topc(k3_topmetr))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers))))).#trainneg
cnf(i_0_900, plain, (v1_xreal_0(u1_pre_topc(k3_topmetr))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers))))).#trainneg
cnf(i_0_729, plain, (v1_relat_1(k7_relset_1(X5,X6,X3,X4))|~v1_relat_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_276, plain, (v1_xreal_0(k3_real_1(X1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_277, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1225, plain, (v1_xboole_0(X1)|v1_xreal_0(k1_funct_1(X3,X4))|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_1003, plain, (m1_subset_1(k2_pcomps_1(k8_metric_1),k1_zfmisc_1(k1_zfmisc_1(k1_numbers))))).#trainneg
cnf(i_0_220, negated_conjecture, (m1_subset_1(X1,k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers))|~r2_hidden(X1,esk2_0))).#trainneg
cnf(i_0_278, plain, (m1_subset_1(k3_real_1(X2,X1),k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_558, plain, (m1_subset_1(esk10_1(X1),k5_numbers)|~v6_membered(X1))).#trainneg
cnf(i_0_1019, plain, (v7_ordinal1(k2_pcomps_1(k8_metric_1))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers))))).#trainneg
cnf(i_0_985, plain, (v7_ordinal1(k2_pcomps_1(k3_topmetr))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_304, plain, (m1_subset_1(X1,u1_struct_0(k15_euclid(X2)))|~r2_hidden(X1,k1_topreal9(X2,X3,X4))|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_280, plain, (v3_valued_0(esk10_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v3_membered(X2))).#trainneg
cnf(i_0_916, plain, (k2_rcomp_1(k5_real_1(X1,X2),k3_real_1(X1,X2))=k9_metric_1(k8_metric_1,X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1276, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))))))).#trainneg
cnf(i_0_1270, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(k7_relset_1(X2,X3,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(esk4_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_1322, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,u1_pre_topc(k3_topmetr)))).#trainneg
cnf(i_0_1263, negated_conjecture, (v3_membered(k7_relat_1(esk2_0,X1)))).#trainneg
cnf(i_0_1330, plain, (u1_metric_1(k8_metric_1)=k7_metric_1)).#trainneg
cnf(i_0_649, negated_conjecture, (v1_xboole_0(esk2_0)|m1_subset_1(X1,k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers))|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_283, plain, (v1_pre_topc(g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1))))|v2_struct_0(k14_euclid(X1))|~l1_pre_topc(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1267, plain, (m1_subset_1(u2_struct_0(k3_topmetr),k5_numbers)|~v6_membered(k1_numbers)|~l2_struct_0(k3_topmetr))).#trainneg
cnf(i_0_1281, plain, (v1_xreal_0(u2_struct_0(k8_metric_1))|~l2_struct_0(k8_metric_1))).#trainneg
cnf(i_0_1320, plain, (v3_membered(u1_pre_topc(k3_topmetr))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1283, negated_conjecture, (v3_membered(k7_relset_1(X1,X2,esk2_0,X3))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_286, plain, (v2_struct_0(k14_euclid(X1))|~l1_pre_topc(k14_euclid(X1))|~v2_struct_0(g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_288, plain, (v1_valued_0(esk10_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_membered(X2))).#trainneg
cnf(i_0_1300, plain, (k2_pcomps_1(k8_metric_1)=u1_pre_topc(k3_topmetr))).#trainneg
cnf(i_0_1230, plain, (k2_rcomp_1(k5_real_1(X1,X2),k2_xcmplx_0(X1,X2))=k9_metric_1(k8_metric_1,X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1130, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(X1)|v2_struct_0(X2)|r1_tarski(k7_relset_1(X4,X5,esk2_0,k1_topreal9(esk1_0,esk7_3(X2,X1,esk2_0),esk5_2(esk7_3(X2,X1,esk2_0),X3))),k2_rcomp_1(k6_xcmplx_0(esk8_3(X2,X1,esk2_0),X3),k2_xcmplx_0(X3,esk8_3(X2,X1,esk2_0))))|v5_pre_topc(esk2_0,k3_pcomps_1(X2),k3_pcomps_1(X1))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~l1_metric_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X1)|~v9_metric_1(X2)|~v8_metric_1(X1)|~v8_metric_1(X2)|~v7_metric_1(X1)|~v7_metric_1(X2)|~v6_metric_1(X1)|~v6_metric_1(X2)|~v2_xxreal_0(X3)|~v1_xreal_0(esk8_3(X2,X1,esk2_0))|~v1_xreal_0(X3)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X1)))))|~m1_subset_1(esk7_3(X2,X1,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X1)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_1321, plain, (v1_relat_1(u1_pre_topc(k3_topmetr))|~v1_relat_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_545, plain, (v2_struct_0(X1)|v5_pre_topc(X2,k3_topmetr,k3_pcomps_1(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~r1_tarski(k7_relat_1(X2,k9_metric_1(k8_metric_1,esk7_3(k8_metric_1,X1,X2),X3)),k9_metric_1(X1,esk8_3(k8_metric_1,X1,X2),esk9_3(k8_metric_1,X1,X2)))|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(k3_pcomps_1(X1)))))|~v1_funct_2(X2,k1_numbers,u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_306, plain, (v2_pre_topc(g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1))))|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1280, plain, (v7_ordinal1(u2_struct_0(k8_metric_1))|~v6_membered(k1_numbers)|~l2_struct_0(k8_metric_1))).#trainneg
cnf(i_0_1086, plain, (v3_valued_0(k7_relset_1(X1,X2,k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5)),X6))|~v3_membered(X5)|~m1_subset_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5)),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_310, plain, (v1_funct_2(u1_rlvect_1(k14_euclid(X1)),k2_zfmisc_1(k1_numbers,u1_struct_0(k15_euclid(X1))),u1_struct_0(k15_euclid(X1)))|~l1_rlvect_1(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1380, plain, (v1_xreal_0(esk10_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_435, plain, (m1_subset_1(k6_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_356, plain, (v2_struct_0(X1)|v5_pre_topc(X2,k3_topmetr,k3_pcomps_1(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~r1_tarski(k7_relset_1(k1_numbers,u1_struct_0(k3_pcomps_1(X1)),X2,k9_metric_1(k8_metric_1,esk7_3(k8_metric_1,X1,X2),X3)),k9_metric_1(X1,esk8_3(k8_metric_1,X1,X2),esk9_3(k8_metric_1,X1,X2)))|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(k3_pcomps_1(X1)))))|~v1_funct_2(X2,k1_numbers,u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1264, negated_conjecture, (v1_relat_1(k7_relat_1(esk2_0,X1))|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1017, plain, (v7_ordinal1(u1_pre_topc(k8_metric_1))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_1087, plain, (v1_valued_0(k7_relset_1(X1,X2,k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5)),X6))|~v1_membered(X5)|~m1_subset_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5)),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_362, plain, (v3_membered(k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|v1_xboole_0(X1)|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_314, plain, (v1_funct_2(u1_metric_1(k14_euclid(X1)),k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X1))),k1_numbers)|~l1_metric_1(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1224, plain, (m1_subset_1(k3_funct_2(X1,X2,X3,X4),k1_numbers)|v1_xboole_0(X1)|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_1425, negated_conjecture, (r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1),X2))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1428, negated_conjecture, (r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1)))),k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1),k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1))|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))|~m1_subset_1(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1),k1_numbers))).#trainneg
cnf(i_0_1424, negated_conjecture, (r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1),X2))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1427, negated_conjecture, (r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1)))),k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1),k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1))|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))|~m1_subset_1(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1),k1_numbers))).#trainneg
cnf(i_0_1421, negated_conjecture, (~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~m1_subset_1(esk4_0,k1_numbers))).#trainneg
cnf(i_0_316, plain, (v7_ordinal1(u1_rlvect_1(X1))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(X1)),u1_struct_0(X1))))|~l1_rlvect_1(X1))).#trainneg
cnf(i_0_1487, plain, (m1_subset_1(u2_struct_0(k8_metric_1),k5_numbers)|~v6_membered(k1_numbers)|~l2_struct_0(k8_metric_1))).#trainneg
cnf(i_0_317, plain, (v1_xreal_0(u1_rlvect_1(X1))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(X1)),u1_struct_0(X1))))|~l1_rlvect_1(X1))).#trainneg
cnf(i_0_319, plain, (v3_membered(u1_rlvect_1(X1))|~v3_membered(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(X1)),u1_struct_0(X1)))|~l1_rlvect_1(X1))).#trainneg
cnf(i_0_1411, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,X1),X2))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))|~v1_xreal_0(k1_funct_1(esk2_0,X1))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1404, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(X1)|v2_struct_0(X2)|r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk7_3(X2,X1,esk2_0),esk5_2(esk7_3(X2,X1,esk2_0),X3))),k9_metric_1(k8_metric_1,esk8_3(X2,X1,esk2_0),X3))|v5_pre_topc(esk2_0,k3_pcomps_1(X2),k3_pcomps_1(X1))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~l1_metric_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X1)|~v9_metric_1(X2)|~v8_metric_1(X1)|~v8_metric_1(X2)|~v7_metric_1(X1)|~v7_metric_1(X2)|~v6_metric_1(X1)|~v6_metric_1(X2)|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X1)))))|~m1_subset_1(esk7_3(X2,X1,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X1)))|~v1_xreal_0(esk8_3(X2,X1,esk2_0))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_1407, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(X1)|v2_struct_0(X2)|r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk7_3(X2,X1,esk2_0),esk5_2(esk7_3(X2,X1,esk2_0),esk8_3(X2,X1,esk2_0)))),k9_metric_1(k8_metric_1,esk8_3(X2,X1,esk2_0),esk8_3(X2,X1,esk2_0)))|v5_pre_topc(esk2_0,k3_pcomps_1(X2),k3_pcomps_1(X1))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~l1_metric_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X1)|~v9_metric_1(X2)|~v8_metric_1(X1)|~v8_metric_1(X2)|~v7_metric_1(X1)|~v7_metric_1(X2)|~v6_metric_1(X1)|~v6_metric_1(X2)|~v2_xxreal_0(esk8_3(X2,X1,esk2_0))|~v1_xreal_0(esk8_3(X2,X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X1)))))|~m1_subset_1(esk7_3(X2,X1,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X1)))|~m1_subset_1(esk8_3(X2,X1,esk2_0),k1_numbers))).#trainneg
cnf(i_0_1528, plain, (v1_xboole_0(X1)|m1_subset_1(k1_funct_1(X3,X4),k1_numbers)|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_1484, plain, (v2_struct_0(X1)|v5_pre_topc(X2,k3_topmetr,k3_pcomps_1(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~r1_tarski(k7_relset_1(X4,X5,X2,k9_metric_1(k8_metric_1,esk7_3(k8_metric_1,X1,X2),X3)),k9_metric_1(X1,esk8_3(k8_metric_1,X1,X2),esk9_3(k8_metric_1,X1,X2)))|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(k3_pcomps_1(X1)))))|~v1_funct_2(X2,k1_numbers,u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_320, plain, (v1_relat_1(u1_rlvect_1(X1))|~l1_rlvect_1(X1))).#trainneg
cnf(i_0_1018, plain, (v1_xreal_0(u1_pre_topc(k8_metric_1))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_363, plain, (v1_relat_1(k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|v1_xboole_0(X1)|~v1_relat_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_321, plain, (m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(X2)),u1_struct_0(X2)))|~r2_hidden(X1,u1_rlvect_1(X2))|~l1_rlvect_1(X2))).#trainneg
cnf(i_0_322, plain, (v3_valued_0(u1_rlvect_1(X1))|~v3_membered(u1_struct_0(X1))|~l1_rlvect_1(X1))).#trainneg
cnf(i_0_1472, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(X1)|v2_struct_0(X2)|r1_tarski(k7_relset_1(X3,X4,esk2_0,k1_topreal9(esk1_0,esk7_3(X1,X2,esk2_0),esk5_2(esk7_3(X1,X2,esk2_0),X5))),k2_rcomp_1(k6_xcmplx_0(esk8_3(X1,X2,esk2_0),X5),k2_xcmplx_0(esk8_3(X1,X2,esk2_0),X5)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k3_pcomps_1(X1),k3_pcomps_1(X2))|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v2_xxreal_0(X5)|~v1_xreal_0(esk8_3(X1,X2,esk2_0))|~v1_xreal_0(X5)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(esk7_3(X1,X2,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2))))).#trainneg
cnf(i_0_1406, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(X1)|v2_struct_0(X2)|r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk7_3(X2,X1,esk2_0),esk5_2(esk7_3(X2,X1,esk2_0),esk8_3(X2,X1,esk2_0)))),k9_metric_1(k8_metric_1,esk8_3(X2,X1,esk2_0),esk8_3(X2,X1,esk2_0)))|v5_pre_topc(esk2_0,k3_pcomps_1(X2),k3_pcomps_1(X1))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~l1_metric_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X1)|~v9_metric_1(X2)|~v8_metric_1(X1)|~v8_metric_1(X2)|~v7_metric_1(X1)|~v7_metric_1(X2)|~v6_metric_1(X1)|~v6_metric_1(X2)|~v2_xxreal_0(esk8_3(X2,X1,esk2_0))|~v1_xreal_0(esk8_3(X2,X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X1)))))|~m1_subset_1(esk7_3(X2,X1,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X1)))|~m1_subset_1(esk8_3(X2,X1,esk2_0),k1_numbers))).#trainneg
cnf(i_0_1429, negated_conjecture, (r1_tarski(k7_relset_1(X1,X2,esk2_0,k1_topreal9(esk1_0,X3,esk5_2(X3,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X3)))),k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X3),k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X3)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X3))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X3))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(esk1_0)))|~m1_subset_1(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X3),k1_numbers))).#trainneg
cnf(i_0_1426, negated_conjecture, (r1_tarski(k7_relset_1(X1,X2,esk2_0,k1_topreal9(esk1_0,X3,esk5_2(X3,X4))),k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X3),X4))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X4)|~v1_xreal_0(X4)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(esk1_0)))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X3))|~m1_subset_1(X4,k1_numbers))).#trainneg
cnf(i_0_1420, negated_conjecture, (~r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~m1_subset_1(esk4_0,k1_numbers))).#trainneg
cnf(i_0_1412, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,X1),X2))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))|~v1_xreal_0(k1_funct_1(esk2_0,X1))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1482, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(X1)|v2_struct_0(X2)|r1_tarski(k7_relset_1(X4,X5,esk2_0,k1_topreal9(esk1_0,esk7_3(X1,X2,esk2_0),esk5_2(esk7_3(X1,X2,esk2_0),X3))),k2_rcomp_1(k6_xcmplx_0(esk8_3(X1,X2,esk2_0),X3),k2_xcmplx_0(esk8_3(X1,X2,esk2_0),X3)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k3_pcomps_1(X1),k3_pcomps_1(X2))|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(esk7_3(X1,X2,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_1415, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,k1_funct_1(esk2_0,X1)))),k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X1)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(k1_funct_1(esk2_0,X1))|~v1_xreal_0(k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))|~m1_subset_1(k1_funct_1(esk2_0,X1),k1_numbers))).#trainneg
cnf(i_0_1474, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(X1)|v2_struct_0(X2)|r1_tarski(k7_relset_1(X3,X4,esk2_0,k1_topreal9(esk1_0,esk7_3(X1,X2,esk2_0),esk5_2(esk7_3(X1,X2,esk2_0),esk8_3(X1,X2,esk2_0)))),k9_metric_1(k8_metric_1,esk8_3(X1,X2,esk2_0),esk8_3(X1,X2,esk2_0)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k3_pcomps_1(X1),k3_pcomps_1(X2))|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v2_xxreal_0(esk8_3(X1,X2,esk2_0))|~v1_xreal_0(esk8_3(X1,X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(esk7_3(X1,X2,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(esk8_3(X1,X2,esk2_0),k1_numbers))).#trainneg
cnf(i_0_1419, negated_conjecture, (~r1_tarski(k7_relset_1(X1,X2,esk2_0,k1_topreal9(esk1_0,esk3_0,X3)),k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~m1_subset_1(esk4_0,k1_numbers))).#trainneg
cnf(i_0_878, plain, (g1_pre_topc(k1_numbers,k2_pcomps_1(k3_topmetr))=k3_pcomps_1(k3_topmetr)|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_355, plain, (v2_struct_0(X1)|v5_pre_topc(X2,k3_pcomps_1(X1),k3_topmetr)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),k1_numbers,X2,k9_metric_1(X1,esk7_3(X1,k8_metric_1,X2),X3)),k9_metric_1(k8_metric_1,esk8_3(X1,k8_metric_1,X2),esk9_3(X1,k8_metric_1,X2)))|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),k1_numbers)))|~v1_funct_2(X2,u1_struct_0(k3_pcomps_1(X1)),k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_323, plain, (v1_valued_0(u1_rlvect_1(X1))|~v1_membered(u1_struct_0(X1))|~l1_rlvect_1(X1))).#trainneg
cnf(i_0_1567, plain, (v3_membered(u1_rlvect_1(k8_metric_1))|~v3_membered(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))|~l1_rlvect_1(k8_metric_1))).#trainneg
cnf(i_0_1569, plain, (v3_membered(u1_rlvect_1(k3_topmetr))|~v3_membered(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))|~l1_rlvect_1(k3_topmetr))).#trainneg
cnf(i_0_790, plain, (v7_ordinal1(k7_relset_1(X5,X6,X3,X4))|~v6_membered(k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_324, plain, (m1_subset_1(u1_rlvect_1(k14_euclid(X1)),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(k15_euclid(X1))),u1_struct_0(k15_euclid(X1)))))|~l1_rlvect_1(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1549, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk4_0,k1_numbers)|~m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_952, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k7_relat_1(u1_metric_1(X2),X3))|~l1_metric_1(X2))).#trainneg
cnf(i_0_330, plain, (v1_relat_1(u1_metric_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_999, plain, (v1_pre_topc(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_372, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|m1_subset_1(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1),X2),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1),X2))))|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_326, plain, (v7_ordinal1(u1_metric_1(X1))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)))|~l1_metric_1(X1))).#trainneg
cnf(i_0_559, plain, (m1_subset_1(esk10_1(X1),k1_numbers)|~v6_membered(X1))).#trainneg
cnf(i_0_824, plain, (v1_xreal_0(k7_relset_1(X5,X6,X3,X4))|~v3_membered(k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_327, plain, (v1_xreal_0(u1_metric_1(X1))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)))|~l1_metric_1(X1))).#trainneg
cnf(i_0_329, plain, (v3_membered(u1_metric_1(X1))|~v3_membered(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers))|~l1_metric_1(X1))).#trainneg
cnf(i_0_1413, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|r1_tarski(k7_relset_1(X1,X2,esk2_0,k1_topreal9(esk1_0,X3,esk5_2(X3,X4))),k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,X3),X4))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X4)|~v1_xreal_0(X4)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(esk1_0)))|~v1_xreal_0(k1_funct_1(esk2_0,X3))|~m1_subset_1(X4,k1_numbers))).#trainneg
cnf(i_0_941, plain, (k2_rcomp_1(k5_real_1(X1,X2),k3_real_1(X2,X1))=k9_metric_1(k8_metric_1,X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1079, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(esk4_0,k1_funct_1(esk2_0,esk3_0))))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_332, plain, (v3_valued_0(u1_metric_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_880, plain, (g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))=k3_topmetr|~v1_metric_1(k3_topmetr)|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_1668, plain, (u1_struct_0(k15_euclid(X1))=k1_numbers|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k3_topmetr)|~v7_ordinal1(X1)|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_331, plain, (m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2)),k1_numbers))|~r2_hidden(X1,u1_metric_1(X2))|~l1_metric_1(X2))).#trainneg
cnf(i_0_640, plain, (m1_subset_1(esk10_1(X1),k1_numbers)|~v3_membered(X1))).#trainneg
cnf(i_0_1383, plain, (m1_subset_1(X1,u1_struct_0(k15_euclid(X2)))|v1_xboole_0(k1_topreal9(X2,X3,X4))|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_topreal9(X2,X3,X4)))).#trainneg
cnf(i_0_333, plain, (v1_valued_0(u1_metric_1(X1))|~v1_membered(k1_numbers)|~l1_metric_1(X1))).#trainneg
cnf(i_0_334, plain, (m1_subset_1(u1_metric_1(k14_euclid(X1)),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X1))),k1_numbers)))|~l1_metric_1(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1769, plain, (k2_rcomp_1(k5_real_1(X1,X2),k3_real_1(X1,X2))=k9_metric_1(k8_metric_1,X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1770, plain, (k2_rcomp_1(k5_real_1(X1,X2),k2_xcmplx_0(X1,X2))=k9_metric_1(k8_metric_1,X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1700, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk4_0,k1_numbers))).#trainneg
cnf(i_0_1701, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relset_1(X2,X3,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk4_0,k1_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1792, plain, (v1_relat_1(u1_metric_1(k14_euclid(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_905, plain, (u1_struct_0(k14_euclid(X1))=k1_numbers|k3_pcomps_1(k14_euclid(X1))!=k3_topmetr|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1779, plain, (u1_metric_1(k3_topmetr)=k7_metric_1|k8_metric_1!=k3_topmetr|~v1_metric_1(k3_topmetr)|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_341, plain, (g1_rlvect_1(u1_struct_0(k15_euclid(X1)),u2_struct_0(k14_euclid(X1)),u1_algstr_0(k14_euclid(X1)),u1_rlvect_1(k14_euclid(X1)))=k14_euclid(X1)|~v1_rlvect_1(k14_euclid(X1))|~l1_rlvect_1(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_342, plain, (g1_rltopsp1(u1_struct_0(k15_euclid(X1)),u2_struct_0(k14_euclid(X1)),u1_algstr_0(k14_euclid(X1)),u1_rlvect_1(k14_euclid(X1)),u1_pre_topc(k14_euclid(X1)))=k14_euclid(X1)|~l1_rltopsp1(k14_euclid(X1))|~v5_rltopsp1(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1768, plain, (k2_rcomp_1(k5_real_1(X1,X2),k2_xcmplx_0(X2,X1))=k9_metric_1(k8_metric_1,X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1775, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relset_1(X2,X3,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(esk4_0,k1_funct_1(esk2_0,esk3_0))))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1759, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|r1_tarski(k7_relset_1(X1,X2,esk2_0,k1_topreal9(esk1_0,esk7_3(X3,X4,esk2_0),esk5_2(esk7_3(X3,X4,esk2_0),X5))),k9_metric_1(k8_metric_1,esk8_3(X3,X4,esk2_0),X5))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v2_struct_0(X3)|v2_struct_0(X4)|v5_pre_topc(esk2_0,k3_pcomps_1(X3),k3_pcomps_1(X4))|~v2_xxreal_0(X5)|~v1_xreal_0(esk8_3(X3,X4,esk2_0))|~v1_xreal_0(X5)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(esk7_3(X3,X4,esk2_0),u1_struct_0(k15_euclid(esk1_0)))|~l1_metric_1(X4)|~l1_metric_1(X3)|~v9_metric_1(X4)|~v9_metric_1(X3)|~v8_metric_1(X4)|~v8_metric_1(X3)|~v7_metric_1(X4)|~v7_metric_1(X3)|~v6_metric_1(X4)|~v6_metric_1(X3)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X3)),u1_struct_0(k3_pcomps_1(X4)))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(X3)),u1_struct_0(k3_pcomps_1(X4))))).#trainneg
cnf(i_0_1670, plain, (v2_struct_0(k14_euclid(X1))|v5_pre_topc(X2,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X1))|~r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),k1_numbers,X2,k1_topreal9(X1,esk7_3(k14_euclid(X1),k8_metric_1,X2),X3)),k9_metric_1(k8_metric_1,esk8_3(k14_euclid(X1),k8_metric_1,X2),esk9_3(k14_euclid(X1),k8_metric_1,X2)))|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),k1_numbers)))|~v1_funct_2(X2,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(esk7_3(k14_euclid(X1),k8_metric_1,X2),u1_struct_0(k14_euclid(X1)))|~m1_subset_1(esk7_3(k14_euclid(X1),k8_metric_1,X2),u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1776, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(esk4_0,k1_funct_1(esk2_0,esk3_0))))))).#trainneg
cnf(i_0_1794, plain, (v3_valued_0(u1_metric_1(k14_euclid(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1753, plain, (v3_membered(u1_metric_1(k3_topmetr))|~v3_membered(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_892, plain, (k15_euclid(X1)=k3_topmetr|g1_rlvect_1(k1_numbers,u2_struct_0(k3_topmetr),u1_algstr_0(k3_topmetr),u1_rlvect_1(k3_topmetr))!=k10_funcsdom(k2_finseq_1(X1))|g1_pre_topc(k1_numbers,u1_pre_topc(k3_topmetr))!=k3_pcomps_1(k14_euclid(X1))|~l1_rltopsp1(k3_topmetr)|~v5_rltopsp1(k3_topmetr)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_349, plain, (l1_metric_1(g1_metric_1(u1_struct_0(X1),u1_metric_1(X1)))|~m1_subset_1(u1_metric_1(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)))|~v1_funct_1(u1_metric_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_1780, plain, (u1_struct_0(X1)=k1_numbers|X1!=k3_topmetr|~v1_metric_1(X1)|~l1_metric_1(X1)|~v1_metric_1(k3_topmetr)|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_852, plain, (m1_subset_1(X1,X2)|v1_xboole_0(k7_relset_1(X3,X2,X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~m1_subset_1(X1,k7_relset_1(X3,X2,X4,X5)))).#trainneg
cnf(i_0_351, plain, (v1_metric_1(g1_metric_1(u1_struct_0(X1),u1_metric_1(X1)))|~m1_subset_1(u1_metric_1(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)))|~v1_funct_1(u1_metric_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_358, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v2_xxreal_0(X4)|~v1_xreal_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~m1_subset_1(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)),X3,k9_metric_1(X1,esk7_3(X1,X2,X3),X4)),k1_zfmisc_1(k9_metric_1(X2,esk8_3(X1,X2,X3),esk9_3(X1,X2,X3)))))).#trainneg
cnf(i_0_1720, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|m1_subset_1(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1),X2),k2_xcmplx_0(X2,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1)))))|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1)))).#trainneg
cnf(i_0_1081, negated_conjecture, (~r1_tarski(k7_relset_1(X2,X3,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(esk4_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))))|~v2_xxreal_0(X1)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1010, plain, (k15_euclid(X1)=k8_metric_1|g1_rlvect_1(k1_numbers,u2_struct_0(k8_metric_1),u1_algstr_0(k8_metric_1),u1_rlvect_1(k8_metric_1))!=k10_funcsdom(k2_finseq_1(X1))|g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1))!=k3_pcomps_1(k14_euclid(X1))|~l1_rltopsp1(k8_metric_1)|~v5_rltopsp1(k8_metric_1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_359, plain, (v3_valued_0(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|v1_xboole_0(X1)|~v3_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_365, plain, (v3_valued_0(k3_funct_2(X1,u1_struct_0(k15_euclid(X2)),X3,X4))|v1_xboole_0(X1)|~v7_ordinal1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_366, plain, (v1_finseq_1(k3_funct_2(X1,u1_struct_0(k15_euclid(X2)),X3,X4))|v1_xboole_0(X1)|~v7_ordinal1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_360, plain, (v1_valued_0(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|v1_xboole_0(X1)|~v1_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_364, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X3)|~r2_hidden(X1,k3_funct_2(X3,k1_zfmisc_1(X2),X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,k1_zfmisc_1(X2))))|~m1_subset_1(X5,X3)|~v1_funct_2(X4,X3,k1_zfmisc_1(X2))|~v1_funct_1(X4))).#trainneg
cnf(i_0_1931, negated_conjecture, (~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))|~v2_xxreal_0(X1)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_1693, plain, (m1_subset_1(X1,u1_struct_0(k15_euclid(X2)))|~r2_hidden(X1,k7_relat_1(u1_rlvect_1(k14_euclid(X2)),X3))|~l1_rlvect_1(k14_euclid(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1342, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1889, plain, (l1_metric_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_369, plain, (v1_xboole_0(X1)|m1_subset_1(k1_funct_1(X3,X4),X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_1555, plain, (v7_ordinal1(u1_rlvect_1(k8_metric_1))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~l1_rlvect_1(k8_metric_1))).#trainneg
cnf(i_0_1892, plain, (v1_xboole_0(k7_relset_1(X6,X7,X3,X4))|m1_subset_1(X5,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k7_relset_1(X6,X7,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))).#trainneg
cnf(i_0_370, plain, (l1_rlvect_1(g1_rlvect_1(k1_numbers,X1,X2,k7_metric_1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_371, plain, (l1_rlvect_1(g1_rlvect_1(u1_struct_0(X1),X2,X3,u1_rlvect_1(X1)))|~m1_subset_1(u1_rlvect_1(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(X1)),u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~v1_funct_2(X3,k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~v1_funct_1(u1_rlvect_1(X1))|~v1_funct_1(X3)|~l1_rlvect_1(X1))).#trainneg
cnf(i_0_1960, negated_conjecture, (~r1_tarski(k7_relset_1(X1,X2,esk2_0,k1_topreal9(esk1_0,esk3_0,X3)),k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))|~v2_xxreal_0(X3)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(X3)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_1930, negated_conjecture, (~r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))|~v2_xxreal_0(X1)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_1343, plain, (v1_xreal_0(k3_real_1(X2,X1))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1901, plain, (v1_metric_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_1659, plain, (X1=k1_numbers|g1_pre_topc(X1,X2)!=k3_pcomps_1(k3_topmetr)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_377, plain, (v1_rlvect_1(g1_rlvect_1(k1_numbers,X1,X2,k7_metric_1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1557, plain, (v7_ordinal1(u1_rlvect_1(k3_topmetr))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~l1_rlvect_1(k3_topmetr))).#trainneg
cnf(i_0_923, plain, (v3_valued_0(k7_relset_1(X6,X7,X4,X5))|~v3_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))).#trainneg
cnf(i_0_378, plain, (v1_rlvect_1(g1_rlvect_1(u1_struct_0(X1),X2,X3,u1_rlvect_1(X1)))|~m1_subset_1(u1_rlvect_1(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(X1)),u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~v1_funct_2(X3,k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~v1_funct_1(u1_rlvect_1(X1))|~v1_funct_1(X3)|~l1_rlvect_1(X1))).#trainneg
cnf(i_0_379, plain, (v7_ordinal1(esk7_3(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v6_membered(u1_struct_0(X1))|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1944, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~r2_hidden(X1,k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,X2,esk5_2(X2,X3))))|~v1_xboole_0(k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X2),X3),k2_xcmplx_0(X3,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X2))))|~v2_xxreal_0(X3)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X2))|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_1925, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,k1_funct_1(esk2_0,X1)))),k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X1)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(k1_funct_1(esk2_0,X1))|~v1_xreal_0(k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_1783, plain, (v7_ordinal1(esk10_1(X1))|~v6_membered(k1_numbers)|~v3_membered(X1))).#trainneg
cnf(i_0_997, plain, (g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1))=k8_metric_1|~v1_pre_topc(k8_metric_1)|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_2022, plain, (u1_struct_0(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))=k1_numbers|k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))!=k3_topmetr|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_380, plain, (v1_xreal_0(esk7_3(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v3_membered(u1_struct_0(X1))|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1561, plain, (v1_xreal_0(u1_rlvect_1(k8_metric_1))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~l1_rlvect_1(k8_metric_1))).#trainneg
cnf(i_0_975, plain, (v1_valued_0(k7_relset_1(X6,X7,X4,X5))|~v1_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))).#trainneg
cnf(i_0_381, plain, (v3_valued_0(esk7_3(k15_euclid(X1),X2,X3))|v2_struct_0(k15_euclid(X1))|v2_struct_0(X2)|v5_pre_topc(X3,k3_pcomps_1(k15_euclid(X1)),k3_pcomps_1(X2))|~v7_ordinal1(X1)|~l1_metric_1(X2)|~l1_metric_1(k15_euclid(X1))|~v9_metric_1(X2)|~v9_metric_1(k15_euclid(X1))|~v8_metric_1(X2)|~v8_metric_1(k15_euclid(X1))|~v7_metric_1(X2)|~v7_metric_1(k15_euclid(X1))|~v6_metric_1(X2)|~v6_metric_1(k15_euclid(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k15_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k15_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_382, plain, (v1_finseq_1(esk7_3(k15_euclid(X1),X2,X3))|v2_struct_0(k15_euclid(X1))|v2_struct_0(X2)|v5_pre_topc(X3,k3_pcomps_1(k15_euclid(X1)),k3_pcomps_1(X2))|~v7_ordinal1(X1)|~l1_metric_1(X2)|~l1_metric_1(k15_euclid(X1))|~v9_metric_1(X2)|~v9_metric_1(k15_euclid(X1))|~v8_metric_1(X2)|~v8_metric_1(k15_euclid(X1))|~v7_metric_1(X2)|~v7_metric_1(k15_euclid(X1))|~v6_metric_1(X2)|~v6_metric_1(k15_euclid(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k15_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k15_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1999, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_1796, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k7_relat_1(u1_metric_1(k14_euclid(X2)),X3))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1000, plain, (v2_pre_topc(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))|~l1_pre_topc(k8_metric_1)|~v2_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_2024, plain, (u1_struct_0(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))=u1_struct_0(X1)|k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))!=k3_pcomps_1(X1)|~l1_metric_1(X1)|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_1563, plain, (v1_xreal_0(u1_rlvect_1(k3_topmetr))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~l1_rlvect_1(k3_topmetr))).#trainneg
cnf(i_0_850, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relset_1(X6,X7,X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))).#trainneg
cnf(i_0_384, plain, (v7_ordinal1(esk8_3(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v6_membered(u1_struct_0(X2))|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_385, plain, (v1_xreal_0(esk8_3(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v3_membered(u1_struct_0(X2))|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1998, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0)))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_1994, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0)))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_1993, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(k7_relset_1(X2,X3,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_2101, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~r2_hidden(X1,k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,X2,esk5_2(X2,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X2)))))|~v1_xboole_0(k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X2),k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X2)))|~v2_xxreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X2))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X2))|~m1_subset_1(X2,u1_struct_0(k15_euclid(esk1_0)))|~m1_subset_1(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X2),k1_numbers))).#trainneg
cnf(i_0_386, plain, (v3_valued_0(esk8_3(X1,k15_euclid(X2),X3))|v2_struct_0(X1)|v2_struct_0(k15_euclid(X2))|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(k15_euclid(X2)))|~v7_ordinal1(X2)|~l1_metric_1(k15_euclid(X2))|~l1_metric_1(X1)|~v9_metric_1(k15_euclid(X2))|~v9_metric_1(X1)|~v8_metric_1(k15_euclid(X2))|~v8_metric_1(X1)|~v7_metric_1(k15_euclid(X2))|~v7_metric_1(X1)|~v6_metric_1(k15_euclid(X2))|~v6_metric_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(k15_euclid(X2))))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(k15_euclid(X2))))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1666, plain, (u1_struct_0(X1)=k1_numbers|k3_pcomps_1(X1)!=k3_pcomps_1(k3_topmetr)|~l1_metric_1(X1)|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_1729, plain, (v7_ordinal1(u1_metric_1(k3_topmetr))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_387, plain, (v1_finseq_1(esk8_3(X1,k15_euclid(X2),X3))|v2_struct_0(X1)|v2_struct_0(k15_euclid(X2))|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(k15_euclid(X2)))|~v7_ordinal1(X2)|~l1_metric_1(k15_euclid(X2))|~l1_metric_1(X1)|~v9_metric_1(k15_euclid(X2))|~v9_metric_1(X1)|~v8_metric_1(k15_euclid(X2))|~v8_metric_1(X1)|~v7_metric_1(k15_euclid(X2))|~v7_metric_1(X1)|~v6_metric_1(k15_euclid(X2))|~v6_metric_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(k15_euclid(X2))))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(k15_euclid(X2))))|~v1_funct_1(X3))).#trainneg
cnf(i_0_388, plain, (v2_struct_0(X1)|v2_struct_0(k14_euclid(X2))|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(k14_euclid(X2)))|m1_subset_1(esk8_3(X1,k14_euclid(X2),X3),u1_struct_0(k15_euclid(X2)))|~l1_metric_1(k14_euclid(X2))|~l1_metric_1(X1)|~v9_metric_1(k14_euclid(X2))|~v9_metric_1(X1)|~v8_metric_1(k14_euclid(X2))|~v8_metric_1(X1)|~v7_metric_1(k14_euclid(X2))|~v7_metric_1(X1)|~v6_metric_1(k14_euclid(X2))|~v6_metric_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(k14_euclid(X2))))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(k14_euclid(X2))))|~v1_funct_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1671, plain, (v5_pre_topc(X1,k3_topmetr,k3_topmetr)|~r1_tarski(k7_relset_1(k1_numbers,k1_numbers,X1,k9_metric_1(k8_metric_1,esk7_3(k8_metric_1,k8_metric_1,X1),X2)),k9_metric_1(k8_metric_1,esk8_3(k8_metric_1,k8_metric_1,X1),esk9_3(k8_metric_1,k8_metric_1,X1)))|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1831, plain, (k2_rcomp_1(k5_real_1(X1,X2),k2_xcmplx_0(X2,X1))=k9_metric_1(k8_metric_1,X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1057, plain, (u1_pre_topc(k3_topmetr)=X1|g1_pre_topc(X2,X1)!=k3_topmetr|~m1_subset_1(u1_pre_topc(k3_topmetr),k1_zfmisc_1(k1_zfmisc_1(k1_numbers))))).#trainneg
cnf(i_0_1204, negated_conjecture, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k7_relset_1(X3,X4,esk2_0,X2))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_885, plain, (m1_subset_1(k2_pcomps_1(k3_topmetr),k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_902, plain, (k1_numbers=X1|g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))!=g1_metric_1(X1,X2)|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_1747, plain, (v1_xreal_0(u1_metric_1(k3_topmetr))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_391, plain, (m1_subset_1(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)),X3,k9_metric_1(X1,X4,esk6_6(X1,X2,X3,X4,X5,X6))),k1_zfmisc_1(k9_metric_1(X2,X5,X6)))|v2_struct_0(X1)|v2_struct_0(X2)|X5!=k1_funct_1(X3,X4)|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_395, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r1_tarski(k7_relat_1(X3,k9_metric_1(X1,X4,esk6_6(X1,X2,X3,X4,X5,X6))),k9_metric_1(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_418, plain, (u1_struct_0(X1)=u1_struct_0(k15_euclid(X2))|g1_metric_1(u1_struct_0(X1),u1_metric_1(X1))!=k14_euclid(X2)|~l1_metric_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1002, plain, (m1_subset_1(u1_pre_topc(k8_metric_1),k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_901, plain, (u1_metric_1(k3_topmetr)=X1|g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))!=g1_metric_1(X2,X1)|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_1667, plain, (u1_struct_0(X1)=k1_numbers|X1!=k3_pcomps_1(k3_topmetr)|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_2196, plain, (v3_membered(k2_pcomps_1(k3_topmetr))|~v3_membered(k1_zfmisc_1(k1_numbers))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_398, plain, (u1_struct_0(k15_euclid(X1))=X2|g1_metric_1(u1_struct_0(k15_euclid(X1)),u1_metric_1(k14_euclid(X1)))!=g1_metric_1(X2,X3)|~l1_metric_1(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_400, plain, (u1_pre_topc(k15_euclid(X1))=X2|k3_pcomps_1(k14_euclid(X1))!=g1_pre_topc(X3,X2)|~m1_subset_1(u1_pre_topc(k15_euclid(X1)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2225, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr))).#trainneg
cnf(i_0_2132, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,esk3_0),esk4_0))))|~m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_2151, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,esk3_0),esk4_0)))|~m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_2198, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,k2_pcomps_1(k3_topmetr))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_2182, plain, (k2_pcomps_1(k3_topmetr)=u1_pre_topc(k3_topmetr)|k3_pcomps_1(k3_topmetr)!=k3_topmetr|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_401, plain, (X1=u1_pre_topc(k15_euclid(X2))|g1_pre_topc(X3,X1)!=k3_pcomps_1(k14_euclid(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X3)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2197, plain, (v1_relat_1(k2_pcomps_1(k3_topmetr))|~v1_relat_1(k1_zfmisc_1(k1_numbers))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_2240, plain, (v3_membered(u1_pre_topc(k8_metric_1))|~v3_membered(k1_zfmisc_1(k1_numbers))|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_404, plain, (l1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(u1_pre_topc(k15_euclid(X1)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2242, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,u1_pre_topc(k8_metric_1))|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_1795, plain, (v1_valued_0(u1_metric_1(k14_euclid(X1)))|~v1_membered(k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2241, plain, (v1_relat_1(u1_pre_topc(k8_metric_1))|~v1_relat_1(k1_zfmisc_1(k1_numbers))|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_405, plain, (v1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(u1_pre_topc(k15_euclid(X1)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2194, plain, (v1_xreal_0(k2_pcomps_1(k3_topmetr))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_1075, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relset_1(X5,X6,k2_zfmisc_1(X3,X2),X4))|~m1_subset_1(k2_zfmisc_1(X3,X2),k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_1713, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~r2_hidden(X1,k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,X2,esk5_2(X2,X3))))|~v1_xboole_0(k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X2),X3),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X2),X3)))|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_2155, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,esk3_0),esk4_0)))|~m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_1926, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|r1_tarski(k7_relset_1(X1,X2,esk2_0,k1_topreal9(esk1_0,X3,esk5_2(X3,k1_funct_1(esk2_0,X3)))),k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,X3),k1_funct_1(esk2_0,X3)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(k1_funct_1(esk2_0,X3))|~v1_xreal_0(k1_funct_1(esk2_0,X3))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_2224, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r1_tarski(k7_relset_1(X7,X8,X3,k9_metric_1(X2,X4,esk6_6(X2,X1,X3,X4,X5,X6))),k9_metric_1(X1,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X1)|~v9_metric_1(X2)|~v8_metric_1(X1)|~v8_metric_1(X2)|~v7_metric_1(X1)|~v7_metric_1(X2)|~v6_metric_1(X1)|~v6_metric_1(X2)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(X2),k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X1)))))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X2))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_1661, plain, (u1_struct_0(k3_pcomps_1(k3_topmetr))=k1_numbers|~v1_pre_topc(k3_pcomps_1(k3_topmetr))|~l1_pre_topc(k3_pcomps_1(k3_topmetr))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_410, plain, (v2_struct_0(k15_euclid(X1))|~l1_pre_topc(k15_euclid(X1))|~v2_struct_0(k3_pcomps_1(k14_euclid(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_427, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk10_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,esk10_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_980, plain, (u1_struct_0(X1)=u1_struct_0(k15_euclid(X2))|X1!=k14_euclid(X2)|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v1_pre_topc(k14_euclid(X2))|~l1_pre_topc(k14_euclid(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1723, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|m1_subset_1(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k1_zfmisc_1(k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1),X2)))|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1875, plain, (v2_struct_0(k14_euclid(X1))|v5_pre_topc(X2,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X1))|~r1_tarski(k7_relat_1(X2,k1_topreal9(X1,esk7_3(k14_euclid(X1),k8_metric_1,X2),X3)),k9_metric_1(k8_metric_1,esk8_3(k14_euclid(X1),k8_metric_1,X2),esk9_3(k14_euclid(X1),k8_metric_1,X2)))|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),k1_numbers)))|~m1_subset_1(esk7_3(k14_euclid(X1),k8_metric_1,X2),u1_struct_0(k14_euclid(X1)))|~m1_subset_1(esk7_3(k14_euclid(X1),k8_metric_1,X2),u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X2,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_887, plain, (v1_funct_2(u1_rlvect_1(k3_topmetr),k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~l1_rlvect_1(k3_topmetr))).#trainneg
cnf(i_0_2200, plain, (u1_struct_0(k15_euclid(X1))=k1_numbers|g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))!=k14_euclid(X1)|~l1_metric_1(k3_topmetr)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_428, plain, (u1_metric_1(k14_euclid(X1))=k7_metric_1|k14_euclid(X1)!=k8_metric_1|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2119, plain, (u1_struct_0(X1)=k1_numbers|X1!=k8_metric_1|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v1_pre_topc(k8_metric_1)|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_888, plain, (v1_funct_2(u1_metric_1(k3_topmetr),k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_429, plain, (u1_metric_1(X1)=k7_metric_1|X1!=k8_metric_1|~v1_metric_1(X1)|~l1_metric_1(X1))).#trainneg
cnf(i_0_430, plain, (m1_subset_1(esk9_3(X1,X2,X3),k1_numbers)|v2_struct_0(X1)|v2_struct_0(X2)|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_2258, plain, (u1_struct_0(k15_euclid(X1))=X2|k14_euclid(X1)!=g1_metric_1(X2,X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2328, plain, (v2_struct_0(X1)|v2_struct_0(k14_euclid(X2))|r1_tarski(k7_relset_1(X3,X4,X5,k9_metric_1(X1,X6,esk6_6(X1,k14_euclid(X2),X5,X6,X7,X8))),k1_topreal9(X2,X7,X8))|X7!=k1_funct_1(X5,X6)|~l1_metric_1(k14_euclid(X2))|~l1_metric_1(X1)|~v9_metric_1(k14_euclid(X2))|~v9_metric_1(X1)|~v8_metric_1(k14_euclid(X2))|~v8_metric_1(X1)|~v7_metric_1(k14_euclid(X2))|~v7_metric_1(X1)|~v6_metric_1(k14_euclid(X2))|~v6_metric_1(X1)|~v2_xxreal_0(X8)|~v1_xreal_0(X8)|~v5_pre_topc(X5,k3_pcomps_1(X1),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(k14_euclid(X2))))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X7,u1_struct_0(k14_euclid(X2)))|~m1_subset_1(X6,u1_struct_0(X1))|~v1_funct_2(X5,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(k14_euclid(X2))))|~v1_funct_1(X5)|~m1_subset_1(X7,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2337, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~r2_hidden(X1,k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,X2,esk5_2(X2,X3))))|~v1_xboole_0(k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X2),X3))|~v2_xxreal_0(X3)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X2))|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_1177, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(esk4_0,k1_funct_1(esk2_0,esk3_0))))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_2245, plain, (u1_metric_1(k3_topmetr)=u1_metric_1(X1)|g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))!=X1|~l1_metric_1(k3_topmetr)|~v1_metric_1(X1)|~l1_metric_1(X1))).#trainneg
cnf(i_0_433, plain, (m1_subset_1(k1_funct_1(X1,X2),k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_439, plain, (k2_pcomps_1(k14_euclid(X1))=X2|k3_pcomps_1(k14_euclid(X1))!=g1_pre_topc(X3,X2)|~m1_subset_1(k2_pcomps_1(k14_euclid(X1)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1060, plain, (k1_numbers=X1|g1_pre_topc(X1,X2)!=k3_topmetr|~m1_subset_1(u1_pre_topc(k3_topmetr),k1_zfmisc_1(k1_zfmisc_1(k1_numbers))))).#trainneg
cnf(i_0_1954, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|m1_subset_1(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1)))),k1_zfmisc_1(k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1),k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1))))|~v2_xxreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1))|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0)))|~m1_subset_1(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X1),k1_numbers))).#trainneg
cnf(i_0_2222, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(X2)|r1_tarski(k7_relat_1(X3,k9_metric_1(X2,X4,esk6_6(X2,k14_euclid(X1),X3,X4,X5,X6))),k1_topreal9(X1,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X1))|~l1_metric_1(X2)|~v9_metric_1(k14_euclid(X1))|~v9_metric_1(X2)|~v8_metric_1(k14_euclid(X1))|~v8_metric_1(X2)|~v7_metric_1(k14_euclid(X1))|~v7_metric_1(X2)|~v6_metric_1(k14_euclid(X1))|~v6_metric_1(X2)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(X2),k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(k14_euclid(X1))))))|~m1_subset_1(X5,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X4,u1_struct_0(X2))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(k14_euclid(X1))))|~v1_funct_1(X3)|~m1_subset_1(X5,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2244, plain, (u1_metric_1(k14_euclid(X1))=u1_metric_1(k3_topmetr)|g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))!=k14_euclid(X1)|~l1_metric_1(k3_topmetr)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2043, plain, (v1_xboole_0(k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5))|m1_subset_1(X6,X4)|~m1_subset_1(X6,k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5))|~m1_subset_1(k2_zfmisc_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2204, plain, (m1_subset_1(esk8_3(X1,X2,X3),k1_numbers)|v2_struct_0(X2)|v2_struct_0(X1)|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~v3_valued_0(X3)|~v1_relat_1(X3)|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1063, plain, (k2_pcomps_1(X1)=u1_pre_topc(k3_topmetr)|k3_pcomps_1(X1)!=k3_topmetr|~l1_metric_1(X1))).#trainneg
cnf(i_0_2299, plain, (m1_subset_1(k2_pcomps_1(k3_topmetr),k1_numbers)|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_2203, plain, (k1_numbers=X1|g1_metric_1(X1,X2)!=k3_topmetr|~l1_metric_1(k3_topmetr)|~v1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_1005, plain, (v1_funct_2(u1_rlvect_1(k8_metric_1),k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~l1_rlvect_1(k8_metric_1))).#trainneg
cnf(i_0_2353, plain, (l1_metric_1(g1_metric_1(k1_numbers,u1_rlvect_1(k3_topmetr)))|~m1_subset_1(u1_rlvect_1(k3_topmetr),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~v1_funct_1(u1_rlvect_1(k3_topmetr))|~l1_rlvect_1(k3_topmetr))).#trainneg
cnf(i_0_1288, plain, (v3_membered(k7_relset_1(X1,X2,k7_relset_1(X3,k2_zfmisc_1(X4,X5),X6,X7),X8))|~v3_membered(X5)|~m1_subset_1(k7_relset_1(X3,k2_zfmisc_1(X4,X5),X6,X7),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5)))))).#trainneg
cnf(i_0_447, plain, (u1_metric_1(X1)=u1_metric_1(k14_euclid(X2))|g1_metric_1(u1_struct_0(X1),u1_metric_1(X1))!=k14_euclid(X2)|~l1_metric_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_448, plain, (u1_metric_1(X1)=u1_metric_1(X2)|g1_metric_1(u1_struct_0(X1),u1_metric_1(X1))!=X2|~l1_metric_1(X1)|~v1_metric_1(X2)|~l1_metric_1(X2))).#trainneg
cnf(i_0_973, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(k7_relset_1(X2,X3,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2023, plain, (u1_struct_0(k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))=u1_struct_0(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_2520, plain, (u1_struct_0(g1_metric_1(k1_numbers,u1_rlvect_1(k3_topmetr)))=k1_numbers|k3_pcomps_1(g1_metric_1(k1_numbers,u1_rlvect_1(k3_topmetr)))!=k3_topmetr|~l1_rlvect_1(k3_topmetr)|~m1_subset_1(u1_rlvect_1(k3_topmetr),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))))).#trainneg
cnf(i_0_449, plain, (u1_metric_1(k14_euclid(X1))=X2|g1_metric_1(u1_struct_0(k15_euclid(X1)),u1_metric_1(k14_euclid(X1)))!=g1_metric_1(X3,X2)|~l1_metric_1(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1338, plain, (v1_relat_1(k7_relset_1(X1,X2,k7_relset_1(X3,k2_zfmisc_1(X4,X5),X6,X7),X8))|~v1_relat_1(X5)|~m1_subset_1(k7_relset_1(X3,k2_zfmisc_1(X4,X5),X6,X7),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5)))))).#trainneg
cnf(i_0_1682, plain, (v7_ordinal1(k7_relset_1(X1,X2,k7_relset_1(X3,k2_zfmisc_1(X4,X5),X6,X7),X8))|~v6_membered(k1_zfmisc_1(X5))|~m1_subset_1(k7_relset_1(X3,k2_zfmisc_1(X4,X5),X6,X7),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5)))))).#trainneg
cnf(i_0_450, plain, (u1_metric_1(X1)=X2|X1!=g1_metric_1(X3,X2)|~l1_metric_1(X1)|~v1_metric_1(X1))).#trainneg
cnf(i_0_2432, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~m1_subset_1(esk9_3(k14_euclid(esk1_0),k8_metric_1,esk2_0),k1_numbers))).#trainneg
cnf(i_0_2552, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0))))))).#trainneg
cnf(i_0_2429, plain, (u1_struct_0(k15_euclid(X1))=u1_struct_0(k15_euclid(X2))|k14_euclid(X1)!=k14_euclid(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_704, plain, (v7_ordinal1(k7_relat_1(k7_metric_1,X1))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2354, plain, (v1_metric_1(g1_metric_1(k1_numbers,u1_rlvect_1(k3_topmetr)))|~m1_subset_1(u1_rlvect_1(k3_topmetr),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~v1_funct_1(u1_rlvect_1(k3_topmetr))|~l1_rlvect_1(k3_topmetr))).#trainneg
cnf(i_0_1742, plain, (v1_xreal_0(k7_relset_1(X1,X2,k7_relset_1(X3,k2_zfmisc_1(X4,X5),X6,X7),X8))|~v3_membered(k1_zfmisc_1(X5))|~m1_subset_1(k7_relset_1(X3,k2_zfmisc_1(X4,X5),X6,X7),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5)))))).#trainneg
cnf(i_0_688, plain, (m1_subset_1(k7_relset_1(X1,X2,k7_relset_1(X3,k2_zfmisc_1(X4,X5),X6,X7),X8),k1_zfmisc_1(X5))|~m1_subset_1(k7_relset_1(X3,k2_zfmisc_1(X4,X5),X6,X7),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5)))))).#trainneg
cnf(i_0_460, plain, (X1=u1_struct_0(k15_euclid(X2))|g1_pre_topc(X1,X3)!=k14_euclid(X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v1_pre_topc(k14_euclid(X2))|~l1_pre_topc(k14_euclid(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_461, plain, (u1_struct_0(k15_euclid(X1))=X2|k14_euclid(X1)!=g1_pre_topc(X2,X3)|~m1_subset_1(u1_pre_topc(k14_euclid(X1)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))|~v1_pre_topc(k14_euclid(X1))|~l1_pre_topc(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2157, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(k7_relset_1(X2,X3,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,esk3_0),esk4_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_1147, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(esk4_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)))))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_705, plain, (v1_xreal_0(k7_relat_1(k7_metric_1,X1))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_969, plain, (v1_xboole_0(k7_metric_1)|m1_subset_1(esk10_1(k7_metric_1),k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))).#trainneg
cnf(i_0_2514, plain, (l1_metric_1(g1_metric_1(k1_numbers,u1_rlvect_1(k8_metric_1)))|~m1_subset_1(u1_rlvect_1(k8_metric_1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~v1_funct_1(u1_rlvect_1(k8_metric_1))|~l1_rlvect_1(k8_metric_1))).#trainneg
cnf(i_0_1972, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X3)|v1_xboole_0(k3_funct_2(X1,k1_zfmisc_1(X3),X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X3))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(X3))|~v1_funct_1(X4)|~m1_subset_1(X2,k3_funct_2(X1,k1_zfmisc_1(X3),X4,X5)))).#trainneg
cnf(i_0_471, plain, (v2_struct_0(X1)|v2_struct_0(k14_euclid(X2))|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(k14_euclid(X2)))|~l1_metric_1(k14_euclid(X2))|~l1_metric_1(X1)|~v9_metric_1(k14_euclid(X2))|~v9_metric_1(X1)|~v8_metric_1(k14_euclid(X2))|~v8_metric_1(X1)|~v7_metric_1(k14_euclid(X2))|~v7_metric_1(X1)|~v6_metric_1(k14_euclid(X2))|~v6_metric_1(X1)|~r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(k14_euclid(X2))),X3,k9_metric_1(X1,esk7_3(X1,k14_euclid(X2),X3),X4)),k1_topreal9(X2,esk8_3(X1,k14_euclid(X2),X3),esk9_3(X1,k14_euclid(X2),X3)))|~v2_xxreal_0(X4)|~v1_xreal_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(k14_euclid(X2))))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(k14_euclid(X2))))|~v1_funct_1(X3)|~v1_xreal_0(esk9_3(X1,k14_euclid(X2),X3))|~m1_subset_1(esk8_3(X1,k14_euclid(X2),X3),u1_struct_0(k14_euclid(X2)))|~m1_subset_1(esk8_3(X1,k14_euclid(X2),X3),u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_473, plain, (v2_struct_0(X1)|v2_struct_0(k14_euclid(X2))|r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(k14_euclid(X2))),X3,k9_metric_1(X1,X4,esk6_6(X1,k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X2))|~l1_metric_1(X1)|~v9_metric_1(k14_euclid(X2))|~v9_metric_1(X1)|~v8_metric_1(k14_euclid(X2))|~v8_metric_1(X1)|~v7_metric_1(k14_euclid(X2))|~v7_metric_1(X1)|~v6_metric_1(k14_euclid(X2))|~v6_metric_1(X1)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(k14_euclid(X2))))))|~m1_subset_1(X5,u1_struct_0(k14_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(k14_euclid(X2))))|~v1_funct_1(X3)|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2545, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(k7_relset_1(X2,X3,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(esk4_0,k1_funct_1(esk2_0,esk3_0)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_1260, negated_conjecture, (v7_ordinal1(k7_relat_1(esk2_0,X1))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1261, negated_conjecture, (v1_xreal_0(k7_relat_1(esk2_0,X1))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1446, negated_conjecture, (v1_xboole_0(esk2_0)|m1_subset_1(esk10_1(esk2_0),k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers)))).#trainneg
cnf(i_0_2522, plain, (u1_struct_0(g1_metric_1(k1_numbers,u1_rlvect_1(k3_topmetr)))=u1_struct_0(X1)|k3_pcomps_1(g1_metric_1(k1_numbers,u1_rlvect_1(k3_topmetr)))!=k3_pcomps_1(X1)|~l1_metric_1(X1)|~l1_rlvect_1(k3_topmetr)|~m1_subset_1(u1_rlvect_1(k3_topmetr),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))))).#trainneg
cnf(i_0_485, plain, (v7_ordinal1(u2_struct_0(k14_euclid(X1)))|~v6_membered(u1_struct_0(k15_euclid(X1)))|~l2_struct_0(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2515, plain, (v1_metric_1(g1_metric_1(k1_numbers,u1_rlvect_1(k8_metric_1)))|~m1_subset_1(u1_rlvect_1(k8_metric_1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~v1_funct_1(u1_rlvect_1(k8_metric_1))|~l1_rlvect_1(k8_metric_1))).#trainneg
cnf(i_0_1710, plain, (u1_struct_0(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))=k1_numbers|~l1_pre_topc(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_2524, plain, (v3_membered(k7_relset_1(X1,X2,k7_relset_1(X9,X10,X6,X7),X8))|~v3_membered(X5)|~m1_subset_1(k7_relset_1(X9,X10,X6,X7),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10))))).#trainneg
cnf(i_0_486, plain, (v1_xreal_0(u2_struct_0(k14_euclid(X1)))|~v3_membered(u1_struct_0(k15_euclid(X1)))|~l2_struct_0(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2740, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(X2)|r1_tarski(k7_relat_1(X3,k1_topreal9(X1,X4,esk6_6(k14_euclid(X1),X2,X3,X4,X5,X6))),k9_metric_1(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X1))|~l1_metric_1(X2)|~v9_metric_1(k14_euclid(X1))|~v9_metric_1(X2)|~v8_metric_1(k14_euclid(X1))|~v8_metric_1(X2)|~v7_metric_1(k14_euclid(X1))|~v7_metric_1(X2)|~v6_metric_1(k14_euclid(X1))|~v6_metric_1(X2)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_2735, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|~v2_xxreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v2_xxreal_0(esk4_0)|~v1_xreal_0(esk4_0)|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0)))|~m1_subset_1(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers))).#trainneg
cnf(i_0_2741, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(X2)|r1_tarski(k7_relset_1(X7,X8,X3,k1_topreal9(X1,X4,esk6_6(k14_euclid(X1),X2,X3,X4,X5,X6))),k9_metric_1(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X1))|~l1_metric_1(X2)|~v9_metric_1(k14_euclid(X1))|~v9_metric_1(X2)|~v8_metric_1(k14_euclid(X1))|~v8_metric_1(X2)|~v7_metric_1(k14_euclid(X1))|~v7_metric_1(X2)|~v6_metric_1(k14_euclid(X1))|~v6_metric_1(X2)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_2704, plain, (v2_struct_0(k14_euclid(X1))|r1_tarski(k7_relset_1(k1_numbers,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),X2,k9_metric_1(k8_metric_1,X3,esk6_6(k8_metric_1,k14_euclid(X1),X2,X3,X4,X5))),k1_topreal9(X1,X4,X5))|X4!=k1_funct_1(X2,X3)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X5)|~v1_xreal_0(X5)|~v5_pre_topc(X2,k3_topmetr,k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(k3_pcomps_1(k14_euclid(X1))))))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X2,k1_numbers,u1_struct_0(k3_pcomps_1(k14_euclid(X1))))|~v1_funct_1(X2))).#trainneg
cnf(i_0_650, plain, (m1_subset_1(k2_finseq_1(X1),k1_numbers)|~v3_membered(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_958, plain, (v7_ordinal1(k2_xcmplx_0(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2686, plain, (v1_xreal_0(esk10_1(k7_metric_1))|v1_xboole_0(k7_metric_1)|~v3_membered(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))).#trainneg
cnf(i_0_2687, plain, (u1_struct_0(g1_metric_1(k1_numbers,u1_rlvect_1(k8_metric_1)))=k1_numbers|k3_pcomps_1(g1_metric_1(k1_numbers,u1_rlvect_1(k8_metric_1)))!=k3_topmetr|~l1_rlvect_1(k8_metric_1)|~m1_subset_1(u1_rlvect_1(k8_metric_1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))))).#trainneg
cnf(i_0_487, plain, (v3_valued_0(u2_struct_0(k14_euclid(X1)))|~v7_ordinal1(X1)|~l2_struct_0(k14_euclid(X1)))).#trainneg
cnf(i_0_2640, plain, (m1_subset_1(k7_relat_1(k7_metric_1,X1),k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2685, plain, (v7_ordinal1(esk10_1(k7_metric_1))|v1_xboole_0(k7_metric_1)|~v6_membered(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))).#trainneg
cnf(i_0_2719, negated_conjecture, (v7_ordinal1(esk10_1(esk2_0))|v1_xboole_0(esk2_0)|~v6_membered(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers)))).#trainneg
cnf(i_0_488, plain, (v1_finseq_1(u2_struct_0(k14_euclid(X1)))|~v7_ordinal1(X1)|~l2_struct_0(k14_euclid(X1)))).#trainneg
cnf(i_0_2737, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(k14_euclid(X2))|r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(k14_euclid(X2))),X3,k1_topreal9(X1,X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X1))|~l1_metric_1(k14_euclid(X2))|~v9_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X2))|~v8_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X2))|~v7_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X2))|~v6_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X2))|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(k14_euclid(X2))))))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X5,u1_struct_0(k14_euclid(X2)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(k14_euclid(X2))))|~v1_funct_1(X3)|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2738, plain, (v2_struct_0(k14_euclid(X1))|r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),k1_numbers,X2,k1_topreal9(X1,X3,esk6_6(k14_euclid(X1),k8_metric_1,X2,X3,X4,X5))),k9_metric_1(k8_metric_1,X4,X5))|X4!=k1_funct_1(X2,X3)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X5)|~v1_xreal_0(X5)|~v5_pre_topc(X2,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),k1_numbers)))|~m1_subset_1(X3,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X4,k1_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X2,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1168, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relset_1(X1,X2,esk2_0,k1_topreal9(esk1_0,esk3_0,X3)),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0)))|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_1496, plain, (v7_ordinal1(k6_xcmplx_0(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2720, negated_conjecture, (v1_xreal_0(esk10_1(esk2_0))|v1_xboole_0(esk2_0)|~v3_membered(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers)))).#trainneg
cnf(i_0_2689, plain, (u1_struct_0(g1_metric_1(k1_numbers,u1_rlvect_1(k8_metric_1)))=u1_struct_0(X1)|k3_pcomps_1(g1_metric_1(k1_numbers,u1_rlvect_1(k8_metric_1)))!=k3_pcomps_1(X1)|~l1_metric_1(X1)|~l1_rlvect_1(k8_metric_1)|~m1_subset_1(u1_rlvect_1(k8_metric_1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))))).#trainneg
cnf(i_0_510, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1)|~r1_tarski(k7_relset_1(X5,X6,X3,k9_metric_1(X1,esk7_3(X1,X2,X3),X4)),k9_metric_1(X2,esk8_3(X1,X2,X3),esk9_3(X1,X2,X3)))|~v2_xxreal_0(X4)|~v1_xreal_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_2641, plain, (m1_subset_1(k7_relat_1(k7_metric_1,X1),k1_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2726, plain, (u1_struct_0(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))=k1_numbers|~l1_pre_topc(k8_metric_1)|~m1_subset_1(u1_pre_topc(k8_metric_1),k1_zfmisc_1(k1_zfmisc_1(k1_numbers))))).#trainneg
cnf(i_0_2754, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k14_euclid(X1),X2)|g1_pre_topc(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(X2)|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(X2)|~v2_pre_topc(k14_euclid(X1))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2756, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),X2)|g1_pre_topc(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=k3_pcomps_1(k14_euclid(X1))|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(X2)|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(X2)|~v2_pre_topc(k15_euclid(X1))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2757, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=X1|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_pre_topc(X1))).#trainneg
cnf(i_0_2921, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(k14_euclid(esk1_0)),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2)))).#trainneg
cnf(i_0_2901, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(X2)|r1_tarski(k7_relset_1(u1_struct_0(k14_euclid(X1)),u1_struct_0(k3_pcomps_1(X2)),X3,k9_metric_1(k14_euclid(X1),X4,esk6_6(k14_euclid(X1),X2,X3,X4,X5,X6))),k9_metric_1(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(X2)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(X2)|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(X2)|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(X2)|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(X2)|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k14_euclid(X1)),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2896, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(X1)),X2)|g1_pre_topc(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=g1_pre_topc(u1_struct_0(k14_euclid(X1)),u1_pre_topc(k3_pcomps_1(k14_euclid(X1))))|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(X2)|~l1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~v2_pre_topc(X2)|~v2_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2853, plain, (u1_metric_1(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))=u1_metric_1(k14_euclid(X1))|g1_metric_1(k1_numbers,u1_metric_1(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1))))!=k14_euclid(X1)|~l1_metric_1(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))|~v7_ordinal1(X1)|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_2714, negated_conjecture, (m1_subset_1(k7_relat_1(esk2_0,X1),k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2670, plain, (m1_subset_1(k7_relset_1(X1,X2,k7_relset_1(X9,X10,X6,X7),X8),k1_zfmisc_1(X5))|~m1_subset_1(k7_relset_1(X9,X10,X6,X7),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10))))).#trainneg
cnf(i_0_529, plain, (m1_subset_1(k7_relat_1(esk10_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X3),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_531, plain, (m1_subset_1(k7_relat_1(u1_metric_1(X1),X2),k1_zfmisc_1(k1_numbers))|~l1_metric_1(X1))).#trainneg
cnf(i_0_2905, plain, (v2_struct_0(X1)|v2_struct_0(k14_euclid(X2))|r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k14_euclid(X2)),X3,k9_metric_1(X1,X4,esk6_6(X1,k14_euclid(X2),X3,X4,X5,X6))),k9_metric_1(k14_euclid(X2),X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X2))|~l1_metric_1(X1)|~v9_metric_1(k14_euclid(X2))|~v9_metric_1(X1)|~v8_metric_1(k14_euclid(X2))|~v8_metric_1(X1)|~v7_metric_1(k14_euclid(X2))|~v7_metric_1(X1)|~v6_metric_1(k14_euclid(X2))|~v6_metric_1(X1)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k14_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k14_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k14_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2917, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(X2)|r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k14_euclid(X1)),X3,k9_metric_1(X2,X4,esk6_6(X2,k14_euclid(X1),X3,X4,X5,X6))),k1_topreal9(X1,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X1))|~l1_metric_1(X2)|~v9_metric_1(k14_euclid(X1))|~v9_metric_1(X2)|~v8_metric_1(k14_euclid(X1))|~v8_metric_1(X2)|~v7_metric_1(k14_euclid(X1))|~v7_metric_1(X2)|~v6_metric_1(k14_euclid(X1))|~v6_metric_1(X2)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(X2),k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k14_euclid(X1)))))|~m1_subset_1(X5,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k14_euclid(X1)))|~v1_funct_1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2752, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v2_xxreal_0(esk4_0)|~v1_xreal_0(esk4_0)|~m1_subset_1(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),k1_numbers))).#trainneg
cnf(i_0_2918, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(X2)|r1_tarski(k7_relset_1(u1_struct_0(k14_euclid(X1)),u1_struct_0(k3_pcomps_1(X2)),X3,k1_topreal9(X1,X4,esk6_6(k14_euclid(X1),X2,X3,X4,X5,X6))),k9_metric_1(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X1))|~l1_metric_1(X2)|~v9_metric_1(k14_euclid(X1))|~v9_metric_1(X2)|~v8_metric_1(k14_euclid(X1))|~v8_metric_1(X2)|~v7_metric_1(k14_euclid(X1))|~v7_metric_1(X2)|~v6_metric_1(k14_euclid(X1))|~v6_metric_1(X2)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X3,u1_struct_0(k14_euclid(X1)),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2753, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),X1)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_topmetr|~l1_pre_topc(X1)|~l1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0)))|~v2_pre_topc(X1)|~v2_pre_topc(k3_pcomps_1(k14_euclid(esk1_0)))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),u1_struct_0(X1))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),u1_struct_0(X1)))).#trainneg
cnf(i_0_2900, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(X2)|v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(X2))|~l1_metric_1(X2)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(X2)|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(X2)|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(X2)|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(X2)|~v6_metric_1(k14_euclid(X1))|~r1_tarski(k7_relset_1(u1_struct_0(k14_euclid(X1)),u1_struct_0(k3_pcomps_1(X2)),X3,k9_metric_1(k14_euclid(X1),esk7_3(k14_euclid(X1),X2,X3),X4)),k9_metric_1(X2,esk8_3(k14_euclid(X1),X2,X3),esk9_3(k14_euclid(X1),X2,X3)))|~v2_xxreal_0(X4)|~v1_xreal_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(k14_euclid(X1)),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3024, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~v2_xxreal_0(esk4_0)|~v1_xreal_0(esk4_0)|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0)))|~m1_subset_1(k1_funct_1(esk2_0,esk3_0),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers))).#trainneg
cnf(i_0_2976, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relat_1(k7_relat_1(esk10_1(k1_zfmisc_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,X2)))),X5),X6)))).#trainneg
cnf(i_0_2816, plain, (m1_subset_1(u2_struct_0(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1))),k1_numbers)|~l2_struct_0(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_2591, plain, (u1_metric_1(k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))=u1_metric_1(k14_euclid(X1))|g1_metric_1(u1_struct_0(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))),u1_metric_1(k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))))!=k14_euclid(X1)|~l1_metric_1(k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))|~v7_ordinal1(X1)|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_533, plain, (m1_subset_1(k7_relat_1(u1_rlvect_1(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_rlvect_1(X1))).#trainneg
cnf(i_0_2715, negated_conjecture, (m1_subset_1(k7_relat_1(esk2_0,X1),k1_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2815, plain, (v1_finset_1(k1_numbers)|~v8_struct_0(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))|~l1_struct_0(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_2614, plain, (v1_relat_1(k7_relset_1(X1,X2,k7_relset_1(X9,X10,X6,X7),X8))|~v1_relat_1(X5)|~m1_subset_1(k7_relset_1(X9,X10,X6,X7),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10))))).#trainneg
cnf(i_0_534, plain, (m1_subset_1(k7_relat_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5),X6),k1_zfmisc_1(X3))|v1_xboole_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_535, plain, (m1_subset_1(k7_relat_1(k7_relset_1(X1,k2_zfmisc_1(X2,X3),X4,X5),X6),k1_zfmisc_1(X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_2936, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(k15_euclid(esk1_0)),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2)))).#trainneg
cnf(i_0_3036, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,k14_euclid(X2))|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|g1_pre_topc(u1_struct_0(k15_euclid(X2)),u1_pre_topc(k14_euclid(X2)))!=k3_topmetr|~l1_pre_topc(k14_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(k14_euclid(X2))|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3030, plain, (v2_struct_0(X1)|r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(X2)),u1_struct_0(k3_pcomps_1(X1)),X3,k1_topreal9(X2,X4,esk6_6(k14_euclid(X2),X1,X3,X4,X5,X6))),k9_metric_1(X1,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X2)),k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),u1_struct_0(k3_pcomps_1(X1)))))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X5,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(k15_euclid(X2)),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3025, negated_conjecture, (~v2_xxreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~v2_xxreal_0(esk4_0)|~v1_xreal_0(esk4_0)|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0)))|~m1_subset_1(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers))).#trainneg
cnf(i_0_3041, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_pre_topc(X2))).#trainneg
cnf(i_0_3040, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,k15_euclid(X2))|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(k15_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(X2))|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3012, plain, (v2_struct_0(X1)|r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k15_euclid(X2)),X3,k9_metric_1(X1,X4,esk6_6(X1,k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3160, negated_conjecture, (~v2_xxreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~v2_xxreal_0(esk4_0)|~v1_xreal_0(esk4_0)|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0)))|~m1_subset_1(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),k1_numbers))).#trainneg
cnf(i_0_3159, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(esk4_0)|~v1_xreal_0(esk4_0)|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0)))|~m1_subset_1(k1_funct_1(esk2_0,esk3_0),k1_numbers))).#trainneg
cnf(i_0_3189, plain, (r1_tarski(k7_relset_1(k1_numbers,u1_struct_0(k15_euclid(X1)),X2,k9_metric_1(k8_metric_1,X3,esk6_6(k8_metric_1,k14_euclid(X1),X2,X3,X4,X5))),k1_topreal9(X1,X4,X5))|X4!=k1_funct_1(X2,X3)|~v2_xxreal_0(X5)|~v1_xreal_0(X5)|~v5_pre_topc(X2,k3_topmetr,k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(k15_euclid(X1)))))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X2,k1_numbers,u1_struct_0(k15_euclid(X1)))|~v1_funct_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_536, plain, (m1_subset_1(k7_relat_1(k2_zfmisc_1(X1,X2),X3),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_2777, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2781, plain, (m1_subset_1(esk10_1(k7_metric_1),k1_numbers)|v1_xboole_0(k7_metric_1)|~v3_membered(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))).#trainneg
cnf(i_0_555, negated_conjecture, (v7_ordinal1(esk5_2(X1,X2))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v6_membered(k1_numbers)|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_568, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(X2)|v5_pre_topc(X3,k3_pcomps_1(X2),k3_pcomps_1(k14_euclid(X1)))|~l1_metric_1(k14_euclid(X1))|~l1_metric_1(X2)|~v9_metric_1(k14_euclid(X1))|~v9_metric_1(X2)|~v8_metric_1(k14_euclid(X1))|~v8_metric_1(X2)|~v7_metric_1(k14_euclid(X1))|~v7_metric_1(X2)|~v6_metric_1(k14_euclid(X1))|~v6_metric_1(X2)|~r1_tarski(k7_relat_1(X3,k9_metric_1(X2,esk7_3(X2,k14_euclid(X1),X3),X4)),k1_topreal9(X1,esk8_3(X2,k14_euclid(X1),X3),esk9_3(X2,k14_euclid(X1),X3)))|~v2_xxreal_0(X4)|~v1_xreal_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(k14_euclid(X1))))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(k14_euclid(X1))))|~v1_funct_1(X3)|~v1_xreal_0(esk9_3(X2,k14_euclid(X1),X3))|~m1_subset_1(esk8_3(X2,k14_euclid(X1),X3),u1_struct_0(k14_euclid(X1)))|~m1_subset_1(esk8_3(X2,k14_euclid(X1),X3),u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3192, plain, (v2_struct_0(X1)|r1_tarski(k7_relat_1(X3,k9_metric_1(X1,X4,esk6_6(X1,k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3069, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,k3_topmetr)|X1!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers)|~v1_pre_topc(X1))).#trainneg
cnf(i_0_3165, plain, (v2_struct_0(X1)|r1_tarski(k7_relat_1(X3,k1_topreal9(X2,X4,esk6_6(k14_euclid(X2),X1,X3,X4,X5,X6))),k9_metric_1(X1,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X2)),k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),u1_struct_0(k3_pcomps_1(X1)))))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X5,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(k15_euclid(X2)),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3178, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,X2)|X1!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_pre_topc(X1))).#trainneg
cnf(i_0_3166, plain, (v2_struct_0(X1)|r1_tarski(k7_relset_1(X7,X8,X3,k1_topreal9(X2,X4,esk6_6(k14_euclid(X2),X1,X3,X4,X5,X6))),k9_metric_1(X1,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X2)),k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),u1_struct_0(k3_pcomps_1(X1)))))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X5,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(k15_euclid(X2)),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_3064, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k14_euclid(X1),k3_topmetr)|g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k14_euclid(X1))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3156, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(k15_euclid(esk1_0)),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=X1|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_pre_topc(X1))).#trainneg
cnf(i_0_3177, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),X2)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~l1_pre_topc(X2)|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(X2)|~v2_pre_topc(k15_euclid(X1))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3191, plain, (v2_struct_0(k14_euclid(X1))|r1_tarski(k7_relset_1(u1_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X2)),X3,k9_metric_1(k14_euclid(X1),X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3075, plain, (v1_xreal_0(u2_struct_0(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1))))|~l2_struct_0(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_2805, negated_conjecture, (m1_subset_1(esk10_1(esk2_0),k1_numbers)|v1_xboole_0(esk2_0)|~v3_membered(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers)))).#trainneg
cnf(i_0_3074, plain, (v7_ordinal1(u2_struct_0(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1))))|~v6_membered(k1_numbers)|~l2_struct_0(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_2622, plain, (v7_ordinal1(k7_relset_1(X1,X2,k7_relset_1(X9,X10,X6,X7),X8))|~v6_membered(k1_zfmisc_1(X5))|~m1_subset_1(k7_relset_1(X9,X10,X6,X7),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10))))).#trainneg
cnf(i_0_571, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v5_pre_topc(X3,k3_pcomps_1(X2),k3_pcomps_1(X1))|~l1_metric_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X1)|~v9_metric_1(X2)|~v8_metric_1(X1)|~v8_metric_1(X2)|~v7_metric_1(X1)|~v7_metric_1(X2)|~v6_metric_1(X1)|~v6_metric_1(X2)|~v2_xxreal_0(X4)|~v1_xreal_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X1)))))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X3)|~m1_subset_1(k7_relat_1(X3,k9_metric_1(X2,esk7_3(X2,X1,X3),X4)),k1_zfmisc_1(k9_metric_1(X1,esk8_3(X2,X1,X3),esk9_3(X2,X1,X3)))))).#trainneg
cnf(i_0_593, plain, (k10_funcsdom(k2_finseq_1(X1))=k15_euclid(X1)|~v1_rlvect_1(k15_euclid(X1))|~l1_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3152, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k14_euclid(X1),X2)|g1_pre_topc(u1_struct_0(k15_euclid(esk1_0)),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(X2)|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(X2)|~v2_pre_topc(k14_euclid(X1))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3173, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k14_euclid(X1),X2)|g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~l1_pre_topc(X2)|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(X2)|~v2_pre_topc(k14_euclid(X1))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3176, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k3_topmetr,X1)|k3_pcomps_1(k14_euclid(esk1_0))!=k3_topmetr|X1!=k3_topmetr|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,k1_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_3157, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,k14_euclid(X2))|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|k14_euclid(X2)!=k3_topmetr|~l1_pre_topc(k14_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(k14_euclid(X2))|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2)|~v1_pre_topc(k14_euclid(X2)))).#trainneg
cnf(i_0_3193, plain, (v2_struct_0(X1)|r1_tarski(k7_relset_1(X7,X8,X3,k9_metric_1(X1,X4,esk6_6(X1,k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_3155, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),X2)|g1_pre_topc(u1_struct_0(k15_euclid(esk1_0)),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=k3_pcomps_1(k14_euclid(X1))|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(X2)|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(X2)|~v2_pre_topc(k15_euclid(X1))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2912, plain, (v2_struct_0(k14_euclid(X1))|r1_tarski(k7_relset_1(k1_numbers,u1_struct_0(k14_euclid(X1)),X2,k9_metric_1(k8_metric_1,X3,esk6_6(k8_metric_1,k14_euclid(X1),X2,X3,X4,X5))),k1_topreal9(X1,X4,X5))|X4!=k1_funct_1(X2,X3)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X5)|~v1_xreal_0(X5)|~v5_pre_topc(X2,k3_topmetr,k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(k14_euclid(X1)))))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X2,k1_numbers,u1_struct_0(k14_euclid(X1)))|~v1_funct_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3021, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~m1_subset_1(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k1_numbers))).#trainneg
cnf(i_0_2944, plain, (v2_struct_0(X1)|r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(X2)),u1_struct_0(k3_pcomps_1(X1)),X3,k9_metric_1(k14_euclid(X2),X4,esk6_6(k14_euclid(X2),X1,X3,X4,X5,X6))),k9_metric_1(X1,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X2)),k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),u1_struct_0(k3_pcomps_1(X1)))))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X5,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(k15_euclid(X2)),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2775, plain, (v7_ordinal1(k2_finseq_1(X1))|~v6_membered(k1_numbers)|~v3_membered(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_889, plain, (m1_subset_1(u1_rlvect_1(k3_topmetr),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~l1_rlvect_1(k3_topmetr))).#trainneg
cnf(i_0_3283, plain, (k15_euclid(X1)=k3_topmetr|g1_rlvect_1(k1_numbers,u2_struct_0(k3_topmetr),u1_algstr_0(k3_topmetr),u1_rlvect_1(k3_topmetr))!=k15_euclid(X1)|k3_pcomps_1(k14_euclid(X1))!=k3_topmetr|~l1_rltopsp1(k3_topmetr)|~v5_rltopsp1(k3_topmetr)|~v7_ordinal1(X1)|~v1_rlvect_1(k15_euclid(X1))|~l1_rlvect_1(k15_euclid(X1)))).#trainneg
cnf(i_0_594, plain, (k15_euclid(X1)=k15_euclid(X2)|k10_funcsdom(k2_finseq_1(X1))!=k10_funcsdom(k2_finseq_1(X2))|g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k15_euclid(X1)))!=k3_pcomps_1(k14_euclid(X2))|~l1_rltopsp1(k15_euclid(X1))|~v5_rltopsp1(k15_euclid(X1))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3329, plain, (v3_valued_0(u1_rlvect_1(k3_topmetr))|~l1_rlvect_1(k3_topmetr))).#trainneg
cnf(i_0_3330, plain, (v1_valued_0(u1_rlvect_1(k3_topmetr))|~v1_membered(k1_numbers)|~l1_rlvect_1(k3_topmetr))).#trainneg
cnf(i_0_2646, plain, (v1_xreal_0(k7_relset_1(X1,X2,k7_relset_1(X9,X10,X6,X7),X8))|~v3_membered(k1_zfmisc_1(X5))|~m1_subset_1(k7_relset_1(X9,X10,X6,X7),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10))))).#trainneg
cnf(i_0_595, plain, (k2_pcomps_1(X1)=u1_pre_topc(k3_pcomps_1(X1))|~v1_pre_topc(k3_pcomps_1(X1))|~l1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_3234, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,k3_topmetr)|X1!=k3_pcomps_1(k14_euclid(esk1_0))|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers))).#trainneg
cnf(i_0_3295, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k14_euclid(X1),X2)|g1_pre_topc(u1_struct_0(k15_euclid(esk1_0)),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))|X2!=k3_topmetr|~l1_pre_topc(k14_euclid(X1))|~l1_pre_topc(X2)|~v2_pre_topc(k14_euclid(X1))|~v2_pre_topc(X2)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1)|~v1_pre_topc(X2))).#trainneg
cnf(i_0_3293, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k14_euclid(X1),k3_topmetr)|g1_pre_topc(u1_struct_0(k15_euclid(esk1_0)),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k14_euclid(X1))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3254, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k14_euclid(X1),k3_topmetr)|k14_euclid(X1)!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k14_euclid(X1))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1)|~v1_pre_topc(k14_euclid(X1)))).#trainneg
cnf(i_0_3297, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k3_topmetr,X1)|k3_pcomps_1(k14_euclid(esk1_0))!=k3_topmetr|X1!=k3_topmetr|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,k1_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_3242, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,X2)|X1!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~v1_pre_topc(X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2)))).#trainneg
cnf(i_0_3212, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relat_1(k7_relat_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,X2)),X5),X6)))).#trainneg
cnf(i_0_3078, plain, (u1_metric_1(k14_euclid(X1))=u1_metric_1(k3_pcomps_1(k3_topmetr))|g1_metric_1(k1_numbers,u1_metric_1(k3_pcomps_1(k3_topmetr)))!=k14_euclid(X1)|~l1_metric_1(k3_pcomps_1(k3_topmetr))|~l1_metric_1(k3_topmetr)|~v7_ordinal1(X1)|~v1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_614, plain, (m1_subset_1(u2_struct_0(X1),k5_numbers)|~v6_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_2118, plain, (u1_struct_0(X1)=k1_numbers|k3_pcomps_1(X1)!=k8_metric_1|~l1_metric_1(X1)|~v1_pre_topc(k8_metric_1)|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_615, plain, (m1_subset_1(u2_struct_0(X1),k1_numbers)|~v6_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_624, negated_conjecture, (m1_subset_1(k7_relat_1(esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,X1),X2),k2_xcmplx_0(k1_funct_1(esk2_0,X1),X2))))|v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_3296, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k14_euclid(X1),X2)|k14_euclid(X1)!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~l1_pre_topc(k14_euclid(X1))|~l1_pre_topc(X2)|~v2_pre_topc(k14_euclid(X1))|~v2_pre_topc(X2)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1)|~v1_pre_topc(k14_euclid(X1)))).#trainneg
cnf(i_0_3258, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),X2)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3199, plain, (r1_tarski(k7_relat_1(X2,k9_metric_1(k8_metric_1,X3,esk6_6(k8_metric_1,k14_euclid(X1),X2,X3,X4,X5))),k1_topreal9(X1,X4,X5))|X4!=k1_funct_1(X2,X3)|~v2_xxreal_0(X5)|~v1_xreal_0(X5)|~v5_pre_topc(X2,k3_topmetr,k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(k15_euclid(X1)))))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X2,k1_numbers,u1_struct_0(k15_euclid(X1)))|~v1_funct_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3290, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k14_euclid(X1),k14_euclid(X2))|g1_pre_topc(u1_struct_0(k15_euclid(esk1_0)),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))|g1_pre_topc(u1_struct_0(k15_euclid(X2)),u1_pre_topc(k14_euclid(X2)))!=k3_topmetr|~l1_pre_topc(k14_euclid(X1))|~l1_pre_topc(k14_euclid(X2))|~v2_pre_topc(k14_euclid(X1))|~v2_pre_topc(k14_euclid(X2))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3261, plain, (r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)),X3,k9_metric_1(k14_euclid(X1),X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3294, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k14_euclid(X1),k15_euclid(X2))|g1_pre_topc(u1_struct_0(k15_euclid(esk1_0)),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(k14_euclid(X1))|~l1_pre_topc(k15_euclid(X2))|~v2_pre_topc(k14_euclid(X1))|~v2_pre_topc(k15_euclid(X2))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3318, plain, (v2_struct_0(X1)|r1_tarski(k7_relat_1(X3,k9_metric_1(k14_euclid(X2),X4,esk6_6(k14_euclid(X2),X1,X3,X4,X5,X6))),k9_metric_1(X1,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X2)),k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),u1_struct_0(k3_pcomps_1(X1)))))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X5,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(k15_euclid(X2)),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3332, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k7_relat_1(u1_rlvect_1(k3_topmetr),X2))|~l1_rlvect_1(k3_topmetr))).#trainneg
cnf(i_0_644, plain, (v1_xboole_0(k5_numbers)|m1_subset_1(k3_funct_2(X1,k5_numbers,X2,X3),k1_numbers)|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k5_numbers)))|~m1_subset_1(X3,X1)|~v1_funct_2(X2,X1,k5_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2683, plain, (m1_subset_1(k7_relat_1(k7_metric_1,X1),k1_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_651, plain, (m1_subset_1(u2_struct_0(X1),k1_numbers)|~l2_struct_0(X1)|~v3_membered(u1_struct_0(X1)))).#trainneg
cnf(i_0_682, plain, (m1_subset_1(k7_relset_1(X1,X2,esk10_1(k1_zfmisc_1(k2_zfmisc_1(X3,X4))),X5),k1_zfmisc_1(X4))|~m1_subset_1(esk10_1(k1_zfmisc_1(k2_zfmisc_1(X3,X4))),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2928, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(k14_euclid(esk1_0)),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=X1|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~v1_pre_topc(X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2)))).#trainneg
cnf(i_0_2772, plain, (v2_struct_0(k14_euclid(X1))|r1_tarski(k7_relat_1(X2,k9_metric_1(k8_metric_1,X3,esk6_6(k8_metric_1,k14_euclid(X1),X2,X3,X4,X5))),k1_topreal9(X1,X4,X5))|X4!=k1_funct_1(X2,X3)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X5)|~v1_xreal_0(X5)|~v5_pre_topc(X2,k3_topmetr,k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(k3_pcomps_1(k14_euclid(X1))))))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X2,k1_numbers,u1_struct_0(k3_pcomps_1(k14_euclid(X1))))|~v1_funct_1(X2))).#trainneg
cnf(i_0_3188, plain, (v2_struct_0(k14_euclid(X1))|r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k15_euclid(X2)),X3,k1_topreal9(X1,X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v1_xreal_0(esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3310, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v2_xxreal_0(esk4_0)|~v1_xreal_0(esk4_0)|~m1_subset_1(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),k1_numbers))).#trainneg
cnf(i_0_2914, plain, (v2_struct_0(k14_euclid(X1))|r1_tarski(k7_relset_1(u1_struct_0(k14_euclid(X1)),k1_numbers,X2,k1_topreal9(X1,X3,esk6_6(k14_euclid(X1),k8_metric_1,X2,X3,X4,X5))),k9_metric_1(k8_metric_1,X4,X5))|X4!=k1_funct_1(X2,X3)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X5)|~v1_xreal_0(X5)|~v5_pre_topc(X2,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),k1_numbers)))|~m1_subset_1(X3,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X4,k1_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X2,u1_struct_0(k14_euclid(X1)),k1_numbers)|~v1_funct_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3367, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,k3_topmetr)|X1!=k3_pcomps_1(k14_euclid(esk1_0))|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers))).#trainneg
cnf(i_0_3370, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k3_topmetr,X1)|k3_pcomps_1(k14_euclid(esk1_0))!=k3_topmetr|X1!=k3_topmetr|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,k1_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_3358, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,u1_pre_topc(k3_pcomps_1(k3_topmetr)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_3284, plain, (k15_euclid(X1)=k8_metric_1|g1_rlvect_1(k1_numbers,u2_struct_0(k8_metric_1),u1_algstr_0(k8_metric_1),u1_rlvect_1(k8_metric_1))!=k15_euclid(X1)|g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1))!=k3_pcomps_1(k14_euclid(X1))|~l1_rltopsp1(k8_metric_1)|~v5_rltopsp1(k8_metric_1)|~v7_ordinal1(X1)|~v1_rlvect_1(k15_euclid(X1))|~l1_rlvect_1(k15_euclid(X1)))).#trainneg
cnf(i_0_684, plain, (m1_subset_1(k7_relset_1(X1,X2,u1_metric_1(X3),X4),k1_zfmisc_1(k1_numbers))|~m1_subset_1(u1_metric_1(X3),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~l1_metric_1(X3))).#trainneg
cnf(i_0_1786, plain, (v1_xboole_0(k1_topreal9(X1,X2,X3))|m1_subset_1(k3_funct_2(X4,k1_topreal9(X1,X2,X3),X5,X6),u1_struct_0(k15_euclid(X1)))|v1_xboole_0(X4)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1)))|~v7_ordinal1(X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,k1_topreal9(X1,X2,X3))))|~m1_subset_1(X6,X4)|~v1_funct_2(X5,X4,k1_topreal9(X1,X2,X3))|~v1_funct_1(X5))).#trainneg
cnf(i_0_685, plain, (m1_subset_1(k7_relset_1(X1,X2,k7_metric_1,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_686, plain, (m1_subset_1(k7_relset_1(X1,X2,u1_rlvect_1(X3),X4),k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(u1_rlvect_1(X3),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~l1_rlvect_1(X3))).#trainneg
cnf(i_0_3371, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,X2)|X1!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~v1_pre_topc(X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2)))).#trainneg
cnf(i_0_3404, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),X2)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~l1_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3411, plain, (r1_tarski(k7_relat_1(X3,k9_metric_1(k14_euclid(X1),X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3468, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~m1_subset_1(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),k1_numbers))).#trainneg
cnf(i_0_3469, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_3419, plain, (v2_struct_0(X1)|r1_tarski(k7_relset_1(X7,X8,X2,k9_metric_1(k14_euclid(X3),X4,esk6_6(k14_euclid(X3),X1,X2,X4,X5,X6))),k9_metric_1(X1,X5,X6))|X5!=k1_funct_1(X2,X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X2,k3_pcomps_1(k14_euclid(X3)),k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),u1_struct_0(k3_pcomps_1(X1)))))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X3)))|~m1_subset_1(X5,u1_struct_0(X1))|~v1_funct_2(X2,u1_struct_0(k15_euclid(X3)),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_3462, plain, (v2_struct_0(k14_euclid(X1))|r1_tarski(k7_relset_1(u1_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X2)),X3,k1_topreal9(X1,X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X6)|~v1_xreal_0(esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X3,u1_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3467, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v7_ordinal1(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)))).#trainneg
cnf(i_0_3463, plain, (v2_struct_0(k14_euclid(X1))|r1_tarski(k7_relat_1(X3,k1_topreal9(X1,X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X6)|~v1_xreal_0(esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1802, plain, (m1_subset_1(k7_relat_1(u1_metric_1(k14_euclid(X1)),X2),k1_zfmisc_1(k1_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3356, plain, (v3_membered(u1_pre_topc(k3_pcomps_1(k3_topmetr)))|~v3_membered(k1_zfmisc_1(k1_numbers))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_687, plain, (m1_subset_1(k7_relset_1(X1,X2,k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)),X6,X7),X8),k1_zfmisc_1(X5))|v1_xboole_0(X3)|~m1_subset_1(k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)),X6,X7),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))))|~m1_subset_1(X7,X3)|~v1_funct_2(X6,X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~v1_funct_1(X6))).#trainneg
cnf(i_0_2717, negated_conjecture, (m1_subset_1(k7_relat_1(esk2_0,X1),k1_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3357, plain, (v1_relat_1(u1_pre_topc(k3_pcomps_1(k3_topmetr)))|~v1_relat_1(k1_zfmisc_1(k1_numbers))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_3355, plain, (v1_xreal_0(u1_pre_topc(k3_pcomps_1(k3_topmetr)))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_1895, plain, (v1_xboole_0(k7_relset_1(X1,X2,X3,X4))|m1_subset_1(k3_funct_2(X5,k7_relset_1(X1,X2,X3,X4),X6,X7),X2)|v1_xboole_0(X5)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,k7_relset_1(X1,X2,X3,X4))))|~m1_subset_1(X7,X5)|~v1_funct_2(X6,X5,k7_relset_1(X1,X2,X3,X4))|~v1_funct_1(X6))).#trainneg
cnf(i_0_3494, plain, (v3_valued_0(k3_funct_2(X1,k1_topreal9(X2,X3,X4),X5,X6))|v1_xboole_0(k1_topreal9(X2,X3,X4))|v1_xboole_0(X1)|~v7_ordinal1(X2)|~v1_xreal_0(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k1_topreal9(X2,X3,X4))))|~m1_subset_1(X3,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X6,X1)|~v1_funct_2(X5,X1,k1_topreal9(X2,X3,X4))|~v1_funct_1(X5))).#trainneg
cnf(i_0_690, negated_conjecture, (m1_subset_1(k7_relset_1(X1,X2,esk2_0,k1_topreal9(esk1_0,X3,esk5_2(X3,X4))),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,X3),X4),k2_xcmplx_0(k1_funct_1(esk2_0,X3),X4))))|v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X4)|~v1_xreal_0(X4)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_3477, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,k3_topmetr)|X1!=k3_pcomps_1(k14_euclid(esk1_0))|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers))).#trainneg
cnf(i_0_3478, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k3_topmetr,X1)|k3_pcomps_1(k14_euclid(esk1_0))!=k3_topmetr|X1!=k3_topmetr|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,k1_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_3521, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,X2)|X1!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~v1_pre_topc(X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2)))).#trainneg
cnf(i_0_3523, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),X2)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~l1_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3545, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v7_ordinal1(k1_funct_1(esk2_0,esk3_0))|~m1_subset_1(esk3_0,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_3541, plain, (r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)),X3,k1_topreal9(X1,X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~v2_xxreal_0(X6)|~v1_xreal_0(esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3407, plain, (r1_tarski(k7_relset_1(X6,X7,X1,k9_metric_1(k8_metric_1,X2,esk6_6(k8_metric_1,k14_euclid(X3),X1,X2,X4,X5))),k1_topreal9(X3,X4,X5))|X4!=k1_funct_1(X1,X2)|~v2_xxreal_0(X5)|~v1_xreal_0(X5)|~v5_pre_topc(X1,k3_topmetr,k3_pcomps_1(k14_euclid(X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(k15_euclid(X3)))))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X3)))|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k1_numbers,u1_struct_0(k15_euclid(X3)))|~v1_funct_1(X1)|~v7_ordinal1(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))).#trainneg
cnf(i_0_3473, plain, (r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(X1)),k1_numbers,X2,k1_topreal9(X1,X3,esk6_6(k14_euclid(X1),k8_metric_1,X2,X3,X4,X5))),k9_metric_1(k8_metric_1,X4,X5))|X4!=k1_funct_1(X2,X3)|~v2_xxreal_0(X5)|~v1_xreal_0(X5)|~v5_pre_topc(X2,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X4,k1_numbers)|~v1_funct_2(X2,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v1_funct_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3417, plain, (v2_struct_0(k14_euclid(X1))|r1_tarski(k7_relat_1(X2,k9_metric_1(k14_euclid(X3),X4,esk6_6(k14_euclid(X3),k14_euclid(X1),X2,X4,X5,X6))),k1_topreal9(X1,X5,X6))|X5!=k1_funct_1(X2,X4)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X2,k3_pcomps_1(k14_euclid(X3)),k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),u1_struct_0(k3_pcomps_1(k14_euclid(X1))))))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X3)))|~m1_subset_1(X5,u1_struct_0(k14_euclid(X1)))|~v1_funct_2(X2,u1_struct_0(k15_euclid(X3)),u1_struct_0(k3_pcomps_1(k14_euclid(X1))))|~v1_funct_1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X5,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3412, plain, (r1_tarski(k7_relset_1(X7,X8,X3,k9_metric_1(k14_euclid(X1),X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_3552, plain, (v3_membered(k7_relat_1(u1_metric_1(k14_euclid(X1)),X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3354, plain, (g1_pre_topc(k1_numbers,u1_pre_topc(k3_pcomps_1(k3_topmetr)))=k3_pcomps_1(k3_topmetr)|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_2831, plain, (g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1))=k15_euclid(X1)|g1_rlvect_1(k1_numbers,u2_struct_0(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1))),u1_algstr_0(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1))),u1_rlvect_1(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1))))!=k10_funcsdom(k2_finseq_1(X1))|g1_pre_topc(k1_numbers,u1_pre_topc(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1))))!=k3_pcomps_1(k14_euclid(X1))|~l1_rltopsp1(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))|~v5_rltopsp1(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))|~v7_ordinal1(X1)|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_2107, plain, (m1_subset_1(esk10_1(X1),k5_numbers)|~v6_membered(k1_numbers)|~v3_membered(X1))).#trainneg
cnf(i_0_3359, plain, (m1_subset_1(u1_pre_topc(k3_pcomps_1(k3_topmetr)),k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_2825, plain, (v2_struct_0(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))|~l1_pre_topc(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))|~v2_struct_0(g1_pre_topc(k1_numbers,u1_pre_topc(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))))|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_3495, plain, (v1_finseq_1(k3_funct_2(X1,k1_topreal9(X2,X3,X4),X5,X6))|v1_xboole_0(k1_topreal9(X2,X3,X4))|v1_xboole_0(X1)|~v7_ordinal1(X2)|~v1_xreal_0(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k1_topreal9(X2,X3,X4))))|~m1_subset_1(X3,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X6,X1)|~v1_funct_2(X5,X1,k1_topreal9(X2,X3,X4))|~v1_funct_1(X5))).#trainneg
cnf(i_0_3492, plain, (v7_ordinal1(k3_funct_2(X1,k1_topreal9(X2,X3,X4),X5,X6))|v1_xboole_0(k1_topreal9(X2,X3,X4))|v1_xboole_0(X1)|~v6_membered(u1_struct_0(k15_euclid(X2)))|~v1_xreal_0(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k1_topreal9(X2,X3,X4))))|~m1_subset_1(X3,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X6,X1)|~v1_funct_2(X5,X1,k1_topreal9(X2,X3,X4))|~v1_funct_1(X5)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3368, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k14_euclid(X1),X2)|g1_pre_topc(u1_struct_0(k15_euclid(esk1_0)),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=k14_euclid(X1)|X2!=k3_topmetr|~v1_pre_topc(X2)|~l1_pre_topc(k14_euclid(X1))|~l1_pre_topc(X2)|~v2_pre_topc(k14_euclid(X1))|~v2_pre_topc(X2)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1)|~v1_pre_topc(k14_euclid(X1)))).#trainneg
cnf(i_0_3360, plain, (g1_pre_topc(u1_struct_0(X1),u1_pre_topc(k3_pcomps_1(X1)))=k3_pcomps_1(X1)|~l1_metric_1(X1))).#trainneg
cnf(i_0_3369, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k14_euclid(X1),k3_topmetr)|g1_pre_topc(u1_struct_0(k15_euclid(esk1_0)),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=k14_euclid(X1)|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k14_euclid(X1))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1)|~v1_pre_topc(k14_euclid(X1)))).#trainneg
cnf(i_0_3464, plain, (v2_struct_0(k14_euclid(X1))|r1_tarski(k7_relset_1(X7,X8,X3,k1_topreal9(X1,X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X6)|~v1_xreal_0(esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_2934, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),X2)|g1_pre_topc(u1_struct_0(k14_euclid(esk1_0)),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=k3_pcomps_1(k14_euclid(X1))|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(X2)|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(X2)|~v2_pre_topc(k15_euclid(X1))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3066, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|g1_pre_topc(u1_struct_0(k14_euclid(X1)),u1_pre_topc(k3_pcomps_1(k14_euclid(X1))))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~v2_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3408, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k14_euclid(X1),k14_euclid(X2))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|g1_pre_topc(u1_struct_0(k15_euclid(esk1_0)),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))|k14_euclid(X2)!=k3_topmetr|~l1_pre_topc(k14_euclid(X1))|~l1_pre_topc(k14_euclid(X2))|~v2_pre_topc(k14_euclid(X1))|~v2_pre_topc(k14_euclid(X2))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_pre_topc(k14_euclid(X2)))).#trainneg
cnf(i_0_1047, negated_conjecture, (m1_subset_1(k7_relset_1(X1,X2,esk2_0,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_890, plain, (m1_subset_1(u1_metric_1(k3_topmetr),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_711, plain, (u1_struct_0(g1_metric_1(X1,X2))=X1|~v1_metric_1(g1_metric_1(X1,X2))|~l1_metric_1(g1_metric_1(X1,X2)))).#trainneg
cnf(i_0_3493, plain, (v1_xreal_0(k3_funct_2(X1,k1_topreal9(X2,X3,X4),X5,X6))|v1_xboole_0(k1_topreal9(X2,X3,X4))|v1_xboole_0(X1)|~v3_membered(u1_struct_0(k15_euclid(X2)))|~v1_xreal_0(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k1_topreal9(X2,X3,X4))))|~m1_subset_1(X3,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X6,X1)|~v1_funct_2(X5,X1,k1_topreal9(X2,X3,X4))|~v1_funct_1(X5)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_716, plain, (k2_pcomps_1(X1)=k2_pcomps_1(k14_euclid(X2))|k3_pcomps_1(X1)!=k3_pcomps_1(k14_euclid(X2))|~l1_metric_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3732, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,X2)|k3_pcomps_1(k14_euclid(esk1_0))!=X1|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~v1_pre_topc(X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2)))).#trainneg
cnf(i_0_3731, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2)))).#trainneg
cnf(i_0_3645, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,k3_topmetr)|X1!=k3_pcomps_1(k14_euclid(esk1_0))|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers))).#trainneg
cnf(i_0_3647, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,X2)|X1!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~v1_pre_topc(X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2)))).#trainneg
cnf(i_0_719, plain, (k2_pcomps_1(X1)=u1_pre_topc(k15_euclid(X2))|k3_pcomps_1(X1)!=k3_pcomps_1(k14_euclid(X2))|~l1_metric_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3620, plain, (v7_ordinal1(k3_funct_2(X1,k7_relset_1(X2,X3,X4,X5),X6,X7))|v1_xboole_0(k7_relset_1(X2,X3,X4,X5))|v1_xboole_0(X1)|~v6_membered(X3)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k7_relset_1(X2,X3,X4,X5))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X7,X1)|~v1_funct_2(X6,X1,k7_relset_1(X2,X3,X4,X5))|~v1_funct_1(X6))).#trainneg
cnf(i_0_3621, plain, (v1_xreal_0(k3_funct_2(X1,k7_relset_1(X2,X3,X4,X5),X6,X7))|v1_xboole_0(k7_relset_1(X2,X3,X4,X5))|v1_xboole_0(X1)|~v3_membered(X3)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k7_relset_1(X2,X3,X4,X5))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X7,X1)|~v1_funct_2(X6,X1,k7_relset_1(X2,X3,X4,X5))|~v1_funct_1(X6))).#trainneg
cnf(i_0_725, plain, (v3_membered(k7_relat_1(X3,X4))|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_728, plain, (v1_relat_1(k7_relat_1(X3,X4))|~v1_relat_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3747, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(X1)),k1_numbers)|~v7_ordinal1(X1)|~l1_metric_1(k14_euclid(X1)))).#trainneg
cnf(i_0_3648, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),X2)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~l1_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3745, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),X2)|k3_pcomps_1(k14_euclid(esk1_0))!=k3_pcomps_1(k14_euclid(X1))|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3646, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k3_topmetr,X1)|k3_pcomps_1(k14_euclid(esk1_0))!=k3_topmetr|X1!=k3_topmetr|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,k1_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_3746, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k3_pcomps_1(k14_euclid(X1))))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3306, plain, (r1_tarski(k7_relat_1(X2,k9_metric_1(k8_metric_1,X3,esk6_6(k8_metric_1,k14_euclid(X1),X2,X3,X4,X5))),k1_topreal9(X1,X4,X5))|X4!=k1_funct_1(X2,X3)|~v2_xxreal_0(X5)|~v1_xreal_0(X5)|~v5_pre_topc(X2,k3_topmetr,k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(k14_euclid(X1)))))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X2,k1_numbers,u1_struct_0(k14_euclid(X1)))|~v1_funct_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_732, plain, (m1_subset_1(X1,k5_numbers)|v1_xboole_0(k2_finseq_1(X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k2_finseq_1(X2)))).#trainneg
cnf(i_0_777, plain, (v3_membered(k7_relset_1(X2,X3,k7_metric_1,X1))|~m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_781, plain, (v7_ordinal1(u1_pre_topc(k14_euclid(X1)))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))|~l1_pre_topc(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2909, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(k14_euclid(X2))|r1_tarski(k7_relset_1(u1_struct_0(k14_euclid(X2)),u1_struct_0(k3_pcomps_1(k14_euclid(X1))),X3,k9_metric_1(k14_euclid(X2),X4,esk6_6(k14_euclid(X2),k14_euclid(X1),X3,X4,X5,X6))),k1_topreal9(X1,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X1))|~l1_metric_1(k14_euclid(X2))|~v9_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X2))|~v8_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X2))|~v7_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X2))|~v6_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X2))|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X2)),k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X2)),u1_struct_0(k3_pcomps_1(k14_euclid(X1))))))|~m1_subset_1(X5,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X2)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X3,u1_struct_0(k14_euclid(X2)),u1_struct_0(k3_pcomps_1(k14_euclid(X1))))|~v1_funct_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3424, plain, (v1_xboole_0(k5_numbers)|v1_xboole_0(X1)|m1_subset_1(k1_funct_1(X2,X3),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k5_numbers)))|~m1_subset_1(X3,X1)|~v1_funct_2(X2,X1,k5_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2879, plain, (v2_struct_0(k3_pcomps_1(k14_euclid(X1)))|~l1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~v1_xboole_0(u1_struct_0(k14_euclid(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2923, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k14_euclid(X1),X2)|g1_pre_topc(u1_struct_0(k14_euclid(esk1_0)),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(k14_euclid(X1))|~l1_pre_topc(X2)|~v2_pre_topc(k14_euclid(X1))|~v2_pre_topc(X2)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3651, plain, (r1_tarski(k7_relat_1(X3,k1_topreal9(X1,X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~v2_xxreal_0(X6)|~v1_xreal_0(esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3052, plain, (v5_pre_topc(X1,k3_pcomps_1(k14_euclid(X2)),k3_topmetr)|~r1_tarski(k7_relset_1(u1_struct_0(k14_euclid(X2)),k1_numbers,X1,k9_metric_1(k14_euclid(X2),esk7_3(k14_euclid(X2),k8_metric_1,X1),X3)),k9_metric_1(k8_metric_1,esk8_3(k14_euclid(X2),k8_metric_1,X1),esk9_3(k14_euclid(X2),k8_metric_1,X1)))|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X2)),k1_numbers)))|~v1_funct_2(X1,u1_struct_0(k14_euclid(X2)),k1_numbers)|~v1_funct_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2913, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(k14_euclid(X2))|r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(k14_euclid(X2))),u1_struct_0(k14_euclid(X1)),X3,k1_topreal9(X2,X4,esk6_6(k14_euclid(X2),k14_euclid(X1),X3,X4,X5,X6))),k1_topreal9(X1,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X2))|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X2))|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X2))|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X2))|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X2))|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X2)),k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X2))),u1_struct_0(k14_euclid(X1)))))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X5,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X2))),u1_struct_0(k14_euclid(X1)))|~v1_funct_1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_784, plain, (v1_xboole_0(k5_numbers)|v7_ordinal1(k3_funct_2(X1,k5_numbers,X2,X3))|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k5_numbers)))|~m1_subset_1(X3,X1)|~v1_funct_2(X2,X1,k5_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1289, plain, (v3_membered(k7_relset_1(X1,X2,k7_relset_1(X3,X4,k2_zfmisc_1(X5,k2_zfmisc_1(X6,X7)),X8),X9))|~v3_membered(X7)|~m1_subset_1(k7_relset_1(X3,X4,k2_zfmisc_1(X5,k2_zfmisc_1(X6,X7)),X8),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_zfmisc_1(X5,k2_zfmisc_1(X6,X7)),k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_1339, plain, (v1_relat_1(k7_relset_1(X1,X2,k7_relset_1(X3,X4,k2_zfmisc_1(X5,k2_zfmisc_1(X6,X7)),X8),X9))|~v1_relat_1(X7)|~m1_subset_1(k7_relset_1(X3,X4,k2_zfmisc_1(X5,k2_zfmisc_1(X6,X7)),X8),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_zfmisc_1(X5,k2_zfmisc_1(X6,X7)),k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_787, plain, (m1_subset_1(k7_relset_1(X1,X2,X3,X4),k5_numbers)|~v6_membered(k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_788, plain, (m1_subset_1(k7_relset_1(X1,X2,X3,X4),k1_numbers)|~v6_membered(k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3804, plain, (k2_pcomps_1(k14_euclid(X1))=u1_pre_topc(k15_euclid(X1))|~l1_metric_1(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3870, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),k15_euclid(X2))|k3_pcomps_1(k14_euclid(esk1_0))!=k3_pcomps_1(k14_euclid(X1))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(k15_euclid(X2))|~v2_pre_topc(k15_euclid(X2))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3791, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,k15_euclid(X2))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|k3_pcomps_1(k14_euclid(esk1_0))!=X1|~v1_pre_topc(X1)|~l1_pre_topc(k15_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(X2))|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3787, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,k14_euclid(X2))|g1_pre_topc(u1_struct_0(k15_euclid(X2)),u1_pre_topc(k14_euclid(X2)))!=k3_topmetr|k3_pcomps_1(k14_euclid(esk1_0))!=X1|~v1_pre_topc(X1)|~l1_pre_topc(k14_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(k14_euclid(X2))|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3866, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),k14_euclid(X2))|k3_pcomps_1(k14_euclid(esk1_0))!=k3_pcomps_1(k14_euclid(X1))|g1_pre_topc(u1_struct_0(k15_euclid(X2)),u1_pre_topc(k14_euclid(X2)))!=k3_topmetr|~l1_pre_topc(k14_euclid(X2))|~v2_pre_topc(k14_euclid(X2))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3798, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,k3_topmetr)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers))).#trainneg
cnf(i_0_3795, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,k14_euclid(X2))|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|g1_pre_topc(u1_struct_0(k15_euclid(X2)),u1_pre_topc(k14_euclid(X2)))!=k3_topmetr|~l1_pre_topc(k14_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(k14_euclid(X2))|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3695, plain, (g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1))=k15_euclid(X1)|g1_rlvect_1(k1_numbers,u2_struct_0(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1))),u1_algstr_0(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1))),u1_rlvect_1(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1))))!=k15_euclid(X1)|g1_pre_topc(k1_numbers,u1_pre_topc(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1))))!=k3_pcomps_1(k14_euclid(X1))|~l1_rltopsp1(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))|~v5_rltopsp1(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))|~l1_pre_topc(k8_metric_1)|~v7_ordinal1(X1)|~v1_rlvect_1(k15_euclid(X1))|~l1_rlvect_1(k15_euclid(X1)))).#trainneg
cnf(i_0_789, plain, (v7_ordinal1(k7_relat_1(X3,X4))|~v6_membered(k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3703, plain, (v7_ordinal1(u1_pre_topc(k3_pcomps_1(k3_topmetr)))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_2081, plain, (v3_valued_0(k7_relset_1(X1,X2,k7_relset_1(X3,k2_zfmisc_1(X4,k2_zfmisc_1(X5,X6)),X7,X8),X9))|~v3_membered(X6)|~m1_subset_1(k7_relset_1(X3,k2_zfmisc_1(X4,k2_zfmisc_1(X5,X6)),X7,X8),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,k2_zfmisc_1(X5,X6))))))).#trainneg
cnf(i_0_2129, plain, (v1_valued_0(k7_relset_1(X1,X2,k7_relset_1(X3,k2_zfmisc_1(X4,k2_zfmisc_1(X5,X6)),X7,X8),X9))|~v1_membered(X6)|~m1_subset_1(k7_relset_1(X3,k2_zfmisc_1(X4,k2_zfmisc_1(X5,X6)),X7,X8),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,k2_zfmisc_1(X5,X6))))))).#trainneg
cnf(i_0_793, plain, (v1_xreal_0(u1_pre_topc(k14_euclid(X1)))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))|~l1_pre_topc(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_794, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|v1_xboole_0(u1_pre_topc(X2))|~l1_pre_topc(X2)|~m1_subset_1(X1,u1_pre_topc(X2)))).#trainneg
cnf(i_0_3872, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k3_topmetr,X1)|k3_pcomps_1(k14_euclid(esk1_0))!=k3_topmetr|X1!=k3_topmetr|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,k1_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_3897, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k14_euclid(X1),k3_topmetr)|g1_pre_topc(u1_struct_0(k14_euclid(esk1_0)),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k14_euclid(X1))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3801, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,k3_topmetr)|X1!=k3_pcomps_1(k14_euclid(esk1_0))|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers))).#trainneg
cnf(i_0_3799, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,k15_euclid(X2))|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(k15_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(X2))|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3802, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,X2)|X1!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~v1_pre_topc(X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2)))).#trainneg
cnf(i_0_3800, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_pre_topc(X2))).#trainneg
cnf(i_0_3789, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,k3_pcomps_1(k14_euclid(X2)))|g1_pre_topc(u1_struct_0(k14_euclid(X2)),u1_pre_topc(k3_pcomps_1(k14_euclid(X2))))!=k3_topmetr|k3_pcomps_1(k14_euclid(esk1_0))!=X1|~v1_pre_topc(X1)|~l1_pre_topc(k3_pcomps_1(k14_euclid(X2)))|~l1_pre_topc(X1)|~v2_pre_topc(k3_pcomps_1(k14_euclid(X2)))|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k14_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k14_euclid(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3898, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k14_euclid(X1),k15_euclid(X2))|g1_pre_topc(u1_struct_0(k14_euclid(esk1_0)),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(k14_euclid(X1))|~l1_pre_topc(k15_euclid(X2))|~v2_pre_topc(k14_euclid(X1))|~v2_pre_topc(k15_euclid(X2))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3863, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),X2)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~l1_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3868, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k15_euclid(X1),k3_pcomps_1(k14_euclid(X2)))|k3_pcomps_1(k14_euclid(esk1_0))!=k3_pcomps_1(k14_euclid(X1))|g1_pre_topc(u1_struct_0(k14_euclid(X2)),u1_pre_topc(k3_pcomps_1(k14_euclid(X2))))!=k3_topmetr|~l1_pre_topc(k3_pcomps_1(k14_euclid(X2)))|~v2_pre_topc(k3_pcomps_1(k14_euclid(X2)))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k14_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(k14_euclid(X2)))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2991, plain, (v3_membered(k7_relat_1(u1_metric_1(X1),X2))|~l1_metric_1(X1))).#trainneg
cnf(i_0_804, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k7_relset_1(X3,X4,k7_metric_1,X2))|~m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_2838, plain, (v7_ordinal1(u1_pre_topc(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1))))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_pre_topc(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_2839, plain, (v1_xreal_0(u1_pre_topc(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1))))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_pre_topc(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_805, plain, (m1_subset_1(X1,k1_numbers)|v1_xboole_0(k7_relat_1(k7_metric_1,X2))|~m1_subset_1(X1,k7_relat_1(k7_metric_1,X2)))).#trainneg
cnf(i_0_3781, plain, (u1_struct_0(g1_metric_1(u1_struct_0(X1),u1_metric_1(X1)))=u1_struct_0(X1)|~l1_metric_1(g1_metric_1(u1_struct_0(X1),u1_metric_1(X1)))|~l1_metric_1(X1))).#trainneg
cnf(i_0_3931, plain, (g1_pre_topc(u1_struct_0(k14_euclid(X1)),u1_pre_topc(k15_euclid(X1)))=k3_pcomps_1(k14_euclid(X1))|~l1_metric_1(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4018, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,X1,k3_pcomps_1(k14_euclid(X2)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|k3_pcomps_1(k14_euclid(esk1_0))!=X1|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k14_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k14_euclid(X2)))|~v7_ordinal1(X2)|~l1_metric_1(k14_euclid(X2)))).#trainneg
cnf(i_0_4005, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k3_topmetr,k15_euclid(X1))|k3_pcomps_1(k14_euclid(esk1_0))!=k3_topmetr|k3_pcomps_1(k14_euclid(X1))!=k3_topmetr|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(k15_euclid(X1)))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,k1_numbers,u1_struct_0(k15_euclid(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3943, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,k15_euclid(X2))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|k3_pcomps_1(k14_euclid(esk1_0))!=X1|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3999, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,k3_topmetr)|X1!=k3_pcomps_1(k14_euclid(esk1_0))|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers))).#trainneg
cnf(i_0_3948, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k14_euclid(X1),k3_topmetr)|g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k14_euclid(X1))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4113, plain, (u1_struct_0(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))=k1_numbers|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_4008, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,X2)|X1!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~v1_pre_topc(X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2)))).#trainneg
cnf(i_0_4017, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,X1,k3_pcomps_1(k14_euclid(X2)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|g1_pre_topc(u1_struct_0(k15_euclid(X2)),u1_pre_topc(k3_pcomps_1(k14_euclid(X2))))!=k3_topmetr|k3_pcomps_1(k14_euclid(esk1_0))!=X1|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4098, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,g1_metric_1(u1_struct_0(X2),u1_metric_1(X2)))|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(g1_metric_1(u1_struct_0(X2),u1_metric_1(X2))))!=k3_topmetr|k3_pcomps_1(k14_euclid(esk1_0))!=X1|~v1_pre_topc(X1)|~l1_pre_topc(g1_metric_1(u1_struct_0(X2),u1_metric_1(X2)))|~l1_pre_topc(X1)|~v2_pre_topc(g1_metric_1(u1_struct_0(X2),u1_metric_1(X2)))|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~l1_metric_1(X2))).#trainneg
cnf(i_0_3951, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k3_topmetr,k3_topmetr)|k3_pcomps_1(k14_euclid(esk1_0))!=k3_topmetr|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_3944, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,k14_euclid(X2))|k14_euclid(X2)!=k3_topmetr|k3_pcomps_1(k14_euclid(esk1_0))!=X1|~v1_pre_topc(X1)|~l1_pre_topc(k14_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(k14_euclid(X2))|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2)|~v1_pre_topc(k14_euclid(X2)))).#trainneg
cnf(i_0_3996, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k14_euclid(X1),k3_topmetr)|g1_pre_topc(u1_struct_0(k14_euclid(esk1_0)),u1_pre_topc(k3_pcomps_1(k14_euclid(esk1_0))))!=k14_euclid(X1)|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k14_euclid(X1))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1)|~v1_pre_topc(k14_euclid(X1)))).#trainneg
cnf(i_0_3954, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,X1,k14_euclid(X2))|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|k14_euclid(X2)!=k3_topmetr|~l1_pre_topc(k14_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(k14_euclid(X2))|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2)|~v1_pre_topc(k14_euclid(X2)))).#trainneg
cnf(i_0_812, plain, (v3_membered(u1_pre_topc(k14_euclid(X1)))|~v3_membered(k1_zfmisc_1(u1_struct_0(k15_euclid(X1))))|~l1_pre_topc(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2840, plain, (v7_ordinal1(k2_pcomps_1(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1))))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_4170, plain, (~l1_metric_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|~v2_struct_0(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_1683, plain, (v7_ordinal1(k7_relset_1(X1,X2,k7_relset_1(X3,X4,k2_zfmisc_1(X5,k2_zfmisc_1(X6,X7)),X8),X9))|~v6_membered(k1_zfmisc_1(X7))|~m1_subset_1(k7_relset_1(X3,X4,k2_zfmisc_1(X5,k2_zfmisc_1(X6,X7)),X8),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_zfmisc_1(X5,k2_zfmisc_1(X6,X7)),k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_1743, plain, (v1_xreal_0(k7_relset_1(X1,X2,k7_relset_1(X3,X4,k2_zfmisc_1(X5,k2_zfmisc_1(X6,X7)),X8),X9))|~v3_membered(k1_zfmisc_1(X7))|~m1_subset_1(k7_relset_1(X3,X4,k2_zfmisc_1(X5,k2_zfmisc_1(X6,X7)),X8),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_zfmisc_1(X5,k2_zfmisc_1(X6,X7)),k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_813, plain, (v1_relat_1(u1_pre_topc(k14_euclid(X1)))|~v1_relat_1(k1_zfmisc_1(u1_struct_0(k15_euclid(X1))))|~l1_pre_topc(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_814, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(X2))))|~r2_hidden(X1,u1_pre_topc(k14_euclid(X2)))|~l1_pre_topc(k14_euclid(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4118, plain, (u1_struct_0(k14_euclid(X1))=X2|k3_pcomps_1(k14_euclid(X1))!=g1_pre_topc(X2,X3)|~m1_subset_1(u1_pre_topc(k15_euclid(X1)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k14_euclid(X1)))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4241, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k14_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(esk1_0))!=k14_euclid(X1)|~v1_pre_topc(k14_euclid(X1))|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k14_euclid(X1))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4128, plain, (u1_struct_0(X1)=u1_struct_0(k14_euclid(X2))|X1!=k3_pcomps_1(k14_euclid(X2))|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3995, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|v5_pre_topc(esk2_0,k3_topmetr,X1)|k3_pcomps_1(k14_euclid(esk1_0))!=k3_topmetr|X1!=k3_topmetr|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k1_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_4024, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|v5_pre_topc(esk2_0,k15_euclid(X1),k3_pcomps_1(k14_euclid(X2)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|k3_pcomps_1(k14_euclid(esk1_0))!=k3_pcomps_1(k14_euclid(X1))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k14_euclid(X2)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(k14_euclid(X2)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~l1_metric_1(k14_euclid(X2)))).#trainneg
cnf(i_0_4129, plain, (u1_struct_0(k15_euclid(X1))=u1_struct_0(k14_euclid(X2))|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_816, plain, (u1_struct_0(k14_euclid(X1))=k1_numbers|k14_euclid(X1)!=k8_metric_1|~l1_metric_1(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2111, plain, (X1=k1_numbers|g1_pre_topc(X1,X2)!=k8_metric_1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v1_pre_topc(k8_metric_1)|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_4151, plain, (v1_finset_1(k1_numbers)|~v8_struct_0(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|~l1_struct_0(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_822, plain, (m1_subset_1(k7_relset_1(X1,X2,X3,X4),k1_numbers)|~v3_membered(k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_823, plain, (v1_xreal_0(k7_relat_1(X3,X4))|~v3_membered(k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4114, plain, (u1_pre_topc(k15_euclid(X1))=X2|k3_pcomps_1(k14_euclid(X1))!=g1_pre_topc(X3,X2)|~m1_subset_1(u1_pre_topc(k15_euclid(X1)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k14_euclid(X1)))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2910, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(k14_euclid(X2))|r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k14_euclid(X2)),X3,k1_topreal9(X1,X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k9_metric_1(k14_euclid(X2),X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X1))|~l1_metric_1(k14_euclid(X2))|~v9_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X2))|~v8_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X2))|~v7_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X2))|~v6_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X2))|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k14_euclid(X2)))))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X5,u1_struct_0(k14_euclid(X2)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k14_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4125, plain, (u1_pre_topc(X1)=u1_pre_topc(k15_euclid(X2))|X1!=k3_pcomps_1(k14_euclid(X2))|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4050, plain, (g1_pre_topc(u1_struct_0(X1),k2_pcomps_1(g1_metric_1(u1_struct_0(X1),u1_metric_1(X1))))=k3_pcomps_1(g1_metric_1(u1_struct_0(X1),u1_metric_1(X1)))|~l1_metric_1(g1_metric_1(u1_struct_0(X1),u1_metric_1(X1)))|~l1_metric_1(X1))).#trainneg
cnf(i_0_3262, plain, (r1_tarski(k7_relat_1(X3,k9_metric_1(k14_euclid(X1),X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3366, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,u1_pre_topc(k3_pcomps_1(X2)))|~l1_metric_1(X2))).#trainneg
cnf(i_0_3827, plain, (v3_membered(k7_relat_1(u1_rlvect_1(k3_topmetr),X1))|~l1_rlvect_1(k3_topmetr))).#trainneg
cnf(i_0_1089, plain, (m1_subset_1(k7_relset_1(X1,X2,k7_relset_1(X3,X4,k2_zfmisc_1(X5,k2_zfmisc_1(X6,X7)),X8),X9),k1_zfmisc_1(X7))|~m1_subset_1(k7_relset_1(X3,X4,k2_zfmisc_1(X5,k2_zfmisc_1(X6,X7)),X8),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_zfmisc_1(X5,k2_zfmisc_1(X6,X7)),k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_829, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(k14_euclid(X2))|v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~l1_metric_1(k14_euclid(X1))|~l1_metric_1(k14_euclid(X2))|~v9_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X2))|~v8_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X2))|~v7_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X2))|~v6_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X2))|~r1_tarski(k7_relat_1(X3,k1_topreal9(X1,esk7_3(k14_euclid(X1),k14_euclid(X2),X3),X4)),k1_topreal9(X2,esk8_3(k14_euclid(X1),k14_euclid(X2),X3),esk9_3(k14_euclid(X1),k14_euclid(X2),X3)))|~v2_xxreal_0(X4)|~v1_xreal_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(k14_euclid(X2))))))|~m1_subset_1(esk7_3(k14_euclid(X1),k14_euclid(X2),X3),u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(k14_euclid(X2))))|~v1_funct_1(X3)|~v1_xreal_0(esk9_3(k14_euclid(X1),k14_euclid(X2),X3))|~m1_subset_1(esk8_3(k14_euclid(X1),k14_euclid(X2),X3),u1_struct_0(k14_euclid(X2)))|~m1_subset_1(esk8_3(k14_euclid(X1),k14_euclid(X2),X3),u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_830, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(X2)|v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(X2))|~l1_metric_1(k14_euclid(X1))|~l1_metric_1(X2)|~v9_metric_1(k14_euclid(X1))|~v9_metric_1(X2)|~v8_metric_1(k14_euclid(X1))|~v8_metric_1(X2)|~v7_metric_1(k14_euclid(X1))|~v7_metric_1(X2)|~v6_metric_1(k14_euclid(X1))|~v6_metric_1(X2)|~r1_tarski(k7_relset_1(X5,X6,X3,k1_topreal9(X1,esk7_3(k14_euclid(X1),X2,X3),X4)),k9_metric_1(X2,esk8_3(k14_euclid(X1),X2,X3),esk9_3(k14_euclid(X1),X2,X3)))|~v2_xxreal_0(X4)|~v1_xreal_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(esk7_3(k14_euclid(X1),X2,X3),u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_3474, plain, (r1_tarski(k7_relat_1(X2,k1_topreal9(X1,X3,esk6_6(k14_euclid(X1),k8_metric_1,X2,X3,X4,X5))),k9_metric_1(k8_metric_1,X4,X5))|X4!=k1_funct_1(X2,X3)|~v2_xxreal_0(X5)|~v1_xreal_0(X5)|~v5_pre_topc(X2,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),k1_numbers)))|~m1_subset_1(X3,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X4,k1_numbers)|~v1_funct_2(X2,u1_struct_0(k14_euclid(X1)),k1_numbers)|~v1_funct_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3016, plain, (v2_struct_0(X1)|r1_tarski(k7_relat_1(X3,k9_metric_1(X1,X4,esk6_6(X1,k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k14_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k14_euclid(X2)))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k14_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4195, plain, (u1_metric_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))=u1_metric_1(k14_euclid(X1))|g1_metric_1(k1_numbers,u1_metric_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))!=k14_euclid(X1)|~l1_metric_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|~v7_ordinal1(X1)|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_834, plain, (u1_struct_0(k15_euclid(X1))=u1_struct_0(k15_euclid(X2))|k3_pcomps_1(k14_euclid(X1))!=k14_euclid(X2)|~v7_ordinal1(X1)|~v1_pre_topc(k14_euclid(X2))|~l1_pre_topc(k14_euclid(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_831, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(X2)|v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(X2))|~l1_metric_1(k14_euclid(X1))|~l1_metric_1(X2)|~v9_metric_1(k14_euclid(X1))|~v9_metric_1(X2)|~v8_metric_1(k14_euclid(X1))|~v8_metric_1(X2)|~v7_metric_1(k14_euclid(X1))|~v7_metric_1(X2)|~v6_metric_1(k14_euclid(X1))|~v6_metric_1(X2)|~v2_xxreal_0(X4)|~v1_xreal_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(esk7_3(k14_euclid(X1),X2,X3),u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~m1_subset_1(k7_relat_1(X3,k1_topreal9(X1,esk7_3(k14_euclid(X1),X2,X3),X4)),k1_zfmisc_1(k9_metric_1(X2,esk8_3(k14_euclid(X1),X2,X3),esk9_3(k14_euclid(X1),X2,X3)))))).#trainneg
cnf(i_0_853, plain, (v1_relat_1(k7_relset_1(X2,X3,k7_metric_1,X1))|~v1_relat_1(k1_numbers)|~m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_3460, plain, (v2_struct_0(k14_euclid(X1))|r1_tarski(k7_relset_1(X6,X7,X2,k9_metric_1(k8_metric_1,X3,esk6_6(k8_metric_1,k14_euclid(X1),X2,X3,X4,X5))),k1_topreal9(X1,X4,X5))|X4!=k1_funct_1(X2,X3)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X5)|~v1_xreal_0(X5)|~v5_pre_topc(X2,k3_topmetr,k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(k3_pcomps_1(k14_euclid(X1))))))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X2,k1_numbers,u1_struct_0(k3_pcomps_1(k14_euclid(X1))))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))).#trainneg
cnf(i_0_3055, plain, (v2_struct_0(X1)|v5_pre_topc(X2,k3_pcomps_1(k14_euclid(X3)),k3_pcomps_1(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(X3)),u1_struct_0(k3_pcomps_1(X1)),X2,k9_metric_1(k14_euclid(X3),esk7_3(k14_euclid(X3),X1,X2),X4)),k9_metric_1(X1,esk8_3(k14_euclid(X3),X1,X2),esk9_3(k14_euclid(X3),X1,X2)))|~v2_xxreal_0(X4)|~v1_xreal_0(X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),u1_struct_0(k3_pcomps_1(X1)))))|~v1_funct_2(X2,u1_struct_0(k15_euclid(X3)),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_4520, plain, (v2_struct_0(X1)|r1_tarski(k7_relset_1(X7,X8,X2,k9_metric_1(X1,X3,esk6_6(X1,k14_euclid(X4),X2,X3,X5,X6))),k1_topreal9(X4,X5,X6))|X5!=k1_funct_1(X2,X3)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X2,k3_pcomps_1(X1),k3_pcomps_1(k14_euclid(X4)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k14_euclid(X4)))))|~m1_subset_1(X5,u1_struct_0(k14_euclid(X4)))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X4)))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k14_euclid(X4)))|~v1_funct_1(X2)|~v7_ordinal1(X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_4115, plain, (u1_pre_topc(k15_euclid(X1))=u1_pre_topc(k15_euclid(X2))|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(X2))|~m1_subset_1(u1_pre_topc(k15_euclid(X1)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k14_euclid(X1)))))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3664, plain, (r1_tarski(k7_relat_1(X2,k1_topreal9(X1,X3,esk6_6(k14_euclid(X1),k8_metric_1,X2,X3,X4,X5))),k9_metric_1(k8_metric_1,X4,X5))|X4!=k1_funct_1(X2,X3)|~v2_xxreal_0(X5)|~v1_xreal_0(X5)|~v5_pre_topc(X2,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X4,k1_numbers)|~v1_funct_2(X2,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v1_funct_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4454, plain, (r1_tarski(k7_relset_1(X7,X8,X1,k9_metric_1(k14_euclid(X2),X3,esk6_6(k14_euclid(X2),k14_euclid(X4),X1,X3,X5,X6))),k1_topreal9(X4,X5,X6))|X5!=k1_funct_1(X1,X3)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X1,k3_pcomps_1(k14_euclid(X2)),k3_pcomps_1(k14_euclid(X4)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X2)),u1_struct_0(k15_euclid(X4)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X4)))|~m1_subset_1(X3,u1_struct_0(k14_euclid(X2)))|~v1_funct_2(X1,u1_struct_0(k14_euclid(X2)),u1_struct_0(k15_euclid(X4)))|~v1_funct_1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_3315, plain, (r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(X1)),k1_numbers,X2,k9_metric_1(k14_euclid(X1),X3,esk6_6(k14_euclid(X1),k8_metric_1,X2,X3,X4,X5))),k9_metric_1(k8_metric_1,X4,X5))|X4!=k1_funct_1(X2,X3)|~v2_xxreal_0(X5)|~v1_xreal_0(X5)|~v5_pre_topc(X2,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X4,k1_numbers)|~v1_funct_2(X2,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v1_funct_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3003, plain, (v2_struct_0(X1)|r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k15_euclid(X2)),X3,k9_metric_1(X1,X4,esk6_6(X1,k14_euclid(X2),X3,X4,X5,X6))),k9_metric_1(k14_euclid(X2),X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4425, plain, (u1_struct_0(k15_euclid(X1))=u1_struct_0(X2)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(g1_metric_1(u1_struct_0(X2),u1_metric_1(X2)))|~v7_ordinal1(X1)|~l1_metric_1(X2))).#trainneg
cnf(i_0_4438, plain, (g1_pre_topc(k1_numbers,k2_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))=k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_4208, plain, (u1_metric_1(k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))=u1_metric_1(k14_euclid(X1))|g1_metric_1(k1_numbers,u1_metric_1(k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))))!=k14_euclid(X1)|~l1_metric_1(k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))|~l1_metric_1(k3_topmetr)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_854, plain, (k10_funcsdom(k2_finseq_1(X1))=k15_euclid(X1)|k10_funcsdom(k2_finseq_1(X1))!=k3_pcomps_1(k14_euclid(X1))|~v1_rlvect_1(k10_funcsdom(k2_finseq_1(X1)))|~l1_rltopsp1(k10_funcsdom(k2_finseq_1(X1)))|~v5_rltopsp1(k10_funcsdom(k2_finseq_1(X1)))|~v7_ordinal1(X1)|~v1_pre_topc(k10_funcsdom(k2_finseq_1(X1)))|~l1_pre_topc(k10_funcsdom(k2_finseq_1(X1))))).#trainneg
cnf(i_0_4152, plain, (m1_subset_1(u2_struct_0(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))),k1_numbers)|~l2_struct_0(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_4161, plain, (v2_struct_0(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|~l1_pre_topc(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|~v2_struct_0(g1_pre_topc(k1_numbers,u1_pre_topc(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_855, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(k14_euclid(X2))|v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~l1_metric_1(k14_euclid(X1))|~l1_metric_1(k14_euclid(X2))|~v9_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X2))|~v8_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X2))|~v7_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X2))|~v6_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X2))|~r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(k14_euclid(X2))),X3,k1_topreal9(X1,esk7_3(k14_euclid(X1),k14_euclid(X2),X3),X4)),k1_topreal9(X2,esk8_3(k14_euclid(X1),k14_euclid(X2),X3),esk9_3(k14_euclid(X1),k14_euclid(X2),X3)))|~v2_xxreal_0(X4)|~v1_xreal_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(k14_euclid(X2))))))|~m1_subset_1(esk7_3(k14_euclid(X1),k14_euclid(X2),X3),u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(k14_euclid(X2))))|~v1_funct_1(X3)|~v1_xreal_0(esk9_3(k14_euclid(X1),k14_euclid(X2),X3))|~m1_subset_1(esk8_3(k14_euclid(X1),k14_euclid(X2),X3),u1_struct_0(k14_euclid(X2)))|~m1_subset_1(esk8_3(k14_euclid(X1),k14_euclid(X2),X3),u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2911, plain, (v2_struct_0(X1)|v2_struct_0(k14_euclid(X2))|m1_subset_1(k7_relset_1(u1_struct_0(k14_euclid(X2)),u1_struct_0(k3_pcomps_1(X1)),X3,k9_metric_1(k14_euclid(X2),X4,esk6_6(k14_euclid(X2),X1,X3,X4,X5,X6))),k1_zfmisc_1(k9_metric_1(X1,X5,X6)))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(X1)|~l1_metric_1(k14_euclid(X2))|~v9_metric_1(X1)|~v9_metric_1(k14_euclid(X2))|~v8_metric_1(X1)|~v8_metric_1(k14_euclid(X2))|~v7_metric_1(X1)|~v7_metric_1(k14_euclid(X2))|~v6_metric_1(X1)|~v6_metric_1(k14_euclid(X2))|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X2)),k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X2)),u1_struct_0(k3_pcomps_1(X1)))))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X2)))|~v1_funct_2(X3,u1_struct_0(k14_euclid(X2)),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2164, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~r2_hidden(X1,k7_relset_1(X3,X4,esk2_0,k1_topreal9(esk1_0,X2,esk5_2(X2,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X2)))))|~v1_xboole_0(k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X2),k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X2)))|~v2_xxreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X2))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,X2))|~m1_subset_1(X2,u1_struct_0(k15_euclid(esk1_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_3428, plain, (v7_ordinal1(u2_struct_0(X1))|~v6_membered(k1_numbers)|~l2_struct_0(X1)|~v3_membered(u1_struct_0(X1)))).#trainneg
cnf(i_0_4167, plain, (g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))=k15_euclid(X1)|g1_rlvect_1(k1_numbers,u2_struct_0(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))),u1_algstr_0(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))),u1_rlvect_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))!=k10_funcsdom(k2_finseq_1(X1))|g1_pre_topc(k1_numbers,u1_pre_topc(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))!=k3_pcomps_1(k14_euclid(X1))|~l1_rltopsp1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|~v5_rltopsp1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|~v7_ordinal1(X1)|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_861, negated_conjecture, (m1_subset_1(k7_relat_1(esk2_0,k1_topreal9(esk1_0,X1,esk5_2(X1,X2))),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,X1),X2),k2_xcmplx_0(X2,k1_funct_1(esk2_0,X1)))))|v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v5_pre_topc(esk2_0,k15_euclid(esk1_0),k3_topmetr)|~v2_xxreal_0(X2)|~v1_xreal_0(k1_funct_1(esk2_0,X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_3980, plain, (v3_valued_0(k7_relset_1(X1,X2,k7_relset_1(X10,X11,X7,X8),X9))|~v3_membered(X6)|~m1_subset_1(k7_relset_1(X10,X11,X7,X8),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,k2_zfmisc_1(X5,X6)))))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X10,X11))))).#trainneg
cnf(i_0_860, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(X2)|v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(X2))|~l1_metric_1(k14_euclid(X1))|~l1_metric_1(X2)|~v9_metric_1(k14_euclid(X1))|~v9_metric_1(X2)|~v8_metric_1(k14_euclid(X1))|~v8_metric_1(X2)|~v7_metric_1(k14_euclid(X1))|~v7_metric_1(X2)|~v6_metric_1(k14_euclid(X1))|~v6_metric_1(X2)|~v2_xxreal_0(X4)|~v1_xreal_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(esk7_3(k14_euclid(X1),X2,X3),u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~m1_subset_1(k7_relset_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k3_pcomps_1(X2)),X3,k1_topreal9(X1,esk7_3(k14_euclid(X1),X2,X3),X4)),k1_zfmisc_1(k9_metric_1(X2,esk8_3(k14_euclid(X1),X2,X3),esk9_3(k14_euclid(X1),X2,X3)))))).#trainneg
cnf(i_0_4577, plain, (v2_struct_0(k14_euclid(X1))|r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k15_euclid(X2)),X3,k1_topreal9(X1,X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k9_metric_1(k14_euclid(X2),X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v1_xreal_0(esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4549, plain, (v2_struct_0(k14_euclid(X1))|r1_tarski(k7_relset_1(X2,X3,X4,k1_topreal9(X1,X5,esk6_6(k14_euclid(X1),k14_euclid(X6),X4,X5,X7,X8))),k1_topreal9(X6,X7,X8))|X7!=k1_funct_1(X4,X5)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X8)|~v1_xreal_0(X8)|~v5_pre_topc(X4,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X6)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k14_euclid(X6)))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X7,u1_struct_0(k14_euclid(X6)))|~m1_subset_1(X7,u1_struct_0(k15_euclid(X6)))|~m1_subset_1(X5,u1_struct_0(k14_euclid(X1)))|~v1_funct_2(X4,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),u1_struct_0(k14_euclid(X6)))|~v1_funct_1(X4)|~v7_ordinal1(X6)|~v1_xreal_0(esk6_6(k14_euclid(X1),k14_euclid(X6),X4,X5,X7,X8))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4563, plain, (r1_tarski(k7_relset_1(X1,X2,X3,k1_topreal9(X4,X5,esk6_6(k14_euclid(X4),k14_euclid(X6),X3,X5,X7,X8))),k1_topreal9(X6,X7,X8))|X7!=k1_funct_1(X3,X5)|~v2_xxreal_0(X8)|~v1_xreal_0(X8)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X4)),k3_pcomps_1(k14_euclid(X6)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X4)),u1_struct_0(k15_euclid(X6)))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X7,u1_struct_0(k15_euclid(X6)))|~m1_subset_1(X5,u1_struct_0(k14_euclid(X4)))|~v1_funct_2(X3,u1_struct_0(k14_euclid(X4)),u1_struct_0(k15_euclid(X6)))|~v1_funct_1(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v1_xreal_0(esk6_6(k14_euclid(X4),k14_euclid(X6),X3,X5,X7,X8))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X4)))|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_4578, plain, (r1_tarski(k7_relset_1(k1_numbers,u1_struct_0(k15_euclid(X1)),X2,k9_metric_1(k8_metric_1,X3,esk6_6(k8_metric_1,k14_euclid(X1),X2,X3,X4,X5))),k9_metric_1(k14_euclid(X1),X4,X5))|X4!=k1_funct_1(X2,X3)|~v2_xxreal_0(X5)|~v1_xreal_0(X5)|~v5_pre_topc(X2,k3_topmetr,k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(k15_euclid(X1)))))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X2,k1_numbers,u1_struct_0(k15_euclid(X1)))|~v1_funct_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4580, plain, (v2_struct_0(k14_euclid(X1))|r1_tarski(k7_relset_1(u1_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X2)),X3,k9_metric_1(k14_euclid(X1),X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k9_metric_1(k14_euclid(X2),X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4537, plain, (v5_pre_topc(X1,k3_pcomps_1(k14_euclid(X2)),k3_topmetr)|~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(X2)),k1_numbers,X1,k9_metric_1(k14_euclid(X2),esk7_3(k14_euclid(X2),k8_metric_1,X1),X3)),k9_metric_1(k8_metric_1,esk8_3(k14_euclid(X2),k8_metric_1,X1),esk9_3(k14_euclid(X2),k8_metric_1,X1)))|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~v1_funct_2(X1,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2803, plain, (m1_subset_1(k6_xcmplx_0(X1,X2),k5_numbers)|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4609, plain, (u1_struct_0(X1)=k1_numbers|X1!=k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_891, plain, (g1_rlvect_1(k1_numbers,u2_struct_0(k3_topmetr),u1_algstr_0(k3_topmetr),u1_rlvect_1(k3_topmetr))=k3_topmetr|~v1_rlvect_1(k3_topmetr)|~l1_rlvect_1(k3_topmetr))).#trainneg
cnf(i_0_4654, plain, (m1_subset_1(u2_struct_0(X1),k5_numbers)|~v6_membered(k1_numbers)|~l2_struct_0(X1)|~v3_membered(u1_struct_0(X1)))).#trainneg
cnf(i_0_4620, plain, (v1_xreal_0(u2_struct_0(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))|~l2_struct_0(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_3986, plain, (v1_valued_0(k7_relset_1(X1,X2,k7_relset_1(X10,X11,X7,X8),X9))|~v1_membered(X6)|~m1_subset_1(k7_relset_1(X10,X11,X7,X8),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,k2_zfmisc_1(X5,X6)))))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X10,X11))))).#trainneg
cnf(i_0_894, plain, (g1_rltopsp1(k1_numbers,u2_struct_0(k3_topmetr),u1_algstr_0(k3_topmetr),u1_rlvect_1(k3_topmetr),u1_pre_topc(k3_topmetr))=k3_topmetr|~l1_rltopsp1(k3_topmetr)|~v5_rltopsp1(k3_topmetr))).#trainneg
cnf(i_0_906, plain, (u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))=k1_numbers|v2_struct_0(X1)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_topmetr|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5, negated_conjecture, (v2_xxreal_0(esk4_0))).#trainneg
cnf(i_0_6, negated_conjecture, (v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_3160, negated_conjecture, (k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~v2_xxreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))|~v1_xreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr))).#trainneg
cnf(i_0_416, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0)))))).#trainneg
cnf(i_0_4, negated_conjecture, (~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_175, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0)))))).#trainneg
cnf(i_0_3159, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v1_xreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr))).#trainneg
cnf(i_0_1276, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))))))).#trainneg
cnf(i_0_338, negated_conjecture, (~r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1999, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0)))))).#trainneg
cnf(i_0_907, plain, (u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))=k1_numbers|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_topmetr|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4619, plain, (v7_ordinal1(u2_struct_0(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))|~v6_membered(k1_numbers)|~l2_struct_0(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_4174, plain, (v7_ordinal1(u1_pre_topc(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_pre_topc(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_3623, plain, (v1_xboole_0(k7_relset_1(X8,X9,X3,X4))|v1_xboole_0(X5)|m1_subset_1(k3_funct_2(X5,k7_relset_1(X8,X9,X3,X4),X6,X7),X2)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,k7_relset_1(X8,X9,X3,X4))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X7,X5)|~v1_funct_2(X6,X5,k7_relset_1(X8,X9,X3,X4))|~v1_funct_1(X6)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X8,X9))))).#trainneg
cnf(i_0_908, plain, (u1_struct_0(g1_pre_topc(X1,X2))=k1_numbers|g1_pre_topc(X1,X2)!=k3_topmetr|~l1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_909, plain, (u1_struct_0(k3_pcomps_1(X1))=k1_numbers|k3_pcomps_1(X1)!=k3_topmetr|~l1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_4939, plain, (u1_struct_0(k3_pcomps_1(k14_euclid(X1)))=k1_numbers|k3_pcomps_1(k14_euclid(X1))!=k3_topmetr|~m1_subset_1(k2_pcomps_1(k14_euclid(X1)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1147, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(X1)|~m1_subset_1(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(esk4_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))))))).#trainneg
cnf(i_0_4750, plain, (k15_euclid(X1)=k3_topmetr|k10_funcsdom(k2_finseq_1(X1))!=k3_topmetr|k3_pcomps_1(k14_euclid(X1))!=k3_topmetr|~l1_rltopsp1(k3_topmetr)|~v5_rltopsp1(k3_topmetr)|~v7_ordinal1(X1)|~v1_rlvect_1(k3_topmetr)|~l1_rlvect_1(k3_topmetr))).#trainneg
cnf(i_0_4782, negated_conjecture, (v3_valued_0(esk3_0))).#trainneg
cnf(i_0_1930, negated_conjecture, (~r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))|~v2_xxreal_0(X1)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_912, plain, (m1_subset_1(k9_metric_1(k8_metric_1,X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(k3_real_1(X2,X1))|~v1_xxreal_0(k5_real_1(X1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4783, negated_conjecture, (v1_finseq_1(esk3_0))).#trainneg
cnf(i_0_4780, negated_conjecture, (v7_ordinal1(esk3_0)|~v6_membered(u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_3615, plain, (v3_membered(k3_funct_2(X1,k7_relset_1(X2,k1_zfmisc_1(X3),X4,X5),X6,X7))|v1_xboole_0(k7_relset_1(X2,k1_zfmisc_1(X3),X4,X5))|v1_xboole_0(X1)|~v3_membered(X3)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k7_relset_1(X2,k1_zfmisc_1(X3),X4,X5))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k1_zfmisc_1(X3))))|~m1_subset_1(X7,X1)|~v1_funct_2(X6,X1,k7_relset_1(X2,k1_zfmisc_1(X3),X4,X5))|~v1_funct_1(X6))).#trainneg
cnf(i_0_913, plain, (v3_membered(k9_metric_1(k8_metric_1,X1,X2))|~v1_xxreal_0(k3_real_1(X2,X1))|~v1_xxreal_0(k5_real_1(X1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4730, plain, (r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)),X3,k9_metric_1(k14_euclid(X1),X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k9_metric_1(k14_euclid(X2),X5,X6))|X5!=k1_funct_1(X3,X4)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4773, plain, (u1_struct_0(g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1))))=k1_numbers|v2_struct_0(k14_euclid(X1))|g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))!=k3_topmetr|~l1_pre_topc(g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1))))|~l1_pre_topc(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4787, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(esk1_0),X1)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_topmetr|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(X1))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(X1)))).#trainneg
cnf(i_0_1931, negated_conjecture, (~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))|~v2_xxreal_0(X1)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4784, negated_conjecture, (m1_subset_1(esk4_0,k1_numbers))).#trainneg
cnf(i_0_1994, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(X1)|~m1_subset_1(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))))).#trainneg
cnf(i_0_921, plain, (m1_subset_1(k1_topreal9(X1,X2,X3),k1_numbers)|~v3_membered(k1_zfmisc_1(u1_struct_0(k15_euclid(X1))))|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4781, negated_conjecture, (v1_xreal_0(esk3_0)|~v3_membered(u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_3616, plain, (v1_relat_1(k3_funct_2(X1,k7_relset_1(X2,k1_zfmisc_1(X3),X4,X5),X6,X7))|v1_xboole_0(k7_relset_1(X2,k1_zfmisc_1(X3),X4,X5))|v1_xboole_0(X1)|~v1_relat_1(X3)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k7_relset_1(X2,k1_zfmisc_1(X3),X4,X5))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k1_zfmisc_1(X3))))|~m1_subset_1(X7,X1)|~v1_funct_2(X6,X1,k7_relset_1(X2,k1_zfmisc_1(X3),X4,X5))|~v1_funct_1(X6))).#trainneg
cnf(i_0_3808, plain, (v1_xboole_0(k7_relset_1(X8,X9,X3,X4))|v1_xboole_0(X5)|v7_ordinal1(k3_funct_2(X5,k7_relset_1(X8,X9,X3,X4),X6,X7))|~v6_membered(X2)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,k7_relset_1(X8,X9,X3,X4))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X7,X5)|~v1_funct_2(X6,X5,k7_relset_1(X8,X9,X3,X4))|~v1_funct_1(X6)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X8,X9))))).#trainneg
cnf(i_0_922, plain, (v3_valued_0(k7_relat_1(X4,X5))|~v3_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_925, plain, (u1_pre_topc(g1_pre_topc(X1,X2))=X2|~v1_pre_topc(g1_pre_topc(X1,X2))|~l1_pre_topc(g1_pre_topc(X1,X2)))).#trainneg
cnf(i_0_5083, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(esk1_0),k14_euclid(X1))|g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))!=k3_topmetr|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(k14_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k15_euclid(X1)))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k15_euclid(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4796, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(k15_euclid(esk1_0)),u1_pre_topc(k15_euclid(esk1_0)))!=X1|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_pre_topc(X1))).#trainneg
cnf(i_0_5088, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(esk1_0),k15_euclid(X1))|k3_pcomps_1(k14_euclid(X1))!=k3_topmetr|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(k15_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k15_euclid(X1)))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k15_euclid(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_497, negated_conjecture, (~r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(esk4_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))))|~v2_xxreal_0(X1)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_496, negated_conjecture, (~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(esk4_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))))|~v2_xxreal_0(X1)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4786, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~v1_xreal_0(esk6_6(k14_euclid(esk1_0),k8_metric_1,esk2_0,esk3_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),k1_numbers)|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers))).#trainneg
cnf(i_0_4423, plain, (u1_struct_0(X1)=u1_struct_0(X2)|k3_pcomps_1(X1)!=k3_pcomps_1(g1_metric_1(u1_struct_0(X2),u1_metric_1(X2)))|~l1_metric_1(X1)|~l1_metric_1(X2))).#trainneg
cnf(i_0_4611, plain, (k2_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))=u1_pre_topc(k3_topmetr)|k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))!=k3_topmetr|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_1998, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(X1)|~m1_subset_1(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))))).#trainneg
cnf(i_0_926, plain, (u1_pre_topc(X1)=k2_pcomps_1(k14_euclid(X2))|X1!=k3_pcomps_1(k14_euclid(X2))|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3813, plain, (v1_xboole_0(k7_relset_1(X8,X9,X3,X4))|v1_xboole_0(X5)|v1_xreal_0(k3_funct_2(X5,k7_relset_1(X8,X9,X3,X4),X6,X7))|~v3_membered(X2)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,k7_relset_1(X8,X9,X3,X4))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X7,X5)|~v1_funct_2(X6,X5,k7_relset_1(X8,X9,X3,X4))|~v1_funct_1(X6)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X8,X9))))).#trainneg
cnf(i_0_5170, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|k3_pcomps_1(k14_euclid(esk1_0))!=X1|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~v1_pre_topc(X1)|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2)))).#trainneg
cnf(i_0_5167, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(esk1_0),k14_euclid(X1))|k14_euclid(X1)!=k3_topmetr|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(k14_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k15_euclid(X1)))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(k15_euclid(X1)))|~v7_ordinal1(X1)|~v1_pre_topc(k14_euclid(X1)))).#trainneg
cnf(i_0_5169, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|k15_euclid(esk1_0)!=X1|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~v1_pre_topc(X1)|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_pre_topc(k15_euclid(esk1_0)))).#trainneg
cnf(i_0_5160, negated_conjecture, (v5_pre_topc(esk2_0,X1,k14_euclid(X2))|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|g1_pre_topc(u1_struct_0(k15_euclid(X2)),u1_pre_topc(k14_euclid(X2)))!=k3_topmetr|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k14_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(k14_euclid(X2))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5166, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_pre_topc(X2))).#trainneg
cnf(i_0_5165, negated_conjecture, (v5_pre_topc(esk2_0,X1,k15_euclid(X2))|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k15_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(k15_euclid(X2))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5098, negated_conjecture, (v7_ordinal1(esk4_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1960, negated_conjecture, (~r1_tarski(k7_relset_1(X1,X2,esk2_0,k1_topreal9(esk1_0,esk3_0,X3)),k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))|~v2_xxreal_0(X3)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(X3)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5220, negated_conjecture, (v1_xboole_0(k7_relset_1(X1,X2,esk2_0,X3))|v1_xboole_0(X4)|v1_xreal_0(k3_funct_2(X4,k7_relset_1(X1,X2,esk2_0,X3),X5,X6))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,k7_relset_1(X1,X2,esk2_0,X3))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,X4)|~v1_funct_2(X5,X4,k7_relset_1(X1,X2,esk2_0,X3))|~v1_funct_1(X5))).#trainneg
cnf(i_0_937, plain, (m1_subset_1(k9_metric_1(k8_metric_1,X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(k3_real_1(X1,X2))|~v1_xxreal_0(k5_real_1(X1,X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_5403, plain, (v2_struct_0(X1)|v2_struct_0(k14_euclid(X2))|r1_tarski(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k15_euclid(X2)),X3,k9_metric_1(X1,X4,esk6_6(X1,k14_euclid(X2),X3,X4,X5,X6))),k9_metric_1(k14_euclid(X2),X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X2))|~l1_metric_1(X1)|~v9_metric_1(k14_euclid(X2))|~v9_metric_1(X1)|~v8_metric_1(k14_euclid(X2))|~v8_metric_1(X1)|~v7_metric_1(k14_euclid(X2))|~v7_metric_1(X1)|~v6_metric_1(k14_euclid(X2))|~v6_metric_1(X1)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k14_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5387, plain, (v2_struct_0(k14_euclid(X1))|v2_struct_0(X2)|r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k3_pcomps_1(X2)),X3,k9_metric_1(k14_euclid(X1),X4,esk6_6(k14_euclid(X1),X2,X3,X4,X5,X6))),k9_metric_1(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(X2)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(X2)|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(X2)|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(X2)|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(X2)|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k15_euclid(X1)),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5264, negated_conjecture, (v5_pre_topc(esk2_0,X1,k3_topmetr)|k15_euclid(esk1_0)!=X1|~v1_pre_topc(k15_euclid(esk1_0))|~v1_pre_topc(X1)|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers))).#trainneg
cnf(i_0_5149, plain, (u1_pre_topc(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))=u1_pre_topc(k8_metric_1)|~l1_pre_topc(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_1993, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(X1)|~m1_subset_1(k7_relset_1(X2,X3,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k9_metric_1(k8_metric_1,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_938, plain, (v3_membered(k9_metric_1(k8_metric_1,X1,X2))|~v1_xxreal_0(k3_real_1(X1,X2))|~v1_xxreal_0(k5_real_1(X1,X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_5207, plain, (v3_membered(u1_pre_topc(k3_topmetr))|~v3_membered(k1_zfmisc_1(u1_struct_0(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))))|~l1_metric_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))!=k3_topmetr|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_5208, plain, (v1_relat_1(u1_pre_topc(k3_topmetr))|~v1_relat_1(k1_zfmisc_1(u1_struct_0(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))))|~l1_metric_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))!=k3_topmetr|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_955, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X3)|~r2_hidden(X1,k7_relat_1(k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X2)),X5,X6),X7))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))))|~m1_subset_1(X6,X3)|~v1_funct_2(X5,X3,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))|~v1_funct_1(X5))).#trainneg
cnf(i_0_956, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relat_1(k7_relset_1(X3,k2_zfmisc_1(X4,X2),X5,X6),X7))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,X2)))))).#trainneg
cnf(i_0_4832, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4797, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k15_euclid(esk1_0)|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_pre_topc(k15_euclid(esk1_0)))).#trainneg
cnf(i_0_4806, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5089, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(esk1_0),X1)|X1!=k3_topmetr|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(X1))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(X1))|~v1_pre_topc(X1))).#trainneg
cnf(i_0_4826, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0))))))).#trainneg
cnf(i_0_3940, plain, (m1_subset_1(u1_pre_topc(k15_euclid(X1)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4947, plain, (u1_struct_0(k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))=k1_numbers|k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))!=k3_topmetr|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_477, negated_conjecture, (~r1_tarski(k7_relset_1(X1,X2,esk2_0,k1_topreal9(esk1_0,esk3_0,X3)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0)))|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_970, plain, (v1_xboole_0(k7_metric_1)|m1_subset_1(k3_funct_2(X1,k7_metric_1,X2,X3),k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_metric_1)))|~m1_subset_1(X3,X1)|~v1_funct_2(X2,X1,k7_metric_1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_5573, plain, (~l1_metric_1(k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))|~v2_struct_0(k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))|k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))!=k3_topmetr|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_1287, plain, (v3_membered(k7_relset_1(X1,X2,k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)),X6,X7),X8))|v1_xboole_0(X3)|~v3_membered(X5)|~m1_subset_1(k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)),X6,X7),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))))|~m1_subset_1(X7,X3)|~v1_funct_2(X6,X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~v1_funct_1(X6))).#trainneg
cnf(i_0_1337, plain, (v1_relat_1(k7_relset_1(X1,X2,k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)),X6,X7),X8))|v1_xboole_0(X3)|~v1_relat_1(X5)|~m1_subset_1(k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)),X6,X7),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))))|~m1_subset_1(X7,X3)|~v1_funct_2(X6,X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~v1_funct_1(X6))).#trainneg
cnf(i_0_974, plain, (v1_valued_0(k7_relat_1(X4,X5))|~v1_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_984, plain, (v7_ordinal1(k2_pcomps_1(k14_euclid(X1)))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))|~l1_metric_1(k14_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5532, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(esk1_0),X1)|X1!=k3_topmetr|~v1_pre_topc(X1)|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(X1))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(esk1_0)),u1_struct_0(X1)))).#trainneg
cnf(i_0_5521, negated_conjecture, (v5_pre_topc(esk2_0,X1,k3_topmetr)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k15_euclid(esk1_0)|~v1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers))).#trainneg
cnf(i_0_5516, negated_conjecture, (v5_pre_topc(esk2_0,X1,k14_euclid(X2))|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k15_euclid(esk1_0)|g1_pre_topc(u1_struct_0(k15_euclid(X2)),u1_pre_topc(k14_euclid(X2)))!=k3_topmetr|~v1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k14_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(k14_euclid(X2))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5523, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k15_euclid(esk1_0)|X2!=k3_topmetr|~v1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_pre_topc(X2))).#trainneg
cnf(i_0_5522, negated_conjecture, (v5_pre_topc(esk2_0,X1,k15_euclid(X2))|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k15_euclid(esk1_0)|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~v1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k15_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(k15_euclid(X2))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5266, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|X2!=k3_topmetr|k15_euclid(esk1_0)!=X1|~v1_pre_topc(k15_euclid(esk1_0))|~v1_pre_topc(X1)|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_pre_topc(X2))).#trainneg
cnf(i_0_5280, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|X1!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_pre_topc(X1))).#trainneg
cnf(i_0_5498, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relat_1(k7_relat_1(X5,X6),X7))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,X2)))))).#trainneg
cnf(i_0_1081, negated_conjecture, (~r1_tarski(k7_relset_1(X1,X2,esk2_0,k1_topreal9(esk1_0,esk3_0,X3)),k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(esk4_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))))|~v2_xxreal_0(X3)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(X3)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1007, plain, (m1_subset_1(u1_rlvect_1(k8_metric_1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~l1_rlvect_1(k8_metric_1))).#trainneg
cnf(i_0_5649, plain, (k3_pcomps_1(k3_topmetr)!=k3_topmetr|~l1_metric_1(k3_pcomps_1(k3_topmetr))|~l1_metric_1(k3_topmetr)|~v2_struct_0(k3_pcomps_1(k3_topmetr))|~v1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_1009, plain, (g1_rlvect_1(k1_numbers,u2_struct_0(k8_metric_1),u1_algstr_0(k8_metric_1),u1_rlvect_1(k8_metric_1))=k8_metric_1|~v1_rlvect_1(k8_metric_1)|~l1_rlvect_1(k8_metric_1))).#trainneg
cnf(i_0_1012, plain, (g1_rltopsp1(k1_numbers,u2_struct_0(k8_metric_1),u1_algstr_0(k8_metric_1),u1_rlvect_1(k8_metric_1),u1_pre_topc(k8_metric_1))=k8_metric_1|~l1_rltopsp1(k8_metric_1)|~v5_rltopsp1(k8_metric_1))).#trainneg
cnf(i_0_5292, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(k15_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4722, plain, (r1_tarski(k7_relat_1(X2,k9_metric_1(k8_metric_1,X3,esk6_6(k8_metric_1,k14_euclid(X1),X2,X3,X4,X5))),k9_metric_1(k14_euclid(X1),X4,X5))|X4!=k1_funct_1(X2,X3)|~v2_xxreal_0(X5)|~v1_xreal_0(X5)|~v5_pre_topc(X2,k3_topmetr,k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(k15_euclid(X1)))))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X2,k1_numbers,u1_struct_0(k15_euclid(X1)))|~v1_funct_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5699, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k15_euclid(esk1_0)|~v1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(k15_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5144, plain, (k2_pcomps_1(k14_euclid(X1))=u1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~v1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~l1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5733, negated_conjecture, (v5_pre_topc(esk2_0,X1,k15_euclid(X2))|X1!=k15_euclid(esk1_0)|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~v1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2)|~v1_pre_topc(X1))).#trainneg
cnf(i_0_5395, plain, (v2_struct_0(k14_euclid(X1))|r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(X2)),u1_struct_0(k15_euclid(X1)),X3,k9_metric_1(k14_euclid(X2),X4,esk6_6(k14_euclid(X2),k14_euclid(X1),X3,X4,X5,X6))),k9_metric_1(k14_euclid(X1),X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X2)),k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),u1_struct_0(k15_euclid(X1)))))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X5,u1_struct_0(k14_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k15_euclid(X2)),u1_struct_0(k15_euclid(X1)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5635, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relset_1(X1,X2,esk2_0,k1_topreal9(esk1_0,esk3_0,X3)),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0)))|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5762, plain, (v3_valued_0(u1_rlvect_1(k8_metric_1))|~l1_rlvect_1(k8_metric_1))).#trainneg
cnf(i_0_513, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(k7_relset_1(X2,X3,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_5763, plain, (v1_valued_0(u1_rlvect_1(k8_metric_1))|~v1_membered(k1_numbers)|~l1_rlvect_1(k8_metric_1))).#trainneg
cnf(i_0_5574, plain, (~l1_pre_topc(k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))|~v2_struct_0(k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))|k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))!=k3_topmetr|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_1681, plain, (v7_ordinal1(k7_relset_1(X1,X2,k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)),X6,X7),X8))|v1_xboole_0(X3)|~v6_membered(k1_zfmisc_1(X5))|~m1_subset_1(k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)),X6,X7),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))))|~m1_subset_1(X7,X3)|~v1_funct_2(X6,X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~v1_funct_1(X6))).#trainneg
cnf(i_0_1741, plain, (v1_xreal_0(k7_relset_1(X1,X2,k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)),X6,X7),X8))|v1_xboole_0(X3)|~v3_membered(k1_zfmisc_1(X5))|~m1_subset_1(k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)),X6,X7),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))))|~m1_subset_1(X7,X3)|~v1_funct_2(X6,X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~v1_funct_1(X6))).#trainneg
cnf(i_0_1076, plain, (m1_subset_1(X1,X2)|v1_xboole_0(k7_relat_1(k2_zfmisc_1(X3,X2),X4))|~m1_subset_1(X1,k7_relat_1(k2_zfmisc_1(X3,X2),X4)))).#trainneg
cnf(i_0_4731, plain, (r1_tarski(k7_relat_1(X3,k9_metric_1(k14_euclid(X1),X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k9_metric_1(k14_euclid(X2),X5,X6))|X5!=k1_funct_1(X3,X4)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1270, negated_conjecture, (~v2_xxreal_0(X1)|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v1_xreal_0(X1)|~m1_subset_1(k7_relset_1(X2,X3,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),esk4_0),k2_xcmplx_0(esk4_0,k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1986, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5543, plain, (u1_pre_topc(k15_euclid(X1))=X2|k3_pcomps_1(k14_euclid(X1))!=g1_pre_topc(X3,X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5768, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k7_relat_1(u1_rlvect_1(k8_metric_1),X2))|~l1_rlvect_1(k8_metric_1))).#trainneg
cnf(i_0_5739, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relat_1(k7_relat_1(k7_relat_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,k2_zfmisc_1(X5,X2))),X6),X7),X8)))).#trainneg
cnf(i_0_3310, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers))).#trainneg
cnf(i_0_1077, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relset_1(X5,X6,esk10_1(k1_zfmisc_1(k2_zfmisc_1(X3,X2))),X4))|~m1_subset_1(esk10_1(k1_zfmisc_1(k2_zfmisc_1(X3,X2))),k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_1078, plain, (m1_subset_1(X1,X2)|v1_xboole_0(k7_relat_1(esk10_1(k1_zfmisc_1(k2_zfmisc_1(X3,X2))),X4))|~m1_subset_1(X1,k7_relat_1(esk10_1(k1_zfmisc_1(k2_zfmisc_1(X3,X2))),X4)))).#trainneg
cnf(i_0_1083, plain, (m1_subset_1(k2_pcomps_1(X1),k1_numbers)|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_metric_1(X1))).#trainneg
cnf(i_0_5781, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5889, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~m1_subset_1(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),k1_numbers))).#trainneg
cnf(i_0_5888, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v7_ordinal1(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)))).#trainneg
cnf(i_0_5890, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers))).#trainneg
cnf(i_0_5848, negated_conjecture, (v5_pre_topc(esk2_0,k3_topmetr,k3_topmetr)|k3_pcomps_1(k14_euclid(esk1_0))!=k3_topmetr|~l1_pre_topc(k15_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(esk2_0,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_840, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0)))))).#trainneg
cnf(i_0_5736, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relat_1(k7_relat_1(k7_relat_1(esk10_1(k1_zfmisc_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,k2_zfmisc_1(X5,X2))))),X6),X7),X8)))).#trainneg
cnf(i_0_1088, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relat_1(k7_relset_1(X3,X4,k2_zfmisc_1(X5,k2_zfmisc_1(X6,X2)),X7),X8))|~m1_subset_1(k2_zfmisc_1(X5,k2_zfmisc_1(X6,X2)),k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_5824, plain, (k3_pcomps_1(k3_topmetr)!=k3_topmetr|~l1_pre_topc(k3_pcomps_1(k3_topmetr))|~l1_metric_1(k3_topmetr)|~v2_struct_0(k3_pcomps_1(k3_topmetr))|~v1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_5825, plain, (k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))!=k3_topmetr|~l1_metric_1(k3_topmetr)|~v2_struct_0(k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))|~l1_metric_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))).#trainneg
cnf(i_0_2082, plain, (v3_valued_0(k7_relset_1(X1,X2,k7_relset_1(X3,X4,k2_zfmisc_1(X5,k2_zfmisc_1(X6,k2_zfmisc_1(X7,X8))),X9),X10))|~v3_membered(X8)|~m1_subset_1(k7_relset_1(X3,X4,k2_zfmisc_1(X5,k2_zfmisc_1(X6,k2_zfmisc_1(X7,X8))),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_zfmisc_1(X5,k2_zfmisc_1(X6,k2_zfmisc_1(X7,X8))),k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_1090, plain, (m1_subset_1(k7_relat_1(k7_relset_1(X1,X2,k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5)),X6),X7),k1_zfmisc_1(X5))|~m1_subset_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5)),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5716, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(X1),X2)|k3_pcomps_1(k14_euclid(X1))!=k15_euclid(esk1_0)|X2!=k3_topmetr|~v1_pre_topc(k15_euclid(esk1_0))|~v1_pre_topc(X2)|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X2)|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X2)|~v2_pre_topc(k15_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5692, negated_conjecture, (v5_pre_topc(esk2_0,k14_euclid(X1),k3_topmetr)|g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))!=k15_euclid(esk1_0)|~v1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(k14_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1776, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1)|~m1_subset_1(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(esk4_0,k1_funct_1(esk2_0,esk3_0))))))).#trainneg
cnf(i_0_5390, plain, (v2_struct_0(k14_euclid(X1))|r1_tarski(k7_relset_1(u1_struct_0(k14_euclid(X2)),u1_struct_0(k15_euclid(X1)),X3,k9_metric_1(k14_euclid(X2),X4,esk6_6(k14_euclid(X2),k14_euclid(X1),X3,X4,X5,X6))),k9_metric_1(k14_euclid(X1),X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X2)),k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X2)),u1_struct_0(k15_euclid(X1)))))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X2)))|~m1_subset_1(X5,u1_struct_0(k14_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k14_euclid(X2)),u1_struct_0(k15_euclid(X1)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2225, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5777, plain, (v3_membered(k7_relat_1(u1_rlvect_1(k8_metric_1),X1))|~l1_rlvect_1(k8_metric_1))).#trainneg
cnf(i_0_2132, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,esk3_0),esk4_0)))))).#trainneg
cnf(i_0_1132, plain, (m1_subset_1(X1,u1_struct_0(X2))|~l1_rlvect_1(X2)|~r2_hidden(X1,k7_relset_1(X4,X5,u1_rlvect_1(X2),X3))|~m1_subset_1(u1_rlvect_1(X2),k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_5901, plain, (m1_subset_1(u1_pre_topc(k3_topmetr),k1_numbers)|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers))))).#trainneg
cnf(i_0_2130, plain, (v1_valued_0(k7_relset_1(X1,X2,k7_relset_1(X3,X4,k2_zfmisc_1(X5,k2_zfmisc_1(X6,k2_zfmisc_1(X7,X8))),X9),X10))|~v1_membered(X8)|~m1_subset_1(k7_relset_1(X3,X4,k2_zfmisc_1(X5,k2_zfmisc_1(X6,k2_zfmisc_1(X7,X8))),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_zfmisc_1(X5,k2_zfmisc_1(X6,k2_zfmisc_1(X7,X8))),k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_1133, plain, (m1_subset_1(X1,u1_struct_0(X2))|v1_xboole_0(k7_relat_1(u1_rlvect_1(X2),X3))|~l1_rlvect_1(X2)|~m1_subset_1(X1,k7_relat_1(u1_rlvect_1(X2),X3)))).#trainneg
cnf(i_0_1151, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|v1_xboole_0(k2_pcomps_1(X2))|~l1_metric_1(X2)|~m1_subset_1(X1,k2_pcomps_1(X2)))).#trainneg
cnf(i_0_5260, negated_conjecture, (v5_pre_topc(esk2_0,X1,k14_euclid(X2))|g1_pre_topc(u1_struct_0(k15_euclid(X2)),u1_pre_topc(k14_euclid(X2)))!=k3_topmetr|k15_euclid(esk1_0)!=X1|~v1_pre_topc(k15_euclid(esk1_0))|~v1_pre_topc(X1)|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k14_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(k14_euclid(X2))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1079, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(esk4_0,k1_funct_1(esk2_0,esk3_0))))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1325, plain, (v7_ordinal1(k2_xcmplx_0(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1152, plain, (v3_membered(k7_relat_1(k2_zfmisc_1(X3,X4),X5))|~v3_membered(X4)|~m1_subset_1(k2_zfmisc_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1155, plain, (v1_relat_1(k7_relat_1(k2_zfmisc_1(X3,X4),X5))|~v1_relat_1(X4)|~m1_subset_1(k2_zfmisc_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5732, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(X1),k15_euclid(X2))|k3_pcomps_1(k14_euclid(X1))!=k15_euclid(esk1_0)|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~v1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(k15_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5287, negated_conjecture, (v5_pre_topc(esk2_0,k14_euclid(X1),k3_topmetr)|g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(k14_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1177, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(esk4_0,k1_funct_1(esk2_0,esk3_0))))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5393, plain, (v2_struct_0(k14_euclid(X1))|r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k14_euclid(X2)),X3,k9_metric_1(k14_euclid(X1),X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k14_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k14_euclid(X2)))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k15_euclid(X1)),u1_struct_0(k14_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5912, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6042, plain, (r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)),X3,k9_metric_1(k14_euclid(X1),X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1326, plain, (v7_ordinal1(k3_real_1(X2,X1))|~v6_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2155, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1)|~m1_subset_1(k7_relat_1(esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,esk3_0),esk4_0))))).#trainneg
cnf(i_0_1159, plain, (v1_xreal_0(k2_pcomps_1(k14_euclid(X1)))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6027, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v6_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1161, plain, (v3_membered(k2_pcomps_1(k14_euclid(X1)))|~v3_membered(k1_zfmisc_1(u1_struct_0(k15_euclid(X1))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1162, plain, (v1_relat_1(k2_pcomps_1(k14_euclid(X1)))|~v1_relat_1(k1_zfmisc_1(u1_struct_0(k15_euclid(X1))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6033, negated_conjecture, (v5_pre_topc(esk2_0,k14_euclid(X1),k3_topmetr)|k14_euclid(X1)!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(k14_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1)|~v1_pre_topc(k14_euclid(X1)))).#trainneg
cnf(i_0_5981, negated_conjecture, (v5_pre_topc(esk2_0,k14_euclid(X1),k3_topmetr)|k14_euclid(X1)!=k15_euclid(esk1_0)|~v1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(k14_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1)|~v1_pre_topc(k14_euclid(X1)))).#trainneg
cnf(i_0_5979, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),k1_numbers)|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~l1_pre_topc(k15_euclid(esk1_0)))).#trainneg
cnf(i_0_5980, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~l1_pre_topc(k15_euclid(esk1_0)))).#trainneg
cnf(i_0_6044, plain, (r1_tarski(k7_relset_1(X7,X8,X3,k9_metric_1(k14_euclid(X1),X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k14_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k14_euclid(X2)))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k15_euclid(X1)),u1_struct_0(k14_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_6023, negated_conjecture, (v5_pre_topc(esk2_0,X1,k14_euclid(X2))|k14_euclid(X2)!=k3_topmetr|k15_euclid(esk1_0)!=X1|~v1_pre_topc(k15_euclid(esk1_0))|~v1_pre_topc(X1)|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k14_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(k14_euclid(X2))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2)|~v1_pre_topc(k14_euclid(X2)))).#trainneg
cnf(i_0_5915, negated_conjecture, (v5_pre_topc(esk2_0,k3_topmetr,k3_topmetr)|k3_pcomps_1(k14_euclid(esk1_0))!=k3_topmetr|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(esk2_0,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_2151, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1)|~m1_subset_1(k7_relset_1(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,esk3_0),esk4_0))))).#trainneg
cnf(i_0_1163, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(X2))))|~r2_hidden(X1,k2_pcomps_1(k14_euclid(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5472, plain, (v3_membered(u1_pre_topc(k8_metric_1))|~v3_membered(k1_zfmisc_1(u1_struct_0(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))))|~l1_pre_topc(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_5063, plain, (v3_membered(k3_funct_2(X1,k7_relset_1(X8,X9,X4,X5),X6,X7))|v1_xboole_0(k7_relset_1(X8,X9,X4,X5))|v1_xboole_0(X1)|~v3_membered(X3)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k7_relset_1(X8,X9,X4,X5))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k1_zfmisc_1(X3))))|~m1_subset_1(X7,X1)|~v1_funct_2(X6,X1,k7_relset_1(X8,X9,X4,X5))|~v1_funct_1(X6)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X8,X9))))).#trainneg
cnf(i_0_1184, plain, (m1_subset_1(k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5),k5_numbers)|~v6_membered(k1_zfmisc_1(X4))|~m1_subset_1(k2_zfmisc_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1185, plain, (m1_subset_1(k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5),k1_numbers)|~v6_membered(k1_zfmisc_1(X4))|~m1_subset_1(k2_zfmisc_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6043, plain, (r1_tarski(k7_relat_1(X3,k9_metric_1(k14_euclid(X1),X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k14_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k14_euclid(X2)))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k15_euclid(X1)),u1_struct_0(k14_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5990, plain, (r1_tarski(k7_relat_1(X3,k9_metric_1(k14_euclid(X1),X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k9_metric_1(k14_euclid(X2),X5,X6))|X5!=k1_funct_1(X3,X4)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X5,u1_struct_0(k14_euclid(X2)))|~v1_funct_2(X3,u1_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5991, plain, (r1_tarski(k7_relset_1(X7,X8,X3,k9_metric_1(k14_euclid(X1),X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k9_metric_1(k14_euclid(X2),X5,X6))|X5!=k1_funct_1(X3,X4)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X5,u1_struct_0(k14_euclid(X2)))|~v1_funct_2(X3,u1_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_5913, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(esk1_0)),k3_topmetr)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v7_ordinal1(k1_funct_1(esk2_0,esk3_0)))).#trainneg
cnf(i_0_6025, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relset_1(X2,X3,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(esk4_0,k1_funct_1(esk2_0,esk3_0))))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2416, plain, (u1_struct_0(X1)=X2|g1_metric_1(u1_struct_0(X1),k7_metric_1)!=g1_metric_1(X2,X3)|~l1_metric_1(X1)|X1!=k8_metric_1|~v1_metric_1(X1))).#trainneg
cnf(i_0_4175, plain, (v1_xreal_0(u1_pre_topc(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_pre_topc(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_2226, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~r1_tarski(k7_relset_1(X1,X2,esk2_0,k1_topreal9(esk1_0,esk3_0,X3)),k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,esk3_0),esk4_0))|~v2_xxreal_0(X3)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X3)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1186, plain, (v7_ordinal1(k7_relat_1(k2_zfmisc_1(X3,X4),X5))|~v6_membered(k1_zfmisc_1(X4))|~m1_subset_1(k2_zfmisc_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6116, plain, (v1_xreal_0(k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5))|~v6_membered(k1_zfmisc_1(X4))|~m1_subset_1(k2_zfmisc_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1189, plain, (m1_subset_1(k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5),k1_numbers)|~v3_membered(k1_zfmisc_1(X4))|~m1_subset_1(k2_zfmisc_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1190, plain, (v1_xreal_0(k7_relat_1(k2_zfmisc_1(X3,X4),X5))|~v3_membered(k1_zfmisc_1(X4))|~m1_subset_1(k2_zfmisc_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6078, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~l1_pre_topc(k15_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0))))).#trainneg
cnf(i_0_6046, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|g1_pre_topc(u1_struct_0(k15_euclid(X3)),u1_pre_topc(k15_euclid(X3)))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~l1_pre_topc(k15_euclid(X3))|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(k15_euclid(X3))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_6049, plain, (r1_tarski(k7_relat_1(X3,k9_metric_1(k14_euclid(X1),X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2157, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1)|~m1_subset_1(k7_relset_1(X2,X3,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k9_metric_1(k8_metric_1,k1_funct_1(esk2_0,esk3_0),esk4_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1205, negated_conjecture, (m1_subset_1(X1,k1_numbers)|v1_xboole_0(k7_relat_1(esk2_0,X2))|~m1_subset_1(X1,k7_relat_1(esk2_0,X2)))).#trainneg
cnf(i_0_5473, plain, (v1_relat_1(u1_pre_topc(k8_metric_1))|~v1_relat_1(k1_zfmisc_1(u1_struct_0(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))))|~l1_pre_topc(g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1)))|~l1_pre_topc(k8_metric_1))).#trainneg
cnf(i_0_1221, plain, (m1_subset_1(k3_funct_2(X1,X2,X3,X4),k5_numbers)|v1_xboole_0(X1)|~v6_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_1222, plain, (m1_subset_1(k3_funct_2(X1,X2,X3,X4),k1_numbers)|v1_xboole_0(X1)|~v6_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_6074, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(k1_funct_1(esk2_0,esk3_0),k1_numbers)|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers))).#trainneg
cnf(i_0_6050, plain, (r1_tarski(k7_relset_1(X7,X8,X3,k9_metric_1(k14_euclid(X1),X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_6040, plain, (r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k14_euclid(X2)),X3,k1_topreal9(X1,X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k1_topreal9(X2,X5,X6))|X5!=k1_funct_1(X3,X4)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k14_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k14_euclid(X2)))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k15_euclid(X1)),u1_struct_0(k14_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xreal_0(esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_6146, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v7_ordinal1(k1_funct_1(esk2_0,esk3_0))|~l1_pre_topc(k15_euclid(esk1_0)))).#trainneg
cnf(i_0_2508, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,u1_pre_topc(k3_topmetr))|~l1_metric_1(X2)|k3_pcomps_1(X2)!=k3_topmetr)).#trainneg
cnf(i_0_4176, plain, (v7_ordinal1(k2_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_1223, plain, (v1_xboole_0(X1)|v7_ordinal1(k1_funct_1(X3,X4))|~v6_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_3425, plain, (v7_ordinal1(k7_relat_1(k7_metric_1,X1))|~v6_membered(k1_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5554, plain, (v1_finset_1(k1_numbers)|~v8_struct_0(k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))|~l1_struct_0(k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))|k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))!=k3_topmetr|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_5206, plain, (v1_xreal_0(u1_pre_topc(k3_topmetr))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))))|~l1_metric_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))!=k3_topmetr|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_5103, plain, (v1_relat_1(k3_funct_2(X1,k7_relset_1(X8,X9,X4,X5),X6,X7))|v1_xboole_0(k7_relset_1(X8,X9,X4,X5))|v1_xboole_0(X1)|~v1_relat_1(X3)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k7_relset_1(X8,X9,X4,X5))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k1_zfmisc_1(X3))))|~m1_subset_1(X7,X1)|~v1_funct_2(X6,X1,k7_relset_1(X8,X9,X4,X5))|~v1_funct_1(X6)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X8,X9))))).#trainneg
cnf(i_0_4795, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(X1),X2)|g1_pre_topc(u1_struct_0(k15_euclid(esk1_0)),u1_pre_topc(k15_euclid(esk1_0)))!=k3_pcomps_1(k14_euclid(X1))|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X2)|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X2)|~v2_pre_topc(k15_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5842, negated_conjecture, (v5_pre_topc(esk2_0,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),k3_topmetr)|v2_struct_0(X1)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),k1_numbers)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6298, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(X1),X2)|k3_pcomps_1(k14_euclid(esk1_0))!=k3_pcomps_1(k14_euclid(X1))|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X2)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X2)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6206, negated_conjecture, (v5_pre_topc(esk2_0,X1,k14_euclid(X2))|g1_pre_topc(u1_struct_0(k15_euclid(X3)),u1_pre_topc(k15_euclid(X3)))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|g1_pre_topc(u1_struct_0(k15_euclid(X2)),u1_pre_topc(k14_euclid(X2)))!=k3_topmetr|~l1_pre_topc(k14_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(k14_euclid(X2))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6213, negated_conjecture, (v5_pre_topc(esk2_0,X1,k15_euclid(X2))|g1_pre_topc(u1_struct_0(k15_euclid(X3)),u1_pre_topc(k15_euclid(X3)))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(k15_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(X2))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1277, plain, (m1_subset_1(k2_rcomp_1(X1,X2),k1_numbers)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6278, plain, (v1_finset_1(k1_numbers)|k3_pcomps_1(k3_topmetr)!=k3_topmetr|~v8_struct_0(k3_pcomps_1(k3_topmetr))|~l1_struct_0(k3_pcomps_1(k3_topmetr))|~l1_metric_1(k3_topmetr)|~v1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_1278, plain, (v1_xreal_0(k9_metric_1(k8_metric_1,X1,X2))|~v1_xxreal_0(k3_real_1(X1,X2))|~v1_xxreal_0(k5_real_1(X1,X2))|~v3_membered(k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1279, plain, (v1_xreal_0(k9_metric_1(k8_metric_1,X1,X2))|~v1_xxreal_0(k3_real_1(X2,X1))|~v1_xxreal_0(k5_real_1(X1,X2))|~v3_membered(k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6297, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(X1),X2)|k3_pcomps_1(k14_euclid(X1))!=k15_euclid(esk1_0)|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X2)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X2)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1)|~v1_pre_topc(k15_euclid(esk1_0)))).#trainneg
cnf(i_0_6238, plain, (v1_xboole_0(X1)|m1_subset_1(k1_funct_1(X3,X4),k1_numbers)|~v6_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_6231, plain, (v1_xreal_0(k3_funct_2(X1,X2,X3,X4))|v1_xboole_0(X1)|~v6_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_6212, negated_conjecture, (v5_pre_topc(esk2_0,X1,k3_topmetr)|g1_pre_topc(u1_struct_0(k15_euclid(X2)),u1_pre_topc(k15_euclid(X2)))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6214, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(k15_euclid(X3)),u1_pre_topc(k15_euclid(X3)))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v7_ordinal1(X3)|~v1_pre_topc(X2))).#trainneg
cnf(i_0_6210, negated_conjecture, (v5_pre_topc(esk2_0,X1,k3_pcomps_1(k14_euclid(X2)))|g1_pre_topc(u1_struct_0(k15_euclid(X3)),u1_pre_topc(k15_euclid(X3)))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|g1_pre_topc(u1_struct_0(k15_euclid(X2)),u1_pre_topc(k3_pcomps_1(k14_euclid(X2))))!=k3_topmetr|~l1_pre_topc(k3_pcomps_1(k14_euclid(X2)))|~l1_pre_topc(X1)|~v2_pre_topc(k3_pcomps_1(k14_euclid(X2)))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6211, negated_conjecture, (v5_pre_topc(esk2_0,X1,k3_pcomps_1(k14_euclid(X2)))|g1_pre_topc(u1_struct_0(k15_euclid(X3)),u1_pre_topc(k15_euclid(X3)))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|g1_pre_topc(u1_struct_0(k14_euclid(X2)),u1_pre_topc(k3_pcomps_1(k14_euclid(X2))))!=k3_topmetr|~l1_pre_topc(k3_pcomps_1(k14_euclid(X2)))|~l1_pre_topc(X1)|~v2_pre_topc(k3_pcomps_1(k14_euclid(X2)))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k14_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k14_euclid(X2)))|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4704, plain, (v2_struct_0(k14_euclid(X1))|r1_tarski(k7_relset_1(u1_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X2)),X3,k1_topreal9(X1,X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k9_metric_1(k14_euclid(X2),X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(k14_euclid(X1))|~v9_metric_1(k14_euclid(X1))|~v8_metric_1(k14_euclid(X1))|~v7_metric_1(k14_euclid(X1))|~v6_metric_1(k14_euclid(X1))|~v2_xxreal_0(X6)|~v1_xreal_0(esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X3,u1_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3340, plain, (k15_euclid(X1)=k15_euclid(X2)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(X2))|k10_funcsdom(k2_finseq_1(X1))!=k10_funcsdom(k2_finseq_1(X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_973, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(k7_relset_1(X2,X3,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1282, plain, (v3_membered(k7_relset_1(X1,X2,esk10_1(k1_zfmisc_1(k2_zfmisc_1(X3,X4))),X5))|~v3_membered(X4)|~m1_subset_1(esk10_1(k1_zfmisc_1(k2_zfmisc_1(X3,X4))),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6391, plain, (v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|m1_subset_1(k1_funct_1(k7_metric_1,X1),k1_numbers)|~v6_membered(k1_numbers)|~m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1284, plain, (v3_membered(k7_relset_1(X1,X2,u1_metric_1(X3),X4))|~m1_subset_1(u1_metric_1(X3),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~l1_metric_1(X3))).#trainneg
cnf(i_0_6406, negated_conjecture, (v5_pre_topc(esk2_0,k14_euclid(X1),k3_topmetr)|g1_pre_topc(u1_struct_0(k15_euclid(X2)),u1_pre_topc(k15_euclid(X2)))!=g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k14_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6410, negated_conjecture, (v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|g1_pre_topc(u1_struct_0(k15_euclid(X2)),u1_pre_topc(k15_euclid(X2)))!=g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k3_pcomps_1(k14_euclid(X1))))|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~v2_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6411, negated_conjecture, (v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|g1_pre_topc(u1_struct_0(k15_euclid(X2)),u1_pre_topc(k15_euclid(X2)))!=g1_pre_topc(u1_struct_0(k14_euclid(X1)),u1_pre_topc(k3_pcomps_1(k14_euclid(X1))))|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~v2_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(X1)),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6435, negated_conjecture, (v5_pre_topc(esk2_0,X1,k3_pcomps_1(k14_euclid(X2)))|g1_pre_topc(u1_struct_0(k15_euclid(X3)),u1_pre_topc(k15_euclid(X3)))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k14_euclid(X2)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k14_euclid(X2)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l1_metric_1(k14_euclid(X2)))).#trainneg
cnf(i_0_6348, negated_conjecture, (v5_pre_topc(esk2_0,X1,k15_euclid(X2))|g1_pre_topc(u1_struct_0(k15_euclid(X3)),u1_pre_topc(k15_euclid(X3)))!=X1|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v1_pre_topc(X1))).#trainneg
cnf(i_0_6414, negated_conjecture, (v5_pre_topc(esk2_0,X1,k3_topmetr)|k15_euclid(X2)!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers)|~v7_ordinal1(X2)|~v1_pre_topc(k15_euclid(X2))|~l1_pre_topc(k15_euclid(X2)))).#trainneg
cnf(i_0_6350, negated_conjecture, (v5_pre_topc(esk2_0,X1,k15_euclid(X2))|k3_pcomps_1(k14_euclid(X3))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6430, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(k15_euclid(X3)),u1_pre_topc(k15_euclid(X3)))!=X1|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v7_ordinal1(X3)|~v1_pre_topc(X1))).#trainneg
cnf(i_0_6258, plain, (m1_subset_1(k2_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))),k5_numbers)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_2545, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~v1_xreal_0(X1)|~m1_subset_1(k7_relset_1(X2,X3,esk2_0,k1_topreal9(esk1_0,esk3_0,X1)),k1_zfmisc_1(k2_rcomp_1(k6_xcmplx_0(k1_funct_1(esk2_0,esk3_0),esk4_0),k2_xcmplx_0(esk4_0,k1_funct_1(esk2_0,esk3_0)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1286, plain, (v3_membered(k7_relset_1(X1,X2,u1_rlvect_1(X3),X4))|~v3_membered(u1_struct_0(X3))|~m1_subset_1(u1_rlvect_1(X3),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~l1_rlvect_1(X3))).#trainneg
cnf(i_0_6274, plain, (m1_subset_1(k7_relat_1(k7_metric_1,X1),k5_numbers)|~v6_membered(k1_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5127, plain, (v1_xboole_0(k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5))|v1_xboole_0(X6)|v7_ordinal1(k3_funct_2(X6,k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5),X7,X8))|~v6_membered(X4)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X6,k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5))))|~m1_subset_1(k2_zfmisc_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X8,X6)|~v1_funct_2(X7,X6,k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5))|~v1_funct_1(X7))).#trainneg
cnf(i_0_5240, plain, (v1_xboole_0(k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5))|v1_xboole_0(X6)|v1_xreal_0(k3_funct_2(X6,k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5),X7,X8))|~v3_membered(X4)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X6,k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5))))|~m1_subset_1(k2_zfmisc_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X8,X6)|~v1_funct_2(X7,X6,k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5))|~v1_funct_1(X7))).#trainneg
cnf(i_0_1291, plain, (v1_relat_1(k9_metric_1(k8_metric_1,X1,X2))|~v1_xxreal_0(k3_real_1(X1,X2))|~v1_xxreal_0(k5_real_1(X1,X2))|~v1_relat_1(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1292, plain, (v1_relat_1(k9_metric_1(k8_metric_1,X1,X2))|~v1_xxreal_0(k3_real_1(X2,X1))|~v1_xxreal_0(k5_real_1(X1,X2))|~v1_relat_1(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6416, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(X1),X2)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~l1_pre_topc(X2)|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(X2)|~v2_pre_topc(k15_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6341, negated_conjecture, (v5_pre_topc(esk2_0,k14_euclid(X1),k15_euclid(X2))|g1_pre_topc(u1_struct_0(k15_euclid(X3)),u1_pre_topc(k15_euclid(X3)))!=g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k14_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6334, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(X1),k15_euclid(X2))|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(k15_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6440, plain, (r1_tarski(k7_relset_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)),X3,k1_topreal9(X1,X4,esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))),k9_metric_1(k14_euclid(X2),X5,X6))|X5!=k1_funct_1(X3,X4)|~v2_xxreal_0(X6)|~v1_xreal_0(esk6_6(k14_euclid(X1),k14_euclid(X2),X3,X4,X5,X6))|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X1)),k3_pcomps_1(k14_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(X5,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))|~v1_funct_2(X3,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6413, negated_conjecture, (v5_pre_topc(esk2_0,X1,k3_topmetr)|g1_pre_topc(u1_struct_0(k15_euclid(X2)),u1_pre_topc(k15_euclid(X2)))!=X1|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers)|~v7_ordinal1(X2)|~v1_pre_topc(X1))).#trainneg
cnf(i_0_6415, negated_conjecture, (v5_pre_topc(esk2_0,X1,k3_topmetr)|k3_pcomps_1(k14_euclid(X2))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6412, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|g1_pre_topc(u1_struct_0(k15_euclid(X2)),u1_pre_topc(k15_euclid(X2)))!=k3_pcomps_1(k14_euclid(X1))|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(k15_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6573, negated_conjecture, (v5_pre_topc(esk2_0,X1,k3_topmetr)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers))).#trainneg
cnf(i_0_3220, plain, (v3_membered(k7_relat_1(k2_zfmisc_1(X1,X2),X3))|~v3_membered(X2))).#trainneg
cnf(i_0_1293, plain, (m1_subset_1(X1,k1_numbers)|~v1_xxreal_0(k3_real_1(X2,X3))|~v1_xxreal_0(k5_real_1(X2,X3))|~r2_hidden(X1,k9_metric_1(k8_metric_1,X2,X3))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_6520, plain, (m1_subset_1(k2_pcomps_1(k3_topmetr),k5_numbers)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(k3_topmetr)|~v1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_6519, plain, (m1_subset_1(u1_pre_topc(k3_topmetr),k5_numbers)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(k3_topmetr)|k8_metric_1!=k3_topmetr|~v1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_2492, plain, (v1_xboole_0(k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5))|m1_subset_1(k3_funct_2(X6,k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5),X7,X8),X4)|v1_xboole_0(X6)|~m1_subset_1(k2_zfmisc_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X6,k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5))))|~m1_subset_1(X8,X6)|~v1_funct_2(X7,X6,k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5))|~v1_funct_1(X7))).#trainneg
cnf(i_0_1294, plain, (m1_subset_1(X1,k1_numbers)|~v1_xxreal_0(k3_real_1(X2,X3))|~v1_xxreal_0(k5_real_1(X3,X2))|~r2_hidden(X1,k9_metric_1(k8_metric_1,X3,X2))|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1295, plain, (m1_subset_1(X1,k1_numbers)|v1_xboole_0(k2_rcomp_1(X3,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~m1_subset_1(X1,k2_rcomp_1(X3,X2)))).#trainneg
cnf(i_0_6467, negated_conjecture, (v5_pre_topc(esk2_0,k14_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X2))!=g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k14_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6474, negated_conjecture, (v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|k3_pcomps_1(k14_euclid(X2))!=g1_pre_topc(u1_struct_0(k14_euclid(X1)),u1_pre_topc(k3_pcomps_1(k14_euclid(X1))))|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(X1)),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6469, negated_conjecture, (v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|k3_pcomps_1(k14_euclid(X2))!=g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k3_pcomps_1(k14_euclid(X1))))|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6575, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(X1))|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6510, negated_conjecture, (v5_pre_topc(esk2_0,X1,k3_topmetr)|k15_euclid(X2)!=X1|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|~v1_pre_topc(k15_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers)|~v7_ordinal1(X2)|~v1_pre_topc(X1))).#trainneg
cnf(i_0_6513, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|k15_euclid(X3)!=X1|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~v1_pre_topc(X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v7_ordinal1(X3)|~v1_pre_topc(k15_euclid(X3))|~l1_pre_topc(k15_euclid(X3)))).#trainneg
cnf(i_0_6563, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(k15_euclid(X3)),u1_pre_topc(k15_euclid(X3)))!=g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))|g1_pre_topc(u1_struct_0(k15_euclid(X4)),u1_pre_topc(k15_euclid(X4)))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X2)|~l1_pre_topc(k15_euclid(X3))|~l1_pre_topc(X1)|~l1_pre_topc(k15_euclid(X4))|~v2_pre_topc(X2)|~v2_pre_topc(k15_euclid(X3))|~v2_pre_topc(X1)|~v2_pre_topc(k15_euclid(X4))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X4)),u1_struct_0(k15_euclid(X3)))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X4)),u1_struct_0(k15_euclid(X3)))|k3_pcomps_1(k14_euclid(X4))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X3))!=k3_topmetr|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X4)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X4)),k1_numbers)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_6572, negated_conjecture, (v5_pre_topc(esk2_0,X1,k3_topmetr)|k3_pcomps_1(k14_euclid(X2))!=X1|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6511, negated_conjecture, (v5_pre_topc(esk2_0,X1,k15_euclid(X2))|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4424, plain, (u1_struct_0(X1)=u1_struct_0(X2)|X1!=k3_pcomps_1(g1_metric_1(u1_struct_0(X2),u1_metric_1(X2)))|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~l1_metric_1(X2))).#trainneg
cnf(i_0_3221, plain, (v1_relat_1(k7_relat_1(k2_zfmisc_1(X1,X2),X3))|~v1_relat_1(X2))).#trainneg
cnf(i_0_6516, plain, (v1_xreal_0(k2_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_6523, plain, (m1_subset_1(u1_pre_topc(k3_topmetr),k5_numbers)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(k3_topmetr)|k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))!=k3_topmetr)).#trainneg
cnf(i_0_1307, plain, (u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))=u1_struct_0(X1)|v2_struct_0(X1)|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6693, negated_conjecture, (v5_pre_topc(esk2_0,k3_topmetr,k15_euclid(X1))|k3_pcomps_1(k14_euclid(esk1_0))!=k3_topmetr|k3_pcomps_1(k14_euclid(X1))!=k3_topmetr|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(k15_euclid(X1)))))|~v1_funct_2(esk2_0,k1_numbers,u1_struct_0(k15_euclid(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6584, negated_conjecture, (v5_pre_topc(esk2_0,k14_euclid(X1),k3_topmetr)|g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k14_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6679, negated_conjecture, (v5_pre_topc(esk2_0,X1,k3_topmetr)|k3_pcomps_1(k14_euclid(esk1_0))!=X1|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),k1_numbers))).#trainneg
cnf(i_0_6677, negated_conjecture, (v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(X1))|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(X1)),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~l1_metric_1(k14_euclid(X1)))).#trainneg
cnf(i_0_1308, plain, (u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))=u1_struct_0(X1)|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1309, plain, (u1_struct_0(g1_pre_topc(X1,X2))=X1|~l1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_1323, plain, (m1_subset_1(k3_real_1(X1,X2),k5_numbers)|~v6_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6739, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(k15_euclid(X3)),u1_pre_topc(k15_euclid(X3)))!=g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_topmetr|~l1_pre_topc(X2)|~l1_pre_topc(k15_euclid(X3))|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(k15_euclid(X3))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(k15_euclid(X3)))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(esk2_0,k1_numbers,u1_struct_0(k15_euclid(X3)))|k3_pcomps_1(k14_euclid(esk1_0))!=k3_topmetr|k3_pcomps_1(k14_euclid(X3))!=k3_topmetr|~v7_ordinal1(X3))).#trainneg
cnf(i_0_6750, negated_conjecture, (v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|k3_pcomps_1(k14_euclid(esk1_0))!=k3_pcomps_1(k14_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6589, negated_conjecture, (v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|g1_pre_topc(u1_struct_0(k14_euclid(X1)),u1_pre_topc(k3_pcomps_1(k14_euclid(X1))))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~v2_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6717, plain, (v1_xreal_0(k2_pcomps_1(k3_topmetr))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(k3_topmetr)|~v1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_5780, plain, (k15_euclid(X1)=k8_metric_1|k10_funcsdom(k2_finseq_1(X1))!=k8_metric_1|g1_pre_topc(k1_numbers,u1_pre_topc(k8_metric_1))!=k3_pcomps_1(k14_euclid(X1))|~l1_rltopsp1(k8_metric_1)|~v5_rltopsp1(k8_metric_1)|~v7_ordinal1(X1)|~v1_rlvect_1(k8_metric_1)|~l1_rlvect_1(k8_metric_1))).#trainneg
cnf(i_0_3585, negated_conjecture, (v7_ordinal1(k7_relat_1(esk2_0,X1))|~v6_membered(k1_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6716, plain, (v1_xreal_0(u1_pre_topc(k3_topmetr))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(k3_topmetr)|k8_metric_1!=k3_topmetr|~v1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_6720, plain, (v1_xreal_0(u1_pre_topc(k3_topmetr))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(k3_topmetr)|k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))!=k3_topmetr)).#trainneg
cnf(i_0_3618, plain, (v3_valued_0(k3_funct_2(X1,k7_relset_1(X2,u1_struct_0(k15_euclid(X3)),X4,X5),X6,X7))|v1_xboole_0(k7_relset_1(X2,u1_struct_0(k15_euclid(X3)),X4,X5))|v1_xboole_0(X1)|~v7_ordinal1(X3)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k7_relset_1(X2,u1_struct_0(k15_euclid(X3)),X4,X5))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k15_euclid(X3)))))|~m1_subset_1(X7,X1)|~v1_funct_2(X6,X1,k7_relset_1(X2,u1_struct_0(k15_euclid(X3)),X4,X5))|~v1_funct_1(X6))).#trainneg
cnf(i_0_6694, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(X1),k15_euclid(X2))|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(k15_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6743, negated_conjecture, (v5_pre_topc(esk2_0,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),k3_topmetr)|v2_struct_0(X1)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),k1_numbers)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6740, negated_conjecture, (v5_pre_topc(esk2_0,k14_euclid(X1),k3_topmetr)|k14_euclid(X1)!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k14_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1)|~v1_pre_topc(k14_euclid(X1)))).#trainneg
cnf(i_0_6744, negated_conjecture, (v5_pre_topc(esk2_0,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),k3_topmetr)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),k1_numbers)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_6588, negated_conjecture, (v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k3_pcomps_1(k14_euclid(X1))))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~v2_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6687, negated_conjecture, (v5_pre_topc(esk2_0,k14_euclid(X1),k15_euclid(X2))|g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k14_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6508, negated_conjecture, (v5_pre_topc(esk2_0,k3_topmetr,k3_topmetr)|k15_euclid(X1)!=k3_topmetr|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~v1_pre_topc(k15_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v1_funct_2(esk2_0,k1_numbers,k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6472, negated_conjecture, (v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|g1_pre_topc(u1_struct_0(k15_euclid(X2)),u1_pre_topc(k15_euclid(X2)))!=k3_pcomps_1(k14_euclid(X1))|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(X1)),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~l1_metric_1(k14_euclid(X1)))).#trainneg
cnf(i_0_6509, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(X1),k3_topmetr)|k15_euclid(X2)!=k3_pcomps_1(k14_euclid(X1))|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|~v1_pre_topc(k15_euclid(X2))|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(k15_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6493, negated_conjecture, (v5_pre_topc(esk2_0,X1,k15_euclid(X2))|k15_euclid(X3)!=X1|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v1_pre_topc(k15_euclid(X3))|~l1_pre_topc(k15_euclid(X3)))).#trainneg
cnf(i_0_3088, plain, (m1_subset_1(k7_relat_1(u1_rlvect_1(k8_metric_1),X1),k1_zfmisc_1(k1_numbers))|~l1_rlvect_1(k8_metric_1))).#trainneg
cnf(i_0_6722, plain, (m1_subset_1(u1_pre_topc(k3_topmetr),k5_numbers)|k3_pcomps_1(k3_topmetr)!=k3_topmetr|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(k3_topmetr)|~v1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_3619, plain, (v1_finseq_1(k3_funct_2(X1,k7_relset_1(X2,u1_struct_0(k15_euclid(X3)),X4,X5),X6,X7))|v1_xboole_0(k7_relset_1(X2,u1_struct_0(k15_euclid(X3)),X4,X5))|v1_xboole_0(X1)|~v7_ordinal1(X3)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k7_relset_1(X2,u1_struct_0(k15_euclid(X3)),X4,X5))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k15_euclid(X3)))))|~m1_subset_1(X7,X1)|~v1_funct_2(X6,X1,k7_relset_1(X2,u1_struct_0(k15_euclid(X3)),X4,X5))|~v1_funct_1(X6))).#trainneg
cnf(i_0_2692, plain, (v1_xboole_0(k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|v1_xboole_0(X1)|m1_subset_1(k3_funct_2(X5,k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4),X6,X7),X2)|v1_xboole_0(X5)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))))|~m1_subset_1(X7,X5)|~v1_funct_2(X6,X5,k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|~v1_funct_1(X6))).#trainneg
cnf(i_0_1332, plain, (v1_relat_1(k7_relset_1(X1,X2,esk10_1(k1_zfmisc_1(k2_zfmisc_1(X3,X4))),X5))|~v1_relat_1(X4)|~m1_subset_1(esk10_1(k1_zfmisc_1(k2_zfmisc_1(X3,X4))),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1333, negated_conjecture, (v1_relat_1(k7_relset_1(X1,X2,esk2_0,X3))|~v1_relat_1(k1_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6889, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),k1_numbers)|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers))).#trainneg
cnf(i_0_6887, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers))).#trainneg
cnf(i_0_6890, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers))).#trainneg
cnf(i_0_6888, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v7_ordinal1(k1_funct_1(esk2_0,esk3_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers))).#trainneg
cnf(i_0_6907, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(k15_euclid(X3)),u1_pre_topc(k15_euclid(X3)))!=g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))|g1_pre_topc(u1_struct_0(k15_euclid(X4)),u1_pre_topc(k15_euclid(X4)))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X2)|~l1_pre_topc(k15_euclid(X3))|~l1_pre_topc(X1)|~l1_pre_topc(k15_euclid(X4))|~v2_pre_topc(X2)|~v2_pre_topc(k15_euclid(X3))|~v2_pre_topc(X1)|~v2_pre_topc(k15_euclid(X4))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X4)),u1_struct_0(k15_euclid(X3)))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X4)),u1_struct_0(k15_euclid(X3)))|k3_pcomps_1(k14_euclid(X4))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X3))!=k3_topmetr|~v7_ordinal1(X3)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_6695, negated_conjecture, (v5_pre_topc(esk2_0,X1,k15_euclid(X2))|X1!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2)|~v1_pre_topc(X1))).#trainneg
cnf(i_0_1334, plain, (v1_relat_1(k7_relset_1(X1,X2,u1_metric_1(X3),X4))|~v1_relat_1(k1_numbers)|~m1_subset_1(u1_metric_1(X3),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~l1_metric_1(X3))).#trainneg
cnf(i_0_6902, plain, (v1_xreal_0(u1_pre_topc(k3_topmetr))|k3_pcomps_1(k3_topmetr)!=k3_topmetr|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(k3_topmetr)|~v1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_6904, plain, (v3_valued_0(k3_funct_2(X1,k7_relset_1(X8,X9,X4,X5),X6,X7))|v1_xboole_0(k7_relset_1(X8,X9,X4,X5))|v1_xboole_0(X1)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k7_relset_1(X8,X9,X4,X5))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k15_euclid(X3)))))|~m1_subset_1(X7,X1)|~v1_funct_2(X6,X1,k7_relset_1(X8,X9,X4,X5))|~v1_funct_1(X6)|~v7_ordinal1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X8,X9))))).#trainneg
cnf(i_0_1336, plain, (v1_relat_1(k7_relset_1(X1,X2,u1_rlvect_1(X3),X4))|~v1_relat_1(u1_struct_0(X3))|~m1_subset_1(u1_rlvect_1(X3),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~l1_rlvect_1(X3))).#trainneg
cnf(i_0_6945, negated_conjecture, (v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|k15_euclid(X2)!=k3_pcomps_1(k14_euclid(X1))|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(X1)),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_pre_topc(k15_euclid(X2))|~l1_pre_topc(k15_euclid(X2)))).#trainneg
cnf(i_0_6514, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|k3_pcomps_1(k14_euclid(X3))!=X1|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~v1_pre_topc(X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_6923, negated_conjecture, (v2_struct_0(k15_euclid(X1))|v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(X1)),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(X1))),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(X1))),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6562, negated_conjecture, (v5_pre_topc(esk2_0,k14_euclid(X1),k15_euclid(X2))|k3_pcomps_1(k14_euclid(X3))!=g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k14_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6502, negated_conjecture, (v5_pre_topc(esk2_0,k14_euclid(X1),k3_topmetr)|k15_euclid(X2)!=g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|~v1_pre_topc(k15_euclid(X2))|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k14_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6349, negated_conjecture, (v5_pre_topc(esk2_0,X1,k15_euclid(X2))|k15_euclid(X3)!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v1_pre_topc(k15_euclid(X3))|~l1_pre_topc(k15_euclid(X3)))).#trainneg
cnf(i_0_6465, negated_conjecture, (v5_pre_topc(esk2_0,k14_euclid(X1),k3_topmetr)|g1_pre_topc(u1_struct_0(k15_euclid(X2)),u1_pre_topc(k15_euclid(X2)))!=k14_euclid(X1)|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k14_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_pre_topc(k14_euclid(X1)))).#trainneg
cnf(i_0_6259, plain, (m1_subset_1(k2_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr))),k1_numbers)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_4608, plain, (u1_struct_0(X1)=k1_numbers|k3_pcomps_1(X1)!=k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|~l1_metric_1(X1)|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_1351, plain, (m1_subset_1(k9_metric_1(k8_metric_1,X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(k3_real_1(X1,X2))|~v1_xxreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6955, plain, (v1_finseq_1(k3_funct_2(X1,k7_relset_1(X8,X9,X4,X5),X6,X7))|v1_xboole_0(k7_relset_1(X8,X9,X4,X5))|v1_xboole_0(X1)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k7_relset_1(X8,X9,X4,X5))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k15_euclid(X3)))))|~m1_subset_1(X7,X1)|~v1_funct_2(X6,X1,k7_relset_1(X8,X9,X4,X5))|~v1_funct_1(X6)|~v7_ordinal1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X8,X9))))).#trainneg
cnf(i_0_1352, plain, (v1_xreal_0(k9_metric_1(k8_metric_1,X1,X2))|~v1_xxreal_0(k3_real_1(X1,X2))|~v1_xxreal_0(k6_xcmplx_0(X1,X2))|~v3_membered(k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1353, plain, (v3_membered(k9_metric_1(k8_metric_1,X1,X2))|~v1_xxreal_0(k3_real_1(X1,X2))|~v1_xxreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6494, negated_conjecture, (v5_pre_topc(esk2_0,X1,k15_euclid(X2))|k3_pcomps_1(k14_euclid(X3))!=X1|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6491, negated_conjecture, (v5_pre_topc(esk2_0,X1,k3_pcomps_1(k14_euclid(X2)))|k3_pcomps_1(k14_euclid(X3))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k14_euclid(X2)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k14_euclid(X2)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6432, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|k3_pcomps_1(k14_euclid(X3))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_6333, negated_conjecture, (v5_pre_topc(esk2_0,X1,k14_euclid(X2))|g1_pre_topc(u1_struct_0(k15_euclid(X3)),u1_pre_topc(k15_euclid(X3)))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|k14_euclid(X2)!=k3_topmetr|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|~l1_pre_topc(k14_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(k14_euclid(X2))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v1_pre_topc(k14_euclid(X2)))).#trainneg
cnf(i_0_7134, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2)))).#trainneg
cnf(i_0_7080, negated_conjecture, (v5_pre_topc(esk2_0,k14_euclid(X1),k3_topmetr)|k15_euclid(X2)!=k14_euclid(X1)|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|~v1_pre_topc(k15_euclid(X2))|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k14_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_pre_topc(k14_euclid(X1)))).#trainneg
cnf(i_0_7077, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0),k1_numbers)|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers))).#trainneg
cnf(i_0_2505, plain, (v1_xreal_0(u1_pre_topc(k3_topmetr))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_metric_1(X1)|k3_pcomps_1(X1)!=k3_topmetr)).#trainneg
cnf(i_0_3093, plain, (m1_subset_1(k7_relat_1(u1_rlvect_1(k3_topmetr),X1),k1_zfmisc_1(k1_numbers))|~l1_rlvect_1(k3_topmetr))).#trainneg
cnf(i_0_1354, plain, (v1_relat_1(k9_metric_1(k8_metric_1,X1,X2))|~v1_xxreal_0(k3_real_1(X1,X2))|~v1_xxreal_0(k6_xcmplx_0(X1,X2))|~v1_relat_1(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_7103, plain, (m1_subset_1(k2_pcomps_1(k3_topmetr),k1_numbers)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(k3_topmetr)|~v1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_7102, plain, (m1_subset_1(u1_pre_topc(k3_topmetr),k1_numbers)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(k3_topmetr)|k8_metric_1!=k3_topmetr|~v1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_7013, plain, (v7_ordinal1(k3_funct_2(X1,k3_funct_2(X2,k1_zfmisc_1(X3),X4,X5),X6,X7))|v1_xboole_0(k3_funct_2(X2,k1_zfmisc_1(X3),X4,X5))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v6_membered(X3)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k3_funct_2(X2,k1_zfmisc_1(X3),X4,X5))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k1_zfmisc_1(X3))))|~m1_subset_1(X5,X2)|~m1_subset_1(X7,X1)|~v1_funct_2(X6,X1,k3_funct_2(X2,k1_zfmisc_1(X3),X4,X5))|~v1_funct_2(X4,X2,k1_zfmisc_1(X3))|~v1_funct_1(X4)|~v1_funct_1(X6))).#trainneg
cnf(i_0_7014, plain, (v1_xreal_0(k3_funct_2(X1,k3_funct_2(X2,k1_zfmisc_1(X3),X4,X5),X6,X7))|v1_xboole_0(k3_funct_2(X2,k1_zfmisc_1(X3),X4,X5))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v3_membered(X3)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k3_funct_2(X2,k1_zfmisc_1(X3),X4,X5))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k1_zfmisc_1(X3))))|~m1_subset_1(X5,X2)|~m1_subset_1(X7,X1)|~v1_funct_2(X6,X1,k3_funct_2(X2,k1_zfmisc_1(X3),X4,X5))|~v1_funct_2(X4,X2,k1_zfmisc_1(X3))|~v1_funct_1(X4)|~v1_funct_1(X6))).#trainneg
cnf(i_0_1355, plain, (m1_subset_1(X1,k1_numbers)|~v1_xxreal_0(k3_real_1(X2,X3))|~v1_xxreal_0(k6_xcmplx_0(X2,X3))|~r2_hidden(X1,k9_metric_1(k8_metric_1,X2,X3))|~v1_xreal_0(X2)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_7095, negated_conjecture, (v5_pre_topc(esk2_0,k3_topmetr,k15_euclid(X1))|k15_euclid(X2)!=k3_topmetr|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X1))!=k3_topmetr|~v1_pre_topc(k15_euclid(X2))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(k15_euclid(X1)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~v1_funct_2(esk2_0,k1_numbers,u1_struct_0(k15_euclid(X1)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7032, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers))).#trainneg
cnf(i_0_7029, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers))).#trainneg
cnf(i_0_7075, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers))).#trainneg
cnf(i_0_7099, negated_conjecture, (v5_pre_topc(esk2_0,k14_euclid(X1),k3_topmetr)|k3_pcomps_1(k14_euclid(X2))!=k14_euclid(X1)|k3_pcomps_1(k14_euclid(X2))!=k3_pcomps_1(k14_euclid(esk1_0))|~v1_pre_topc(k14_euclid(X1))|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k14_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X2)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7073, negated_conjecture, (v5_pre_topc(esk2_0,X1,X2)|k3_pcomps_1(k14_euclid(esk1_0))!=X1|X2!=k3_topmetr|~v1_pre_topc(X2)|~v1_pre_topc(X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(X2)))).#trainneg
cnf(i_0_7078, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers))).#trainneg
cnf(i_0_7133, negated_conjecture, (v5_pre_topc(esk2_0,X1,k3_pcomps_1(k14_euclid(X2)))|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k14_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k14_euclid(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_7035, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers)|~v7_ordinal1(k1_funct_1(esk2_0,esk3_0)))).#trainneg
cnf(i_0_3321, plain, (m1_subset_1(k2_finseq_1(X1),k5_numbers)|~v6_membered(k1_numbers)|~v3_membered(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7106, plain, (m1_subset_1(u1_pre_topc(k3_topmetr),k1_numbers)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(k3_topmetr)|k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))!=k3_topmetr)).#trainneg
cnf(i_0_2668, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relset_1(X3,X4,k7_relset_1(X5,k2_zfmisc_1(X6,X2),X7,X8),X9))|~m1_subset_1(k7_relset_1(X5,k2_zfmisc_1(X6,X2),X7,X8),k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,k2_zfmisc_1(X6,X2)))))).#trainneg
cnf(i_0_1361, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v1_xreal_0(k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_7076, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)|~v7_ordinal1(k1_funct_1(esk2_0,esk3_0)))).#trainneg
cnf(i_0_7152, negated_conjecture, (v5_pre_topc(esk2_0,X1,k14_euclid(X2))|k3_pcomps_1(k14_euclid(X3))!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|k14_euclid(X2)!=k3_topmetr|~v1_pre_topc(k14_euclid(X2))|~l1_pre_topc(k14_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(k14_euclid(X2))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6378, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(X1),k14_euclid(X2))|g1_pre_topc(u1_struct_0(k15_euclid(X2)),u1_pre_topc(k14_euclid(X2)))!=k3_topmetr|k3_pcomps_1(k14_euclid(X1))!=k15_euclid(esk1_0)|~v1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k14_euclid(X2))|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(k14_euclid(X2))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5701, negated_conjecture, (v5_pre_topc(esk2_0,X1,k14_euclid(X2))|k14_euclid(X2)!=k3_topmetr|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k15_euclid(esk1_0)|~v1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k14_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(k14_euclid(X2))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2)|~v1_pre_topc(k14_euclid(X2)))).#trainneg
cnf(i_0_7096, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(X1),k15_euclid(X2))|k15_euclid(X3)!=k3_pcomps_1(k14_euclid(X1))|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~v1_pre_topc(k15_euclid(X3))|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(k15_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5454, plain, (v2_struct_0(X1)|r1_tarski(k7_relat_1(X3,k9_metric_1(k14_euclid(X2),X4,esk6_6(k14_euclid(X2),X1,X3,X4,X5,X6))),k9_metric_1(X1,X5,X6))|X5!=k1_funct_1(X3,X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_xxreal_0(X6)|~v1_xreal_0(X6)|~v5_pre_topc(X3,k3_pcomps_1(k14_euclid(X2)),k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X2)),u1_struct_0(k3_pcomps_1(X1)))))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X2)))|~m1_subset_1(X5,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(k15_euclid(X2)),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_7222, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|k3_funct_2(u1_struct_0(k15_euclid(esk1_0)),k1_numbers,esk2_0,esk3_0)!=k1_funct_1(esk2_0,esk3_0)|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0))))).#trainneg
cnf(i_0_7208, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(esk1_0)),k1_numbers)))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(esk1_0)),k1_numbers))).#trainneg
cnf(i_0_1362, plain, (v1_xboole_0(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|v1_xreal_0(k1_funct_1(u1_metric_1(X1),X2))|~m1_subset_1(u1_metric_1(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)))|~m1_subset_1(X2,k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~v1_funct_1(u1_metric_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_7247, plain, (m1_subset_1(u1_pre_topc(k3_topmetr),k1_numbers)|k3_pcomps_1(k3_topmetr)!=k3_topmetr|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)))|~l1_metric_1(k3_topmetr)|~v1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_7249, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relset_1(X3,X4,k7_relset_1(X10,X11,X7,X8),X9))|~m1_subset_1(k7_relset_1(X10,X11,X7,X8),k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,k2_zfmisc_1(X6,X2))))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X10,X11))))).#trainneg
cnf(i_0_1363, plain, (v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|v1_xreal_0(k1_funct_1(k7_metric_1,X1))|~m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1364, plain, (v1_xboole_0(k2_zfmisc_1(k1_numbers,u1_struct_0(X1)))|v1_xreal_0(k1_funct_1(u1_rlvect_1(X1),X2))|~v3_membered(u1_struct_0(X1))|~m1_subset_1(u1_rlvect_1(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(X1)),u1_struct_0(X1))))|~m1_subset_1(X2,k2_zfmisc_1(k1_numbers,u1_struct_0(X1)))|~v1_funct_1(u1_rlvect_1(X1))|~l1_rlvect_1(X1))).#trainneg
cnf(i_0_7210, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0))))).#trainneg
cnf(i_0_7168, negated_conjecture, (v5_pre_topc(esk2_0,k3_topmetr,X1)|k3_pcomps_1(k14_euclid(esk1_0))!=k3_topmetr|X1!=k3_topmetr|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k1_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_7169, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(X1),X2)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~l1_pre_topc(X2)|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(X2)|~v2_pre_topc(k15_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7255, negated_conjecture, (m1_subset_1(k1_funct_1(esk2_0,X1),k1_numbers)|v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(esk1_0))))).#trainneg
cnf(i_0_7259, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|v2_struct_0(k14_euclid(esk1_0))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v7_ordinal1(k1_funct_1(esk2_0,esk3_0)))).#trainneg
cnf(i_0_7238, negated_conjecture, (v5_pre_topc(esk2_0,k3_topmetr,k3_pcomps_1(k14_euclid(X1)))|k3_pcomps_1(k14_euclid(esk1_0))!=k3_topmetr|k3_pcomps_1(k14_euclid(X1))!=k3_topmetr|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(k14_euclid(X1)))))|~v1_funct_2(esk2_0,k1_numbers,u1_struct_0(k14_euclid(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7171, negated_conjecture, (v2_struct_0(k14_euclid(esk1_0))|v1_xboole_0(u1_struct_0(k15_euclid(esk1_0)))|~l1_metric_1(k14_euclid(esk1_0))|~v9_metric_1(k14_euclid(esk1_0))|~v8_metric_1(k14_euclid(esk1_0))|~v7_metric_1(k14_euclid(esk1_0))|~v6_metric_1(k14_euclid(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers)))|~m1_subset_1(k1_funct_1(esk2_0,esk3_0),k1_numbers)|~m1_subset_1(esk3_0,u1_struct_0(k14_euclid(esk1_0)))|~v1_funct_2(esk2_0,u1_struct_0(k3_pcomps_1(k14_euclid(esk1_0))),k1_numbers))).#trainneg
cnf(i_0_4595, plain, (u1_struct_0(k15_euclid(X1))=k1_numbers|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(g1_metric_1(k1_numbers,u1_metric_1(k3_topmetr)))|~l1_metric_1(k3_topmetr)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6898, negated_conjecture, (m1_subset_1(k7_relat_1(esk2_0,X1),k5_numbers)|~v6_membered(k1_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7261, negated_conjecture, (v5_pre_topc(esk2_0,X1,k14_euclid(X2))|g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|k14_euclid(X2)!=k3_topmetr|~v1_pre_topc(k14_euclid(X2))|~l1_pre_topc(k14_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(k14_euclid(X2))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_7162, negated_conjecture, (v5_pre_topc(esk2_0,k14_euclid(X1),X2)|g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))!=k3_pcomps_1(k14_euclid(esk1_0))|X2!=k3_topmetr|~v1_pre_topc(X2)|~l1_pre_topc(X2)|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(X2)|~v2_pre_topc(k14_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7240, negated_conjecture, (v5_pre_topc(esk2_0,X1,k3_pcomps_1(k14_euclid(X2)))|X1!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k14_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k14_euclid(X2)))|~v7_ordinal1(X2)|~v1_pre_topc(X1))).#trainneg
cnf(i_0_7262, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(X1),k14_euclid(X2))|k14_euclid(X2)!=k3_topmetr|k3_pcomps_1(k14_euclid(X1))!=k15_euclid(esk1_0)|~v1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(k14_euclid(X2))|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(k14_euclid(X2))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_pre_topc(k14_euclid(X2)))).#trainneg
cnf(i_0_6570, negated_conjecture, (v5_pre_topc(esk2_0,g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k15_euclid(X1))),k3_topmetr)|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|~v1_pre_topc(g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k15_euclid(X1))))|~l1_pre_topc(g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k15_euclid(X1))))|~v2_pre_topc(g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k15_euclid(X1))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k15_euclid(X1)))),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v1_funct_2(esk2_0,u1_struct_0(g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k15_euclid(X1)))),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6345, negated_conjecture, (v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(X1)),k15_euclid(X2))|g1_pre_topc(u1_struct_0(k15_euclid(X3)),u1_pre_topc(k15_euclid(X3)))!=g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k3_pcomps_1(k14_euclid(X1))))|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~v2_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1384, plain, (m1_subset_1(k9_metric_1(k8_metric_1,X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(k3_real_1(X1,X2))|~v1_xxreal_0(k5_real_1(X1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1385, plain, (v1_xreal_0(k9_metric_1(k8_metric_1,X1,X2))|~v1_xxreal_0(k3_real_1(X1,X2))|~v1_xxreal_0(k5_real_1(X1,X2))|~v3_membered(k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1386, plain, (v3_membered(k9_metric_1(k8_metric_1,X1,X2))|~v1_xxreal_0(k3_real_1(X1,X2))|~v1_xxreal_0(k5_real_1(X1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4792, negated_conjecture, (v5_pre_topc(esk2_0,k14_euclid(X1),X2)|g1_pre_topc(u1_struct_0(k15_euclid(X1)),u1_pre_topc(k14_euclid(X1)))!=g1_pre_topc(u1_struct_0(k15_euclid(esk1_0)),u1_pre_topc(k15_euclid(esk1_0)))|g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))!=k3_topmetr|~l1_pre_topc(k15_euclid(esk1_0))|~l1_pre_topc(X2)|~l1_pre_topc(k14_euclid(X1))|~v2_pre_topc(k15_euclid(esk1_0))|~v2_pre_topc(X2)|~v2_pre_topc(k14_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6346, negated_conjecture, (v5_pre_topc(esk2_0,k3_pcomps_1(k14_euclid(X1)),k15_euclid(X2))|g1_pre_topc(u1_struct_0(k15_euclid(X3)),u1_pre_topc(k15_euclid(X3)))!=g1_pre_topc(u1_struct_0(k14_euclid(X1)),u1_pre_topc(k3_pcomps_1(k14_euclid(X1))))|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~v2_pre_topc(k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k14_euclid(X1)),u1_struct_0(k15_euclid(X2)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6475, negated_conjecture, (v5_pre_topc(esk2_0,k15_euclid(X1),k3_pcomps_1(k14_euclid(X2)))|k3_pcomps_1(k14_euclid(X1))!=k3_pcomps_1(k14_euclid(esk1_0))|k3_pcomps_1(k14_euclid(X2))!=k3_topmetr|~l1_pre_topc(k15_euclid(X1))|~v2_pre_topc(k15_euclid(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),u1_struct_0(k14_euclid(X2)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X1)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),u1_struct_0(k14_euclid(X2)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X1)),k1_numbers)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5436, plain, (r1_tarski(k7_relset_1(k1_numbers,u1_struct_0(k15_euclid(X1)),X2,k9_metric_1(k8_metric_1,X3,esk6_6(k8_metric_1,k14_euclid(X1),X2,X3,X4,X5))),k9_metric_1(k14_euclid(X1),X4,X5))|X4!=k1_funct_1(X2,X3)|~v2_xxreal_0(X5)|~v1_xreal_0(X5)|~v5_pre_topc(X2,k3_topmetr,k3_pcomps_1(k14_euclid(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,u1_struct_0(k15_euclid(X1)))))|~m1_subset_1(X4,u1_struct_0(k14_euclid(X1)))|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X2,k1_numbers,u1_struct_0(k15_euclid(X1)))|~v1_funct_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7151, negated_conjecture, (v5_pre_topc(esk2_0,X1,k14_euclid(X2))|k15_euclid(X3)!=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|k3_pcomps_1(k14_euclid(X3))!=k3_pcomps_1(k14_euclid(esk1_0))|k14_euclid(X2)!=k3_topmetr|~v1_pre_topc(k14_euclid(X2))|~l1_pre_topc(k14_euclid(X2))|~l1_pre_topc(X1)|~v2_pre_topc(k14_euclid(X2))|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k15_euclid(X3)),k1_numbers)))|~v1_funct_2(esk2_0,u1_struct_0(X1),u1_struct_0(k15_euclid(X2)))|~v1_funct_2(esk2_0,u1_struct_0(k15_euclid(X3)),k1_numbers)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v1_pre_topc(k15_euclid(X3))|~l1_pre_topc(k15_euclid(X3)))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 101
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 173
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 173
# Processed clauses                    : 2722
# ...of these trivial                  : 22
# ...subsumed                          : 1222
# ...remaining for further processing  : 1478
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 192
# Backward-rewritten                   : 148
# Generated clauses                    : 7257
# ...of the previous two non-trivial   : 6984
# Contextual simplify-reflections      : 457
# Paramodulations                      : 7193
# Factorizations                       : 0
# Equation resolutions                 : 64
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1136
#    Positive orientable unit clauses  : 56
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 1076
# Current number of unprocessed clauses: 4161
# ...number of literals in the above   : 37890
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1818
# Proof object given clauses           : 184
# Proof search given clauses           : 1478
# Clause-clause subsumption calls (NU) : 2113492
# Rec. Clause-clause subsumption calls : 73949
# Non-unit clause-clause subsumptions  : 1851
# Unit Clause-clause subsumption calls : 4129
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 35
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 345806

# -------------------------------------------------
# User time                : 22.307 s
# System time              : 0.813 s
# Total time               : 23.119 s
# Maximum resident set size: 362064 pages
22.30user 0.86system 0:23.47elapsed 98%CPU (0avgtext+0avgdata 362676maxresident)k
0inputs+1288outputs (0major+206123minor)pagefaults 0swaps
