# ENIGMA: Connected to the TF server '127.0.0.1:8888'.
# Preprocessing time       : 0.030 s
#TF#SERVER: Sending data to the TF server (conj=15, context=92, query=96, total=203).
#TF#SERVER: Sent 145665 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (756 bytes expected).
#TF#SERVER: Received 756 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=139, query=130, total=284).
#TF#SERVER: Sent 250530 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1080 bytes expected).
#TF#SERVER: Received 1080 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=176, query=128, total=319).
#TF#SERVER: Sent 227739 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1220 bytes expected).
#TF#SERVER: Received 1220 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=220, query=137, total=372).
#TF#SERVER: Sent 275179 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1432 bytes expected).
#TF#SERVER: Received 1432 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=244, query=129, total=388).
#TF#SERVER: Sent 265209 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1496 bytes expected).
#TF#SERVER: Received 1496 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=270, query=131, total=416).
#TF#SERVER: Sent 310717 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1608 bytes expected).
#TF#SERVER: Received 1608 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=297, query=153, total=465).
#TF#SERVER: Sent 397248 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1804 bytes expected).
#TF#SERVER: Received 1804 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=313, query=128, total=456).
#TF#SERVER: Sent 326887 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1768 bytes expected).
#TF#SERVER: Received 1768 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=336, query=148, total=499).
#TF#SERVER: Sent 392942 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1940 bytes expected).
#TF#SERVER: Received 1940 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=358, query=131, total=504).
#TF#SERVER: Sent 394464 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1960 bytes expected).
#TF#SERVER: Received 1960 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=381, query=131, total=527).
#TF#SERVER: Sent 456228 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2052 bytes expected).
#TF#SERVER: Received 2052 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=398, query=133, total=546).
#TF#SERVER: Sent 418208 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2128 bytes expected).
#TF#SERVER: Received 2128 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=426, query=128, total=569).
#TF#SERVER: Sent 504096 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2220 bytes expected).
#TF#SERVER: Received 2220 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=446, query=144, total=605).
#TF#SERVER: Sent 586516 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2364 bytes expected).
#TF#SERVER: Received 2364 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=465, query=134, total=614).
#TF#SERVER: Sent 605655 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2400 bytes expected).
#TF#SERVER: Received 2400 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=485, query=141, total=641).
#TF#SERVER: Sent 511976 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2508 bytes expected).
#TF#SERVER: Received 2508 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=493, query=136, total=644).
#TF#SERVER: Sent 507038 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2520 bytes expected).
#TF#SERVER: Received 2520 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=501, query=129, total=645).
#TF#SERVER: Sent 477899 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2524 bytes expected).
#TF#SERVER: Received 2524 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=510, query=136, total=661).
#TF#SERVER: Sent 513488 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2588 bytes expected).
#TF#SERVER: Received 2588 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=512, query=135, total=662).
#TF#SERVER: Sent 487962 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2592 bytes expected).
#TF#SERVER: Received 2592 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=512, query=135, total=662).
#TF#SERVER: Sent 481824 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2592 bytes expected).
#TF#SERVER: Received 2592 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=512, query=141, total=668).
#TF#SERVER: Sent 604411 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2616 bytes expected).
#TF#SERVER: Received 2616 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=512, query=132, total=659).
#TF#SERVER: Sent 519955 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2580 bytes expected).
#TF#SERVER: Received 2580 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=512, query=134, total=661).
#TF#SERVER: Sent 459862 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2588 bytes expected).
#TF#SERVER: Received 2588 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=512, query=128, total=655).
#TF#SERVER: Sent 581667 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=512, query=107, total=634).
#TF#SERVER: Sent 490536 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2480 bytes expected).
#TF#SERVER: Received 2480 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=512, query=128, total=655).
#TF#SERVER: Sent 574213 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=512, query=110, total=637).
#TF#SERVER: Sent 507975 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2492 bytes expected).
#TF#SERVER: Received 2492 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=512, query=128, total=655).
#TF#SERVER: Sent 466346 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=512, query=128, total=655).
#TF#SERVER: Sent 587130 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=15, context=512, query=18, total=545).
#TF#SERVER: Sent 414854 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2124 bytes expected).
#TF#SERVER: Received 2124 bytes from the TF server.

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t43_tops_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>((k2_struct_0(X2)=k1_xboole_0=>k2_struct_0(X1)=k1_xboole_0)=>(v5_pre_topc(X3,X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>(v3_pre_topc(X4,X2)=>v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X1)))))))), file('/scratch/mptp/predsm40/t25_topgrp_1__0', t43_tops_2)).
fof(t54_tops_2, axiom, ![X1]:(l1_struct_0(X1)=>![X2]:(l1_struct_0(X2)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>((k2_relset_1(u1_struct_0(X2),X3)=k2_struct_0(X2)&v2_funct_1(X3))=>k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4)=k8_relset_1(u1_struct_0(X2),u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X4)))))), file('/scratch/mptp/predsm40/t25_topgrp_1__0', t54_tops_2)).
fof(t25_topgrp_1, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&l1_pre_topc(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(v3_tops_2(X3,X1,X2)<=>(k1_relset_1(u1_struct_0(X1),X3)=k2_struct_0(X1)&(k2_relset_1(u1_struct_0(X2),X3)=k2_struct_0(X2)&(v2_funct_1(X3)&![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X4,X1)<=>v3_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X2)))))))))), file('/scratch/mptp/predsm40/t25_topgrp_1__0', t25_topgrp_1)).
fof(dt_k2_tops_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(v1_funct_1(k2_tops_2(X1,X2,X3))&(v1_funct_2(k2_tops_2(X1,X2,X3),X2,X1)&m1_subset_1(k2_tops_2(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X2,X1)))))), file('/scratch/mptp/predsm40/t25_topgrp_1__0', dt_k2_tops_2)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('/scratch/mptp/predsm40/t25_topgrp_1__0', dt_l1_pre_topc)).
fof(d5_tops_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(v3_tops_2(X3,X1,X2)<=>(k1_relset_1(u1_struct_0(X1),X3)=k2_struct_0(X1)&(k2_relset_1(u1_struct_0(X2),X3)=k2_struct_0(X2)&(v2_funct_1(X3)&(v5_pre_topc(X3,X1,X2)&v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1))))))))), file('/scratch/mptp/predsm40/t25_topgrp_1__0', d5_tops_2)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('/scratch/mptp/predsm40/t25_topgrp_1__0', cc2_relset_1)).
fof(redefinition_k7_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k7_relset_1(X1,X2,X3,X4)=k7_relat_1(X3,X4)), file('/scratch/mptp/predsm40/t25_topgrp_1__0', redefinition_k7_relset_1)).
fof(dt_k8_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>m1_subset_1(k8_relset_1(X1,X2,X3,X4),k1_zfmisc_1(X1))), file('/scratch/mptp/predsm40/t25_topgrp_1__0', dt_k8_relset_1)).
fof(redefinition_k8_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k8_relset_1(X1,X2,X3,X4)=k8_relat_1(X3,X4)), file('/scratch/mptp/predsm40/t25_topgrp_1__0', redefinition_k8_relset_1)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v5_relat_1(X2,X1))=>k2_relset_1(X1,X2)=k10_xtuple_0(X2)), file('/scratch/mptp/predsm40/t25_topgrp_1__0', redefinition_k2_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/scratch/mptp/predsm40/t25_topgrp_1__0', cc1_relset_1)).
fof(redefinition_r2_funct_2, axiom, ![X1, X2, X3, X4]:(((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&(v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))=>(r2_funct_2(X1,X2,X3,X4)<=>X3=X4)), file('/scratch/mptp/predsm40/t25_topgrp_1__0', redefinition_r2_funct_2)).
fof(t51_tops_2, axiom, ![X1]:(l1_struct_0(X1)=>![X2]:((~(v2_struct_0(X2))&l1_struct_0(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>((k2_relset_1(u1_struct_0(X2),X3)=k2_struct_0(X2)&v2_funct_1(X3))=>r2_funct_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),X3))))), file('/scratch/mptp/predsm40/t25_topgrp_1__0', t51_tops_2)).
fof(t77_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r1_tarski(X1,k10_xtuple_0(X2))=>k7_relat_1(X2,k8_relat_1(X2,X1))=X1)), file('/scratch/mptp/predsm40/t25_topgrp_1__0', t77_funct_1)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('/scratch/mptp/predsm40/t25_topgrp_1__0', d3_struct_0)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/scratch/mptp/predsm40/t25_topgrp_1__0', t3_subset)).
fof(t49_tops_2, axiom, ![X1]:(l1_struct_0(X1)=>![X2]:((~(v2_struct_0(X2))&l1_struct_0(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>((k2_relset_1(u1_struct_0(X2),X3)=k2_struct_0(X2)&v2_funct_1(X3))=>(k1_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))=k2_struct_0(X2)&k2_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))=k2_struct_0(X1)))))), file('/scratch/mptp/predsm40/t25_topgrp_1__0', t49_tops_2)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('/scratch/mptp/predsm40/t25_topgrp_1__0', d10_xboole_0)).
fof(t50_tops_2, axiom, ![X1]:(l1_struct_0(X1)=>![X2]:(l1_struct_0(X2)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>((k2_relset_1(u1_struct_0(X2),X3)=k2_struct_0(X2)&v2_funct_1(X3))=>v2_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)))))), file('/scratch/mptp/predsm40/t25_topgrp_1__0', t50_tops_2)).
fof(t82_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(v2_funct_1(X2)=>r1_tarski(k8_relat_1(X2,k7_relat_1(X2,X1)),X1))), file('/scratch/mptp/predsm40/t25_topgrp_1__0', t82_funct_1)).
fof(t76_funct_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(X1,k9_xtuple_0(X2))=>r1_tarski(X1,k8_relat_1(X2,k7_relat_1(X2,X1))))), file('/scratch/mptp/predsm40/t25_topgrp_1__0', t76_funct_1)).
fof(dt_k7_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>m1_subset_1(k7_relset_1(X1,X2,X3,X4),k1_zfmisc_1(X2))), file('/scratch/mptp/predsm40/t25_topgrp_1__0', dt_k7_relset_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>k1_relset_1(X1,X2)=k9_xtuple_0(X2)), file('/scratch/mptp/predsm40/t25_topgrp_1__0', redefinition_k1_relset_1)).
fof(d1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((~(X2=k1_xboole_0)=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X3)))&(X2=k1_xboole_0=>(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0)))), file('/scratch/mptp/predsm40/t25_topgrp_1__0', d1_funct_2)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('/scratch/mptp/predsm40/t25_topgrp_1__0', fc2_struct_0)).
fof(dt_o_0_0_xboole_0, axiom, v1_xboole_0(o_0_0_xboole_0), file('/scratch/mptp/predsm40/t25_topgrp_1__0', dt_o_0_0_xboole_0)).
fof(d2_xboole_0, axiom, k1_xboole_0=o_0_0_xboole_0, file('/scratch/mptp/predsm40/t25_topgrp_1__0', d2_xboole_0)).
fof(c_0_28, plain, ![X80, X81, X82, X83]:(((~v5_pre_topc(X82,X80,X81)|(~m1_subset_1(X83,k1_zfmisc_1(u1_struct_0(X81)))|(~v3_pre_topc(X83,X81)|v3_pre_topc(k8_relset_1(u1_struct_0(X80),u1_struct_0(X81),X82,X83),X80)))|k2_struct_0(X81)=k1_xboole_0|(~v1_funct_1(X82)|(~v1_funct_2(X82,u1_struct_0(X80),u1_struct_0(X81))|~m1_subset_1(X82,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X80),u1_struct_0(X81))))))|~l1_pre_topc(X81)|~l1_pre_topc(X80))&((m1_subset_1(esk8_3(X80,X81,X82),k1_zfmisc_1(u1_struct_0(X81)))|v5_pre_topc(X82,X80,X81)|k2_struct_0(X81)=k1_xboole_0|(~v1_funct_1(X82)|(~v1_funct_2(X82,u1_struct_0(X80),u1_struct_0(X81))|~m1_subset_1(X82,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X80),u1_struct_0(X81))))))|~l1_pre_topc(X81)|~l1_pre_topc(X80))&((v3_pre_topc(esk8_3(X80,X81,X82),X81)|v5_pre_topc(X82,X80,X81)|k2_struct_0(X81)=k1_xboole_0|(~v1_funct_1(X82)|(~v1_funct_2(X82,u1_struct_0(X80),u1_struct_0(X81))|~m1_subset_1(X82,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X80),u1_struct_0(X81))))))|~l1_pre_topc(X81)|~l1_pre_topc(X80))&(~v3_pre_topc(k8_relset_1(u1_struct_0(X80),u1_struct_0(X81),X82,esk8_3(X80,X81,X82)),X80)|v5_pre_topc(X82,X80,X81)|k2_struct_0(X81)=k1_xboole_0|(~v1_funct_1(X82)|(~v1_funct_2(X82,u1_struct_0(X80),u1_struct_0(X81))|~m1_subset_1(X82,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X80),u1_struct_0(X81))))))|~l1_pre_topc(X81)|~l1_pre_topc(X80)))))&((~v5_pre_topc(X82,X80,X81)|(~m1_subset_1(X83,k1_zfmisc_1(u1_struct_0(X81)))|(~v3_pre_topc(X83,X81)|v3_pre_topc(k8_relset_1(u1_struct_0(X80),u1_struct_0(X81),X82,X83),X80)))|k2_struct_0(X80)!=k1_xboole_0|(~v1_funct_1(X82)|(~v1_funct_2(X82,u1_struct_0(X80),u1_struct_0(X81))|~m1_subset_1(X82,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X80),u1_struct_0(X81))))))|~l1_pre_topc(X81)|~l1_pre_topc(X80))&((m1_subset_1(esk8_3(X80,X81,X82),k1_zfmisc_1(u1_struct_0(X81)))|v5_pre_topc(X82,X80,X81)|k2_struct_0(X80)!=k1_xboole_0|(~v1_funct_1(X82)|(~v1_funct_2(X82,u1_struct_0(X80),u1_struct_0(X81))|~m1_subset_1(X82,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X80),u1_struct_0(X81))))))|~l1_pre_topc(X81)|~l1_pre_topc(X80))&((v3_pre_topc(esk8_3(X80,X81,X82),X81)|v5_pre_topc(X82,X80,X81)|k2_struct_0(X80)!=k1_xboole_0|(~v1_funct_1(X82)|(~v1_funct_2(X82,u1_struct_0(X80),u1_struct_0(X81))|~m1_subset_1(X82,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X80),u1_struct_0(X81))))))|~l1_pre_topc(X81)|~l1_pre_topc(X80))&(~v3_pre_topc(k8_relset_1(u1_struct_0(X80),u1_struct_0(X81),X82,esk8_3(X80,X81,X82)),X80)|v5_pre_topc(X82,X80,X81)|k2_struct_0(X80)!=k1_xboole_0|(~v1_funct_1(X82)|(~v1_funct_2(X82,u1_struct_0(X80),u1_struct_0(X81))|~m1_subset_1(X82,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X80),u1_struct_0(X81))))))|~l1_pre_topc(X81)|~l1_pre_topc(X80)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_tops_2])])])])])).
fof(c_0_29, plain, ![X94, X95, X96, X97]:(~l1_struct_0(X94)|(~l1_struct_0(X95)|(~v1_funct_1(X96)|(~v1_funct_2(X96,u1_struct_0(X94),u1_struct_0(X95))|~m1_subset_1(X96,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X94),u1_struct_0(X95)))))|(~m1_subset_1(X97,k1_zfmisc_1(u1_struct_0(X94)))|(k2_relset_1(u1_struct_0(X95),X96)!=k2_struct_0(X95)|~v2_funct_1(X96)|k7_relset_1(u1_struct_0(X94),u1_struct_0(X95),X96,X97)=k8_relset_1(u1_struct_0(X95),u1_struct_0(X94),k2_tops_2(u1_struct_0(X94),u1_struct_0(X95),X96),X97)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t54_tops_2])])])).
fof(c_0_30, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&l1_pre_topc(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(v3_tops_2(X3,X1,X2)<=>(k1_relset_1(u1_struct_0(X1),X3)=k2_struct_0(X1)&(k2_relset_1(u1_struct_0(X2),X3)=k2_struct_0(X2)&(v2_funct_1(X3)&![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X4,X1)<=>v3_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X2))))))))))), inference(assume_negation,[status(cth)],[t25_topgrp_1])).
cnf(c_0_31, plain, (v5_pre_topc(X3,X1,X2)|k2_struct_0(X2)=k1_xboole_0|~v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk8_3(X1,X2,X3)),X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_32, plain, (k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4)=k8_relset_1(u1_struct_0(X2),u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X4)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v2_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_33, plain, ![X37, X38, X39]:((v1_funct_1(k2_tops_2(X37,X38,X39))|(~v1_funct_1(X39)|(~v1_funct_2(X39,X37,X38)|~m1_subset_1(X39,k1_zfmisc_1(k2_zfmisc_1(X37,X38))))))&((v1_funct_2(k2_tops_2(X37,X38,X39),X38,X37)|(~v1_funct_1(X39)|(~v1_funct_2(X39,X37,X38)|~m1_subset_1(X39,k1_zfmisc_1(k2_zfmisc_1(X37,X38))))))&(m1_subset_1(k2_tops_2(X37,X38,X39),k1_zfmisc_1(k2_zfmisc_1(X38,X37)))|(~v1_funct_1(X39)|(~v1_funct_2(X39,X37,X38)|~m1_subset_1(X39,k1_zfmisc_1(k2_zfmisc_1(X37,X38)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_tops_2])])])).
fof(c_0_34, plain, ![X48]:(~l1_pre_topc(X48)|l1_struct_0(X48)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_35, plain, ![X34, X35, X36]:(((k1_relset_1(u1_struct_0(X34),X36)=k2_struct_0(X34)|~v3_tops_2(X36,X34,X35)|(~v1_funct_1(X36)|(~v1_funct_2(X36,u1_struct_0(X34),u1_struct_0(X35))|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X34),u1_struct_0(X35))))))|~l1_pre_topc(X35)|~l1_pre_topc(X34))&((k2_relset_1(u1_struct_0(X35),X36)=k2_struct_0(X35)|~v3_tops_2(X36,X34,X35)|(~v1_funct_1(X36)|(~v1_funct_2(X36,u1_struct_0(X34),u1_struct_0(X35))|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X34),u1_struct_0(X35))))))|~l1_pre_topc(X35)|~l1_pre_topc(X34))&((v2_funct_1(X36)|~v3_tops_2(X36,X34,X35)|(~v1_funct_1(X36)|(~v1_funct_2(X36,u1_struct_0(X34),u1_struct_0(X35))|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X34),u1_struct_0(X35))))))|~l1_pre_topc(X35)|~l1_pre_topc(X34))&((v5_pre_topc(X36,X34,X35)|~v3_tops_2(X36,X34,X35)|(~v1_funct_1(X36)|(~v1_funct_2(X36,u1_struct_0(X34),u1_struct_0(X35))|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X34),u1_struct_0(X35))))))|~l1_pre_topc(X35)|~l1_pre_topc(X34))&(v5_pre_topc(k2_tops_2(u1_struct_0(X34),u1_struct_0(X35),X36),X35,X34)|~v3_tops_2(X36,X34,X35)|(~v1_funct_1(X36)|(~v1_funct_2(X36,u1_struct_0(X34),u1_struct_0(X35))|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X34),u1_struct_0(X35))))))|~l1_pre_topc(X35)|~l1_pre_topc(X34))))))&(k1_relset_1(u1_struct_0(X34),X36)!=k2_struct_0(X34)|(k2_relset_1(u1_struct_0(X35),X36)!=k2_struct_0(X35)|(~v2_funct_1(X36)|(~v5_pre_topc(X36,X34,X35)|~v5_pre_topc(k2_tops_2(u1_struct_0(X34),u1_struct_0(X35),X36),X35,X34))))|v3_tops_2(X36,X34,X35)|(~v1_funct_1(X36)|(~v1_funct_2(X36,u1_struct_0(X34),u1_struct_0(X35))|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X34),u1_struct_0(X35))))))|~l1_pre_topc(X35)|~l1_pre_topc(X34))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_tops_2])])])])).
fof(c_0_36, negated_conjecture, ![X9]:((~v2_struct_0(esk1_0)&l1_pre_topc(esk1_0))&((~v2_struct_0(esk2_0)&l1_pre_topc(esk2_0))&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))))&(((m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_funct_1(esk3_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))&((~v3_pre_topc(esk4_0,esk1_0)|~v3_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk4_0),esk2_0)|~v2_funct_1(esk3_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))&(v3_pre_topc(esk4_0,esk1_0)|v3_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk4_0),esk2_0)|~v2_funct_1(esk3_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))))&((k1_relset_1(u1_struct_0(esk1_0),esk3_0)=k2_struct_0(esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0))&((k2_relset_1(u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0))&((v2_funct_1(esk3_0)|v3_tops_2(esk3_0,esk1_0,esk2_0))&((~v3_pre_topc(X9,esk1_0)|v3_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X9),esk2_0)|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(esk1_0)))|v3_tops_2(esk3_0,esk1_0,esk2_0))&(~v3_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X9),esk2_0)|v3_pre_topc(X9,esk1_0)|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(esk1_0)))|v3_tops_2(esk3_0,esk1_0,esk2_0)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_30])])])])])])).
fof(c_0_37, plain, ![X14, X15, X16]:((v4_relat_1(X16,X14)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X14,X15))))&(v5_relat_1(X16,X15)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X14,X15))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
cnf(c_0_38, plain, (k2_struct_0(X2)=k1_xboole_0|v5_pre_topc(X3,X1,X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk8_3(X1,X2,X3)),X1)), inference(evalgc,[status(thm)],[c_0_31])).
cnf(c_0_39, plain, (k8_relset_1(u1_struct_0(X2),u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X4)=k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v1_funct_1(X3)|~v2_funct_1(X3)|~l1_struct_0(X2)|~l1_struct_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_32])).
cnf(c_0_40, plain, (v1_funct_1(k2_tops_2(X1,X2,X3))|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_41, plain, (v1_funct_2(k2_tops_2(X1,X2,X3),X2,X1)|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42, plain, (m1_subset_1(k2_tops_2(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_43, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_44, plain, (v2_funct_1(X1)|~v3_tops_2(X1,X2,X3)|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(X3)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_45, negated_conjecture, (v2_funct_1(esk3_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_46, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_47, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_48, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_49, negated_conjecture, (l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_50, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_51, plain, ![X61, X62, X63, X64]:(~m1_subset_1(X63,k1_zfmisc_1(k2_zfmisc_1(X61,X62)))|k7_relset_1(X61,X62,X63,X64)=k7_relat_1(X63,X64)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_relset_1])])).
fof(c_0_52, plain, ![X44, X45, X46, X47]:(~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(X44,X45)))|m1_subset_1(k8_relset_1(X44,X45,X46,X47),k1_zfmisc_1(X44))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_relset_1])])).
fof(c_0_53, plain, ![X65, X66, X67, X68]:(~m1_subset_1(X67,k1_zfmisc_1(k2_zfmisc_1(X65,X66)))|k8_relset_1(X65,X66,X67,X68)=k8_relat_1(X67,X68)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_relset_1])])).
fof(c_0_54, plain, ![X59, X60]:(~v1_relat_1(X60)|~v5_relat_1(X60,X59)|k2_relset_1(X59,X60)=k10_xtuple_0(X60)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
cnf(c_0_55, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_56, plain, ![X11, X12, X13]:(~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12)))|v1_relat_1(X13)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_57, plain, ![X69, X70, X71, X72]:((~r2_funct_2(X69,X70,X71,X72)|X71=X72|(~v1_funct_1(X71)|(~v1_funct_2(X71,X69,X70)|~m1_subset_1(X71,k1_zfmisc_1(k2_zfmisc_1(X69,X70))))|(~v1_funct_1(X72)|(~v1_funct_2(X72,X69,X70)|~m1_subset_1(X72,k1_zfmisc_1(k2_zfmisc_1(X69,X70)))))))&(X71!=X72|r2_funct_2(X69,X70,X71,X72)|(~v1_funct_1(X71)|(~v1_funct_2(X71,X69,X70)|~m1_subset_1(X71,k1_zfmisc_1(k2_zfmisc_1(X69,X70))))|(~v1_funct_1(X72)|(~v1_funct_2(X72,X69,X70)|~m1_subset_1(X72,k1_zfmisc_1(k2_zfmisc_1(X69,X70)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r2_funct_2])])])).
fof(c_0_58, plain, ![X91, X92, X93]:(~l1_struct_0(X91)|(v2_struct_0(X92)|~l1_struct_0(X92)|(~v1_funct_1(X93)|(~v1_funct_2(X93,u1_struct_0(X91),u1_struct_0(X92))|~m1_subset_1(X93,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X91),u1_struct_0(X92)))))|(k2_relset_1(u1_struct_0(X92),X93)!=k2_struct_0(X92)|~v2_funct_1(X93)|r2_funct_2(u1_struct_0(X91),u1_struct_0(X92),k2_tops_2(u1_struct_0(X92),u1_struct_0(X91),k2_tops_2(u1_struct_0(X91),u1_struct_0(X92),X93)),X93))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t51_tops_2])])])])).
cnf(c_0_59, plain, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|~v3_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk8_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))),X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk8_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39])])).
cnf(c_0_60, plain, (v1_funct_1(k2_tops_2(X1,X2,X3))|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_40])).
cnf(c_0_61, plain, (v1_funct_2(k2_tops_2(X1,X2,X3),X2,X1)|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_41])).
cnf(c_0_62, plain, (m1_subset_1(k2_tops_2(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_42])).
cnf(c_0_63, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[c_0_43])).
cnf(c_0_64, negated_conjecture, (v3_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1),esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_65, plain, (v2_funct_1(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v1_funct_1(X1)|~v3_tops_2(X1,X2,X3)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))), inference(evalgc,[status(thm)],[c_0_44])).
cnf(c_0_66, negated_conjecture, (v2_funct_1(esk3_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[c_0_45])).
cnf(c_0_67, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))), inference(evalgc,[status(thm)],[c_0_46])).
cnf(c_0_68, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))), inference(evalgc,[status(thm)],[c_0_47])).
cnf(c_0_69, negated_conjecture, (v1_funct_1(esk3_0)), inference(evalgc,[status(thm)],[c_0_48])).
cnf(c_0_70, negated_conjecture, (l1_pre_topc(esk2_0)), inference(evalgc,[status(thm)],[c_0_49])).
cnf(c_0_71, negated_conjecture, (l1_pre_topc(esk1_0)), inference(evalgc,[status(thm)],[c_0_50])).
cnf(c_0_72, negated_conjecture, (v3_pre_topc(X1,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_73, plain, (k7_relset_1(X2,X3,X1,X4)=k7_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_51])).
fof(c_0_74, plain, ![X105, X106]:(~v1_relat_1(X106)|~v1_funct_1(X106)|(~r1_tarski(X105,k10_xtuple_0(X106))|k7_relat_1(X106,k8_relat_1(X106,X105))=X105)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t77_funct_1])])).
cnf(c_0_75, plain, (m1_subset_1(k8_relset_1(X2,X3,X1,X4),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_76, plain, (k8_relset_1(X2,X3,X1,X4)=k8_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_77, plain, (k2_relset_1(X2,X1)=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_78, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(evalgc,[status(thm)],[c_0_55])).
cnf(c_0_79, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_80, plain, (X3=X4|~r2_funct_2(X1,X2,X3,X4)|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X4)|~v1_funct_2(X4,X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_81, plain, (v2_struct_0(X2)|r2_funct_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),X3)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v2_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_82, plain, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v3_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk8_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))),X2)|~v2_funct_1(X3)|~m1_subset_1(esk8_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_59, c_0_60]), c_0_61]), c_0_62]), c_0_63]), c_0_63])).
cnf(c_0_83, negated_conjecture, (v3_tops_2(esk3_0,esk1_0,esk2_0)|v3_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1),esk2_0)|~v3_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[c_0_64])).
cnf(c_0_84, negated_conjecture, (v2_funct_1(esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67]), c_0_68]), c_0_69]), c_0_70]), c_0_71])])])).
cnf(c_0_85, negated_conjecture, (k2_relset_1(u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_86, negated_conjecture, (v3_pre_topc(X1,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v3_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1),esk2_0)), inference(evalgc,[status(thm)],[c_0_72])).
cnf(c_0_87, plain, (k7_relset_1(X2,X3,X1,X4)=k7_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_73])).
cnf(c_0_88, plain, (k7_relat_1(X1,k8_relat_1(X1,X2))=X2|~v1_relat_1(X1)|~v1_funct_1(X1)|~r1_tarski(X2,k10_xtuple_0(X1))), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_89, plain, (m1_subset_1(k8_relset_1(X2,X3,X1,X4),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_75])).
cnf(c_0_90, plain, (k8_relset_1(X2,X3,X1,X4)=k8_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_76])).
cnf(c_0_91, plain, (k2_relset_1(X2,X1)=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_77])).
cnf(c_0_92, negated_conjecture, (v5_relat_1(esk3_0,u1_struct_0(esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_67])])).
cnf(c_0_93, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_79])).
cnf(c_0_94, plain, (X3=X4|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_2(X4,X1,X2)|~v1_funct_2(X3,X1,X2)|~r2_funct_2(X1,X2,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_80])).
cnf(c_0_95, plain, (v2_struct_0(X2)|r2_funct_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),X3)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v1_funct_1(X3)|~v2_funct_1(X3)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_81])).
cnf(c_0_96, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v5_pre_topc(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),esk2_0,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~m1_subset_1(esk8_3(esk2_0,esk1_0,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v3_pre_topc(esk8_3(esk2_0,esk1_0,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_84]), c_0_67]), c_0_68]), c_0_69]), c_0_71]), c_0_70])])])).
cnf(c_0_97, negated_conjecture, (k2_relset_1(u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[c_0_85])).
cnf(c_0_98, plain, (v3_pre_topc(esk8_3(X1,X2,X3),X2)|v5_pre_topc(X3,X1,X2)|k2_struct_0(X2)=k1_xboole_0|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_99, negated_conjecture, (v3_pre_topc(X1,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(k7_relat_1(esk3_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_67])])])).
cnf(c_0_100, plain, (k7_relat_1(X1,k8_relat_1(X1,X2))=X2|~v1_funct_1(X1)|~v1_relat_1(X1)|~r1_tarski(X2,k10_xtuple_0(X1))), inference(evalgc,[status(thm)],[c_0_88])).
cnf(c_0_101, plain, (m1_subset_1(k8_relat_1(X3,X4),k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90])])).
cnf(c_0_102, negated_conjecture, (k10_xtuple_0(esk3_0)=k2_relset_1(u1_struct_0(esk2_0),esk3_0)|~v1_relat_1(esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92])])).
cnf(c_0_103, negated_conjecture, (v1_relat_1(esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_67])])).
cnf(c_0_104, plain, (k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))=X3|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3)),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3)))|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95])])).
cnf(c_0_105, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v5_pre_topc(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),esk2_0,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(esk8_3(esk2_0,esk1_0,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)),esk1_0)|~m1_subset_1(esk8_3(esk2_0,esk1_0,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_106, plain, (k2_struct_0(X2)=k1_xboole_0|v5_pre_topc(X3,X1,X2)|v3_pre_topc(esk8_3(X1,X2,X3),X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_98])).
cnf(c_0_107, plain, (m1_subset_1(esk8_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X2)))|v5_pre_topc(X3,X1,X2)|k2_struct_0(X2)=k1_xboole_0|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_108, negated_conjecture, (v3_pre_topc(k8_relat_1(esk3_0,X1),esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(X1,esk2_0)|~m1_subset_1(k8_relat_1(esk3_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_tarski(X1,k10_xtuple_0(esk3_0))|~v1_relat_1(esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_69])])])).
cnf(c_0_109, negated_conjecture, (m1_subset_1(k8_relat_1(esk3_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_67])])).
cnf(c_0_110, negated_conjecture, (k10_xtuple_0(esk3_0)=k2_relset_1(u1_struct_0(esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_102, c_0_103])])).
fof(c_0_111, plain, ![X30]:(~l1_struct_0(X30)|k2_struct_0(X30)=u1_struct_0(X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_112, plain, (k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))=X3|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3)),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3)))|~v1_funct_1(X3)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_62])])).
cnf(c_0_113, plain, (v3_tops_2(X2,X1,X3)|k1_relset_1(u1_struct_0(X1),X2)!=k2_struct_0(X1)|k2_relset_1(u1_struct_0(X3),X2)!=k2_struct_0(X3)|~v2_funct_1(X2)|~v5_pre_topc(X2,X1,X3)|~v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X3),X2),X3,X1)|~v1_funct_1(X2)|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))|~l1_pre_topc(X3)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_114, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v5_pre_topc(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),esk2_0,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk8_3(esk2_0,esk1_0,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_71]), c_0_70])])])).
cnf(c_0_115, plain, (k2_struct_0(X2)=k1_xboole_0|v5_pre_topc(X3,X1,X2)|m1_subset_1(esk8_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_107])).
cnf(c_0_116, plain, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(X2,X3,X1)|~v3_pre_topc(k8_relat_1(X2,esk8_3(X3,X1,X2)),X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_pre_topc(X1)|~l1_pre_topc(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_90])])).
cnf(c_0_117, negated_conjecture, (v3_pre_topc(k8_relat_1(esk3_0,X1),esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~r1_tarski(X1,k2_relset_1(u1_struct_0(esk2_0),esk3_0))|~v3_pre_topc(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_108, c_0_109]), c_0_110]), c_0_103])])).
cnf(c_0_118, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_119, plain, (k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))=X3|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3)),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3)))|~v1_funct_1(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_112, c_0_60]), c_0_61]), c_0_62])).
cnf(c_0_120, plain, (v3_tops_2(X2,X1,X3)|k1_relset_1(u1_struct_0(X1),X2)!=k2_struct_0(X1)|k2_relset_1(u1_struct_0(X3),X2)!=k2_struct_0(X3)|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~v1_funct_1(X2)|~v2_funct_1(X2)|~v5_pre_topc(X2,X1,X3)|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))|~v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X3),X2),X3,X1)), inference(evalgc,[status(thm)],[c_0_113])).
cnf(c_0_121, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v5_pre_topc(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),esk2_0,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_71]), c_0_70])])])).
cnf(c_0_122, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)=k2_struct_0(esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_123, negated_conjecture, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(esk3_0,esk1_0,X1)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(X1))|~l1_pre_topc(X1)|~r1_tarski(esk8_3(esk1_0,X1,esk3_0),k2_relset_1(u1_struct_0(esk2_0),esk3_0))|~v3_pre_topc(esk8_3(esk1_0,X1,esk3_0),esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_69]), c_0_71])])])).
cnf(c_0_124, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_118])).
fof(c_0_125, plain, ![X78, X79]:((~m1_subset_1(X78,k1_zfmisc_1(X79))|r1_tarski(X78,X79))&(~r1_tarski(X78,X79)|m1_subset_1(X78,k1_zfmisc_1(X79)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_126, plain, (k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))=X3|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3)))|~v1_funct_1(X3)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_61])])).
cnf(c_0_127, negated_conjecture, (v3_tops_2(esk3_0,esk1_0,esk2_0)|k2_struct_0(esk1_0)=k1_xboole_0|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~m1_subset_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121]), c_0_84]), c_0_67]), c_0_68]), c_0_69]), c_0_70]), c_0_71])])])).
cnf(c_0_128, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)=k2_struct_0(esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[c_0_122])).
cnf(c_0_129, negated_conjecture, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(esk3_0,esk1_0,X1)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~r1_tarski(esk8_3(esk1_0,X1,esk3_0),k2_struct_0(esk2_0))|~v3_pre_topc(esk8_3(esk1_0,X1,esk3_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(X1))|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_97])])).
cnf(c_0_130, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_63])])).
cnf(c_0_131, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_125])).
fof(c_0_132, plain, ![X85, X86, X87]:((k1_relset_1(u1_struct_0(X86),k2_tops_2(u1_struct_0(X85),u1_struct_0(X86),X87))=k2_struct_0(X86)|(k2_relset_1(u1_struct_0(X86),X87)!=k2_struct_0(X86)|~v2_funct_1(X87))|(~v1_funct_1(X87)|(~v1_funct_2(X87,u1_struct_0(X85),u1_struct_0(X86))|~m1_subset_1(X87,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X85),u1_struct_0(X86))))))|(v2_struct_0(X86)|~l1_struct_0(X86))|~l1_struct_0(X85))&(k2_relset_1(u1_struct_0(X85),k2_tops_2(u1_struct_0(X85),u1_struct_0(X86),X87))=k2_struct_0(X85)|(k2_relset_1(u1_struct_0(X86),X87)!=k2_struct_0(X86)|~v2_funct_1(X87))|(~v1_funct_1(X87)|(~v1_funct_2(X87,u1_struct_0(X85),u1_struct_0(X86))|~m1_subset_1(X87,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X85),u1_struct_0(X86))))))|(v2_struct_0(X86)|~l1_struct_0(X86))|~l1_struct_0(X85))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t49_tops_2])])])])])).
cnf(c_0_133, plain, (k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))=X3|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3)))|~v1_funct_1(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_126, c_0_60]), c_0_61]), c_0_62])).
cnf(c_0_134, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~m1_subset_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_127, c_0_97]), c_0_128])).
cnf(c_0_135, negated_conjecture, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(esk3_0,esk1_0,X1)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~r1_tarski(esk8_3(esk1_0,X1,esk3_0),u1_struct_0(esk2_0))|~v3_pre_topc(esk8_3(esk1_0,X1,esk3_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(X1))|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130]), c_0_70])])])).
cnf(c_0_136, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_131])).
fof(c_0_137, plain, ![X23, X24]:(((r1_tarski(X23,X24)|X23!=X24)&(r1_tarski(X24,X23)|X23!=X24))&(~r1_tarski(X23,X24)|~r1_tarski(X24,X23)|X23=X24)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_138, plain, (k1_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))=k2_struct_0(X1)|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~v2_funct_1(X3)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~l1_struct_0(X1)|~l1_struct_0(X2)), inference(split_conjunct,[status(thm)],[c_0_132])).
cnf(c_0_139, plain, (k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))=X3|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_60])])).
fof(c_0_140, plain, ![X88, X89, X90]:(~l1_struct_0(X88)|(~l1_struct_0(X89)|(~v1_funct_1(X90)|(~v1_funct_2(X90,u1_struct_0(X88),u1_struct_0(X89))|~m1_subset_1(X90,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X88),u1_struct_0(X89)))))|(k2_relset_1(u1_struct_0(X89),X90)!=k2_struct_0(X89)|~v2_funct_1(X90)|v2_funct_1(k2_tops_2(u1_struct_0(X88),u1_struct_0(X89),X90)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t50_tops_2])])])).
cnf(c_0_141, plain, (v3_pre_topc(k8_relset_1(u1_struct_0(X2),u1_struct_0(X3),X1,X4),X2)|k2_struct_0(X3)=k1_xboole_0|~v5_pre_topc(X1,X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~v3_pre_topc(X4,X3)|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(X3)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_142, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~v1_funct_2(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_62]), c_0_67]), c_0_68]), c_0_69])])])).
cnf(c_0_143, negated_conjecture, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(esk3_0,esk1_0,X1)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(esk8_3(esk1_0,X1,esk3_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(X1))|~l1_pre_topc(X1)|~m1_subset_1(esk8_3(esk1_0,X1,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_136])])).
cnf(c_0_144, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_137])).
fof(c_0_145, plain, ![X107, X108]:(~v1_relat_1(X108)|~v1_funct_1(X108)|(~v2_funct_1(X108)|r1_tarski(k8_relat_1(X108,k7_relat_1(X108,X107)),X107))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t82_funct_1])])).
fof(c_0_146, plain, ![X103, X104]:(~v1_relat_1(X104)|(~r1_tarski(X103,k9_xtuple_0(X104))|r1_tarski(X103,k8_relat_1(X104,k7_relat_1(X104,X103))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t76_funct_1])])).
fof(c_0_147, plain, ![X40, X41, X42, X43]:(~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(X40,X41)))|m1_subset_1(k7_relset_1(X40,X41,X42,X43),k1_zfmisc_1(X41))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_relset_1])])).
cnf(c_0_148, plain, (k1_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))=k2_struct_0(X1)|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~v1_funct_1(X3)|~v2_funct_1(X3)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_138])).
cnf(c_0_149, plain, (k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))=X3|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_139, c_0_60]), c_0_61]), c_0_62])).
cnf(c_0_150, plain, (v2_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~l1_struct_0(X1)|~l1_struct_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v2_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_140])).
cnf(c_0_151, plain, (k2_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))=k2_struct_0(X1)|v2_struct_0(X2)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v2_funct_1(X3)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_struct_0(X2)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_132])).
cnf(c_0_152, plain, (k2_struct_0(X3)=k1_xboole_0|v3_pre_topc(k8_relset_1(u1_struct_0(X2),u1_struct_0(X3),X1,X4),X2)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v1_funct_1(X1)|~v3_pre_topc(X4,X3)|~v5_pre_topc(X1,X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))), inference(evalgc,[status(thm)],[c_0_141])).
cnf(c_0_153, plain, (v5_pre_topc(X1,X2,X3)|~v3_tops_2(X1,X2,X3)|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(X3)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_154, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~v1_funct_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_61]), c_0_67]), c_0_68]), c_0_69])])])).
cnf(c_0_155, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(esk3_0,esk1_0,esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk8_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_106]), c_0_67]), c_0_68]), c_0_70]), c_0_69]), c_0_71])])])).
cnf(c_0_156, plain, (X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_144])).
cnf(c_0_157, plain, (r1_tarski(k8_relat_1(X1,k7_relat_1(X1,X2)),X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v2_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_145])).
cnf(c_0_158, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_125])).
cnf(c_0_159, plain, (r1_tarski(X2,k8_relat_1(X1,k7_relat_1(X1,X2)))|~v1_relat_1(X1)|~r1_tarski(X2,k9_xtuple_0(X1))), inference(split_conjunct,[status(thm)],[c_0_146])).
cnf(c_0_160, plain, (m1_subset_1(k7_relset_1(X2,X3,X1,X4),k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_147])).
fof(c_0_161, plain, ![X57, X58]:(~v1_relat_1(X58)|~v4_relat_1(X58,X57)|k1_relset_1(X57,X58)=k9_xtuple_0(X58)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_162, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_163, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_funct_1(esk3_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_164, plain, (k1_relset_1(u1_struct_0(X1),X3)=k2_struct_0(X1)|v2_struct_0(X1)|v2_struct_0(X2)|k2_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X1)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v2_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_149])])).
cnf(c_0_165, plain, (v2_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v1_funct_1(X3)|~v2_funct_1(X3)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_150])).
cnf(c_0_166, plain, (k2_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))=k2_struct_0(X1)|v2_struct_0(X2)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v1_funct_1(X3)|~v2_funct_1(X3)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_151])).
cnf(c_0_167, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_168, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_169, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X2)|~v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|~v3_pre_topc(X4,X1)|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_39])])).
cnf(c_0_170, plain, (v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|~v3_tops_2(X3,X1,X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_171, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(k8_relat_1(X3,X4),X2)|~v5_pre_topc(X3,X2,X1)|~v3_pre_topc(X4,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_90])])).
cnf(c_0_172, plain, (v5_pre_topc(X1,X2,X3)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v1_funct_1(X1)|~v3_tops_2(X1,X2,X3)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))), inference(evalgc,[status(thm)],[c_0_153])).
cnf(c_0_173, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_60]), c_0_67]), c_0_68]), c_0_69])])])).
cnf(c_0_174, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(esk3_0,esk1_0,esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_115]), c_0_67]), c_0_68]), c_0_69]), c_0_70]), c_0_71])])])).
cnf(c_0_175, plain, (X1=X2|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_136])])).
cnf(c_0_176, plain, (r1_tarski(k8_relat_1(X1,k7_relat_1(X1,X2)),X2)|~v1_funct_1(X1)|~v2_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[c_0_157])).
cnf(c_0_177, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_158])).
cnf(c_0_178, plain, (r1_tarski(X2,k8_relat_1(X1,k7_relat_1(X1,X2)))|~v1_relat_1(X1)|~r1_tarski(X2,k9_xtuple_0(X1))), inference(evalgc,[status(thm)],[c_0_159])).
cnf(c_0_179, plain, (m1_subset_1(k7_relset_1(X2,X3,X1,X4),k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_160])).
cnf(c_0_180, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_161])).
cnf(c_0_181, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_162])).
cnf(c_0_182, negated_conjecture, (v3_pre_topc(esk4_0,esk1_0)|v3_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk4_0),esk2_0)|~v2_funct_1(esk3_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_183, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v2_funct_1(esk3_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[c_0_163])).
cnf(c_0_184, plain, (k1_relset_1(u1_struct_0(X1),X2)=k2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X3),X2)!=k2_struct_0(X3)|~l1_struct_0(X3)|~l1_struct_0(X1)|~v2_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(X3))|~v1_funct_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_164, c_0_60]), c_0_61]), c_0_62]), c_0_165]), c_0_166])).
cnf(c_0_185, negated_conjecture, (~v2_struct_0(esk2_0)), inference(evalgc,[status(thm)],[c_0_167])).
cnf(c_0_186, negated_conjecture, (~v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_168])).
cnf(c_0_187, negated_conjecture, (~v3_pre_topc(esk4_0,esk1_0)|~v3_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk4_0),esk2_0)|~v2_funct_1(esk3_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_188, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X2)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|~v3_pre_topc(X4,X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_169, c_0_60]), c_0_61]), c_0_62]), c_0_63]), c_0_63])).
cnf(c_0_189, plain, (v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v1_funct_1(X3)|~v3_tops_2(X3,X1,X2)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_170])).
cnf(c_0_190, plain, (k2_relset_1(u1_struct_0(X1),X2)=k2_struct_0(X1)|~v3_tops_2(X2,X3,X1)|~v1_funct_1(X2)|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~l1_pre_topc(X1)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_191, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(k8_relat_1(X2,X3),X4)|~v3_pre_topc(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X2,u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_pre_topc(X1)|~l1_pre_topc(X4)|~v3_tops_2(X2,X4,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_172])])).
cnf(c_0_192, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v3_tops_2(esk3_0,esk1_0,esk2_0)|k2_struct_0(esk2_0)=k1_xboole_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_174])])).
cnf(c_0_193, plain, (k8_relat_1(X1,k7_relat_1(X1,X2))=X2|~m1_subset_1(X2,k1_zfmisc_1(k8_relat_1(X1,k7_relat_1(X1,X2))))|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_176])])).
cnf(c_0_194, plain, (m1_subset_1(X1,k1_zfmisc_1(k8_relat_1(X2,k7_relat_1(X2,X1))))|~r1_tarski(X1,k9_xtuple_0(X2))|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_178])])).
cnf(c_0_195, plain, (m1_subset_1(k7_relat_1(X3,X4),k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_87])])).
cnf(c_0_196, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_180])).
cnf(c_0_197, negated_conjecture, (v4_relat_1(esk3_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_67])])).
cnf(c_0_198, negated_conjecture, (v3_pre_topc(esk4_0,esk1_0)|v3_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk4_0),esk2_0)|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v2_funct_1(esk3_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[c_0_182])).
cnf(c_0_199, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_183, c_0_84])])])).
cnf(c_0_200, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~l1_struct_0(esk2_0)|~l1_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_67]), c_0_84]), c_0_68]), c_0_69])]), c_0_185]), c_0_186])])).
cnf(c_0_201, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v2_funct_1(esk3_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk4_0),esk2_0)), inference(evalgc,[status(thm)],[c_0_187])).
cnf(c_0_202, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X2)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v3_pre_topc(X4,X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v3_tops_2(X3,X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_189])])).
cnf(c_0_203, plain, (k2_relset_1(u1_struct_0(X1),X2)=k2_struct_0(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~v1_funct_1(X2)|~v3_tops_2(X2,X3,X1)|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_190])).
cnf(c_0_204, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X1),esk1_0)|k2_struct_0(esk1_0)=k1_xboole_0|~v3_pre_topc(X1,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_192]), c_0_67]), c_0_68]), c_0_69]), c_0_70]), c_0_71])])])).
cnf(c_0_205, plain, (k8_relat_1(X1,k7_relat_1(X1,X2))=X2|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)|~r1_tarski(X2,k9_xtuple_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_194])])).
cnf(c_0_206, negated_conjecture, (m1_subset_1(k7_relat_1(esk3_0,X1),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_67])])).
cnf(c_0_207, negated_conjecture, (k9_xtuple_0(esk3_0)=k1_relset_1(u1_struct_0(esk1_0),esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_197]), c_0_103])])])).
cnf(c_0_208, negated_conjecture, (v3_pre_topc(k7_relat_1(esk3_0,esk4_0),esk2_0)|v3_pre_topc(esk4_0,esk1_0)|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v2_funct_1(esk3_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_87]), c_0_67])])])).
cnf(c_0_209, plain, (k1_relset_1(u1_struct_0(X1),X2)=k2_struct_0(X1)|~v3_tops_2(X2,X1,X3)|~v1_funct_1(X2)|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))|~l1_pre_topc(X3)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_210, plain, ![X25, X26, X27]:(((~v1_funct_2(X27,X25,X26)|X25=k1_relset_1(X25,X27)|X26=k1_xboole_0|~m1_subset_1(X27,k1_zfmisc_1(k2_zfmisc_1(X25,X26))))&(X25!=k1_relset_1(X25,X27)|v1_funct_2(X27,X25,X26)|X26=k1_xboole_0|~m1_subset_1(X27,k1_zfmisc_1(k2_zfmisc_1(X25,X26)))))&((~v1_funct_2(X27,X25,X26)|X27=k1_xboole_0|X26!=k1_xboole_0|~m1_subset_1(X27,k1_zfmisc_1(k2_zfmisc_1(X25,X26))))&(X27!=k1_xboole_0|v1_funct_2(X27,X25,X26)|X26!=k1_xboole_0|~m1_subset_1(X27,k1_zfmisc_1(k2_zfmisc_1(X25,X26)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_funct_2])])])])).
cnf(c_0_211, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)|~l1_struct_0(esk2_0)|~l1_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_200])])).
cnf(c_0_212, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk4_0),esk2_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_201, c_0_84])])])).
cnf(c_0_213, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X2)|~v3_pre_topc(X4,X1)|~v3_tops_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_202, c_0_65]), c_0_203])).
cnf(c_0_214, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,esk1_0)|~v3_pre_topc(k7_relat_1(esk3_0,X1),esk2_0)|~r1_tarski(X1,k1_relset_1(u1_struct_0(esk1_0),esk3_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_205]), c_0_206]), c_0_207]), c_0_103]), c_0_84]), c_0_69])])])).
cnf(c_0_215, negated_conjecture, (v3_pre_topc(k7_relat_1(esk3_0,esk4_0),esk2_0)|v3_pre_topc(esk4_0,esk1_0)|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_208, c_0_84])])])).
cnf(c_0_216, plain, (k1_relset_1(u1_struct_0(X1),X2)=k2_struct_0(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~v1_funct_1(X2)|~v3_tops_2(X2,X1,X3)|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))), inference(evalgc,[status(thm)],[c_0_209])).
cnf(c_0_217, plain, (X2=k1_relset_1(X2,X1)|X3=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_210])).
cnf(c_0_218, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~l1_struct_0(esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_63]), c_0_70])])])).
cnf(c_0_219, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_213]), c_0_67]), c_0_68]), c_0_69]), c_0_71]), c_0_70])])])).
cnf(c_0_220, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0|v3_pre_topc(esk4_0,esk1_0)|~r1_tarski(esk4_0,k1_relset_1(u1_struct_0(esk1_0),esk3_0))|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_215])])).
cnf(c_0_221, negated_conjecture, (k2_relset_1(u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk2_0)|k2_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_192]), c_0_67]), c_0_68]), c_0_69]), c_0_71]), c_0_70])])])).
cnf(c_0_222, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)=k2_struct_0(esk1_0)|k2_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_192]), c_0_67]), c_0_68]), c_0_69]), c_0_70]), c_0_71])])])).
cnf(c_0_223, plain, (X3=k1_xboole_0|k1_relset_1(X2,X1)=X2|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_217])).
cnf(c_0_224, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_63]), c_0_71])])])).
cnf(c_0_225, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(csr,[status(thm)],[c_0_219, c_0_199])).
cnf(c_0_226, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(esk4_0,esk1_0)|~r1_tarski(esk4_0,k1_relset_1(u1_struct_0(esk1_0),esk3_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_220, c_0_192]), c_0_221]), c_0_222])).
cnf(c_0_227, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)=u1_struct_0(esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_67]), c_0_68])])])).
cnf(c_0_228, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|~v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_224, c_0_221])])).
cnf(c_0_229, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)|~l1_struct_0(esk2_0)|~l1_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_200])])).
cnf(c_0_230, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(esk4_0,esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0|~r1_tarski(esk4_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_227])])).
cnf(c_0_231, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0|m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[c_0_228, c_0_192])).
cnf(c_0_232, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~l1_struct_0(esk1_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_63]), c_0_70])])])).
cnf(c_0_233, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0|v3_pre_topc(esk4_0,esk1_0)|~m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_136])])).
cnf(c_0_234, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0|m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_231]), c_0_70])])])).
cnf(c_0_235, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_63]), c_0_71])])])).
cnf(c_0_236, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(esk4_0,esk1_0)), inference(csr,[status(thm)],[c_0_233, c_0_234])).
cnf(c_0_237, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|u1_struct_0(esk2_0)=k1_xboole_0|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_236])])).
fof(c_0_238, plain, ![X49]:(v2_struct_0(X49)|~l1_struct_0(X49)|~v1_xboole_0(u1_struct_0(X49))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_239, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_237, c_0_192]), c_0_221])).
cnf(c_0_240, plain, (v1_xboole_0(o_0_0_xboole_0)), inference(split_conjunct,[status(thm)],[dt_o_0_0_xboole_0])).
cnf(c_0_241, plain, (k1_xboole_0=o_0_0_xboole_0), inference(split_conjunct,[status(thm)],[d2_xboole_0])).
cnf(c_0_242, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_238])).
cnf(c_0_243, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_239]), c_0_70])])])).
cnf(c_0_244, plain, (v1_xboole_0(o_0_0_xboole_0)), inference(evalgc,[status(thm)],[c_0_240])).
cnf(c_0_245, plain, (o_0_0_xboole_0=k1_xboole_0), inference(evalgc,[status(thm)],[c_0_241])).
cnf(c_0_246, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_242])).
cnf(c_0_247, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|u1_struct_0(esk2_0)=k1_xboole_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_243]), c_0_71])])])).
cnf(c_0_248, plain, (v1_xboole_0(k1_xboole_0)), inference(rw,[status(thm)],[c_0_244, c_0_245])).
cnf(c_0_249, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|~l1_struct_0(esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_247]), c_0_248])]), c_0_185])])).
cnf(c_0_250, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_249, c_0_63]), c_0_70])])])).
cnf(c_0_251, negated_conjecture, (~l1_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_250]), c_0_248])]), c_0_186])])).
cnf(c_0_252, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_63]), c_0_71])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 253
# Proof object clause steps            : 198
# Proof object formula steps           : 55
# Proof object conjectures             : 95
# Proof object clause conjectures      : 92
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 55
# Proof object initial formulas used   : 28
# Proof object generating inferences   : 68
# Proof object simplifying inferences  : 175
# Training examples: 125 positive, 484 negative
# Training: Positive examples begin
cnf(c_0_245, plain, (o_0_0_xboole_0=k1_xboole_0)).# trainpos
cnf(c_0_71, negated_conjecture, (l1_pre_topc(esk1_0))).# trainpos
cnf(c_0_70, negated_conjecture, (l1_pre_topc(esk2_0))).# trainpos
cnf(c_0_69, negated_conjecture, (v1_funct_1(esk3_0))).# trainpos
cnf(c_0_244, plain, (v1_xboole_0(o_0_0_xboole_0))).# trainpos
cnf(c_0_186, negated_conjecture, (~v2_struct_0(esk1_0))).# trainpos
cnf(c_0_185, negated_conjecture, (~v2_struct_0(esk2_0))).# trainpos
cnf(c_0_68, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0)))).# trainpos
cnf(c_0_67, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))))).# trainpos
cnf(c_0_66, negated_conjecture, (v2_funct_1(esk3_0)|v3_tops_2(esk3_0,esk1_0,esk2_0))).# trainpos
cnf(c_0_201, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v2_funct_1(esk3_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk4_0),esk2_0))).# trainpos
cnf(c_0_63, plain, (l1_struct_0(X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_183, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v2_funct_1(esk3_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))).# trainpos
cnf(c_0_128, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)=k2_struct_0(esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0))).# trainpos
cnf(c_0_124, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1))).# trainpos
cnf(c_0_246, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).# trainpos
cnf(c_0_177, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_136, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_156, plain, (X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_93, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_97, negated_conjecture, (k2_relset_1(u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0))).# trainpos
cnf(c_0_86, negated_conjecture, (v3_pre_topc(X1,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v3_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1),esk2_0))).# trainpos
cnf(c_0_83, negated_conjecture, (v3_tops_2(esk3_0,esk1_0,esk2_0)|v3_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1),esk2_0)|~v3_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_198, negated_conjecture, (v3_pre_topc(esk4_0,esk1_0)|v3_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk4_0),esk2_0)|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v2_funct_1(esk3_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))).# trainpos
cnf(c_0_196, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).# trainpos
cnf(c_0_91, plain, (k2_relset_1(X2,X1)=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).# trainpos
cnf(c_0_181, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_78, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).# trainpos
cnf(c_0_87, plain, (k7_relset_1(X2,X3,X1,X4)=k7_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_90, plain, (k8_relset_1(X2,X3,X1,X4)=k8_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_178, plain, (r1_tarski(X2,k8_relat_1(X1,k7_relat_1(X1,X2)))|~v1_relat_1(X1)|~r1_tarski(X2,k9_xtuple_0(X1)))).# trainpos
cnf(c_0_179, plain, (m1_subset_1(k7_relset_1(X2,X3,X1,X4),k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_89, plain, (m1_subset_1(k8_relset_1(X2,X3,X1,X4),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_100, plain, (k7_relat_1(X1,k8_relat_1(X1,X2))=X2|~v1_funct_1(X1)|~v1_relat_1(X1)|~r1_tarski(X2,k10_xtuple_0(X1)))).# trainpos
cnf(c_0_176, plain, (r1_tarski(k8_relat_1(X1,k7_relat_1(X1,X2)),X2)|~v1_funct_1(X1)|~v2_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_223, plain, (X3=k1_xboole_0|k1_relset_1(X2,X1)=X2|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_65, plain, (v2_funct_1(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v1_funct_1(X1)|~v3_tops_2(X1,X2,X3)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)))))).# trainpos
cnf(c_0_60, plain, (v1_funct_1(k2_tops_2(X1,X2,X3))|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_94, plain, (X3=X4|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_2(X4,X1,X2)|~v1_funct_2(X3,X1,X2)|~r2_funct_2(X1,X2,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_61, plain, (v1_funct_2(k2_tops_2(X1,X2,X3),X2,X1)|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_172, plain, (v5_pre_topc(X1,X2,X3)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v1_funct_1(X1)|~v3_tops_2(X1,X2,X3)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)))))).# trainpos
cnf(c_0_216, plain, (k1_relset_1(u1_struct_0(X1),X2)=k2_struct_0(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~v1_funct_1(X2)|~v3_tops_2(X2,X1,X3)|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3)))))).# trainpos
cnf(c_0_62, plain, (m1_subset_1(k2_tops_2(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_203, plain, (k2_relset_1(u1_struct_0(X1),X2)=k2_struct_0(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~v1_funct_1(X2)|~v3_tops_2(X2,X3,X1)|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1)))))).# trainpos
cnf(c_0_106, plain, (k2_struct_0(X2)=k1_xboole_0|v5_pre_topc(X3,X1,X2)|v3_pre_topc(esk8_3(X1,X2,X3),X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).# trainpos
cnf(c_0_189, plain, (v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v1_funct_1(X3)|~v3_tops_2(X3,X1,X2)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).# trainpos
cnf(c_0_165, plain, (v2_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v1_funct_1(X3)|~v2_funct_1(X3)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).# trainpos
cnf(c_0_115, plain, (k2_struct_0(X2)=k1_xboole_0|v5_pre_topc(X3,X1,X2)|m1_subset_1(esk8_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).# trainpos
cnf(c_0_38, plain, (k2_struct_0(X2)=k1_xboole_0|v5_pre_topc(X3,X1,X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk8_3(X1,X2,X3)),X1))).# trainpos
cnf(c_0_148, plain, (k1_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))=k2_struct_0(X1)|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~v1_funct_1(X3)|~v2_funct_1(X3)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1)))))).# trainpos
cnf(c_0_166, plain, (k2_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))=k2_struct_0(X1)|v2_struct_0(X2)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v1_funct_1(X3)|~v2_funct_1(X3)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).# trainpos
cnf(c_0_120, plain, (v3_tops_2(X2,X1,X3)|k1_relset_1(u1_struct_0(X1),X2)!=k2_struct_0(X1)|k2_relset_1(u1_struct_0(X3),X2)!=k2_struct_0(X3)|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~v1_funct_1(X2)|~v2_funct_1(X2)|~v5_pre_topc(X2,X1,X3)|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))|~v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X3),X2),X3,X1))).# trainpos
cnf(c_0_152, plain, (k2_struct_0(X3)=k1_xboole_0|v3_pre_topc(k8_relset_1(u1_struct_0(X2),u1_struct_0(X3),X1,X4),X2)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v1_funct_1(X1)|~v3_pre_topc(X4,X3)|~v5_pre_topc(X1,X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)))))).# trainpos
cnf(c_0_95, plain, (v2_struct_0(X2)|r2_funct_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),X3)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v1_funct_1(X3)|~v2_funct_1(X3)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).# trainpos
cnf(c_0_39, plain, (k8_relset_1(u1_struct_0(X2),u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X4)=k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v1_funct_1(X3)|~v2_funct_1(X3)|~l1_struct_0(X2)|~l1_struct_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).# trainpos
cnf(c_0_99, negated_conjecture, (v3_pre_topc(X1,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(k7_relat_1(esk3_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_130, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_208, negated_conjecture, (v3_pre_topc(k7_relat_1(esk3_0,esk4_0),esk2_0)|v3_pre_topc(esk4_0,esk1_0)|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v2_funct_1(esk3_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))).# trainpos
cnf(c_0_84, negated_conjecture, (v2_funct_1(esk3_0))).# trainpos
cnf(c_0_92, negated_conjecture, (v5_relat_1(esk3_0,u1_struct_0(esk2_0)))).# trainpos
cnf(c_0_103, negated_conjecture, (v1_relat_1(esk3_0))).# trainpos
cnf(c_0_197, negated_conjecture, (v4_relat_1(esk3_0,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_175, plain, (X1=X2|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_227, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)=u1_struct_0(esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0)).# trainpos
cnf(c_0_104, plain, (k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))=X3|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3)),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3)))|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3))).# trainpos
cnf(c_0_169, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X2)|~v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|~v3_pre_topc(X4,X1)|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3))).# trainpos
cnf(c_0_59, plain, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|~v3_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk8_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))),X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk8_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3))).# trainpos
cnf(c_0_171, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(k8_relat_1(X3,X4),X2)|~v5_pre_topc(X3,X2,X1)|~v3_pre_topc(X4,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2))).# trainpos
cnf(c_0_195, plain, (m1_subset_1(k7_relat_1(X3,X4),k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_101, plain, (m1_subset_1(k8_relat_1(X3,X4),k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_194, plain, (m1_subset_1(X1,k1_zfmisc_1(k8_relat_1(X2,k7_relat_1(X2,X1))))|~r1_tarski(X1,k9_xtuple_0(X2))|~v1_relat_1(X2))).# trainpos
cnf(c_0_212, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk4_0),esk2_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))).# trainpos
cnf(c_0_199, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))).# trainpos
cnf(c_0_109, negated_conjecture, (m1_subset_1(k8_relat_1(esk3_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_207, negated_conjecture, (k9_xtuple_0(esk3_0)=k1_relset_1(u1_struct_0(esk1_0),esk3_0))).# trainpos
cnf(c_0_215, negated_conjecture, (v3_pre_topc(k7_relat_1(esk3_0,esk4_0),esk2_0)|v3_pre_topc(esk4_0,esk1_0)|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))).# trainpos
cnf(c_0_102, negated_conjecture, (k10_xtuple_0(esk3_0)=k2_relset_1(u1_struct_0(esk2_0),esk3_0)|~v1_relat_1(esk3_0))).# trainpos
cnf(c_0_202, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X2)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v3_pre_topc(X4,X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v3_tops_2(X3,X1,X2))).# trainpos
cnf(c_0_193, plain, (k8_relat_1(X1,k7_relat_1(X1,X2))=X2|~m1_subset_1(X2,k1_zfmisc_1(k8_relat_1(X1,k7_relat_1(X1,X2))))|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1))).# trainpos
cnf(c_0_96, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v5_pre_topc(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),esk2_0,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~m1_subset_1(esk8_3(esk2_0,esk1_0,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v3_pre_topc(esk8_3(esk2_0,esk1_0,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)),esk1_0))).# trainpos
cnf(c_0_206, negated_conjecture, (m1_subset_1(k7_relat_1(esk3_0,X1),k1_zfmisc_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_219, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_114, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v5_pre_topc(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),esk2_0,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk8_3(esk2_0,esk1_0,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)))).# trainpos
cnf(c_0_205, plain, (k8_relat_1(X1,k7_relat_1(X1,X2))=X2|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)|~r1_tarski(X2,k9_xtuple_0(X1)))).# trainpos
cnf(c_0_191, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(k8_relat_1(X2,X3),X4)|~v3_pre_topc(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X2,u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_pre_topc(X1)|~l1_pre_topc(X4)|~v3_tops_2(X2,X4,X1))).# trainpos
cnf(c_0_108, negated_conjecture, (v3_pre_topc(k8_relat_1(esk3_0,X1),esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(X1,esk2_0)|~m1_subset_1(k8_relat_1(esk3_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_tarski(X1,k10_xtuple_0(esk3_0))|~v1_relat_1(esk3_0))).# trainpos
cnf(c_0_112, plain, (k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))=X3|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3)),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3)))|~v1_funct_1(X3)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3)))).# trainpos
cnf(c_0_121, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v5_pre_topc(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),esk2_0,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)))).# trainpos
cnf(c_0_126, plain, (k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))=X3|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3)))|~v1_funct_1(X3)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3)))).# trainpos
cnf(c_0_127, negated_conjecture, (v3_tops_2(esk3_0,esk1_0,esk2_0)|k2_struct_0(esk1_0)=k1_xboole_0|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~m1_subset_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)))).# trainpos
cnf(c_0_139, plain, (k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))=X3|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3)))).# trainpos
cnf(c_0_164, plain, (k1_relset_1(u1_struct_0(X1),X3)=k2_struct_0(X1)|v2_struct_0(X1)|v2_struct_0(X2)|k2_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X1)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v2_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3))).# trainpos
cnf(c_0_142, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~v1_funct_2(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)))).# trainpos
cnf(c_0_154, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~v1_funct_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)))).# trainpos
cnf(c_0_200, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~l1_struct_0(esk2_0)|~l1_struct_0(esk1_0))).# trainpos
cnf(c_0_229, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)|~l1_struct_0(esk2_0)|~l1_struct_0(esk1_0))).# trainpos
cnf(c_0_211, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)|~l1_struct_0(esk2_0)|~l1_struct_0(esk1_0))).# trainpos
cnf(c_0_116, plain, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(X2,X3,X1)|~v3_pre_topc(k8_relat_1(X2,esk8_3(X3,X1,X2)),X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_pre_topc(X1)|~l1_pre_topc(X3))).# trainpos
cnf(c_0_173, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v5_pre_topc(esk3_0,esk1_0,esk2_0))).# trainpos
cnf(c_0_232, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~l1_struct_0(esk1_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))).# trainpos
cnf(c_0_218, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~l1_struct_0(esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))).# trainpos
cnf(c_0_123, negated_conjecture, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(esk3_0,esk1_0,X1)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(X1))|~l1_pre_topc(X1)|~r1_tarski(esk8_3(esk1_0,X1,esk3_0),k2_relset_1(u1_struct_0(esk2_0),esk3_0))|~v3_pre_topc(esk8_3(esk1_0,X1,esk3_0),esk2_0))).# trainpos
cnf(c_0_235, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))).# trainpos
cnf(c_0_224, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))).# trainpos
cnf(c_0_129, negated_conjecture, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(esk3_0,esk1_0,X1)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~r1_tarski(esk8_3(esk1_0,X1,esk3_0),k2_struct_0(esk2_0))|~v3_pre_topc(esk8_3(esk1_0,X1,esk3_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(X1))|~l1_pre_topc(X1))).# trainpos
cnf(c_0_135, negated_conjecture, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(esk3_0,esk1_0,X1)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~r1_tarski(esk8_3(esk1_0,X1,esk3_0),u1_struct_0(esk2_0))|~v3_pre_topc(esk8_3(esk1_0,X1,esk3_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(X1))|~l1_pre_topc(X1))).# trainpos
cnf(c_0_143, negated_conjecture, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(esk3_0,esk1_0,X1)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(esk8_3(esk1_0,X1,esk3_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(X1))|~l1_pre_topc(X1)|~m1_subset_1(esk8_3(esk1_0,X1,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_155, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(esk3_0,esk1_0,esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk8_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_174, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(esk3_0,esk1_0,esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0))).# trainpos
cnf(c_0_192, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v3_tops_2(esk3_0,esk1_0,esk2_0)|k2_struct_0(esk2_0)=k1_xboole_0)).# trainpos
cnf(c_0_221, negated_conjecture, (k2_relset_1(u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk2_0)|k2_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0)).# trainpos
cnf(c_0_222, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)=k2_struct_0(esk1_0)|k2_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0)).# trainpos
cnf(c_0_204, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X1),esk1_0)|k2_struct_0(esk1_0)=k1_xboole_0|~v3_pre_topc(X1,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_228, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|~v3_tops_2(esk3_0,esk1_0,esk2_0))).# trainpos
cnf(c_0_214, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,esk1_0)|~v3_pre_topc(k7_relat_1(esk3_0,X1),esk2_0)|~r1_tarski(X1,k1_relset_1(u1_struct_0(esk1_0),esk3_0)))).# trainpos
cnf(c_0_234, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0|m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_220, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0|v3_pre_topc(esk4_0,esk1_0)|~r1_tarski(esk4_0,k1_relset_1(u1_struct_0(esk1_0),esk3_0))|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))).# trainpos
cnf(c_0_230, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(esk4_0,esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0|~r1_tarski(esk4_0,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_233, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0|v3_pre_topc(esk4_0,esk1_0)|~m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_237, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|u1_struct_0(esk2_0)=k1_xboole_0|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))).# trainpos
cnf(c_0_243, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0)).# trainpos
cnf(c_0_247, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|u1_struct_0(esk2_0)=k1_xboole_0)).# trainpos
cnf(c_0_249, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|~l1_struct_0(esk2_0))).# trainpos
cnf(c_0_250, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0)).# trainpos
cnf(c_0_251, negated_conjecture, (~l1_struct_0(esk1_0))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_95, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_55, plain, (v1_funct_1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_63, plain, (v1_xboole_0(esk6_2(X1,X2)))).#trainneg
cnf(i_0_58, plain, (v1_relat_1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_62, plain, (v1_relat_1(esk6_2(X1,X2)))).#trainneg
cnf(i_0_16, plain, (v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_26, plain, (r1_tarski(X1,X2)|X1!=X2)).#trainneg
cnf(i_0_27, plain, (r1_tarski(X1,X2)|X1!=X2)).#trainneg
cnf(i_0_57, plain, (v4_relat_1(esk5_2(X1,X2),X1))).#trainneg
cnf(i_0_61, plain, (v4_relat_1(esk6_2(X1,X2),X1))).#trainneg
cnf(i_0_56, plain, (v5_relat_1(esk5_2(X1,X2),X2))).#trainneg
cnf(i_0_60, plain, (v5_relat_1(esk6_2(X1,X2),X2))).#trainneg
cnf(i_0_54, plain, (v1_funct_2(esk5_2(X1,X2),X1,X2))).#trainneg
cnf(i_0_59, plain, (m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_64, plain, (m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_29, plain, (X1=k1_xboole_0|X3!=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_52, plain, (v1_funct_1(k2_funct_1(X1))|~v1_funct_1(X1)|~v2_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_34, plain, (k2_relset_1(X2,X1)=X2|~v1_relat_1(X1)|~v5_relat_1(X1,X2)|~v2_funct_2(X1,X2))).#trainneg
cnf(i_0_51, plain, (v2_funct_1(k2_funct_1(X1))|~v1_funct_1(X1)|~v2_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_53, plain, (v1_relat_1(k2_funct_1(X1))|~v1_funct_1(X1)|~v2_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_33, plain, (v2_funct_2(X2,X1)|k2_relset_1(X1,X2)!=X1|~v1_relat_1(X2)|~v5_relat_1(X2,X1))).#trainneg
cnf(i_0_20, plain, (v3_funct_2(X1,X2,X3)|~v1_funct_1(X1)|~v2_funct_1(X1)|~v2_funct_2(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_28, plain, (v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_30, plain, (X3=k1_xboole_0|v1_funct_2(X2,X1,X3)|k1_relset_1(X1,X2)!=X1|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_23, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|~v1_funct_1(X1)|~v1_xboole_0(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_36, plain, (k2_tops_2(X2,X3,X1)=k2_funct_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~v3_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_69, plain, (r2_funct_2(X3,X4,X1,X2)|X1!=X2|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,X3,X4)|~v1_funct_2(X1,X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_82, plain, (v5_pre_topc(X3,X1,X2)|v3_pre_topc(esk8_3(X1,X2,X3),X2)|k2_struct_0(X1)!=k1_xboole_0|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_83, plain, (v5_pre_topc(X3,X1,X2)|m1_subset_1(esk8_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X2)))|k2_struct_0(X1)!=k1_xboole_0|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_81, plain, (v5_pre_topc(X3,X1,X2)|k2_struct_0(X1)!=k1_xboole_0|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk8_3(X1,X2,X3)),X1))).#trainneg
cnf(i_0_74, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v4_pre_topc(X4,X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v1_funct_1(X3)|~v3_tops_2(X3,X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v4_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X1))).#trainneg
cnf(i_0_84, plain, (v3_pre_topc(k8_relset_1(u1_struct_0(X2),u1_struct_0(X3),X1,X4),X2)|k2_struct_0(X2)!=k1_xboole_0|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v1_funct_1(X1)|~v3_pre_topc(X4,X3)|~v5_pre_topc(X1,X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)))))).#trainneg
cnf(i_0_73, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v3_tops_2(X3,X1,X2)|m1_subset_1(esk7_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X2)))|k1_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v1_funct_1(X3)|~v2_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_75, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v4_pre_topc(k8_relset_1(u1_struct_0(X3),u1_struct_0(X2),X4,X1),X3)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v1_funct_1(X4)|~v4_pre_topc(X1,X2)|~v3_tops_2(X4,X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))).#trainneg
cnf(i_0_72, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v3_tops_2(X3,X1,X2)|k1_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v1_funct_1(X3)|~v2_funct_1(X3)|~v4_pre_topc(esk7_3(X1,X2,X3),X2)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v4_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk7_3(X1,X2,X3)),X1))).#trainneg
cnf(i_0_94, plain, (k7_relset_1(u1_struct_0(X2),u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X4)=k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v1_funct_1(X3)|~v2_funct_1(X3)|~l1_struct_0(X2)|~l1_struct_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_71, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v3_tops_2(X3,X1,X2)|v4_pre_topc(esk7_3(X1,X2,X3),X2)|v4_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk7_3(X1,X2,X3)),X1)|k1_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v1_funct_1(X3)|~v2_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_99, plain, (esk6_2(X1,X2)=k1_xboole_0)).#trainneg
cnf(i_0_131, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_pre_topc(k7_relat_1(esk3_0,esk4_0),esk2_0)|~v3_pre_topc(esk4_0,esk1_0)|~v2_funct_1(esk3_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))).#trainneg
cnf(i_0_132, negated_conjecture, (v3_pre_topc(k7_relat_1(esk3_0,X1),esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_180, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)=k2_struct_0(esk1_0)|v2_funct_1(esk3_0))).#trainneg
cnf(i_0_189, negated_conjecture, (k2_relset_1(u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk2_0)|v2_funct_1(esk3_0))).#trainneg
cnf(i_0_203, plain, (v3_tops_2(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X3)|k1_relset_1(u1_struct_0(X2),X1)!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X3),X1)!=k2_struct_0(X3)|~v4_pre_topc(k8_relat_1(X1,esk7_3(X2,X3,X1)),X2)|~v4_pre_topc(esk7_3(X2,X3,X1),X3)|~v2_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_219, plain, (v4_pre_topc(k8_relat_1(X3,esk7_3(X1,X2,X3)),X1)|v4_pre_topc(esk7_3(X1,X2,X3),X2)|v3_tops_2(X3,X1,X2)|v2_struct_0(X1)|v2_struct_0(X2)|k1_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_101, plain, (v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_117, negated_conjecture, (k1_xboole_0=esk3_0|u1_struct_0(esk2_0)!=k1_xboole_0)).#trainneg
cnf(i_0_26, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_175, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk2_0))|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_142, plain, (v1_relat_1(k7_relset_1(X1,k2_zfmisc_1(X2,X3),X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_143, plain, (v4_relat_1(k7_relset_1(X1,k2_zfmisc_1(X2,X3),X4,X5),X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_116, plain, (esk5_2(X1,X2)=k1_xboole_0|X2!=k1_xboole_0)).#trainneg
cnf(i_0_120, plain, (k1_relset_1(X1,esk5_2(X1,X2))=k9_xtuple_0(esk5_2(X1,X2)))).#trainneg
cnf(i_0_196, plain, (v3_tops_2(X1,X2,X3)|k1_relset_1(u1_struct_0(X2),X1)!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X3),X1)!=k2_struct_0(X3)|~v5_pre_topc(X1,X2,X3)|~v2_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v3_tops_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X3),X1),X3,X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X3),X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X3),X1),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X3),X1)))).#trainneg
cnf(i_0_220, plain, (v4_pre_topc(k7_relset_1(u1_struct_0(X2),u1_struct_0(X1),X3,esk7_3(X1,X2,k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))),X1)|v4_pre_topc(esk7_3(X1,X2,k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3)),X2)|v3_tops_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),X1,X2)|v2_struct_0(X1)|v2_struct_0(X2)|k1_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))!=k2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))!=k2_struct_0(X2)|~v2_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(esk7_3(X1,X2,k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3)),k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_216, plain, (v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|v2_struct_0(X2)|v2_struct_0(X1)|k1_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X1)|~v4_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk7_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))),X2)|~v4_pre_topc(esk7_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),X1)|~v2_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk7_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_202, plain, (v4_pre_topc(k8_relat_1(X3,X4),X1)|v2_struct_0(X2)|v2_struct_0(X1)|~v4_pre_topc(X4,X2)|~v3_tops_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_192, plain, (v4_pre_topc(X1,X2)|v2_struct_0(X3)|v2_struct_0(X2)|~v4_pre_topc(k8_relat_1(X4,X1),X3)|~v3_tops_2(X4,X3,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X4)|~l1_pre_topc(X2)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_214, plain, (v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|k2_struct_0(X2)!=k1_xboole_0|~v3_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk8_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))),X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk8_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_212, plain, (v4_pre_topc(X1,X2)|v2_struct_0(X3)|v2_struct_0(X2)|~v4_pre_topc(k7_relset_1(u1_struct_0(X2),u1_struct_0(X3),X4,X1),X3)|~v3_tops_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X3),X4),X3,X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X3),X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X3),X4),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X3),X4))|~l1_pre_topc(X2)|~l1_pre_topc(X3)|k2_relset_1(u1_struct_0(X3),X4)!=k2_struct_0(X3)|~l1_struct_0(X3)|~l1_struct_0(X2)|~v2_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X4))).#trainneg
cnf(i_0_211, plain, (v4_pre_topc(k7_relset_1(u1_struct_0(X2),u1_struct_0(X1),X3,X4),X1)|v2_struct_0(X2)|v2_struct_0(X1)|~v4_pre_topc(X4,X2)|~v3_tops_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),X1,X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_174, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(esk5_2(X1,X2)))).#trainneg
cnf(i_0_122, plain, (k2_relset_1(X1,esk5_2(X2,X1))=k10_xtuple_0(esk5_2(X2,X1)))).#trainneg
cnf(i_0_168, plain, (k1_relset_1(X1,esk5_2(X1,X2))=X1|X2=k1_xboole_0)).#trainneg
cnf(i_0_69, plain, (r2_funct_2(X1,X2,X3,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_129, plain, (k7_relset_1(X1,X2,X3,X4)=k7_relset_1(X5,X6,X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_182, plain, (v1_relat_1(k2_tops_2(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_145, plain, (v5_relat_1(k7_relset_1(X1,k2_zfmisc_1(X2,X3),X4,X5),X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_177, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(k7_relset_1(X3,k2_zfmisc_1(X1,X2),X4,X5))|~v1_funct_2(k7_relset_1(X3,k2_zfmisc_1(X1,X2),X4,X5),X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_134, plain, (k8_relset_1(X1,X2,X3,X4)=k8_relset_1(X5,X6,X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_6, negated_conjecture, (v3_pre_topc(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk4_0),esk2_0)|v3_pre_topc(esk4_0,esk1_0)|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))).#trainneg
cnf(i_0_325, plain, (v4_pre_topc(k8_relset_1(X5,X6,X3,X4),X1)|v2_struct_0(X2)|v2_struct_0(X1)|~v4_pre_topc(X4,X2)|~v3_tops_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_240, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_62, plain, (v1_relat_1(k1_xboole_0))).#trainneg
cnf(i_0_244, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v1_xboole_0(u1_struct_0(esk1_0))|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_131, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_pre_topc(k7_relat_1(esk3_0,esk4_0),esk2_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))).#trainneg
cnf(i_0_136, plain, (k8_relat_1(X1,k7_relat_1(X1,X2))=X2|~r1_tarski(k8_relat_1(X1,k7_relat_1(X1,X2)),X2)|~r1_tarski(X2,k9_xtuple_0(X1))|~v1_relat_1(X1))).#trainneg
cnf(i_0_246, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~l1_struct_0(esk2_0)|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_245, negated_conjecture, (v1_funct_1(u1_struct_0(esk2_0))|v1_xboole_0(u1_struct_0(esk1_0))|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_243, plain, (X1=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_230, negated_conjecture, (v3_pre_topc(k7_relset_1(X2,X3,esk3_0,X1),esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_137, plain, (r1_tarski(X1,k8_relat_1(X2,k7_relset_1(X3,X4,X2,X1)))|~r1_tarski(X1,k9_xtuple_0(X2))|~v1_relat_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_138, plain, (r1_tarski(X1,k8_relset_1(X3,X4,X2,k7_relat_1(X2,X1)))|~r1_tarski(X1,k9_xtuple_0(X2))|~v1_relat_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_247, negated_conjecture, (k2_struct_0(esk1_0)=u1_struct_0(esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|u1_struct_0(esk2_0)=k1_xboole_0)).#trainneg
cnf(i_0_235, negated_conjecture, (v4_pre_topc(k8_relat_1(esk3_0,esk7_3(esk1_0,esk2_0,esk3_0)),esk1_0)|v4_pre_topc(esk7_3(esk1_0,esk2_0,esk3_0),esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0))).#trainneg
cnf(i_0_222, negated_conjecture, (v3_pre_topc(X1,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(k7_relset_1(X2,X3,esk3_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_229, negated_conjecture, (v3_pre_topc(X1,esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(k8_relat_1(esk3_0,X1),esk1_0)|~m1_subset_1(k8_relat_1(esk3_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_tarski(X1,k10_xtuple_0(esk3_0))|~v1_relat_1(esk3_0))).#trainneg
cnf(i_0_233, plain, (v3_tops_2(X1,X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|k1_relset_1(u1_struct_0(X2),X1)!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X3),X1)!=k2_struct_0(X3)|~v4_pre_topc(k8_relset_1(X4,X5,X1,esk7_3(X2,X3,X1)),X2)|~v4_pre_topc(esk7_3(X2,X3,X1),X3)|~v2_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_60, plain, (v5_relat_1(k1_xboole_0,X2))).#trainneg
cnf(i_0_61, plain, (v4_relat_1(k1_xboole_0,X1))).#trainneg
cnf(i_0_64, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_144, plain, (v1_funct_2(k7_relset_1(X1,k2_zfmisc_1(X2,X3),X4,X5),X2,X3)|X3!=k1_xboole_0|k7_relset_1(X1,k2_zfmisc_1(X2,X3),X4,X5)!=k1_xboole_0|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_270, negated_conjecture, (v5_pre_topc(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),esk2_0,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|k2_struct_0(esk2_0)!=k1_xboole_0|~m1_subset_1(esk8_3(esk2_0,esk1_0,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v3_pre_topc(esk8_3(esk2_0,esk1_0,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)),esk1_0))).#trainneg
cnf(i_0_146, plain, (k7_relset_1(X1,k2_zfmisc_1(X2,X3),X4,X5)=k1_xboole_0|X3!=k1_xboole_0|~v1_funct_2(k7_relset_1(X1,k2_zfmisc_1(X2,X3),X4,X5),X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_148, plain, (v1_relat_1(k8_relset_1(k2_zfmisc_1(X1,X2),X3,X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_265, plain, (v4_pre_topc(X1,X2)|v2_struct_0(X2)|v2_struct_0(X3)|~v4_pre_topc(k8_relset_1(X5,X6,X4,X1),X3)|~v3_tops_2(X4,X3,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X4)|~l1_pre_topc(X2)|~l1_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_430, negated_conjecture, (k2_relset_1(u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk2_0)|k2_struct_0(esk1_0)=u1_struct_0(esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0)).#trainneg
cnf(i_0_428, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)=k2_struct_0(esk1_0)|k2_struct_0(esk1_0)=u1_struct_0(esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0)).#trainneg
cnf(i_0_363, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|u1_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)!=u1_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))).#trainneg
cnf(i_0_415, plain, (r1_tarski(X1,k8_relat_1(X2,k7_relat_1(X2,X1)))|~r1_tarski(X1,k9_xtuple_0(X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_355, negated_conjecture, (v3_pre_topc(k7_relset_1(X1,X2,esk3_0,esk4_0),esk2_0)|v3_pre_topc(esk4_0,esk1_0)|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_378, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|u1_struct_0(esk2_0)=k1_xboole_0|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_257, plain, (k1_relset_1(X1,k1_xboole_0)=k9_xtuple_0(k1_xboole_0)|X2!=k1_xboole_0)).#trainneg
cnf(i_0_149, plain, (v4_relat_1(k8_relset_1(k2_zfmisc_1(X1,X2),X3,X4,X5),X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_390, negated_conjecture, (~l1_struct_0(esk1_0)|~l1_struct_0(esk2_0)|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_389, negated_conjecture, (v1_funct_1(u1_struct_0(esk1_0))|~l1_struct_0(esk2_0)|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_368, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_392, negated_conjecture, (v1_funct_1(u1_struct_0(esk1_0))|v1_funct_1(u1_struct_0(esk2_0))|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_150, plain, (v1_funct_2(k8_relset_1(k2_zfmisc_1(X1,X2),X3,X4,X5),X1,X2)|X2!=k1_xboole_0|k8_relset_1(k2_zfmisc_1(X1,X2),X3,X4,X5)!=k1_xboole_0|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_151, plain, (v5_relat_1(k8_relset_1(k2_zfmisc_1(X1,X2),X3,X4,X5),X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_496, plain, (v4_pre_topc(X1,X2)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~v3_tops_2(X5,X3,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X5)|~l1_pre_topc(X2)|~l1_pre_topc(X3)|~v4_pre_topc(X1,X4)|~v3_tops_2(X5,X3,X4)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X4)))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X4))|~l1_pre_topc(X4))).#trainneg
cnf(i_0_483, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r1_tarski(X1,k2_relset_1(u1_struct_0(esk2_0),esk3_0)))).#trainneg
cnf(i_0_431, negated_conjecture, (v4_pre_topc(k8_relat_1(esk3_0,X1),esk1_0)|k2_struct_0(esk1_0)=u1_struct_0(esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0|~v4_pre_topc(X1,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_461, plain, (k2_relset_1(X1,k1_xboole_0)=k10_xtuple_0(k1_xboole_0))).#trainneg
cnf(i_0_447, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(k7_relat_1(X3,X4),X2)|~v3_pre_topc(X4,X1)|~v3_tops_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_398, negated_conjecture, (k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))=esk3_0|~m1_subset_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_380, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|~l1_struct_0(esk1_0)|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_152, plain, (k8_relset_1(k2_zfmisc_1(X1,X2),X3,X4,X5)=k1_xboole_0|X2!=k1_xboole_0|~v1_funct_2(k8_relset_1(k2_zfmisc_1(X1,X2),X3,X4,X5),X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_276, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|X1!=k1_xboole_0)).#trainneg
cnf(i_0_393, negated_conjecture, (v1_funct_1(u1_struct_0(esk2_0))|~l1_struct_0(esk1_0)|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_388, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|~l1_struct_0(esk2_0)|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_369, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_371, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_154, plain, (k7_relset_1(X1,X2,X3,k8_relat_1(X3,X4))=X4|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~r1_tarski(X4,k10_xtuple_0(X3))|~v1_relat_1(X3)|~v1_funct_1(X3))).#trainneg
cnf(i_0_195, plain, (v3_pre_topc(k8_relat_1(X3,X4),X1)|k2_struct_0(X1)!=k1_xboole_0|~v5_pre_topc(X3,X1,X2)|~v3_pre_topc(X4,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_413, plain, (k8_relat_1(X1,k7_relset_1(X2,X3,X1,X4))=X4|~r1_tarski(k8_relat_1(X1,k7_relset_1(X2,X3,X1,X4)),X4)|~r1_tarski(X4,k9_xtuple_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_161, plain, (r1_tarski(k8_relset_1(X3,X4,X1,k7_relat_1(X1,X2)),X2)|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_456, plain, (v3_tops_2(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|k1_relset_1(u1_struct_0(X2),X1)!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X3),X1)!=k2_struct_0(X3)|~v4_pre_topc(esk7_3(X2,X3,X1),X3)|~v2_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X4))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v4_pre_topc(esk7_3(X2,X3,X1),X4)|~v3_tops_2(X1,X2,X4)|~m1_subset_1(esk7_3(X2,X3,X1),k1_zfmisc_1(u1_struct_0(X4)))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X4))|~l1_pre_topc(X4))).#trainneg
cnf(i_0_242, plain, (k8_relat_1(X2,k7_relat_1(X2,X1))=X1|~m1_subset_1(k8_relat_1(X2,k7_relat_1(X2,X1)),k1_zfmisc_1(X1))|~r1_tarski(X1,k9_xtuple_0(X2))|~v1_relat_1(X2))).#trainneg
cnf(i_0_418, plain, (r1_tarski(X1,k8_relset_1(X5,X6,X2,k7_relset_1(X3,X4,X2,X1)))|~r1_tarski(X1,k9_xtuple_0(X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_462, plain, (k1_relset_1(X1,k1_xboole_0)=k9_xtuple_0(k1_xboole_0))).#trainneg
cnf(i_0_208, plain, (k8_relat_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),X4)=k7_relset_1(u1_struct_0(X2),u1_struct_0(X1),X3,X4)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_450, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(k8_relset_1(u1_struct_0(X2),u1_struct_0(X1),X3,X4),X2)|~v3_pre_topc(X4,X1)|~v3_tops_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),X1,X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_473, plain, (m1_subset_1(k8_relat_1(k1_xboole_0,X1),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_498, negated_conjecture, (k2_struct_0(esk1_0)=u1_struct_0(esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0)).#trainneg
cnf(i_0_501, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|u1_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)!=u1_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))).#trainneg
cnf(i_0_156, plain, (k7_relat_1(X1,k8_relset_1(X3,X4,X1,X2))=X2|~r1_tarski(X2,k10_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_255, plain, (v1_funct_2(k1_xboole_0,X1,X2)|X2!=k1_xboole_0)).#trainneg
cnf(i_0_379, negated_conjecture, (v1_funct_1(u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_391, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|v1_funct_1(u1_struct_0(esk2_0))|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_545, negated_conjecture, (m1_subset_1(k2_relset_1(u1_struct_0(esk2_0),esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_541, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|v3_tops_2(esk3_0,esk1_0,esk2_0)|~r1_tarski(X1,k2_struct_0(esk2_0)))).#trainneg
cnf(i_0_157, plain, (m1_subset_1(k8_relat_1(X1,k7_relat_1(X1,X2)),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_158, plain, (k8_relat_1(X2,k7_relat_1(X2,X1))=X1|~r1_tarski(X1,k8_relat_1(X2,k7_relat_1(X2,X1)))|~v1_relat_1(X2)|~v2_funct_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_594, plain, (r1_tarski(k8_relset_1(X1,X2,X3,X4),k8_relat_1(X3,X4))|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~r1_tarski(X4,k10_xtuple_0(X3))|~v1_relat_1(X3))).#trainneg
cnf(i_0_508, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_tarski(X1,k1_relset_1(u1_struct_0(esk1_0),esk3_0)))).#trainneg
cnf(i_0_514, plain, (k8_relat_1(X1,k7_relset_1(X3,X4,X1,X2))=X2|~r1_tarski(X2,k9_xtuple_0(X1))|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_554, negated_conjecture, (v3_pre_topc(X1,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~r1_tarski(k7_relat_1(esk3_0,X1),k2_relset_1(u1_struct_0(esk2_0),esk3_0))|~v3_pre_topc(k7_relat_1(esk3_0,X1),esk2_0)|~r1_tarski(X1,k1_relset_1(u1_struct_0(esk1_0),esk3_0)))).#trainneg
cnf(i_0_509, plain, (k8_relset_1(X1,X2,X3,k7_relat_1(X3,X4))=X4|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~r1_tarski(X4,k9_xtuple_0(X3))|~v1_relat_1(X3)|~v2_funct_1(X3)|~v1_funct_1(X3))).#trainneg
cnf(i_0_581, plain, (k7_relset_1(X1,X2,X3,k8_relset_1(X5,X6,X3,X4))=X4|~r1_tarski(X4,k10_xtuple_0(X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_646, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,k9_xtuple_0(k1_xboole_0))|~v2_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_550, plain, (k2_relset_1(X1,k1_xboole_0)=k2_relset_1(X2,k1_xboole_0))).#trainneg
cnf(i_0_602, negated_conjecture, (k8_relat_1(esk3_0,k7_relat_1(esk3_0,u1_struct_0(esk1_0)))=u1_struct_0(esk1_0)|~r1_tarski(u1_struct_0(esk1_0),k1_relset_1(u1_struct_0(esk1_0),esk3_0)))).#trainneg
cnf(i_0_358, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_pre_topc(k7_relset_1(X1,X2,esk3_0,esk4_0),esk2_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_481, negated_conjecture, (v5_pre_topc(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),esk2_0,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|k2_struct_0(esk2_0)!=k1_xboole_0|~m1_subset_1(esk8_3(esk2_0,esk1_0,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)))).#trainneg
cnf(i_0_484, negated_conjecture, (m1_subset_1(k7_relset_1(X2,X3,esk3_0,X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_666, negated_conjecture, (m1_subset_1(k2_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|v3_tops_2(esk3_0,esk1_0,esk2_0))).#trainneg
cnf(i_0_160, plain, (r1_tarski(k8_relat_1(X1,k7_relset_1(X3,X4,X1,X2)),X2)|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_165, plain, (X1=k1_xboole_0|v1_funct_2(k8_relset_1(k2_zfmisc_1(X2,X1),X3,X4,X5),X2,X1)|k1_relset_1(X2,k8_relset_1(k2_zfmisc_1(X2,X1),X3,X4,X5))!=X2|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X1),X3))))).#trainneg
cnf(i_0_708, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|v3_tops_2(esk3_0,esk1_0,esk2_0)|~r1_tarski(X1,k2_struct_0(esk1_0)))).#trainneg
cnf(i_0_785, negated_conjecture, (v5_pre_topc(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),esk2_0,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|k2_struct_0(esk2_0)!=k1_xboole_0|~m1_subset_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)))).#trainneg
cnf(i_0_713, negated_conjecture, (m1_subset_1(k1_relset_1(u1_struct_0(esk1_0),esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_756, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~v2_funct_1(k1_xboole_0)|~m1_subset_1(X1,k1_zfmisc_1(k9_xtuple_0(k1_xboole_0))))).#trainneg
cnf(i_0_712, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(u1_struct_0(esk1_0),esk3_0))))).#trainneg
cnf(i_0_779, negated_conjecture, (k2_struct_0(X1)=k1_xboole_0|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v3_pre_topc(esk4_0,X1)|~v3_tops_2(esk3_0,X1,esk2_0)|~m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_166, plain, (X1=k1_xboole_0|v1_funct_2(k7_relset_1(X2,k2_zfmisc_1(X3,X1),X4,X5),X3,X1)|k1_relset_1(X3,k7_relset_1(X2,k2_zfmisc_1(X3,X1),X4,X5))!=X3|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k2_zfmisc_1(X3,X1)))))).#trainneg
cnf(i_0_466, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_funct_2(k1_xboole_0,X1,X2))).#trainneg
cnf(i_0_665, negated_conjecture, (k2_relset_1(u1_struct_0(esk2_0),esk3_0)=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(k2_relset_1(u1_struct_0(esk2_0),esk3_0))))).#trainneg
cnf(i_0_670, negated_conjecture, (v3_tops_2(esk3_0,esk1_0,esk2_0)|m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(k2_struct_0(esk2_0))))).#trainneg
cnf(i_0_367, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2))).#trainneg
cnf(i_0_170, plain, (k1_relset_1(X1,k8_relset_1(k2_zfmisc_1(X1,X2),X3,X4,X5))=X1|X2=k1_xboole_0|~v1_funct_2(k8_relset_1(k2_zfmisc_1(X1,X2),X3,X4,X5),X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_171, plain, (k1_relset_1(X1,k7_relset_1(X2,k2_zfmisc_1(X1,X3),X4,X5))=X1|X3=k1_xboole_0|~v1_funct_2(k7_relset_1(X2,k2_zfmisc_1(X1,X3),X4,X5),X1,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k2_zfmisc_1(X1,X3)))))).#trainneg
cnf(i_0_729, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),X1),esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(X1,esk1_0)|~m1_subset_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_2(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)))).#trainneg
cnf(i_0_653, plain, (r1_tarski(k8_relset_1(X1,X2,X3,X4),k8_relat_1(X3,X4))|~r1_tarski(k8_relset_1(X1,X2,X3,X4),k9_xtuple_0(X3))|~v1_relat_1(X3)|~r1_tarski(X4,k10_xtuple_0(X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_266, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X2)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v3_pre_topc(X4,X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)))).#trainneg
cnf(i_0_593, plain, (k8_relset_1(X2,X3,X4,k7_relat_1(X4,X1))=X1|~r1_tarski(X1,k8_relset_1(X2,X3,X4,k7_relat_1(X4,X1)))|~v2_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_291, plain, (m1_subset_1(k7_relset_1(X5,X6,X3,X4),k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_451, negated_conjecture, (v3_pre_topc(X1,esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~r1_tarski(X1,k2_relset_1(u1_struct_0(esk2_0),esk3_0))|~v3_pre_topc(k8_relset_1(X2,X3,esk3_0,X1),esk1_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_753, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,k1_relset_1(X3,k1_xboole_0))|~v2_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_786, negated_conjecture, (k7_relset_1(X1,X2,esk3_0,X3)=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(k7_relset_1(X1,X2,esk3_0,X3)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_826, negated_conjecture, (v3_tops_2(esk3_0,esk1_0,esk2_0)|m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_176, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(k8_relset_1(k2_zfmisc_1(X1,X2),X3,X4,X5))|~v1_funct_2(k8_relset_1(k2_zfmisc_1(X1,X2),X3,X4,X5),X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_855, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)=u1_struct_0(esk1_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_relset_1(u1_struct_0(esk1_0),esk3_0))))).#trainneg
cnf(i_0_791, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)=k2_struct_0(esk1_0)|m1_subset_1(k2_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_178, plain, (k2_tops_2(X1,X2,X3)=k2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)|~v2_funct_2(X3,X2)|~v2_funct_1(X3))).#trainneg
cnf(i_0_181, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(k2_tops_2(X2,X1,X3))|~v1_funct_2(k2_tops_2(X2,X1,X3),X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v1_funct_2(X3,X2,X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_844, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),X4),X2)|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))!=k2_struct_0(X2)|~v5_pre_topc(X3,X2,X1)|~v3_pre_topc(X4,X1)|~v2_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_840, plain, (k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),X4)=k8_relat_1(X3,X4)|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))!=k2_struct_0(X2)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v2_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_853, plain, (v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|v2_struct_0(X2)|k1_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X1)|~v5_pre_topc(X3,X1,X2)|~v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|~v2_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_872, negated_conjecture, (k2_struct_0(X1)=k1_xboole_0|u1_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)!=u1_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_pre_topc(esk4_0,X1)|~v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_tops_2(esk3_0,X1,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_825, negated_conjecture, (v3_tops_2(esk3_0,esk1_0,esk2_0)|m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(k2_struct_0(esk1_0))))).#trainneg
cnf(i_0_793, negated_conjecture, (k2_relset_1(u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk2_0)|m1_subset_1(k2_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_183, plain, (v4_relat_1(k2_tops_2(X1,X2,X3),X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_471, plain, (k1_relset_1(X1,k1_xboole_0)=X1|X2=k1_xboole_0|~v1_funct_2(k1_xboole_0,X1,X2))).#trainneg
cnf(i_0_878, negated_conjecture, (k2_struct_0(esk2_0)=u1_struct_0(esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(k2_struct_0(esk2_0))))).#trainneg
cnf(i_0_773, negated_conjecture, (v3_pre_topc(k7_relat_1(esk3_0,u1_struct_0(esk1_0)),esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~r1_tarski(k7_relat_1(esk3_0,u1_struct_0(esk1_0)),k2_relset_1(u1_struct_0(esk2_0),esk3_0))|~v3_pre_topc(u1_struct_0(esk1_0),esk1_0)|~r1_tarski(u1_struct_0(esk1_0),k1_relset_1(u1_struct_0(esk1_0),esk3_0)))).#trainneg
cnf(i_0_544, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(k2_relset_1(u1_struct_0(esk2_0),esk3_0))))).#trainneg
cnf(i_0_185, plain, (v5_relat_1(k2_tops_2(X1,X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_912, plain, (k8_relat_1(X3,k7_relat_1(X3,X4))=X4|~r1_tarski(X4,k8_relat_1(X3,k7_relat_1(X3,X4)))|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_842, plain, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(X3,X2,X1)|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))!=k2_struct_0(X2)|~v3_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),esk8_3(X2,X1,X3)),X2)|~v2_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~m1_subset_1(esk8_3(X2,X1,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_561, plain, (k1_relset_1(X1,k2_zfmisc_1(X1,X2))=k9_xtuple_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_838, plain, (v5_pre_topc(X3,X2,X1)|v2_struct_0(X1)|~v3_tops_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),X1,X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_597, plain, (r1_tarski(k8_relat_1(X3,k7_relat_1(X3,X4)),X4)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_846, plain, (v4_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),esk7_3(X2,X1,X3)),X2)|v4_pre_topc(esk7_3(X2,X1,X3),X1)|v3_tops_2(X3,X2,X1)|v2_struct_0(X1)|v2_struct_0(X2)|k1_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))!=k2_struct_0(X2)|~v2_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~m1_subset_1(esk7_3(X2,X1,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_751, plain, (k7_relset_1(X1,X2,X3,k8_relset_1(X4,X5,X3,X6))=X6|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~v1_funct_1(X3)|~m1_subset_1(X6,k1_zfmisc_1(k10_xtuple_0(X3))))).#trainneg
cnf(i_0_858, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~v2_funct_1(k1_xboole_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(X3,k1_xboole_0))))).#trainneg
cnf(i_0_186, plain, (k2_tops_2(X1,X2,X3)=k1_xboole_0|X1!=k1_xboole_0|~v1_funct_2(k2_tops_2(X1,X2,X3),X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_953, negated_conjecture, (k2_struct_0(esk1_0)=u1_struct_0(esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k2_struct_0(esk1_0))))).#trainneg
cnf(i_0_1014, negated_conjecture, (m1_subset_1(k2_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_941, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)=k2_struct_0(esk1_0)|m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_943, negated_conjecture, (k2_relset_1(u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk2_0)|m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_187, plain, (k1_relset_1(X1,k2_tops_2(X2,X1,X3))=X1|X2=k1_xboole_0|~v1_funct_2(k2_tops_2(X2,X1,X3),X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v1_funct_2(X3,X2,X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_1039, plain, (k2_relset_1(X1,k2_tops_2(X1,X2,X3))=k10_xtuple_0(k2_tops_2(X1,X2,X3))|~v1_relat_1(k2_tops_2(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_1059, negated_conjecture, (r1_tarski(k8_relat_1(esk3_0,k7_relat_1(esk3_0,X1)),X1))).#trainneg
cnf(i_0_977, plain, (k1_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))=k2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))!=k2_struct_0(X2)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v2_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1101, plain, (k7_relset_1(X1,X2,X3,k8_relat_1(X3,X6))=X6|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k10_xtuple_0(X3)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_948, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v3_pre_topc(k7_relat_1(esk3_0,X1),esk2_0)|m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v3_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1094, plain, (r1_tarski(k8_relset_1(X1,X2,X3,X4),k8_relat_1(X3,X4))|~r1_tarski(k8_relset_1(X1,X2,X3,X4),k9_xtuple_0(X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,k1_zfmisc_1(k10_xtuple_0(X3)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1087, plain, (k7_relat_1(X3,k8_relset_1(X4,X5,X3,X6))=X6|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k10_xtuple_0(X3)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1175, negated_conjecture, (m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1160, negated_conjecture, (k2_struct_0(esk2_0)=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(k2_struct_0(esk2_0))))).#trainneg
cnf(i_0_1152, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)=k2_struct_0(esk1_0)|k2_struct_0(esk1_0)=u1_struct_0(esk1_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k2_struct_0(esk1_0))))).#trainneg
cnf(i_0_1154, negated_conjecture, (k2_relset_1(u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk2_0)|k2_struct_0(esk1_0)=u1_struct_0(esk1_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k2_struct_0(esk1_0))))).#trainneg
cnf(i_0_191, plain, (v5_pre_topc(X1,X2,X3)|k2_struct_0(X2)!=k1_xboole_0|~v3_pre_topc(k8_relat_1(X1,esk8_3(X2,X3,X1)),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_206, plain, (k7_relat_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),X4)=k8_relset_1(u1_struct_0(X2),u1_struct_0(X1),X3,X4)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1208, negated_conjecture, (r1_tarski(k8_relat_1(esk3_0,k7_relset_1(X2,X3,esk3_0,X1)),X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_589, plain, (k8_relat_1(X1,k7_relset_1(X2,X3,X1,X4))=X4|~r1_tarski(X4,k9_xtuple_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k8_relat_1(X1,k7_relset_1(X2,X3,X1,X4)),k1_zfmisc_1(X4)))).#trainneg
cnf(i_0_845, plain, (v3_tops_2(X3,X2,X1)|v2_struct_0(X1)|v2_struct_0(X2)|k1_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))!=k2_struct_0(X2)|~v4_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),esk7_3(X2,X1,X3)),X2)|~v4_pre_topc(esk7_3(X2,X1,X3),X1)|~v2_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~m1_subset_1(esk7_3(X2,X1,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1189, plain, (k10_xtuple_0(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))=k2_struct_0(X1)|v2_struct_0(X2)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_854, plain, (v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|v2_struct_0(X2)|k1_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X1)|~v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|~v2_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~v3_tops_2(X3,X1,X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_334, plain, (m1_subset_1(k8_relset_1(X5,X6,X3,X4),k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_492, plain, (v4_pre_topc(X1,X2)|v2_struct_0(X3)|v2_struct_0(X2)|~v4_pre_topc(k7_relset_1(u1_struct_0(X5),u1_struct_0(X4),X6,X1),X3)|~v3_tops_2(k2_tops_2(u1_struct_0(X5),u1_struct_0(X4),X6),X3,X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X5),u1_struct_0(X4),X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(k2_tops_2(u1_struct_0(X5),u1_struct_0(X4),X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X5))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(k2_tops_2(u1_struct_0(X5),u1_struct_0(X4),X6),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X5),u1_struct_0(X4),X6))|~l1_pre_topc(X2)|~l1_pre_topc(X3)|k2_relset_1(u1_struct_0(X4),X6)!=k2_struct_0(X4)|~l1_struct_0(X4)|~l1_struct_0(X5)|~v2_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X4))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X5)))|~v1_funct_2(X6,u1_struct_0(X5),u1_struct_0(X4))|~v1_funct_1(X6))).#trainneg
cnf(i_0_615, plain, (k1_relset_1(X1,k1_xboole_0)=k1_relset_1(X2,k1_xboole_0))).#trainneg
cnf(i_0_1114, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v5_pre_topc(k1_xboole_0,esk2_0,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v1_funct_2(k1_xboole_0,u1_struct_0(esk2_0),u1_struct_0(esk1_0))|u1_struct_0(esk1_0)!=k1_xboole_0)).#trainneg
cnf(i_0_1097, negated_conjecture, (X5=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(X5))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_relset_1(u1_struct_0(esk2_0),esk3_0))))).#trainneg
cnf(i_0_1111, negated_conjecture, (v5_pre_topc(k1_xboole_0,esk2_0,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|k2_struct_0(esk2_0)!=k1_xboole_0|~v1_funct_2(k1_xboole_0,u1_struct_0(esk2_0),u1_struct_0(esk1_0))|u1_struct_0(esk1_0)!=k1_xboole_0)).#trainneg
cnf(i_0_1197, negated_conjecture, (m1_subset_1(k8_relat_1(esk3_0,k7_relat_1(esk3_0,k2_relset_1(u1_struct_0(esk2_0),esk3_0))),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_925, plain, (m1_subset_1(k7_relset_1(X1,X2,k2_zfmisc_1(X3,X4),X5),k1_zfmisc_1(X4))|~m1_subset_1(k2_zfmisc_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_210, plain, (v3_pre_topc(k7_relset_1(u1_struct_0(X2),u1_struct_0(X1),X3,X4),X1)|k2_struct_0(X1)!=k1_xboole_0|~v5_pre_topc(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),X1,X2)|~v3_pre_topc(X4,X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1249, negated_conjecture, (k2_struct_0(esk1_0)=u1_struct_0(esk1_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k2_struct_0(esk1_0))))).#trainneg
cnf(i_0_1338, plain, (v3_tops_2(X3,X2,X1)|v2_struct_0(X2)|v2_struct_0(X1)|k1_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~v2_funct_1(X3)|~v3_tops_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),X1,X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~l1_struct_0(X1)|~l1_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1277, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,X1)|v3_tops_2(esk3_0,esk1_0,esk2_0)|k2_struct_0(esk1_0)!=k1_xboole_0|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(X1))|~l1_pre_topc(X1)|~r1_tarski(esk8_3(esk1_0,X1,esk3_0),k2_relset_1(u1_struct_0(esk2_0),esk3_0))|~v3_pre_topc(esk8_3(esk1_0,X1,esk3_0),esk2_0))).#trainneg
cnf(i_0_1313, negated_conjecture, (r1_tarski(k8_relset_1(X4,X5,esk3_0,k7_relset_1(X1,X2,esk3_0,X3)),X3)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_1198, negated_conjecture, (v3_tops_2(esk3_0,esk1_0,esk2_0)|m1_subset_1(k8_relat_1(esk3_0,k7_relat_1(esk3_0,k2_struct_0(esk2_0))),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_248, plain, (v1_relat_1(k7_relat_1(X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_474, plain, (X1=k1_xboole_0|v1_funct_2(k1_xboole_0,X2,X1)|k1_relset_1(X2,k1_xboole_0)!=X2)).#trainneg
cnf(i_0_1373, negated_conjecture, (k8_relat_1(esk3_0,k7_relat_1(esk3_0,k2_relset_1(u1_struct_0(esk2_0),esk3_0)))=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(k8_relat_1(esk3_0,k7_relat_1(esk3_0,k2_relset_1(u1_struct_0(esk2_0),esk3_0))))))).#trainneg
cnf(i_0_883, negated_conjecture, (v3_tops_2(esk3_0,esk1_0,esk2_0)|m1_subset_1(k7_relset_1(X1,k2_struct_0(esk2_0),X2,X3),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k2_struct_0(esk2_0)))))).#trainneg
cnf(i_0_1010, negated_conjecture, (v3_tops_2(esk3_0,esk1_0,esk2_0)|m1_subset_1(k7_relset_1(X1,k2_struct_0(esk1_0),X2,X3),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k2_struct_0(esk1_0)))))).#trainneg
cnf(i_0_234, plain, (v4_pre_topc(k8_relat_1(esk5_2(u1_struct_0(X1),u1_struct_0(X2)),esk7_3(X1,X2,esk5_2(u1_struct_0(X1),u1_struct_0(X2)))),X1)|v4_pre_topc(esk7_3(X1,X2,esk5_2(u1_struct_0(X1),u1_struct_0(X2))),X2)|v3_tops_2(esk5_2(u1_struct_0(X1),u1_struct_0(X2)),X1,X2)|v2_struct_0(X2)|v2_struct_0(X1)|k1_relset_1(u1_struct_0(X1),esk5_2(u1_struct_0(X1),u1_struct_0(X2)))!=k2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),esk5_2(u1_struct_0(X1),u1_struct_0(X2)))!=k2_struct_0(X2)|~v2_funct_1(esk5_2(u1_struct_0(X1),u1_struct_0(X2)))|~l1_pre_topc(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_236, plain, (v4_pre_topc(k8_relat_1(k8_relset_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)),X3,X4,X5),esk7_3(X1,X2,k8_relset_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)),X3,X4,X5))),X1)|v4_pre_topc(esk7_3(X1,X2,k8_relset_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)),X3,X4,X5)),X2)|v3_tops_2(k8_relset_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)),X3,X4,X5),X1,X2)|v2_struct_0(X2)|v2_struct_0(X1)|k1_relset_1(u1_struct_0(X1),k8_relset_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)),X3,X4,X5))!=k2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),k8_relset_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)),X3,X4,X5))!=k2_struct_0(X2)|~v2_funct_1(k8_relset_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)),X3,X4,X5))|~v1_funct_2(k8_relset_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)),X3,X4,X5),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k8_relset_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)),X3,X4,X5))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)),X3))))).#trainneg
cnf(i_0_726, plain, (k8_relat_1(X1,k7_relat_1(X1,X4))=X4|~r1_tarski(X4,k9_xtuple_0(X1))|~v2_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1))).#trainneg
cnf(i_0_967, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(k7_relat_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),X4),X2)|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~v5_pre_topc(X3,X2,X1)|~v3_pre_topc(X4,X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_1003, plain, (v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|v2_struct_0(X2)|k1_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v5_pre_topc(X3,X1,X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v3_tops_2(X3,X1,X2))).#trainneg
cnf(i_0_905, plain, (r1_tarski(k8_relset_1(X1,X2,X3,X4),k8_relat_1(X3,X4))|~r1_tarski(X4,k10_xtuple_0(X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(k8_relset_1(X1,X2,X3,X4),k1_zfmisc_1(k9_xtuple_0(X3))))).#trainneg
cnf(i_0_399, negated_conjecture, (k8_relat_1(esk3_0,X1)=u1_struct_0(esk1_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k8_relat_1(esk3_0,X1))))).#trainneg
cnf(i_0_249, plain, (k1_relset_1(X1,k7_relset_1(X2,k2_zfmisc_1(X1,X3),X4,X5))=k9_xtuple_0(k7_relset_1(X2,k2_zfmisc_1(X1,X3),X4,X5))|~v1_relat_1(k7_relset_1(X2,k2_zfmisc_1(X1,X3),X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k2_zfmisc_1(X1,X3)))))).#trainneg
cnf(i_0_1460, plain, (v1_funct_2(k1_xboole_0,X1,X2)|k9_xtuple_0(k1_xboole_0)!=X1)).#trainneg
cnf(i_0_237, plain, (v4_pre_topc(k8_relat_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),esk7_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))),X2)|v4_pre_topc(esk7_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),X1)|v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|v2_struct_0(X1)|v2_struct_0(X2)|k1_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X1)|~v2_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_238, plain, (v4_pre_topc(k8_relat_1(k7_relset_1(X1,k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)),X4,X5),esk7_3(X2,X3,k7_relset_1(X1,k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)),X4,X5))),X2)|v4_pre_topc(esk7_3(X2,X3,k7_relset_1(X1,k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)),X4,X5)),X3)|v3_tops_2(k7_relset_1(X1,k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)),X4,X5),X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|k1_relset_1(u1_struct_0(X2),k7_relset_1(X1,k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)),X4,X5))!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X3),k7_relset_1(X1,k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)),X4,X5))!=k2_struct_0(X3)|~v2_funct_1(k7_relset_1(X1,k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)),X4,X5))|~v1_funct_2(k7_relset_1(X1,k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)),X4,X5),u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(k7_relset_1(X1,k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)),X4,X5))|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))))).#trainneg
cnf(i_0_1415, plain, (v3_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X2)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|k2_struct_0(X2)!=k1_xboole_0|~v3_pre_topc(X4,X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v3_tops_2(X3,X1,X2))).#trainneg
cnf(i_0_1514, plain, (v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|v2_struct_0(X2)|k2_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X1)|~v3_tops_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3))).#trainneg
cnf(i_0_516, negated_conjecture, (r1_tarski(X1,k8_relat_1(esk3_0,k7_relat_1(esk3_0,X1)))|~r1_tarski(X1,k1_relset_1(u1_struct_0(esk1_0),esk3_0)))).#trainneg
cnf(i_0_1412, plain, (v3_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X2)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|k2_struct_0(X2)!=k1_xboole_0|~v3_pre_topc(X4,X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)))).#trainneg
cnf(i_0_1017, plain, (k1_relset_1(X1,k2_tops_2(X2,X1,X3))=k9_xtuple_0(k2_tops_2(X2,X1,X3))|~v1_relat_1(k2_tops_2(X2,X1,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v1_funct_2(X3,X2,X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_1420, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,X1)|v3_tops_2(esk3_0,esk1_0,esk2_0)|k2_struct_0(esk1_0)!=k1_xboole_0|~r1_tarski(esk8_3(esk1_0,X1,esk3_0),k2_struct_0(esk2_0))|~v3_pre_topc(esk8_3(esk1_0,X1,esk3_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1506, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(k7_relat_1(X3,X4),X2)|v2_struct_0(X1)|v2_struct_0(X2)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X1)|~v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|~v3_pre_topc(X4,X1)|~v2_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_482, negated_conjecture, (k7_relat_1(esk3_0,X1)=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(k7_relat_1(esk3_0,X1))))).#trainneg
cnf(i_0_250, plain, (v4_relat_1(k7_relat_1(X4,X5),X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_633, plain, (v1_relat_1(k8_relat_1(k1_xboole_0,X1)))).#trainneg
cnf(i_0_1446, negated_conjecture, (v3_tops_2(esk3_0,esk1_0,esk2_0)|m1_subset_1(k8_relat_1(esk3_0,k7_relat_1(esk3_0,u1_struct_0(esk2_0))),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_1527, negated_conjecture, (X1=u1_struct_0(esk1_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(X1))|~r1_tarski(X1,k1_relset_1(u1_struct_0(esk1_0),esk3_0)))).#trainneg
cnf(i_0_1387, plain, (v1_relat_1(k7_relset_1(X1,X2,k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5)),X6))|~m1_subset_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5)),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_260, plain, (v4_pre_topc(k7_relat_1(X3,esk7_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))),X2)|v4_pre_topc(esk7_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),X1)|v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|v2_struct_0(X1)|v2_struct_0(X2)|k1_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(esk7_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_565, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(X4,X2)|~v3_pre_topc(k8_relat_1(X3,X4),X1)|~v3_tops_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(k8_relat_1(X3,X4),k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~r1_tarski(X4,k10_xtuple_0(X3)))).#trainneg
cnf(i_0_693, plain, (k8_relat_1(X1,k7_relset_1(X3,X4,X1,X2))=X2|~r1_tarski(X2,k8_relat_1(X1,k7_relset_1(X3,X4,X1,X2)))|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_843, plain, (v5_pre_topc(X3,X2,X1)|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))!=k2_struct_0(X2)|k2_struct_0(X2)!=k1_xboole_0|~v3_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),esk8_3(X2,X1,X3)),X2)|~v2_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~m1_subset_1(esk8_3(X2,X1,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1588, plain, (v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|v2_struct_0(X2)|~v3_tops_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v2_funct_1(X3))).#trainneg
cnf(i_0_1571, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)|k2_struct_0(esk2_0)!=k1_xboole_0|~m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1607, plain, (k1_relset_1(u1_struct_0(X1),X3)=k9_xtuple_0(X3)|v2_struct_0(X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1635, negated_conjecture, (k8_relat_1(esk3_0,k7_relat_1(esk3_0,u1_struct_0(esk2_0)))=u1_struct_0(esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~r1_tarski(u1_struct_0(esk2_0),k1_relset_1(u1_struct_0(esk1_0),esk3_0)))).#trainneg
cnf(i_0_1631, plain, (v1_relat_1(X1)|~r1_tarski(X1,k9_xtuple_0(k1_xboole_0))|~v2_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_1028, negated_conjecture, (v3_pre_topc(k7_relat_1(esk3_0,u1_struct_0(esk1_0)),esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~r1_tarski(k7_relat_1(esk3_0,u1_struct_0(esk1_0)),k2_struct_0(esk2_0))|~r1_tarski(u1_struct_0(esk1_0),k1_relset_1(u1_struct_0(esk1_0),esk3_0))|~v3_pre_topc(u1_struct_0(esk1_0),esk1_0))).#trainneg
cnf(i_0_1639, negated_conjecture, (X1=u1_struct_0(esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~r1_tarski(X1,k2_struct_0(esk1_0))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1619, negated_conjecture, (X1=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(X1))|~r1_tarski(X1,k2_relset_1(u1_struct_0(esk2_0),esk3_0)))).#trainneg
cnf(i_0_262, plain, (v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|v2_struct_0(X1)|v2_struct_0(X2)|k1_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v4_pre_topc(k7_relat_1(X3,esk7_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))),X2)|~v4_pre_topc(esk7_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),X1)|~v2_funct_1(X3)|~m1_subset_1(esk7_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_268, plain, (v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|k2_struct_0(X2)!=k1_xboole_0|~v3_pre_topc(k7_relat_1(X3,esk8_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))),X2)|~v2_funct_1(X3)|~m1_subset_1(esk8_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1611, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,X1)|v3_tops_2(esk3_0,esk1_0,esk2_0)|k2_struct_0(esk1_0)!=k1_xboole_0|~r1_tarski(esk8_3(esk1_0,X1,esk3_0),u1_struct_0(esk2_0))|~v3_pre_topc(esk8_3(esk1_0,X1,esk3_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_330, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(k8_relset_1(X5,X6,X3,X4),X2)|~v5_pre_topc(X3,X2,X1)|~v3_pre_topc(X4,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_1276, negated_conjecture, (v5_pre_topc(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),esk2_0,X1)|k2_struct_0(esk1_0)=k1_xboole_0|v3_tops_2(esk3_0,esk1_0,esk2_0)|k2_struct_0(esk2_0)!=k1_xboole_0|~m1_subset_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),u1_struct_0(esk2_0),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0))|~l1_pre_topc(X1)|~v3_pre_topc(esk8_3(esk2_0,X1,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)),esk1_0)|~m1_subset_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~m1_subset_1(esk8_3(esk2_0,X1,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_2(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1209, negated_conjecture, (r1_tarski(k8_relset_1(X2,X3,esk3_0,k7_relat_1(esk3_0,X1)),X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1336, plain, (k10_xtuple_0(X3)=k2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))!=k2_struct_0(X2)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v2_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1591, negated_conjecture, (k8_relat_1(esk3_0,k7_relat_1(esk3_0,X1))=X1|~r1_tarski(X1,k1_relset_1(u1_struct_0(esk1_0),esk3_0)))).#trainneg
cnf(i_0_1706, negated_conjecture, (k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|k2_struct_0(esk2_0)!=k1_xboole_0|~v3_pre_topc(esk4_0,esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)|~l1_struct_0(esk2_0)|~l1_struct_0(esk1_0))).#trainneg
cnf(i_0_1744, plain, (v1_relat_1(k9_xtuple_0(k1_xboole_0))|~v2_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_1705, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)!=u1_struct_0(esk1_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|k2_struct_0(esk2_0)!=k1_xboole_0|~v3_pre_topc(esk4_0,esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))).#trainneg
cnf(i_0_1737, negated_conjecture, (v1_relat_1(k8_relat_1(esk3_0,k7_relat_1(esk3_0,k9_xtuple_0(k1_xboole_0))))|~v2_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_1758, negated_conjecture, (X1=u1_struct_0(esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~r1_tarski(X1,k2_struct_0(esk2_0))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_271, plain, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v3_pre_topc(k7_relat_1(X3,esk8_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))),X2)|~v2_funct_1(X3)|~m1_subset_1(esk8_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_274, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(k8_relat_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X4),X2)|~v3_pre_topc(X4,X1)|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v3_tops_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1703, plain, (v3_tops_2(X3,X2,X1)|v2_struct_0(X2)|v2_struct_0(X1)|~v3_tops_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),X1,X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1775, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,X1)|v3_tops_2(esk3_0,esk1_0,esk2_0)|k2_struct_0(esk1_0)!=k1_xboole_0|~v3_pre_topc(esk8_3(esk1_0,X1,esk3_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(X1))|~l1_pre_topc(X1)|~m1_subset_1(esk8_3(esk1_0,X1,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_1699, plain, (v4_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X2)|v2_struct_0(X2)|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v4_pre_topc(X4,X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v3_tops_2(X3,X1,X2))).#trainneg
cnf(i_0_1687, plain, (k2_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))=k2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v3_tops_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1697, plain, (v4_pre_topc(X1,X2)|v2_struct_0(X2)|v2_struct_0(X3)|k2_relset_1(u1_struct_0(X3),X4)!=k2_struct_0(X3)|~v4_pre_topc(k7_relset_1(u1_struct_0(X2),u1_struct_0(X3),X4,X1),X3)|~v2_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X4)|~l1_pre_topc(X2)|~l1_pre_topc(X3)|~v3_tops_2(X4,X2,X3))).#trainneg
cnf(i_0_1169, negated_conjecture, (k2_struct_0(X1)=k1_xboole_0|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_pre_topc(esk4_0,X1)|~v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_tops_2(esk3_0,X1,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_pre_topc(X1)|~l1_struct_0(esk2_0)|~l1_struct_0(esk1_0))).#trainneg
cnf(i_0_1756, negated_conjecture, (X1=u1_struct_0(esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_struct_0(esk1_0))))).#trainneg
cnf(i_0_1747, negated_conjecture, (v3_pre_topc(k7_relat_1(esk3_0,u1_struct_0(esk1_0)),esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~r1_tarski(u1_struct_0(esk1_0),k1_relset_1(u1_struct_0(esk1_0),esk3_0))|~v3_pre_topc(u1_struct_0(esk1_0),esk1_0)|~m1_subset_1(k7_relat_1(esk3_0,u1_struct_0(esk1_0)),k1_zfmisc_1(k2_struct_0(esk2_0))))).#trainneg
cnf(i_0_279, plain, (k9_xtuple_0(esk5_2(X1,X2))=X1|X2=k1_xboole_0)).#trainneg
cnf(i_0_1745, negated_conjecture, (v3_pre_topc(k7_relat_1(esk3_0,u1_struct_0(esk1_0)),esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~r1_tarski(k7_relat_1(esk3_0,u1_struct_0(esk1_0)),u1_struct_0(esk2_0))|~r1_tarski(u1_struct_0(esk1_0),k1_relset_1(u1_struct_0(esk1_0),esk3_0))|~v3_pre_topc(u1_struct_0(esk1_0),esk1_0))).#trainneg
cnf(i_0_1031, negated_conjecture, (v3_pre_topc(k7_relat_1(esk3_0,u1_struct_0(esk1_0)),esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~r1_tarski(u1_struct_0(esk1_0),k1_relset_1(u1_struct_0(esk1_0),esk3_0))|~v3_pre_topc(u1_struct_0(esk1_0),esk1_0)|~m1_subset_1(k7_relat_1(esk3_0,u1_struct_0(esk1_0)),k1_zfmisc_1(k2_relset_1(u1_struct_0(esk2_0),esk3_0))))).#trainneg
cnf(i_0_283, plain, (v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|v2_struct_0(X1)|v2_struct_0(X2)|k1_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v4_pre_topc(k7_relset_1(X4,X5,X3,esk7_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))),X2)|~v4_pre_topc(esk7_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),X1)|~v2_funct_1(X3)|~m1_subset_1(esk7_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_284, plain, (v4_pre_topc(k7_relset_1(X4,X5,X3,esk7_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))),X2)|v4_pre_topc(esk7_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),X1)|v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|v2_struct_0(X1)|v2_struct_0(X2)|k1_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(esk7_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_573, plain, (k7_relat_1(X3,k8_relat_1(X3,X4))=X4|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~r1_tarski(X4,k10_xtuple_0(X3))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1696, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(k8_relset_1(u1_struct_0(X2),u1_struct_0(X1),X3,X4),X2)|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~v3_pre_topc(X4,X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v3_tops_2(X3,X2,X1))).#trainneg
cnf(i_0_562, plain, (k2_relset_1(X1,k2_zfmisc_1(X2,X1))=k10_xtuple_0(k2_zfmisc_1(X2,X1)))).#trainneg
cnf(i_0_1327, plain, (v3_tops_2(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X3)|k2_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X3),X1))!=k2_struct_0(X2)|k1_relset_1(u1_struct_0(X2),X1)!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X3),X1)!=k2_struct_0(X3)|~v4_pre_topc(k7_relat_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X3),X1),esk7_3(X2,X3,X1)),X2)|~v4_pre_topc(esk7_3(X2,X3,X1),X3)|~v2_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X3),X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))).#trainneg
cnf(i_0_585, plain, (k8_relat_1(X1,k7_relat_1(X1,X4))=X4|~r1_tarski(k8_relat_1(X1,k7_relat_1(X1,X4)),X4)|~r1_tarski(X4,k9_xtuple_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1075, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(X6,X2)|~v3_pre_topc(k8_relset_1(X4,X5,X3,X6),X1)|~v3_tops_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(k8_relset_1(X4,X5,X3,X6),k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k10_xtuple_0(X3))))).#trainneg
cnf(i_0_324, plain, (v3_pre_topc(k8_relset_1(X5,X6,X3,X4),X1)|k2_struct_0(X1)!=k1_xboole_0|~v5_pre_topc(X3,X1,X2)|~v3_pre_topc(X4,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_1879, plain, (v4_pre_topc(k7_relat_1(X3,X4),X2)|v2_struct_0(X1)|v2_struct_0(X2)|~v4_pre_topc(X4,X1)|~v3_tops_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1648, negated_conjecture, (X1=u1_struct_0(esk1_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(u1_struct_0(esk1_0),esk3_0))))).#trainneg
cnf(i_0_288, plain, (v1_relat_1(k7_relset_1(X6,X7,X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))).#trainneg
cnf(i_0_1545, plain, (v1_funct_2(k1_xboole_0,k9_xtuple_0(k1_xboole_0),X1))).#trainneg
cnf(i_0_1842, negated_conjecture, (k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|k2_struct_0(esk2_0)!=k1_xboole_0|~l1_struct_0(esk1_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))).#trainneg
cnf(i_0_1857, negated_conjecture, (X1=u1_struct_0(esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_struct_0(esk2_0))))).#trainneg
cnf(i_0_285, plain, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v3_pre_topc(k7_relset_1(X4,X5,X3,esk8_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))),X2)|~v2_funct_1(X3)|~m1_subset_1(esk8_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_286, plain, (v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|k2_struct_0(X2)!=k1_xboole_0|~v3_pre_topc(k7_relset_1(X4,X5,X3,esk8_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))),X2)|~v2_funct_1(X3)|~m1_subset_1(esk8_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_1890, plain, (v4_pre_topc(X1,X2)|v2_struct_0(X3)|v2_struct_0(X2)|~v4_pre_topc(k7_relat_1(X4,X1),X3)|~v3_tops_2(X4,X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X4)|~l1_pre_topc(X2)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_1887, plain, (k1_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))=k2_struct_0(X1)|v2_struct_0(X1)|v2_struct_0(X2)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v2_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~v3_tops_2(X3,X2,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1882, plain, (v4_pre_topc(X4,X2)|v2_struct_0(X1)|v2_struct_0(X2)|~v4_pre_topc(k8_relat_1(X3,X4),X1)|~v3_tops_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(k8_relat_1(X3,X4),k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~r1_tarski(X4,k10_xtuple_0(X3)))).#trainneg
cnf(i_0_1886, plain, (k10_xtuple_0(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))=k2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~v3_tops_2(X3,X1,X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1900, negated_conjecture, (k2_struct_0(X1)=k1_xboole_0|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~l1_struct_0(esk1_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_pre_topc(esk4_0,X1)|~v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_tops_2(esk3_0,X1,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1893, plain, (v4_pre_topc(k8_relat_1(X1,X2),X3)|v2_struct_0(X4)|v2_struct_0(X3)|~v4_pre_topc(X2,X4)|~v3_tops_2(X1,X3,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4))))|~m1_subset_1(k8_relat_1(X1,X2),k1_zfmisc_1(u1_struct_0(X3)))|~v1_funct_2(X1,u1_struct_0(X3),u1_struct_0(X4))|~v1_funct_1(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X4)|~r1_tarski(X2,k10_xtuple_0(X1)))).#trainneg
cnf(i_0_1917, negated_conjecture, (v3_pre_topc(k7_relat_1(esk3_0,u1_struct_0(esk1_0)),esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~r1_tarski(u1_struct_0(esk1_0),k1_relset_1(u1_struct_0(esk1_0),esk3_0))|~v3_pre_topc(u1_struct_0(esk1_0),esk1_0))).#trainneg
cnf(i_0_1765, negated_conjecture, (X1=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_relset_1(u1_struct_0(esk2_0),esk3_0))))).#trainneg
cnf(i_0_289, plain, (v4_relat_1(k7_relset_1(X6,X7,X4,X5),X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))).#trainneg
cnf(i_0_1455, plain, (v1_relat_1(k7_relat_1(k1_xboole_0,X1)))).#trainneg
cnf(i_0_1876, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|k2_struct_0(esk1_0)!=k1_xboole_0|~m1_subset_1(esk8_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_287, plain, (k7_relset_1(X5,X6,k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),X4)=k8_relset_1(u1_struct_0(X2),u1_struct_0(X1),X3,X4)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_2034, plain, (k9_xtuple_0(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))=k2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~v2_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~v3_tops_2(X3,X2,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_2051, negated_conjecture, (k2_struct_0(X1)=k1_xboole_0|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_pre_topc(esk4_0,esk1_0)|~v3_pre_topc(esk4_0,X1)|~v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_tops_2(esk3_0,X1,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1997, negated_conjecture, (k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|k2_struct_0(esk2_0)!=k1_xboole_0|~v3_pre_topc(esk4_0,esk1_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0))).#trainneg
cnf(i_0_2075, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|k2_struct_0(esk1_0)!=k1_xboole_0)).#trainneg
cnf(i_0_2032, plain, (v4_pre_topc(X1,X2)|v2_struct_0(X2)|v2_struct_0(X3)|~v4_pre_topc(k7_relset_1(X5,X6,X4,X1),X3)|~v3_tops_2(X4,X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X4)|~l1_pre_topc(X2)|~l1_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_1596, negated_conjecture, (r1_tarski(X1,k8_relat_1(esk3_0,k7_relset_1(X2,X3,esk3_0,X1)))|~r1_tarski(X1,k1_relset_1(u1_struct_0(esk1_0),esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2071, plain, (v1_relat_1(X1)|~r1_tarski(X1,k10_xtuple_0(k1_xboole_0)))).#trainneg
cnf(i_0_1685, plain, (k1_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))=k2_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v3_tops_2(X3,X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1047, plain, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|v2_struct_0(X1)|v2_struct_0(X2)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))!=k2_struct_0(X1)|~v3_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk8_3(X2,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))),X2)|~v2_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1884, plain, (v4_pre_topc(k8_relset_1(u1_struct_0(X2),u1_struct_0(X1),X3,X4),X2)|v2_struct_0(X1)|v2_struct_0(X2)|~v4_pre_topc(X4,X1)|~v3_tops_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),X1,X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1073, negated_conjecture, (v3_pre_topc(X3,esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(k8_relset_1(X1,X2,esk3_0,X3),esk1_0)|~m1_subset_1(k8_relset_1(X1,X2,esk3_0,X3),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_relset_1(u1_struct_0(esk2_0),esk3_0))))).#trainneg
cnf(i_0_404, negated_conjecture, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(esk2_0),esk3_0),esk2_0,X1)|v3_tops_2(esk3_0,esk1_0,esk2_0)|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~m1_subset_1(esk8_3(esk2_0,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(esk2_0),esk3_0)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_pre_topc(X1)|~v3_pre_topc(esk8_3(esk2_0,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(esk2_0),esk3_0)),esk1_0)|~m1_subset_1(esk8_3(esk2_0,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(esk2_0),esk3_0)),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1880, plain, (v4_pre_topc(k7_relset_1(X5,X6,X3,X4),X2)|v2_struct_0(X1)|v2_struct_0(X2)|~v4_pre_topc(X4,X1)|~v3_tops_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_1445, negated_conjecture, (k8_relat_1(esk3_0,k7_relat_1(esk3_0,k2_struct_0(esk2_0)))=u1_struct_0(esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(k8_relat_1(esk3_0,k7_relat_1(esk3_0,k2_struct_0(esk2_0))))))).#trainneg
cnf(i_0_1295, plain, (k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),X4)=k7_relat_1(X3,X4)|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))!=k2_struct_0(X2)|~l1_struct_0(X2)|~l1_struct_0(X1)|~v2_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1256, negated_conjecture, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),esk2_0,X1)|k2_struct_0(esk1_0)=k1_xboole_0|v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),u1_struct_0(esk2_0),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0))|~l1_pre_topc(X1)|~v3_pre_topc(esk8_3(esk2_0,X1,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)),esk1_0)|~m1_subset_1(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~m1_subset_1(esk8_3(esk2_0,X1,k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_2(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_497, plain, (v4_pre_topc(esk7_3(X1,X2,X3),X4)|v2_struct_0(X1)|v2_struct_0(X4)|v4_pre_topc(esk7_3(X1,X2,X3),X2)|v3_tops_2(X3,X1,X2)|v2_struct_0(X2)|~v3_tops_2(X3,X1,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X4))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk7_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X4)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X4))|~v1_funct_1(X3)|~l1_pre_topc(X4)|~l1_pre_topc(X1)|k1_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v2_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1411, plain, (v3_pre_topc(k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),X4),X2)|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))!=k2_struct_0(X2)|k2_struct_0(X2)!=k1_xboole_0|~v5_pre_topc(X3,X2,X1)|~v3_pre_topc(X4,X1)|~v2_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_2139, plain, (v4_pre_topc(X1,X2)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~v3_tops_2(X5,X2,X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X5,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X5)|~l1_pre_topc(X2)|~l1_pre_topc(X3)|~v4_pre_topc(X1,X4)|~v3_tops_2(X5,X4,X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X4)))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X3))|~l1_pre_topc(X4))).#trainneg
cnf(i_0_2161, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k10_xtuple_0(k1_xboole_0))))).#trainneg
cnf(i_0_2162, plain, (v1_relat_1(k10_xtuple_0(k1_xboole_0)))).#trainneg
cnf(i_0_2154, negated_conjecture, (v1_relat_1(k8_relat_1(esk3_0,k7_relat_1(esk3_0,k10_xtuple_0(k1_xboole_0)))))).#trainneg
cnf(i_0_1979, negated_conjecture, (k8_relat_1(k1_xboole_0,X1)=u1_struct_0(esk1_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k8_relat_1(k1_xboole_0,X1))))).#trainneg
cnf(i_0_1459, plain, (v1_funct_2(k1_xboole_0,X1,X2)|k1_relset_1(X3,k1_xboole_0)!=X1)).#trainneg
cnf(i_0_869, negated_conjecture, (m1_subset_1(k7_relset_1(X1,k1_relset_1(u1_struct_0(esk1_0),esk3_0),X2,X3),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_relset_1(u1_struct_0(esk1_0),esk3_0)))))).#trainneg
cnf(i_0_2166, plain, (k9_xtuple_0(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))=k2_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~v3_tops_2(X3,X2,X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2127, negated_conjecture, (v3_pre_topc(k8_relat_1(esk3_0,X1),esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|k2_struct_0(esk1_0)!=k1_xboole_0|~v3_pre_topc(X1,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_2251, plain, (v3_pre_topc(k7_relat_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),X4),X2)|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|k2_struct_0(X2)!=k1_xboole_0|~v5_pre_topc(X3,X2,X1)|~v3_pre_topc(X4,X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_2192, negated_conjecture, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(k2_tops_2(u1_struct_0(X1),u1_struct_0(esk2_0),esk3_0),esk2_0,X1)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(esk8_3(esk2_0,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(esk2_0),esk3_0)),esk1_0)|~m1_subset_1(esk8_3(esk2_0,X1,k2_tops_2(u1_struct_0(X1),u1_struct_0(esk2_0),esk3_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_pre_topc(X1)|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(esk2_0),esk3_0),u1_struct_0(esk2_0),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(esk2_0),esk3_0)))).#trainneg
cnf(i_0_2233, negated_conjecture, (v4_pre_topc(esk7_3(esk1_0,X1,esk3_0),X1)|v4_pre_topc(esk7_3(esk1_0,X1,esk3_0),esk2_0)|v3_tops_2(esk3_0,esk1_0,X1)|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X1),esk3_0)!=k2_struct_0(X1)|~v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~m1_subset_1(esk7_3(esk1_0,X1,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1509, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(k7_relset_1(X5,X6,k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),X4),X2)|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))!=k2_struct_0(X2)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~v5_pre_topc(X3,X2,X1)|~v3_pre_topc(X4,X1)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_1743, plain, (v1_relat_1(X1)|~v2_funct_1(k1_xboole_0)|~m1_subset_1(X1,k1_zfmisc_1(k9_xtuple_0(k1_xboole_0))))).#trainneg
cnf(i_0_1634, negated_conjecture, (k8_relat_1(esk3_0,k7_relat_1(esk3_0,u1_struct_0(esk2_0)))=u1_struct_0(esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(k8_relat_1(esk3_0,k7_relat_1(esk3_0,u1_struct_0(esk2_0))))))).#trainneg
cnf(i_0_2056, negated_conjecture, (k8_relat_1(k1_xboole_0,X1)=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(k8_relat_1(k1_xboole_0,X1))))).#trainneg
cnf(i_0_1037, negated_conjecture, (m1_subset_1(k7_relset_1(X1,k2_relset_1(u1_struct_0(esk2_0),esk3_0),X2,X3),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k2_relset_1(u1_struct_0(esk2_0),esk3_0)))))).#trainneg
cnf(i_0_1388, plain, (v4_relat_1(k7_relset_1(X1,X2,k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5)),X6),X4)|~m1_subset_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5)),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_307, plain, (m1_subset_1(k7_relat_1(esk5_2(X1,X2),X3),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_309, plain, (m1_subset_1(k7_relat_1(k8_relset_1(k2_zfmisc_1(X1,X2),X3,X4,X5),X6),k1_zfmisc_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_2259, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X1),esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(X1,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_2265, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_relset_1(X2,k1_xboole_0))))).#trainneg
cnf(i_0_2306, negated_conjecture, (v3_pre_topc(X1,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|k2_struct_0(esk1_0)!=k1_xboole_0|~v3_pre_topc(k7_relat_1(esk3_0,X1),esk2_0)|~r1_tarski(X1,k1_relset_1(u1_struct_0(esk1_0),esk3_0)))).#trainneg
cnf(i_0_2272, plain, (v1_relat_1(k2_relset_1(X1,k1_xboole_0)))).#trainneg
cnf(i_0_310, plain, (m1_subset_1(k7_relat_1(k2_tops_2(X1,X2,X3),X4),k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_2281, negated_conjecture, (k8_relat_1(k1_xboole_0,k7_relat_1(k1_xboole_0,u1_struct_0(esk1_0)))=u1_struct_0(esk1_0)|~r1_tarski(u1_struct_0(esk1_0),k9_xtuple_0(k1_xboole_0)))).#trainneg
cnf(i_0_2554, negated_conjecture, (v1_relat_1(u1_struct_0(esk1_0))|~r1_tarski(u1_struct_0(esk1_0),k9_xtuple_0(k1_xboole_0)))).#trainneg
cnf(i_0_311, plain, (m1_subset_1(k7_relat_1(k7_relset_1(X1,k2_zfmisc_1(X2,X3),X4,X5),X6),k1_zfmisc_1(X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_312, plain, (m1_subset_1(k8_relat_1(esk5_2(X1,X2),X3),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2463, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|v3_pre_topc(k7_relat_1(esk3_0,X1),esk2_0)|k2_struct_0(esk2_0)=k1_xboole_0|~v3_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2471, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(k7_relat_1(esk3_0,X1),esk2_0)|~r1_tarski(X1,k1_relset_1(u1_struct_0(esk1_0),esk3_0)))).#trainneg
cnf(i_0_2680, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0|v3_pre_topc(X1,esk2_0)|~v3_pre_topc(k8_relat_1(esk3_0,X1),esk1_0)|~r1_tarski(X1,k2_relset_1(u1_struct_0(esk2_0),esk3_0)))).#trainneg
cnf(i_0_2349, negated_conjecture, (k8_relat_1(k1_xboole_0,k7_relat_1(k1_xboole_0,u1_struct_0(esk2_0)))=u1_struct_0(esk2_0)|~r1_tarski(u1_struct_0(esk2_0),k9_xtuple_0(k1_xboole_0)))).#trainneg
cnf(i_0_314, plain, (m1_subset_1(k8_relat_1(k8_relset_1(k2_zfmisc_1(X1,X2),X3,X4,X5),X6),k1_zfmisc_1(X1))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_2558, negated_conjecture, (v1_relat_1(u1_struct_0(esk1_0))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k9_xtuple_0(k1_xboole_0))))).#trainneg
cnf(i_0_2555, negated_conjecture, (u1_struct_0(esk2_0)=u1_struct_0(esk1_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_tarski(u1_struct_0(esk1_0),k9_xtuple_0(k1_xboole_0)))).#trainneg
cnf(i_0_1390, plain, (v5_relat_1(k7_relset_1(X1,X2,k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5)),X6),X5)|~m1_subset_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5)),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_315, plain, (m1_subset_1(k8_relat_1(k2_tops_2(X1,X2,X3),X4),k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_316, plain, (m1_subset_1(k8_relat_1(k7_relset_1(X1,k2_zfmisc_1(X2,X3),X4,X5),X6),k1_zfmisc_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_2683, negated_conjecture, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(esk3_0,esk1_0,X1)|k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(X1))|~l1_pre_topc(X1)|~v3_pre_topc(esk8_3(esk1_0,X1,esk3_0),esk2_0)|~m1_subset_1(esk8_3(esk1_0,X1,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_2710, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X1),esk1_0)|~r1_tarski(k8_relat_1(esk3_0,X1),k1_relset_1(u1_struct_0(esk1_0),esk3_0))|~v3_pre_topc(X1,esk2_0)|~r1_tarski(X1,k2_relset_1(u1_struct_0(esk2_0),esk3_0)))).#trainneg
cnf(i_0_2472, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X3,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(k7_relset_1(X1,X2,esk3_0,X3),esk2_0)|~m1_subset_1(k7_relset_1(X1,X2,esk3_0,X3),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r1_tarski(X3,k1_relset_1(u1_struct_0(esk1_0),esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2461, negated_conjecture, (v4_pre_topc(k7_relat_1(esk3_0,X1),esk2_0)|k2_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0|~v4_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2717, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k7_relat_1(esk3_0,X1),esk2_0)|~r1_tarski(k7_relat_1(esk3_0,X1),k2_relset_1(u1_struct_0(esk2_0),esk3_0))|~v3_pre_topc(X1,esk1_0)|~r1_tarski(X1,k1_relset_1(u1_struct_0(esk1_0),esk3_0)))).#trainneg
cnf(i_0_2738, negated_conjecture, (u1_struct_0(esk2_0)=u1_struct_0(esk1_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r1_tarski(u1_struct_0(esk2_0),k9_xtuple_0(k1_xboole_0)))).#trainneg
cnf(i_0_634, plain, (v4_relat_1(k8_relat_1(k1_xboole_0,X1),X2))).#trainneg
cnf(i_0_2704, negated_conjecture, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(u1_struct_0(esk2_0)))|k2_struct_0(esk1_0)=k1_xboole_0|m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2791, negated_conjecture, (u1_struct_0(esk2_0)=u1_struct_0(esk1_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k9_xtuple_0(k1_xboole_0))))).#trainneg
cnf(i_0_317, plain, (k2_relset_1(X1,k7_relset_1(X2,k2_zfmisc_1(X3,X1),X4,X5))=k10_xtuple_0(k7_relset_1(X2,k2_zfmisc_1(X3,X1),X4,X5))|~v1_relat_1(k7_relset_1(X2,k2_zfmisc_1(X3,X1),X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k2_zfmisc_1(X3,X1)))))).#trainneg
cnf(i_0_2737, negated_conjecture, (v1_relat_1(u1_struct_0(esk2_0))|~r1_tarski(u1_struct_0(esk2_0),k9_xtuple_0(k1_xboole_0)))).#trainneg
cnf(i_0_2953, negated_conjecture, (u1_struct_0(esk2_0)=u1_struct_0(esk1_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(k9_xtuple_0(k1_xboole_0))))).#trainneg
cnf(i_0_1382, negated_conjecture, (v3_tops_2(esk3_0,esk1_0,esk2_0)|m1_subset_1(k7_relset_1(X1,X2,k2_zfmisc_1(X3,k2_struct_0(esk2_0)),X4),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(k2_zfmisc_1(X3,k2_struct_0(esk2_0)),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1383, negated_conjecture, (v3_tops_2(esk3_0,esk1_0,esk2_0)|m1_subset_1(k7_relset_1(X1,X2,k2_zfmisc_1(X3,k2_struct_0(esk1_0)),X4),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_zfmisc_1(X3,k2_struct_0(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_318, plain, (v5_relat_1(k7_relat_1(X4,X5),X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_319, plain, (v5_relat_1(k7_relset_1(X6,X7,X4,X5),X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))).#trainneg
cnf(i_0_2913, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0|v5_pre_topc(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk8_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_2911, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(esk4_0,esk1_0)|~r1_tarski(esk4_0,k2_struct_0(esk1_0)))).#trainneg
cnf(i_0_2943, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|v4_pre_topc(X1,esk2_0)|~v4_pre_topc(k8_relat_1(esk3_0,X1),esk1_0)|~r1_tarski(X1,k2_relset_1(u1_struct_0(esk2_0),esk3_0)))).#trainneg
cnf(i_0_2899, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|u1_struct_0(esk2_0)=k1_xboole_0|m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2918, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X1),esk1_0)|~r1_tarski(k8_relat_1(esk3_0,X1),k2_struct_0(esk1_0))|~r1_tarski(X1,k2_relset_1(u1_struct_0(esk2_0),esk3_0))|~v3_pre_topc(X1,esk2_0))).#trainneg
cnf(i_0_2916, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X1),esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0|~r1_tarski(k8_relat_1(esk3_0,X1),u1_struct_0(esk1_0))|~r1_tarski(X1,k2_relset_1(u1_struct_0(esk2_0),esk3_0))|~v3_pre_topc(X1,esk2_0))).#trainneg
cnf(i_0_2185, negated_conjecture, (v3_pre_topc(X1,esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(k8_relat_1(esk3_0,X1),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_relset_1(u1_struct_0(esk2_0),esk3_0)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_636, plain, (v5_relat_1(k8_relat_1(k1_xboole_0,X1),X2))).#trainneg
cnf(i_0_1843, plain, (v1_relat_1(k1_relset_1(X1,k1_xboole_0))|~v2_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_2912, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(esk4_0,esk1_0)|~m1_subset_1(esk4_0,k1_zfmisc_1(k1_relset_1(u1_struct_0(esk1_0),esk3_0))))).#trainneg
cnf(i_0_2904, negated_conjecture, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|u1_struct_0(esk2_0)=k1_xboole_0|m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_321, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(k7_relat_1(X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,k2_zfmisc_1(X2,X1))))|~v1_funct_2(k7_relat_1(X4,X5),X2,X1))).#trainneg
cnf(i_0_322, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(k7_relset_1(X6,X7,X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,k2_zfmisc_1(X2,X1))))|~v1_funct_2(k7_relset_1(X6,X7,X4,X5),X2,X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))).#trainneg
cnf(i_0_3038, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(esk3_0,esk1_0,esk2_0))).#trainneg
cnf(i_0_3042, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0|v3_pre_topc(esk4_0,esk1_0)|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_struct_0(esk1_0))))).#trainneg
cnf(i_0_3065, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X1),esk1_0)|~r1_tarski(X1,k2_relset_1(u1_struct_0(esk2_0),esk3_0))|~v3_pre_topc(X1,esk2_0))).#trainneg
cnf(i_0_3041, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|u1_struct_0(esk1_0)=k1_xboole_0|v3_pre_topc(esk4_0,esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0|~r1_tarski(esk4_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3040, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|u1_struct_0(esk1_0)=k1_xboole_0|v3_pre_topc(esk4_0,esk1_0)|~r1_tarski(esk4_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3070, negated_conjecture, (v3_pre_topc(X1,esk2_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(k8_relat_1(esk3_0,X1),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_relset_1(u1_struct_0(esk2_0),esk3_0))))).#trainneg
cnf(i_0_2960, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|m1_subset_1(k1_xboole_0,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_3049, negated_conjecture, (u1_struct_0(esk1_0)=esk4_0|u1_struct_0(esk2_0)=k1_xboole_0|u1_struct_0(esk1_0)=k1_xboole_0|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(esk4_0)))).#trainneg
cnf(i_0_2984, negated_conjecture, (v1_relat_1(u1_struct_0(esk2_0))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(k9_xtuple_0(k1_xboole_0))))).#trainneg
cnf(i_0_326, plain, (v4_pre_topc(k8_relset_1(X4,X5,X3,esk7_3(X1,X2,X3)),X1)|v4_pre_topc(esk7_3(X1,X2,X3),X2)|v3_tops_2(X3,X1,X2)|v2_struct_0(X1)|v2_struct_0(X2)|k1_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),X3)!=k2_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_328, plain, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(X2,X3,X1)|~v3_pre_topc(k8_relset_1(X4,X5,X2,esk8_3(X3,X1,X2)),X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_pre_topc(X1)|~l1_pre_topc(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_2928, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~r1_tarski(X1,k2_struct_0(esk1_0))|~v3_pre_topc(k7_relset_1(X2,X3,esk3_0,X1),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_3056, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X1),esk1_0)|~r1_tarski(X1,k2_relset_1(u1_struct_0(esk2_0),esk3_0))|~v3_pre_topc(X1,esk2_0)|~m1_subset_1(k8_relat_1(esk3_0,X1),k1_zfmisc_1(k2_struct_0(esk1_0))))).#trainneg
cnf(i_0_2929, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|u1_struct_0(esk2_0)=k1_xboole_0|~r1_tarski(X1,u1_struct_0(esk1_0))|~v3_pre_topc(k7_relset_1(X2,X3,esk3_0,X1),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2947, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0|v3_pre_topc(k7_relat_1(esk3_0,X1),esk2_0)|~r1_tarski(k7_relat_1(esk3_0,X1),k2_struct_0(esk2_0))|~r1_tarski(X1,k1_relset_1(u1_struct_0(esk1_0),esk3_0))|~v3_pre_topc(X1,esk1_0))).#trainneg
cnf(i_0_2940, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k1_relset_1(u1_struct_0(esk1_0),esk3_0),esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(k7_relset_1(X1,X2,esk3_0,k1_relset_1(u1_struct_0(esk1_0),esk3_0)),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2790, negated_conjecture, (u1_struct_0(esk2_0)=u1_struct_0(esk1_0)|~r1_tarski(u1_struct_0(esk1_0),k1_relset_1(X1,k1_xboole_0))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1627, plain, (v4_relat_1(k7_relat_1(k1_xboole_0,X1),X2))).#trainneg
cnf(i_0_2666, negated_conjecture, (k8_relat_1(esk3_0,k7_relat_1(esk3_0,k2_struct_0(esk2_0)))=u1_struct_0(esk2_0)|k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(k8_relat_1(esk3_0,k7_relat_1(esk3_0,k2_struct_0(esk2_0))))))).#trainneg
cnf(i_0_329, plain, (v5_pre_topc(X1,X2,X3)|k2_struct_0(X2)!=k1_xboole_0|~v3_pre_topc(k8_relset_1(X4,X5,X1,esk8_3(X2,X3,X1)),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_3084, negated_conjecture, (u1_struct_0(esk1_0)=esk4_0|u1_struct_0(esk2_0)=k1_xboole_0|m1_subset_1(k1_xboole_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(esk4_0)))).#trainneg
cnf(i_0_332, plain, (k8_relset_1(X5,X6,k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),X4)=k7_relset_1(u1_struct_0(X2),u1_struct_0(X1),X3,X4)|k2_relset_1(u1_struct_0(X1),X3)!=k2_struct_0(X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_3140, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(esk4_0,esk1_0)|~m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3129, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|k2_relset_1(u1_struct_0(esk2_0),esk3_0)!=k2_struct_0(esk2_0)|~v3_tops_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_struct_0(esk1_0))))).#trainneg
cnf(i_0_3176, plain, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(X2,X3,X1)|k2_struct_0(X4)=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_pre_topc(X1)|~l1_pre_topc(X3)|~v5_pre_topc(X2,X3,X4)|~v3_pre_topc(esk8_3(X3,X1,X2),X4)|~m1_subset_1(esk8_3(X3,X1,X2),k1_zfmisc_1(u1_struct_0(X4)))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X4))|~l1_pre_topc(X4))).#trainneg
cnf(i_0_3264, plain, (k2_struct_0(X1)=k1_xboole_0|v3_pre_topc(k7_relset_1(u1_struct_0(X3),u1_struct_0(X4),X5,X6),X2)|~v5_pre_topc(k2_tops_2(u1_struct_0(X3),u1_struct_0(X4),X5),X2,X1)|~v3_pre_topc(X6,X1)|~m1_subset_1(k2_tops_2(u1_struct_0(X3),u1_struct_0(X4),X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(u1_struct_0(X3),u1_struct_0(X4),X5),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X3),u1_struct_0(X4),X5))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|k2_relset_1(u1_struct_0(X4),X5)!=k2_struct_0(X4)|~l1_struct_0(X4)|~l1_struct_0(X3)|~v2_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4))))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X3)))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X4))|~v1_funct_1(X5))).#trainneg
cnf(i_0_3244, plain, (v5_pre_topc(X1,X2,X3)|k2_struct_0(X2)!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X4))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v5_pre_topc(X1,X2,X4)|~v3_pre_topc(esk8_3(X2,X3,X1),X4)|~m1_subset_1(esk8_3(X2,X3,X1),k1_zfmisc_1(u1_struct_0(X4)))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X4))|~l1_pre_topc(X4))).#trainneg
cnf(i_0_3258, plain, (v3_pre_topc(k7_relset_1(u1_struct_0(X3),u1_struct_0(X4),X5,X6),X1)|k2_struct_0(X1)!=k1_xboole_0|~v5_pre_topc(k2_tops_2(u1_struct_0(X3),u1_struct_0(X4),X5),X1,X2)|~v3_pre_topc(X6,X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X3),u1_struct_0(X4),X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(k2_tops_2(u1_struct_0(X3),u1_struct_0(X4),X5),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X3),u1_struct_0(X4),X5))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|k2_relset_1(u1_struct_0(X4),X5)!=k2_struct_0(X4)|~l1_struct_0(X4)|~l1_struct_0(X3)|~v2_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4))))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X3)))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X4))|~v1_funct_1(X5))).#trainneg
cnf(i_0_3051, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|u1_struct_0(esk1_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X1),esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0|~r1_tarski(k8_relat_1(esk3_0,X1),u1_struct_0(esk1_0))|~r1_tarski(X1,k2_relset_1(u1_struct_0(esk2_0),esk3_0))|~v3_pre_topc(X1,esk2_0))).#trainneg
cnf(i_0_3017, plain, (v5_relat_1(k7_relat_1(k1_xboole_0,X1),X2))).#trainneg
cnf(i_0_2965, negated_conjecture, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|m1_subset_1(k1_xboole_0,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_3234, negated_conjecture, (k8_relat_1(esk3_0,k7_relat_1(esk3_0,u1_struct_0(esk2_0)))=u1_struct_0(esk2_0)|u1_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(k8_relat_1(esk3_0,k7_relat_1(esk3_0,u1_struct_0(esk2_0))))))).#trainneg
cnf(i_0_3231, plain, (v4_relat_1(X1,X2)|~r1_tarski(X1,k10_xtuple_0(k1_xboole_0)))).#trainneg
cnf(i_0_3150, negated_conjecture, (u1_struct_0(esk1_0)=esk4_0|u1_struct_0(esk1_0)=k1_xboole_0|m1_subset_1(k1_xboole_0,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(esk4_0)))).#trainneg
cnf(i_0_1469, negated_conjecture, (v3_tops_2(esk3_0,esk1_0,esk2_0)|m1_subset_1(k7_relset_1(X4,X5,X2,X3),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k2_struct_0(esk2_0))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_3050, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|u1_struct_0(esk1_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X1),esk1_0)|~r1_tarski(k8_relat_1(esk3_0,X1),u1_struct_0(esk1_0))|~r1_tarski(X1,k2_relset_1(u1_struct_0(esk2_0),esk3_0))|~v3_pre_topc(X1,esk2_0))).#trainneg
cnf(i_0_2670, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|~r1_tarski(X1,k2_struct_0(esk1_0)))).#trainneg
cnf(i_0_3215, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k2_struct_0(esk1_0),esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(k7_relset_1(X1,X2,esk3_0,k2_struct_0(esk1_0)),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3170, plain, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(k2_tops_2(u1_struct_0(X2),u1_struct_0(X3),X4),X5,X1)|~v3_pre_topc(k7_relset_1(u1_struct_0(X2),u1_struct_0(X3),X4,esk8_3(X5,X1,k2_tops_2(u1_struct_0(X2),u1_struct_0(X3),X4))),X5)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X3),X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X1))))|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X3),X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X3),X4),u1_struct_0(X5),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X3),X4))|~l1_pre_topc(X1)|~l1_pre_topc(X5)|k2_relset_1(u1_struct_0(X3),X4)!=k2_struct_0(X3)|~l1_struct_0(X3)|~l1_struct_0(X2)|~v2_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(esk8_3(X5,X1,k2_tops_2(u1_struct_0(X2),u1_struct_0(X3),X4)),k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X4))).#trainneg
cnf(i_0_2457, negated_conjecture, (v4_pre_topc(k8_relat_1(esk3_0,X1),esk1_0)|k2_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0|~v4_pre_topc(X1,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_3177, plain, (k2_struct_0(X1)=k1_xboole_0|v5_pre_topc(X2,X3,X1)|k2_struct_0(X4)=k1_xboole_0|v2_struct_0(X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_pre_topc(X1)|~l1_pre_topc(X3)|~v3_pre_topc(esk8_3(X3,X1,X2),X4)|~v3_tops_2(X2,X3,X4)|~m1_subset_1(esk8_3(X3,X1,X2),k1_zfmisc_1(u1_struct_0(X4)))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X4))|~l1_pre_topc(X4))).#trainneg
cnf(i_0_2473, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k8_relset_1(X2,X3,esk3_0,X1),esk1_0)|v3_tops_2(esk3_0,esk1_0,esk2_0)|~v3_pre_topc(X1,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2952, negated_conjecture, (u1_struct_0(esk2_0)=u1_struct_0(esk1_0)|~r1_tarski(u1_struct_0(esk2_0),k1_relset_1(X1,k1_xboole_0))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_3349, plain, (v4_relat_1(k10_xtuple_0(k1_xboole_0),X1))).#trainneg
cnf(i_0_3305, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|u1_struct_0(esk1_0)=k1_xboole_0|~m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2969, negated_conjecture, (u1_struct_0(esk2_0)=u1_struct_0(esk1_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_relset_1(X1,k1_xboole_0))))).#trainneg
cnf(i_0_3336, negated_conjecture, (u1_struct_0(esk1_0)=esk4_0|m1_subset_1(k1_xboole_0,k1_zfmisc_1(u1_struct_0(esk2_0)))|m1_subset_1(k1_xboole_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(esk4_0)))).#trainneg
cnf(i_0_1477, negated_conjecture, (v3_tops_2(esk3_0,esk1_0,esk2_0)|m1_subset_1(k7_relset_1(X4,X5,X2,X3),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k2_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_3416, negated_conjecture, (v3_tops_2(esk3_0,esk1_0,X1)|v2_struct_0(X1)|k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|k1_relset_1(u1_struct_0(esk1_0),esk3_0)!=k2_struct_0(esk1_0)|k2_relset_1(u1_struct_0(X1),esk3_0)!=k2_struct_0(X1)|~v4_pre_topc(esk7_3(esk1_0,X1,esk3_0),X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(X1))|~l1_pre_topc(X1)|~v4_pre_topc(esk7_3(esk1_0,X1,esk3_0),esk2_0)|~m1_subset_1(esk7_3(esk1_0,X1,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_3420, negated_conjecture, (k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|v4_pre_topc(X1,esk1_0)|~v4_pre_topc(k7_relat_1(esk3_0,X1),esk2_0)|~r1_tarski(X1,k1_relset_1(u1_struct_0(esk1_0),esk3_0)))).#trainneg
cnf(i_0_3374, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X1),esk1_0)|~r1_tarski(X1,k2_relset_1(u1_struct_0(esk2_0),esk3_0))|~v3_pre_topc(X1,esk2_0))).#trainneg
cnf(i_0_2682, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0|v3_pre_topc(k7_relset_1(X2,X3,esk3_0,X1),esk2_0)|~v3_pre_topc(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2466, negated_conjecture, (v4_pre_topc(esk7_3(esk1_0,X1,esk3_0),esk2_0)|v4_pre_topc(esk7_3(esk1_0,X1,esk3_0),X1)|v3_tops_2(esk3_0,esk1_0,X1)|v2_struct_0(X1)|k2_struct_0(esk2_0)=k1_xboole_0|k2_struct_0(esk1_0)=k1_xboole_0|k2_relset_1(u1_struct_0(X1),esk3_0)!=k2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~m1_subset_1(esk7_3(esk1_0,X1,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_475, plain, (m1_subset_1(k7_relat_1(k1_xboole_0,X1),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_2662, negated_conjecture, (m1_subset_1(k8_relat_1(esk3_0,k7_relat_1(esk3_0,k2_struct_0(esk2_0))),k1_zfmisc_1(u1_struct_0(esk2_0)))|k2_struct_0(esk1_0)=k1_xboole_0|k2_struct_0(esk2_0)=k1_xboole_0)).#trainneg
cnf(i_0_2985, negated_conjecture, (u1_struct_0(esk2_0)=u1_struct_0(esk1_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(k1_relset_1(X1,k1_xboole_0))))).#trainneg
cnf(i_0_3333, plain, (v5_relat_1(X1,X2)|~r1_tarski(X1,k10_xtuple_0(k1_xboole_0)))).#trainneg
cnf(i_0_351, plain, (m1_subset_1(X1,k1_zfmisc_1(k8_relat_1(X2,k7_relset_1(X3,X4,X2,X1))))|~r1_tarski(X1,k9_xtuple_0(X2))|~v1_relat_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_3780, negated_conjecture, (k1_xboole_0=esk3_0|u1_struct_0(esk1_0)=k1_xboole_0)).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 39
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 98
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 95
# Processed clauses                    : 945
# ...of these trivial                  : 11
# ...subsumed                          : 325
# ...remaining for further processing  : 609
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 69
# Backward-rewritten                   : 139
# Generated clauses                    : 4161
# ...of the previous two non-trivial   : 3947
# Contextual simplify-reflections      : 277
# Paramodulations                      : 4154
# Factorizations                       : 0
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 398
#    Positive orientable unit clauses  : 47
#    Positive unorientable unit clauses: 8
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 340
# Current number of unprocessed clauses: 3020
# ...number of literals in the above   : 23868
# Current number of archived formulas  : 0
# Current number of archived clauses   : 817
# Proof object given clauses           : 125
# Proof search given clauses           : 609
# Clause-clause subsumption calls (NU) : 172103
# Rec. Clause-clause subsumption calls : 8378
# Non-unit clause-clause subsumptions  : 655
# Unit Clause-clause subsumption calls : 1560
# Rewrite failures with RHS unbound    : 39
# BW rewrite match attempts            : 66
# BW rewrite match successes           : 13
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 160462

# -------------------------------------------------
# User time                : 2.306 s
# System time              : 0.156 s
# Total time               : 2.462 s
# Maximum resident set size: 48112 pages
