2020-06-03 00:03:05.686101: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_35
2020-06-03 00:03:05.737879: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-06-03 00:03:05.866929: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-06-03 00:03:06.328083: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-06-03 00:03:09.066846: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 3380740 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_35' loaded (query=512; context=768; weigths=real; len_mult=0.000000).
# Preprocessing time       : 4.351 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(spc7_arithm, axiom, ![X1, X2, X3]:(((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))&v1_xcmplx_0(X3))=>k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', spc7_arithm)).
fof(commutativity_k3_xcmplx_0, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', commutativity_k3_xcmplx_0)).
fof(cc17_membered, axiom, ![X1]:(v6_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v7_ordinal1(X2))), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', cc17_membered)).
fof(fc6_membered, axiom, v6_membered(k4_ordinal1), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', fc6_membered)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', redefinition_k5_numbers)).
fof(t3_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k3_xcmplx_0(np__1,X1)=X1), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', t3_arithm)).
fof(cc2_xcmplx_0, axiom, ![X1]:(v7_ordinal1(X1)=>v1_xcmplx_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', cc2_xcmplx_0)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', spc1_numerals)).
fof(t28_sin_cos3, conjecture, ![X1]:(m1_subset_1(X1,k5_numbers)=>k15_sin_cos(k5_binop_2(k11_binop_2(k11_binop_2(np__2,k32_sin_cos),X1),k7_complex1))=np__1), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', t28_sin_cos3)).
fof(dt_k11_binop_2, axiom, ![X1, X2]:((v1_xreal_0(X1)&v1_xreal_0(X2))=>m1_subset_1(k11_binop_2(X1,X2),k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', dt_k11_binop_2)).
fof(redefinition_k11_binop_2, axiom, ![X1, X2]:((v1_xreal_0(X1)&v1_xreal_0(X2))=>k11_binop_2(X1,X2)=k3_xcmplx_0(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', redefinition_k11_binop_2)).
fof(fc3_xcmplx_0, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>v1_xcmplx_0(k3_xcmplx_0(X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', fc3_xcmplx_0)).
fof(commutativity_k11_binop_2, axiom, ![X1, X2]:((v1_xreal_0(X1)&v1_xreal_0(X2))=>k11_binop_2(X1,X2)=k11_binop_2(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', commutativity_k11_binop_2)).
fof(cc14_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xreal_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', cc14_membered)).
fof(cc21_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v3_membered(X2))), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', cc21_membered)).
fof(dt_k1_real_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>m1_subset_1(k1_real_1(X1),k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', dt_k1_real_1)).
fof(dt_k31_sin_cos, axiom, v1_xreal_0(k31_sin_cos), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', dt_k31_sin_cos)).
fof(redefinition_k32_sin_cos, axiom, k32_sin_cos=k31_sin_cos, file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', redefinition_k32_sin_cos)).
fof(dt_k5_numbers, axiom, m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', dt_k5_numbers)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', fc3_membered)).
fof(involutiveness_k1_real_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>k1_real_1(k1_real_1(X1))=X1), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', involutiveness_k1_real_1)).
fof(spc2_numerals, axiom, (((v2_xxreal_0(np__2)&m2_subset_1(np__2,k1_numbers,k5_numbers))&m1_subset_1(np__2,k5_numbers))&m1_subset_1(np__2,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', spc2_numerals)).
fof(cc12_membered, axiom, ![X1]:(v1_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xcmplx_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', cc12_membered)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', t6_boole)).
fof(redefinition_k7_complex1, axiom, k7_complex1=k1_xcmplx_0, file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', redefinition_k7_complex1)).
fof(commutativity_k5_binop_2, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k5_binop_2(X1,X2)=k5_binop_2(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', commutativity_k5_binop_2)).
fof(t10_sin_cos2, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v7_ordinal1(X2)=>k1_seq_1(k16_sin_cos,X1)=k1_seq_1(k16_sin_cos,k7_real_1(k8_real_1(k8_real_1(np__2,k32_sin_cos),X2),X1)))), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', t10_sin_cos2)).
fof(redefinition_k8_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', redefinition_k8_real_1)).
fof(dt_k32_sin_cos, axiom, m1_subset_1(k32_sin_cos,k1_numbers), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', dt_k32_sin_cos)).
fof(cc5_membered, axiom, ![X1]:(v3_membered(X1)=>v1_membered(X1)), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', cc5_membered)).
fof(redefinition_k7_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k7_real_1(X1,X2)=k2_xcmplx_0(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', redefinition_k7_real_1)).
fof(commutativity_k7_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k7_real_1(X1,X2)=k7_real_1(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', commutativity_k7_real_1)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', redefinition_k6_numbers)).
fof(fc1_xcmplx_0, axiom, v1_xcmplx_0(k1_xcmplx_0), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', fc1_xcmplx_0)).
fof(redefinition_k5_binop_2, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k5_binop_2(X1,X2)=k3_xcmplx_0(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', redefinition_k5_binop_2)).
fof(t11_sin_cos2, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v7_ordinal1(X2)=>k1_seq_1(k19_sin_cos,X1)=k1_seq_1(k19_sin_cos,k7_real_1(k8_real_1(k8_real_1(np__2,k32_sin_cos),X2),X1)))), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', t11_sin_cos2)).
fof(t1_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k2_xcmplx_0(X1,k6_numbers)=X1), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', t1_arithm)).
fof(spc0_boole, axiom, v1_xboole_0(np__0), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', spc0_boole)).
fof(spc0_numerals, axiom, ((m2_subset_1(np__0,k1_numbers,k5_numbers)&m1_subset_1(np__0,k5_numbers))&m1_subset_1(np__0,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', spc0_numerals)).
fof(t30_sin_cos, axiom, ![X1]:(v1_xreal_0(X1)=>(((k1_seq_1(k19_sin_cos,k6_numbers)=np__1&k1_seq_1(k16_sin_cos,k6_numbers)=k6_numbers)&k1_seq_1(k19_sin_cos,k4_xcmplx_0(X1))=k1_seq_1(k19_sin_cos,X1))&k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1))=k1_real_1(k1_seq_1(k16_sin_cos,X1)))), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', t30_sin_cos)).
fof(t25_sin_cos, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>k15_sin_cos(k3_xcmplx_0(X1,k7_complex1))=k2_xcmplx_0(k21_sin_cos(X1),k3_xcmplx_0(k18_sin_cos(X1),k7_complex1))), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', t25_sin_cos)).
fof(d19_sin_cos, axiom, ![X1]:(v1_xreal_0(X1)=>k20_sin_cos(X1)=k1_seq_1(k19_sin_cos,X1)), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', d19_sin_cos)).
fof(redefinition_k21_sin_cos, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>k21_sin_cos(X1)=k20_sin_cos(X1)), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', redefinition_k21_sin_cos)).
fof(commutativity_k2_xcmplx_0, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', commutativity_k2_xcmplx_0)).
fof(d17_sin_cos, axiom, ![X1]:(v1_xreal_0(X1)=>k17_sin_cos(X1)=k1_seq_1(k16_sin_cos,X1)), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', d17_sin_cos)).
fof(redefinition_k18_sin_cos, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>k18_sin_cos(X1)=k17_sin_cos(X1)), file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', redefinition_k18_sin_cos)).
fof(rqRealMult__k3_xcmplx_0__cr1r0_r0_r0, axiom, k3_xcmplx_0(k1_xcmplx_0,np__0)=np__0, file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', rqRealMult__k3_xcmplx_0__cr1r0_r0_r0)).
fof(rqRealAdd__k2_xcmplx_0__r1_r0_r1, axiom, k2_xcmplx_0(np__1,np__0)=np__1, file('/scratch/mptp/allpred/preds__0.005/t28_sin_cos3', rqRealAdd__k2_xcmplx_0__r1_r0_r1)).
fof(c_0_48, plain, ![X104, X105, X106]:(~v1_xcmplx_0(X104)|~v1_xcmplx_0(X105)|~v1_xcmplx_0(X106)|k3_xcmplx_0(k3_xcmplx_0(X104,X105),X106)=k3_xcmplx_0(X104,k3_xcmplx_0(X105,X106))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc7_arithm])])).
fof(c_0_49, plain, ![X73, X74]:(~v1_xcmplx_0(X73)|~v1_xcmplx_0(X74)|k3_xcmplx_0(X73,X74)=k3_xcmplx_0(X74,X73)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k3_xcmplx_0])])).
fof(c_0_50, plain, ![X12, X13]:(~v6_membered(X12)|(~m1_subset_1(X13,X12)|v7_ordinal1(X13))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc17_membered])])])).
cnf(c_0_51, plain, (v6_membered(k4_ordinal1)), inference(split_conjunct,[status(thm)],[fc6_membered])).
cnf(c_0_52, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
cnf(c_0_53, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_54, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_49])).
fof(c_0_55, plain, ![X48]:(~v1_xcmplx_0(X48)|k3_xcmplx_0(np__1,X48)=X48), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_arithm])])).
fof(c_0_56, plain, ![X52]:(~v7_ordinal1(X52)|v1_xcmplx_0(X52)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_xcmplx_0])])).
cnf(c_0_57, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_58, plain, (m1_subset_1(np__1,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_59, plain, (v6_membered(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_51])).
cnf(c_0_60, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_52])).
cnf(c_0_61, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_53])).
cnf(c_0_62, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_54])).
fof(c_0_63, negated_conjecture, ~(![X1]:(m1_subset_1(X1,k5_numbers)=>k15_sin_cos(k5_binop_2(k11_binop_2(k11_binop_2(np__2,k32_sin_cos),X1),k7_complex1))=np__1)), inference(assume_negation,[status(cth)],[t28_sin_cos3])).
fof(c_0_64, plain, ![X42, X43]:(~v1_xreal_0(X42)|~v1_xreal_0(X43)|m1_subset_1(k11_binop_2(X42,X43),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k11_binop_2])])).
fof(c_0_65, plain, ![X36, X37]:(~v1_xreal_0(X36)|~v1_xreal_0(X37)|k11_binop_2(X36,X37)=k3_xcmplx_0(X36,X37)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k11_binop_2])])).
cnf(c_0_66, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_67, plain, (v1_xcmplx_0(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_68, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_57])).
cnf(c_0_69, plain, (m1_subset_1(np__1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_58])).
cnf(c_0_70, plain, (v6_membered(k5_numbers)), inference(rw,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_71, plain, (k3_xcmplx_0(k3_xcmplx_0(X2,X1),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62])])).
fof(c_0_72, plain, ![X58, X59]:(~v1_xcmplx_0(X58)|~v1_xcmplx_0(X59)|v1_xcmplx_0(k3_xcmplx_0(X58,X59))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc3_xcmplx_0])])).
fof(c_0_73, plain, ![X29, X30]:(~v1_xreal_0(X29)|~v1_xreal_0(X30)|k11_binop_2(X29,X30)=k11_binop_2(X30,X29)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k11_binop_2])])).
fof(c_0_74, plain, ![X40, X41]:(~v3_membered(X40)|(~m1_subset_1(X41,X40)|v1_xreal_0(X41))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc14_membered])])])).
fof(c_0_75, negated_conjecture, (m1_subset_1(esk1_0,k5_numbers)&k15_sin_cos(k5_binop_2(k11_binop_2(k11_binop_2(np__2,k32_sin_cos),esk1_0),k7_complex1))!=np__1), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_63])])])).
fof(c_0_76, plain, ![X96, X97]:(~v3_membered(X96)|(~m1_subset_1(X97,k1_zfmisc_1(X96))|v3_membered(X97))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc21_membered])])])).
fof(c_0_77, plain, ![X33]:(~m1_subset_1(X33,k1_numbers)|m1_subset_1(k1_real_1(X33),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_real_1])])).
cnf(c_0_78, plain, (m1_subset_1(k11_binop_2(X1,X2),k1_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_79, plain, (k11_binop_2(X1,X2)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_80, plain, (v1_xreal_0(k31_sin_cos)), inference(split_conjunct,[status(thm)],[dt_k31_sin_cos])).
cnf(c_0_81, plain, (k32_sin_cos=k31_sin_cos), inference(split_conjunct,[status(thm)],[redefinition_k32_sin_cos])).
cnf(c_0_82, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_66])).
cnf(c_0_83, plain, (v1_xcmplx_0(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_67])).
cnf(c_0_84, plain, (v7_ordinal1(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70])])])).
cnf(c_0_85, plain, (k3_xcmplx_0(X2,k3_xcmplx_0(X1,X3))=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_71])])).
cnf(c_0_86, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_87, plain, (k11_binop_2(X1,X2)=k11_binop_2(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_88, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_89, negated_conjecture, (m1_subset_1(esk1_0,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_90, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_91, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[dt_k5_numbers])).
cnf(c_0_92, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
cnf(c_0_93, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_77])).
fof(c_0_94, plain, ![X25]:(~m1_subset_1(X25,k1_numbers)|k1_real_1(k1_real_1(X25))=X25), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k1_real_1])])).
cnf(c_0_95, plain, (m1_subset_1(k11_binop_2(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_78])).
cnf(c_0_96, plain, (k11_binop_2(X1,X2)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_79])).
cnf(c_0_97, plain, (v1_xreal_0(k31_sin_cos)), inference(evalgc,[status(thm)],[c_0_80])).
cnf(c_0_98, plain, (k31_sin_cos=k32_sin_cos), inference(evalgc,[status(thm)],[c_0_81])).
cnf(c_0_99, plain, (m1_subset_1(np__2,k1_numbers)), inference(split_conjunct,[status(thm)],[spc2_numerals])).
cnf(c_0_100, plain, (k3_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_62])])).
cnf(c_0_101, plain, (v1_xcmplx_0(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84])])).
cnf(c_0_102, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X3,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_85])])).
cnf(c_0_103, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_86])).
cnf(c_0_104, plain, (k11_binop_2(X1,X2)=k11_binop_2(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_87])).
cnf(c_0_105, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_88])).
cnf(c_0_106, negated_conjecture, (m1_subset_1(esk1_0,k5_numbers)), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_107, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_90])).
cnf(c_0_108, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_91])).
cnf(c_0_109, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_92])).
cnf(c_0_110, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_93])).
cnf(c_0_111, plain, (k1_real_1(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_112, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96])])).
cnf(c_0_113, plain, (v1_xreal_0(k32_sin_cos)), inference(rw,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_114, plain, (m1_subset_1(np__2,k1_numbers)), inference(evalgc,[status(thm)],[c_0_99])).
fof(c_0_115, plain, ![X61, X62]:(~v1_membered(X61)|(~m1_subset_1(X62,X61)|v1_xcmplx_0(X62))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc12_membered])])])).
fof(c_0_116, plain, ![X57]:(~v1_xboole_0(X57)|X57=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
cnf(c_0_117, negated_conjecture, (k15_sin_cos(k5_binop_2(k11_binop_2(k11_binop_2(np__2,k32_sin_cos),esk1_0),k7_complex1))!=np__1), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_118, plain, (k7_complex1=k1_xcmplx_0), inference(split_conjunct,[status(thm)],[redefinition_k7_complex1])).
fof(c_0_119, plain, ![X27, X28]:(~v1_xcmplx_0(X27)|~v1_xcmplx_0(X28)|k5_binop_2(X27,X28)=k5_binop_2(X28,X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k5_binop_2])])).
fof(c_0_120, plain, ![X9, X10]:(~v1_xreal_0(X9)|(~v7_ordinal1(X10)|k1_seq_1(k16_sin_cos,X9)=k1_seq_1(k16_sin_cos,k7_real_1(k8_real_1(k8_real_1(np__2,k32_sin_cos),X10),X9)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_sin_cos2])])])).
fof(c_0_121, plain, ![X16, X17]:(~m1_subset_1(X16,k1_numbers)|~v1_xreal_0(X17)|k8_real_1(X16,X17)=k3_xcmplx_0(X16,X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_real_1])])).
cnf(c_0_122, plain, (k3_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_100, c_0_101])])])).
cnf(c_0_123, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X3,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_124, plain, (k11_binop_2(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_96])])).
cnf(c_0_125, negated_conjecture, (v1_xreal_0(esk1_0)|~v3_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106])])).
cnf(c_0_126, plain, (v3_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_109])])])).
cnf(c_0_127, plain, (v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_110]), c_0_109])])])).
cnf(c_0_128, plain, (k1_real_1(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_111])).
cnf(c_0_129, plain, (m1_subset_1(k3_xcmplx_0(X1,k32_sin_cos),k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113])])).
cnf(c_0_130, plain, (v1_xreal_0(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_114]), c_0_109])])])).
cnf(c_0_131, plain, (m1_subset_1(np__2,k5_numbers)), inference(split_conjunct,[status(thm)],[spc2_numerals])).
cnf(c_0_132, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_115])).
cnf(c_0_133, plain, (m1_subset_1(k32_sin_cos,k1_numbers)), inference(split_conjunct,[status(thm)],[dt_k32_sin_cos])).
fof(c_0_134, plain, ![X75]:(~v3_membered(X75)|v1_membered(X75)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc5_membered])])).
fof(c_0_135, plain, ![X82, X83]:(~m1_subset_1(X82,k1_numbers)|~v1_xreal_0(X83)|k7_real_1(X82,X83)=k2_xcmplx_0(X82,X83)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_real_1])])).
fof(c_0_136, plain, ![X38, X39]:(~m1_subset_1(X38,k1_numbers)|~v1_xreal_0(X39)|k7_real_1(X38,X39)=k7_real_1(X39,X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k7_real_1])])).
cnf(c_0_137, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_116])).
cnf(c_0_138, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
cnf(c_0_139, negated_conjecture, (k15_sin_cos(k5_binop_2(k11_binop_2(k11_binop_2(np__2,k32_sin_cos),esk1_0),k7_complex1))!=np__1), inference(evalgc,[status(thm)],[c_0_117])).
cnf(c_0_140, plain, (k7_complex1=k1_xcmplx_0), inference(evalgc,[status(thm)],[c_0_118])).
cnf(c_0_141, plain, (k5_binop_2(X1,X2)=k5_binop_2(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_142, plain, (v1_xcmplx_0(k1_xcmplx_0)), inference(split_conjunct,[status(thm)],[fc1_xcmplx_0])).
fof(c_0_143, plain, ![X31, X32]:(~v1_xcmplx_0(X31)|~v1_xcmplx_0(X32)|k5_binop_2(X31,X32)=k3_xcmplx_0(X31,X32)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k5_binop_2])])).
cnf(c_0_144, plain, (k1_seq_1(k16_sin_cos,X1)=k1_seq_1(k16_sin_cos,k7_real_1(k8_real_1(k8_real_1(np__2,k32_sin_cos),X2),X1))|~v1_xreal_0(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_145, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_146, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(np__1,X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_101])])])).
cnf(c_0_147, plain, (k3_xcmplx_0(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_124])])).
cnf(c_0_148, negated_conjecture, (v1_xreal_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_125, c_0_126])])).
cnf(c_0_149, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128])])).
cnf(c_0_150, plain, (m1_subset_1(k3_xcmplx_0(np__2,k32_sin_cos),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130])])).
cnf(c_0_151, plain, (m1_subset_1(np__2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_131])).
cnf(c_0_152, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_132])).
cnf(c_0_153, plain, (m1_subset_1(k32_sin_cos,k1_numbers)), inference(evalgc,[status(thm)],[c_0_133])).
cnf(c_0_154, plain, (v1_membered(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_134])).
fof(c_0_155, plain, ![X7, X8]:(~v1_xreal_0(X7)|(~v7_ordinal1(X8)|k1_seq_1(k19_sin_cos,X7)=k1_seq_1(k19_sin_cos,k7_real_1(k8_real_1(k8_real_1(np__2,k32_sin_cos),X8),X7)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_sin_cos2])])])).
cnf(c_0_156, plain, (k7_real_1(X1,X2)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_135])).
cnf(c_0_157, plain, (k7_real_1(X1,X2)=k7_real_1(X2,X1)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_136])).
fof(c_0_158, plain, ![X56]:(~v1_xcmplx_0(X56)|k2_xcmplx_0(X56,k6_numbers)=X56), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_arithm])])).
cnf(c_0_159, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_137])).
cnf(c_0_160, plain, (k1_xboole_0=k6_numbers), inference(evalgc,[status(thm)],[c_0_138])).
cnf(c_0_161, plain, (v1_xboole_0(np__0)), inference(split_conjunct,[status(thm)],[spc0_boole])).
cnf(c_0_162, plain, (m1_subset_1(np__0,k5_numbers)), inference(split_conjunct,[status(thm)],[spc0_numerals])).
fof(c_0_163, plain, ![X6]:((((k1_seq_1(k19_sin_cos,k6_numbers)=np__1|~v1_xreal_0(X6))&(k1_seq_1(k16_sin_cos,k6_numbers)=k6_numbers|~v1_xreal_0(X6)))&(k1_seq_1(k19_sin_cos,k4_xcmplx_0(X6))=k1_seq_1(k19_sin_cos,X6)|~v1_xreal_0(X6)))&(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X6))=k1_real_1(k1_seq_1(k16_sin_cos,X6))|~v1_xreal_0(X6))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_sin_cos])])])).
cnf(c_0_164, negated_conjecture, (k15_sin_cos(k5_binop_2(k11_binop_2(k11_binop_2(np__2,k32_sin_cos),esk1_0),k1_xcmplx_0))!=np__1), inference(rw,[status(thm)],[c_0_139, c_0_140])).
cnf(c_0_165, plain, (k5_binop_2(X1,X2)=k5_binop_2(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_141])).
cnf(c_0_166, plain, (v1_xcmplx_0(k1_xcmplx_0)), inference(evalgc,[status(thm)],[c_0_142])).
cnf(c_0_167, plain, (k5_binop_2(X1,X2)=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_143])).
fof(c_0_168, plain, ![X15]:(~m1_subset_1(X15,k1_numbers)|k15_sin_cos(k3_xcmplx_0(X15,k7_complex1))=k2_xcmplx_0(k21_sin_cos(X15),k3_xcmplx_0(k18_sin_cos(X15),k7_complex1))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t25_sin_cos])])).
fof(c_0_169, plain, ![X11]:(~v1_xreal_0(X11)|k20_sin_cos(X11)=k1_seq_1(k19_sin_cos,X11)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d19_sin_cos])])).
fof(c_0_170, plain, ![X20]:(~m1_subset_1(X20,k1_numbers)|k21_sin_cos(X20)=k20_sin_cos(X20)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k21_sin_cos])])).
cnf(c_0_171, plain, (k1_seq_1(k16_sin_cos,k7_real_1(k8_real_1(k8_real_1(np__2,k32_sin_cos),X2),X1))=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[c_0_144])).
cnf(c_0_172, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_145])).
cnf(c_0_173, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(np__1,X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_146, c_0_103])).
cnf(c_0_174, negated_conjecture, (k3_xcmplx_0(esk1_0,X1)=k3_xcmplx_0(X1,esk1_0)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_148])])).
cnf(c_0_175, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_149, c_0_110])).
cnf(c_0_176, negated_conjecture, (v7_ordinal1(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_106]), c_0_70])])])).
cnf(c_0_177, plain, (m1_subset_1(np__1,k1_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_178, plain, (v1_xreal_0(k3_xcmplx_0(np__2,k32_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_150]), c_0_109])])])).
cnf(c_0_179, plain, (v7_ordinal1(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_151]), c_0_70])])])).
cnf(c_0_180, plain, (v1_xcmplx_0(k32_sin_cos)|~v1_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_153])])).
cnf(c_0_181, plain, (v1_membered(X1)|~v3_membered(X1)), inference(evalgc,[status(thm)],[c_0_154])).
cnf(c_0_182, negated_conjecture, (m1_subset_1(k3_xcmplx_0(X1,esk1_0),k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_148])])).
cnf(c_0_183, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))=k3_xcmplx_0(X3,k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_61])])).
cnf(c_0_184, plain, (k3_xcmplx_0(k32_sin_cos,X1)=k3_xcmplx_0(X1,k32_sin_cos)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_113])])).
cnf(c_0_185, plain, (k1_seq_1(k19_sin_cos,X1)=k1_seq_1(k19_sin_cos,k7_real_1(k8_real_1(k8_real_1(np__2,k32_sin_cos),X2),X1))|~v1_xreal_0(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_155])).
cnf(c_0_186, plain, (k7_real_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_156])).
cnf(c_0_187, plain, (k7_real_1(X1,X2)=k7_real_1(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_157])).
fof(c_0_188, plain, ![X78, X79]:(~v1_xcmplx_0(X78)|~v1_xcmplx_0(X79)|k2_xcmplx_0(X78,X79)=k2_xcmplx_0(X79,X78)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k2_xcmplx_0])])).
cnf(c_0_189, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_158])).
cnf(c_0_190, plain, (X1=k6_numbers|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_159, c_0_160])).
cnf(c_0_191, plain, (v1_xboole_0(np__0)), inference(evalgc,[status(thm)],[c_0_161])).
cnf(c_0_192, plain, (m1_subset_1(np__0,k5_numbers)), inference(evalgc,[status(thm)],[c_0_162])).
cnf(c_0_193, plain, (k1_seq_1(k19_sin_cos,k6_numbers)=np__1|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_163])).
cnf(c_0_194, negated_conjecture, (k15_sin_cos(k5_binop_2(k1_xcmplx_0,k11_binop_2(k11_binop_2(np__2,k32_sin_cos),esk1_0)))!=np__1|~v1_xcmplx_0(k11_binop_2(k11_binop_2(np__2,k32_sin_cos),esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_165]), c_0_166])])])).
cnf(c_0_195, plain, (k5_binop_2(X1,X2)=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_167])).
cnf(c_0_196, plain, (k15_sin_cos(k3_xcmplx_0(X1,k7_complex1))=k2_xcmplx_0(k21_sin_cos(X1),k3_xcmplx_0(k18_sin_cos(X1),k7_complex1))|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_168])).
cnf(c_0_197, plain, (k20_sin_cos(X1)=k1_seq_1(k19_sin_cos,X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_169])).
cnf(c_0_198, plain, (k21_sin_cos(X1)=k20_sin_cos(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_170])).
fof(c_0_199, plain, ![X14]:(~v1_xreal_0(X14)|k17_sin_cos(X14)=k1_seq_1(k16_sin_cos,X14)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d17_sin_cos])])).
fof(c_0_200, plain, ![X21]:(~m1_subset_1(X21,k1_numbers)|k18_sin_cos(X21)=k17_sin_cos(X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k18_sin_cos])])).
cnf(c_0_201, plain, (k1_seq_1(k16_sin_cos,k7_real_1(k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1),X2))=k1_seq_1(k16_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_172]), c_0_113]), c_0_114])])])).
cnf(c_0_202, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__1))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_62]), c_0_101])])])).
cnf(c_0_203, negated_conjecture, (k3_xcmplx_0(esk1_0,X1)=k3_xcmplx_0(X1,esk1_0)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_175])])).
cnf(c_0_204, negated_conjecture, (v1_xcmplx_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_176])])).
cnf(c_0_205, plain, (m1_subset_1(np__1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_177])).
cnf(c_0_206, negated_conjecture, (k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),esk1_0)=k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_178])])).
cnf(c_0_207, plain, (v1_xcmplx_0(np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_179])])).
cnf(c_0_208, plain, (v1_xcmplx_0(k32_sin_cos)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_181]), c_0_109])])])).
cnf(c_0_209, negated_conjecture, (m1_subset_1(k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),esk1_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_178])])).
cnf(c_0_210, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))=k3_xcmplx_0(X3,k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_183, c_0_103])).
cnf(c_0_211, plain, (k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),X1)=k3_xcmplx_0(X1,k3_xcmplx_0(np__2,k32_sin_cos))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_178])])).
cnf(c_0_212, plain, (k3_xcmplx_0(k32_sin_cos,np__2)=k3_xcmplx_0(np__2,k32_sin_cos)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_130])])).
cnf(c_0_213, plain, (k1_seq_1(k19_sin_cos,k7_real_1(k8_real_1(k8_real_1(np__2,k32_sin_cos),X2),X1))=k1_seq_1(k19_sin_cos,X1)|~v1_xreal_0(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[c_0_185])).
cnf(c_0_214, plain, (k7_real_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_187])])).
cnf(c_0_215, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_188])).
cnf(c_0_216, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_189])).
cnf(c_0_217, plain, (k6_numbers=np__0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_190, c_0_191])])).
cnf(c_0_218, plain, (v7_ordinal1(np__0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_192]), c_0_70])])])).
cnf(c_0_219, plain, (k1_seq_1(k19_sin_cos,k6_numbers)=np__1|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_193])).
cnf(c_0_220, negated_conjecture, (k15_sin_cos(k5_binop_2(k1_xcmplx_0,k3_xcmplx_0(k11_binop_2(np__2,k32_sin_cos),esk1_0)))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(k11_binop_2(np__2,k32_sin_cos),esk1_0))|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_194, c_0_96]), c_0_148])])])).
cnf(c_0_221, plain, (k5_binop_2(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_195])])).
cnf(c_0_222, plain, (k2_xcmplx_0(k21_sin_cos(X1),k3_xcmplx_0(k18_sin_cos(X1),k7_complex1))=k15_sin_cos(k3_xcmplx_0(X1,k7_complex1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_196])).
cnf(c_0_223, plain, (k20_sin_cos(X1)=k1_seq_1(k19_sin_cos,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_197])).
cnf(c_0_224, plain, (k20_sin_cos(X1)=k21_sin_cos(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_198])).
cnf(c_0_225, plain, (k17_sin_cos(X1)=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_199])).
cnf(c_0_226, plain, (k18_sin_cos(X1)=k17_sin_cos(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_200])).
cnf(c_0_227, plain, (k1_seq_1(k16_sin_cos,k7_real_1(k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),X1),X2))=k1_seq_1(k16_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(k3_xcmplx_0(np__2,k32_sin_cos),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_172])])).
cnf(c_0_228, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,np__1),X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_202])])).
cnf(c_0_229, negated_conjecture, (k3_xcmplx_0(np__1,esk1_0)=esk1_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_203]), c_0_204]), c_0_205])])])).
cnf(c_0_230, negated_conjecture, (k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos))=k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_206]), c_0_204]), c_0_207]), c_0_208])])])).
cnf(c_0_231, negated_conjecture, (k3_xcmplx_0(esk1_0,k32_sin_cos)=k3_xcmplx_0(k32_sin_cos,esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_113])])).
cnf(c_0_232, negated_conjecture, (m1_subset_1(k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos)),k1_numbers)), inference(rw,[status(thm)],[c_0_209, c_0_206])).
cnf(c_0_233, plain, (k3_xcmplx_0(X2,k3_xcmplx_0(X1,np__1))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_173]), c_0_101])])])).
cnf(c_0_234, plain, (k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),X1)=k3_xcmplx_0(X1,k3_xcmplx_0(np__2,k32_sin_cos))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_175])])).
cnf(c_0_235, plain, (v1_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_212]), c_0_207]), c_0_208])])])).
cnf(c_0_236, plain, (k1_seq_1(k16_sin_cos,k6_numbers)=k6_numbers|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_163])).
cnf(c_0_237, plain, (k1_seq_1(k19_sin_cos,k2_xcmplx_0(X2,k8_real_1(k8_real_1(np__2,k32_sin_cos),X1)))=k1_seq_1(k19_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(k8_real_1(k8_real_1(np__2,k32_sin_cos),X1))|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_214])])).
cnf(c_0_238, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_215])).
cnf(c_0_239, plain, (k2_xcmplx_0(X1,np__0)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_216, c_0_217])])).
cnf(c_0_240, plain, (v1_xcmplx_0(np__0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_218])])).
cnf(c_0_241, plain, (k1_seq_1(k19_sin_cos,np__0)=np__1|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_219, c_0_217])])).
cnf(c_0_242, plain, (m1_subset_1(np__0,k1_numbers)), inference(split_conjunct,[status(thm)],[spc0_numerals])).
cnf(c_0_243, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k3_xcmplx_0(k11_binop_2(np__2,k32_sin_cos),esk1_0),k1_xcmplx_0))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(k11_binop_2(np__2,k32_sin_cos),esk1_0))|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_221]), c_0_166])])])).
cnf(c_0_244, plain, (k2_xcmplx_0(k21_sin_cos(X1),k3_xcmplx_0(k18_sin_cos(X1),k1_xcmplx_0))=k15_sin_cos(k3_xcmplx_0(X1,k1_xcmplx_0))|~m1_subset_1(X1,k1_numbers)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_222, c_0_140]), c_0_140])).
cnf(c_0_245, plain, (k21_sin_cos(X1)=k1_seq_1(k19_sin_cos,X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_224])])).
cnf(c_0_246, plain, (k17_sin_cos(X1)=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_225])).
cnf(c_0_247, plain, (k17_sin_cos(X1)=k18_sin_cos(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_226])).
cnf(c_0_248, plain, (k1_seq_1(k16_sin_cos,k7_real_1(k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),X1),X2))=k1_seq_1(k16_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_227, c_0_150])])).
cnf(c_0_249, negated_conjecture, (k3_xcmplx_0(esk1_0,X1)=k3_xcmplx_0(X1,esk1_0)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_203]), c_0_229]), c_0_229]), c_0_204]), c_0_204]), c_0_205])])])).
cnf(c_0_250, negated_conjecture, (k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,esk1_0))=k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_230]), c_0_231]), c_0_208]), c_0_204]), c_0_207])])])).
cnf(c_0_251, negated_conjecture, (m1_subset_1(k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,esk1_0)),k1_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_232, c_0_230])])).
cnf(c_0_252, plain, (k3_xcmplx_0(X1,np__1)=k3_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_82]), c_0_101])])])).
cnf(c_0_253, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(np__2,k32_sin_cos))=k3_xcmplx_0(np__2,k32_sin_cos)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_234]), c_0_235]), c_0_205])])])).
cnf(c_0_254, plain, (v1_xreal_0(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_205]), c_0_109])])])).
cnf(c_0_255, plain, (k1_seq_1(k16_sin_cos,k6_numbers)=k6_numbers|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_236])).
cnf(c_0_256, plain, (k1_seq_1(k19_sin_cos,k2_xcmplx_0(X1,k8_real_1(k8_real_1(np__2,k32_sin_cos),X2)))=k1_seq_1(k19_sin_cos,X1)|~v7_ordinal1(X2)|~v1_xreal_0(k8_real_1(k8_real_1(np__2,k32_sin_cos),X2))|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_237, c_0_175])).
cnf(c_0_257, plain, (k2_xcmplx_0(np__0,X1)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_239]), c_0_240])])])).
cnf(c_0_258, negated_conjecture, (k1_seq_1(k19_sin_cos,np__0)=np__1), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_148])])).
cnf(c_0_259, plain, (m1_subset_1(np__0,k1_numbers)), inference(evalgc,[status(thm)],[c_0_242])).
cnf(c_0_260, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),esk1_0),k1_xcmplx_0))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_96]), c_0_178]), c_0_113]), c_0_130])])])).
cnf(c_0_261, plain, (k2_xcmplx_0(k1_seq_1(k19_sin_cos,X1),k3_xcmplx_0(k18_sin_cos(X1),k1_xcmplx_0))=k15_sin_cos(k3_xcmplx_0(X1,k1_xcmplx_0))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_244, c_0_245])])).
cnf(c_0_262, plain, (k18_sin_cos(X1)=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_247])])).
cnf(c_0_263, negated_conjecture, (k1_seq_1(k16_sin_cos,k7_real_1(k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)),X1))=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_248, c_0_249]), c_0_230]), c_0_250]), c_0_176]), c_0_148]), c_0_235])])])).
cnf(c_0_264, negated_conjecture, (m1_subset_1(k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)),k1_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_251, c_0_250])])).
cnf(c_0_265, plain, (k1_seq_1(k16_sin_cos,k7_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1))=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_248, c_0_252]), c_0_253]), c_0_84]), c_0_254]), c_0_235])])])).
cnf(c_0_266, plain, (k1_seq_1(k16_sin_cos,np__0)=np__0|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_255, c_0_217]), c_0_217])])).
cnf(c_0_267, plain, (k1_seq_1(k19_sin_cos,k8_real_1(k8_real_1(np__2,k32_sin_cos),X1))=np__1|~v7_ordinal1(X1)|~v1_xreal_0(k8_real_1(k8_real_1(np__2,k32_sin_cos),X1))|~v1_xcmplx_0(k8_real_1(k8_real_1(np__2,k32_sin_cos),X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_256, c_0_257]), c_0_258]), c_0_259])])])).
cnf(c_0_268, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos)),k1_xcmplx_0))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_260, c_0_206]), c_0_206])).
cnf(c_0_269, plain, (k2_xcmplx_0(k1_seq_1(k19_sin_cos,X1),k3_xcmplx_0(k1_seq_1(k16_sin_cos,X1),k1_xcmplx_0))=k15_sin_cos(k3_xcmplx_0(X1,k1_xcmplx_0))|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_261, c_0_262])])).
cnf(c_0_270, negated_conjecture, (k1_seq_1(k16_sin_cos,k7_real_1(X1,k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0))))=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_187]), c_0_264])])])).
cnf(c_0_271, negated_conjecture, (v1_xreal_0(k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_264]), c_0_109])])])).
cnf(c_0_272, plain, (k1_seq_1(k16_sin_cos,k2_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),X1))=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_186]), c_0_150])])])).
cnf(c_0_273, negated_conjecture, (k1_seq_1(k16_sin_cos,np__0)=np__0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_148])])).
cnf(c_0_274, plain, (v1_xreal_0(np__0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_259]), c_0_109])])])).
cnf(c_0_275, plain, (k1_seq_1(k19_sin_cos,k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1))=np__1|~v1_xcmplx_0(k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_172]), c_0_113]), c_0_114])])])).
cnf(c_0_276, negated_conjecture, (v1_xcmplx_0(k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_230]), c_0_235]), c_0_204])])])).
cnf(c_0_277, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,esk1_0)),k1_xcmplx_0))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,esk1_0)))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_268, c_0_230]), c_0_230])])).
cnf(c_0_278, plain, (k2_xcmplx_0(k1_seq_1(k19_sin_cos,X1),k3_xcmplx_0(k1_seq_1(k16_sin_cos,X1),k1_xcmplx_0))=k15_sin_cos(k3_xcmplx_0(X1,k1_xcmplx_0))|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_269, c_0_175])).
cnf(c_0_279, negated_conjecture, (k1_seq_1(k16_sin_cos,k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)))=k1_seq_1(k16_sin_cos,k3_xcmplx_0(np__2,k32_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_270]), c_0_271]), c_0_178])])])).
cnf(c_0_280, plain, (k1_seq_1(k16_sin_cos,k3_xcmplx_0(np__2,k32_sin_cos))=np__0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_272, c_0_239]), c_0_273]), c_0_274]), c_0_235])])])).
cnf(c_0_281, plain, (k1_seq_1(k19_sin_cos,k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),X1))=np__1|~v1_xcmplx_0(k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_172]), c_0_150])])])).
cnf(c_0_282, negated_conjecture, (v1_xcmplx_0(k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_276, c_0_250])])).
cnf(c_0_283, plain, (k3_xcmplx_0(k1_xcmplx_0,np__0)=np__0), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__cr1r0_r0_r0])).
cnf(c_0_284, plain, (k2_xcmplx_0(np__1,np__0)=np__1), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r0_r1])).
cnf(c_0_285, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,esk1_0)),k1_xcmplx_0))!=np__1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_277, c_0_276])])).
cnf(c_0_286, plain, (k2_xcmplx_0(k1_seq_1(k19_sin_cos,X1),k3_xcmplx_0(k1_xcmplx_0,k1_seq_1(k16_sin_cos,X1)))=k15_sin_cos(k3_xcmplx_0(X1,k1_xcmplx_0))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_62]), c_0_166])])])).
cnf(c_0_287, negated_conjecture, (k1_seq_1(k16_sin_cos,k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)))=np__0), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_279, c_0_280])])).
cnf(c_0_288, negated_conjecture, (k1_seq_1(k19_sin_cos,k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)))=np__1), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_281, c_0_249]), c_0_230]), c_0_250]), c_0_230]), c_0_250]), c_0_282]), c_0_176]), c_0_230]), c_0_250]), c_0_271]), c_0_148]), c_0_235])])])).
cnf(c_0_289, plain, (k3_xcmplx_0(k1_xcmplx_0,np__0)=np__0), inference(evalgc,[status(thm)],[c_0_283])).
cnf(c_0_290, plain, (k2_xcmplx_0(np__1,np__0)=np__1), inference(evalgc,[status(thm)],[c_0_284])).
cnf(c_0_291, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)),k1_xcmplx_0))!=np__1), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_285, c_0_250])])).
cnf(c_0_292, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_287]), c_0_288]), c_0_289]), c_0_290]), c_0_240]), c_0_264])]), c_0_291]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 293
# Proof object clause steps            : 212
# Proof object formula steps           : 81
# Proof object conjectures             : 44
# Proof object clause conjectures      : 41
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 53
# Proof object initial formulas used   : 48
# Proof object generating inferences   : 80
# Proof object simplifying inferences  : 159
# Training examples: 142 positive, 1159 negative
# Training: Positive examples begin
cnf(c_0_60, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_140, plain, (k7_complex1=k1_xcmplx_0)).# trainpos
cnf(c_0_98, plain, (k31_sin_cos=k32_sin_cos)).# trainpos
cnf(c_0_160, plain, (k1_xboole_0=k6_numbers)).# trainpos
cnf(c_0_97, plain, (v1_xreal_0(k31_sin_cos))).# trainpos
cnf(c_0_59, plain, (v6_membered(k4_ordinal1))).# trainpos
cnf(c_0_191, plain, (v1_xboole_0(np__0))).# trainpos
cnf(c_0_166, plain, (v1_xcmplx_0(k1_xcmplx_0))).# trainpos
cnf(c_0_109, plain, (v3_membered(k1_numbers))).# trainpos
cnf(c_0_139, negated_conjecture, (k15_sin_cos(k5_binop_2(k11_binop_2(k11_binop_2(np__2,k32_sin_cos),esk1_0),k7_complex1))!=np__1)).# trainpos
cnf(c_0_106, negated_conjecture, (m1_subset_1(esk1_0,k5_numbers))).# trainpos
cnf(c_0_151, plain, (m1_subset_1(np__2,k5_numbers))).# trainpos
cnf(c_0_69, plain, (m1_subset_1(np__1,k5_numbers))).# trainpos
cnf(c_0_114, plain, (m1_subset_1(np__2,k1_numbers))).# trainpos
cnf(c_0_153, plain, (m1_subset_1(k32_sin_cos,k1_numbers))).# trainpos
cnf(c_0_159, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_205, plain, (m1_subset_1(np__1,k1_numbers))).# trainpos
cnf(c_0_181, plain, (v1_membered(X1)|~v3_membered(X1))).# trainpos
cnf(c_0_289, plain, (k3_xcmplx_0(k1_xcmplx_0,np__0)=np__0)).# trainpos
cnf(c_0_290, plain, (k2_xcmplx_0(np__1,np__0)=np__1)).# trainpos
cnf(c_0_192, plain, (m1_subset_1(np__0,k5_numbers))).# trainpos
cnf(c_0_259, plain, (m1_subset_1(np__0,k1_numbers))).# trainpos
cnf(c_0_108, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_83, plain, (v1_xcmplx_0(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_104, plain, (k11_binop_2(X1,X2)=k11_binop_2(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_219, plain, (k1_seq_1(k19_sin_cos,k6_numbers)=np__1|~v1_xreal_0(X1))).# trainpos
cnf(c_0_255, plain, (k1_seq_1(k16_sin_cos,k6_numbers)=k6_numbers|~v1_xreal_0(X1))).# trainpos
cnf(c_0_216, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_82, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_223, plain, (k20_sin_cos(X1)=k1_seq_1(k19_sin_cos,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_105, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_165, plain, (k5_binop_2(X1,X2)=k5_binop_2(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_68, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_246, plain, (k17_sin_cos(X1)=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_224, plain, (k20_sin_cos(X1)=k21_sin_cos(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_247, plain, (k17_sin_cos(X1)=k18_sin_cos(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_128, plain, (k1_real_1(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_110, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_152, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_107, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_96, plain, (k11_binop_2(X1,X2)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_172, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_195, plain, (k5_binop_2(X1,X2)=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_62, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_238, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_103, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_95, plain, (m1_subset_1(k11_binop_2(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_222, plain, (k2_xcmplx_0(k21_sin_cos(X1),k3_xcmplx_0(k18_sin_cos(X1),k7_complex1))=k15_sin_cos(k3_xcmplx_0(X1,k7_complex1))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_186, plain, (k7_real_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_187, plain, (k7_real_1(X1,X2)=k7_real_1(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_213, plain, (k1_seq_1(k19_sin_cos,k7_real_1(k8_real_1(k8_real_1(np__2,k32_sin_cos),X2),X1))=k1_seq_1(k19_sin_cos,X1)|~v1_xreal_0(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_171, plain, (k1_seq_1(k16_sin_cos,k7_real_1(k8_real_1(k8_real_1(np__2,k32_sin_cos),X2),X1))=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_61, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_254, plain, (v1_xreal_0(np__1))).# trainpos
cnf(c_0_217, plain, (k6_numbers=np__0)).# trainpos
cnf(c_0_130, plain, (v1_xreal_0(np__2))).# trainpos
cnf(c_0_176, negated_conjecture, (v7_ordinal1(esk1_0))).# trainpos
cnf(c_0_126, plain, (v3_membered(k5_numbers))).# trainpos
cnf(c_0_262, plain, (k18_sin_cos(X1)=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_245, plain, (k21_sin_cos(X1)=k1_seq_1(k19_sin_cos,X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_221, plain, (k5_binop_2(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_125, negated_conjecture, (v1_xreal_0(esk1_0)|~v3_membered(k5_numbers))).# trainpos
cnf(c_0_194, negated_conjecture, (k15_sin_cos(k5_binop_2(k1_xcmplx_0,k11_binop_2(k11_binop_2(np__2,k32_sin_cos),esk1_0)))!=np__1|~v1_xcmplx_0(k11_binop_2(k11_binop_2(np__2,k32_sin_cos),esk1_0)))).# trainpos
cnf(c_0_274, plain, (v1_xreal_0(np__0))).# trainpos
cnf(c_0_218, plain, (v7_ordinal1(np__0))).# trainpos
cnf(c_0_84, plain, (v7_ordinal1(np__1))).# trainpos
cnf(c_0_179, plain, (v7_ordinal1(np__2))).# trainpos
cnf(c_0_124, plain, (k11_binop_2(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_214, plain, (k7_real_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_100, plain, (k3_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).# trainpos
cnf(c_0_127, plain, (v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_201, plain, (k1_seq_1(k16_sin_cos,k7_real_1(k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1),X2))=k1_seq_1(k16_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_71, plain, (k3_xcmplx_0(k3_xcmplx_0(X2,X1),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_112, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_180, plain, (v1_xcmplx_0(k32_sin_cos)|~v1_membered(k1_numbers))).# trainpos
cnf(c_0_183, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))=k3_xcmplx_0(X3,k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_204, negated_conjecture, (v1_xcmplx_0(esk1_0))).# trainpos
cnf(c_0_241, plain, (k1_seq_1(k19_sin_cos,np__0)=np__1|~v1_xreal_0(X1))).# trainpos
cnf(c_0_240, plain, (v1_xcmplx_0(np__0))).# trainpos
cnf(c_0_101, plain, (v1_xcmplx_0(np__1))).# trainpos
cnf(c_0_207, plain, (v1_xcmplx_0(np__2))).# trainpos
cnf(c_0_208, plain, (v1_xcmplx_0(k32_sin_cos))).# trainpos
cnf(c_0_266, plain, (k1_seq_1(k16_sin_cos,np__0)=np__0|~v1_xreal_0(X1))).# trainpos
cnf(c_0_147, plain, (k3_xcmplx_0(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_261, plain, (k2_xcmplx_0(k1_seq_1(k19_sin_cos,X1),k3_xcmplx_0(k18_sin_cos(X1),k1_xcmplx_0))=k15_sin_cos(k3_xcmplx_0(X1,k1_xcmplx_0))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_239, plain, (k2_xcmplx_0(X1,np__0)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_182, negated_conjecture, (m1_subset_1(k3_xcmplx_0(X1,esk1_0),k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_149, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_85, plain, (k3_xcmplx_0(X2,k3_xcmplx_0(X1,X3))=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_129, plain, (m1_subset_1(k3_xcmplx_0(X1,k32_sin_cos),k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_220, negated_conjecture, (k15_sin_cos(k5_binop_2(k1_xcmplx_0,k3_xcmplx_0(k11_binop_2(np__2,k32_sin_cos),esk1_0)))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(k11_binop_2(np__2,k32_sin_cos),esk1_0))|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos)))).# trainpos
cnf(c_0_122, plain, (k3_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_150, plain, (m1_subset_1(k3_xcmplx_0(np__2,k32_sin_cos),k1_numbers))).# trainpos
cnf(c_0_257, plain, (k2_xcmplx_0(np__0,X1)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_258, negated_conjecture, (k1_seq_1(k19_sin_cos,np__0)=np__1)).# trainpos
cnf(c_0_273, negated_conjecture, (k1_seq_1(k16_sin_cos,np__0)=np__0)).# trainpos
cnf(c_0_174, negated_conjecture, (k3_xcmplx_0(esk1_0,X1)=k3_xcmplx_0(X1,esk1_0)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_184, plain, (k3_xcmplx_0(k32_sin_cos,X1)=k3_xcmplx_0(X1,k32_sin_cos)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_102, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X3,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_212, plain, (k3_xcmplx_0(k32_sin_cos,np__2)=k3_xcmplx_0(np__2,k32_sin_cos))).# trainpos
cnf(c_0_231, negated_conjecture, (k3_xcmplx_0(esk1_0,k32_sin_cos)=k3_xcmplx_0(k32_sin_cos,esk1_0))).# trainpos
cnf(c_0_178, plain, (v1_xreal_0(k3_xcmplx_0(np__2,k32_sin_cos)))).# trainpos
cnf(c_0_146, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(np__1,X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_203, negated_conjecture, (k3_xcmplx_0(esk1_0,X1)=k3_xcmplx_0(X1,esk1_0)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_243, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k3_xcmplx_0(k11_binop_2(np__2,k32_sin_cos),esk1_0),k1_xcmplx_0))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(k11_binop_2(np__2,k32_sin_cos),esk1_0))|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos)))).# trainpos
cnf(c_0_206, negated_conjecture, (k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),esk1_0)=k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos)))).# trainpos
cnf(c_0_235, plain, (v1_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos)))).# trainpos
cnf(c_0_229, negated_conjecture, (k3_xcmplx_0(np__1,esk1_0)=esk1_0)).# trainpos
cnf(c_0_209, negated_conjecture, (m1_subset_1(k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),esk1_0),k1_numbers))).# trainpos
cnf(c_0_260, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),esk1_0),k1_xcmplx_0))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),esk1_0)))).# trainpos
cnf(c_0_211, plain, (k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),X1)=k3_xcmplx_0(X1,k3_xcmplx_0(np__2,k32_sin_cos))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_202, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__1))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_233, plain, (k3_xcmplx_0(X2,k3_xcmplx_0(X1,np__1))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_230, negated_conjecture, (k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos))=k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,esk1_0)))).# trainpos
cnf(c_0_234, plain, (k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),X1)=k3_xcmplx_0(X1,k3_xcmplx_0(np__2,k32_sin_cos))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_228, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,np__1),X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_276, negated_conjecture, (v1_xcmplx_0(k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,esk1_0))))).# trainpos
cnf(c_0_251, negated_conjecture, (m1_subset_1(k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,esk1_0)),k1_numbers))).# trainpos
cnf(c_0_277, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,esk1_0)),k1_xcmplx_0))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,esk1_0))))).# trainpos
cnf(c_0_249, negated_conjecture, (k3_xcmplx_0(esk1_0,X1)=k3_xcmplx_0(X1,esk1_0)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_253, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(np__2,k32_sin_cos))=k3_xcmplx_0(np__2,k32_sin_cos))).# trainpos
cnf(c_0_250, negated_conjecture, (k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,esk1_0))=k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)))).# trainpos
cnf(c_0_269, plain, (k2_xcmplx_0(k1_seq_1(k19_sin_cos,X1),k3_xcmplx_0(k1_seq_1(k16_sin_cos,X1),k1_xcmplx_0))=k15_sin_cos(k3_xcmplx_0(X1,k1_xcmplx_0))|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_291, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)),k1_xcmplx_0))!=np__1)).# trainpos
cnf(c_0_282, negated_conjecture, (v1_xcmplx_0(k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0))))).# trainpos
cnf(c_0_264, negated_conjecture, (m1_subset_1(k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)),k1_numbers))).# trainpos
cnf(c_0_286, plain, (k2_xcmplx_0(k1_seq_1(k19_sin_cos,X1),k3_xcmplx_0(k1_xcmplx_0,k1_seq_1(k16_sin_cos,X1)))=k15_sin_cos(k3_xcmplx_0(X1,k1_xcmplx_0))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,X1)))).# trainpos
cnf(c_0_271, negated_conjecture, (v1_xreal_0(k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0))))).# trainpos
cnf(c_0_252, plain, (k3_xcmplx_0(X1,np__1)=k3_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_237, plain, (k1_seq_1(k19_sin_cos,k2_xcmplx_0(X2,k8_real_1(k8_real_1(np__2,k32_sin_cos),X1)))=k1_seq_1(k19_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(k8_real_1(k8_real_1(np__2,k32_sin_cos),X1))|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_267, plain, (k1_seq_1(k19_sin_cos,k8_real_1(k8_real_1(np__2,k32_sin_cos),X1))=np__1|~v7_ordinal1(X1)|~v1_xreal_0(k8_real_1(k8_real_1(np__2,k32_sin_cos),X1))|~v1_xcmplx_0(k8_real_1(k8_real_1(np__2,k32_sin_cos),X1)))).# trainpos
cnf(c_0_275, plain, (k1_seq_1(k19_sin_cos,k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1))=np__1|~v1_xcmplx_0(k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1)))).# trainpos
cnf(c_0_281, plain, (k1_seq_1(k19_sin_cos,k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),X1))=np__1|~v1_xcmplx_0(k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_288, negated_conjecture, (k1_seq_1(k19_sin_cos,k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)))=np__1)).# trainpos
cnf(c_0_227, plain, (k1_seq_1(k16_sin_cos,k7_real_1(k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),X1),X2))=k1_seq_1(k16_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(k3_xcmplx_0(np__2,k32_sin_cos),k1_numbers))).# trainpos
cnf(c_0_263, negated_conjecture, (k1_seq_1(k16_sin_cos,k7_real_1(k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)),X1))=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_265, plain, (k1_seq_1(k16_sin_cos,k7_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1))=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_270, negated_conjecture, (k1_seq_1(k16_sin_cos,k7_real_1(X1,k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0))))=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_272, plain, (k1_seq_1(k16_sin_cos,k2_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),X1))=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_279, negated_conjecture, (k1_seq_1(k16_sin_cos,k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)))=k1_seq_1(k16_sin_cos,k3_xcmplx_0(np__2,k32_sin_cos)))).# trainpos
cnf(c_0_280, plain, (k1_seq_1(k16_sin_cos,k3_xcmplx_0(np__2,k32_sin_cos))=np__0)).# trainpos
cnf(c_0_287, negated_conjecture, (k1_seq_1(k16_sin_cos,k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)))=np__0)).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_76, plain, (k4_xcmplx_0(np__0)=np__0)).#trainneg
cnf(i_0_38, plain, (v2_xxreal_0(np__2))).#trainneg
cnf(i_0_32, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_53, plain, (v1_membered(k2_numbers))).#trainneg
cnf(i_0_119, plain, (~v1_xboole_0(np__2))).#trainneg
cnf(i_0_92, plain, (v1_funct_1(k19_sin_cos))).#trainneg
cnf(i_0_37, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_148, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_78, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_67, plain, (v1_funct_1(k16_sin_cos))).#trainneg
cnf(i_0_113, plain, (~v1_xboole_0(k2_numbers))).#trainneg
cnf(i_0_160, plain, (~v1_finset_1(k2_numbers))).#trainneg
cnf(i_0_106, plain, (k4_xcmplx_0(k4_xcmplx_0(np__2))=np__2)).#trainneg
cnf(i_0_57, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_116, plain, (k3_xcmplx_0(np__2,np__1)=np__2)).#trainneg
cnf(i_0_126, plain, (k3_xcmplx_0(np__2,np__0)=np__0)).#trainneg
cnf(i_0_146, plain, (k3_xcmplx_0(np__1,np__2)=np__2)).#trainneg
cnf(i_0_102, plain, (k3_xcmplx_0(np__1,np__1)=np__1)).#trainneg
cnf(i_0_151, plain, (k3_xcmplx_0(np__1,np__0)=np__0)).#trainneg
cnf(i_0_110, plain, (k3_xcmplx_0(np__1,k1_xcmplx_0)=k1_xcmplx_0)).#trainneg
cnf(i_0_156, plain, (k3_xcmplx_0(np__0,np__2)=np__0)).#trainneg
cnf(i_0_145, plain, (k3_xcmplx_0(np__0,np__1)=np__0)).#trainneg
cnf(i_0_125, plain, (k3_xcmplx_0(np__0,np__0)=np__0)).#trainneg
cnf(i_0_75, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_88, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_31, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_59, plain, (m1_subset_1(k7_complex1,k2_numbers))).#trainneg
cnf(i_0_197, plain, (v6_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_129, plain, (k3_xcmplx_0(np__0,k1_xcmplx_0)=np__0)).#trainneg
cnf(i_0_127, plain, (k3_xcmplx_0(k1_xcmplx_0,np__1)=k1_xcmplx_0)).#trainneg
cnf(i_0_64, plain, (k2_xcmplx_0(np__1,np__1)=np__2)).#trainneg
cnf(i_0_171, plain, (k2_xcmplx_0(np__0,np__2)=np__2)).#trainneg
cnf(i_0_97, plain, (k2_xcmplx_0(np__0,np__1)=np__1)).#trainneg
cnf(i_0_142, plain, (k2_xcmplx_0(np__0,np__0)=np__0)).#trainneg
cnf(i_0_95, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_120, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_69, plain, (k4_xcmplx_0(np__1)=k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))).#trainneg
cnf(i_0_177, plain, (k2_xcmplx_0(np__2,k4_xcmplx_0(np__2))=np__0)).#trainneg
cnf(i_0_121, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_111, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_137, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_154, plain, (m2_subset_1(k6_numbers,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_48, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_91, plain, (v1_funct_2(k19_sin_cos,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_98, plain, (v2_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_118, plain, (v4_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_103, plain, (v3_membered(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_109, plain, (v5_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_164, plain, (k2_xcmplx_0(np__2,k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_89, plain, (k2_xcmplx_0(np__1,k4_xcmplx_0(np__1))=np__0)).#trainneg
cnf(i_0_200, plain, (k2_xcmplx_0(k4_xcmplx_0(np__2),np__2)=np__0)).#trainneg
cnf(i_0_199, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),np__2)=np__1)).#trainneg
cnf(i_0_163, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),np__1)=np__0)).#trainneg
cnf(i_0_147, plain, (m1_subset_1(esk2_1(X1),X1))).#trainneg
cnf(i_0_66, plain, (v1_funct_2(k16_sin_cos,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_117, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),k1_xcmplx_0)=k4_xcmplx_0(k1_xcmplx_0))).#trainneg
cnf(i_0_158, plain, (k2_xcmplx_0(np__0,k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_161, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),np__0)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_130, plain, (k3_xcmplx_0(k1_xcmplx_0,np__2)=k3_xcmplx_0(np__2,k1_xcmplx_0))).#trainneg
cnf(i_0_124, plain, (k2_xcmplx_0(k1_xcmplx_0,np__1)=k2_xcmplx_0(np__1,k1_xcmplx_0))).#trainneg
cnf(i_0_72, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_141, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_198, plain, (k2_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0)=k3_xcmplx_0(np__2,k1_xcmplx_0))).#trainneg
cnf(i_0_94, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_33, plain, (k14_sin_cos(X1)=k15_sin_cos(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_54, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k2_numbers))).#trainneg
cnf(i_0_90, plain, (m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_65, plain, (m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_42, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_56, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_58, plain, (v1_xcmplx_0(k14_sin_cos(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_162, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(np__1),k1_xcmplx_0))=k1_xcmplx_0)).#trainneg
cnf(i_0_61, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_87, plain, (k3_xcmplx_0(X1,k6_numbers)=k6_numbers|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_170, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_136, plain, (k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(k4_xcmplx_0(np__1),k1_xcmplx_0))=np__1)).#trainneg
cnf(i_0_174, plain, (k3_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(np__2),k1_xcmplx_0),k1_xcmplx_0)=np__2)).#trainneg
cnf(i_0_144, plain, (k3_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(np__1),k1_xcmplx_0),k1_xcmplx_0)=np__1)).#trainneg
cnf(i_0_105, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_159, plain, (k3_xcmplx_0(k3_xcmplx_0(np__2,k1_xcmplx_0),np__1)=k3_xcmplx_0(np__2,k1_xcmplx_0))).#trainneg
cnf(i_0_134, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_140, plain, (v1_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_numbers)))).#trainneg
cnf(i_0_165, plain, (k3_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(np__1),k1_xcmplx_0),k3_xcmplx_0(k4_xcmplx_0(np__1),k1_xcmplx_0))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_99, plain, (m1_subset_1(k15_sin_cos(X1),k2_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_112, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_27, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_93, plain, (m1_subset_1(k21_sin_cos(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_96, plain, (m1_subset_1(k18_sin_cos(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_138, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_153, plain, (v1_xboole_0(X1)|~v1_xcmplx_0(X1)|~v1_xboole_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_115, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_166, plain, (v1_int_1(X2)|~v5_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_150, plain, (v1_membered(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4, plain, (k1_seq_1(k19_sin_cos,k4_xcmplx_0(X1))=k1_seq_1(k19_sin_cos,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_135, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_172, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3, plain, (k1_real_1(k1_seq_1(k16_sin_cos,X1))=k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_132, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_82, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_123, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_122, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_84, plain, (k9_binop_2(X1,X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_51, plain, (k9_binop_2(X1,X2)=k9_binop_2(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_101, plain, (k3_binop_2(X1,X2)=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_74, plain, (k3_binop_2(X1,X2)=k3_binop_2(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_50, plain, (m1_subset_1(k5_binop_2(X1,X2),k2_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_86, plain, (m1_subset_1(k9_binop_2(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_63, plain, (m1_subset_1(k3_binop_2(X1,X2),k2_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_26, plain, (m1_subset_1(k8_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_104, plain, (m1_subset_1(k7_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_175, plain, (v1_partfun1(X1,X2)|~v1_funct_2(X1,X2,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2))))).#trainneg
cnf(i_0_77, plain, (k1_funct_1(X1,X2)=k1_seq_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_114, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_157, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xboole_0(k3_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_133, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_155, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_21, plain, (k8_real_1(X1,X2)=k8_real_1(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_176, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_143, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_131, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_173, plain, (v1_xboole_0(X1)|v1_partfun1(X2,X3)|~v1_funct_2(X2,X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_149, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,X3),k3_xcmplx_0(X2,X3))=k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_168, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|~v1_xboole_0(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_79, plain, (k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_128, plain, (v1_xboole_0(X1)|m1_subset_1(k3_funct_2(X1,X3,X2,X4),X3)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_207, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk2_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_212, plain, (v1_relat_1(k19_sin_cos))).#trainneg
cnf(i_0_57, plain, (k4_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))=np__1)).#trainneg
cnf(i_0_203, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_481, plain, (v4_relat_1(k16_sin_cos,k1_numbers))).#trainneg
cnf(i_0_732, plain, (~v1_xboole_0(k16_sin_cos))).#trainneg
cnf(i_0_733, plain, (~v1_xboole_0(k19_sin_cos))).#trainneg
cnf(i_0_204, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_205, plain, (v3_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_354, negated_conjecture, (k15_sin_cos(k5_binop_2(k11_binop_2(k3_xcmplx_0(np__2,k32_sin_cos),esk1_0),k1_xcmplx_0))!=np__1|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_274, plain, (v1_xreal_0(k21_sin_cos(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_506, plain, (v1_xreal_0(k9_binop_2(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_215, plain, (v1_relat_1(k16_sin_cos))).#trainneg
cnf(i_0_218, negated_conjecture, (k15_sin_cos(k5_binop_2(k11_binop_2(esk1_0,k11_binop_2(np__2,k32_sin_cos)),k1_xcmplx_0))!=np__1|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos))|~v1_xreal_0(esk1_0))).#trainneg
cnf(i_0_208, plain, (X1=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_482, plain, (v4_relat_1(k19_sin_cos,k1_numbers))).#trainneg
cnf(i_0_283, plain, (~v1_xcmplx_0(np__1)|~v1_xboole_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0)))).#trainneg
cnf(i_0_303, negated_conjecture, (v1_xxreal_0(esk1_0)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_209, plain, (v1_relat_1(esk2_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_211, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k19_sin_cos))).#trainneg
cnf(i_0_516, plain, (v1_xcmplx_0(k3_binop_2(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_348, plain, (k1_seq_1(k19_sin_cos,k4_xcmplx_0(X1))=k1_seq_1(k19_sin_cos,X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_735, plain, (k3_funct_2(k1_numbers,k1_numbers,k16_sin_cos,X1)=k1_funct_1(k16_sin_cos,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_490, plain, (v1_xreal_0(k11_binop_2(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_577, plain, (m1_subset_1(k8_real_1(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_364, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k11_binop_2(k11_binop_2(np__2,k32_sin_cos),esk1_0),k1_xcmplx_0))!=np__1|~v1_xcmplx_0(k11_binop_2(k11_binop_2(np__2,k32_sin_cos),esk1_0)))).#trainneg
cnf(i_0_232, plain, (v1_xreal_0(k1_xcmplx_0)|~v3_membered(k2_numbers))).#trainneg
cnf(i_0_484, plain, (v5_relat_1(k16_sin_cos,k1_numbers))).#trainneg
cnf(i_0_214, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k16_sin_cos))).#trainneg
cnf(i_0_485, plain, (v5_relat_1(k19_sin_cos,k1_numbers))).#trainneg
cnf(i_0_314, plain, (v1_xxreal_0(np__1)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_318, plain, (v1_xxreal_0(np__2)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_221, plain, (v1_xcmplx_0(np__2)|~v1_xcmplx_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_222, plain, (v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_353, negated_conjecture, (k15_sin_cos(k5_binop_2(k3_xcmplx_0(k11_binop_2(np__2,k32_sin_cos),esk1_0),k1_xcmplx_0))!=np__1|~v1_xreal_0(esk1_0)|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_522, plain, (v1_xreal_0(k8_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_276, plain, (v1_xreal_0(k18_sin_cos(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_736, plain, (k3_funct_2(k1_numbers,k1_numbers,k19_sin_cos,X1)=k1_funct_1(k19_sin_cos,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_252, plain, (v7_ordinal1(k1_xcmplx_0)|~v6_membered(k2_numbers))).#trainneg
cnf(i_0_537, plain, (v1_partfun1(k16_sin_cos,k1_numbers))).#trainneg
cnf(i_0_210, plain, (~r2_hidden(X1,k19_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_223, plain, (v1_xcmplx_0(k15_sin_cos(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_538, plain, (v1_partfun1(k19_sin_cos,k1_numbers))).#trainneg
cnf(i_0_261, plain, (v7_ordinal1(k32_sin_cos)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_321, negated_conjecture, (v1_int_1(esk1_0)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_332, plain, (v1_int_1(np__1)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_230, plain, (v1_xreal_0(esk2_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_574, plain, (k8_real_1(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_576, plain, (k8_real_1(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_534, plain, (k2_xcmplx_0(k21_sin_cos(X1),k3_xcmplx_0(k1_xcmplx_0,k18_sin_cos(X1)))=k15_sin_cos(k3_xcmplx_0(X1,k1_xcmplx_0))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(k18_sin_cos(X1)))).#trainneg
cnf(i_0_285, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_268, plain, (k4_xcmplx_0(k1_real_1(X1))=X1|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_465, plain, (v1_xcmplx_0(k3_xcmplx_0(np__2,k1_xcmplx_0)))).#trainneg
cnf(i_0_306, plain, (v1_xxreal_0(k1_xcmplx_0)|~v2_membered(k2_numbers))).#trainneg
cnf(i_0_213, plain, (~r2_hidden(X1,k16_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_235, plain, (v1_xreal_0(k16_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_289, plain, (v1_xcmplx_0(np__0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_290, plain, (v1_xcmplx_0(np__0)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_295, plain, (v1_xcmplx_0(np__1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_336, plain, (v1_int_1(np__2)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_474, plain, (v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_237, plain, (v1_xreal_0(k19_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_243, plain, (v1_xreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_486, plain, (k9_binop_2(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_527, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_500, plain, (v1_xcmplx_0(k5_binop_2(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_273, plain, (m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_598, plain, (k1_seq_1(k19_sin_cos,k7_real_1(k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1),X2))=k1_seq_1(k19_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_296, plain, (v1_xcmplx_0(np__1)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_307, plain, (v1_xxreal_0(np__0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_308, plain, (v1_xxreal_0(np__0)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_244, plain, (v1_xboole_0(esk2_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_299, plain, (v1_xcmplx_0(np__2)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_300, plain, (v1_xcmplx_0(np__2)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_313, plain, (v1_xxreal_0(np__1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_316, plain, (v1_xxreal_0(k32_sin_cos)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_247, plain, (v1_xreal_0(k15_sin_cos(X1))|~v3_membered(k2_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_250, plain, (v7_ordinal1(esk2_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_390, plain, (k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0)))).#trainneg
cnf(i_0_511, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_539, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_699, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,X2),X2)=k3_xcmplx_0(k2_xcmplx_0(X1,np__1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_488, plain, (k3_binop_2(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_581, plain, (m1_subset_1(k7_real_1(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_588, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X2)))=k2_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_724, plain, (k2_xcmplx_0(X1,k3_xcmplx_0(X2,X1))=k3_xcmplx_0(k2_xcmplx_0(np__1,X2),X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_18, plain, (k1_xboole_0=np__0)).#trainneg
cnf(i_0_255, plain, (v7_ordinal1(k16_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_257, plain, (v7_ordinal1(k19_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_260, plain, (v7_ordinal1(k15_sin_cos(X1))|~v6_membered(k2_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_860, negated_conjecture, (k15_sin_cos(k5_binop_2(k11_binop_2(k3_xcmplx_0(k32_sin_cos,np__2),esk1_0),k1_xcmplx_0))!=np__1)).#trainneg
cnf(i_0_780, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_317, plain, (v1_xxreal_0(np__2)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_264, plain, (v7_ordinal1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_324, plain, (v1_int_1(k1_xcmplx_0)|~v5_membered(k2_numbers))).#trainneg
cnf(i_0_325, plain, (v1_int_1(np__0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_790, plain, (v3_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_269, plain, (k1_real_1(k4_xcmplx_0(X1))=X1|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_271, plain, (v7_ordinal1(k1_real_1(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_994, plain, (m1_subset_1(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),k1_numbers))).#trainneg
cnf(i_0_962, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_907, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_326, plain, (v1_int_1(np__0)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_87, plain, (k3_xcmplx_0(X1,np__0)=np__0|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_275, plain, (v7_ordinal1(k21_sin_cos(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_331, plain, (v1_int_1(np__1)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_334, plain, (v1_int_1(k32_sin_cos)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_335, plain, (v1_int_1(np__2)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_277, plain, (v7_ordinal1(k18_sin_cos(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_280, plain, (v1_xboole_0(k4_xcmplx_0(k1_xcmplx_0))|~v1_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0))|~v1_xboole_0(k1_xcmplx_0))).#trainneg
cnf(i_0_993, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_345, plain, (v1_membered(k5_numbers)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_812, negated_conjecture, (k15_sin_cos(k5_binop_2(k1_xcmplx_0,k11_binop_2(esk1_0,k11_binop_2(np__2,k32_sin_cos))))!=np__1|~v1_xcmplx_0(k11_binop_2(esk1_0,k11_binop_2(np__2,k32_sin_cos)))|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_360, plain, (v2_membered(k5_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_301, plain, (v1_xcmplx_0(k5_numbers)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_768, plain, (m1_subset_1(X1,k1_numbers)|v1_xboole_0(k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_282, plain, (v1_xboole_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xboole_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_284, plain, (v1_xcmplx_0(esk2_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_931, plain, (k2_xcmplx_0(k1_seq_1(k19_sin_cos,X1),k3_xcmplx_0(k1_xcmplx_0,k18_sin_cos(X1)))=k15_sin_cos(k3_xcmplx_0(X1,k1_xcmplx_0))|~v1_xcmplx_0(k18_sin_cos(X1))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1198, plain, (k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(k1_xcmplx_0,X1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_938, plain, (k1_seq_1(k19_sin_cos,k1_real_1(X1))=k1_seq_1(k19_sin_cos,X1)|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_319, plain, (v1_xxreal_0(k5_numbers)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_778, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_286, plain, (v1_xcmplx_0(k18_sin_cos(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1093, plain, (m1_subset_1(k3_xcmplx_0(np__2,k1_xcmplx_0),k1_numbers)|~v1_xreal_0(k1_xcmplx_0))).#trainneg
cnf(i_0_337, plain, (v1_int_1(k5_numbers)|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1050, plain, (m1_subset_1(k3_xcmplx_0(X1,np__1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_287, plain, (v1_xcmplx_0(k21_sin_cos(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_291, plain, (v1_xcmplx_0(k1_real_1(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_797, plain, (k2_xcmplx_0(k21_sin_cos(X1),k3_xcmplx_0(k1_seq_1(k16_sin_cos,X1),k1_xcmplx_0))=k15_sin_cos(k3_xcmplx_0(X1,k1_xcmplx_0))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_920, plain, (k1_seq_1(k19_sin_cos,k7_real_1(k8_real_1(k3_xcmplx_0(k32_sin_cos,np__2),X1),X2))=k1_seq_1(k19_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_794, negated_conjecture, (k15_sin_cos(k5_binop_2(k1_xcmplx_0,k11_binop_2(k3_xcmplx_0(np__2,k32_sin_cos),esk1_0)))!=np__1|~v1_xcmplx_0(k11_binop_2(k3_xcmplx_0(np__2,k32_sin_cos),esk1_0)))).#trainneg
cnf(i_0_791, negated_conjecture, (k15_sin_cos(k5_binop_2(k11_binop_2(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos)),k1_xcmplx_0))!=np__1|~v1_xreal_0(esk1_0)|~v1_xreal_0(k3_xcmplx_0(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_782, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_292, plain, (v1_xcmplx_0(k16_sin_cos)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_339, plain, (v3_membered(k16_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_340, plain, (v3_membered(k19_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1053, plain, (m1_subset_1(k3_xcmplx_0(X1,np__2),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_315, plain, (v1_xxreal_0(k15_sin_cos(X1))|~v2_membered(k2_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_294, plain, (v1_xcmplx_0(k19_sin_cos)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_302, plain, (v1_xxreal_0(esk2_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_1215, plain, (k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(X1,k1_xcmplx_0))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_921, plain, (k1_seq_1(k16_sin_cos,k7_real_1(k8_real_1(k3_xcmplx_0(k32_sin_cos,np__2),X1),X2))=k1_seq_1(k16_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_802, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k1_xcmplx_0,k11_binop_2(k11_binop_2(np__2,k32_sin_cos),esk1_0)))!=np__1|~v1_xcmplx_0(k11_binop_2(k11_binop_2(np__2,k32_sin_cos),esk1_0)))).#trainneg
cnf(i_0_878, negated_conjecture, (k15_sin_cos(k5_binop_2(k3_xcmplx_0(esk1_0,k11_binop_2(np__2,k32_sin_cos)),k1_xcmplx_0))!=np__1|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos))|~v1_xcmplx_0(k11_binop_2(np__2,k32_sin_cos))|~v1_xcmplx_0(esk1_0))).#trainneg
cnf(i_0_770, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_968, negated_conjecture, (m1_subset_1(k3_xcmplx_0(X1,esk1_0),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_437, plain, (v1_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0))|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0)))).#trainneg
cnf(i_0_304, plain, (v1_xxreal_0(k18_sin_cos(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_991, plain, (m1_subset_1(k1_xcmplx_0,k1_numbers)|~m1_subset_1(k4_xcmplx_0(k1_xcmplx_0),k1_numbers))).#trainneg
cnf(i_0_343, plain, (v1_membered(k16_sin_cos)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_344, plain, (v1_membered(k19_sin_cos)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_333, plain, (v1_int_1(k15_sin_cos(X1))|~v5_membered(k2_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_305, plain, (v1_xxreal_0(k21_sin_cos(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_309, plain, (v1_xxreal_0(k1_real_1(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_862, negated_conjecture, (k15_sin_cos(k5_binop_2(k1_xcmplx_0,k11_binop_2(k3_xcmplx_0(k32_sin_cos,np__2),esk1_0)))!=np__1|~v1_xcmplx_0(k11_binop_2(k3_xcmplx_0(k32_sin_cos,np__2),esk1_0)))).#trainneg
cnf(i_0_801, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k1_xcmplx_0,k11_binop_2(k3_xcmplx_0(np__2,k32_sin_cos),esk1_0)))!=np__1|~v1_xcmplx_0(k11_binop_2(k3_xcmplx_0(np__2,k32_sin_cos),esk1_0)))).#trainneg
cnf(i_0_795, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k11_binop_2(k3_xcmplx_0(np__2,k32_sin_cos),esk1_0),k1_xcmplx_0))!=np__1|~v1_xcmplx_0(k11_binop_2(k3_xcmplx_0(np__2,k32_sin_cos),esk1_0)))).#trainneg
cnf(i_0_841, plain, (v1_xreal_0(k8_real_1(X1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_771, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_972, plain, (m1_subset_1(k3_xcmplx_0(X1,np__0),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_283, plain, (~v1_xboole_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0)))).#trainneg
cnf(i_0_310, plain, (v1_xxreal_0(k16_sin_cos)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1817, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers))).#trainneg
cnf(i_0_1515, plain, (~v1_xcmplx_0(k4_xcmplx_0(np__2))|~v1_xboole_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_1368, plain, (v3_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_312, plain, (v1_xxreal_0(k19_sin_cos)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_320, plain, (v1_int_1(esk2_1(X1))|~v5_membered(X1))).#trainneg
cnf(i_0_1753, plain, (m1_subset_1(k3_xcmplx_0(k32_sin_cos,np__2),k1_numbers))).#trainneg
cnf(i_0_1556, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1379, plain, (k1_real_1(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))=np__1)).#trainneg
cnf(i_0_1802, negated_conjecture, (k15_sin_cos(k5_binop_2(k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos)),k1_xcmplx_0))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos))|~v1_xreal_0(k3_xcmplx_0(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_1398, plain, (v1_xreal_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0)))).#trainneg
cnf(i_0_1467, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1375, plain, (k1_real_1(np__0)=np__0)).#trainneg
cnf(i_0_724, plain, (k2_xcmplx_0(X1,k3_xcmplx_0(X2,X1))=k3_xcmplx_0(k2_xcmplx_0(np__1,X2),X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_322, plain, (v1_int_1(k18_sin_cos(X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1428, plain, (v1_xreal_0(k1_xcmplx_0)|~m1_subset_1(k4_xcmplx_0(k1_xcmplx_0),k1_numbers))).#trainneg
cnf(i_0_1366, plain, (v7_ordinal1(k5_numbers)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1485, plain, (v1_xcmplx_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1508, plain, (X1=k4_xcmplx_0(X2)|~v1_xboole_0(X1)|~v1_xcmplx_0(X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1507, plain, (k4_xcmplx_0(X1)=np__0|~v1_xcmplx_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_323, plain, (v1_int_1(k21_sin_cos(X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_327, plain, (v1_int_1(k1_real_1(X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1731, negated_conjecture, (k15_sin_cos(k5_binop_2(k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos)),k1_xcmplx_0))!=np__1|~v1_xreal_0(k3_xcmplx_0(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_1804, negated_conjecture, (k15_sin_cos(k5_binop_2(k1_xcmplx_0,k3_xcmplx_0(esk1_0,k11_binop_2(np__2,k32_sin_cos))))!=np__1|~v1_xcmplx_0(k11_binop_2(np__2,k32_sin_cos))|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos))|~v1_xcmplx_0(k3_xcmplx_0(esk1_0,k11_binop_2(np__2,k32_sin_cos))))).#trainneg
cnf(i_0_1481, negated_conjecture, (k15_sin_cos(k5_binop_2(k1_xcmplx_0,k3_xcmplx_0(esk1_0,k11_binop_2(np__2,k32_sin_cos))))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(esk1_0,k11_binop_2(np__2,k32_sin_cos)))|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_1727, negated_conjecture, (k15_sin_cos(k5_binop_2(k1_xcmplx_0,k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos))))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos)))|~v1_xreal_0(k3_xcmplx_0(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_772, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_222, plain, (v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0)))).#trainneg
cnf(i_0_699, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,X2),X2)=k3_xcmplx_0(k2_xcmplx_0(X1,np__1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_328, plain, (v1_int_1(k16_sin_cos)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1663, plain, (v1_xxreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1666, plain, (v1_int_1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1869, plain, (v7_ordinal1(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_330, plain, (v1_int_1(k19_sin_cos)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_338, plain, (v3_membered(esk2_1(k1_zfmisc_1(X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_1345, negated_conjecture, (k15_sin_cos(k5_binop_2(k11_binop_2(esk1_0,k3_xcmplx_0(k32_sin_cos,np__2)),k1_xcmplx_0))!=np__1|~v1_xreal_0(k3_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_1553, plain, (k1_seq_1(k19_sin_cos,k4_xcmplx_0(X1))=k1_seq_1(k19_sin_cos,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1806, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(esk1_0,k11_binop_2(np__2,k32_sin_cos))))!=np__1|~v1_xcmplx_0(k11_binop_2(np__2,k32_sin_cos))|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos))|~v1_xcmplx_0(k3_xcmplx_0(esk1_0,k11_binop_2(np__2,k32_sin_cos))))).#trainneg
cnf(i_0_1800, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(esk1_0,k11_binop_2(np__2,k32_sin_cos))))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(esk1_0,k11_binop_2(np__2,k32_sin_cos)))|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_1732, negated_conjecture, (k15_sin_cos(k5_binop_2(k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),esk1_0),k1_xcmplx_0))!=np__1|~v1_xreal_0(k3_xcmplx_0(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_1803, negated_conjecture, (k15_sin_cos(k5_binop_2(k3_xcmplx_0(esk1_0,k3_xcmplx_0(k32_sin_cos,np__2)),k1_xcmplx_0))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(k32_sin_cos,np__2))|~v1_xreal_0(k3_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_1480, negated_conjecture, (k15_sin_cos(k5_binop_2(k1_xcmplx_0,k11_binop_2(esk1_0,k3_xcmplx_0(k32_sin_cos,np__2))))!=np__1|~v1_xcmplx_0(k11_binop_2(esk1_0,k3_xcmplx_0(k32_sin_cos,np__2)))|~v1_xreal_0(k3_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_773, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1494, negated_conjecture, (m1_subset_1(k3_xcmplx_0(esk1_0,esk1_0),k1_numbers))).#trainneg
cnf(i_0_1376, plain, (k1_real_1(k1_xcmplx_0)=k4_xcmplx_0(k1_xcmplx_0)|~m1_subset_1(k4_xcmplx_0(k1_xcmplx_0),k1_numbers))).#trainneg
cnf(i_0_819, negated_conjecture, (k15_sin_cos(k5_binop_2(k3_xcmplx_0(esk1_0,k11_binop_2(np__2,k32_sin_cos)),k1_xcmplx_0))!=np__1|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_342, plain, (v1_membered(esk2_1(k1_zfmisc_1(X1)))|~v1_membered(X1))).#trainneg
cnf(i_0_358, plain, (v2_membered(k16_sin_cos)|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_359, plain, (v2_membered(k19_sin_cos)|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_869, plain, (v1_xreal_0(k1_xcmplx_0)|~m1_subset_1(k2_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1393, plain, (k3_xcmplx_0(np__1,X1)=k3_xcmplx_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_346, plain, (k1_seq_1(k19_sin_cos,k4_xcmplx_0(k1_xcmplx_0))=k1_seq_1(k19_sin_cos,k1_xcmplx_0)|~v1_xreal_0(k4_xcmplx_0(k1_xcmplx_0)))).#trainneg
cnf(i_0_347, plain, (k1_seq_1(k19_sin_cos,k4_xcmplx_0(np__2))=k1_seq_1(k19_sin_cos,np__2)|~v1_xreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_1479, negated_conjecture, (k15_sin_cos(k5_binop_2(k1_xcmplx_0,k11_binop_2(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos))))!=np__1|~v1_xcmplx_0(k11_binop_2(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos)))|~v1_xreal_0(k3_xcmplx_0(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_1347, negated_conjecture, (k15_sin_cos(k5_binop_2(k3_xcmplx_0(esk1_0,k3_xcmplx_0(k32_sin_cos,np__2)),k1_xcmplx_0))!=np__1|~v1_xreal_0(k3_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_1841, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos))))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos)))|~v1_xreal_0(k3_xcmplx_0(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_774, plain, (v1_int_1(X1)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1498, negated_conjecture, (m1_subset_1(k3_xcmplx_0(np__0,esk1_0),k1_numbers))).#trainneg
cnf(i_0_1377, plain, (k1_real_1(np__2)=k4_xcmplx_0(np__2)|~m1_subset_1(k4_xcmplx_0(np__2),k1_numbers))).#trainneg
cnf(i_0_349, plain, (k1_seq_1(k19_sin_cos,k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))=k1_seq_1(k19_sin_cos,np__1)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_1822, plain, (v1_membered(k16_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1824, plain, (v1_membered(k19_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1397, plain, (k3_xcmplx_0(np__2,X1)=k3_xcmplx_0(X1,np__2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1406, negated_conjecture, (k2_xcmplx_0(esk1_0,X1)=k2_xcmplx_0(X1,esk1_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_350, plain, (v1_xboole_0(esk2_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_357, plain, (v2_membered(esk2_1(k1_zfmisc_1(X1)))|~v2_membered(X1))).#trainneg
cnf(i_0_1726, negated_conjecture, (k15_sin_cos(k5_binop_2(k1_xcmplx_0,k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),esk1_0)))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),esk1_0))|~v1_xreal_0(k3_xcmplx_0(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_1845, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos)),k1_xcmplx_0))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos)))|~v1_xreal_0(k3_xcmplx_0(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_1799, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(k11_binop_2(np__2,k32_sin_cos),esk1_0)))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(k11_binop_2(np__2,k32_sin_cos),esk1_0))|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_1734, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k11_binop_2(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos)),k1_xcmplx_0))!=np__1|~v1_xreal_0(k3_xcmplx_0(np__2,k32_sin_cos))|~v1_xcmplx_0(k11_binop_2(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos))))).#trainneg
cnf(i_0_1369, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2034, plain, (k2_xcmplx_0(np__2,np__0)=np__2)).#trainneg
cnf(i_0_1892, negated_conjecture, (v1_xxreal_0(esk1_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_390, plain, (k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_361, plain, (k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1))=k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1893, negated_conjecture, (v1_int_1(esk1_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_362, plain, (k1_real_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_363, plain, (m1_subset_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_437, plain, (v1_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0)))).#trainneg
cnf(i_0_2213, negated_conjecture, (k3_xcmplx_0(esk1_0,np__2)=k3_xcmplx_0(np__2,esk1_0))).#trainneg
cnf(i_0_2238, negated_conjecture, (k15_sin_cos(k5_binop_2(k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos)),k1_xcmplx_0))!=np__1|~v1_xreal_0(k3_xcmplx_0(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_2041, plain, (k2_xcmplx_0(X1,X1)=k3_xcmplx_0(np__2,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2222, negated_conjecture, (v1_xreal_0(k3_xcmplx_0(esk1_0,esk1_0)))).#trainneg
cnf(i_0_818, negated_conjecture, (k15_sin_cos(k5_binop_2(k11_binop_2(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos)),k1_xcmplx_0))!=np__1|~v1_xreal_0(k3_xcmplx_0(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_367, plain, (k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))=k4_xcmplx_0(k1_xcmplx_0)|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0)))).#trainneg
cnf(i_0_1413, plain, (k2_xcmplx_0(np__1,X1)=k2_xcmplx_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_368, plain, (k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(k4_xcmplx_0(np__2),k1_xcmplx_0))=np__2|~v1_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(np__2),k1_xcmplx_0)))).#trainneg
cnf(i_0_369, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(np__2,k1_xcmplx_0))=k3_xcmplx_0(np__2,k1_xcmplx_0)|~v1_xcmplx_0(k3_xcmplx_0(np__2,k1_xcmplx_0))|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_2031, plain, (k2_xcmplx_0(k1_xcmplx_0,np__0)=k1_xcmplx_0)).#trainneg
cnf(i_0_1953, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1955, plain, (k1_real_1(np__1)=k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))).#trainneg
cnf(i_0_2367, negated_conjecture, (v1_xreal_0(k3_xcmplx_0(np__0,esk1_0)))).#trainneg
cnf(i_0_2155, plain, (v1_xxreal_0(k5_numbers)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_384, plain, (k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k4_xcmplx_0(np__2)),k1_xcmplx_0)=np__2|~v1_xcmplx_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_2156, plain, (v1_int_1(k5_numbers)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1899, plain, (v1_xreal_0(k5_numbers)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2066, plain, (v1_xcmplx_0(k5_numbers)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1416, plain, (k2_xcmplx_0(k32_sin_cos,X1)=k2_xcmplx_0(X1,k32_sin_cos)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1417, plain, (k2_xcmplx_0(np__2,X1)=k2_xcmplx_0(X1,np__2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1559, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1850, negated_conjecture, (k15_sin_cos(k5_binop_2(k1_xcmplx_0,k3_xcmplx_0(k3_xcmplx_0(k32_sin_cos,np__2),esk1_0)))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(k3_xcmplx_0(k32_sin_cos,np__2),esk1_0))|~v1_xreal_0(k3_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_2724, plain, (m1_subset_1(k1_xcmplx_0,k1_numbers)|~v1_xreal_0(k1_xcmplx_0))).#trainneg
cnf(i_0_2131, plain, (k2_xcmplx_0(np__0,k1_xcmplx_0)=k1_xcmplx_0)).#trainneg
cnf(i_0_2625, plain, (m1_subset_1(k3_xcmplx_0(np__2,X1),k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_491, plain, (v7_ordinal1(k11_binop_2(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2606, negated_conjecture, (k15_sin_cos(k5_binop_2(k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)),k1_xcmplx_0))!=np__1)).#trainneg
cnf(i_0_2605, negated_conjecture, (k15_sin_cos(k5_binop_2(k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,esk1_0)),k1_xcmplx_0))!=np__1)).#trainneg
cnf(i_0_1726, negated_conjecture, (k15_sin_cos(k5_binop_2(k1_xcmplx_0,k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),esk1_0)))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),esk1_0)))).#trainneg
cnf(i_0_1900, plain, (v1_xreal_0(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2501, negated_conjecture, (m1_subset_1(k4_xcmplx_0(esk1_0),k1_numbers))).#trainneg
cnf(i_0_2224, negated_conjecture, (v1_xcmplx_0(k3_xcmplx_0(esk1_0,esk1_0))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_2655, plain, (v1_xboole_0(k1_xcmplx_0)|~v1_xboole_0(k4_xcmplx_0(k1_xcmplx_0)))).#trainneg
cnf(i_0_492, plain, (v1_xcmplx_0(k11_binop_2(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_493, plain, (v1_xxreal_0(k11_binop_2(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_461, plain, (v1_xcmplx_0(k4_xcmplx_0(np__2))|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0)))).#trainneg
cnf(i_0_1841, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos))))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos))))).#trainneg
cnf(i_0_2106, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2489, negated_conjecture, (v1_xcmplx_0(k3_xcmplx_0(k32_sin_cos,esk1_0)))).#trainneg
cnf(i_0_1479, negated_conjecture, (k15_sin_cos(k5_binop_2(k1_xcmplx_0,k11_binop_2(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos))))!=np__1|~v1_xcmplx_0(k11_binop_2(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos))))).#trainneg
cnf(i_0_2369, negated_conjecture, (v1_xcmplx_0(k3_xcmplx_0(np__0,esk1_0))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1895, plain, (v3_membered(k16_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1896, plain, (v3_membered(k19_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_494, plain, (v1_int_1(k11_binop_2(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_498, plain, (v1_xreal_0(k5_binop_2(X1,X2))|~v3_membered(k2_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_466, plain, (v1_xcmplx_0(k2_xcmplx_0(np__1,k1_xcmplx_0))|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_1484, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k11_binop_2(esk1_0,k11_binop_2(np__2,k32_sin_cos)),k1_xcmplx_0))!=np__1|~v1_xcmplx_0(k11_binop_2(esk1_0,k11_binop_2(np__2,k32_sin_cos)))|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_2157, plain, (v1_xcmplx_0(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2599, negated_conjecture, (v1_xcmplx_0(k3_xcmplx_0(np__2,esk1_0)))).#trainneg
cnf(i_0_2930, negated_conjecture, (k3_xcmplx_0(np__0,esk1_0)=np__0)).#trainneg
cnf(i_0_1734, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k11_binop_2(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos)),k1_xcmplx_0))!=np__1|~v1_xcmplx_0(k11_binop_2(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos))))).#trainneg
cnf(i_0_2244, plain, (v2_membered(k16_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_2245, plain, (v2_membered(k19_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_499, plain, (v7_ordinal1(k5_binop_2(X1,X2))|~v6_membered(k2_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_501, plain, (v1_xxreal_0(k5_binop_2(X1,X2))|~v2_membered(k2_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1483, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k1_xcmplx_0,k11_binop_2(esk1_0,k11_binop_2(np__2,k32_sin_cos))))!=np__1|~v1_xcmplx_0(k11_binop_2(esk1_0,k11_binop_2(np__2,k32_sin_cos)))|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_1555, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1839, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k1_xcmplx_0,k11_binop_2(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos))))!=np__1|~v1_xcmplx_0(k11_binop_2(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos)))|~v1_xreal_0(k3_xcmplx_0(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_2221, plain, (v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_280, plain, (v1_xboole_0(k4_xcmplx_0(k1_xcmplx_0))|~v1_xboole_0(k1_xcmplx_0))).#trainneg
cnf(i_0_480, plain, (v4_relat_1(esk2_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_2386, plain, (v1_membered(k16_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_2388, plain, (v1_membered(k19_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_502, plain, (v1_int_1(k5_binop_2(X1,X2))|~v5_membered(k2_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2072, plain, (k4_xcmplx_0(X1)=k4_xcmplx_0(X2)|~v1_xcmplx_0(X2)|~v1_xboole_0(X2)|~v1_xcmplx_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_483, plain, (v5_relat_1(esk2_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_1554, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2285, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2555, plain, (k1_real_1(k1_seq_1(k16_sin_cos,X1))=k1_seq_1(k16_sin_cos,k1_real_1(X1))|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_real_1(X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2366, plain, (v1_int_1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1505, negated_conjecture, (m1_subset_1(k3_xcmplx_0(k32_sin_cos,esk1_0),k1_numbers))).#trainneg
cnf(i_0_2040, plain, (k2_xcmplx_0(np__2,np__2)=k3_xcmplx_0(np__2,np__2))).#trainneg
cnf(i_0_504, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k2_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2473, plain, (v1_xboole_0(k32_sin_cos)|~v1_xboole_0(k3_xcmplx_0(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_2596, negated_conjecture, (v1_xboole_0(esk1_0)|~v1_xboole_0(k3_xcmplx_0(np__2,esk1_0)))).#trainneg
cnf(i_0_1066, plain, (v1_xcmplx_0(esk2_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_1047, plain, (m1_subset_1(k3_xcmplx_0(X1,np__0),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_507, plain, (v7_ordinal1(k9_binop_2(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2907, negated_conjecture, (k3_xcmplx_0(esk1_0,k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))=k4_xcmplx_0(esk1_0))).#trainneg
cnf(i_0_3152, negated_conjecture, (k15_sin_cos(k5_binop_2(k1_xcmplx_0,k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0))))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0))))).#trainneg
cnf(i_0_3590, negated_conjecture, (m1_subset_1(esk1_0,k2_numbers))).#trainneg
cnf(i_0_3579, plain, (m1_subset_1(np__0,k2_numbers))).#trainneg
cnf(i_0_3180, negated_conjecture, (v1_xreal_0(k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_1850, negated_conjecture, (k15_sin_cos(k5_binop_2(k1_xcmplx_0,k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos))))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos))))).#trainneg
cnf(i_0_508, plain, (v1_xcmplx_0(k9_binop_2(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1515, plain, (~v1_xboole_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_3646, negated_conjecture, (v1_xboole_0(esk1_0)|~v1_xboole_0(k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_3596, plain, (m1_subset_1(X1,k2_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_509, plain, (v1_xxreal_0(k9_binop_2(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_510, plain, (v1_int_1(k9_binop_2(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3154, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0))))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0))))).#trainneg
cnf(i_0_3153, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)),k1_xcmplx_0))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0))))).#trainneg
cnf(i_0_3160, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,esk1_0)),k1_xcmplx_0))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,esk1_0))))).#trainneg
cnf(i_0_3161, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,esk1_0))))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,esk1_0))))).#trainneg
cnf(i_0_3580, plain, (m1_subset_1(np__1,k2_numbers))).#trainneg
cnf(i_0_384, plain, (k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k4_xcmplx_0(np__2)),k1_xcmplx_0)=np__2)).#trainneg
cnf(i_0_514, plain, (v1_xreal_0(k3_binop_2(X1,X2))|~v3_membered(k2_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3325, plain, (v2_membered(k16_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_3327, plain, (v2_membered(k19_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_3445, plain, (v1_xxreal_0(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_515, plain, (v7_ordinal1(k3_binop_2(X1,X2))|~v6_membered(k2_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_517, plain, (v1_xxreal_0(k3_binop_2(X1,X2))|~v2_membered(k2_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3597, plain, (m1_subset_1(k3_xcmplx_0(np__2,k32_sin_cos),k2_numbers))).#trainneg
cnf(i_0_3372, plain, (k3_xcmplx_0(X1,k1_xcmplx_0)=k3_xcmplx_0(k1_xcmplx_0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_518, plain, (v1_int_1(k3_binop_2(X1,X2))|~v5_membered(k2_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3595, plain, (m1_subset_1(np__2,k2_numbers))).#trainneg
cnf(i_0_3547, plain, (v1_int_1(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_519, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k2_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_523, plain, (v7_ordinal1(k8_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3260, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k3_xcmplx_0(esk1_0,k11_binop_2(np__2,k32_sin_cos)),k1_xcmplx_0))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(esk1_0,k11_binop_2(np__2,k32_sin_cos)))|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_3649, negated_conjecture, (v1_xcmplx_0(k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_524, plain, (v1_xcmplx_0(k8_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3446, plain, (k4_xcmplx_0(k1_xcmplx_0)=np__0|~v1_xboole_0(k1_xcmplx_0))).#trainneg
cnf(i_0_961, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1340, plain, (v7_ordinal1(k16_sin_cos)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3138, plain, (v1_xcmplx_0(k11_binop_2(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_525, plain, (v1_xxreal_0(k8_real_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_526, plain, (v1_int_1(k8_real_1(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2538, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)))=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3305, negated_conjecture, (k2_xcmplx_0(esk1_0,np__0)=esk1_0)).#trainneg
cnf(i_0_3315, negated_conjecture, (k2_xcmplx_0(np__0,esk1_0)=esk1_0)).#trainneg
cnf(i_0_528, plain, (v1_xreal_0(k7_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1341, plain, (v7_ordinal1(k19_sin_cos)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1486, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k2_numbers)))).#trainneg
cnf(i_0_529, plain, (v7_ordinal1(k7_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_530, plain, (v1_xcmplx_0(k7_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2405, plain, (k3_xcmplx_0(np__2,X1)=k3_xcmplx_0(X1,np__2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2591, negated_conjecture, (k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),esk1_0)=k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,esk1_0)))).#trainneg
cnf(i_0_1726, negated_conjecture, (k15_sin_cos(k5_binop_2(k1_xcmplx_0,k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,esk1_0))))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,esk1_0))))).#trainneg
cnf(i_0_3758, plain, (m1_subset_1(X1,k2_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3689, negated_conjecture, (v1_xxreal_0(esk1_0)|~v2_membered(k2_numbers))).#trainneg
cnf(i_0_531, plain, (v1_xxreal_0(k7_real_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3756, plain, (m1_subset_1(k32_sin_cos,k2_numbers))).#trainneg
cnf(i_0_532, plain, (v1_int_1(k7_real_1(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_533, plain, (v1_xcmplx_0(k15_sin_cos(k3_xcmplx_0(X1,k1_xcmplx_0)))|~v1_xcmplx_0(k3_xcmplx_0(k18_sin_cos(X1),k1_xcmplx_0))|~v1_xcmplx_0(k21_sin_cos(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3690, negated_conjecture, (v1_int_1(esk1_0)|~v5_membered(k2_numbers))).#trainneg
cnf(i_0_536, plain, (v1_partfun1(esk2_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))),X1)|~v1_funct_2(esk2_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))),X1,X1))).#trainneg
cnf(i_0_3694, plain, (v1_xxreal_0(np__0)|~v2_membered(k2_numbers))).#trainneg
cnf(i_0_3695, plain, (v1_int_1(np__0)|~v5_membered(k2_numbers))).#trainneg
cnf(i_0_3794, plain, (v1_xxreal_0(np__1)|~v2_membered(k2_numbers))).#trainneg
cnf(i_0_3795, plain, (v1_int_1(np__1)|~v5_membered(k2_numbers))).#trainneg
cnf(i_0_4163, plain, (k3_xcmplx_0(X2,k3_xcmplx_0(k3_xcmplx_0(X1,np__1),X3))=k3_xcmplx_0(k3_xcmplx_0(X2,X1),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4025, plain, (v1_xxreal_0(np__2)|~v2_membered(k2_numbers))).#trainneg
cnf(i_0_4026, plain, (v1_int_1(np__2)|~v5_membered(k2_numbers))).#trainneg
cnf(i_0_1737, plain, (v1_xcmplx_0(k16_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1756, plain, (v1_xcmplx_0(k19_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4009, plain, (v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4107, plain, (k1_real_1(k1_seq_1(k16_sin_cos,X1))=k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2422, negated_conjecture, (k2_xcmplx_0(esk1_0,X1)=k2_xcmplx_0(X1,esk1_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2577, plain, (m1_subset_1(k3_xcmplx_0(X1,k3_xcmplx_0(np__2,k32_sin_cos)),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3568, plain, (m1_subset_1(k4_xcmplx_0(k1_xcmplx_0),k2_numbers))).#trainneg
cnf(i_0_3181, negated_conjecture, (v7_ordinal1(k4_xcmplx_0(esk1_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2204, negated_conjecture, (k3_xcmplx_0(esk1_0,np__0)=k3_xcmplx_0(np__0,esk1_0))).#trainneg
cnf(i_0_1888, plain, (v1_xxreal_0(k16_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1897, plain, (v1_xreal_0(k16_sin_cos)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3344, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,np__1),X2)=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2535, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,k1_real_1(X1)))=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_real_1(X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4184, plain, (k3_xcmplx_0(np__2,X1)=k3_xcmplx_0(X1,np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3650, negated_conjecture, (m1_subset_1(k4_xcmplx_0(esk1_0),k2_numbers))).#trainneg
cnf(i_0_3183, negated_conjecture, (v1_xxreal_0(k4_xcmplx_0(esk1_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_573, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_2209, negated_conjecture, (k3_xcmplx_0(esk1_0,np__1)=k3_xcmplx_0(np__1,esk1_0))).#trainneg
cnf(i_0_1898, plain, (v1_xreal_0(k19_sin_cos)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1902, plain, (v1_xxreal_0(k19_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_582, plain, (k2_xcmplx_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3095, negated_conjecture, (k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k32_sin_cos))=k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)))).#trainneg
cnf(i_0_4598, negated_conjecture, (k15_sin_cos(np__0)=np__1)).#trainneg
cnf(i_0_2591, negated_conjecture, (k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),esk1_0)=k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)))).#trainneg
cnf(i_0_1726, negated_conjecture, (k15_sin_cos(k5_binop_2(k1_xcmplx_0,k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0))))!=np__1)).#trainneg
cnf(i_0_4346, plain, (v7_ordinal1(k32_sin_cos)|~v6_membered(k2_numbers))).#trainneg
cnf(i_0_3750, plain, (m1_subset_1(k4_xcmplx_0(np__2),k2_numbers))).#trainneg
cnf(i_0_4348, plain, (v1_xxreal_0(k32_sin_cos)|~v2_membered(k2_numbers))).#trainneg
cnf(i_0_583, plain, (v1_xcmplx_0(k4_xcmplx_0(k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4349, plain, (v1_int_1(k32_sin_cos)|~v5_membered(k2_numbers))).#trainneg
cnf(i_0_3161, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,esk1_0))))!=np__1)).#trainneg
cnf(i_0_2067, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_585, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),np__0)=k4_xcmplx_0(k2_xcmplx_0(X1,np__0))|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4735, plain, (k1_seq_1(k16_sin_cos,k1_real_1(X1))=k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_real_1(X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4584, negated_conjecture, (k2_xcmplx_0(esk1_0,esk1_0)=k3_xcmplx_0(np__2,esk1_0))).#trainneg
cnf(i_0_4603, plain, (k2_xcmplx_0(k1_seq_1(k19_sin_cos,X1),k3_xcmplx_0(k1_xcmplx_0,k1_seq_1(k16_sin_cos,X1)))=k15_sin_cos(k3_xcmplx_0(X1,k1_xcmplx_0))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_xcmplx_0,k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_4639, plain, (v1_xreal_0(k4_xcmplx_0(k1_xcmplx_0))|~v3_membered(k2_numbers))).#trainneg
cnf(i_0_586, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(k1_xcmplx_0)))=k2_xcmplx_0(k4_xcmplx_0(X1),k1_xcmplx_0)|~v1_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2153, plain, (v1_int_1(k16_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2154, plain, (v1_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_587, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(np__2)))=k2_xcmplx_0(k4_xcmplx_0(X1),np__2)|~v1_xcmplx_0(k4_xcmplx_0(np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_589, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))=k4_xcmplx_0(k2_xcmplx_0(X1,np__1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4743, plain, (k1_seq_1(k19_sin_cos,k1_seq_1(k16_sin_cos,k1_real_1(X1)))=k1_seq_1(k19_sin_cos,k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_real_1(X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4562, plain, (k1_real_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)),k1_numbers))).#trainneg
cnf(i_0_4326, plain, (v7_ordinal1(X1)|~v6_membered(k2_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3184, negated_conjecture, (v1_int_1(k4_xcmplx_0(esk1_0))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_590, plain, (k2_xcmplx_0(np__0,k4_xcmplx_0(X1))=k4_xcmplx_0(k2_xcmplx_0(np__0,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_2260, plain, (k3_xcmplx_0(k32_sin_cos,np__1)=k3_xcmplx_0(np__1,k32_sin_cos))).#trainneg
cnf(i_0_2158, plain, (v1_int_1(k19_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4530, plain, (v1_xcmplx_0(k16_sin_cos)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_591, plain, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0),X1))=k2_xcmplx_0(k1_xcmplx_0,k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0)))).#trainneg
cnf(i_0_592, plain, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(np__2),X1))=k2_xcmplx_0(np__2,k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_2279, plain, (k3_xcmplx_0(k32_sin_cos,X1)=k3_xcmplx_0(X1,k32_sin_cos)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3589, plain, (m1_subset_1(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),k2_numbers))).#trainneg
cnf(i_0_3276, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,np__1),X2)=k3_xcmplx_0(X2,X1)|~m1_subset_1(k3_xcmplx_0(X1,np__1),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_5257, plain, (k3_xcmplx_0(np__1,k32_sin_cos)=k32_sin_cos)).#trainneg
cnf(i_0_593, plain, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X1),X2))=k2_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4936, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(np__2,k3_xcmplx_0(k3_xcmplx_0(k32_sin_cos,esk1_0),k1_xcmplx_0)))!=np__1)).#trainneg
cnf(i_0_4937, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k3_xcmplx_0(k32_sin_cos,esk1_0),k3_xcmplx_0(np__2,k1_xcmplx_0)))!=np__1)).#trainneg
cnf(i_0_594, plain, (k2_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),k4_xcmplx_0(X1))=k4_xcmplx_0(k2_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_595, plain, (v1_funct_2(esk2_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1,X2)|~v1_partfun1(esk2_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_5020, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,k1_real_1(X1)))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4328, plain, (v1_xxreal_0(X1)|~v2_membered(k2_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5365, plain, (k3_xcmplx_0(np__0,k32_sin_cos)=np__0)).#trainneg
cnf(i_0_4941, plain, (v1_xreal_0(k4_xcmplx_0(np__2))|~v3_membered(k2_numbers))).#trainneg
cnf(i_0_4938, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(np__2,k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(k32_sin_cos,esk1_0))))!=np__1)).#trainneg
cnf(i_0_599, plain, (k1_seq_1(k19_sin_cos,k7_real_1(k8_real_1(X1,k8_real_1(np__2,k32_sin_cos)),X2))=k1_seq_1(k19_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(k8_real_1(np__2,k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_600, plain, (k1_seq_1(k19_sin_cos,k7_real_1(k3_xcmplx_0(k8_real_1(np__2,k32_sin_cos),X1),X2))=k1_seq_1(k19_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(k8_real_1(np__2,k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_5021, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5027, plain, (k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1))=k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4329, plain, (v1_int_1(X1)|~v5_membered(k2_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3186, negated_conjecture, (v1_xcmplx_0(k3_xcmplx_0(esk1_0,esk1_0)))).#trainneg
cnf(i_0_4640, plain, (v7_ordinal1(k4_xcmplx_0(k1_xcmplx_0))|~v6_membered(k2_numbers))).#trainneg
cnf(i_0_619, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2270, plain, (k3_xcmplx_0(k32_sin_cos,np__0)=k3_xcmplx_0(np__0,k32_sin_cos))).#trainneg
cnf(i_0_3548, negated_conjecture, (v1_xreal_0(k3_xcmplx_0(k32_sin_cos,esk1_0)))).#trainneg
cnf(i_0_3555, plain, (v1_xcmplx_0(k3_xcmplx_0(np__2,np__2)))).#trainneg
cnf(i_0_601, plain, (k1_seq_1(k19_sin_cos,k7_real_1(k8_real_1(X1,k8_real_1(np__2,k32_sin_cos)),X2))=k1_seq_1(k19_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(k8_real_1(np__2,k32_sin_cos))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_602, plain, (k1_seq_1(k19_sin_cos,k7_real_1(X2,k8_real_1(k8_real_1(np__2,k32_sin_cos),X1)))=k1_seq_1(k19_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~m1_subset_1(k8_real_1(k8_real_1(np__2,k32_sin_cos),X1),k1_numbers))).#trainneg
cnf(i_0_5168, plain, (k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1))=k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_5157, plain, (k1_seq_1(k19_sin_cos,k1_seq_1(k16_sin_cos,k1_real_1(X1)))=k1_seq_1(k19_sin_cos,k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2260, plain, (k3_xcmplx_0(k32_sin_cos,np__1)=k32_sin_cos)).#trainneg
cnf(i_0_5662, plain, (k2_xcmplx_0(k32_sin_cos,np__0)=k32_sin_cos)).#trainneg
cnf(i_0_5674, plain, (k2_xcmplx_0(np__0,k32_sin_cos)=k32_sin_cos)).#trainneg
cnf(i_0_5871, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,np__0))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_603, plain, (k1_seq_1(k19_sin_cos,k2_xcmplx_0(k8_real_1(k8_real_1(np__2,k32_sin_cos),X1),X2))=k1_seq_1(k19_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~m1_subset_1(k8_real_1(k8_real_1(np__2,k32_sin_cos),X1),k1_numbers))).#trainneg
cnf(i_0_604, plain, (k1_seq_1(k19_sin_cos,k7_real_1(X2,k8_real_1(k8_real_1(np__2,k32_sin_cos),X1)))=k1_seq_1(k19_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(k8_real_1(k8_real_1(np__2,k32_sin_cos),X1))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_5951, negated_conjecture, (k3_xcmplx_0(k3_xcmplx_0(k32_sin_cos,esk1_0),np__2)=k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)))).#trainneg
cnf(i_0_1808, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1506, negated_conjecture, (m1_subset_1(k3_xcmplx_0(np__2,esk1_0),k1_numbers))).#trainneg
cnf(i_0_4642, plain, (v1_xxreal_0(k4_xcmplx_0(k1_xcmplx_0))|~v2_membered(k2_numbers))).#trainneg
cnf(i_0_621, plain, (k3_xcmplx_0(k4_xcmplx_0(np__2),k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))=np__2|~v1_xcmplx_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_5554, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(esk1_0,k1_xcmplx_0))))!=np__1)).#trainneg
cnf(i_0_777, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_605, plain, (k1_seq_1(k16_sin_cos,k4_xcmplx_0(k7_real_1(k8_real_1(k8_real_1(np__2,k32_sin_cos),X1),X2)))=k1_real_1(k1_seq_1(k16_sin_cos,X2))|~v1_xreal_0(k7_real_1(k8_real_1(k8_real_1(np__2,k32_sin_cos),X1),X2))|~v7_ordinal1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_607, plain, (k1_seq_1(k16_sin_cos,k7_real_1(k8_real_1(X1,k8_real_1(np__2,k32_sin_cos)),X2))=k1_seq_1(k16_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(k8_real_1(np__2,k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_5479, plain, (k2_xcmplx_0(k32_sin_cos,k32_sin_cos)=k3_xcmplx_0(np__2,k32_sin_cos))).#trainneg
cnf(i_0_5806, negated_conjecture, (v1_xcmplx_0(k3_xcmplx_0(X1,esk1_0))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6064, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))=k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))).#trainneg
cnf(i_0_5599, negated_conjecture, (m1_subset_1(k3_xcmplx_0(X1,k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0))),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5734, plain, (k1_real_1(k1_seq_1(k16_sin_cos,X1))=k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5600, negated_conjecture, (m1_subset_1(k3_xcmplx_0(X1,k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0))),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2107, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k2_numbers)))).#trainneg
cnf(i_0_4643, plain, (v1_int_1(k4_xcmplx_0(k1_xcmplx_0))|~v5_membered(k2_numbers))).#trainneg
cnf(i_0_623, plain, (k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),k3_xcmplx_0(k1_xcmplx_0,X1))=k3_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0),X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0)))).#trainneg
cnf(i_0_2417, negated_conjecture, (k2_xcmplx_0(esk1_0,np__1)=k2_xcmplx_0(np__1,esk1_0))).#trainneg
cnf(i_0_5555, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(np__2,k3_xcmplx_0(esk1_0,k3_xcmplx_0(k32_sin_cos,k1_xcmplx_0))))!=np__1)).#trainneg
cnf(i_0_5556, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(k1_xcmplx_0,esk1_0))))!=np__1)).#trainneg
cnf(i_0_781, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2624, plain, (v1_xcmplx_0(k3_xcmplx_0(np__2,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_608, plain, (k1_seq_1(k16_sin_cos,k7_real_1(k3_xcmplx_0(k8_real_1(np__2,k32_sin_cos),X1),X2))=k1_seq_1(k16_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(k8_real_1(np__2,k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_609, plain, (k1_seq_1(k16_sin_cos,k7_real_1(k8_real_1(X1,k8_real_1(np__2,k32_sin_cos)),X2))=k1_seq_1(k16_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(k8_real_1(np__2,k32_sin_cos))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5624, plain, (k1_seq_1(k16_sin_cos,k1_real_1(X1))=k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,k1_real_1(X1)))|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5766, plain, (k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1))=k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)),k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5843, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(np__2,k32_sin_cos)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5784, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_624, plain, (k3_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(np__2),k1_xcmplx_0),k3_xcmplx_0(k1_xcmplx_0,X1))=k3_xcmplx_0(np__2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(np__2),k1_xcmplx_0)))).#trainneg
cnf(i_0_5557, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k3_xcmplx_0(np__2,k1_xcmplx_0),k3_xcmplx_0(k32_sin_cos,esk1_0)))!=np__1)).#trainneg
cnf(i_0_6290, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_6550, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_610, plain, (k1_seq_1(k16_sin_cos,k7_real_1(X2,k8_real_1(k8_real_1(np__2,k32_sin_cos),X1)))=k1_seq_1(k16_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~m1_subset_1(k8_real_1(k8_real_1(np__2,k32_sin_cos),X1),k1_numbers))).#trainneg
cnf(i_0_611, plain, (k1_seq_1(k16_sin_cos,k2_xcmplx_0(k8_real_1(k8_real_1(np__2,k32_sin_cos),X1),X2))=k1_seq_1(k16_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~m1_subset_1(k8_real_1(k8_real_1(np__2,k32_sin_cos),X1),k1_numbers))).#trainneg
cnf(i_0_6416, plain, (k1_real_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6643, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)))=k1_seq_1(k16_sin_cos,X1)|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)),k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6703, plain, (k1_real_1(k1_seq_1(k16_sin_cos,X1))=k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)),k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_6518, plain, (k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),np__1)=k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))).#trainneg
cnf(i_0_5729, plain, (k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1))=k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)))|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6138, plain, (v1_xcmplx_0(k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_625, plain, (k3_xcmplx_0(k3_xcmplx_0(np__2,k1_xcmplx_0),k3_xcmplx_0(np__1,X1))=k3_xcmplx_0(k3_xcmplx_0(np__2,k1_xcmplx_0),X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(k3_xcmplx_0(np__2,k1_xcmplx_0)))).#trainneg
cnf(i_0_6146, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,np__0))|~v1_xcmplx_0(X1)|~v1_membered(X2))).#trainneg
cnf(i_0_6147, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,np__0))|~v1_xcmplx_0(X1)|~v6_membered(X2))).#trainneg
cnf(i_0_612, plain, (k1_seq_1(k16_sin_cos,k7_real_1(X2,k8_real_1(k8_real_1(np__2,k32_sin_cos),X1)))=k1_seq_1(k16_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(k8_real_1(k8_real_1(np__2,k32_sin_cos),X1))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_613, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0)))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6013, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)))=k1_seq_1(k16_sin_cos,X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6503, plain, (k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))=np__1)).#trainneg
cnf(i_0_5530, plain, (k2_xcmplx_0(k1_real_1(X1),k4_xcmplx_0(X2))=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7214, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_626, plain, (k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(np__0,X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_616, plain, (v1_xboole_0(k3_xcmplx_0(X1,X2))|v1_xboole_0(X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xboole_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_617, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)),X4)=k3_xcmplx_0(k3_xcmplx_0(X1,X2),k3_xcmplx_0(X3,X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6813, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(esk1_0,k3_xcmplx_0(k3_xcmplx_0(np__2,k1_xcmplx_0),k32_sin_cos)))!=np__1)).#trainneg
cnf(i_0_627, plain, (k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0),X1))=k3_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0)))).#trainneg
cnf(i_0_5565, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k1_xcmplx_0))))!=np__1)).#trainneg
cnf(i_0_628, plain, (k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(np__1,X1))=k3_xcmplx_0(k1_xcmplx_0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_629, plain, (k3_xcmplx_0(k3_xcmplx_0(np__2,k1_xcmplx_0),X1)=k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(np__2,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_7741, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),k3_xcmplx_0(esk1_0,k1_xcmplx_0)))!=np__1)).#trainneg
cnf(i_0_6972, plain, (k1_real_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)),k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_7048, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)))=k1_seq_1(k16_sin_cos,X1)|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)),k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7883, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(esk1_0,k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(np__2,k32_sin_cos))))!=np__1)).#trainneg
cnf(i_0_7262, plain, (k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1))=k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6923, plain, (k1_real_1(k1_seq_1(k16_sin_cos,X1))=k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)),k1_numbers)|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6814, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(k3_xcmplx_0(np__2,k1_xcmplx_0),esk1_0)))!=np__1)).#trainneg
cnf(i_0_7403, negated_conjecture, (k4_xcmplx_0(k4_xcmplx_0(esk1_0))=esk1_0)).#trainneg
cnf(i_0_7060, plain, (m1_subset_1(k4_xcmplx_0(k32_sin_cos),k2_numbers))).#trainneg
cnf(i_0_4811, negated_conjecture, (v7_ordinal1(k4_xcmplx_0(esk1_0))|~v6_membered(k2_numbers))).#trainneg
cnf(i_0_630, plain, (k3_xcmplx_0(np__0,k3_xcmplx_0(k1_xcmplx_0,X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_2420, negated_conjecture, (k2_xcmplx_0(esk1_0,k32_sin_cos)=k2_xcmplx_0(k32_sin_cos,esk1_0))).#trainneg
cnf(i_0_5566, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(esk1_0,k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,k1_xcmplx_0))))!=np__1)).#trainneg
cnf(i_0_631, plain, (k3_xcmplx_0(np__0,k3_xcmplx_0(np__0,X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_632, plain, (k3_xcmplx_0(np__0,k3_xcmplx_0(np__1,X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_6961, plain, (k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1))=k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)))|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7367, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6242, negated_conjecture, (v1_xreal_0(k3_xcmplx_0(np__2,esk1_0)))).#trainneg
cnf(i_0_8264, negated_conjecture, (k1_real_1(esk1_0)=k4_xcmplx_0(esk1_0))).#trainneg
cnf(i_0_633, plain, (k3_xcmplx_0(np__0,k3_xcmplx_0(np__2,X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_8296, plain, (v1_xreal_0(k4_xcmplx_0(k32_sin_cos))|~v3_membered(k2_numbers))).#trainneg
cnf(i_0_634, plain, (k3_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0),k3_xcmplx_0(k1_xcmplx_0,X1))=k3_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0)))).#trainneg
cnf(i_0_635, plain, (k3_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0),k3_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0),X1))=k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0)))).#trainneg
cnf(i_0_8553, plain, (k3_xcmplx_0(X1,k1_real_1(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k1_real_1(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6529, negated_conjecture, (v1_xcmplx_0(k2_xcmplx_0(np__1,esk1_0)))).#trainneg
cnf(i_0_4813, negated_conjecture, (v1_xxreal_0(k4_xcmplx_0(esk1_0))|~v2_membered(k2_numbers))).#trainneg
cnf(i_0_636, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(k1_xcmplx_0,X1))=k3_xcmplx_0(k1_xcmplx_0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_2421, negated_conjecture, (k2_xcmplx_0(esk1_0,np__2)=k2_xcmplx_0(np__2,esk1_0))).#trainneg
cnf(i_0_7108, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,np__0))|~v1_xcmplx_0(X1)|~v3_membered(X2))).#trainneg
cnf(i_0_7111, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,np__0))|~v1_xcmplx_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_637, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(np__0,X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_638, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(np__1,X1))=k3_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_7944, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(np__2,esk1_0))))!=np__1)).#trainneg
cnf(i_0_8558, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_8607, negated_conjecture, (k1_real_1(k4_xcmplx_0(esk1_0))=esk1_0)).#trainneg
cnf(i_0_8362, negated_conjecture, (v1_xcmplx_0(k2_xcmplx_0(k32_sin_cos,esk1_0)))).#trainneg
cnf(i_0_4814, negated_conjecture, (v1_int_1(k4_xcmplx_0(esk1_0))|~v5_membered(k2_numbers))).#trainneg
cnf(i_0_639, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(np__2,X1))=k3_xcmplx_0(np__2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_6534, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(np__2,k3_xcmplx_0(esk1_0,k3_xcmplx_0(k1_xcmplx_0,k32_sin_cos))))!=np__1)).#trainneg
cnf(i_0_640, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_641, plain, (k3_xcmplx_0(np__2,k3_xcmplx_0(np__0,X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_9078, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(k3_xcmplx_0(np__2,esk1_0),k32_sin_cos)))!=np__1)).#trainneg
cnf(i_0_8839, plain, (k3_xcmplx_0(k1_real_1(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8599, negated_conjecture, (k3_xcmplx_0(k3_xcmplx_0(np__2,esk1_0),k32_sin_cos)=k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)))).#trainneg
cnf(i_0_8070, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(esk1_0,k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),k1_xcmplx_0)))!=np__1)).#trainneg
cnf(i_0_3221, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8943, negated_conjecture, (v1_xcmplx_0(k2_xcmplx_0(np__2,esk1_0)))).#trainneg
cnf(i_0_4942, plain, (v7_ordinal1(k4_xcmplx_0(np__2))|~v6_membered(k2_numbers))).#trainneg
cnf(i_0_642, plain, (k3_xcmplx_0(np__2,k3_xcmplx_0(np__1,X1))=k3_xcmplx_0(np__2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_2687, plain, (k2_xcmplx_0(k32_sin_cos,np__1)=k2_xcmplx_0(np__1,k32_sin_cos))).#trainneg
cnf(i_0_643, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_648, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),X4)=k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9779, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9348, negated_conjecture, (k3_xcmplx_0(np__1,k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)))=k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)))).#trainneg
cnf(i_0_9699, plain, (v1_xcmplx_0(k2_xcmplx_0(np__1,k32_sin_cos)))).#trainneg
cnf(i_0_649, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2688, plain, (k2_xcmplx_0(np__2,np__1)=k2_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_650, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9557, plain, (k1_real_1(X1)=k4_xcmplx_0(k3_xcmplx_0(np__1,X1))|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3444, plain, (v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4944, plain, (v1_xxreal_0(k4_xcmplx_0(np__2))|~v2_membered(k2_numbers))).#trainneg
cnf(i_0_653, plain, (k2_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),k2_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),X1))=k2_xcmplx_0(k4_xcmplx_0(np__2),X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0)))).#trainneg
cnf(i_0_652, plain, (k2_xcmplx_0(k21_sin_cos(X1),k2_xcmplx_0(k3_xcmplx_0(k18_sin_cos(X1),k1_xcmplx_0),X2))=k2_xcmplx_0(k15_sin_cos(k3_xcmplx_0(X1,k1_xcmplx_0)),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k3_xcmplx_0(k18_sin_cos(X1),k1_xcmplx_0))|~v1_xcmplx_0(k21_sin_cos(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_654, plain, (k2_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),k2_xcmplx_0(np__0,X1))=k2_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0)))).#trainneg
cnf(i_0_9942, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_9902, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10346, plain, (k1_real_1(X1)=k4_xcmplx_0(k3_xcmplx_0(np__1,X1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers))).#trainneg
cnf(i_0_10016, plain, (k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),np__2)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_9965, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=k3_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10035, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))=k4_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10248, plain, (v1_xcmplx_0(k2_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_4945, plain, (v1_int_1(k4_xcmplx_0(np__2))|~v5_membered(k2_numbers))).#trainneg
cnf(i_0_655, plain, (k2_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),k2_xcmplx_0(np__1,X1))=k2_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0)))).#trainneg
cnf(i_0_656, plain, (k2_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),k2_xcmplx_0(np__2,X1))=k2_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0)))).#trainneg
cnf(i_0_657, plain, (k2_xcmplx_0(k1_xcmplx_0,k2_xcmplx_0(k1_xcmplx_0,X1))=k2_xcmplx_0(k3_xcmplx_0(np__2,k1_xcmplx_0),X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10053, plain, (k4_xcmplx_0(k3_xcmplx_0(X2,X1))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10592, plain, (k3_xcmplx_0(k1_real_1(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11007, plain, (k3_xcmplx_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_10636, plain, (k3_xcmplx_0(np__2,k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_10701, plain, (m1_subset_1(k4_xcmplx_0(np__2),k1_numbers))).#trainneg
cnf(i_0_658, plain, (k2_xcmplx_0(k2_xcmplx_0(np__1,k1_xcmplx_0),X1)=k2_xcmplx_0(k1_xcmplx_0,k2_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_659, plain, (k2_xcmplx_0(np__0,k2_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),X1))=k2_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_660, plain, (k2_xcmplx_0(np__0,k2_xcmplx_0(np__0,X1))=k2_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_10597, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11010, plain, (k3_xcmplx_0(k4_xcmplx_0(X2),X1)=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_11346, plain, (v1_xreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_1377, plain, (k1_real_1(np__2)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_661, plain, (k2_xcmplx_0(np__0,k2_xcmplx_0(np__1,X1))=k2_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_3013, plain, (k2_xcmplx_0(k32_sin_cos,np__2)=k2_xcmplx_0(np__2,k32_sin_cos))).#trainneg
cnf(i_0_6289, plain, (k2_zfmisc_1(X1,X2)=k4_xcmplx_0(X3)|~v1_xcmplx_0(X3)|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_6549, plain, (k2_zfmisc_1(X1,X2)=k4_xcmplx_0(X3)|~v1_xcmplx_0(X3)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_662, plain, (k2_xcmplx_0(np__0,k2_xcmplx_0(np__2,X1))=k2_xcmplx_0(np__2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_663, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k2_xcmplx_0(k4_xcmplx_0(X2),X3))=k2_xcmplx_0(k4_xcmplx_0(k2_xcmplx_0(X1,X2)),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_11171, plain, (k3_xcmplx_0(X1,k1_real_1(X2))=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k1_real_1(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_10517, plain, (k3_xcmplx_0(k1_real_1(X1),X2)=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10490, plain, (k3_xcmplx_0(X1,k1_real_1(X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k1_real_1(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_10521, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_11858, plain, (k1_real_1(k4_xcmplx_0(np__2))=np__2)).#trainneg
cnf(i_0_11905, plain, (v1_xcmplx_0(k2_xcmplx_0(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_8261, negated_conjecture, (v1_xboole_0(k4_xcmplx_0(esk1_0))|~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_347, plain, (k1_seq_1(k19_sin_cos,k4_xcmplx_0(np__2))=k1_seq_1(k19_sin_cos,np__2))).#trainneg
cnf(i_0_664, plain, (k2_xcmplx_0(k4_xcmplx_0(np__2),k2_xcmplx_0(np__2,X1))=k2_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_665, plain, (k2_xcmplx_0(np__1,k2_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),X1))=k2_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_666, plain, (k2_xcmplx_0(np__1,k2_xcmplx_0(np__0,X1))=k2_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_12371, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k3_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_11248, plain, (k4_xcmplx_0(k3_xcmplx_0(X2,X1))=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_12429, plain, (k3_xcmplx_0(k1_real_1(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_667, plain, (k2_xcmplx_0(np__1,k2_xcmplx_0(np__1,X1))=k2_xcmplx_0(np__2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_7882, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k3_xcmplx_0(k1_xcmplx_0,k32_sin_cos))))!=np__1)).#trainneg
cnf(i_0_7884, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(esk1_0,k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,k1_xcmplx_0))))!=np__1)).#trainneg
cnf(i_0_12478, negated_conjecture, (X1=k4_xcmplx_0(esk1_0)|~v1_xboole_0(X1)|~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_668, plain, (k2_xcmplx_0(np__2,k2_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),X1))=k2_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_669, plain, (k2_xcmplx_0(np__2,k2_xcmplx_0(k4_xcmplx_0(np__2),X1))=k2_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__2))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_12176, plain, (k1_real_1(X1)=k3_xcmplx_0(X1,k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11788, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k1_real_1(X2))=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k1_real_1(X2))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_12217, plain, (k3_xcmplx_0(k1_real_1(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10049, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),np__1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_12081, plain, (k1_real_1(X1)=k4_xcmplx_0(k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4532, plain, (v1_xcmplx_0(k19_sin_cos)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_12471, negated_conjecture, (k4_xcmplx_0(esk1_0)=np__0|~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_671, plain, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7945, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,k3_xcmplx_0(esk1_0,k1_xcmplx_0))))!=np__1)).#trainneg
cnf(i_0_12473, negated_conjecture, (k4_xcmplx_0(esk1_0)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)|~v1_xboole_0(X1)|~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_673, plain, (v1_partfun1(esk2_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1)|v1_xboole_0(X2)|~v1_funct_2(esk2_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1,X2))).#trainneg
cnf(i_0_9907, plain, (k1_real_1(k3_xcmplx_0(k4_xcmplx_0(X1),X2))=k3_xcmplx_0(X1,X2)|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13389, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13453, plain, (k3_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13447, plain, (k3_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4668, plain, (v1_xxreal_0(k16_sin_cos)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1694, plain, (m1_subset_1(k3_xcmplx_0(k32_sin_cos,k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_8297, plain, (v7_ordinal1(k4_xcmplx_0(k32_sin_cos))|~v6_membered(k2_numbers))).#trainneg
cnf(i_0_678, plain, (v1_xcmplx_0(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2))|~v1_xcmplx_0(k3_xcmplx_0(X3,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7946, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,k3_xcmplx_0(k1_xcmplx_0,esk1_0))))!=np__1)).#trainneg
cnf(i_0_676, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,X2),k2_xcmplx_0(k3_xcmplx_0(X3,X2),X4))=k2_xcmplx_0(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X4)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k3_xcmplx_0(X3,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_677, plain, (k2_xcmplx_0(k3_xcmplx_0(X3,X2),k3_xcmplx_0(X1,X2))=k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(k3_xcmplx_0(X3,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13843, plain, (k3_xcmplx_0(X1,np__1)=X1|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_13763, plain, (k3_xcmplx_0(X1,k1_real_1(X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4846, plain, (v1_xxreal_0(k19_sin_cos)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_13876, plain, (v1_xreal_0(k3_xcmplx_0(k32_sin_cos,k32_sin_cos)))).#trainneg
cnf(i_0_8299, plain, (v1_xxreal_0(k4_xcmplx_0(k32_sin_cos))|~v2_membered(k2_numbers))).#trainneg
cnf(i_0_680, plain, (k3_xcmplx_0(k3_xcmplx_0(np__2,k1_xcmplx_0),k1_xcmplx_0)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_13878, plain, (v1_xcmplx_0(k3_xcmplx_0(k32_sin_cos,k32_sin_cos))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_8062, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),k3_xcmplx_0(k1_xcmplx_0,esk1_0)))!=np__1)).#trainneg
cnf(i_0_4383, negated_conjecture, (v1_xcmplx_0(k3_xcmplx_0(esk1_0,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_681, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,k1_xcmplx_0),k4_xcmplx_0(k1_xcmplx_0))=k3_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0)),k1_xcmplx_0)|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13291, plain, (k3_xcmplx_0(X1,k1_real_1(X2))=k3_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(k1_real_1(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13269, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_14852, plain, (v1_xcmplx_0(k3_xcmplx_0(k32_sin_cos,k32_sin_cos)))).#trainneg
cnf(i_0_8300, plain, (v1_int_1(k4_xcmplx_0(k32_sin_cos))|~v5_membered(k2_numbers))).#trainneg
cnf(i_0_682, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(k4_xcmplx_0(np__2),k1_xcmplx_0)),k1_xcmplx_0)=k2_xcmplx_0(k3_xcmplx_0(X1,k1_xcmplx_0),np__2)|~v1_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(np__2),k1_xcmplx_0))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_683, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,np__1),k3_xcmplx_0(np__2,k1_xcmplx_0))=k3_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(np__2,k1_xcmplx_0)),np__1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(k3_xcmplx_0(np__2,k1_xcmplx_0))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_684, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,X2),k3_xcmplx_0(X3,k3_xcmplx_0(X4,X2)))=k3_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(X3,X4)),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k3_xcmplx_0(X3,X4))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_14717, plain, (k1_real_1(X1)=k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),X1)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13195, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))=X1|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10780, plain, (k3_xcmplx_0(np__1,k1_real_1(X1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9776, plain, (k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),k3_xcmplx_0(X1,X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_14644, plain, (k3_xcmplx_0(np__1,X1)=X1|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5091, plain, (v1_int_1(k16_sin_cos)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_14877, negated_conjecture, (v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,esk1_0)))).#trainneg
cnf(i_0_685, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,np__0),np__0)=k3_xcmplx_0(k2_xcmplx_0(X1,k1_xcmplx_0),np__0)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4800, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_686, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(k1_xcmplx_0)),np__1)=k3_xcmplx_0(k2_xcmplx_0(X1,k1_xcmplx_0),k4_xcmplx_0(k1_xcmplx_0))|~v1_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_687, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,np__1),k1_xcmplx_0)=k3_xcmplx_0(k2_xcmplx_0(X1,k1_xcmplx_0),np__1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_15683, plain, (k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),k4_xcmplx_0(X1))=X1|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_688, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,np__2),k3_xcmplx_0(np__2,k1_xcmplx_0))=k3_xcmplx_0(k2_xcmplx_0(X1,k1_xcmplx_0),np__2)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2486, negated_conjecture, (v1_xboole_0(esk1_0)|v1_xboole_0(k32_sin_cos)|~v1_xboole_0(k3_xcmplx_0(k32_sin_cos,esk1_0)))).#trainneg
cnf(i_0_5829, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_689, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,k1_xcmplx_0),np__0)=k3_xcmplx_0(k2_xcmplx_0(X1,np__0),k1_xcmplx_0)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_690, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,np__0),np__0)=k3_xcmplx_0(k2_xcmplx_0(X1,np__0),np__0)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_15676, plain, (k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15805, plain, (k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15661, plain, (k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),k1_real_1(X1))=X1|~v1_xcmplx_0(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1754, plain, (m1_subset_1(k3_xcmplx_0(np__2,np__2),k1_numbers))).#trainneg
cnf(i_0_691, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,np__1),np__0)=k3_xcmplx_0(k2_xcmplx_0(X1,np__0),np__1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_692, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,np__2),np__0)=k3_xcmplx_0(k2_xcmplx_0(X1,np__0),np__2)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_693, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(k1_xcmplx_0)),k1_xcmplx_0)=k2_xcmplx_0(k3_xcmplx_0(X1,k1_xcmplx_0),np__1)|~v1_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_16925, plain, (k2_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),np__0)=k3_xcmplx_0(np__2,k32_sin_cos))).#trainneg
cnf(i_0_16810, plain, (v1_xreal_0(k3_xcmplx_0(np__2,np__2)))).#trainneg
cnf(i_0_694, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(k1_xcmplx_0)),k4_xcmplx_0(k1_xcmplx_0))=k2_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(k1_xcmplx_0)),k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))|~v1_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5847, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,k32_sin_cos))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_695, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,k1_xcmplx_0),k1_xcmplx_0)=k3_xcmplx_0(k2_xcmplx_0(X1,np__1),k1_xcmplx_0)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_696, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,np__0),np__0)=k3_xcmplx_0(k2_xcmplx_0(X1,np__1),np__0)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_16754, plain, (k3_xcmplx_0(k1_real_1(X1),np__1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11347, plain, (v7_ordinal1(k4_xcmplx_0(np__2))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_697, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,np__1),np__1)=k3_xcmplx_0(k2_xcmplx_0(X1,np__1),np__1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_698, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,np__2),np__2)=k3_xcmplx_0(k2_xcmplx_0(X1,np__1),np__2)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_700, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,np__0),np__0)=k3_xcmplx_0(k2_xcmplx_0(X1,np__2),np__0)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_11008, plain, (k3_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_16992, plain, (k2_xcmplx_0(np__0,k3_xcmplx_0(np__2,k32_sin_cos))=k3_xcmplx_0(np__2,k32_sin_cos))).#trainneg
cnf(i_0_11349, plain, (v1_xxreal_0(k4_xcmplx_0(np__2))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_701, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,np__1),np__2)=k3_xcmplx_0(k2_xcmplx_0(X1,np__2),np__1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_702, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0)),k4_xcmplx_0(X2))=k3_xcmplx_0(k2_xcmplx_0(X1,X2),k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_704, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,X2),k3_xcmplx_0(X2,X3))=k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_17285, plain, (k3_xcmplx_0(X1,k1_real_1(X2))=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(k1_real_1(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_18127, plain, (k2_xcmplx_0(k3_xcmplx_0(X2,X1),k3_xcmplx_0(X2,X3))=k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5261, plain, (v1_int_1(k19_sin_cos)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3564, negated_conjecture, (m1_subset_1(k3_xcmplx_0(k32_sin_cos,esk1_0),k2_numbers))).#trainneg
cnf(i_0_11350, plain, (v1_int_1(k4_xcmplx_0(np__2))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_3565, negated_conjecture, (m1_subset_1(k3_xcmplx_0(np__2,esk1_0),k2_numbers))).#trainneg
cnf(i_0_6091, plain, (v1_xcmplx_0(k3_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_706, plain, (k2_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0),k3_xcmplx_0(X1,k1_xcmplx_0))=k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),X1),k1_xcmplx_0)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0)))).#trainneg
cnf(i_0_707, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(np__2),k1_xcmplx_0),X1),k1_xcmplx_0)=k2_xcmplx_0(np__2,k3_xcmplx_0(X1,k1_xcmplx_0))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(np__2),k1_xcmplx_0)))).#trainneg
cnf(i_0_11646, plain, (k3_xcmplx_0(np__0,X1)=np__0|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_18272, plain, (k3_xcmplx_0(k2_xcmplx_0(X2,X2),X1)=k3_xcmplx_0(k2_xcmplx_0(X1,X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1739, plain, (v3_membered(k16_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3570, plain, (m1_subset_1(k3_xcmplx_0(np__2,k1_xcmplx_0),k2_numbers))).#trainneg
cnf(i_0_710, plain, (k2_xcmplx_0(np__0,k3_xcmplx_0(X1,np__0))=k3_xcmplx_0(k2_xcmplx_0(k1_xcmplx_0,X1),np__0)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_708, plain, (k2_xcmplx_0(k3_xcmplx_0(np__2,k1_xcmplx_0),k3_xcmplx_0(X1,np__1))=k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(np__2,k1_xcmplx_0),X1),np__1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(np__2,k1_xcmplx_0)))).#trainneg
cnf(i_0_709, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)),k3_xcmplx_0(X4,X3))=k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),X4),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_17632, plain, (k3_xcmplx_0(k3_xcmplx_0(np__1,X1),X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1740, plain, (v3_membered(k19_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_711, plain, (k2_xcmplx_0(np__1,k3_xcmplx_0(X1,k4_xcmplx_0(k1_xcmplx_0)))=k3_xcmplx_0(k2_xcmplx_0(k1_xcmplx_0,X1),k4_xcmplx_0(k1_xcmplx_0))|~v1_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_19155, plain, (v1_xreal_0(k3_xcmplx_0(np__2,k1_xcmplx_0))|~v3_membered(k2_numbers))).#trainneg
cnf(i_0_712, plain, (k2_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(X1,np__1))=k3_xcmplx_0(k2_xcmplx_0(k1_xcmplx_0,X1),np__1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_713, plain, (k2_xcmplx_0(k3_xcmplx_0(np__2,k1_xcmplx_0),k3_xcmplx_0(X1,np__2))=k3_xcmplx_0(k2_xcmplx_0(k1_xcmplx_0,X1),np__2)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_15756, plain, (k3_xcmplx_0(np__0,k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))=np__0)).#trainneg
cnf(i_0_16787, plain, (k1_real_1(X1)=k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1823, plain, (v1_membered(k16_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k2_numbers)))).#trainneg
cnf(i_0_714, plain, (k2_xcmplx_0(np__0,k3_xcmplx_0(X1,k1_xcmplx_0))=k3_xcmplx_0(k2_xcmplx_0(np__0,X1),k1_xcmplx_0)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_715, plain, (k2_xcmplx_0(np__0,k3_xcmplx_0(X1,np__0))=k3_xcmplx_0(k2_xcmplx_0(np__0,X1),np__0)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_716, plain, (k2_xcmplx_0(np__0,k3_xcmplx_0(X1,np__1))=k3_xcmplx_0(k2_xcmplx_0(np__0,X1),np__1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_20100, plain, (k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),X1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_19834, negated_conjecture, (k3_xcmplx_0(k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)),X1)=k3_xcmplx_0(X1,k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_20038, plain, (k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),np__0)=np__0)).#trainneg
cnf(i_0_19825, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,np__1),X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1825, plain, (v1_membered(k19_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k2_numbers)))).#trainneg
cnf(i_0_3736, plain, (m1_subset_1(k2_xcmplx_0(np__1,k1_xcmplx_0),k2_numbers))).#trainneg
cnf(i_0_717, plain, (k2_xcmplx_0(np__0,k3_xcmplx_0(X1,np__2))=k3_xcmplx_0(k2_xcmplx_0(np__0,X1),np__2)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_718, plain, (k3_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0),X1),k1_xcmplx_0)=k2_xcmplx_0(np__1,k3_xcmplx_0(X1,k1_xcmplx_0))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0)))).#trainneg
cnf(i_0_719, plain, (k3_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0),X1),k4_xcmplx_0(k1_xcmplx_0))=k2_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),k3_xcmplx_0(X1,k4_xcmplx_0(k1_xcmplx_0)))|~v1_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_20281, plain, (k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),np__1)=k3_xcmplx_0(np__2,k32_sin_cos))).#trainneg
cnf(i_0_2385, plain, (v1_membered(k16_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5058, plain, (v1_xreal_0(k4_xcmplx_0(k1_xcmplx_0))|~m1_subset_1(k2_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_720, plain, (k2_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(X1,k1_xcmplx_0))=k3_xcmplx_0(k2_xcmplx_0(np__1,X1),k1_xcmplx_0)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_721, plain, (k2_xcmplx_0(np__0,k3_xcmplx_0(X1,np__0))=k3_xcmplx_0(k2_xcmplx_0(np__1,X1),np__0)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_722, plain, (k2_xcmplx_0(np__1,k3_xcmplx_0(X1,np__1))=k3_xcmplx_0(k2_xcmplx_0(np__1,X1),np__1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_19836, plain, (k3_xcmplx_0(k3_xcmplx_0(np__2,k32_sin_cos),X1)=k3_xcmplx_0(X1,k3_xcmplx_0(np__2,k32_sin_cos))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_20088, plain, (k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),k4_xcmplx_0(X1))=X1|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_20121, plain, (k1_seq_1(k19_sin_cos,k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),X1))=k1_seq_1(k19_sin_cos,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_19841, plain, (k3_xcmplx_0(X2,k1_real_1(X1))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k1_real_1(X1))|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2387, plain, (v1_membered(k19_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8678, plain, (v1_xreal_0(k4_xcmplx_0(k32_sin_cos))|~m1_subset_1(k2_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_723, plain, (k2_xcmplx_0(np__2,k3_xcmplx_0(X1,np__2))=k3_xcmplx_0(k2_xcmplx_0(np__1,X1),np__2)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_725, plain, (k2_xcmplx_0(np__0,k3_xcmplx_0(X1,np__0))=k3_xcmplx_0(k2_xcmplx_0(np__2,X1),np__0)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_726, plain, (k2_xcmplx_0(np__2,k3_xcmplx_0(X1,np__1))=k3_xcmplx_0(k2_xcmplx_0(np__2,X1),np__1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_3324, plain, (v2_membered(k16_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1487, plain, (v1_xboole_0(k5_numbers)|m1_subset_1(esk2_1(k5_numbers),k1_numbers))).#trainneg
cnf(i_0_729, plain, (k2_xcmplx_0(k3_xcmplx_0(X2,X1),k3_xcmplx_0(X3,X2))=k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_22014, plain, (v1_xreal_0(esk2_1(k5_numbers))|v1_xboole_0(k5_numbers))).#trainneg
cnf(i_0_727, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k3_xcmplx_0(X2,k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0)))=k3_xcmplx_0(k2_xcmplx_0(X1,X2),k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_21661, plain, (k1_seq_1(k19_sin_cos,k4_xcmplx_0(X1))=k1_seq_1(k19_sin_cos,X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_22015, plain, (v7_ordinal1(esk2_1(k5_numbers))|v1_xboole_0(k5_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_759, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk2_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,esk2_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_4061, plain, (m1_subset_1(k3_xcmplx_0(np__2,np__2),k2_numbers))).#trainneg
cnf(i_0_731, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_funct_2(esk2_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1,X2)|~v1_funct_1(esk2_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(esk2_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_734, plain, (k3_funct_2(X1,X2,esk2_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X3)=k1_funct_1(esk2_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X3)|v1_xboole_0(X1)|~v1_funct_2(esk2_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1,X2)|~v1_funct_1(esk2_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_21689, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))=X1|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(k1_real_1(X1)))).#trainneg
cnf(i_0_3326, plain, (v2_membered(k19_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_22016, plain, (v1_xcmplx_0(esk2_1(k5_numbers))|v1_xboole_0(k5_numbers)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_760, plain, (k4_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),X1))=k2_xcmplx_0(np__1,k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0)))).#trainneg
cnf(i_0_7660, plain, (v1_xboole_0(k1_xcmplx_0)|~v1_xboole_0(k3_xcmplx_0(np__2,k1_xcmplx_0)))).#trainneg
cnf(i_0_10259, negated_conjecture, (v1_xcmplx_0(k2_xcmplx_0(X1,esk1_0))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10274, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_739, plain, (v1_xboole_0(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|v1_xboole_0(X1)|~v1_xboole_0(X3)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1))).#trainneg
cnf(i_0_740, plain, (v5_relat_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5),X3)|v1_xboole_0(X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1))).#trainneg
cnf(i_0_8030, plain, (k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(np__2,k1_xcmplx_0))=k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(np__2,k32_sin_cos))|~m1_subset_1(k3_xcmplx_0(np__2,k1_xcmplx_0),k1_numbers))).#trainneg
cnf(i_0_10003, negated_conjecture, (k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),esk1_0)=k4_xcmplx_0(esk1_0))).#trainneg
cnf(i_0_19843, plain, (k3_xcmplx_0(k32_sin_cos,X1)=k3_xcmplx_0(X1,k32_sin_cos)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_22723, plain, (v1_xcmplx_0(k2_xcmplx_0(k1_xcmplx_0,np__2)))).#trainneg
cnf(i_0_22017, plain, (v1_xxreal_0(esk2_1(k5_numbers))|v1_xboole_0(k5_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_741, plain, (v1_partfun1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)),X3,X4),X2)|v1_xboole_0(X1)|~v1_funct_2(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)),X3,X4),X2,X2)|~v1_funct_2(X3,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))))|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_742, plain, (v1_relat_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|v1_xboole_0(X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1))).#trainneg
cnf(i_0_22702, plain, (k2_xcmplx_0(np__1,k4_xcmplx_0(np__2))=k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))).#trainneg
cnf(i_0_22591, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))=X1|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers))).#trainneg
cnf(i_0_22822, negated_conjecture, (k3_xcmplx_0(k4_xcmplx_0(esk1_0),k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))=esk1_0)).#trainneg
cnf(i_0_17834, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,X2),X2)=k3_xcmplx_0(k2_xcmplx_0(X2,X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_23033, plain, (v1_xcmplx_0(k2_xcmplx_0(np__2,k1_xcmplx_0)))).#trainneg
cnf(i_0_743, plain, (v1_xboole_0(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1))).#trainneg
cnf(i_0_744, plain, (v4_relat_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5),X2)|v1_xboole_0(X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1))).#trainneg
cnf(i_0_12277, plain, (k3_xcmplx_0(k1_real_1(X1),X2)=k3_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_22018, plain, (v1_int_1(esk2_1(k5_numbers))|v1_xboole_0(k5_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_765, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0)))=k2_xcmplx_0(k4_xcmplx_0(X1),np__1)|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5792, negated_conjecture, (m1_subset_1(k3_xcmplx_0(esk1_0,esk1_0),k2_numbers))).#trainneg
cnf(i_0_10285, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,k32_sin_cos))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10286, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_745, plain, (v1_funct_2(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5),X2,X3)|v1_xboole_0(X1)|~v1_partfun1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5),X2)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1))).#trainneg
cnf(i_0_746, plain, (v1_partfun1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5),X2)|v1_xboole_0(X3)|v1_xboole_0(X1)|~v1_funct_2(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5),X2,X3)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1))).#trainneg
cnf(i_0_10630, plain, (k4_xcmplx_0(k3_xcmplx_0(X2,k4_xcmplx_0(X1)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_20824, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_22551, plain, (v1_xcmplx_0(esk2_1(k5_numbers))|v1_xboole_0(k5_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_775, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v1_funct_1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_6530, negated_conjecture, (m1_subset_1(k2_xcmplx_0(np__1,esk1_0),k1_numbers))).#trainneg
cnf(i_0_748, plain, (v3_membered(k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|v1_xboole_0(X1)|~v3_membered(X2)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_749, plain, (v1_membered(k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|v1_xboole_0(X1)|~v1_membered(X2)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_24310, negated_conjecture, (k3_xcmplx_0(k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)),np__1)=k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)))).#trainneg
cnf(i_0_6531, negated_conjecture, (m1_subset_1(k2_xcmplx_0(np__1,esk1_0),k2_numbers))).#trainneg
cnf(i_0_2120, plain, (k2_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0),k1_xcmplx_0)=np__0)).#trainneg
cnf(i_0_779, plain, (v1_partfun1(k2_zfmisc_1(X1,X1),X1)|~v1_funct_2(k2_zfmisc_1(X1,X1),X1,X1))).#trainneg
cnf(i_0_750, plain, (v2_membered(k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|v1_xboole_0(X1)|~v2_membered(X2)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_751, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X3)|~r2_hidden(X1,k3_funct_2(X3,k1_zfmisc_1(X2),X4,X5))|~v1_funct_2(X4,X3,k1_zfmisc_1(X2))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,k1_zfmisc_1(X2))))|~m1_subset_1(X5,X3))).#trainneg
cnf(i_0_15163, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k3_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10541, plain, (k1_real_1(k3_xcmplx_0(X1,X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~m1_subset_1(k3_xcmplx_0(k4_xcmplx_0(X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_24374, negated_conjecture, (v1_xreal_0(k2_xcmplx_0(np__1,esk1_0)))).#trainneg
cnf(i_0_783, plain, (v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v1_partfun1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_8363, negated_conjecture, (m1_subset_1(k2_xcmplx_0(k32_sin_cos,esk1_0),k1_numbers))).#trainneg
cnf(i_0_752, plain, (v1_xreal_0(k3_funct_2(X1,X2,X3,X4))|v1_xboole_0(X1)|~v3_membered(X2)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_753, plain, (v7_ordinal1(k3_funct_2(X1,X2,X3,X4))|v1_xboole_0(X1)|~v6_membered(X2)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_20115, plain, (k4_xcmplx_0(k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0),X1))=X1|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10382, plain, (k2_xcmplx_0(k4_xcmplx_0(np__2),np__1)=k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))).#trainneg
cnf(i_0_13795, plain, (k3_xcmplx_0(k1_real_1(X1),X2)=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_17548, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_21774, plain, (k4_xcmplx_0(k3_xcmplx_0(X2,X1))=k3_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xcmplx_0(k1_real_1(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1430, plain, (v1_xreal_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_25032, negated_conjecture, (v1_xreal_0(k2_xcmplx_0(k32_sin_cos,esk1_0)))).#trainneg
cnf(i_0_7439, negated_conjecture, (k3_xcmplx_0(np__0,k4_xcmplx_0(esk1_0))=np__0)).#trainneg
cnf(i_0_784, plain, (v1_partfun1(k2_zfmisc_1(X1,X2),X1)|v1_xboole_0(X2)|~v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2))).#trainneg
cnf(i_0_8364, negated_conjecture, (m1_subset_1(k2_xcmplx_0(k32_sin_cos,esk1_0),k2_numbers))).#trainneg
cnf(i_0_10914, plain, (v1_xcmplx_0(k2_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_754, plain, (v1_xcmplx_0(k3_funct_2(X1,X2,X3,X4))|v1_xboole_0(X1)|~v1_membered(X2)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_755, plain, (v1_xxreal_0(k3_funct_2(X1,X2,X3,X4))|v1_xboole_0(X1)|~v2_membered(X2)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_25396, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_25371, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))=k4_xcmplx_0(X1)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_953, plain, (v1_xreal_0(k16_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7440, plain, (k3_xcmplx_0(np__0,k4_xcmplx_0(k1_xcmplx_0))=np__0)).#trainneg
cnf(i_0_798, plain, (m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_11940, plain, (v1_xcmplx_0(k2_xcmplx_0(np__2,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_756, plain, (v1_int_1(k3_funct_2(X1,X2,X3,X4))|v1_xboole_0(X1)|~v5_membered(X2)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_776, plain, (k3_funct_2(X1,X2,k2_zfmisc_1(X1,X2),X3)=k1_funct_1(k2_zfmisc_1(X1,X2),X3)|v1_xboole_0(X1)|~v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v1_funct_1(k2_zfmisc_1(X1,X2))|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_6036, plain, (k1_seq_1(k19_sin_cos,k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)))=k1_seq_1(k19_sin_cos,k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10578, plain, (k3_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8944, negated_conjecture, (m1_subset_1(k2_xcmplx_0(np__2,esk1_0),k1_numbers))).#trainneg
cnf(i_0_7443, plain, (k3_xcmplx_0(np__0,k4_xcmplx_0(k32_sin_cos))=np__0)).#trainneg
cnf(i_0_800, plain, (m1_subset_1(k1_seq_1(k19_sin_cos,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8945, negated_conjecture, (m1_subset_1(k2_xcmplx_0(np__2,esk1_0),k2_numbers))).#trainneg
cnf(i_0_807, plain, (v1_xreal_0(k1_seq_1(k19_sin_cos,X1))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_25996, plain, (k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1))=k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_25988, plain, (v1_xreal_0(k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_25993, plain, (k1_real_1(k1_seq_1(k16_sin_cos,X1))=k1_seq_1(k16_sin_cos,k1_real_1(X1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers))).#trainneg
cnf(i_0_25994, plain, (k1_seq_1(k16_sin_cos,k1_real_1(X1))=k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers))).#trainneg
cnf(i_0_25995, plain, (k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1))=k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_960, plain, (v1_xreal_0(k19_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7444, plain, (k3_xcmplx_0(np__0,k4_xcmplx_0(np__2))=np__0)).#trainneg
cnf(i_0_3046, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k11_binop_2(np__2,k32_sin_cos),k3_xcmplx_0(esk1_0,k1_xcmplx_0)))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(k11_binop_2(np__2,k32_sin_cos),esk1_0))|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos))|~v1_xcmplx_0(k11_binop_2(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_808, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_26281, negated_conjecture, (v1_xreal_0(k2_xcmplx_0(np__2,esk1_0)))).#trainneg
cnf(i_0_26504, plain, (k1_real_1(k1_seq_1(k16_sin_cos,k1_real_1(X1)))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_26452, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,k1_real_1(X1)))=k1_seq_1(k16_sin_cos,X1)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_26653, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_26456, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)))=k1_seq_1(k16_sin_cos,X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_25935, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers))).#trainneg
cnf(i_0_26390, plain, (v1_xreal_0(k1_seq_1(k19_sin_cos,np__1)))).#trainneg
cnf(i_0_26401, plain, (v1_xreal_0(k1_seq_1(k19_sin_cos,np__2)))).#trainneg
cnf(i_0_26713, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k11_binop_2(np__2,k32_sin_cos),k3_xcmplx_0(k1_xcmplx_0,esk1_0)))!=np__1|~v1_xcmplx_0(k11_binop_2(np__2,k32_sin_cos))|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_26859, plain, (v1_xreal_0(k2_xcmplx_0(np__1,k32_sin_cos)))).#trainneg
cnf(i_0_26861, plain, (v1_xreal_0(k2_xcmplx_0(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_839, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_26969, plain, (k1_real_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1738, plain, (v1_xcmplx_0(k16_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k2_numbers)))).#trainneg
cnf(i_0_26869, plain, (v1_xreal_0(k2_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_842, plain, (v7_ordinal1(k8_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_9700, plain, (m1_subset_1(k2_xcmplx_0(np__1,k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_843, plain, (v1_xcmplx_0(k8_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_844, plain, (v1_xxreal_0(k8_real_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_26076, plain, (k1_seq_1(k19_sin_cos,k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)))=k1_seq_1(k19_sin_cos,k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_26548, plain, (k1_seq_1(k16_sin_cos,k1_real_1(X1))=k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_27223, plain, (k1_real_1(k1_seq_1(k16_sin_cos,X1))=k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)),k1_numbers)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1757, plain, (v1_xcmplx_0(k19_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k2_numbers)))).#trainneg
cnf(i_0_9701, plain, (m1_subset_1(k2_xcmplx_0(np__1,k32_sin_cos),k2_numbers))).#trainneg
cnf(i_0_8728, plain, (k3_xcmplx_0(k4_xcmplx_0(k1_xcmplx_0),np__0)=np__0)).#trainneg
cnf(i_0_845, plain, (v1_int_1(k8_real_1(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_27563, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),np__1)=k3_xcmplx_0(np__1,k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_27564, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),k32_sin_cos)=k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_27613, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_27759, plain, (k1_real_1(k1_seq_1(k16_sin_cos,X1))=k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_27758, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_27597, plain, (v1_xreal_0(k4_xcmplx_0(k1_xcmplx_0))|~v1_xreal_0(k1_xcmplx_0))).#trainneg
cnf(i_0_10249, plain, (m1_subset_1(k2_xcmplx_0(np__1,np__2),k1_numbers))).#trainneg
cnf(i_0_27565, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),np__2)=k3_xcmplx_0(np__2,k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_27571, negated_conjecture, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),esk1_0)=k3_xcmplx_0(esk1_0,k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_885, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_28252, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10250, plain, (m1_subset_1(k2_xcmplx_0(np__1,np__2),k2_numbers))).#trainneg
cnf(i_0_26726, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(k11_binop_2(np__2,k32_sin_cos),esk1_0)))!=np__1|~v1_xcmplx_0(k11_binop_2(np__2,k32_sin_cos))|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_28451, plain, (m1_subset_1(k4_xcmplx_0(k1_xcmplx_0),k1_numbers)|~v1_xreal_0(k1_xcmplx_0))).#trainneg
cnf(i_0_28250, plain, (v1_xreal_0(k1_xcmplx_0)|~v1_xreal_0(k4_xcmplx_0(k1_xcmplx_0)))).#trainneg
cnf(i_0_28111, negated_conjecture, (k3_xcmplx_0(k3_xcmplx_0(X1,esk1_0),np__1)=k3_xcmplx_0(np__1,k3_xcmplx_0(X1,esk1_0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_28148, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,np__1),np__1)=k3_xcmplx_0(np__1,k3_xcmplx_0(X1,np__1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_27798, plain, (k1_seq_1(k16_sin_cos,k1_real_1(k4_xcmplx_0(X1)))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_27829, plain, (k1_real_1(k1_seq_1(k16_sin_cos,X1))=k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1))|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_27827, plain, (k1_real_1(k1_seq_1(k16_sin_cos,X1))=k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1))|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_26727, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(esk1_0,k3_xcmplx_0(k11_binop_2(np__2,k32_sin_cos),k1_xcmplx_0)))!=np__1|~v1_xcmplx_0(k11_binop_2(np__2,k32_sin_cos))|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_28152, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,k32_sin_cos),np__1)=k3_xcmplx_0(np__1,k3_xcmplx_0(X1,k32_sin_cos))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_28153, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,np__2),np__1)=k3_xcmplx_0(np__1,k3_xcmplx_0(X1,np__2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_889, plain, (k1_seq_1(k16_sin_cos,k2_xcmplx_0(X2,k8_real_1(k8_real_1(np__2,k32_sin_cos),X1)))=k1_seq_1(k16_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(k8_real_1(k8_real_1(np__2,k32_sin_cos),X1))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_906, plain, (v1_xreal_0(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)))|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_29080, plain, (k1_seq_1(k16_sin_cos,k1_real_1(X1))=k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11906, plain, (m1_subset_1(k2_xcmplx_0(np__2,k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_26728, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(esk1_0,k3_xcmplx_0(k1_xcmplx_0,k11_binop_2(np__2,k32_sin_cos))))!=np__1|~v1_xcmplx_0(k11_binop_2(np__2,k32_sin_cos))|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_28158, negated_conjecture, (k3_xcmplx_0(k3_xcmplx_0(X1,esk1_0),k32_sin_cos)=k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(X1,esk1_0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_28195, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,np__1),k32_sin_cos)=k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(X1,np__1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_918, plain, (k1_seq_1(k19_sin_cos,k7_real_1(k3_xcmplx_0(X1,k8_real_1(np__2,k32_sin_cos)),X2))=k1_seq_1(k19_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(k8_real_1(np__2,k32_sin_cos))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_919, plain, (k1_seq_1(k16_sin_cos,k7_real_1(k3_xcmplx_0(X1,k8_real_1(np__2,k32_sin_cos)),X2))=k1_seq_1(k16_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(k8_real_1(np__2,k32_sin_cos))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_26585, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,k1_real_1(X1)))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_26072, plain, (k1_seq_1(k19_sin_cos,k1_seq_1(k16_sin_cos,k1_real_1(X1)))=k1_seq_1(k19_sin_cos,k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_25584, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_11907, plain, (m1_subset_1(k2_xcmplx_0(np__2,k32_sin_cos),k2_numbers))).#trainneg
cnf(i_0_8075, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(np__2,k3_xcmplx_0(k3_xcmplx_0(esk1_0,k1_xcmplx_0),k32_sin_cos)))!=np__1|~v1_xcmplx_0(k3_xcmplx_0(esk1_0,k1_xcmplx_0)))).#trainneg
cnf(i_0_28199, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,k32_sin_cos),k32_sin_cos)=k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(X1,k32_sin_cos))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_28200, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,np__2),k32_sin_cos)=k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(X1,np__2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_928, plain, (k2_xcmplx_0(k21_sin_cos(X1),k2_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k18_sin_cos(X1)),X2))=k2_xcmplx_0(k15_sin_cos(k3_xcmplx_0(X1,k1_xcmplx_0)),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k18_sin_cos(X1)))|~v1_xcmplx_0(k21_sin_cos(X1))|~v1_xcmplx_0(k18_sin_cos(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_29815, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(np__2,k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,esk1_0),k32_sin_cos)))!=np__1)).#trainneg
cnf(i_0_4529, plain, (v1_xcmplx_0(k16_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_27495, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,esk1_0),k11_binop_2(np__2,k32_sin_cos)))!=np__1|~v1_xcmplx_0(k11_binop_2(np__2,k32_sin_cos))|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_929, plain, (v1_xcmplx_0(k15_sin_cos(k3_xcmplx_0(X1,k1_xcmplx_0)))|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k18_sin_cos(X1)))|~v1_xcmplx_0(k21_sin_cos(X1))|~v1_xcmplx_0(k18_sin_cos(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15175, plain, (m1_subset_1(k3_xcmplx_0(k32_sin_cos,k32_sin_cos),k2_numbers))).#trainneg
cnf(i_0_930, plain, (k2_xcmplx_0(k21_sin_cos(X1),k3_xcmplx_0(k1_xcmplx_0,k1_seq_1(k16_sin_cos,X1)))=k15_sin_cos(k3_xcmplx_0(X1,k1_xcmplx_0))|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_934, plain, (k4_xcmplx_0(k2_xcmplx_0(k1_real_1(X1),X2))=k2_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_29732, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_29402, plain, (k1_seq_1(k16_sin_cos,k4_xcmplx_0(k1_real_1(X1)))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_29426, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k1_xcmplx_0,k3_xcmplx_0(esk1_0,k11_binop_2(np__2,k32_sin_cos))))!=np__1|~v1_xcmplx_0(k11_binop_2(np__2,k32_sin_cos))|~v1_xreal_0(k11_binop_2(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_28210, plain, (m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4531, plain, (v1_xcmplx_0(k19_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_16124, negated_conjecture, (m1_subset_1(k3_xcmplx_0(k1_xcmplx_0,esk1_0),k2_numbers))).#trainneg
cnf(i_0_937, plain, (v1_xboole_0(k1_real_1(X1))|~v1_xcmplx_0(k1_real_1(X1))|~v1_xboole_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_23031, plain, (m1_subset_1(k2_xcmplx_0(k1_xcmplx_0,np__2),k2_numbers))).#trainneg
cnf(i_0_28476, negated_conjecture, (k3_xcmplx_0(k3_xcmplx_0(X1,esk1_0),np__2)=k3_xcmplx_0(np__2,k3_xcmplx_0(X1,esk1_0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_940, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,k1_real_1(X2)))=k2_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k1_real_1(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_942, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1)))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_29986, plain, (k1_seq_1(k19_sin_cos,k8_real_1(X1,k3_xcmplx_0(np__2,k32_sin_cos)))=np__1|~v1_xcmplx_0(k8_real_1(X1,k3_xcmplx_0(np__2,k32_sin_cos)))|~v7_ordinal1(X1)|~v1_xreal_0(k8_real_1(X1,k3_xcmplx_0(np__2,k32_sin_cos)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_29983, plain, (k1_seq_1(k19_sin_cos,k8_real_1(X1,k3_xcmplx_0(np__2,k32_sin_cos)))=np__1|~v1_xcmplx_0(k8_real_1(X1,k3_xcmplx_0(np__2,k32_sin_cos)))|~v7_ordinal1(X1)|~v1_xreal_0(k8_real_1(X1,k3_xcmplx_0(np__2,k32_sin_cos)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_30432, plain, (m1_subset_1(k4_xcmplx_0(k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_30436, negated_conjecture, (v1_xreal_0(k3_xcmplx_0(k1_xcmplx_0,esk1_0))|~v3_membered(k2_numbers))).#trainneg
cnf(i_0_28515, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,np__1),np__2)=k3_xcmplx_0(np__2,k3_xcmplx_0(X1,np__1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_28519, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,k32_sin_cos),np__2)=k3_xcmplx_0(np__2,k3_xcmplx_0(X1,k32_sin_cos))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_944, plain, (k1_seq_1(k16_sin_cos,k4_xcmplx_0(k7_real_1(k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1),X2)))=k1_real_1(k1_seq_1(k16_sin_cos,X2))|~v1_xreal_0(k7_real_1(k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1),X2))|~v7_ordinal1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_945, plain, (k1_seq_1(k16_sin_cos,k7_real_1(k8_real_1(X1,k3_xcmplx_0(np__2,k32_sin_cos)),X2))=k1_seq_1(k16_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(k3_xcmplx_0(np__2,k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_30811, plain, (k1_seq_1(k19_sin_cos,k3_xcmplx_0(np__2,k32_sin_cos))=np__1)).#trainneg
cnf(i_0_30851, plain, (v1_xreal_0(k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_30852, plain, (v7_ordinal1(k4_xcmplx_0(k32_sin_cos))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_30417, plain, (m1_subset_1(k1_xcmplx_0,k1_numbers)|~m1_subset_1(k2_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_28520, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,np__2),np__2)=k3_xcmplx_0(np__2,k3_xcmplx_0(X1,np__2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_947, plain, (k1_seq_1(k16_sin_cos,k7_real_1(k3_xcmplx_0(X1,k3_xcmplx_0(np__2,k32_sin_cos)),X2))=k1_seq_1(k16_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(k3_xcmplx_0(np__2,k32_sin_cos))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_29985, plain, (k1_seq_1(k19_sin_cos,k3_xcmplx_0(X1,k3_xcmplx_0(np__2,k32_sin_cos)))=np__1|~v1_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(np__2,k32_sin_cos)))|~v7_ordinal1(X1)|~v1_xreal_0(k3_xcmplx_0(X1,k3_xcmplx_0(np__2,k32_sin_cos)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_30188, plain, (k2_xcmplx_0(k1_real_1(X1),k4_xcmplx_0(X2))=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_30854, plain, (v1_xxreal_0(k4_xcmplx_0(k32_sin_cos))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_30078, negated_conjecture, (k2_xcmplx_0(k21_sin_cos(np__0),np__0)=np__1)).#trainneg
cnf(i_0_30855, plain, (v1_int_1(k4_xcmplx_0(k32_sin_cos))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_23635, plain, (m1_subset_1(k2_xcmplx_0(np__2,k1_xcmplx_0),k2_numbers))).#trainneg
cnf(i_0_28525, negated_conjecture, (k3_xcmplx_0(k3_xcmplx_0(X1,esk1_0),esk1_0)=k3_xcmplx_0(esk1_0,k3_xcmplx_0(X1,esk1_0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_28564, negated_conjecture, (k3_xcmplx_0(k3_xcmplx_0(X1,np__1),esk1_0)=k3_xcmplx_0(esk1_0,k3_xcmplx_0(X1,np__1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_948, plain, (k1_seq_1(k16_sin_cos,k7_real_1(k8_real_1(X1,k3_xcmplx_0(np__2,k32_sin_cos)),X2))=k1_seq_1(k16_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(k3_xcmplx_0(np__2,k32_sin_cos))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_949, plain, (k1_seq_1(k16_sin_cos,k7_real_1(X2,k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1)))=k1_seq_1(k16_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~m1_subset_1(k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1),k1_numbers))).#trainneg
cnf(i_0_31415, negated_conjecture, (k21_sin_cos(np__0)=np__1|~v1_xcmplx_0(k21_sin_cos(np__0)))).#trainneg
cnf(i_0_26341, plain, (m1_subset_1(k1_seq_1(k19_sin_cos,np__1),k1_numbers))).#trainneg
cnf(i_0_967, plain, (m1_subset_1(k3_xcmplx_0(X1,esk2_1(X2)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v3_membered(X2))).#trainneg
cnf(i_0_26352, plain, (m1_subset_1(k1_seq_1(k19_sin_cos,np__2),k1_numbers))).#trainneg
cnf(i_0_28568, negated_conjecture, (k3_xcmplx_0(k3_xcmplx_0(X1,k32_sin_cos),esk1_0)=k3_xcmplx_0(esk1_0,k3_xcmplx_0(X1,k32_sin_cos))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_28569, negated_conjecture, (k3_xcmplx_0(k3_xcmplx_0(X1,np__2),esk1_0)=k3_xcmplx_0(esk1_0,k3_xcmplx_0(X1,np__2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_950, plain, (k1_seq_1(k16_sin_cos,k2_xcmplx_0(k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1),X2))=k1_seq_1(k16_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~m1_subset_1(k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1),k1_numbers))).#trainneg
cnf(i_0_951, plain, (k1_seq_1(k16_sin_cos,k2_xcmplx_0(X2,k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1)))=k1_seq_1(k16_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_31364, plain, (k2_xcmplx_0(k1_real_1(X1),X2)=k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_31176, plain, (k1_seq_1(k16_sin_cos,k7_real_1(k3_xcmplx_0(np__2,k3_xcmplx_0(np__2,k32_sin_cos)),X1))=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_29019, plain, (k1_seq_1(k16_sin_cos,k1_real_1(X1))=k1_seq_1(k16_sin_cos,k4_xcmplx_0(X1))|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_30380, plain, (k1_seq_1(k16_sin_cos,k4_xcmplx_0(k4_xcmplx_0(X1)))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4667, plain, (v1_xxreal_0(k16_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_969, plain, (m1_subset_1(k3_xcmplx_0(X1,k9_binop_2(X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_29987, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,esk1_0),k3_xcmplx_0(np__2,k32_sin_cos)))!=np__1)).#trainneg
cnf(i_0_952, plain, (k1_seq_1(k16_sin_cos,k7_real_1(X2,k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1)))=k1_seq_1(k16_sin_cos,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_970, plain, (m1_subset_1(k3_xcmplx_0(X1,k18_sin_cos(X2)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_31610, plain, (k1_seq_1(k16_sin_cos,k7_real_1(X1,k3_xcmplx_0(np__2,k32_sin_cos)))=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_31815, plain, (k1_seq_1(k16_sin_cos,k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1))=np__0|~v7_ordinal1(X1)|~v1_xreal_0(k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1))|~v1_xcmplx_0(k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1)))).#trainneg
cnf(i_0_31617, negated_conjecture, (k21_sin_cos(np__0)=np__1|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_31620, plain, (v1_xcmplx_0(k1_seq_1(k19_sin_cos,np__1))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_971, plain, (m1_subset_1(k3_xcmplx_0(X1,k21_sin_cos(X2)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_31662, plain, (v1_xcmplx_0(k1_seq_1(k19_sin_cos,np__2))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_32051, negated_conjecture, (k15_sin_cos(k3_xcmplx_0(k32_sin_cos,k3_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,esk1_0),np__2)))!=np__1)).#trainneg
cnf(i_0_13929, plain, (v1_xcmplx_0(k3_xcmplx_0(k32_sin_cos,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_973, plain, (m1_subset_1(k3_xcmplx_0(X1,k8_real_1(X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_974, plain, (m1_subset_1(k3_xcmplx_0(X1,k1_real_1(X2)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_31788, plain, (k1_seq_1(k16_sin_cos,k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1))=np__0|~v7_ordinal1(X1)|~m1_subset_1(k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1),k1_numbers)|~v1_xcmplx_0(k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1)))).#trainneg
cnf(i_0_31592, negated_conjecture, (k1_seq_1(k16_sin_cos,k2_xcmplx_0(k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0)),X1))=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_31593, negated_conjecture, (k1_seq_1(k16_sin_cos,k2_xcmplx_0(X1,k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0))))=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_31395, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_31615, negated_conjecture, (k21_sin_cos(np__0)=np__1|~m1_subset_1(k21_sin_cos(np__0),k1_numbers))).#trainneg
cnf(i_0_9947, negated_conjecture, (k3_xcmplx_0(k4_xcmplx_0(esk1_0),np__0)=np__0)).#trainneg
cnf(i_0_975, plain, (m1_subset_1(k3_xcmplx_0(X1,np__1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_976, plain, (m1_subset_1(k3_xcmplx_0(X1,k11_binop_2(X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_977, plain, (m1_subset_1(k3_xcmplx_0(X1,k32_sin_cos),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_32496, negated_conjecture, (k1_seq_1(k16_sin_cos,k2_xcmplx_0(X1,k3_xcmplx_0(np__2,k3_xcmplx_0(k32_sin_cos,esk1_0))))=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_32645, negated_conjecture, (k21_sin_cos(np__0)=np__1)).#trainneg
cnf(i_0_4845, plain, (v1_xxreal_0(k19_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_32248, plain, (v1_xcmplx_0(k1_seq_1(k19_sin_cos,np__1)))).#trainneg
cnf(i_0_32289, plain, (v1_xcmplx_0(k1_seq_1(k19_sin_cos,np__2)))).#trainneg
cnf(i_0_978, plain, (m1_subset_1(k3_xcmplx_0(X1,np__2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_984, plain, (v7_ordinal1(k4_xcmplx_0(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_985, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_32222, plain, (k1_seq_1(k16_sin_cos,k2_xcmplx_0(X1,k3_xcmplx_0(np__2,k32_sin_cos)))=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_32462, plain, (k1_seq_1(k16_sin_cos,k1_real_1(k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1)))=np__0|~m1_subset_1(k8_real_1(k3_xcmplx_0(np__2,k32_sin_cos),X1),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3571, plain, (m1_subset_1(k4_xcmplx_0(X1),k2_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_32316, plain, (v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k32_sin_cos)))).#trainneg
cnf(i_0_986, plain, (v1_xxreal_0(k4_xcmplx_0(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1122, plain, (m1_subset_1(k2_xcmplx_0(np__1,k1_xcmplx_0),k1_numbers)|~m1_subset_1(k1_xcmplx_0,k1_numbers))).#trainneg
cnf(i_0_987, plain, (v1_int_1(k4_xcmplx_0(X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_995, plain, (k3_xcmplx_0(X2,k3_xcmplx_0(X1,k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0)))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(k1_xcmplx_0,k1_xcmplx_0))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 178
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 200
# Removed in clause preprocessing      : 21
# Initial clauses in saturation        : 179
# Processed clauses                    : 2834
# ...of these trivial                  : 184
# ...subsumed                          : 1349
# ...remaining for further processing  : 1301
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 47
# Backward-rewritten                   : 85
# Generated clauses                    : 33232
# ...of the previous two non-trivial   : 30047
# Contextual simplify-reflections      : 102
# Paramodulations                      : 33232
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1169
#    Positive orientable unit clauses  : 275
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 40
#    Non-unit-clauses                  : 854
# Current number of unprocessed clauses: 27068
# ...number of literals in the above   : 104990
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1433
# Proof object given clauses           : 142
# Proof search given clauses           : 1301
# Clause-clause subsumption calls (NU) : 144598
# Rec. Clause-clause subsumption calls : 79422
# Non-unit clause-clause subsumptions  : 1244
# Unit Clause-clause subsumption calls : 8529
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 44
# BW rewrite match successes           : 31
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 596256

# -------------------------------------------------
# User time                : 45.267 s
# System time              : 1.464 s
# Total time               : 46.732 s
# Maximum resident set size: 399760 pages
45.26user 1.54system 0:47.54elapsed 98%CPU (0avgtext+0avgdata 400052maxresident)k
0inputs+464outputs (0major+441049minor)pagefaults 0swaps
