2020-06-04 13:10:08.795625: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_8_bb1/premsel_enigma_01_2020_T30_loop02_epoch_339
2020-06-04 13:10:08.877495: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-06-04 13:10:09.014618: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-06-04 13:10:09.504649: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-06-04 13:10:11.100653: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 2305063 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_8_bb1/premsel_enigma_01_2020_T30_loop02_epoch_339' loaded (query=256; context=512; weigths=real; len_mult=0.000000).
# Preprocessing time       : 3.114 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t12_ordinal1, axiom, ![X1]:(v1_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>![X3]:(v3_ordinal1(X3)=>((r1_tarski(X1,X2)&r2_hidden(X2,X3))=>r2_hidden(X1,X3))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', t12_ordinal1)).
fof(redefinition_r1_ordinal1, axiom, ![X1, X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>(r1_ordinal1(X1,X2)<=>r1_tarski(X1,X2))), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', redefinition_r1_ordinal1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', t2_xboole_1)).
fof(cc3_ordinal1, axiom, ![X1]:(v1_xboole_0(X1)=>v3_ordinal1(X1)), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', cc3_ordinal1)).
fof(cc1_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>(v1_ordinal1(X1)&v2_ordinal1(X1))), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', cc1_ordinal1)).
fof(t33_ordinal2, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>![X3]:(v3_ordinal1(X3)=>(r1_ordinal1(X1,X2)=>r1_ordinal1(k10_ordinal2(X3,X1),k10_ordinal2(X3,X2)))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', t33_ordinal2)).
fof(t27_ordinal2, axiom, ![X1]:(v3_ordinal1(X1)=>k10_ordinal2(X1,k1_xboole_0)=X1), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', t27_ordinal2)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', fc1_xboole_0)).
fof(t34_ordinal2, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>![X3]:(v3_ordinal1(X3)=>(r1_ordinal1(X1,X2)=>r1_ordinal1(k10_ordinal2(X1,X3),k10_ordinal2(X2,X3)))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', t34_ordinal2)).
fof(dt_k10_ordinal2, axiom, ![X1, X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>v3_ordinal1(k10_ordinal2(X1,X2))), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', dt_k10_ordinal2)).
fof(t32_ordinal2, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>![X3]:(v3_ordinal1(X3)=>(r2_hidden(X1,X2)=>r2_hidden(k10_ordinal2(X3,X1),k10_ordinal2(X3,X2)))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', t32_ordinal2)).
fof(cc5_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(m1_subset_1(X2,X1)=>v3_ordinal1(X2))), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', cc5_ordinal1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', t1_subset)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', t2_subset)).
fof(existence_m1_subset_1, axiom, ![X1]:?[X2]:m1_subset_1(X2,X1), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', existence_m1_subset_1)).
fof(t28_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>(v4_ordinal1(X1)<=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X2,X1)=>r2_hidden(k1_ordinal1(X2),X1))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', t28_ordinal1)).
fof(d1_ordinal1, axiom, ![X1]:k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1)), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', d1_ordinal1)).
fof(t21_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,X2)<=>r1_ordinal1(k1_ordinal1(X1),X2)))), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', t21_ordinal1)).
fof(reflexivity_r1_ordinal1, axiom, ![X1, X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>r1_ordinal1(X1,X1)), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', reflexivity_r1_ordinal1)).
fof(rc1_ordinal2, axiom, ?[X1]:(((v1_ordinal1(X1)&v2_ordinal1(X1))&v3_ordinal1(X1))&v4_ordinal1(X1)), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', rc1_ordinal2)).
fof(antisymmetry_r2_hidden, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>~(r2_hidden(X2,X1))), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', antisymmetry_r2_hidden)).
fof(fc2_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>(~(v1_xboole_0(k1_ordinal1(X1)))&v3_ordinal1(k1_ordinal1(X1)))), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', fc2_ordinal1)).
fof(t16_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r1_ordinal1(X1,X2)|r2_hidden(X2,X1)))), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', t16_ordinal1)).
fof(t11_ordinal1, axiom, ![X1]:(v1_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_xboole_0(X1,X2)=>r2_hidden(X1,X2)))), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', t11_ordinal1)).
fof(d8_xboole_0, axiom, ![X1, X2]:(r2_xboole_0(X1,X2)<=>(r1_tarski(X1,X2)&X1!=X2)), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', d8_xboole_0)).
fof(t28_ordinal2, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>k10_ordinal2(X1,k1_ordinal1(X2))=k1_ordinal1(k10_ordinal2(X1,X2)))), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', t28_ordinal2)).
fof(t8_boole, axiom, ![X1, X2]:~(((v1_xboole_0(X1)&X1!=X2)&v1_xboole_0(X2))), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', t8_boole)).
fof(t27_ordinal3, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>~((r1_ordinal1(X1,X2)&![X3]:(v3_ordinal1(X3)=>X2!=k10_ordinal2(X1,X3)))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', t27_ordinal3)).
fof(t29_ordinal3, conjecture, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(v4_ordinal1(X1)=>(X1=k1_xboole_0|v4_ordinal1(k10_ordinal2(X2,X1)))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', t29_ordinal3)).
fof(fc1_ordinal1, axiom, ![X1]:~(v1_xboole_0(k1_ordinal1(X1))), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', fc1_ordinal1)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', t1_boole)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('/local1/mptp/loop2/hardpred/preds__0.05/t29_ordinal3', commutativity_k2_xboole_0)).
fof(c_0_32, plain, ![X49, X50, X51]:(~v1_ordinal1(X49)|(~v3_ordinal1(X50)|(~v3_ordinal1(X51)|(~r1_tarski(X49,X50)|~r2_hidden(X50,X51)|r2_hidden(X49,X51))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_ordinal1])])])).
fof(c_0_33, plain, ![X58, X59]:((~r1_ordinal1(X58,X59)|r1_tarski(X58,X59)|(~v3_ordinal1(X58)|~v3_ordinal1(X59)))&(~r1_tarski(X58,X59)|r1_ordinal1(X58,X59)|(~v3_ordinal1(X58)|~v3_ordinal1(X59)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r1_ordinal1])])])).
fof(c_0_34, plain, ![X19]:r1_tarski(k1_xboole_0,X19), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_35, plain, ![X66]:(~v1_xboole_0(X66)|v3_ordinal1(X66)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_ordinal1])])).
cnf(c_0_36, plain, (r2_hidden(X1,X3)|~v1_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~r1_tarski(X1,X2)|~r2_hidden(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_37, plain, (r1_tarski(X1,X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_38, plain, ![X121]:((v1_ordinal1(X121)|~v3_ordinal1(X121))&(v2_ordinal1(X121)|~v3_ordinal1(X121))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_ordinal1])])])).
fof(c_0_39, plain, ![X55, X56, X57]:(~v3_ordinal1(X55)|(~v3_ordinal1(X56)|(~v3_ordinal1(X57)|(~r1_ordinal1(X55,X56)|r1_ordinal1(k10_ordinal2(X57,X55),k10_ordinal2(X57,X56)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t33_ordinal2])])])).
fof(c_0_40, plain, ![X41]:(~v3_ordinal1(X41)|k10_ordinal2(X41,k1_xboole_0)=X41), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t27_ordinal2])])).
cnf(c_0_41, plain, (r1_ordinal1(X1,X2)|~r1_tarski(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_43, plain, (v3_ordinal1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_44, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_45, plain, (r2_hidden(X1,X3)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v1_ordinal1(X1)|~r1_tarski(X1,X2)|~r2_hidden(X2,X3)), inference(evalgc,[status(thm)],[c_0_36])).
cnf(c_0_46, plain, (r1_tarski(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(X1,X2)), inference(evalgc,[status(thm)],[c_0_37])).
cnf(c_0_47, plain, (v1_ordinal1(X1)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
fof(c_0_48, plain, ![X52, X53, X54]:(~v3_ordinal1(X52)|(~v3_ordinal1(X53)|(~v3_ordinal1(X54)|(~r1_ordinal1(X52,X53)|r1_ordinal1(k10_ordinal2(X52,X54),k10_ordinal2(X53,X54)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t34_ordinal2])])])).
cnf(c_0_49, plain, (r1_ordinal1(k10_ordinal2(X3,X1),k10_ordinal2(X3,X2))|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~r1_ordinal1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_50, plain, (k10_ordinal2(X1,k1_xboole_0)=X1|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_51, plain, (r1_ordinal1(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_41])).
cnf(c_0_52, plain, (r1_tarski(k1_xboole_0,X1)), inference(evalgc,[status(thm)],[c_0_42])).
cnf(c_0_53, plain, (v3_ordinal1(X1)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_43])).
cnf(c_0_54, plain, (v1_xboole_0(k1_xboole_0)), inference(evalgc,[status(thm)],[c_0_44])).
cnf(c_0_55, plain, (r2_hidden(X1,X2)|~v1_ordinal1(X1)|~r2_hidden(X3,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~r1_ordinal1(X1,X3)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46])])).
cnf(c_0_56, plain, (v1_ordinal1(X1)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_47])).
cnf(c_0_57, plain, (r1_ordinal1(k10_ordinal2(X1,X3),k10_ordinal2(X2,X3))|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~r1_ordinal1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_58, plain, ![X61, X62]:(~v3_ordinal1(X61)|~v3_ordinal1(X62)|v3_ordinal1(k10_ordinal2(X61,X62))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k10_ordinal2])])).
cnf(c_0_59, plain, (r1_ordinal1(k10_ordinal2(X3,X1),k10_ordinal2(X3,X2))|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(X1,X2)), inference(evalgc,[status(thm)],[c_0_49])).
cnf(c_0_60, plain, (k10_ordinal2(X1,k1_xboole_0)=X1|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_50])).
cnf(c_0_61, plain, (r1_ordinal1(k1_xboole_0,X1)|~v3_ordinal1(X1)|~v3_ordinal1(k1_xboole_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52])])).
cnf(c_0_62, plain, (v3_ordinal1(k1_xboole_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54])])).
cnf(c_0_63, plain, (r2_hidden(X1,X2)|~r1_ordinal1(X1,X3)|~r2_hidden(X3,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_64, plain, (r1_ordinal1(k10_ordinal2(X1,X3),k10_ordinal2(X2,X3))|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(X1,X2)), inference(evalgc,[status(thm)],[c_0_57])).
cnf(c_0_65, plain, (v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_58])).
fof(c_0_66, plain, ![X28, X29, X30]:(~v3_ordinal1(X28)|(~v3_ordinal1(X29)|(~v3_ordinal1(X30)|(~r2_hidden(X28,X29)|r2_hidden(k10_ordinal2(X30,X28),k10_ordinal2(X30,X29)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t32_ordinal2])])])).
fof(c_0_67, plain, ![X82, X83]:(~v3_ordinal1(X82)|(~m1_subset_1(X83,X82)|v3_ordinal1(X83))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc5_ordinal1])])])).
fof(c_0_68, plain, ![X72, X73]:(~r2_hidden(X72,X73)|m1_subset_1(X72,X73)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_69, plain, (r1_ordinal1(X1,k10_ordinal2(X1,X2))|~r1_ordinal1(k1_xboole_0,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(k1_xboole_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60])])).
cnf(c_0_70, plain, (r1_ordinal1(k1_xboole_0,X1)|~v3_ordinal1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_62])])).
cnf(c_0_71, plain, (r2_hidden(k10_ordinal2(X1,X2),X3)|~r2_hidden(k10_ordinal2(X4,X2),X3)|~v3_ordinal1(X3)|~v3_ordinal1(k10_ordinal2(X4,X2))|~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(X1,X4)|~v3_ordinal1(X2)|~v3_ordinal1(X4)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64])])).
cnf(c_0_72, plain, (v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_65])).
cnf(c_0_73, plain, (r2_hidden(k10_ordinal2(X3,X1),k10_ordinal2(X3,X2))|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_74, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_75, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_76, plain, (r1_ordinal1(X1,k10_ordinal2(X1,X2))|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_62])]), c_0_70])).
cnf(c_0_77, plain, (r2_hidden(k10_ordinal2(X1,X2),X3)|~r1_ordinal1(X1,X4)|~r2_hidden(k10_ordinal2(X4,X2),X3)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X4)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_71, c_0_72]), c_0_72])).
cnf(c_0_78, plain, (r2_hidden(k10_ordinal2(X3,X1),k10_ordinal2(X3,X2))|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_73])).
cnf(c_0_79, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_74])).
cnf(c_0_80, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_75])).
cnf(c_0_81, plain, (r2_hidden(X1,X2)|~r2_hidden(k10_ordinal2(X1,X3),X2)|~v3_ordinal1(X2)|~v3_ordinal1(k10_ordinal2(X1,X3))|~v3_ordinal1(X1)|~v3_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_76])])).
cnf(c_0_82, plain, (r2_hidden(k10_ordinal2(X1,X2),k10_ordinal2(X3,X4))|~r1_ordinal1(X1,X3)|~v3_ordinal1(k10_ordinal2(X3,X4))|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~r2_hidden(X2,X4)|~v3_ordinal1(X4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78])])).
cnf(c_0_83, plain, (v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80])])).
fof(c_0_84, plain, ![X91, X92]:(~m1_subset_1(X91,X92)|(v1_xboole_0(X92)|r2_hidden(X91,X92))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_85, plain, ![X106]:m1_subset_1(esk10_1(X106),X106), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[existence_m1_subset_1])])).
fof(c_0_86, plain, ![X25, X26]:((~v4_ordinal1(X25)|(~v3_ordinal1(X26)|(~r2_hidden(X26,X25)|r2_hidden(k1_ordinal1(X26),X25)))|~v3_ordinal1(X25))&((v3_ordinal1(esk6_1(X25))|v4_ordinal1(X25)|~v3_ordinal1(X25))&((r2_hidden(esk6_1(X25),X25)|v4_ordinal1(X25)|~v3_ordinal1(X25))&(~r2_hidden(k1_ordinal1(esk6_1(X25)),X25)|v4_ordinal1(X25)|~v3_ordinal1(X25))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_ordinal1])])])])])).
fof(c_0_87, plain, ![X44]:k1_ordinal1(X44)=k2_xboole_0(X44,k1_tarski(X44)), inference(variable_rename,[status(thm)],[d1_ordinal1])).
fof(c_0_88, plain, ![X17, X18]:((~r2_hidden(X17,X18)|r1_ordinal1(k1_ordinal1(X17),X18)|~v3_ordinal1(X18)|~v3_ordinal1(X17))&(~r1_ordinal1(k1_ordinal1(X17),X18)|r2_hidden(X17,X18)|~v3_ordinal1(X18)|~v3_ordinal1(X17))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t21_ordinal1])])])])).
fof(c_0_89, plain, ![X74, X75]:(~v3_ordinal1(X74)|~v3_ordinal1(X75)|r1_ordinal1(X74,X74)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[reflexivity_r1_ordinal1])])).
fof(c_0_90, plain, (((v1_ordinal1(esk15_0)&v2_ordinal1(esk15_0))&v3_ordinal1(esk15_0))&v4_ordinal1(esk15_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc1_ordinal2])])).
cnf(c_0_91, plain, (r2_hidden(X1,X2)|~r2_hidden(k10_ordinal2(X1,X3),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X3)), inference(csr,[status(thm)],[c_0_81, c_0_72])).
cnf(c_0_92, plain, (r2_hidden(k10_ordinal2(X1,X2),k10_ordinal2(X3,X4))|~r1_ordinal1(X1,X3)|~r2_hidden(X2,X4)|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~v3_ordinal1(X4)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_82, c_0_83]), c_0_72])).
cnf(c_0_93, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_94, plain, (m1_subset_1(esk10_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_85])).
fof(c_0_95, plain, ![X93, X94]:(~r2_hidden(X93,X94)|~r2_hidden(X94,X93)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[antisymmetry_r2_hidden])])])).
cnf(c_0_96, plain, (r2_hidden(k1_ordinal1(X2),X1)|~v4_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(X2,X1)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_97, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_98, plain, (r2_hidden(X1,X2)|~r1_ordinal1(k1_ordinal1(X1),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_99, plain, (r1_ordinal1(X1,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_100, plain, (v3_ordinal1(esk15_0)), inference(split_conjunct,[status(thm)],[c_0_90])).
fof(c_0_101, plain, ![X65]:((~v1_xboole_0(k1_ordinal1(X65))|~v3_ordinal1(X65))&(v3_ordinal1(k1_ordinal1(X65))|~v3_ordinal1(X65))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_ordinal1])])])])).
cnf(c_0_102, plain, (v4_ordinal1(X1)|~r2_hidden(k1_ordinal1(esk6_1(X1)),X1)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_103, plain, (r2_hidden(X1,k10_ordinal2(X2,X3))|~v3_ordinal1(k10_ordinal2(X2,X3))|~v3_ordinal1(X1)|~v3_ordinal1(X4)|~r1_ordinal1(X1,X2)|~r2_hidden(X4,X3)|~v3_ordinal1(X2)|~v3_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92])])).
cnf(c_0_104, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_93])).
cnf(c_0_105, plain, (m1_subset_1(esk10_1(X1),X1)), inference(evalgc,[status(thm)],[c_0_94])).
cnf(c_0_106, plain, (~r2_hidden(X1,X2)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_107, plain, (r2_hidden(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v4_ordinal1(X1)|~r2_hidden(X2,X1)), inference(rw,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_108, plain, (r2_hidden(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),X2)), inference(rw,[status(thm)],[c_0_98, c_0_97])).
cnf(c_0_109, plain, (r1_ordinal1(X1,X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_99])).
cnf(c_0_110, plain, (v3_ordinal1(esk15_0)), inference(evalgc,[status(thm)],[c_0_100])).
cnf(c_0_111, plain, (v3_ordinal1(k1_ordinal1(X1))|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_112, plain, (v4_ordinal1(X1)|~v3_ordinal1(X1)|~r2_hidden(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))),X1)), inference(rw,[status(thm)],[c_0_102, c_0_97])).
cnf(c_0_113, plain, (r2_hidden(X1,k10_ordinal2(X2,X3))|~r1_ordinal1(X1,X2)|~r2_hidden(X4,X3)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_103, c_0_83]), c_0_72])).
cnf(c_0_114, plain, (v1_xboole_0(X1)|r2_hidden(esk10_1(X1),X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105])])).
fof(c_0_115, plain, ![X20, X21]:(~v3_ordinal1(X20)|(~v3_ordinal1(X21)|(r1_ordinal1(X20,X21)|r2_hidden(X21,X20)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_ordinal1])])])).
cnf(c_0_116, plain, (~r2_hidden(X2,X1)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_106])).
cnf(c_0_117, plain, (r2_hidden(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v4_ordinal1(X1)|~r2_hidden(X2,X1)), inference(evalgc,[status(thm)],[c_0_107])).
cnf(c_0_118, plain, (r2_hidden(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),X2)), inference(evalgc,[status(thm)],[c_0_108])).
cnf(c_0_119, plain, (r1_ordinal1(X1,X1)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110])])).
cnf(c_0_120, plain, (v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1)), inference(rw,[status(thm)],[c_0_111, c_0_97])).
cnf(c_0_121, plain, (v4_ordinal1(X1)|~v3_ordinal1(X1)|~r2_hidden(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))),X1)), inference(evalgc,[status(thm)],[c_0_112])).
cnf(c_0_122, plain, (r2_hidden(X1,k10_ordinal2(X2,X3))|v1_xboole_0(X3)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114])])).
cnf(c_0_123, plain, (r1_ordinal1(k1_ordinal1(X1),X2)|~r2_hidden(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
fof(c_0_124, plain, ![X42, X43]:(~v1_ordinal1(X42)|(~v3_ordinal1(X43)|(~r2_xboole_0(X42,X43)|r2_hidden(X42,X43)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_ordinal1])])])).
fof(c_0_125, plain, ![X7, X8]:(((r1_tarski(X7,X8)|~r2_xboole_0(X7,X8))&(X7!=X8|~r2_xboole_0(X7,X8)))&(~r1_tarski(X7,X8)|X7=X8|r2_xboole_0(X7,X8))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_xboole_0])])])).
cnf(c_0_126, plain, (r1_ordinal1(X1,X2)|r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_115])).
cnf(c_0_127, plain, (~r2_hidden(X2,k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(X1,X2)|~v4_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117])])).
cnf(c_0_128, plain, (r2_hidden(X1,k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119])])).
cnf(c_0_129, plain, (v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_120])).
cnf(c_0_130, plain, (v4_ordinal1(k10_ordinal2(X1,X2))|v1_xboole_0(X2)|~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(k2_xboole_0(esk6_1(k10_ordinal2(X1,X2)),k1_tarski(esk6_1(k10_ordinal2(X1,X2)))),X1)|~v3_ordinal1(k2_xboole_0(esk6_1(k10_ordinal2(X1,X2)),k1_tarski(esk6_1(k10_ordinal2(X1,X2)))))|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122])])).
cnf(c_0_131, plain, (r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_123, c_0_97])).
fof(c_0_132, plain, ![X45, X46]:(~v3_ordinal1(X45)|(~v3_ordinal1(X46)|k10_ordinal2(X45,k1_ordinal1(X46))=k1_ordinal1(k10_ordinal2(X45,X46)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_ordinal2])])])).
cnf(c_0_133, plain, (r2_hidden(X1,X2)|~v1_ordinal1(X1)|~v3_ordinal1(X2)|~r2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_124])).
cnf(c_0_134, plain, (X1=X2|r2_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_125])).
cnf(c_0_135, plain, (r2_hidden(X2,X1)|r1_ordinal1(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_126])).
fof(c_0_136, plain, ![X113, X114]:(~v1_xboole_0(X113)|X113=X114|~v1_xboole_0(X114)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_boole])])).
cnf(c_0_137, plain, (r2_hidden(k1_xboole_0,X1)|~r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_70]), c_0_62])])])).
cnf(c_0_138, plain, (~r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(X2,X1)|~v4_ordinal1(X1)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[c_0_127, c_0_83])).
cnf(c_0_139, plain, (r2_hidden(X1,k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1)), inference(csr,[status(thm)],[c_0_128, c_0_129])).
cnf(c_0_140, plain, (v1_xboole_0(X1)|v4_ordinal1(k10_ordinal2(X2,X1))|~r1_ordinal1(k2_xboole_0(esk6_1(k10_ordinal2(X2,X1)),k1_tarski(esk6_1(k10_ordinal2(X2,X1)))),X2)|~v3_ordinal1(k2_xboole_0(esk6_1(k10_ordinal2(X2,X1)),k1_tarski(esk6_1(k10_ordinal2(X2,X1)))))|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[c_0_130, c_0_72])).
cnf(c_0_141, plain, (r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_131])).
cnf(c_0_142, plain, (k10_ordinal2(X1,k1_ordinal1(X2))=k1_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_132])).
fof(c_0_143, plain, ![X22, X23]:((v3_ordinal1(esk5_2(X22,X23))|~r1_ordinal1(X22,X23)|~v3_ordinal1(X23)|~v3_ordinal1(X22))&(X23=k10_ordinal2(X22,esk5_2(X22,X23))|~r1_ordinal1(X22,X23)|~v3_ordinal1(X23)|~v3_ordinal1(X22))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t27_ordinal3])])])])])).
cnf(c_0_144, plain, (r2_hidden(X1,X2)|~v3_ordinal1(X2)|~v1_ordinal1(X1)|~r2_xboole_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_133])).
cnf(c_0_145, plain, (X1=X2|r2_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_134])).
cnf(c_0_146, plain, (r2_hidden(X1,X2)|r2_hidden(X3,X1)|~r2_hidden(X3,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_135])])).
cnf(c_0_147, plain, (X1=X2|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_136])).
cnf(c_0_148, plain, (r2_hidden(k1_xboole_0,X1)|~r2_hidden(X2,X1)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[c_0_137, c_0_83])).
fof(c_0_149, negated_conjecture, ~(![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(v4_ordinal1(X1)=>(X1=k1_xboole_0|v4_ordinal1(k10_ordinal2(X2,X1))))))), inference(assume_negation,[status(cth)],[t29_ordinal3])).
cnf(c_0_150, plain, (~r2_hidden(X1,X1)|~v4_ordinal1(X1)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139])])).
cnf(c_0_151, plain, (v1_xboole_0(X1)|v4_ordinal1(k10_ordinal2(X2,X1))|~v3_ordinal1(k2_xboole_0(esk6_1(k10_ordinal2(X2,X1)),k1_tarski(esk6_1(k10_ordinal2(X2,X1)))))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k10_ordinal2(X2,X1)),X2)|~v3_ordinal1(esk6_1(k10_ordinal2(X2,X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_141])])).
cnf(c_0_152, plain, (k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2)))=k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2)))|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_142, c_0_97]), c_0_97])).
cnf(c_0_153, plain, (X1=k10_ordinal2(X2,esk5_2(X2,X1))|~r1_ordinal1(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_143])).
fof(c_0_154, plain, ![X115]:~v1_xboole_0(k1_ordinal1(X115)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_ordinal1])])).
cnf(c_0_155, plain, (r2_hidden(X1,X2)|X1=X2|~v1_ordinal1(X1)|~v3_ordinal1(X2)|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_145])])).
cnf(c_0_156, plain, (r2_hidden(X1,X2)|r2_hidden(X2,X3)|~r2_hidden(X1,X3)|~v3_ordinal1(X3)|~v3_ordinal1(X2)), inference(csr,[status(thm)],[c_0_146, c_0_83])).
cnf(c_0_157, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_147])).
cnf(c_0_158, plain, (r2_hidden(k1_xboole_0,X1)|v1_xboole_0(X1)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_114])])).
fof(c_0_159, negated_conjecture, (v3_ordinal1(esk1_0)&(v3_ordinal1(esk2_0)&(v4_ordinal1(esk1_0)&(esk1_0!=k1_xboole_0&~v4_ordinal1(k10_ordinal2(esk2_0,esk1_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_149])])])).
cnf(c_0_160, plain, (~v4_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(k10_ordinal2(X1,X2))|~r2_hidden(X2,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_78])])).
cnf(c_0_161, plain, (v1_xboole_0(X1)|v4_ordinal1(k10_ordinal2(X2,X1))|~r2_hidden(esk6_1(k10_ordinal2(X2,X1)),X2)|~v3_ordinal1(k2_xboole_0(esk6_1(k10_ordinal2(X2,X1)),k1_tarski(esk6_1(k10_ordinal2(X2,X1)))))|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[c_0_151, c_0_83])).
cnf(c_0_162, plain, (k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2)))=k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_152])).
cnf(c_0_163, plain, (k10_ordinal2(X2,esk5_2(X2,X1))=X1|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(X2,X1)), inference(evalgc,[status(thm)],[c_0_153])).
cnf(c_0_164, plain, (v3_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_143])).
cnf(c_0_165, plain, (~v1_xboole_0(k1_ordinal1(X1))), inference(split_conjunct,[status(thm)],[c_0_154])).
cnf(c_0_166, plain, (X1=X2|r2_hidden(X1,X2)|~v1_ordinal1(X1)|~v3_ordinal1(X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_46])])).
cnf(c_0_167, plain, (r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),X2)|r2_hidden(X2,X3)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~r2_hidden(X1,X3)|~v4_ordinal1(X3)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_117])])).
cnf(c_0_168, plain, (X1=X2|r2_hidden(k1_xboole_0,X2)|~v1_xboole_0(X1)|~v3_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_158])])).
cnf(c_0_169, negated_conjecture, (v3_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_159])).
cnf(c_0_170, plain, (~r2_hidden(X1,X1)|~v4_ordinal1(k10_ordinal2(X2,X1))|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[c_0_160, c_0_72])).
cnf(c_0_171, plain, (v4_ordinal1(esk15_0)), inference(split_conjunct,[status(thm)],[c_0_90])).
fof(c_0_172, plain, ![X64]:k2_xboole_0(X64,k1_xboole_0)=X64, inference(variable_rename,[status(thm)],[t1_boole])).
fof(c_0_173, plain, ![X79, X80]:k2_xboole_0(X79,X80)=k2_xboole_0(X80,X79), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
cnf(c_0_174, plain, (v1_xboole_0(X1)|v4_ordinal1(k10_ordinal2(X2,X1))|~r2_hidden(esk6_1(k10_ordinal2(X2,X1)),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k10_ordinal2(X2,X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_129])])).
cnf(c_0_175, plain, (k10_ordinal2(X1,k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))))=k2_xboole_0(X2,k1_tarski(X2))|~v3_ordinal1(esk5_2(X1,X2))|~v3_ordinal1(X1)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_163])])).
cnf(c_0_176, plain, (v3_ordinal1(esk5_2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(X1,X2)), inference(evalgc,[status(thm)],[c_0_164])).
cnf(c_0_177, plain, (~v1_xboole_0(k2_xboole_0(X1,k1_tarski(X1)))), inference(rw,[status(thm)],[c_0_165, c_0_97])).
cnf(c_0_178, plain, (X1=X2|r2_hidden(X1,X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[c_0_166, c_0_56])).
cnf(c_0_179, plain, (r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),X2)|r2_hidden(X2,X3)|~r2_hidden(X1,X3)|~v4_ordinal1(X3)|~v3_ordinal1(X3)|~v3_ordinal1(X2)), inference(csr,[status(thm)],[c_0_167, c_0_83])).
cnf(c_0_180, plain, (X1=X2|r2_hidden(k1_xboole_0,X2)|r2_hidden(k1_xboole_0,X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_158])])).
cnf(c_0_181, negated_conjecture, (v3_ordinal1(esk1_0)), inference(evalgc,[status(thm)],[c_0_169])).
cnf(c_0_182, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_159])).
cnf(c_0_183, plain, (~r2_hidden(k1_xboole_0,k1_xboole_0)|~v4_ordinal1(X1)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_60]), c_0_62])])])).
cnf(c_0_184, plain, (v4_ordinal1(esk15_0)), inference(evalgc,[status(thm)],[c_0_171])).
cnf(c_0_185, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_172])).
cnf(c_0_186, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_173])).
cnf(c_0_187, plain, (v1_xboole_0(X1)|v4_ordinal1(k10_ordinal2(X2,X1))|~r2_hidden(esk6_1(k10_ordinal2(X2,X1)),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[c_0_174, c_0_83])).
cnf(c_0_188, plain, (k10_ordinal2(X1,k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))))=k2_xboole_0(X2,k1_tarski(X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(csr,[status(thm)],[c_0_175, c_0_176])).
cnf(c_0_189, plain, (~v1_xboole_0(k2_xboole_0(X1,k1_tarski(X1)))), inference(evalgc,[status(thm)],[c_0_177])).
cnf(c_0_190, plain, (k2_xboole_0(X1,k1_tarski(X1))=X2|r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),X2)|~v3_ordinal1(X2)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(X1,X2)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_141])])).
cnf(c_0_191, plain, (r2_hidden(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))),X2)|r2_hidden(X2,X3)|~v4_ordinal1(X3)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~r2_hidden(X1,X3)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_117])])).
cnf(c_0_192, negated_conjecture, (X1=esk1_0|r2_hidden(k1_xboole_0,X1)|r2_hidden(k1_xboole_0,esk1_0)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_181])])).
cnf(c_0_193, negated_conjecture, (esk1_0!=k1_xboole_0), inference(evalgc,[status(thm)],[c_0_182])).
cnf(c_0_194, plain, (~r2_hidden(k1_xboole_0,k1_xboole_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_184]), c_0_110])])])).
cnf(c_0_195, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(evalgc,[status(thm)],[c_0_185])).
cnf(c_0_196, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(evalgc,[status(thm)],[c_0_186])).
cnf(c_0_197, negated_conjecture, (v4_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_159])).
cnf(c_0_198, plain, (v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1)|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_188]), c_0_189])])).
cnf(c_0_199, plain, (k2_xboole_0(X1,k1_tarski(X1))=X2|r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),X2)|~r2_hidden(X1,X2)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X2)), inference(csr,[status(thm)],[c_0_190, c_0_83])).
cnf(c_0_200, plain, (r2_hidden(esk6_1(X1),X1)|v4_ordinal1(X1)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_201, plain, (r2_hidden(k2_xboole_0(X2,k1_tarski(X2)),k10_ordinal2(X1,X3))|~r2_hidden(k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))),X3)|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~v3_ordinal1(k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_188])])).
cnf(c_0_202, plain, (r1_ordinal1(k10_ordinal2(X1,X2),X3)|~r1_ordinal1(X2,esk5_2(X1,X3))|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X1,X3))|~v3_ordinal1(X2)|~r1_ordinal1(X1,X3)|~v3_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_163])])).
cnf(c_0_203, plain, (r2_hidden(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))),X2)|r2_hidden(X2,X3)|~r2_hidden(X1,X3)|~v4_ordinal1(X3)|~v3_ordinal1(X3)|~v3_ordinal1(X2)), inference(csr,[status(thm)],[c_0_191, c_0_83])).
cnf(c_0_204, negated_conjecture, (r2_hidden(k1_xboole_0,esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_62]), c_0_193]), c_0_194])])).
cnf(c_0_205, plain, (k2_xboole_0(k1_xboole_0,X1)=X1), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_196])])).
cnf(c_0_206, negated_conjecture, (v4_ordinal1(esk1_0)), inference(evalgc,[status(thm)],[c_0_197])).
cnf(c_0_207, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(X2,X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X2,X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_129])])).
cnf(c_0_208, plain, (v4_ordinal1(X1)|k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1)))=X1|~v3_ordinal1(X1)|~r2_hidden(esk6_1(X1),X1)|~v3_ordinal1(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_199])])).
cnf(c_0_209, plain, (v4_ordinal1(X1)|r2_hidden(esk6_1(X1),X1)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_200])).
cnf(c_0_210, plain, (r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),k10_ordinal2(X2,X3))|~r1_ordinal1(X2,X1)|~r2_hidden(k2_xboole_0(esk5_2(X2,X1),k1_tarski(esk5_2(X2,X1))),X3)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[c_0_201, c_0_83])).
cnf(c_0_211, plain, (r1_ordinal1(k10_ordinal2(X1,X2),X3)|~r1_ordinal1(X2,esk5_2(X1,X3))|~r1_ordinal1(X1,X3)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3)), inference(csr,[status(thm)],[c_0_202, c_0_176])).
cnf(c_0_212, plain, (~r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),X1)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_139])])).
cnf(c_0_213, negated_conjecture, (r2_hidden(k2_xboole_0(k1_tarski(k1_xboole_0),k1_tarski(k1_tarski(k1_xboole_0))),X1)|r2_hidden(X1,esk1_0)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_204]), c_0_205]), c_0_205]), c_0_206]), c_0_181])])])).
cnf(c_0_214, plain, (v3_ordinal1(k1_tarski(k1_xboole_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_205]), c_0_62])])])).
cnf(c_0_215, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(X2,X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[c_0_207, c_0_176])).
cnf(c_0_216, plain, (k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1)))=X1|v4_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))))|~v3_ordinal1(X1)), inference(csr,[status(thm)],[c_0_208, c_0_209])).
cnf(c_0_217, plain, (v3_ordinal1(esk6_1(X1))|v4_ordinal1(X1)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_218, plain, (r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),k10_ordinal2(X2,X3))|~r1_ordinal1(X2,X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~r2_hidden(esk5_2(X2,X1),X3)|~v4_ordinal1(X3)|~v3_ordinal1(esk5_2(X2,X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_117])])).
cnf(c_0_219, plain, (r1_ordinal1(k10_ordinal2(X1,X2),X3)|r2_hidden(esk5_2(X1,X3),X2)|~r1_ordinal1(X1,X3)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(esk5_2(X1,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_135])])).
cnf(c_0_220, negated_conjecture, (r2_hidden(k1_tarski(k1_xboole_0),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_213]), c_0_214])])])).
cnf(c_0_221, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_209])])).
cnf(c_0_222, plain, (k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1)))=X1|v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_129])])).
cnf(c_0_223, plain, (v4_ordinal1(X1)|v3_ordinal1(esk6_1(X1))|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_217])).
cnf(c_0_224, plain, (r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),k10_ordinal2(X2,X3))|~r1_ordinal1(X2,X1)|~r2_hidden(esk5_2(X2,X1),X3)|~v4_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[c_0_218, c_0_83])).
cnf(c_0_225, plain, (r1_ordinal1(k10_ordinal2(X1,X2),X3)|r2_hidden(esk5_2(X1,X3),X2)|~r1_ordinal1(X1,X3)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3)), inference(csr,[status(thm)],[c_0_219, c_0_176])).
cnf(c_0_226, negated_conjecture, (r2_hidden(X1,k10_ordinal2(X2,esk1_0))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_220]), c_0_181])])])).
cnf(c_0_227, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),X1)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[c_0_221, c_0_129])).
cnf(c_0_228, plain, (k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1)))=X1|v4_ordinal1(X1)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[c_0_222, c_0_223])).
cnf(c_0_229, plain, (r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),k10_ordinal2(X2,X3))|r1_ordinal1(k10_ordinal2(X2,X3),X1)|~r1_ordinal1(X2,X1)|~v4_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_224, c_0_225])])).
cnf(c_0_230, negated_conjecture, (v4_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~r1_ordinal1(k2_xboole_0(esk6_1(k10_ordinal2(X1,esk1_0)),k1_tarski(esk6_1(k10_ordinal2(X1,esk1_0)))),X1)|~v3_ordinal1(k2_xboole_0(esk6_1(k10_ordinal2(X1,esk1_0)),k1_tarski(esk6_1(k10_ordinal2(X1,esk1_0)))))|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_226])])).
cnf(c_0_231, plain, (v4_ordinal1(X1)|~r1_ordinal1(X1,esk6_1(X1))|~v3_ordinal1(esk6_1(X1))|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_228])])).
cnf(c_0_232, plain, (v4_ordinal1(k10_ordinal2(X1,X2))|r1_ordinal1(k10_ordinal2(X1,X2),esk6_1(k10_ordinal2(X1,X2)))|~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(X1,esk6_1(k10_ordinal2(X1,X2)))|~v4_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_229])])).
cnf(c_0_233, negated_conjecture, (v4_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(k2_xboole_0(esk6_1(k10_ordinal2(X1,esk1_0)),k1_tarski(esk6_1(k10_ordinal2(X1,esk1_0)))))|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k10_ordinal2(X1,esk1_0)),X1)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_141])])).
cnf(c_0_234, plain, (v4_ordinal1(X1)|~r1_ordinal1(X1,esk6_1(X1))|~v3_ordinal1(X1)), inference(csr,[status(thm)],[c_0_231, c_0_223])).
cnf(c_0_235, plain, (r1_ordinal1(k10_ordinal2(X1,X2),esk6_1(k10_ordinal2(X1,X2)))|v4_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(X1,esk6_1(k10_ordinal2(X1,X2)))|~v4_ordinal1(X2)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,X2)))|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(csr,[status(thm)],[c_0_232, c_0_72])).
cnf(c_0_236, negated_conjecture, (v4_ordinal1(k10_ordinal2(X1,esk1_0))|~r2_hidden(esk6_1(k10_ordinal2(X1,esk1_0)),X1)|~v3_ordinal1(k2_xboole_0(esk6_1(k10_ordinal2(X1,esk1_0)),k1_tarski(esk6_1(k10_ordinal2(X1,esk1_0)))))|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)), inference(csr,[status(thm)],[c_0_233, c_0_83])).
cnf(c_0_237, plain, (v4_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(X1,esk6_1(k10_ordinal2(X1,X2)))|~v4_ordinal1(X2)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,X2)))|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_235])])).
cnf(c_0_238, negated_conjecture, (v4_ordinal1(k10_ordinal2(X1,esk1_0))|~r2_hidden(esk6_1(k10_ordinal2(X1,esk1_0)),X1)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_236, c_0_129])])).
cnf(c_0_239, plain, (v4_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(X1,esk6_1(k10_ordinal2(X1,X2)))|~v4_ordinal1(X2)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,X2)))|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(csr,[status(thm)],[c_0_237, c_0_72])).
cnf(c_0_240, negated_conjecture, (v4_ordinal1(k10_ordinal2(X1,esk1_0))|~r2_hidden(esk6_1(k10_ordinal2(X1,esk1_0)),X1)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)), inference(csr,[status(thm)],[c_0_238, c_0_83])).
cnf(c_0_241, plain, (v4_ordinal1(k10_ordinal2(X1,X2))|r2_hidden(esk6_1(k10_ordinal2(X1,X2)),X1)|~v4_ordinal1(X2)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,X2)))|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_135])])).
cnf(c_0_242, negated_conjecture, (~v4_ordinal1(k10_ordinal2(esk2_0,esk1_0))), inference(split_conjunct,[status(thm)],[c_0_159])).
cnf(c_0_243, negated_conjecture, (v4_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_241]), c_0_206]), c_0_181])])])).
cnf(c_0_244, negated_conjecture, (v3_ordinal1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_159])).
cnf(c_0_245, negated_conjecture, (~v4_ordinal1(k10_ordinal2(esk2_0,esk1_0))), inference(evalgc,[status(thm)],[c_0_242])).
cnf(c_0_246, negated_conjecture, (v4_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)), inference(csr,[status(thm)],[c_0_243, c_0_223])).
cnf(c_0_247, negated_conjecture, (v3_ordinal1(esk2_0)), inference(evalgc,[status(thm)],[c_0_244])).
cnf(c_0_248, negated_conjecture, (~v3_ordinal1(k10_ordinal2(esk2_0,esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_246]), c_0_247])])])).
cnf(c_0_249, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_248, c_0_72]), c_0_181]), c_0_247])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 250
# Proof object clause steps            : 186
# Proof object formula steps           : 64
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 43
# Proof object initial formulas used   : 32
# Proof object generating inferences   : 60
# Proof object simplifying inferences  : 76
# Training examples: 101 positive, 1286 negative
# Training: Positive examples begin
cnf(c_0_193, negated_conjecture, (esk1_0!=k1_xboole_0)).# trainpos
cnf(c_0_181, negated_conjecture, (v3_ordinal1(esk1_0))).# trainpos
cnf(c_0_247, negated_conjecture, (v3_ordinal1(esk2_0))).# trainpos
cnf(c_0_110, plain, (v3_ordinal1(esk15_0))).# trainpos
cnf(c_0_206, negated_conjecture, (v4_ordinal1(esk1_0))).# trainpos
cnf(c_0_184, plain, (v4_ordinal1(esk15_0))).# trainpos
cnf(c_0_54, plain, (v1_xboole_0(k1_xboole_0))).# trainpos
cnf(c_0_245, negated_conjecture, (~v4_ordinal1(k10_ordinal2(esk2_0,esk1_0)))).# trainpos
cnf(c_0_56, plain, (v1_ordinal1(X1)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_52, plain, (r1_tarski(k1_xboole_0,X1))).# trainpos
cnf(c_0_195, plain, (k2_xboole_0(X1,k1_xboole_0)=X1)).# trainpos
cnf(c_0_105, plain, (m1_subset_1(esk10_1(X1),X1))).# trainpos
cnf(c_0_53, plain, (v3_ordinal1(X1)|~v1_xboole_0(X1))).# trainpos
cnf(c_0_196, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1))).# trainpos
cnf(c_0_189, plain, (~v1_xboole_0(k2_xboole_0(X1,k1_tarski(X1))))).# trainpos
cnf(c_0_60, plain, (k10_ordinal2(X1,k1_xboole_0)=X1|~v3_ordinal1(X1))).# trainpos
cnf(c_0_157, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).# trainpos
cnf(c_0_80, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_72, plain, (v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_223, plain, (v4_ordinal1(X1)|v3_ordinal1(esk6_1(X1))|~v3_ordinal1(X1))).# trainpos
cnf(c_0_109, plain, (r1_ordinal1(X1,X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_116, plain, (~r2_hidden(X2,X1)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_129, plain, (v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1))).# trainpos
cnf(c_0_79, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_145, plain, (X1=X2|r2_xboole_0(X1,X2)|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_209, plain, (v4_ordinal1(X1)|r2_hidden(esk6_1(X1),X1)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_104, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_121, plain, (v4_ordinal1(X1)|~v3_ordinal1(X1)|~r2_hidden(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))),X1))).# trainpos
cnf(c_0_51, plain, (r1_ordinal1(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_135, plain, (r2_hidden(X2,X1)|r1_ordinal1(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_46, plain, (r1_tarski(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(X1,X2))).# trainpos
cnf(c_0_144, plain, (r2_hidden(X1,X2)|~v3_ordinal1(X2)|~v1_ordinal1(X1)|~r2_xboole_0(X1,X2))).# trainpos
cnf(c_0_176, plain, (v3_ordinal1(esk5_2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(X1,X2))).# trainpos
cnf(c_0_163, plain, (k10_ordinal2(X2,esk5_2(X2,X1))=X1|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(X2,X1))).# trainpos
cnf(c_0_141, plain, (r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_118, plain, (r2_hidden(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),X2))).# trainpos
cnf(c_0_45, plain, (r2_hidden(X1,X3)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v1_ordinal1(X1)|~r1_tarski(X1,X2)|~r2_hidden(X2,X3))).# trainpos
cnf(c_0_117, plain, (r2_hidden(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v4_ordinal1(X1)|~r2_hidden(X2,X1))).# trainpos
cnf(c_0_78, plain, (r2_hidden(k10_ordinal2(X3,X1),k10_ordinal2(X3,X2))|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_59, plain, (r1_ordinal1(k10_ordinal2(X3,X1),k10_ordinal2(X3,X2))|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(X1,X2))).# trainpos
cnf(c_0_64, plain, (r1_ordinal1(k10_ordinal2(X1,X3),k10_ordinal2(X2,X3))|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(X1,X2))).# trainpos
cnf(c_0_162, plain, (k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2)))=k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_155, plain, (r2_hidden(X1,X2)|X1=X2|~v1_ordinal1(X1)|~v3_ordinal1(X2)|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_55, plain, (r2_hidden(X1,X2)|~v1_ordinal1(X1)|~r2_hidden(X3,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~r1_ordinal1(X1,X3)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_83, plain, (v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_62, plain, (v3_ordinal1(k1_xboole_0))).# trainpos
cnf(c_0_61, plain, (r1_ordinal1(k1_xboole_0,X1)|~v3_ordinal1(X1)|~v3_ordinal1(k1_xboole_0))).# trainpos
cnf(c_0_119, plain, (r1_ordinal1(X1,X1)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_205, plain, (k2_xboole_0(k1_xboole_0,X1)=X1)).# trainpos
cnf(c_0_175, plain, (k10_ordinal2(X1,k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))))=k2_xboole_0(X2,k1_tarski(X2))|~v3_ordinal1(esk5_2(X1,X2))|~v3_ordinal1(X1)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).# trainpos
cnf(c_0_202, plain, (r1_ordinal1(k10_ordinal2(X1,X2),X3)|~r1_ordinal1(X2,esk5_2(X1,X3))|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X1,X3))|~v3_ordinal1(X2)|~r1_ordinal1(X1,X3)|~v3_ordinal1(X3))).# trainpos
cnf(c_0_127, plain, (~r2_hidden(X2,k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(X1,X2)|~v4_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).# trainpos
cnf(c_0_69, plain, (r1_ordinal1(X1,k10_ordinal2(X1,X2))|~r1_ordinal1(k1_xboole_0,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(k1_xboole_0))).# trainpos
cnf(c_0_114, plain, (v1_xboole_0(X1)|r2_hidden(esk10_1(X1),X1))).# trainpos
cnf(c_0_166, plain, (X1=X2|r2_hidden(X1,X2)|~v1_ordinal1(X1)|~v3_ordinal1(X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_146, plain, (r2_hidden(X1,X2)|r2_hidden(X3,X1)|~r2_hidden(X3,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_214, plain, (v3_ordinal1(k1_tarski(k1_xboole_0)))).# trainpos
cnf(c_0_137, plain, (r2_hidden(k1_xboole_0,X1)|~r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).# trainpos
cnf(c_0_81, plain, (r2_hidden(X1,X2)|~r2_hidden(k10_ordinal2(X1,X3),X2)|~v3_ordinal1(X2)|~v3_ordinal1(k10_ordinal2(X1,X3))|~v3_ordinal1(X1)|~v3_ordinal1(X3))).# trainpos
cnf(c_0_219, plain, (r1_ordinal1(k10_ordinal2(X1,X2),X3)|r2_hidden(esk5_2(X1,X3),X2)|~r1_ordinal1(X1,X3)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(esk5_2(X1,X3)))).# trainpos
cnf(c_0_128, plain, (r2_hidden(X1,k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1))).# trainpos
cnf(c_0_158, plain, (r2_hidden(k1_xboole_0,X1)|v1_xboole_0(X1)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_150, plain, (~r2_hidden(X1,X1)|~v4_ordinal1(X1)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_212, plain, (~r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),X1)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_167, plain, (r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),X2)|r2_hidden(X2,X3)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~r2_hidden(X1,X3)|~v4_ordinal1(X3)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_160, plain, (~v4_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(k10_ordinal2(X1,X2))|~r2_hidden(X2,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).# trainpos
cnf(c_0_183, plain, (~r2_hidden(k1_xboole_0,k1_xboole_0)|~v4_ordinal1(X1)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_194, plain, (~r2_hidden(k1_xboole_0,k1_xboole_0))).# trainpos
cnf(c_0_71, plain, (r2_hidden(k10_ordinal2(X1,X2),X3)|~r2_hidden(k10_ordinal2(X4,X2),X3)|~v3_ordinal1(X3)|~v3_ordinal1(k10_ordinal2(X4,X2))|~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(X1,X4)|~v3_ordinal1(X2)|~v3_ordinal1(X4)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_82, plain, (r2_hidden(k10_ordinal2(X1,X2),k10_ordinal2(X3,X4))|~r1_ordinal1(X1,X3)|~v3_ordinal1(k10_ordinal2(X3,X4))|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~r2_hidden(X2,X4)|~v3_ordinal1(X4))).# trainpos
cnf(c_0_103, plain, (r2_hidden(X1,k10_ordinal2(X2,X3))|~v3_ordinal1(k10_ordinal2(X2,X3))|~v3_ordinal1(X1)|~v3_ordinal1(X4)|~r1_ordinal1(X1,X2)|~r2_hidden(X4,X3)|~v3_ordinal1(X2)|~v3_ordinal1(X3))).# trainpos
cnf(c_0_122, plain, (r2_hidden(X1,k10_ordinal2(X2,X3))|v1_xboole_0(X3)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3))).# trainpos
cnf(c_0_130, plain, (v4_ordinal1(k10_ordinal2(X1,X2))|v1_xboole_0(X2)|~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(k2_xboole_0(esk6_1(k10_ordinal2(X1,X2)),k1_tarski(esk6_1(k10_ordinal2(X1,X2)))),X1)|~v3_ordinal1(k2_xboole_0(esk6_1(k10_ordinal2(X1,X2)),k1_tarski(esk6_1(k10_ordinal2(X1,X2)))))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).# trainpos
cnf(c_0_151, plain, (v1_xboole_0(X1)|v4_ordinal1(k10_ordinal2(X2,X1))|~v3_ordinal1(k2_xboole_0(esk6_1(k10_ordinal2(X2,X1)),k1_tarski(esk6_1(k10_ordinal2(X2,X1)))))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k10_ordinal2(X2,X1)),X2)|~v3_ordinal1(esk6_1(k10_ordinal2(X2,X1))))).# trainpos
cnf(c_0_191, plain, (r2_hidden(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))),X2)|r2_hidden(X2,X3)|~v4_ordinal1(X3)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~r2_hidden(X1,X3)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_174, plain, (v1_xboole_0(X1)|v4_ordinal1(k10_ordinal2(X2,X1))|~r2_hidden(esk6_1(k10_ordinal2(X2,X1)),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k10_ordinal2(X2,X1))))).# trainpos
cnf(c_0_198, plain, (v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1)|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).# trainpos
cnf(c_0_207, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(X2,X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X2,X1)))).# trainpos
cnf(c_0_201, plain, (r2_hidden(k2_xboole_0(X2,k1_tarski(X2)),k10_ordinal2(X1,X3))|~r2_hidden(k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))),X3)|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~v3_ordinal1(k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).# trainpos
cnf(c_0_221, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1))).# trainpos
cnf(c_0_218, plain, (r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),k10_ordinal2(X2,X3))|~r1_ordinal1(X2,X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~r2_hidden(esk5_2(X2,X1),X3)|~v4_ordinal1(X3)|~v3_ordinal1(esk5_2(X2,X1)))).# trainpos
cnf(c_0_229, plain, (r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),k10_ordinal2(X2,X3))|r1_ordinal1(k10_ordinal2(X2,X3),X1)|~r1_ordinal1(X2,X1)|~v4_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_232, plain, (v4_ordinal1(k10_ordinal2(X1,X2))|r1_ordinal1(k10_ordinal2(X1,X2),esk6_1(k10_ordinal2(X1,X2)))|~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(X1,esk6_1(k10_ordinal2(X1,X2)))|~v4_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,X2))))).# trainpos
cnf(c_0_168, plain, (X1=X2|r2_hidden(k1_xboole_0,X2)|~v1_xboole_0(X1)|~v3_ordinal1(X2))).# trainpos
cnf(c_0_180, plain, (X1=X2|r2_hidden(k1_xboole_0,X2)|r2_hidden(k1_xboole_0,X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_192, negated_conjecture, (X1=esk1_0|r2_hidden(k1_xboole_0,X1)|r2_hidden(k1_xboole_0,esk1_0)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_204, negated_conjecture, (r2_hidden(k1_xboole_0,esk1_0))).# trainpos
cnf(c_0_213, negated_conjecture, (r2_hidden(k2_xboole_0(k1_tarski(k1_xboole_0),k1_tarski(k1_tarski(k1_xboole_0))),X1)|r2_hidden(X1,esk1_0)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_220, negated_conjecture, (r2_hidden(k1_tarski(k1_xboole_0),esk1_0))).# trainpos
cnf(c_0_226, negated_conjecture, (r2_hidden(X1,k10_ordinal2(X2,esk1_0))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).# trainpos
cnf(c_0_230, negated_conjecture, (v4_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~r1_ordinal1(k2_xboole_0(esk6_1(k10_ordinal2(X1,esk1_0)),k1_tarski(esk6_1(k10_ordinal2(X1,esk1_0)))),X1)|~v3_ordinal1(k2_xboole_0(esk6_1(k10_ordinal2(X1,esk1_0)),k1_tarski(esk6_1(k10_ordinal2(X1,esk1_0)))))|~v3_ordinal1(X1))).# trainpos
cnf(c_0_233, negated_conjecture, (v4_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(k2_xboole_0(esk6_1(k10_ordinal2(X1,esk1_0)),k1_tarski(esk6_1(k10_ordinal2(X1,esk1_0)))))|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k10_ordinal2(X1,esk1_0)),X1)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,esk1_0))))).# trainpos
cnf(c_0_238, negated_conjecture, (v4_ordinal1(k10_ordinal2(X1,esk1_0))|~r2_hidden(esk6_1(k10_ordinal2(X1,esk1_0)),X1)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,esk1_0))))).# trainpos
cnf(c_0_190, plain, (k2_xboole_0(X1,k1_tarski(X1))=X2|r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),X2)|~v3_ordinal1(X2)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(X1,X2)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_208, plain, (v4_ordinal1(X1)|k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1)))=X1|~v3_ordinal1(X1)|~r2_hidden(esk6_1(X1),X1)|~v3_ordinal1(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1)))))).# trainpos
cnf(c_0_222, plain, (k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1)))=X1|v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(X1)))).# trainpos
cnf(c_0_231, plain, (v4_ordinal1(X1)|~r1_ordinal1(X1,esk6_1(X1))|~v3_ordinal1(esk6_1(X1))|~v3_ordinal1(X1))).# trainpos
cnf(c_0_237, plain, (v4_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(X1,esk6_1(k10_ordinal2(X1,X2)))|~v4_ordinal1(X2)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,X2)))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).# trainpos
cnf(c_0_241, plain, (v4_ordinal1(k10_ordinal2(X1,X2))|r2_hidden(esk6_1(k10_ordinal2(X1,X2)),X1)|~v4_ordinal1(X2)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,X2)))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).# trainpos
cnf(c_0_243, negated_conjecture, (v4_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,esk1_0))))).# trainpos
cnf(c_0_248, negated_conjecture, (~v3_ordinal1(k10_ordinal2(esk2_0,esk1_0)))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_43, plain, (o_0_0_xboole_0=k1_xboole_0)).#trainneg
cnf(i_0_92, plain, (v3_ordinal1(esk13_0))).#trainneg
cnf(i_0_93, plain, (v3_ordinal1(esk14_0))).#trainneg
cnf(i_0_95, plain, (v1_ordinal1(esk14_0))).#trainneg
cnf(i_0_101, plain, (v1_ordinal1(esk15_0))).#trainneg
cnf(i_0_85, plain, (~v1_xboole_0(esk12_0))).#trainneg
cnf(i_0_96, plain, (~v1_xboole_0(esk14_0))).#trainneg
cnf(i_0_80, plain, (v1_xboole_0(esk11_0))).#trainneg
cnf(i_0_84, plain, (v7_ordinal1(esk12_0))).#trainneg
cnf(i_0_94, plain, (v2_ordinal1(esk14_0))).#trainneg
cnf(i_0_100, plain, (v2_ordinal1(esk15_0))).#trainneg
cnf(i_0_90, plain, (v2_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_47, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_49, plain, (v3_ordinal1(k3_ordinal2(X1)))).#trainneg
cnf(i_0_55, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_88, plain, (k2_xboole_0(X1,X1)=X1)).#trainneg
cnf(i_0_76, plain, (~v1_xboole_0(k1_tarski(X1)))).#trainneg
cnf(i_0_97, plain, (~r2_xboole_0(X1,X1))).#trainneg
cnf(i_0_89, plain, (v3_ordinal1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_74, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7, plain, (X1!=X2|~r2_xboole_0(X1,X2))).#trainneg
cnf(i_0_56, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_9, plain, (X2=k1_xboole_0|k10_ordinal2(X3,X2)=k4_ordinal2(X1)|k9_xtuple_0(X1)!=X2|k1_funct_1(X1,esk3_3(X2,X3,X1))!=k10_ordinal2(X3,esk3_3(X2,X3,X1))|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v4_ordinal1(X2)|~v5_ordinal1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_ordinal2(X1))).#trainneg
cnf(i_0_8, plain, (r1_tarski(X1,X2)|~r2_xboole_0(X1,X2))).#trainneg
cnf(i_0_68, plain, (~r2_xboole_0(X2,X1)|~r2_xboole_0(X1,X2))).#trainneg
cnf(i_0_72, plain, (v1_xboole_0(X1)|~v1_xboole_0(k2_xboole_0(X2,X1)))).#trainneg
cnf(i_0_73, plain, (v1_xboole_0(X1)|~v1_xboole_0(k2_xboole_0(X1,X2)))).#trainneg
cnf(i_0_61, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_62, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_87, plain, (v3_ordinal1(k10_ordinal2(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_67, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_79, plain, (v3_ordinal1(k4_ordinal2(X1))|~v5_ordinal1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_78, plain, (v3_ordinal1(k9_xtuple_0(X1))|~v5_ordinal1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_60, plain, (k3_ordinal2(k10_xtuple_0(X1))=k4_ordinal2(X1)|~v5_ordinal1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_75, plain, (v3_ordinal1(k2_xboole_0(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_17, plain, (v5_ordinal1(esk4_2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_86, plain, (v7_ordinal1(k10_ordinal2(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_54, plain, (r1_ordinal1(X2,X1)|r1_ordinal1(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_16, plain, (v1_relat_1(esk4_2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_15, plain, (v1_funct_1(esk4_2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_14, plain, (v1_ordinal2(esk4_2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_13, plain, (k9_xtuple_0(esk4_2(X1,X2))=X1|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_39, plain, (r2_hidden(X1,k3_ordinal2(X2))|~v3_ordinal1(X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_64, plain, (v7_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_71, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_81, plain, (v3_ordinal1(k1_funct_1(X1,X2))|~v3_ordinal1(X2)|~v5_ordinal1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_ordinal2(X1))).#trainneg
cnf(i_0_31, plain, (X2=k10_xtuple_0(X1)|esk8_2(X1,X2)!=k1_funct_1(X1,X3)|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X3,k9_xtuple_0(X1))|~r2_hidden(esk8_2(X1,X2),X2))).#trainneg
cnf(i_0_32, plain, (r2_hidden(X3,X4)|X4!=k10_xtuple_0(X2)|X3!=k1_funct_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~r2_hidden(X1,k9_xtuple_0(X2)))).#trainneg
cnf(i_0_12, plain, (k1_funct_1(esk4_2(X2,X3),X1)=k10_ordinal2(X3,X1)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_11, plain, (X1=k1_xboole_0|k10_ordinal2(X2,X1)=k4_ordinal2(X3)|v3_ordinal1(esk3_3(X1,X2,X3))|k9_xtuple_0(X3)!=X1|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v4_ordinal1(X1)|~v5_ordinal1(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v1_ordinal2(X3))).#trainneg
cnf(i_0_30, plain, (X2=k10_xtuple_0(X1)|r2_hidden(esk8_2(X1,X2),X2)|r2_hidden(esk9_2(X1,X2),k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_29, plain, (X2=k10_xtuple_0(X1)|k1_funct_1(X1,esk9_2(X1,X2))=esk8_2(X1,X2)|r2_hidden(esk8_2(X1,X2),X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_33, plain, (k1_funct_1(X2,esk7_3(X2,X3,X1))=X1|X3!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_34, plain, (r2_hidden(esk7_3(X1,X2,X3),k9_xtuple_0(X1))|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_10, plain, (X1=k1_xboole_0|k10_ordinal2(X2,X1)=k4_ordinal2(X3)|r2_hidden(esk3_3(X1,X2,X3),X1)|k9_xtuple_0(X3)!=X1|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v4_ordinal1(X1)|~v5_ordinal1(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v1_ordinal2(X3))).#trainneg
cnf(i_0_140, plain, (X1=X2|~r2_xboole_0(X2,X1)|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_207, plain, (r2_hidden(X2,k10_ordinal2(X1,X3))|~r2_hidden(esk5_2(X1,X2),X3)|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~v3_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_183, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2))).#trainneg
cnf(i_0_175, plain, (v4_ordinal1(k3_ordinal2(X1))|~r2_hidden(k2_xboole_0(esk6_1(k3_ordinal2(X1)),k1_tarski(esk6_1(k3_ordinal2(X1)))),X1)|~v3_ordinal1(k2_xboole_0(esk6_1(k3_ordinal2(X1)),k1_tarski(esk6_1(k3_ordinal2(X1))))))).#trainneg
cnf(i_0_103, plain, (esk11_0=k1_xboole_0)).#trainneg
cnf(i_0_106, plain, (v3_ordinal1(esk12_0))).#trainneg
cnf(i_0_148, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,esk10_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_177, plain, (~v1_xboole_0(k3_ordinal2(X1))|~r2_hidden(X2,X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_239, plain, (r2_hidden(esk7_3(esk4_2(X1,X2),X3,X4),X1)|X3!=k10_xtuple_0(esk4_2(X1,X2))|~r2_hidden(X4,X3)|~v1_funct_1(esk4_2(X1,X2))|~v1_relat_1(esk4_2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_210, plain, (r1_ordinal1(k10_ordinal2(X1,X2),X1)|~r1_ordinal1(X2,k1_xboole_0)|~v3_ordinal1(X1)|~v3_ordinal1(k1_xboole_0)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_228, plain, (r2_hidden(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))),X3)|~r2_hidden(k10_ordinal2(X1,X2),X3)|~v4_ordinal1(X3)|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_229, plain, (r2_hidden(k10_ordinal2(X1,X2),X3)|~r1_ordinal1(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))),X3)|~v3_ordinal1(X3)|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_234, plain, (k10_ordinal2(X2,esk7_3(esk4_2(X1,X2),X3,X4))=X4|~r2_hidden(esk7_3(esk4_2(X1,X2),X3,X4),X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(esk7_3(esk4_2(X1,X2),X3,X4))|X3!=k10_xtuple_0(esk4_2(X1,X2))|~r2_hidden(X4,X3)|~v1_funct_1(esk4_2(X1,X2))|~v1_relat_1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_208, plain, (r2_hidden(X1,k10_ordinal2(X1,X2))|~r2_hidden(k1_xboole_0,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_137, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_136, plain, (v3_ordinal1(esk10_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_146, plain, (v3_ordinal1(k10_ordinal2(X1,esk12_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_164, plain, (r1_ordinal1(X1,esk15_0)|r1_ordinal1(esk15_0,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_150, plain, (~v1_xboole_0(X2)|~r2_hidden(X3,X1)|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_149, plain, (~v1_xboole_0(X2)|~r2_hidden(X3,X1)|~r2_hidden(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_168, negated_conjecture, (r1_ordinal1(X1,esk2_0)|r1_ordinal1(esk2_0,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_169, negated_conjecture, (r1_ordinal1(X1,esk1_0)|r1_ordinal1(esk1_0,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_202, plain, (r2_hidden(X1,X2)|X1!=k10_ordinal2(X4,X5)|X2!=k10_xtuple_0(esk4_2(X3,X4))|~r2_hidden(X5,k9_xtuple_0(esk4_2(X3,X4)))|~v1_funct_1(esk4_2(X3,X4))|~v1_relat_1(esk4_2(X3,X4))|~r2_hidden(X5,X3)|~v3_ordinal1(X4)|~v3_ordinal1(X3)|~v3_ordinal1(X5))).#trainneg
cnf(i_0_190, plain, (r2_hidden(X1,X2)|r2_hidden(X2,k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_211, plain, (r1_ordinal1(X2,k10_ordinal2(X1,X3))|~r1_ordinal1(esk5_2(X1,X2),X3)|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~v3_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_213, plain, (r1_ordinal1(k10_ordinal2(X1,esk5_2(X2,X3)),X3)|~r1_ordinal1(X1,X2)|~v3_ordinal1(esk5_2(X2,X3))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(X2,X3)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_227, plain, (r1_ordinal1(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))),X3)|~r2_hidden(k10_ordinal2(X1,X2),X3)|~v3_ordinal1(X3)|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_205, plain, (r2_hidden(k10_ordinal2(X1,X2),X3)|~r2_hidden(X2,esk5_2(X1,X3))|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X1,X3))|~v3_ordinal1(X2)|~r1_ordinal1(X1,X3)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_193, plain, (r2_hidden(k1_xboole_0,X1)|~v1_ordinal1(k1_xboole_0)|~r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_176, plain, (~r2_hidden(k3_ordinal2(X2),X1)|~r2_hidden(X1,X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_231, plain, (k2_xboole_0(X1,k1_tarski(X1))=k10_ordinal2(X1,k2_xboole_0(k1_xboole_0,k1_tarski(k1_xboole_0)))|~v3_ordinal1(k1_xboole_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_192, plain, (v3_ordinal1(k1_funct_1(esk4_2(X1,X2),X3))|~v1_funct_1(esk4_2(X1,X2))|~v1_relat_1(esk4_2(X1,X2))|~v5_ordinal1(esk4_2(X1,X2))|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_132, plain, (r1_tarski(esk10_1(k1_zfmisc_1(X1)),X1))).#trainneg
cnf(i_0_142, plain, (v4_ordinal1(X1)|~v1_xboole_0(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_238, plain, (~v1_xboole_0(k9_xtuple_0(X1))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_133, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_165, plain, (r1_ordinal1(X1,esk14_0)|r1_ordinal1(esk14_0,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_496, negated_conjecture, (~v1_xboole_0(k10_ordinal2(esk2_0,esk1_0)))).#trainneg
cnf(i_0_396, negated_conjecture, (r1_ordinal1(esk2_0,esk2_0))).#trainneg
cnf(i_0_413, negated_conjecture, (r1_ordinal1(esk1_0,esk1_0))).#trainneg
cnf(i_0_144, plain, (v1_xboole_0(k1_zfmisc_1(X1))|r2_hidden(X2,k1_zfmisc_1(X1))|~r1_tarski(X2,X1))).#trainneg
cnf(i_0_259, plain, (v4_ordinal1(k3_ordinal2(X1))|~v3_ordinal1(k2_xboole_0(esk6_1(k3_ordinal2(X1)),k1_tarski(esk6_1(k3_ordinal2(X1)))))|~r2_hidden(esk6_1(k3_ordinal2(X1)),X1)|~v4_ordinal1(X1)|~v3_ordinal1(esk6_1(k3_ordinal2(X1)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_242, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1))).#trainneg
cnf(i_0_251, plain, (m1_subset_1(esk6_1(X1),X2)|v4_ordinal1(X1)|~r1_tarski(X1,X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_253, plain, (m1_subset_1(k2_xboole_0(X1,k1_tarski(X1)),X2)|~r1_tarski(X3,X2)|~r2_hidden(X1,X3)|~v4_ordinal1(X3)|~v3_ordinal1(X1)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_255, plain, (m1_subset_1(X1,X2)|~r1_tarski(k3_ordinal2(X3),X2)|~r2_hidden(X1,X3)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_291, plain, (~r2_hidden(X1,esk7_3(esk4_2(X1,X2),X3,X4))|X3!=k10_xtuple_0(esk4_2(X1,X2))|~r2_hidden(X4,X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_250, plain, (m1_subset_1(esk7_3(X1,X2,X3),X4)|~r1_tarski(k9_xtuple_0(X1),X4)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_469, plain, (v4_ordinal1(X1)|~r2_hidden(k10_ordinal2(esk6_1(X1),k1_tarski(k1_xboole_0)),X1)|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(X1)))).#trainneg
cnf(i_0_420, plain, (r2_hidden(X1,X2)|X2!=k10_xtuple_0(esk4_2(X3,X4))|X1!=k10_ordinal2(X4,X5)|~r2_hidden(X5,X3)|~v3_ordinal1(X4)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_358, plain, (r1_ordinal1(esk15_0,esk15_0))).#trainneg
cnf(i_0_504, plain, (r1_ordinal1(esk14_0,esk14_0))).#trainneg
cnf(i_0_292, plain, (~v1_xboole_0(X1)|X3!=k10_xtuple_0(esk4_2(X1,X2))|~r2_hidden(X4,X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_147, plain, (v3_ordinal1(k10_ordinal2(X1,X2))|~v7_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_359, plain, (r1_ordinal1(esk15_0,esk14_0)|r1_ordinal1(esk14_0,esk15_0))).#trainneg
cnf(i_0_277, plain, (v7_ordinal1(k1_tarski(k1_xboole_0))|~v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_354, plain, (v3_ordinal1(k10_ordinal2(esk12_0,esk12_0)))).#trainneg
cnf(i_0_397, negated_conjecture, (r1_ordinal1(esk2_0,esk1_0)|r1_ordinal1(esk1_0,esk2_0))).#trainneg
cnf(i_0_160, plain, (v3_ordinal1(k9_xtuple_0(esk4_2(X1,X2)))|~v1_funct_1(esk4_2(X1,X2))|~v1_relat_1(esk4_2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_254, plain, (m1_subset_1(k10_ordinal2(X1,X2),X3)|~r1_tarski(k10_ordinal2(X1,X4),X3)|~r2_hidden(X2,X4)|~v3_ordinal1(X1)|~v3_ordinal1(X4)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_433, plain, (r2_hidden(X2,X1)|~r2_hidden(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_468, plain, (v4_ordinal1(k3_ordinal2(X1))|~r2_hidden(k10_ordinal2(esk6_1(k3_ordinal2(X1)),k1_tarski(k1_xboole_0)),X1)|~v3_ordinal1(k10_ordinal2(esk6_1(k3_ordinal2(X1)),k1_tarski(k1_xboole_0)))|~v3_ordinal1(esk6_1(k3_ordinal2(X1))))).#trainneg
cnf(i_0_466, plain, (r2_hidden(X2,k3_ordinal2(X1))|~r2_hidden(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_340, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(X1,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v4_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_333, plain, (r1_ordinal1(X4,X1)|~r1_ordinal1(esk7_3(esk4_2(X2,X1),X3,X4),k1_xboole_0)|~v3_ordinal1(X1)|~v3_ordinal1(esk7_3(esk4_2(X2,X1),X3,X4))|X3!=k10_xtuple_0(esk4_2(X2,X1))|~r2_hidden(X4,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_442, plain, (r1_ordinal1(X1,X4)|~v3_ordinal1(X1)|~v3_ordinal1(esk7_3(esk4_2(X2,X1),X3,X4))|X3!=k10_xtuple_0(esk4_2(X2,X1))|~r2_hidden(X4,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_289, plain, (m1_subset_1(esk7_3(esk4_2(X1,X2),X3,X4),X5)|~r1_tarski(X1,X5)|X3!=k10_xtuple_0(esk4_2(X1,X2))|~r2_hidden(X4,X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_473, plain, (v1_xboole_0(X1)|~v1_xboole_0(k10_ordinal2(X1,k1_tarski(k1_xboole_0)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_483, plain, (v3_ordinal1(k10_ordinal2(X1,k1_tarski(k1_xboole_0)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_360, plain, (r1_ordinal1(esk15_0,esk13_0)|r1_ordinal1(esk13_0,esk15_0))).#trainneg
cnf(i_0_375, plain, (~v1_xboole_0(X1)|~r1_tarski(X2,X1)|X4!=k10_xtuple_0(esk4_2(X2,X3))|~r2_hidden(X5,X4)|~v3_ordinal1(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_361, plain, (r1_ordinal1(esk15_0,esk12_0)|r1_ordinal1(esk12_0,esk15_0))).#trainneg
cnf(i_0_366, negated_conjecture, (r1_ordinal1(esk15_0,esk2_0)|r1_ordinal1(esk2_0,esk15_0))).#trainneg
cnf(i_0_367, negated_conjecture, (r1_ordinal1(esk15_0,esk1_0)|r1_ordinal1(esk1_0,esk15_0))).#trainneg
cnf(i_0_403, negated_conjecture, (r1_ordinal1(esk2_0,k1_xboole_0)|r1_ordinal1(k1_xboole_0,esk2_0))).#trainneg
cnf(i_0_162, plain, (v3_ordinal1(k10_ordinal2(X1,k10_ordinal2(X2,X3)))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_166, plain, (r1_ordinal1(X1,esk13_0)|r1_ordinal1(esk13_0,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_431, plain, (m1_subset_1(X1,X2)|r2_hidden(X3,X1)|~r1_tarski(k2_xboole_0(X3,k1_tarski(X3)),X2)|~v3_ordinal1(X1)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_459, plain, (r2_hidden(k10_ordinal2(X1,esk6_1(esk5_2(X1,X2))),X2)|v4_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_246, plain, (r2_hidden(k10_ordinal2(X1,X2),X3)|~r2_hidden(k10_ordinal2(X1,X4),X3)|~v3_ordinal1(X3)|~v3_ordinal1(k10_ordinal2(X1,X4))|~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(X2,X4)|~v3_ordinal1(X1)|~v3_ordinal1(X4)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_558, plain, (X1=X2|r2_hidden(X1,X2)|r2_hidden(X2,X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_630, plain, (m1_subset_1(k10_ordinal2(X1,X2),k10_ordinal2(X1,X3))|~r2_hidden(X2,X3)|~v3_ordinal1(X1)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_541, plain, (v4_ordinal1(k3_ordinal2(X1))|~r2_hidden(esk6_1(k3_ordinal2(X1)),X1)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k3_ordinal2(X1))))).#trainneg
cnf(i_0_724, plain, (r2_hidden(X1,k10_ordinal2(X1,X2))|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~r2_hidden(X3,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_697, plain, (r1_ordinal1(esk13_0,esk13_0))).#trainneg
cnf(i_0_776, plain, (r2_hidden(k1_xboole_0,k1_tarski(k1_xboole_0)))).#trainneg
cnf(i_0_373, plain, (r1_ordinal1(esk15_0,k1_xboole_0)|r1_ordinal1(k1_xboole_0,esk15_0))).#trainneg
cnf(i_0_167, plain, (r1_ordinal1(X1,esk6_1(X2))|r1_ordinal1(esk6_1(X2),X1)|v4_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_606, plain, (esk15_0=esk14_0|r2_hidden(esk15_0,esk14_0)|r1_ordinal1(esk14_0,esk15_0))).#trainneg
cnf(i_0_660, plain, (v4_ordinal1(k1_tarski(k1_xboole_0))|~r2_hidden(k1_xboole_0,esk6_1(k1_tarski(k1_xboole_0)))|~v4_ordinal1(esk6_1(k1_tarski(k1_xboole_0))))).#trainneg
cnf(i_0_608, plain, (v3_ordinal1(k10_ordinal2(k1_tarski(k1_xboole_0),esk12_0))|~v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_634, plain, (r2_hidden(k1_xboole_0,X1)|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_419, negated_conjecture, (r1_ordinal1(esk1_0,k1_xboole_0)|r1_ordinal1(k1_xboole_0,esk1_0))).#trainneg
cnf(i_0_170, plain, (r1_ordinal1(X1,k3_ordinal2(X2))|r1_ordinal1(k3_ordinal2(X2),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_171, plain, (r1_ordinal1(X1,k2_xboole_0(X2,k1_tarski(X2)))|r1_ordinal1(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_548, plain, (v1_xboole_0(X1)|r2_hidden(esk6_1(X2),X1)|v4_ordinal1(X2)|~r1_tarski(X2,X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_766, plain, (m1_subset_1(X1,X2)|~r1_tarski(k2_xboole_0(X1,k1_tarski(X1)),X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_576, plain, (r2_hidden(esk6_1(X1),X2)|r2_hidden(X2,X1)|v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_762, plain, (~r2_hidden(k2_xboole_0(k3_ordinal2(X1),k1_tarski(k3_ordinal2(X1))),X1)|~v3_ordinal1(k2_xboole_0(k3_ordinal2(X1),k1_tarski(k3_ordinal2(X1)))))).#trainneg
cnf(i_0_722, plain, (r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))|r2_hidden(X2,k10_ordinal2(X1,X3))|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~v3_ordinal1(k10_ordinal2(X1,X3))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_572, plain, (m1_subset_1(X1,k3_ordinal2(X2))|~r2_hidden(X1,X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_357, plain, (v3_ordinal1(k10_ordinal2(X1,esk12_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_821, plain, (r2_hidden(k1_xboole_0,esk14_0))).#trainneg
cnf(i_0_288, plain, (~v1_xboole_0(X1)|X3!=k10_xtuple_0(esk4_2(esk10_1(k1_zfmisc_1(X1)),X2))|~r2_hidden(X4,X3)|~v3_ordinal1(X2)|~v3_ordinal1(esk10_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_172, plain, (r1_ordinal1(X1,k2_xboole_0(X2,X3))|r1_ordinal1(k2_xboole_0(X2,X3),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_617, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk2_0,esk1_0)|r1_ordinal1(esk1_0,esk2_0))).#trainneg
cnf(i_0_856, plain, (~r2_hidden(k1_tarski(k1_xboole_0),k1_xboole_0))).#trainneg
cnf(i_0_684, negated_conjecture, (esk2_0=k1_xboole_0|r2_hidden(esk2_0,k1_xboole_0)|r1_ordinal1(k1_xboole_0,esk2_0))).#trainneg
cnf(i_0_556, plain, (k1_xboole_0=X1|r2_hidden(k1_xboole_0,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_173, plain, (r1_ordinal1(X1,k4_ordinal2(X2))|r1_ordinal1(k4_ordinal2(X2),X1)|~v3_ordinal1(X1)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v5_ordinal1(X2))).#trainneg
cnf(i_0_174, plain, (r1_ordinal1(X1,k10_ordinal2(X2,X3))|r1_ordinal1(k10_ordinal2(X2,X3),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_837, plain, (r2_hidden(X1,k10_ordinal2(X1,X2))|v1_xboole_0(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_977, plain, (r2_hidden(X1,k10_ordinal2(X1,X2))|r2_hidden(X2,X3)|v4_ordinal1(X3)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_975, plain, (r2_hidden(esk6_1(X1),X2)|r2_hidden(X2,X3)|r2_hidden(X3,X1)|v4_ordinal1(X1)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_840, plain, (r2_hidden(X1,k10_ordinal2(X1,X2))|v4_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_966, plain, (m1_subset_1(X1,k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1032, plain, (r2_hidden(k2_xboole_0(esk7_3(esk4_2(X1,X2),X3,X4),k1_tarski(esk7_3(esk4_2(X1,X2),X3,X4))),X5)|r2_hidden(X5,X1)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(X5)|X3!=k10_xtuple_0(esk4_2(X1,X2))|~r2_hidden(X4,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_961, plain, (v4_ordinal1(k3_ordinal2(X1))|v1_xboole_0(X1)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~r1_tarski(k3_ordinal2(X1),X1))).#trainneg
cnf(i_0_602, plain, (v3_ordinal1(k10_ordinal2(esk12_0,X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_777, plain, (r2_hidden(X1,k10_ordinal2(X1,k1_tarski(k1_xboole_0)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_822, plain, (r2_hidden(k1_xboole_0,esk12_0))).#trainneg
cnf(i_0_178, plain, (r2_hidden(X1,k4_ordinal2(X2))|~r2_hidden(X1,k10_xtuple_0(X2))|~v3_ordinal1(X1)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v5_ordinal1(X2))).#trainneg
cnf(i_0_675, plain, (esk13_0=esk15_0|r2_hidden(esk13_0,esk15_0)|r1_ordinal1(esk15_0,esk13_0))).#trainneg
cnf(i_0_1214, negated_conjecture, (r2_hidden(esk2_0,k10_ordinal2(esk2_0,X1))|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1215, negated_conjecture, (r2_hidden(esk1_0,k10_ordinal2(esk1_0,X1))|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_179, plain, (v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_180, plain, (v3_ordinal1(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1107, plain, (r2_hidden(X1,k10_ordinal2(X1,X2))|k1_xboole_0=X2|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_1146, plain, (r2_hidden(X1,X2)|r2_hidden(X2,k10_ordinal2(X1,X3))|v1_xboole_0(X3)|~v3_ordinal1(k10_ordinal2(X1,X3))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_1144, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X3)|~r1_tarski(k10_ordinal2(X1,X3),X2)|~v3_ordinal1(X1)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_982, plain, (r2_hidden(k1_xboole_0,X1)|r2_hidden(X1,X2)|v4_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_1155, plain, (v1_xboole_0(esk5_2(X1,X2))|r2_hidden(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_1109, plain, (k1_xboole_0=X1|~r2_hidden(X1,k1_xboole_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_829, negated_conjecture, (r2_hidden(k1_xboole_0,k10_ordinal2(esk2_0,esk1_0))|~v3_ordinal1(k10_ordinal2(esk2_0,esk1_0)))).#trainneg
cnf(i_0_1035, plain, (r2_hidden(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))),X2)|r2_hidden(X2,X3)|r2_hidden(X3,X1)|v4_ordinal1(X1)|~v4_ordinal1(X3)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1002, plain, (r2_hidden(X1,k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))))|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1233, plain, (m1_subset_1(k1_xboole_0,k1_tarski(k1_xboole_0)))).#trainneg
cnf(i_0_181, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk10_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_678, plain, (esk12_0=esk15_0|r2_hidden(esk12_0,esk15_0)|r1_ordinal1(esk15_0,esk12_0))).#trainneg
cnf(i_0_1063, plain, (~r2_hidden(esk14_0,k1_xboole_0))).#trainneg
cnf(i_0_1301, plain, (~r2_hidden(esk12_0,k1_xboole_0))).#trainneg
cnf(i_0_1229, plain, (r2_hidden(k1_xboole_0,k10_ordinal2(k1_xboole_0,X1))|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1092, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk1_0,esk2_0)|r2_hidden(esk2_0,esk1_0))).#trainneg
cnf(i_0_182, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X3)|~r2_hidden(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_186, plain, (r1_ordinal1(X1,esk5_2(X2,X3))|r1_ordinal1(esk5_2(X2,X3),X1)|~v3_ordinal1(X1)|~r1_ordinal1(X2,X3)|~v3_ordinal1(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_811, plain, (v1_xboole_0(k10_ordinal2(X1,X2))|r2_hidden(k10_ordinal2(X1,X3),k10_ordinal2(X1,X2))|~r2_hidden(X3,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_1465, plain, (esk5_2(X1,X2)=k1_xboole_0|r2_hidden(X1,X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_1395, plain, (k1_xboole_0=X2|~r2_hidden(k10_ordinal2(X1,X2),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_1414, plain, (r2_hidden(X1,k10_ordinal2(X2,X3))|v1_xboole_0(X3)|r2_hidden(X2,k10_ordinal2(X1,X4))|~v3_ordinal1(k10_ordinal2(X2,X3))|~v3_ordinal1(X1)|~v3_ordinal1(X4)|~v3_ordinal1(k10_ordinal2(X1,X4))|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_1409, plain, (v1_xboole_0(X2)|r2_hidden(X1,k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2))))|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2))))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_1505, plain, (~r2_hidden(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))),X1)|~v3_ordinal1(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1494, plain, (~r2_hidden(k2_xboole_0(k2_xboole_0(k3_ordinal2(X1),k1_tarski(k3_ordinal2(X1))),k1_tarski(k2_xboole_0(k3_ordinal2(X1),k1_tarski(k3_ordinal2(X1))))),X1)|~v3_ordinal1(k2_xboole_0(k2_xboole_0(k3_ordinal2(X1),k1_tarski(k3_ordinal2(X1))),k1_tarski(k2_xboole_0(k3_ordinal2(X1),k1_tarski(k3_ordinal2(X1)))))))).#trainneg
cnf(i_0_851, plain, (m1_subset_1(k1_xboole_0,X1)|~r1_tarski(k1_tarski(k1_xboole_0),X1))).#trainneg
cnf(i_0_389, negated_conjecture, (r1_ordinal1(esk2_0,esk14_0)|r1_ordinal1(esk14_0,esk2_0))).#trainneg
cnf(i_0_997, plain, (~v3_ordinal1(k2_xboole_0(k3_ordinal2(X1),k1_tarski(k3_ordinal2(X1))))|~r2_hidden(k3_ordinal2(X1),X1)|~v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_189, plain, (v7_ordinal1(X2)|~v7_ordinal1(esk5_2(X1,X2))|~v7_ordinal1(X1)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_680, negated_conjecture, (esk2_0=esk15_0|r2_hidden(esk2_0,esk15_0)|r1_ordinal1(esk15_0,esk2_0))).#trainneg
cnf(i_0_1445, negated_conjecture, (r2_hidden(k1_xboole_0,esk2_0)|r2_hidden(esk2_0,X1)|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1446, negated_conjecture, (r2_hidden(k1_xboole_0,esk1_0)|r2_hidden(esk1_0,X1)|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_199, plain, (k3_ordinal2(X1)=k10_xtuple_0(X2)|esk8_2(X2,k3_ordinal2(X1))!=k1_funct_1(X2,X3)|~r2_hidden(X3,k9_xtuple_0(X2))|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(esk8_2(X2,k3_ordinal2(X1)),X1)|~v3_ordinal1(esk8_2(X2,k3_ordinal2(X1))))).#trainneg
cnf(i_0_200, plain, (r2_hidden(k1_funct_1(X1,X2),X3)|X3!=k10_xtuple_0(X1)|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1496, plain, (r2_hidden(X1,k2_xboole_0(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))),k1_tarski(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))))))|~v3_ordinal1(k2_xboole_0(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))),k1_tarski(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))))))|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_1497, plain, (r2_hidden(X1,k2_xboole_0(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2))),k1_tarski(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2))))))|~v3_ordinal1(k2_xboole_0(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2))),k1_tarski(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2))))))|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(k10_ordinal2(X1,X2)))).#trainneg
cnf(i_0_1406, plain, (v4_ordinal1(k10_ordinal2(X1,X2))|v1_xboole_0(X2)|r2_hidden(X1,k2_xboole_0(esk6_1(k10_ordinal2(X1,X2)),k1_tarski(esk6_1(k10_ordinal2(X1,X2)))))|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(k2_xboole_0(esk6_1(k10_ordinal2(X1,X2)),k1_tarski(esk6_1(k10_ordinal2(X1,X2)))))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_1410, plain, (r2_hidden(X1,k3_ordinal2(k10_ordinal2(X2,X3)))|v1_xboole_0(X3)|r2_hidden(X2,k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_1634, plain, (esk6_1(esk5_2(X1,X1))=k1_xboole_0|v4_ordinal1(esk5_2(X1,X1))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(esk5_2(X1,X1)))|~r1_ordinal1(X1,X1))).#trainneg
cnf(i_0_1692, plain, (~r2_hidden(k2_xboole_0(k10_ordinal2(X1,k1_tarski(k1_xboole_0)),k1_tarski(k10_ordinal2(X1,k1_tarski(k1_xboole_0)))),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1696, plain, (~r2_hidden(k10_ordinal2(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k1_xboole_0)),X1)|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_1234, plain, (m1_subset_1(X1,k10_ordinal2(X1,k1_tarski(k1_xboole_0)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_390, negated_conjecture, (r1_ordinal1(esk2_0,esk13_0)|r1_ordinal1(esk13_0,esk2_0))).#trainneg
cnf(i_0_1722, plain, (~r2_hidden(k3_ordinal2(X1),X1)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(k1_tarski(k3_ordinal2(X1))))).#trainneg
cnf(i_0_206, plain, (r2_hidden(k10_ordinal2(X1,X2),X1)|~r2_hidden(X2,k1_xboole_0)|~v3_ordinal1(X1)|~v3_ordinal1(k1_xboole_0)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_682, negated_conjecture, (esk15_0=esk1_0|r2_hidden(esk15_0,esk1_0)|r1_ordinal1(esk1_0,esk15_0))).#trainneg
cnf(i_0_1094, negated_conjecture, (esk2_0=k1_xboole_0|r2_hidden(k1_xboole_0,esk2_0)|r2_hidden(esk2_0,k1_xboole_0))).#trainneg
cnf(i_0_201, plain, (k10_ordinal2(X1,X2)=k4_ordinal2(esk4_2(X3,X4))|X2=k1_xboole_0|k10_ordinal2(X4,esk3_3(X2,X1,esk4_2(X3,X4)))!=k10_ordinal2(X1,esk3_3(X2,X1,esk4_2(X3,X4)))|k9_xtuple_0(esk4_2(X3,X4))!=X2|~v1_ordinal2(esk4_2(X3,X4))|~v1_funct_1(esk4_2(X3,X4))|~v1_relat_1(esk4_2(X3,X4))|~v5_ordinal1(esk4_2(X3,X4))|~v4_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(esk3_3(X2,X1,esk4_2(X3,X4)),X3)|~v3_ordinal1(X4)|~v3_ordinal1(X3)|~v3_ordinal1(esk3_3(X2,X1,esk4_2(X3,X4))))).#trainneg
cnf(i_0_214, plain, (r1_ordinal1(k10_ordinal2(X1,k1_xboole_0),X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(k1_xboole_0)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1621, plain, (k10_ordinal2(X1,k1_tarski(k1_xboole_0))=k2_xboole_0(X2,k1_tarski(X2))|r2_hidden(X1,X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_1628, plain, (esk5_2(X1,X2)=k1_xboole_0|~r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_1415, plain, (r2_hidden(k10_ordinal2(X1,X2),k10_ordinal2(X3,X4))|v1_xboole_0(X4)|r2_hidden(X3,k10_ordinal2(X1,X5))|~r1_ordinal1(X2,X5)|~v3_ordinal1(k10_ordinal2(X3,X4))|~v3_ordinal1(X1)|~v3_ordinal1(X5)|~v3_ordinal1(X2)|~v3_ordinal1(k10_ordinal2(X1,X5))|~v3_ordinal1(X3)|~v3_ordinal1(X4))).#trainneg
cnf(i_0_1418, plain, (r2_hidden(X1,X2)|r2_hidden(X2,k10_ordinal2(X3,X4))|v1_xboole_0(X4)|r2_hidden(X3,X1)|~v3_ordinal1(k10_ordinal2(X3,X4))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~v3_ordinal1(X4))).#trainneg
cnf(i_0_1036, plain, (r2_hidden(k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))),X3)|r2_hidden(X3,X4)|r1_ordinal1(k10_ordinal2(X1,X4),X2)|~v4_ordinal1(X4)|~v3_ordinal1(X4)|~v3_ordinal1(X3)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_1393, plain, (r2_hidden(X1,k10_ordinal2(X1,k10_ordinal2(X2,X3)))|k1_xboole_0=X3|~v3_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(X2,X3))|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_391, negated_conjecture, (r1_ordinal1(esk2_0,esk12_0)|r1_ordinal1(esk12_0,esk2_0))).#trainneg
cnf(i_0_1721, plain, (~r2_hidden(k3_ordinal2(X1),X1)|~v4_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(k3_ordinal2(X1),k1_tarski(k1_xboole_0)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_215, plain, (r1_ordinal1(X2,k10_ordinal2(X3,esk5_2(X1,X2)))|~r1_ordinal1(X1,X3)|~v3_ordinal1(esk5_2(X1,X2))|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_860, plain, (esk15_0=k1_xboole_0|r2_hidden(esk15_0,k1_xboole_0)|r1_ordinal1(k1_xboole_0,esk15_0))).#trainneg
cnf(i_0_1691, plain, (~r2_hidden(k2_xboole_0(k1_tarski(k1_xboole_0),k1_tarski(k1_tarski(k1_xboole_0))),k1_xboole_0))).#trainneg
cnf(i_0_1553, negated_conjecture, (r2_hidden(k1_xboole_0,esk1_0)|esk2_0=esk1_0|r2_hidden(esk1_0,esk2_0))).#trainneg
cnf(i_0_797, negated_conjecture, (X1=esk2_0|r2_hidden(esk2_0,X1)|r2_hidden(X1,esk2_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_216, plain, (r1_ordinal1(X1,k10_ordinal2(X2,k1_xboole_0))|~r1_ordinal1(X1,X2)|~v3_ordinal1(k1_xboole_0)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_217, plain, (r1_ordinal1(X1,esk3_3(X2,X3,X4))|r1_ordinal1(esk3_3(X2,X3,X4),X1)|k10_ordinal2(X3,X2)=k4_ordinal2(X4)|X2=k1_xboole_0|~v3_ordinal1(X1)|k9_xtuple_0(X4)!=X2|~v1_ordinal2(X4)|~v1_funct_1(X4)|~v1_relat_1(X4)|~v5_ordinal1(X4)|~v4_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2091, plain, (r1_ordinal1(k10_ordinal2(X1,X2),X3)|r2_hidden(k2_xboole_0(esk5_2(X1,X3),k1_tarski(esk5_2(X1,X3))),X4)|r2_hidden(X4,X2)|r2_hidden(X3,X1)|~v4_ordinal1(X2)|~v3_ordinal1(X2)|~v3_ordinal1(X4)|~v3_ordinal1(X1)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_1996, plain, (esk5_2(X1,X2)=k1_xboole_0|~r1_ordinal1(X1,X2)|~r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2030, plain, (v1_xboole_0(X1)|r2_hidden(k10_ordinal2(X2,X3),k10_ordinal2(X4,X1))|r2_hidden(X4,k10_ordinal2(X2,X5))|r2_hidden(X5,X3)|~v3_ordinal1(X2)|~v3_ordinal1(X5)|~v3_ordinal1(X3)|~v3_ordinal1(X4)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1892, plain, (esk6_1(esk5_2(X1,X1))=k1_xboole_0|v4_ordinal1(esk5_2(X1,X1))|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X1,X1)))).#trainneg
cnf(i_0_1949, negated_conjecture, (esk2_0=k1_xboole_0|r2_hidden(k1_xboole_0,esk2_0))).#trainneg
cnf(i_0_1041, plain, (r2_hidden(k2_xboole_0(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))),k1_tarski(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))))),X3)|r2_hidden(X3,X4)|~v4_ordinal1(X4)|~v3_ordinal1(X4)|~v3_ordinal1(X3)|~r2_hidden(k10_ordinal2(X1,X2),X4)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_555, plain, (k10_ordinal2(X1,X2)=X1|r2_hidden(k10_ordinal2(X1,X2),X1)|~v3_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(X2,k1_xboole_0)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_788, plain, (r2_hidden(k10_ordinal2(X1,X2),X3)|~r1_ordinal1(X2,k2_xboole_0(X4,k1_tarski(X4)))|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X4,k1_tarski(X4)))|~v3_ordinal1(X2)|~r2_hidden(k10_ordinal2(X1,X4),X3)|~v4_ordinal1(X3)|~v3_ordinal1(X4))).#trainneg
cnf(i_0_812, plain, (v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(k10_ordinal2(X1,X3))|~r2_hidden(X2,X3)|~v3_ordinal1(X1)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_1854, plain, (v1_xboole_0(X2)|v4_ordinal1(k10_ordinal2(X1,X2))|~r2_hidden(esk6_1(k10_ordinal2(X1,X2)),X1)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(esk6_1(k10_ordinal2(X1,X2)),k1_tarski(esk6_1(k10_ordinal2(X1,X2)))))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_1913, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k1_xboole_0)|~v3_ordinal1(esk10_1(k1_zfmisc_1(X1)))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_219, plain, (X1=k10_xtuple_0(esk4_2(X2,X3))|r2_hidden(esk9_2(esk4_2(X2,X3),X1),k9_xtuple_0(esk4_2(X2,X3)))|r2_hidden(esk8_2(esk4_2(X2,X3),X1),X1)|~v1_relat_1(esk4_2(X2,X3))|~v3_ordinal1(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_884, plain, (esk15_0=esk14_0|r2_hidden(esk14_0,esk15_0)|r2_hidden(esk15_0,esk14_0))).#trainneg
cnf(i_0_1901, plain, (~r2_hidden(k10_ordinal2(k1_tarski(k1_xboole_0),k1_tarski(k1_xboole_0)),k1_xboole_0))).#trainneg
cnf(i_0_2148, negated_conjecture, (r2_hidden(k1_xboole_0,esk2_0)|esk2_0=esk1_0|r2_hidden(k1_xboole_0,esk1_0))).#trainneg
cnf(i_0_220, plain, (k1_funct_1(esk4_2(X1,X2),esk9_2(esk4_2(X1,X2),X3))=esk8_2(esk4_2(X1,X2),X3)|X3=k10_xtuple_0(esk4_2(X1,X2))|r2_hidden(esk8_2(esk4_2(X1,X2),X3),X3)|~v1_relat_1(esk4_2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_221, plain, (v1_xboole_0(k10_ordinal2(X1,X2))|~v1_xboole_0(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_2209, plain, (r1_ordinal1(k10_ordinal2(X1,X2),X3)|~r1_ordinal1(X2,k1_xboole_0)|~r1_ordinal1(X1,X3)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~r2_hidden(X3,X1))).#trainneg
cnf(i_0_2199, plain, (r2_hidden(X1,k10_ordinal2(X2,X3))|~r1_ordinal1(X2,X1)|~r2_hidden(k1_xboole_0,X3)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_2204, plain, (k10_ordinal2(X1,k1_tarski(k1_xboole_0))=k2_xboole_0(X2,k1_tarski(X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_2295, plain, (k10_ordinal2(X1,X2)=X1|~r2_hidden(X1,k10_ordinal2(X1,X2))|~r1_ordinal1(X2,k1_xboole_0)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2291, plain, (r2_hidden(k10_ordinal2(X1,X2),X3)|r2_hidden(X3,X1)|k10_ordinal2(X1,X2)=X1|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~r1_ordinal1(X2,k1_xboole_0)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2366, plain, (v1_xboole_0(X1)|v4_ordinal1(k10_ordinal2(X2,X1))|~r2_hidden(esk6_1(k10_ordinal2(X2,X1)),X2)|~v4_ordinal1(X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k10_ordinal2(X2,X1))))).#trainneg
cnf(i_0_2202, plain, (k10_ordinal2(X1,k1_xboole_0)=X2|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_2200, plain, (r1_ordinal1(X1,k10_ordinal2(X2,X3))|~r1_ordinal1(k1_xboole_0,X3)|~r1_ordinal1(X2,X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_2109, plain, (esk5_2(X1,X2)=k1_xboole_0|r2_hidden(X3,k10_ordinal2(X3,X2))|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_1601, plain, (v1_xboole_0(k10_ordinal2(X1,X3))|~r2_hidden(k10_ordinal2(X1,X3),k10_ordinal2(X1,X2))|~r2_hidden(X2,X3)|~v3_ordinal1(X1)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_2339, plain, (v3_ordinal1(k10_ordinal2(X1,k1_xboole_0))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_536, plain, (v1_xboole_0(k1_zfmisc_1(X1))|r2_hidden(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1747, plain, (X1!=k10_xtuple_0(esk4_2(k1_funct_1(X2,X3),X4))|~r2_hidden(X5,X1)|~v3_ordinal1(X4)|~v3_ordinal1(k1_funct_1(X2,X3))|esk7_3(esk4_2(k1_funct_1(X2,X3),X4),X1,X5)!=k10_xtuple_0(X2)|~r2_hidden(X3,k9_xtuple_0(X2))|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_222, plain, (v3_ordinal1(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(k1_tarski(k10_ordinal2(X1,X2)))|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1336, plain, (esk13_0=esk15_0|r2_hidden(esk15_0,esk13_0)|r2_hidden(esk13_0,esk15_0))).#trainneg
cnf(i_0_1483, negated_conjecture, (~r2_hidden(k10_ordinal2(esk2_0,esk1_0),k1_xboole_0)|~v3_ordinal1(k10_ordinal2(esk2_0,esk1_0)))).#trainneg
cnf(i_0_224, plain, (v7_ordinal1(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))))|~v7_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_226, plain, (v3_ordinal1(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_2655, plain, (r1_ordinal1(X1,X2)|~r1_ordinal1(X2,X1)|~r2_hidden(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2653, plain, (r1_ordinal1(X1,X3)|~r1_ordinal1(X2,X1)|~r2_hidden(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(esk5_2(X2,X3))|~r1_ordinal1(X2,X3)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_2600, plain, (v1_xboole_0(X1)|v4_ordinal1(k10_ordinal2(X2,X1))|v1_xboole_0(X2)|~v4_ordinal1(X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_tarski(k10_ordinal2(X2,X1),X2)|~v3_ordinal1(k10_ordinal2(X2,X1)))).#trainneg
cnf(i_0_2552, plain, (k10_ordinal2(X1,k10_ordinal2(X2,X3))=X1|k1_xboole_0=X3|~r1_ordinal1(k10_ordinal2(X2,X3),k1_xboole_0)|~v3_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(X2,X3))|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_2703, plain, (v1_xboole_0(X1)|~r2_hidden(X1,k10_ordinal2(X1,X2))|~r2_hidden(X2,k1_xboole_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_2697, plain, (v1_xboole_0(k10_ordinal2(X1,X2))|~r2_hidden(k10_ordinal2(X1,X2),X3)|~r2_hidden(esk5_2(X1,X3),X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r1_ordinal1(X1,X3)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_2638, plain, (k10_ordinal2(X1,k1_xboole_0)=k10_ordinal2(X1,X2)|~r2_hidden(k10_ordinal2(X1,X2),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2695, plain, (esk5_2(X1,X2)=k1_xboole_0|r2_hidden(X3,k10_ordinal2(X3,X2))|r2_hidden(X2,X1)|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2592, plain, (v1_xboole_0(esk5_2(X1,X2))|v4_ordinal1(X2)|~r2_hidden(esk6_1(X2),X1)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_282, plain, (r2_hidden(k1_xboole_0,X1)|~r1_ordinal1(k1_tarski(k1_xboole_0),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1219, plain, (r2_hidden(k1_tarski(k1_xboole_0),k10_ordinal2(k1_tarski(k1_xboole_0),X1))|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1765, plain, (~v4_ordinal1(k1_funct_1(X1,X2))|~v3_ordinal1(k1_funct_1(X1,X2))|k1_funct_1(X1,X2)!=k10_xtuple_0(X1)|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1537, plain, (esk12_0=esk15_0|r2_hidden(esk15_0,esk12_0)|r2_hidden(esk12_0,esk15_0))).#trainneg
cnf(i_0_2254, negated_conjecture, (esk2_0=k1_xboole_0|~r2_hidden(esk2_0,k1_xboole_0))).#trainneg
cnf(i_0_2839, plain, (r1_ordinal1(esk14_0,esk15_0)|~r2_hidden(esk14_0,esk15_0))).#trainneg
cnf(i_0_798, negated_conjecture, (X1=esk1_0|r2_hidden(esk1_0,X1)|r2_hidden(X1,esk1_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_245, plain, (r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),X2)|~r2_hidden(X3,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(X1,X3)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_2874, plain, (r1_ordinal1(k10_ordinal2(X1,X2),X3)|~r2_hidden(k10_ordinal2(X1,X2),X3)|~v3_ordinal1(X3)|~r1_ordinal1(X1,X3)|~r2_hidden(X3,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2875, plain, (r1_ordinal1(k10_ordinal2(X1,X2),X1)|~r2_hidden(k10_ordinal2(X1,X2),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2922, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|v4_ordinal1(k10_ordinal2(X1,X2))|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r1_ordinal1(k10_ordinal2(X1,X2),X1)|~v3_ordinal1(k10_ordinal2(X1,X2)))).#trainneg
cnf(i_0_2914, plain, (r1_ordinal1(k10_ordinal2(X1,X2),X3)|~r1_ordinal1(X4,X3)|~r2_hidden(k10_ordinal2(X1,X2),X4)|~v3_ordinal1(X4)|~v3_ordinal1(X3)|~r1_ordinal1(X1,X4)|~r2_hidden(X4,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2915, plain, (r1_ordinal1(k10_ordinal2(X1,X2),X3)|~r1_ordinal1(X1,X3)|~r2_hidden(k10_ordinal2(X1,X2),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_3022, plain, (r2_hidden(k1_xboole_0,X1)|r2_hidden(X1,k1_tarski(k1_xboole_0))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_2593, plain, (v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2469, plain, (v4_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(X1,k2_xboole_0(esk6_1(k10_ordinal2(X1,X2)),k1_tarski(esk6_1(k10_ordinal2(X1,X2)))))|~r2_hidden(k1_xboole_0,X2)|~r2_hidden(k2_xboole_0(esk6_1(k10_ordinal2(X1,X2)),k1_tarski(esk6_1(k10_ordinal2(X1,X2)))),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2951, plain, (v1_xboole_0(k10_ordinal2(X1,X2))|~r1_ordinal1(X1,X3)|~r2_hidden(k10_ordinal2(X1,X2),X3)|~r2_hidden(k1_xboole_0,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~r2_hidden(X3,X1))).#trainneg
cnf(i_0_3281, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(X2,X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v4_ordinal1(X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_2748, plain, (v3_ordinal1(X1)|v1_xboole_0(k1_zfmisc_1(X1))|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_405, negated_conjecture, (r1_ordinal1(esk1_0,esk14_0)|r1_ordinal1(esk14_0,esk1_0))).#trainneg
cnf(i_0_1920, plain, (k10_ordinal2(k1_xboole_0,X1)=k1_xboole_0|~v3_ordinal1(k10_ordinal2(k1_xboole_0,X1))|~r2_hidden(X1,k1_xboole_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1717, negated_conjecture, (esk2_0=esk14_0|r2_hidden(esk2_0,esk14_0)|r1_ordinal1(esk14_0,esk2_0))).#trainneg
cnf(i_0_2843, negated_conjecture, (r1_ordinal1(esk1_0,esk15_0)|~r2_hidden(esk1_0,esk15_0))).#trainneg
cnf(i_0_2844, plain, (r1_ordinal1(k1_xboole_0,esk15_0)|~r2_hidden(k1_xboole_0,esk15_0))).#trainneg
cnf(i_0_249, plain, (r2_hidden(X1,k10_ordinal2(X2,k3_ordinal2(X3)))|~r1_ordinal1(X2,X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(esk5_2(X2,X1),X3)|~v3_ordinal1(esk5_2(X2,X1)))).#trainneg
cnf(i_0_252, plain, (m1_subset_1(esk3_3(X1,X2,X3),X4)|k10_ordinal2(X2,X1)=k4_ordinal2(X3)|X1=k1_xboole_0|~r1_tarski(X1,X4)|k9_xtuple_0(X3)!=X1|~v1_ordinal2(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v5_ordinal1(X3)|~v4_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_3197, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|v4_ordinal1(k10_ordinal2(X2,X1))|~v4_ordinal1(X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(k10_ordinal2(X2,X1),X2))).#trainneg
cnf(i_0_3286, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(X2,X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v4_ordinal1(X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X2,X1)))).#trainneg
cnf(i_0_3186, plain, (v1_xboole_0(X1)|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_ordinal1(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2541, plain, (X2=X1|~r1_ordinal1(esk5_2(X1,X2),k1_xboole_0)|~r2_hidden(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2932, plain, (v1_xboole_0(X1)|~r2_hidden(X1,X2)|~r2_hidden(esk5_2(X1,X2),k1_xboole_0)|~v3_ordinal1(X1)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2677, negated_conjecture, (esk5_2(esk1_0,esk1_0)=k1_xboole_0|r2_hidden(X1,k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_2496, plain, (~v4_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(X1,k10_ordinal2(X1,X2))|~r2_hidden(k1_xboole_0,X2)|~r2_hidden(k10_ordinal2(X1,X2),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_1714, plain, (m1_subset_1(k1_xboole_0,X1)|~r1_ordinal1(k1_tarski(k1_xboole_0),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_2836, plain, (v3_ordinal1(k10_ordinal2(X1,k10_ordinal2(X2,k1_tarski(k1_xboole_0))))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_406, negated_conjecture, (r1_ordinal1(esk1_0,esk13_0)|r1_ordinal1(esk13_0,esk1_0))).#trainneg
cnf(i_0_256, plain, (v4_ordinal1(k4_ordinal2(X1))|~r2_hidden(k2_xboole_0(esk6_1(k4_ordinal2(X1)),k1_tarski(esk6_1(k4_ordinal2(X1)))),k10_xtuple_0(X1))|~v3_ordinal1(k2_xboole_0(esk6_1(k4_ordinal2(X1)),k1_tarski(esk6_1(k4_ordinal2(X1)))))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v5_ordinal1(X1))).#trainneg
cnf(i_0_1728, negated_conjecture, (esk2_0=esk15_0|r2_hidden(esk15_0,esk2_0)|r2_hidden(esk2_0,esk15_0))).#trainneg
cnf(i_0_2847, plain, (r1_ordinal1(esk15_0,esk13_0)|~r2_hidden(esk15_0,esk13_0))).#trainneg
cnf(i_0_2849, negated_conjecture, (r1_ordinal1(esk2_0,esk13_0)|~r2_hidden(esk2_0,esk13_0))).#trainneg
cnf(i_0_1546, negated_conjecture, (r2_hidden(esk2_0,X1)|r2_hidden(X1,esk1_0)|esk2_0=esk1_0|r2_hidden(esk1_0,esk2_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_2141, negated_conjecture, (r2_hidden(esk1_0,X1)|r2_hidden(X1,esk2_0)|esk2_0=esk1_0|r2_hidden(k1_xboole_0,esk1_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_258, plain, (v4_ordinal1(k3_ordinal2(k3_ordinal2(X1)))|~v3_ordinal1(k2_xboole_0(esk6_1(k3_ordinal2(k3_ordinal2(X1))),k1_tarski(esk6_1(k3_ordinal2(k3_ordinal2(X1))))))|~r2_hidden(k2_xboole_0(esk6_1(k3_ordinal2(k3_ordinal2(X1))),k1_tarski(esk6_1(k3_ordinal2(k3_ordinal2(X1))))),X1))).#trainneg
cnf(i_0_3405, plain, (r2_hidden(k10_ordinal2(X1,X2),X3)|k10_ordinal2(X3,X2)=X3|~r1_ordinal1(X1,X3)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(X2,k1_xboole_0))).#trainneg
cnf(i_0_3506, plain, (v1_xboole_0(X1)|v4_ordinal1(k10_ordinal2(X2,X1))|k10_ordinal2(X2,X1)=X2|~v4_ordinal1(X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(X1,k1_xboole_0))).#trainneg
cnf(i_0_3402, plain, (r2_hidden(k10_ordinal2(X1,esk6_1(esk5_2(X2,X3))),X3)|v4_ordinal1(esk5_2(X2,X3))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X3)|~v3_ordinal1(esk6_1(esk5_2(X2,X3)))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(X2,X3))).#trainneg
cnf(i_0_646, plain, (r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))|r2_hidden(X2,k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_3824, plain, (r2_hidden(X2,k10_ordinal2(X3,X4))|~r1_ordinal1(X1,X3)|~r2_hidden(esk5_2(X1,X2),X4)|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~v3_ordinal1(X4)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_3826, plain, (r2_hidden(X1,k10_ordinal2(X2,X3))|~r1_ordinal1(X1,X2)|~r2_hidden(k1_xboole_0,X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_3789, plain, (m1_subset_1(k10_ordinal2(X1,X2),X3)|~r1_tarski(k10_ordinal2(X4,X5),X3)|~r1_ordinal1(X1,X4)|~r2_hidden(X2,X5)|~v3_ordinal1(X4)|~v3_ordinal1(X1)|~v3_ordinal1(X5))).#trainneg
cnf(i_0_3796, plain, (~r2_hidden(k10_ordinal2(X3,X4),k10_ordinal2(X1,X2))|~r1_ordinal1(X1,X3)|~r2_hidden(X2,X4)|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~v3_ordinal1(X4))).#trainneg
cnf(i_0_3356, plain, (k1_zfmisc_1(X1)=k1_xboole_0|v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_481, plain, (v7_ordinal1(k10_ordinal2(X1,k1_tarski(k1_xboole_0)))|~v7_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_407, negated_conjecture, (r1_ordinal1(esk1_0,esk12_0)|r1_ordinal1(esk12_0,esk1_0))).#trainneg
cnf(i_0_260, plain, (v3_ordinal1(esk7_3(X1,X2,X3))|~v3_ordinal1(k9_xtuple_0(X1))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1910, negated_conjecture, (esk13_0=esk2_0|r2_hidden(esk13_0,esk2_0)|r1_ordinal1(esk2_0,esk13_0))).#trainneg
cnf(i_0_2850, plain, (r1_ordinal1(esk15_0,esk12_0)|~r2_hidden(esk15_0,esk12_0))).#trainneg
cnf(i_0_2851, negated_conjecture, (r1_ordinal1(esk2_0,esk12_0)|~r2_hidden(esk2_0,esk12_0))).#trainneg
cnf(i_0_768, plain, (m1_subset_1(k2_xboole_0(X1,k1_tarski(X1)),X2)|~r1_tarski(k2_xboole_0(X1,k1_tarski(X1)),X2)|~v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_2208, plain, (r2_hidden(k10_ordinal2(X1,X2),X3)|~r1_ordinal1(X1,X3)|~r2_hidden(X2,k1_xboole_0)|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~r2_hidden(X3,X1))).#trainneg
cnf(i_0_4034, plain, (m1_subset_1(k10_ordinal2(X1,X2),k10_ordinal2(X3,X4))|~r1_ordinal1(X1,X3)|~r2_hidden(X2,X4)|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~v3_ordinal1(X4))).#trainneg
cnf(i_0_4036, plain, (~r1_ordinal1(X1,X2)|~r2_hidden(k10_ordinal2(X2,X3),X4)|~r2_hidden(esk5_2(X1,X4),X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~r1_ordinal1(X1,X4)|~v3_ordinal1(X4))).#trainneg
cnf(i_0_3311, plain, (v1_xboole_0(k10_ordinal2(X1,X2))|~r1_ordinal1(X1,k2_xboole_0(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2))),k1_tarski(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2))))))|~r2_hidden(k1_xboole_0,X2)|~r2_hidden(k2_xboole_0(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2))),k1_tarski(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2))))),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(k2_xboole_0(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2))),k1_tarski(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2))))))|~v3_ordinal1(k10_ordinal2(X1,X2)))).#trainneg
cnf(i_0_2210, plain, (r1_ordinal1(k10_ordinal2(X1,X2),X3)|r2_hidden(k1_xboole_0,X2)|~r1_ordinal1(X1,X3)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~r2_hidden(X3,X1))).#trainneg
cnf(i_0_1436, plain, (m1_subset_1(X1,k10_ordinal2(X1,X2))|v1_xboole_0(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_962, plain, (m1_subset_1(k10_ordinal2(X1,X2),X3)|~r1_tarski(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))),X3)|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_605, plain, (v3_ordinal1(k10_ordinal2(X1,X2))|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4062, plain, (v3_ordinal1(k10_ordinal2(X1,k1_tarski(k1_xboole_0)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_268, plain, (r1_ordinal1(X1,esk12_0)|r1_ordinal1(esk12_0,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1939, negated_conjecture, (esk15_0=esk1_0|r2_hidden(esk1_0,esk15_0)|r2_hidden(esk15_0,esk1_0))).#trainneg
cnf(i_0_4426, plain, (r1_ordinal1(esk12_0,esk12_0))).#trainneg
cnf(i_0_280, plain, (r1_ordinal1(k1_tarski(k1_xboole_0),X1)|~r2_hidden(k1_xboole_0,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_281, plain, (r2_hidden(k1_tarski(k1_xboole_0),X1)|~r2_hidden(k1_xboole_0,X1)|~v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_4220, plain, (v1_xboole_0(X2)|~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2))),X1)|~v3_ordinal1(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2))))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_4354, plain, (m1_subset_1(k10_ordinal2(X1,X2),k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_4222, plain, (r2_hidden(X1,k10_ordinal2(X2,X3))|v1_xboole_0(X3)|~v3_ordinal1(k10_ordinal2(X2,X3))|~v3_ordinal1(X1)|~r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),X2)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_4241, plain, (r2_hidden(X1,X2)|r2_hidden(X2,k10_ordinal2(X3,X4))|v1_xboole_0(X4)|~v3_ordinal1(k10_ordinal2(X3,X4))|~v3_ordinal1(X2)|~r1_ordinal1(X1,X3)|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~v3_ordinal1(X4))).#trainneg
cnf(i_0_581, plain, (r2_hidden(X1,X2)|r2_hidden(X2,k3_ordinal2(X3))|~v3_ordinal1(X2)|~r2_hidden(X1,X3)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1501, plain, (r2_hidden(X1,X2)|r2_hidden(X2,k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))))|~v3_ordinal1(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_4521, plain, (~v4_ordinal1(k1_tarski(k1_xboole_0)))).#trainneg
cnf(i_0_505, plain, (r1_ordinal1(esk14_0,esk13_0)|r1_ordinal1(esk13_0,esk14_0))).#trainneg
cnf(i_0_290, plain, (v3_ordinal1(esk7_3(esk4_2(X1,X2),X3,X4))|~v3_ordinal1(X1)|X3!=k10_xtuple_0(esk4_2(X1,X2))|~r2_hidden(X4,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2112, negated_conjecture, (esk12_0=esk2_0|r2_hidden(esk12_0,esk2_0)|r1_ordinal1(esk2_0,esk12_0))).#trainneg
cnf(i_0_4172, negated_conjecture, (~v3_ordinal1(k10_ordinal2(esk2_0,esk1_0))|~r1_ordinal1(esk2_0,k1_xboole_0)|~r2_hidden(esk1_0,k1_xboole_0)|~r2_hidden(k1_xboole_0,esk2_0))).#trainneg
cnf(i_0_2852, negated_conjecture, (r1_ordinal1(esk15_0,esk2_0)|~r2_hidden(esk15_0,esk2_0))).#trainneg
cnf(i_0_293, plain, (r2_hidden(k10_ordinal2(X1,X2),X3)|~r2_hidden(X1,X3)|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(X2,k1_xboole_0)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_294, plain, (r1_ordinal1(X2,X1)|~r1_ordinal1(esk5_2(X1,X2),k1_xboole_0)|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_4569, plain, (v1_xboole_0(X1)|r2_hidden(X2,k10_ordinal2(X3,X1))|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~r2_hidden(X2,X3))).#trainneg
cnf(i_0_1899, plain, (~v3_ordinal1(X1)|~r2_hidden(k10_ordinal2(X1,k1_tarski(k1_xboole_0)),X1)|~v4_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(X1,k1_tarski(k1_xboole_0))))).#trainneg
cnf(i_0_4793, plain, (r1_ordinal1(X1,X2)|X3!=k10_xtuple_0(esk4_2(X4,X1))|~r2_hidden(X2,X3)|~v3_ordinal1(X1)|~v3_ordinal1(X4))).#trainneg
cnf(i_0_4953, plain, (r2_hidden(k2_xboole_0(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))),k1_tarski(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))))),X2)|r2_hidden(X2,X3)|r2_hidden(X3,X1)|v4_ordinal1(X1)|~v4_ordinal1(X3)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1690, plain, (~r2_hidden(k2_xboole_0(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))),k1_tarski(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))))),k10_ordinal2(X1,X2))|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_4852, plain, (v4_ordinal1(X1)|~v3_ordinal1(X1)|~r1_ordinal1(k1_tarski(k1_xboole_0),k1_xboole_0)|~r2_hidden(esk6_1(X1),X1)|~v3_ordinal1(k10_ordinal2(esk6_1(X1),k1_tarski(k1_xboole_0))))).#trainneg
cnf(i_0_4553, plain, (m1_subset_1(k10_ordinal2(X1,k1_xboole_0),k10_ordinal2(X1,k1_tarski(k1_xboole_0)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_506, plain, (r1_ordinal1(esk14_0,esk12_0)|r1_ordinal1(esk12_0,esk14_0))).#trainneg
cnf(i_0_297, plain, (m1_subset_1(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))),X3)|~r1_tarski(X4,X3)|~r2_hidden(k10_ordinal2(X1,X2),X4)|~v4_ordinal1(X4)|~v3_ordinal1(X4)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_2132, plain, (esk15_0=k1_xboole_0|r2_hidden(k1_xboole_0,esk15_0)|r2_hidden(esk15_0,k1_xboole_0))).#trainneg
cnf(i_0_4849, negated_conjecture, (~v3_ordinal1(k10_ordinal2(esk2_0,esk1_0))|~r1_ordinal1(esk1_0,k1_xboole_0)|~r2_hidden(esk2_0,k1_xboole_0))).#trainneg
cnf(i_0_2409, negated_conjecture, (r2_hidden(k1_xboole_0,X1)|r2_hidden(X1,esk2_0)|esk2_0=esk1_0|r2_hidden(k1_xboole_0,esk1_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1105, plain, (r2_hidden(k1_xboole_0,X1)|r2_hidden(X1,X2)|k1_xboole_0=X2|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_301, plain, (r2_hidden(k10_ordinal2(X1,k10_ordinal2(X2,k2_xboole_0(X3,k1_tarski(X3)))),X4)|~r2_hidden(k10_ordinal2(X1,k10_ordinal2(X2,X3)),X4)|~v4_ordinal1(X4)|~v3_ordinal1(X4)|~v3_ordinal1(k10_ordinal2(X2,X3))|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2487, plain, (r2_hidden(X1,X2)|r2_hidden(X2,k10_ordinal2(X3,X4))|~v3_ordinal1(k10_ordinal2(X3,X4))|~v3_ordinal1(X2)|~r1_ordinal1(X3,X1)|~r2_hidden(k1_xboole_0,X4)|~r2_hidden(X1,X3)|~v3_ordinal1(X3)|~v3_ordinal1(X4))).#trainneg
cnf(i_0_5035, plain, (v1_xboole_0(X1)|v4_ordinal1(k10_ordinal2(X2,X1))|v1_xboole_0(X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_tarski(k10_ordinal2(X2,X1),X2)|~v3_ordinal1(k10_ordinal2(X2,X1)))).#trainneg
cnf(i_0_5021, plain, (v1_xboole_0(esk7_3(esk4_2(X1,X2),X3,X4))|v4_ordinal1(X4)|~r2_hidden(esk6_1(X4),X2)|~v3_ordinal1(X2)|~v3_ordinal1(esk7_3(esk4_2(X1,X2),X3,X4))|X3!=k10_xtuple_0(esk4_2(X1,X2))|~r2_hidden(X4,X3)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_5037, plain, (r1_ordinal1(X1,X2)|~r2_hidden(X2,k10_xtuple_0(esk4_2(X3,X1)))|~v3_ordinal1(X1)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_5316, plain, (r2_hidden(X1,k10_ordinal2(X2,X3))|r2_hidden(X2,X1)|~r2_hidden(k1_xboole_0,X3)|~r2_hidden(X2,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_1399, plain, (esk7_3(esk4_2(X1,X2),X3,X4)=k1_xboole_0|r2_hidden(X2,X4)|~v3_ordinal1(X2)|~v3_ordinal1(esk7_3(esk4_2(X1,X2),X3,X4))|X3!=k10_xtuple_0(esk4_2(X1,X2))|~r2_hidden(X4,X3)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_4234, plain, (r2_hidden(X1,k10_ordinal2(X2,X3))|v1_xboole_0(X3)|~v3_ordinal1(k10_ordinal2(X2,X3))|~v3_ordinal1(X1)|~v3_ordinal1(X4)|~r1_ordinal1(k10_ordinal2(X1,X4),X2)|~v3_ordinal1(k10_ordinal2(X1,X4))|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_2862, plain, (r1_ordinal1(k10_ordinal2(X1,X2),k10_ordinal2(X1,X3))|~r2_hidden(k10_ordinal2(X1,X2),k10_ordinal2(X1,X3))|~v3_ordinal1(k10_ordinal2(X1,X3))|~r1_ordinal1(X3,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_3538, plain, (v1_xboole_0(X1)|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_ordinal1(k2_xboole_0(X2,k1_tarski(X2)),X1)|~r1_ordinal1(X1,X2)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_1995, plain, (k10_ordinal2(X1,k1_tarski(k1_xboole_0))=k2_xboole_0(X2,k1_tarski(X2))|r2_hidden(X1,X2)|r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_5069, plain, (v4_ordinal1(X1)|~r1_ordinal1(k1_tarski(k1_xboole_0),k1_xboole_0)|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(X1)))).#trainneg
cnf(i_0_5137, plain, (esk15_0=k1_xboole_0|r2_hidden(k1_xboole_0,esk15_0))).#trainneg
cnf(i_0_302, plain, (r2_hidden(k10_ordinal2(X1,k1_tarski(k1_xboole_0)),X2)|~r2_hidden(k10_ordinal2(X1,k1_xboole_0),X2)|~v4_ordinal1(X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_2853, negated_conjecture, (r1_ordinal1(esk14_0,esk2_0)|~r2_hidden(esk14_0,esk2_0))).#trainneg
cnf(i_0_2855, negated_conjecture, (r1_ordinal1(esk1_0,esk2_0)|~r2_hidden(esk1_0,esk2_0))).#trainneg
cnf(i_0_5187, negated_conjecture, (esk2_0=k1_xboole_0|r2_hidden(k1_xboole_0,X1)|r2_hidden(X1,esk2_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_304, plain, (v7_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~v7_ordinal1(k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))))|~v7_ordinal1(X1)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_5327, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(X2,X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_5447, plain, (v1_xboole_0(X1)|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_ordinal1(X1,X2)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_5322, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|v4_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r1_ordinal1(k10_ordinal2(X1,X2),X1)|~v3_ordinal1(k10_ordinal2(X1,X2)))).#trainneg
cnf(i_0_5396, plain, (v1_xboole_0(X1)|r2_hidden(X2,k10_ordinal2(X3,X1))|~r1_ordinal1(X4,X3)|~v3_ordinal1(X2)|~v3_ordinal1(esk5_2(X2,X4))|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~r1_ordinal1(X2,X4)|~v3_ordinal1(X4))).#trainneg
cnf(i_0_996, plain, (r2_hidden(X1,k2_xboole_0(k3_ordinal2(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(k3_ordinal2(k2_xboole_0(X1,k1_tarski(X1))))))|~v3_ordinal1(k2_xboole_0(k3_ordinal2(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(k3_ordinal2(k2_xboole_0(X1,k1_tarski(X1))))))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_546, plain, (r1_tarski(esk6_1(X1),X2)|v4_ordinal1(X1)|~r1_tarski(X1,k1_zfmisc_1(X2))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_309, plain, (r1_ordinal1(k2_xboole_0(X2,k1_tarski(X2)),X1)|~r1_ordinal1(k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))),k1_xboole_0)|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2150, negated_conjecture, (esk2_0=esk14_0|r2_hidden(esk14_0,esk2_0)|r2_hidden(esk2_0,esk14_0))).#trainneg
cnf(i_0_5531, plain, (~r2_hidden(k10_ordinal2(k1_tarski(k1_xboole_0),k1_xboole_0),k1_xboole_0)|~v4_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_5586, negated_conjecture, (esk2_0=k1_xboole_0|r2_hidden(esk1_0,esk2_0)|r2_hidden(k1_xboole_0,esk1_0))).#trainneg
cnf(i_0_306, plain, (r1_ordinal1(k2_xboole_0(X2,k1_tarski(X2)),k10_ordinal2(X1,X3))|~r1_ordinal1(k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))),X3)|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~v3_ordinal1(k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_5721, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|v4_ordinal1(k10_ordinal2(X2,X1))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(k10_ordinal2(X2,X1),X2))).#trainneg
cnf(i_0_5745, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|k1_xboole_0=X2|~r1_ordinal1(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),X1)|~v3_ordinal1(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_5620, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(k3_ordinal2(X2),X1)|~r2_hidden(X1,k3_ordinal2(X2))|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))).#trainneg
cnf(i_0_5727, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(k3_ordinal2(X2),X1)|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))).#trainneg
cnf(i_0_5734, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))),X1)|~v3_ordinal1(k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))).#trainneg
cnf(i_0_5739, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(k1_xboole_0)),X1)|~v3_ordinal1(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(k1_xboole_0)))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))).#trainneg
cnf(i_0_5735, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|r2_hidden(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r1_ordinal1(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_355, plain, (v3_ordinal1(k10_ordinal2(k2_xboole_0(X1,k1_tarski(X1)),esk12_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_313, plain, (r2_hidden(k2_xboole_0(X2,k1_tarski(X2)),X3)|~r2_hidden(k10_ordinal2(X1,esk5_2(X1,X2)),X3)|~v4_ordinal1(X3)|~v3_ordinal1(X3)|~v3_ordinal1(esk5_2(X1,X2))|~v3_ordinal1(X1)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2151, negated_conjecture, (esk13_0=esk2_0|r2_hidden(esk13_0,esk2_0)|r2_hidden(esk2_0,esk13_0))).#trainneg
cnf(i_0_6039, plain, (~r1_ordinal1(k1_tarski(k1_xboole_0),k1_xboole_0))).#trainneg
cnf(i_0_307, plain, (r1_ordinal1(k2_xboole_0(X2,k1_tarski(X2)),k10_ordinal2(X3,k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2)))))|~r1_ordinal1(X1,X3)|~v3_ordinal1(k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))))|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_308, plain, (k10_ordinal2(X1,k2_xboole_0(k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))),k1_tarski(k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))))))=k2_xboole_0(k2_xboole_0(X2,k1_tarski(X2)),k1_tarski(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))))|~v3_ordinal1(X1)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_6055, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1))).#trainneg
cnf(i_0_6044, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1)|~r2_hidden(X1,X1))).#trainneg
cnf(i_0_5964, plain, (v1_xboole_0(X1)|v4_ordinal1(k10_ordinal2(X2,X1))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(X1,k1_xboole_0)|~r2_hidden(X2,X2)|~v3_ordinal1(k10_ordinal2(X2,X1)))).#trainneg
cnf(i_0_6075, plain, (r2_hidden(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_6313, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v1_xboole_0(X1)|~v3_ordinal1(X1)|~r1_tarski(k2_xboole_0(X1,k1_tarski(X1)),X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_6354, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_6371, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_6310, plain, (v4_ordinal1(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2))))|v1_xboole_0(X2)|~v3_ordinal1(k10_ordinal2(X1,X2))|~r2_hidden(esk6_1(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2)))),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_6353, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1)|~r2_hidden(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_1930, plain, (r2_hidden(k1_xboole_0,X1)|~v3_ordinal1(X1)|~r2_hidden(X2,k1_xboole_0)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_5542, plain, (~v4_ordinal1(k10_ordinal2(X1,k1_tarski(k1_xboole_0)))|~v3_ordinal1(X1)|~r2_hidden(k10_ordinal2(X1,k1_xboole_0),X1)|~v4_ordinal1(X1))).#trainneg
cnf(i_0_314, plain, (r2_hidden(k10_ordinal2(X1,k10_ordinal2(X2,X3)),X4)|~r1_ordinal1(k10_ordinal2(X1,k10_ordinal2(X2,k2_xboole_0(X3,k1_tarski(X3)))),X4)|~v3_ordinal1(X4)|~v3_ordinal1(k10_ordinal2(X2,X3))|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2152, negated_conjecture, (esk12_0=esk2_0|r2_hidden(esk12_0,esk2_0)|r2_hidden(esk2_0,esk12_0))).#trainneg
cnf(i_0_5889, negated_conjecture, (r2_hidden(esk1_0,X1)|r2_hidden(X1,esk2_0)|esk2_0=k1_xboole_0|r2_hidden(k1_xboole_0,esk1_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_2325, negated_conjecture, (v3_ordinal1(k10_ordinal2(X1,esk2_0))|esk2_0=esk1_0|r2_hidden(esk1_0,esk2_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_310, plain, (r1_ordinal1(k10_ordinal2(X1,k2_xboole_0(esk5_2(X2,X3),k1_tarski(esk5_2(X2,X3)))),k2_xboole_0(X3,k1_tarski(X3)))|~r1_ordinal1(X1,X2)|~v3_ordinal1(k2_xboole_0(esk5_2(X2,X3),k1_tarski(esk5_2(X2,X3))))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(X2,X3)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_311, plain, (r2_hidden(k10_ordinal2(X1,X2),k2_xboole_0(X3,k1_tarski(X3)))|~r2_hidden(X2,k2_xboole_0(esk5_2(X1,X3),k1_tarski(esk5_2(X1,X3))))|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(esk5_2(X1,X3),k1_tarski(esk5_2(X1,X3))))|~v3_ordinal1(X2)|~r1_ordinal1(X1,X3)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_6389, plain, (r2_hidden(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~r2_hidden(k10_ordinal2(X1,X3),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_6426, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|r2_hidden(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_6425, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))))|~r2_hidden(esk6_1(k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_6433, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(k1_xboole_0)),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_6369, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(k2_xboole_0(k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))),k1_tarski(k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))))),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1932, plain, (r2_hidden(X1,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(X2,k1_xboole_0))).#trainneg
cnf(i_0_4060, plain, (v3_ordinal1(k10_ordinal2(k10_ordinal2(X1,k1_tarski(k1_xboole_0)),esk12_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4506, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(k1_xboole_0,k10_ordinal2(k1_tarski(k1_xboole_0),X1))|~v4_ordinal1(k10_ordinal2(k1_tarski(k1_xboole_0),X1))|~v3_ordinal1(k10_ordinal2(k1_tarski(k1_xboole_0),X1)))).#trainneg
cnf(i_0_315, plain, (r2_hidden(k10_ordinal2(X1,k1_xboole_0),X2)|~r1_ordinal1(k10_ordinal2(X1,k1_tarski(k1_xboole_0)),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_2800, plain, (r2_hidden(k1_xboole_0,esk15_0)|esk13_0=esk15_0|r2_hidden(esk15_0,esk13_0))).#trainneg
cnf(i_0_6301, plain, (~r2_hidden(esk6_1(k1_tarski(k1_xboole_0)),k1_xboole_0))).#trainneg
cnf(i_0_6057, plain, (~r1_ordinal1(k10_ordinal2(esk6_1(k1_tarski(k1_xboole_0)),k1_tarski(k1_xboole_0)),k1_xboole_0)|~v3_ordinal1(esk6_1(k1_tarski(k1_xboole_0))))).#trainneg
cnf(i_0_312, plain, (r1_ordinal1(k10_ordinal2(X1,X2),k2_xboole_0(X3,k1_tarski(X3)))|~r1_ordinal1(X2,k2_xboole_0(esk5_2(X1,X3),k1_tarski(esk5_2(X1,X3))))|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(esk5_2(X1,X3),k1_tarski(esk5_2(X1,X3))))|~v3_ordinal1(X2)|~r1_ordinal1(X1,X3)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_316, plain, (r2_hidden(k10_ordinal2(X1,esk5_2(X1,X2)),X3)|~r1_ordinal1(k2_xboole_0(X2,k1_tarski(X2)),X3)|~v3_ordinal1(X3)|~v3_ordinal1(esk5_2(X1,X2))|~v3_ordinal1(X1)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_6363, plain, (v4_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(k2_xboole_0(esk6_1(esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),k1_tarski(esk6_1(esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_6991, plain, (v4_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),X1)|~v4_ordinal1(X1)|~v3_ordinal1(esk6_1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))))).#trainneg
cnf(i_0_6763, plain, (v4_ordinal1(k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v1_xboole_0(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~r1_tarski(k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))),X1)|~v3_ordinal1(k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))))).#trainneg
cnf(i_0_6979, plain, (v4_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k10_ordinal2(esk6_1(esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),k1_tarski(k1_xboole_0)),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))))).#trainneg
cnf(i_0_6743, plain, (r2_hidden(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|v1_xboole_0(X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r1_tarski(k2_xboole_0(X1,k1_tarski(X1)),X2)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_6948, plain, (r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),k10_ordinal2(X2,X3))|~r1_ordinal1(X2,X1)|~r2_hidden(k1_xboole_0,X3)|~v4_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_6717, plain, (r2_hidden(k1_xboole_0,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(esk6_1(k1_tarski(k1_xboole_0)),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_6373, plain, (r2_hidden(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X1)|~r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_6457, plain, (v1_xboole_0(X1)|v4_ordinal1(k10_ordinal2(X2,k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(esk6_1(k10_ordinal2(X2,k2_xboole_0(X1,k1_tarski(X1)))),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_317, plain, (r2_hidden(k10_ordinal2(X1,X2),k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_3059, plain, (r2_hidden(k1_xboole_0,esk15_0)|esk12_0=esk15_0|r2_hidden(esk15_0,esk12_0))).#trainneg
cnf(i_0_6766, plain, (~r2_hidden(k10_ordinal2(esk6_1(k1_tarski(k1_xboole_0)),k1_tarski(k1_xboole_0)),k1_xboole_0)|~v3_ordinal1(esk6_1(k1_tarski(k1_xboole_0))))).#trainneg
cnf(i_0_6046, plain, (~r1_ordinal1(k2_xboole_0(esk6_1(k1_tarski(k1_xboole_0)),k1_tarski(esk6_1(k1_tarski(k1_xboole_0)))),k1_xboole_0)|~v3_ordinal1(esk6_1(k1_tarski(k1_xboole_0))))).#trainneg
cnf(i_0_2326, negated_conjecture, (v3_ordinal1(k10_ordinal2(X1,esk1_0))|esk2_0=esk1_0|r2_hidden(k1_xboole_0,esk1_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_318, plain, (r2_hidden(k10_ordinal2(X1,X2),X3)|r2_hidden(X3,k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2)))))).#trainneg
cnf(i_0_319, plain, (r2_hidden(k10_ordinal2(X1,X2),k10_ordinal2(X1,X3))|~v3_ordinal1(k10_ordinal2(X1,X3))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(k2_xboole_0(X2,k1_tarski(X2)),X3)|~v3_ordinal1(X3)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2))))).#trainneg
cnf(i_0_7006, plain, (v4_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v1_xboole_0(X1)|~v4_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~r1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1))).#trainneg
cnf(i_0_7007, plain, (v4_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|r2_hidden(X1,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v4_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_7090, plain, (r1_ordinal1(k10_ordinal2(X2,X3),X1)|~r2_hidden(k10_ordinal2(X2,X3),k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(X2,X1)|~v4_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_7197, plain, (~r2_hidden(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))),k10_ordinal2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_6023, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|r2_hidden(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r1_ordinal1(k3_ordinal2(k2_xboole_0(X2,k1_tarski(X2))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_6784, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~r2_hidden(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_xboole_0),X1)|~v4_ordinal1(X1))).#trainneg
cnf(i_0_5742, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v1_xboole_0(X3)|~r1_ordinal1(k10_ordinal2(X2,X3),X1)|~v3_ordinal1(k10_ordinal2(X2,X3))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_6706, plain, (r2_hidden(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|v4_ordinal1(esk5_2(X1,X2))|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(esk5_2(X1,X2)))|~v3_ordinal1(X2)|~r1_ordinal1(X1,X2))).#trainneg
cnf(i_0_6384, plain, (k1_xboole_0=X1|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X1)|~r2_hidden(k10_ordinal2(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_519, plain, (r1_ordinal1(esk14_0,k1_xboole_0)|r1_ordinal1(k1_xboole_0,esk14_0))).#trainneg
cnf(i_0_320, plain, (r2_hidden(k10_ordinal2(X1,X2),k10_ordinal2(X3,k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(k10_ordinal2(X3,k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(X1,X3)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(X3))).#trainneg
cnf(i_0_5897, negated_conjecture, (v3_ordinal1(k10_ordinal2(X1,esk1_0))|esk2_0=k1_xboole_0|r2_hidden(k1_xboole_0,esk1_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_321, plain, (r2_hidden(k10_ordinal2(X1,X2),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r1_ordinal1(k2_xboole_0(X2,k1_tarski(X2)),k1_xboole_0)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2))))).#trainneg
cnf(i_0_322, plain, (r1_ordinal1(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))),X3)|~r1_ordinal1(X1,X3)|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(X3)|~r2_hidden(X2,esk5_2(X1,X3))|~v3_ordinal1(esk5_2(X1,X3))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_7383, plain, (v1_xboole_0(X1)|v4_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v4_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1))).#trainneg
cnf(i_0_7450, plain, (~r2_hidden(k10_ordinal2(X1,k1_tarski(k1_xboole_0)),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_7393, plain, (r1_ordinal1(X2,X3)|~r1_ordinal1(X1,X3)|~r2_hidden(X2,k2_xboole_0(X3,k1_tarski(X3)))|~v4_ordinal1(esk5_2(X1,X2))|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X1,X2))|~v3_ordinal1(X3)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_7528, plain, (v1_xboole_0(X1)|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|v1_xboole_0(X3)|~r1_ordinal1(k10_ordinal2(X3,X1),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~r1_tarski(k2_xboole_0(X2,k1_tarski(X2)),X3)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2))))).#trainneg
cnf(i_0_7533, plain, (r2_hidden(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|v4_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_7556, plain, (k2_xboole_0(X1,k1_tarski(X1))=k1_xboole_0|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X2)|~r2_hidden(k10_ordinal2(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1),X2)|~v4_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_7527, plain, (v1_xboole_0(X1)|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_ordinal1(k10_ordinal2(k2_xboole_0(X2,k1_tarski(X2)),X1),X2)|~v3_ordinal1(X2)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_7216, plain, (r2_hidden(k10_ordinal2(X1,k1_xboole_0),k10_ordinal2(X1,k1_tarski(k1_xboole_0)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_7564, plain, (r1_ordinal1(k1_xboole_0,esk14_0))).#trainneg
cnf(i_0_323, plain, (r1_ordinal1(k10_ordinal2(X1,k10_ordinal2(esk5_2(X1,X2),X3)),X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(esk5_2(X1,X2),X3))|~v3_ordinal1(X2)|~r1_ordinal1(X3,k1_xboole_0)|~v3_ordinal1(esk5_2(X1,X2))|~v3_ordinal1(X3))).#trainneg
cnf(i_0_3072, negated_conjecture, (esk14_0=esk1_0|r2_hidden(esk14_0,esk1_0)|r2_hidden(esk1_0,esk14_0))).#trainneg
cnf(i_0_7244, plain, (~v3_ordinal1(esk6_1(k1_tarski(k1_xboole_0)))|~r2_hidden(k10_ordinal2(esk6_1(k1_tarski(k1_xboole_0)),k1_xboole_0),k1_xboole_0)|~v4_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_325, plain, (r1_ordinal1(k10_ordinal2(X1,esk5_2(X1,X2)),X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X1,X2))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_7388, plain, (r2_hidden(X1,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|v4_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_7081, plain, (m1_subset_1(k2_xboole_0(X1,k1_tarski(X1)),X2)|r1_ordinal1(k10_ordinal2(X3,X4),X1)|~r1_tarski(k10_ordinal2(X3,X4),X2)|~r1_ordinal1(X3,X1)|~v4_ordinal1(X4)|~v3_ordinal1(X3)|~v3_ordinal1(X4)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_6446, plain, (v1_xboole_0(esk7_3(esk4_2(X1,X2),X3,X4))|v4_ordinal1(k2_xboole_0(X4,k1_tarski(X4)))|~r2_hidden(esk6_1(k2_xboole_0(X4,k1_tarski(X4))),X2)|~v3_ordinal1(X2)|~v3_ordinal1(esk7_3(esk4_2(X1,X2),X3,X4))|X3!=k10_xtuple_0(esk4_2(X1,X2))|~r2_hidden(X4,X3)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_7069, plain, (~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(X1,k10_ordinal2(X1,X2))|~r2_hidden(k1_xboole_0,X2)|~r2_hidden(k10_ordinal2(X1,X2),X1)|~v4_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_7878, plain, (k2_xboole_0(X1,k1_tarski(X1))=k1_xboole_0|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|k10_ordinal2(X2,X1)=X2|~v4_ordinal1(X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(X1,k1_xboole_0)|~r1_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X2))).#trainneg
cnf(i_0_7897, plain, (v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(X2)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k10_ordinal2(k2_xboole_0(X2,k1_tarski(X2)),X1),X2)|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2))))).#trainneg
cnf(i_0_7895, plain, (v1_xboole_0(esk5_2(X1,X2))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(X2)|~v3_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(k2_xboole_0(X2,k1_tarski(X2)),X1)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2))))).#trainneg
cnf(i_0_8021, plain, (esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1)=k1_xboole_0|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1))|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1))).#trainneg
cnf(i_0_327, plain, (v3_ordinal1(X4)|~v3_ordinal1(esk7_3(esk4_2(X2,X1),X3,X4))|~v3_ordinal1(X1)|X3!=k10_xtuple_0(esk4_2(X2,X1))|~r2_hidden(X4,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_3073, negated_conjecture, (esk13_0=esk1_0|r2_hidden(esk13_0,esk1_0)|r2_hidden(esk1_0,esk13_0))).#trainneg
cnf(i_0_1166, negated_conjecture, (r2_hidden(esk2_0,k10_ordinal2(esk2_0,X1))|r2_hidden(X1,X2)|v4_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_328, plain, (v7_ordinal1(X4)|~v7_ordinal1(esk7_3(esk4_2(X2,X1),X3,X4))|~v7_ordinal1(X1)|X3!=k10_xtuple_0(esk4_2(X2,X1))|~r2_hidden(X4,X3)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_329, plain, (r2_hidden(X4,k10_ordinal2(X1,X5))|~r2_hidden(esk7_3(esk4_2(X2,X1),X3,X4),X5)|~v3_ordinal1(X1)|~v3_ordinal1(X5)|~v3_ordinal1(esk7_3(esk4_2(X2,X1),X3,X4))|X3!=k10_xtuple_0(esk4_2(X2,X1))|~r2_hidden(X4,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_8110, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k1_xboole_0,X2)|~v4_ordinal1(X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X2)|~r2_hidden(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_8077, plain, (v1_xboole_0(esk7_3(esk4_2(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2)))),X3,X4))|v4_ordinal1(k2_xboole_0(X4,k1_tarski(X4)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|X3!=k10_xtuple_0(esk4_2(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2)))))|~r2_hidden(X4,X3)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k2_xboole_0(X4,k1_tarski(X4))),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_8199, plain, (k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_xboole_0)=X1|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_8146, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(k2_xboole_0(X1,k1_tarski(X1)),X2))|~r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),X2)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_8165, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(esk5_2(X2,X1))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(esk5_2(X2,X1)))|~r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),X2)|~r1_ordinal1(X2,X1)|~v3_ordinal1(X2)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_8221, plain, (v3_ordinal1(X1)|~r2_hidden(X1,k10_xtuple_0(esk4_2(X2,X3)))|~v3_ordinal1(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_816, plain, (m1_subset_1(k10_ordinal2(X1,X2),X1)|~r2_hidden(X2,k1_xboole_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_330, plain, (r1_ordinal1(X4,k10_ordinal2(X1,X5))|~r1_ordinal1(esk7_3(esk4_2(X2,X1),X3,X4),X5)|~v3_ordinal1(X1)|~v3_ordinal1(X5)|~v3_ordinal1(esk7_3(esk4_2(X2,X1),X3,X4))|X3!=k10_xtuple_0(esk4_2(X2,X1))|~r2_hidden(X4,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_3074, negated_conjecture, (esk12_0=esk1_0|r2_hidden(esk12_0,esk1_0)|r2_hidden(esk1_0,esk12_0))).#trainneg
cnf(i_0_6428, plain, (~r2_hidden(k2_xboole_0(esk6_1(k1_tarski(k1_xboole_0)),k1_tarski(esk6_1(k1_tarski(k1_xboole_0)))),k1_xboole_0)|~v3_ordinal1(esk6_1(k1_tarski(k1_xboole_0))))).#trainneg
cnf(i_0_8286, negated_conjecture, (r2_hidden(esk2_0,k10_ordinal2(esk2_0,esk2_0))|r2_hidden(esk2_0,X1)|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_8287, negated_conjecture, (r2_hidden(esk2_0,k10_ordinal2(esk2_0,esk1_0))|r2_hidden(esk1_0,X1)|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_331, plain, (r1_ordinal1(X4,k10_ordinal2(X5,esk7_3(esk4_2(X2,X1),X3,X4)))|~r1_ordinal1(X1,X5)|~v3_ordinal1(esk7_3(esk4_2(X2,X1),X3,X4))|~v3_ordinal1(X5)|~v3_ordinal1(X1)|X3!=k10_xtuple_0(esk4_2(X2,X1))|~r2_hidden(X4,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_332, plain, (k10_ordinal2(X1,k2_xboole_0(esk7_3(esk4_2(X2,X1),X3,X4),k1_tarski(esk7_3(esk4_2(X2,X1),X3,X4))))=k2_xboole_0(X4,k1_tarski(X4))|~v3_ordinal1(esk7_3(esk4_2(X2,X1),X3,X4))|~v3_ordinal1(X1)|X3!=k10_xtuple_0(esk4_2(X2,X1))|~r2_hidden(X4,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_8419, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),X2)|~r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_8205, plain, (k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(k1_xboole_0))=k2_xboole_0(X1,k1_tarski(X1))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_8184, plain, (v1_xboole_0(esk5_2(X1,X2))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_8430, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(esk5_2(X2,X1))|~r1_ordinal1(X2,X1)|~v3_ordinal1(esk6_1(esk5_2(X2,X1)))|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_8168, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1)|~r2_hidden(k10_ordinal2(k2_xboole_0(X1,k1_tarski(X1)),k1_xboole_0),X1)|~v4_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_8369, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1))|~r2_hidden(k1_xboole_0,esk6_1(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1)))|~v4_ordinal1(esk6_1(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1)))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(esk6_1(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1)))|~v3_ordinal1(X1)|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1))).#trainneg
cnf(i_0_6975, plain, (v4_ordinal1(esk6_1(k10_ordinal2(X1,k1_tarski(k1_xboole_0))))|v4_ordinal1(k10_ordinal2(X1,k1_tarski(k1_xboole_0)))|~r2_hidden(k2_xboole_0(esk6_1(esk6_1(k10_ordinal2(X1,k1_tarski(k1_xboole_0)))),k1_tarski(esk6_1(esk6_1(k10_ordinal2(X1,k1_tarski(k1_xboole_0)))))),X1)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,k1_tarski(k1_xboole_0))))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_6401, plain, (v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1)|~r2_hidden(X1,X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_8458, plain, (v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X1)|~r2_hidden(X2,k1_xboole_0))).#trainneg
cnf(i_0_698, plain, (r1_ordinal1(esk13_0,esk12_0)|r1_ordinal1(esk12_0,esk13_0))).#trainneg
cnf(i_0_334, plain, (r1_ordinal1(k10_ordinal2(X1,esk7_3(esk4_2(X2,X3),X4,X5)),X5)|~r1_ordinal1(X1,X3)|~v3_ordinal1(esk7_3(esk4_2(X2,X3),X4,X5))|~v3_ordinal1(X3)|~v3_ordinal1(X1)|X4!=k10_xtuple_0(esk4_2(X2,X3))|~r2_hidden(X5,X4)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_3362, negated_conjecture, (esk14_0=esk1_0|r2_hidden(esk14_0,esk1_0)|r1_ordinal1(esk1_0,esk14_0))).#trainneg
cnf(i_0_8247, negated_conjecture, (X1=esk2_0|r2_hidden(k1_xboole_0,X1)|r2_hidden(k1_xboole_0,esk2_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_335, plain, (r2_hidden(k10_ordinal2(X1,X2),X5)|~r2_hidden(X2,esk7_3(esk4_2(X3,X1),X4,X5))|~v3_ordinal1(X1)|~v3_ordinal1(esk7_3(esk4_2(X3,X1),X4,X5))|~v3_ordinal1(X2)|X4!=k10_xtuple_0(esk4_2(X3,X1))|~r2_hidden(X5,X4)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_336, plain, (r1_ordinal1(k10_ordinal2(X1,X2),X5)|~r1_ordinal1(X2,esk7_3(esk4_2(X3,X1),X4,X5))|~v3_ordinal1(X1)|~v3_ordinal1(esk7_3(esk4_2(X3,X1),X4,X5))|~v3_ordinal1(X2)|X4!=k10_xtuple_0(esk4_2(X3,X1))|~r2_hidden(X5,X4)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_8824, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_8765, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(esk5_2(X2,X1))|~r1_ordinal1(X2,X1)|~r2_hidden(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(esk5_2(X2,X1)))).#trainneg
cnf(i_0_8846, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v1_xboole_0(X2)|~r2_hidden(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_493, plain, (r1_ordinal1(esk10_1(k1_zfmisc_1(X1)),X1)|~v3_ordinal1(X1)|~v3_ordinal1(esk10_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_353, plain, (r1_ordinal1(X1,esk10_1(X2))|r1_ordinal1(esk10_1(X2),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2856, negated_conjecture, (r1_ordinal1(k1_xboole_0,esk2_0)|~r2_hidden(k1_xboole_0,esk2_0))).#trainneg
cnf(i_0_1167, negated_conjecture, (r2_hidden(esk1_0,k10_ordinal2(esk1_0,X1))|r2_hidden(X1,X2)|v4_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_356, plain, (v3_ordinal1(k10_ordinal2(k10_ordinal2(X1,X2),esk12_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_362, plain, (r1_ordinal1(esk15_0,esk10_1(X1))|r1_ordinal1(esk10_1(X1),esk15_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_9067, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|v1_xboole_0(X2)|~r2_hidden(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_tarski(k2_xboole_0(X1,k1_tarski(X1)),X2)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_9061, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(k10_ordinal2(X2,X3),k1_tarski(k10_ordinal2(X2,X3))))|v1_xboole_0(X3)|~r2_hidden(esk6_1(k2_xboole_0(k10_ordinal2(X2,X3),k1_tarski(k10_ordinal2(X2,X3)))),X1)|~v3_ordinal1(k10_ordinal2(X2,X3))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_9072, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(esk5_2(X2,X1))|~r1_ordinal1(X2,X1)|~r2_hidden(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_9039, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(esk6_1(k1_tarski(k1_xboole_0)),X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_xboole_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_9066, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))))|~r2_hidden(esk6_1(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1))))),X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_8834, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v1_xboole_0(X2)|~r2_hidden(k2_xboole_0(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),k1_tarski(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),k1_tarski(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_9355, plain, (v1_xboole_0(X1)|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|v1_xboole_0(X2)|~r1_tarski(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r1_tarski(k2_xboole_0(X1,k1_tarski(X1)),X2)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_9486, negated_conjecture, (~r2_hidden(esk1_0,k1_xboole_0))).#trainneg
cnf(i_0_9517, plain, (v4_ordinal1(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v1_xboole_0(X1)|~v3_ordinal1(X1)|~r1_tarski(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))),X1)|~v3_ordinal1(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1))))))).#trainneg
cnf(i_0_9487, negated_conjecture, (~v1_xboole_0(X1)|~r1_tarski(esk1_0,X1))).#trainneg
cnf(i_0_609, plain, (v3_ordinal1(k10_ordinal2(k1_tarski(k1_xboole_0),X1))|~v1_xboole_0(X1)|~v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_9489, negated_conjecture, (~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_363, plain, (r1_ordinal1(esk15_0,esk6_1(X1))|r1_ordinal1(esk6_1(X1),esk15_0)|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_3629, negated_conjecture, (esk13_0=esk1_0|r2_hidden(esk13_0,esk1_0)|r1_ordinal1(esk1_0,esk13_0))).#trainneg
cnf(i_0_2860, negated_conjecture, (r1_ordinal1(k1_xboole_0,esk1_0)|~r2_hidden(k1_xboole_0,esk1_0))).#trainneg
cnf(i_0_9231, negated_conjecture, (r2_hidden(esk1_0,k10_ordinal2(esk1_0,esk2_0))|r2_hidden(esk2_0,X1)|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_364, plain, (r1_ordinal1(esk15_0,esk5_2(X1,X2))|r1_ordinal1(esk5_2(X1,X2),esk15_0)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_365, plain, (r1_ordinal1(esk15_0,esk3_3(X1,X2,X3))|r1_ordinal1(esk3_3(X1,X2,X3),esk15_0)|k10_ordinal2(X2,X1)=k4_ordinal2(X3)|X1=k1_xboole_0|k9_xtuple_0(X3)!=X1|~v1_ordinal2(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v5_ordinal1(X3)|~v4_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_9522, plain, (v1_xboole_0(esk7_3(esk4_2(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2)))),X3,X4))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(k2_xboole_0(X4,k1_tarski(X4)),X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2)|~v3_ordinal1(esk7_3(esk4_2(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2)))),X3,X4))|X3!=k10_xtuple_0(esk4_2(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2)))))|~r2_hidden(X4,X3)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_9357, plain, (v1_xboole_0(esk7_3(esk4_2(X1,X2),X3,X4))|v4_ordinal1(k2_xboole_0(X4,k1_tarski(X4)))|v4_ordinal1(k2_xboole_0(X5,k1_tarski(X5)))|~r2_hidden(esk6_1(k2_xboole_0(X4,k1_tarski(X4))),X5)|~r2_hidden(esk6_1(k2_xboole_0(X5,k1_tarski(X5))),X2)|~v3_ordinal1(X5)|~v3_ordinal1(X2)|~v3_ordinal1(esk7_3(esk4_2(X1,X2),X3,X4))|X3!=k10_xtuple_0(esk4_2(X1,X2))|~r2_hidden(X4,X3)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_9370, plain, (v1_xboole_0(X1)|v4_ordinal1(k10_ordinal2(X2,k2_xboole_0(X1,k1_tarski(X1))))|v4_ordinal1(k2_xboole_0(X3,k1_tarski(X3)))|~r2_hidden(esk6_1(k10_ordinal2(X2,k2_xboole_0(X1,k1_tarski(X1)))),X3)|~r2_hidden(esk6_1(k2_xboole_0(X3,k1_tarski(X3))),X2)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_9060, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(k10_ordinal2(X2,X3),k1_tarski(k10_ordinal2(X2,X3))))|v1_xboole_0(X3)|~r2_hidden(esk6_1(k2_xboole_0(k10_ordinal2(X2,X3),k1_tarski(k10_ordinal2(X2,X3)))),X1)|~v3_ordinal1(k10_ordinal2(X2,X3))|~v3_ordinal1(X1)|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_9558, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_tarski(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),X2)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_9559, plain, (v1_xboole_0(X1)|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))))|~r1_tarski(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))),X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_9097, plain, (v1_xboole_0(esk6_1(esk5_2(X1,X2)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|v4_ordinal1(esk5_2(X1,X2))|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2)|~v3_ordinal1(esk6_1(esk5_2(X1,X2)))|~r1_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_9040, plain, (v4_ordinal1(k10_ordinal2(X1,k1_tarski(k1_xboole_0)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(esk6_1(k10_ordinal2(X1,k1_tarski(k1_xboole_0))),X2)|~r2_hidden(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_8133, plain, (~r2_hidden(k1_xboole_0,X2)|~v4_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r1_ordinal1(X2,k1_xboole_0)|~r2_hidden(X1,X1)|~v3_ordinal1(k10_ordinal2(X1,X2)))).#trainneg
cnf(i_0_368, plain, (r1_ordinal1(esk15_0,k3_ordinal2(X1))|r1_ordinal1(k3_ordinal2(X1),esk15_0))).#trainneg
cnf(i_0_3649, negated_conjecture, (r2_hidden(k1_xboole_0,esk15_0)|esk2_0=esk15_0|r2_hidden(esk15_0,esk2_0))).#trainneg
cnf(i_0_3358, negated_conjecture, (r1_ordinal1(esk1_0,esk14_0)|~r2_hidden(esk1_0,esk14_0))).#trainneg
cnf(i_0_3625, negated_conjecture, (r1_ordinal1(esk1_0,esk13_0)|~r2_hidden(esk1_0,esk13_0))).#trainneg
cnf(i_0_9232, negated_conjecture, (r2_hidden(esk1_0,k10_ordinal2(esk1_0,esk1_0))|r2_hidden(esk1_0,X1)|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_369, plain, (r1_ordinal1(esk15_0,k2_xboole_0(X1,k1_tarski(X1)))|r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),esk15_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_370, plain, (r1_ordinal1(esk15_0,k2_xboole_0(X1,X2))|r1_ordinal1(k2_xboole_0(X1,X2),esk15_0)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_9041, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),k1_tarski(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(esk6_1(k2_xboole_0(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),k1_tarski(esk6_1(k2_xboole_0(X2,k1_tarski(X2)))))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_9874, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_tarski(k2_xboole_0(X1,k1_tarski(X1)),X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_8850, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v1_xboole_0(X2)|~r2_hidden(X2,X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~r1_tarski(k2_xboole_0(X1,k1_tarski(X1)),X2)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_9523, plain, (v1_xboole_0(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_4860, plain, (~r2_hidden(k1_xboole_0,X2)|~v4_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r1_ordinal1(X2,k1_xboole_0)|~r2_hidden(X1,X1)|~v3_ordinal1(k10_ordinal2(X1,X2)))).#trainneg
cnf(i_0_371, plain, (r1_ordinal1(esk15_0,k4_ordinal2(X1))|r1_ordinal1(k4_ordinal2(X1),esk15_0)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v5_ordinal1(X1))).#trainneg
cnf(i_0_4068, negated_conjecture, (esk12_0=esk1_0|r2_hidden(esk12_0,esk1_0)|r1_ordinal1(esk1_0,esk12_0))).#trainneg
cnf(i_0_4064, negated_conjecture, (r1_ordinal1(esk1_0,esk12_0)|~r2_hidden(esk1_0,esk12_0))).#trainneg
cnf(i_0_4750, plain, (r1_ordinal1(esk14_0,esk13_0)|~r2_hidden(esk14_0,esk13_0))).#trainneg
cnf(i_0_1181, plain, (r2_hidden(k1_xboole_0,k10_ordinal2(k1_xboole_0,X1))|r2_hidden(X1,X2)|v4_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2721, plain, (r1_ordinal1(X1,k10_ordinal2(X2,k1_xboole_0))|r1_ordinal1(k10_ordinal2(X2,k1_xboole_0),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_372, plain, (r1_ordinal1(esk15_0,k10_ordinal2(X1,X2))|r1_ordinal1(k10_ordinal2(X1,X2),esk15_0)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_392, negated_conjecture, (r1_ordinal1(esk2_0,esk10_1(X1))|r1_ordinal1(esk10_1(X1),esk2_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_10210, negated_conjecture, (r2_hidden(X1,k10_ordinal2(X1,X2))|r2_hidden(X2,esk1_0)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_9099, plain, (v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k10_ordinal2(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1),X2)|~v4_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_8839, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k10_ordinal2(X2,k3_ordinal2(X3)),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~r1_ordinal1(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(esk5_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_9358, plain, (v1_xboole_0(esk5_2(X1,X2))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|v4_ordinal1(k2_xboole_0(X3,k1_tarski(X3)))|~r2_hidden(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X3)|~r2_hidden(esk6_1(k2_xboole_0(X3,k1_tarski(X3))),X1)|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_10221, negated_conjecture, (r2_hidden(X1,k10_ordinal2(X2,X3))|r2_hidden(X3,esk1_0)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_10603, plain, (v1_xboole_0(esk5_2(X1,X2))|v4_ordinal1(k2_xboole_0(X3,k1_tarski(X3)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|v1_xboole_0(X3)|~r1_ordinal1(X1,X2)|~r2_hidden(esk6_1(k2_xboole_0(X3,k1_tarski(X3))),X1)|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r1_tarski(k2_xboole_0(X2,k1_tarski(X2)),X3)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2))))).#trainneg
cnf(i_0_10531, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(esk5_2(X2,X1))|~v4_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(esk5_2(X2,X1)))|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~r1_ordinal1(X2,X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_10510, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(X4,X1)|~v4_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(esk7_3(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),X3,X4))|X3!=k10_xtuple_0(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))|~r2_hidden(X4,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_1223, plain, (r2_hidden(k10_ordinal2(esk12_0,esk12_0),k10_ordinal2(k10_ordinal2(esk12_0,esk12_0),X1))|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_716, plain, (r1_ordinal1(esk13_0,k1_xboole_0)|r1_ordinal1(k1_xboole_0,esk13_0))).#trainneg
cnf(i_0_393, negated_conjecture, (r1_ordinal1(esk2_0,esk6_1(X1))|r1_ordinal1(esk6_1(X1),esk2_0)|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_5076, plain, (r1_ordinal1(esk14_0,esk12_0)|~r2_hidden(esk14_0,esk12_0))).#trainneg
cnf(i_0_5493, plain, (esk15_0=k1_xboole_0|~r2_hidden(esk15_0,k1_xboole_0))).#trainneg
cnf(i_0_10326, negated_conjecture, (r2_hidden(k1_xboole_0,k10_ordinal2(k1_xboole_0,esk2_0))|r2_hidden(esk2_0,X1)|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_10327, negated_conjecture, (r2_hidden(k1_xboole_0,k10_ordinal2(k1_xboole_0,esk1_0))|r2_hidden(esk1_0,X1)|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_394, negated_conjecture, (r1_ordinal1(esk2_0,esk5_2(X1,X2))|r1_ordinal1(esk5_2(X1,X2),esk2_0)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_395, negated_conjecture, (r1_ordinal1(esk2_0,esk3_3(X1,X2,X3))|r1_ordinal1(esk3_3(X1,X2,X3),esk2_0)|k10_ordinal2(X2,X1)=k4_ordinal2(X3)|X1=k1_xboole_0|k9_xtuple_0(X3)!=X1|~v1_ordinal2(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v5_ordinal1(X3)|~v4_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_10184, plain, (v1_xboole_0(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(X2,X1)|~v4_ordinal1(X1))).#trainneg
cnf(i_0_8770, plain, (v4_ordinal1(k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v4_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(k10_ordinal2(k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))),k1_xboole_0),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_9525, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k2_xboole_0(k2_xboole_0(X2,k1_tarski(X2)),k1_tarski(k2_xboole_0(X2,k1_tarski(X2)))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),k1_tarski(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))))|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_9391, plain, (v1_xboole_0(X1)|v4_ordinal1(k2_xboole_0(k10_ordinal2(X2,X1),k1_tarski(k10_ordinal2(X2,X1))))|v4_ordinal1(k2_xboole_0(k10_ordinal2(X3,X4),k1_tarski(k10_ordinal2(X3,X4))))|v1_xboole_0(X4)|~r2_hidden(esk6_1(k2_xboole_0(k10_ordinal2(X3,X4),k1_tarski(k10_ordinal2(X3,X4)))),X2)|~v3_ordinal1(k10_ordinal2(X3,X4))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k2_xboole_0(k10_ordinal2(X2,X1),k1_tarski(k10_ordinal2(X2,X1)))),X3)|~v3_ordinal1(X3)|~v3_ordinal1(X4))).#trainneg
cnf(i_0_10926, plain, (esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)=k1_xboole_0|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~r2_hidden(X2,X1)|~v4_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_10967, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~r2_hidden(k2_xboole_0(k2_xboole_0(X2,k1_tarski(X2)),k1_tarski(k2_xboole_0(X2,k1_tarski(X2)))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))).#trainneg
cnf(i_0_1726, plain, (v7_ordinal1(X1)|~v7_ordinal1(X2)|~r1_ordinal1(X2,X1)|~v3_ordinal1(X1)|~v1_xboole_0(esk5_2(X2,X1)))).#trainneg
cnf(i_0_909, plain, (r1_ordinal1(k3_ordinal2(X1),k3_ordinal2(X2))|r1_ordinal1(k3_ordinal2(X2),k3_ordinal2(X1)))).#trainneg
cnf(i_0_10766, plain, (X1!=k10_xtuple_0(esk4_2(X2,esk6_1(k1_tarski(k1_xboole_0))))|~r2_hidden(X3,k1_xboole_0)|~r2_hidden(X3,X1)|~v4_ordinal1(k1_xboole_0)|~v3_ordinal1(esk6_1(k1_tarski(k1_xboole_0)))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_398, negated_conjecture, (r1_ordinal1(esk2_0,k3_ordinal2(X1))|r1_ordinal1(k3_ordinal2(X1),esk2_0))).#trainneg
cnf(i_0_2719, plain, (r2_hidden(k1_xboole_0,k10_ordinal2(X1,k1_xboole_0))|r2_hidden(k10_ordinal2(X1,k1_xboole_0),X2)|v4_ordinal1(X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_399, negated_conjecture, (r1_ordinal1(esk2_0,k2_xboole_0(X1,k1_tarski(X1)))|r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),esk2_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_400, negated_conjecture, (r1_ordinal1(esk2_0,k2_xboole_0(X1,X2))|r1_ordinal1(k2_xboole_0(X1,X2),esk2_0)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_10750, plain, (v1_xboole_0(esk5_2(k2_xboole_0(X1,k1_tarski(X1)),X2))|v1_xboole_0(X1)|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),X2)|~r1_tarski(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_10977, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~r2_hidden(k2_xboole_0(k2_xboole_0(X2,k1_tarski(X2)),k1_tarski(k2_xboole_0(X2,k1_tarski(X2)))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)))).#trainneg
cnf(i_0_11156, negated_conjecture, (r1_ordinal1(esk2_0,k3_ordinal2(X1))|~r2_hidden(esk2_0,k3_ordinal2(X1)))).#trainneg
cnf(i_0_10764, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(X2,X1)|~r2_hidden(X2,k10_xtuple_0(esk4_2(X3,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))))|~v4_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_10751, plain, (v1_xboole_0(esk5_2(X1,X2))|v1_xboole_0(X3)|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|v4_ordinal1(k2_xboole_0(X3,k1_tarski(X3)))|v1_xboole_0(X1)|~r1_ordinal1(X1,X2)|~r1_tarski(k2_xboole_0(X2,k1_tarski(X2)),X3)|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r1_tarski(k2_xboole_0(X3,k1_tarski(X3)),X1)|~v3_ordinal1(k2_xboole_0(X3,k1_tarski(X3))))).#trainneg
cnf(i_0_8833, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|r2_hidden(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(k2_xboole_0(k2_xboole_0(X2,k1_tarski(X2)),k1_tarski(k2_xboole_0(X2,k1_tarski(X2)))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(k2_xboole_0(X2,k1_tarski(X2)),k1_tarski(k2_xboole_0(X2,k1_tarski(X2)))))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_11103, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~r2_hidden(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~r2_hidden(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v4_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2))))).#trainneg
cnf(i_0_9098, plain, (v1_xboole_0(esk6_1(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1)))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1)))|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1))).#trainneg
cnf(i_0_9822, plain, (v1_xboole_0(esk5_2(X1,X2))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|v4_ordinal1(k2_xboole_0(X3,k1_tarski(X3)))|~r1_ordinal1(esk6_1(k2_xboole_0(X3,k1_tarski(X3))),X1)|~r2_hidden(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X3)|~v3_ordinal1(esk6_1(k2_xboole_0(X3,k1_tarski(X3))))|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_964, plain, (m1_subset_1(X1,X2)|~r1_tarski(k10_ordinal2(X1,k1_tarski(k1_xboole_0)),X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_11112, plain, (r1_ordinal1(k3_ordinal2(X3),k3_ordinal2(X3)))).#trainneg
cnf(i_0_401, negated_conjecture, (r1_ordinal1(esk2_0,k4_ordinal2(X1))|r1_ordinal1(k4_ordinal2(X1),esk2_0)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v5_ordinal1(X1))).#trainneg
cnf(i_0_4754, plain, (esk13_0=esk14_0|r2_hidden(esk13_0,esk14_0)|r1_ordinal1(esk14_0,esk13_0))).#trainneg
cnf(i_0_8900, plain, (r1_ordinal1(esk12_0,esk13_0)|~r2_hidden(esk12_0,esk13_0))).#trainneg
cnf(i_0_402, negated_conjecture, (r1_ordinal1(esk2_0,k10_ordinal2(X1,X2))|r1_ordinal1(k10_ordinal2(X1,X2),esk2_0)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_408, negated_conjecture, (r1_ordinal1(esk1_0,esk10_1(X1))|r1_ordinal1(esk10_1(X1),esk1_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_11329, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v4_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2))))).#trainneg
cnf(i_0_11294, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_xboole_0)|~r2_hidden(k2_xboole_0(k1_tarski(k1_xboole_0),k1_tarski(k1_tarski(k1_xboole_0))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_11378, plain, (v1_xboole_0(esk5_2(X1,X2))|v1_xboole_0(X1)|v1_xboole_0(X3)|v4_ordinal1(k2_xboole_0(X3,k1_tarski(X3)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_ordinal1(X1,X2)|~r1_tarski(k2_xboole_0(X3,k1_tarski(X3)),X1)|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r1_ordinal1(k2_xboole_0(X2,k1_tarski(X2)),X3)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2))))).#trainneg
cnf(i_0_11424, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(k2_xboole_0(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),k1_tarski(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))),X1)|~v4_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_11291, plain, (v1_xboole_0(esk5_2(k2_xboole_0(X1,k1_tarski(X1)),X2))|v1_xboole_0(X1)|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r1_ordinal1(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2))))).#trainneg
cnf(i_0_11292, plain, (v1_xboole_0(esk5_2(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))),X1))|v4_ordinal1(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))),X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_10819, negated_conjecture, (r1_ordinal1(esk2_0,esk6_1(X1))|v4_ordinal1(X1)|~r2_hidden(esk2_0,esk6_1(X1))|~v3_ordinal1(esk6_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_409, negated_conjecture, (r1_ordinal1(esk1_0,esk6_1(X1))|r1_ordinal1(esk6_1(X1),esk1_0)|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_10802, plain, (r1_ordinal1(k1_xboole_0,esk13_0)|~r2_hidden(k1_xboole_0,esk13_0))).#trainneg
cnf(i_0_410, negated_conjecture, (r1_ordinal1(esk1_0,esk5_2(X1,X2))|r1_ordinal1(esk5_2(X1,X2),esk1_0)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_411, negated_conjecture, (r1_ordinal1(esk1_0,esk3_3(X1,X2,X3))|r1_ordinal1(esk3_3(X1,X2,X3),esk1_0)|k10_ordinal2(X2,X1)=k4_ordinal2(X3)|X1=k1_xboole_0|k9_xtuple_0(X3)!=X1|~v1_ordinal2(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v5_ordinal1(X3)|~v4_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_8999, plain, (r2_hidden(k10_ordinal2(X1,esk6_1(esk7_3(esk4_2(X2,X1),X3,X4))),X4)|v4_ordinal1(esk7_3(esk4_2(X2,X1),X3,X4))|X3!=k10_xtuple_0(esk4_2(X2,X1))|~r2_hidden(X4,X3)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(esk7_3(esk4_2(X2,X1),X3,X4)))).#trainneg
cnf(i_0_11833, plain, (esk6_1(esk7_3(esk4_2(X1,X2),X3,X2))=k1_xboole_0|v4_ordinal1(esk7_3(esk4_2(X1,X2),X3,X2))|~v3_ordinal1(X2)|~v3_ordinal1(esk6_1(esk7_3(esk4_2(X1,X2),X3,X2)))|X3!=k10_xtuple_0(esk4_2(X1,X2))|~r2_hidden(X2,X3)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_11708, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|r1_ordinal1(esk1_0,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(k2_xboole_0(esk1_0,k1_tarski(esk1_0)),X1)|~v4_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_11814, plain, (v4_ordinal1(esk7_3(esk4_2(X2,X1),X3,X4))|~r2_hidden(X4,k10_ordinal2(X1,esk6_1(esk7_3(esk4_2(X2,X1),X3,X4))))|X3!=k10_xtuple_0(esk4_2(X2,X1))|~r2_hidden(X4,X3)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_11623, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k2_xboole_0(X2,k1_tarski(X2))),X1)|~r2_hidden(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v4_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_11583, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|r1_ordinal1(esk2_0,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(k2_xboole_0(esk2_0,k1_tarski(esk2_0)),X1)|~v4_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_8106, plain, (~r2_hidden(X2,X1)|~r2_hidden(k1_xboole_0,esk5_2(X1,X2))|~v4_ordinal1(esk5_2(X1,X2))|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_414, negated_conjecture, (r1_ordinal1(esk1_0,k3_ordinal2(X1))|r1_ordinal1(k3_ordinal2(X1),esk1_0))).#trainneg
cnf(i_0_5080, plain, (esk12_0=esk14_0|r2_hidden(esk12_0,esk14_0)|r1_ordinal1(esk14_0,esk12_0))).#trainneg
cnf(i_0_11744, negated_conjecture, (~r2_hidden(esk1_0,k1_tarski(k1_xboole_0)))).#trainneg
cnf(i_0_11184, negated_conjecture, (r2_hidden(k1_xboole_0,k10_ordinal2(X1,k1_xboole_0))|r2_hidden(k10_ordinal2(X1,k1_xboole_0),esk2_0)|v4_ordinal1(esk2_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_11213, plain, (r2_hidden(k1_xboole_0,k10_ordinal2(X1,k1_xboole_0))|r2_hidden(k10_ordinal2(X1,k1_xboole_0),k1_xboole_0)|v4_ordinal1(k1_xboole_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_415, negated_conjecture, (r1_ordinal1(esk1_0,k2_xboole_0(X1,k1_tarski(X1)))|r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),esk1_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_416, negated_conjecture, (r1_ordinal1(esk1_0,k2_xboole_0(X1,X2))|r1_ordinal1(k2_xboole_0(X1,X2),esk1_0)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_12077, negated_conjecture, (r1_ordinal1(esk1_0,k3_ordinal2(X1))|~r2_hidden(esk1_0,k3_ordinal2(X1)))).#trainneg
cnf(i_0_12002, negated_conjecture, (~r2_hidden(k10_ordinal2(X2,esk1_0),X1)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_11845, plain, (esk6_1(esk7_3(esk4_2(X1,X2),X3,X2))=k1_xboole_0|v4_ordinal1(esk7_3(esk4_2(X1,X2),X3,X2))|X3!=k10_xtuple_0(esk4_2(X1,X2))|~r2_hidden(X2,X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(esk7_3(esk4_2(X1,X2),X3,X2)))).#trainneg
cnf(i_0_11906, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(k10_ordinal2(X2,esk1_0),X1)|~v3_ordinal1(k10_ordinal2(X2,esk1_0))|~v3_ordinal1(X1)|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_11893, plain, (v4_ordinal1(esk7_3(esk4_2(X1,X2),X3,X2))|v1_xboole_0(esk6_1(esk7_3(esk4_2(X1,X2),X3,X2)))|X3!=k10_xtuple_0(esk4_2(X1,X2))|~r2_hidden(X2,X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(esk7_3(esk4_2(X1,X2),X3,X2))))).#trainneg
cnf(i_0_11908, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(k3_ordinal2(k10_ordinal2(X2,esk1_0)),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_11992, negated_conjecture, (v1_xboole_0(k10_ordinal2(X1,X2))|~r2_hidden(esk1_0,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r1_ordinal1(k10_ordinal2(X1,X2),X1)|~v3_ordinal1(k10_ordinal2(X1,X2)))).#trainneg
cnf(i_0_12217, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|r1_ordinal1(esk1_0,k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(esk1_0,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_12304, negated_conjecture, (~r1_ordinal1(k10_ordinal2(X1,esk1_0),X2)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X2)|~r1_ordinal1(k10_ordinal2(X2,esk1_0),X1)|~v3_ordinal1(k10_ordinal2(X2,esk1_0))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_6482, plain, (r2_hidden(k1_xboole_0,X1)|v4_ordinal1(k1_xboole_0)|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k1_xboole_0)))).#trainneg
cnf(i_0_2807, plain, (v7_ordinal1(k10_ordinal2(X1,k1_tarski(k1_xboole_0)))|~v7_ordinal1(k10_ordinal2(X1,k1_xboole_0))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_417, negated_conjecture, (r1_ordinal1(esk1_0,k4_ordinal2(X1))|r1_ordinal1(k4_ordinal2(X1),esk1_0)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v5_ordinal1(X1))).#trainneg
cnf(i_0_5575, negated_conjecture, (esk2_0=k1_xboole_0|r2_hidden(esk15_0,esk2_0)|r2_hidden(k1_xboole_0,esk15_0))).#trainneg
cnf(i_0_12123, negated_conjecture, (r2_hidden(k10_ordinal2(esk2_0,k1_xboole_0),esk2_0)|r2_hidden(k1_xboole_0,k10_ordinal2(esk2_0,k1_xboole_0))|v4_ordinal1(esk2_0))).#trainneg
cnf(i_0_12124, negated_conjecture, (r2_hidden(k10_ordinal2(esk1_0,k1_xboole_0),esk2_0)|r2_hidden(k1_xboole_0,k10_ordinal2(esk1_0,k1_xboole_0))|v4_ordinal1(esk2_0))).#trainneg
cnf(i_0_418, negated_conjecture, (r1_ordinal1(esk1_0,k10_ordinal2(X1,X2))|r1_ordinal1(k10_ordinal2(X1,X2),esk1_0)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_427, plain, (k2_xboole_0(X1,k1_tarski(X1))=k10_xtuple_0(X2)|r2_hidden(X1,esk8_2(X2,k2_xboole_0(X1,k1_tarski(X1))))|esk8_2(X2,k2_xboole_0(X1,k1_tarski(X1)))!=k1_funct_1(X2,X3)|~r2_hidden(X3,k9_xtuple_0(X2))|~v1_funct_1(X2)|~v1_relat_1(X2)|~v3_ordinal1(esk8_2(X2,k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_12320, negated_conjecture, (~r1_ordinal1(k10_ordinal2(X1,X2),X1)|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X1)|~r2_hidden(esk1_0,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_12312, negated_conjecture, (v1_xboole_0(X2)|~r1_ordinal1(k10_ordinal2(k10_ordinal2(X1,esk1_0),X2),X1)|~v3_ordinal1(k10_ordinal2(k10_ordinal2(X1,esk1_0),X2))|~v3_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_11826, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(esk7_3(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),X3,X1))|~v4_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(esk7_3(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),X3,X1)))|X3!=k10_xtuple_0(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))|~r2_hidden(X1,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_11891, plain, (v4_ordinal1(esk7_3(esk4_2(X1,X2),X3,X4))|v1_xboole_0(esk6_1(esk7_3(esk4_2(X1,X2),X3,X4)))|X3!=k10_xtuple_0(esk4_2(X1,X2))|~r2_hidden(X4,X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X4,X2)|~v3_ordinal1(esk6_1(esk7_3(esk4_2(X1,X2),X3,X4))))).#trainneg
cnf(i_0_11834, plain, (v1_xboole_0(k10_ordinal2(X1,esk6_1(esk7_3(esk4_2(X2,X1),X3,X4))))|v4_ordinal1(esk7_3(esk4_2(X2,X1),X3,X4))|~r1_ordinal1(X1,X4)|~r2_hidden(k1_xboole_0,esk6_1(esk7_3(esk4_2(X2,X1),X3,X4)))|~r2_hidden(X4,X1)|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(esk7_3(esk4_2(X2,X1),X3,X4)))|X3!=k10_xtuple_0(esk4_2(X2,X1))|~r2_hidden(X4,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_11823, plain, (v1_xboole_0(esk6_1(esk7_3(esk4_2(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2)))),X3,X2)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|v4_ordinal1(esk7_3(esk4_2(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2)))),X3,X2))|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2)|~v3_ordinal1(esk6_1(esk7_3(esk4_2(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2)))),X3,X2)))|X3!=k10_xtuple_0(esk4_2(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2)))))|~r2_hidden(X2,X3)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_6814, plain, (r2_hidden(X1,X1)|v4_ordinal1(k1_xboole_0)|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k1_xboole_0)))).#trainneg
cnf(i_0_820, plain, (m1_subset_1(X1,k10_ordinal2(X1,X2))|~r2_hidden(k1_xboole_0,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_9931, plain, (~r1_ordinal1(X1,k1_xboole_0)|~r2_hidden(k1_xboole_0,X1)|~v4_ordinal1(X1)|~v3_ordinal1(k1_funct_1(X2,X3))|~v3_ordinal1(X1)|k1_funct_1(X2,X3)!=k10_xtuple_0(X2)|~r2_hidden(X3,k9_xtuple_0(X2))|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_428, plain, (r2_hidden(X1,k10_ordinal2(X2,k2_xboole_0(X3,k1_tarski(X3))))|r2_hidden(X3,esk5_2(X2,X1))|~r1_ordinal1(X2,X1)|~v3_ordinal1(X2)|~v3_ordinal1(k2_xboole_0(X3,k1_tarski(X3)))|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X2,X1))|~v3_ordinal1(X3))).#trainneg
cnf(i_0_5577, negated_conjecture, (esk2_0=k1_xboole_0|r2_hidden(esk13_0,esk2_0)|r2_hidden(k1_xboole_0,esk13_0))).#trainneg
cnf(i_0_12425, plain, (r2_hidden(k1_xboole_0,X1)|v4_ordinal1(k1_xboole_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_12594, plain, (r2_hidden(X1,X1)|v4_ordinal1(k1_xboole_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_429, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|r2_hidden(X1,k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))))|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_430, plain, (v4_ordinal1(k3_ordinal2(k2_xboole_0(X1,k1_tarski(X1))))|r2_hidden(X1,k2_xboole_0(esk6_1(k3_ordinal2(k2_xboole_0(X1,k1_tarski(X1)))),k1_tarski(esk6_1(k3_ordinal2(k2_xboole_0(X1,k1_tarski(X1)))))))|~v3_ordinal1(k2_xboole_0(esk6_1(k3_ordinal2(k2_xboole_0(X1,k1_tarski(X1)))),k1_tarski(esk6_1(k3_ordinal2(k2_xboole_0(X1,k1_tarski(X1)))))))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_12532, negated_conjecture, (~r2_hidden(esk1_0,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(k10_ordinal2(X1,X2),X1))).#trainneg
cnf(i_0_12423, negated_conjecture, (~r1_ordinal1(k10_ordinal2(X1,esk1_0),X1)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)|~r2_hidden(k10_ordinal2(X1,esk1_0),X1))).#trainneg
cnf(i_0_12526, negated_conjecture, (~r2_hidden(esk1_0,k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(k10_ordinal2(X1,X2),X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_12586, negated_conjecture, (v4_ordinal1(k10_ordinal2(X1,esk1_0))|r2_hidden(X1,k2_xboole_0(esk6_1(k10_ordinal2(X1,esk1_0)),k1_tarski(esk6_1(k10_ordinal2(X1,esk1_0)))))|~v3_ordinal1(k2_xboole_0(esk6_1(k10_ordinal2(X1,esk1_0)),k1_tarski(esk6_1(k10_ordinal2(X1,esk1_0)))))|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_12995, negated_conjecture, (~r2_hidden(esk1_0,X3)|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X3)|~r1_ordinal1(k10_ordinal2(X1,X2),X1)|~r2_hidden(X3,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_12997, negated_conjecture, (v4_ordinal1(esk7_3(esk4_2(X2,X1),X3,X1))|~r2_hidden(esk1_0,esk6_1(esk7_3(esk4_2(X2,X1),X3,X1)))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(esk7_3(esk4_2(X2,X1),X3,X1)))|X3!=k10_xtuple_0(esk4_2(X2,X1))|~r2_hidden(X1,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_12757, plain, (r2_hidden(k1_xboole_0,k3_ordinal2(X1))|v4_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_12784, plain, (v4_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_12605, plain, (k1_funct_1(esk4_2(X1,X2),X3)!=k10_xtuple_0(esk4_2(X1,X2))|~r1_ordinal1(X4,k1_xboole_0)|~r2_hidden(X3,k9_xtuple_0(esk4_2(X1,X2)))|~r2_hidden(k1_xboole_0,X4)|~v1_funct_1(esk4_2(X1,X2))|~v1_relat_1(esk4_2(X1,X2))|~v4_ordinal1(X4)|~v3_ordinal1(X4)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_6161, negated_conjecture, (r2_hidden(k1_xboole_0,esk2_0)|esk13_0=esk2_0|r2_hidden(esk2_0,esk13_0))).#trainneg
cnf(i_0_12152, negated_conjecture, (r2_hidden(k10_ordinal2(k1_xboole_0,k1_xboole_0),esk2_0)|r2_hidden(k1_xboole_0,k10_ordinal2(k1_xboole_0,k1_xboole_0))|v4_ordinal1(esk2_0))).#trainneg
cnf(i_0_11547, negated_conjecture, (v1_xboole_0(X1)|v1_xboole_0(esk2_0)|v4_ordinal1(k10_ordinal2(esk2_0,X1))|r1_ordinal1(esk2_0,k10_ordinal2(esk2_0,X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_439, plain, (r2_hidden(X1,k10_ordinal2(k2_xboole_0(X1,k1_tarski(X1)),X2))|~v3_ordinal1(k10_ordinal2(k2_xboole_0(X1,k1_tarski(X1)),X2))|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_13018, negated_conjecture, (~r2_hidden(k10_ordinal2(X1,esk1_0),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_13049, negated_conjecture, (v4_ordinal1(k10_ordinal2(X1,esk1_0))|v1_xboole_0(X1)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)|~r1_tarski(k10_ordinal2(X1,esk1_0),X1))).#trainneg
cnf(i_0_12979, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(esk1_0,X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X2)|~r2_hidden(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_12415, negated_conjecture, (v4_ordinal1(k10_ordinal2(X1,esk1_0))|~r1_ordinal1(k10_ordinal2(esk6_1(k10_ordinal2(X1,esk1_0)),esk1_0),X1)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(k10_ordinal2(esk6_1(k10_ordinal2(X1,esk1_0)),esk1_0))|~v3_ordinal1(esk6_1(k10_ordinal2(X1,esk1_0)))|~v3_ordinal1(X1)|~r1_ordinal1(X1,esk6_1(k10_ordinal2(X1,esk1_0))))).#trainneg
cnf(i_0_12562, plain, (v4_ordinal1(esk7_3(esk4_2(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2)))),X3,X2))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|X3!=k10_xtuple_0(esk4_2(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2)))))|~r2_hidden(X2,X3)|~v4_ordinal1(X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X1)|~v3_ordinal1(esk7_3(esk4_2(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2)))),X3,X2)))).#trainneg
cnf(i_0_12552, negated_conjecture, (~v3_ordinal1(k10_ordinal2(X2,esk1_0))|~v3_ordinal1(X2)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k10_ordinal2(k10_ordinal2(X2,esk1_0),X1),X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_12973, negated_conjecture, (~r2_hidden(X4,X1)|~r2_hidden(esk1_0,esk7_3(esk4_2(X2,X1),X3,X4))|~v3_ordinal1(X1)|~v3_ordinal1(esk7_3(esk4_2(X2,X1),X3,X4))|X3!=k10_xtuple_0(esk4_2(X2,X1))|~r2_hidden(X4,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_3805, plain, (v4_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(esk6_1(k10_ordinal2(X1,X2)),X1)|~r2_hidden(k1_tarski(k1_xboole_0),X2)|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,X2)))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_4455, plain, (r1_ordinal1(esk12_0,k1_xboole_0)|r1_ordinal1(k1_xboole_0,esk12_0))).#trainneg
cnf(i_0_10766, plain, (X1!=k10_xtuple_0(esk4_2(X2,esk6_1(k1_tarski(k1_xboole_0))))|~r2_hidden(X3,k1_xboole_0)|~r2_hidden(X3,X1)|~v3_ordinal1(esk6_1(k1_tarski(k1_xboole_0)))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_440, plain, (r2_hidden(k10_ordinal2(X1,X2),k10_ordinal2(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))),X3))|~v3_ordinal1(k10_ordinal2(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))),X3))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X3))).#trainneg
cnf(i_0_6541, negated_conjecture, (r2_hidden(k1_xboole_0,esk2_0)|esk12_0=esk2_0|r2_hidden(esk2_0,esk12_0))).#trainneg
cnf(i_0_5531, plain, (~r2_hidden(k10_ordinal2(k1_tarski(k1_xboole_0),k1_xboole_0),k1_xboole_0))).#trainneg
cnf(i_0_8272, plain, (X1=k10_ordinal2(X2,k1_xboole_0)|r2_hidden(k1_xboole_0,X1)|r2_hidden(k1_xboole_0,k10_ordinal2(X2,k1_xboole_0))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_444, plain, (r1_ordinal1(X1,k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_446, plain, (r1_ordinal1(X1,k10_ordinal2(X2,k10_ordinal2(esk5_2(X2,X1),X3)))|~r1_ordinal1(X2,X1)|~v3_ordinal1(X2)|~v3_ordinal1(k10_ordinal2(esk5_2(X2,X1),X3))|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X2,X1))|~v3_ordinal1(X3))).#trainneg
cnf(i_0_13324, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(esk5_2(X2,X1))|~r2_hidden(esk1_0,esk6_1(esk5_2(X2,X1)))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(esk6_1(esk5_2(X2,X1)))|~v3_ordinal1(X1)|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~r1_ordinal1(X2,X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_13326, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1))|~r2_hidden(esk1_0,esk6_1(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1)))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(esk6_1(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1)))|~v3_ordinal1(X1)|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1))).#trainneg
cnf(i_0_13524, plain, (r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_13297, negated_conjecture, (v1_xboole_0(X1)|v4_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)|~r1_ordinal1(k10_ordinal2(X1,esk1_0),X1))).#trainneg
cnf(i_0_13523, plain, (r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k2_xboole_0(X2,k1_tarski(X2)))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X2)|~r2_hidden(X2,X1)|~v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_13533, plain, (r1_ordinal1(X1,k2_xboole_0(X2,k1_tarski(X2)))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_13274, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),esk1_0),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_2320, plain, (v3_ordinal1(k10_ordinal2(X1,esk6_1(X2)))|v4_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_13399, plain, (r1_ordinal1(k1_xboole_0,esk12_0))).#trainneg
cnf(i_0_7244, plain, (~r2_hidden(k10_ordinal2(esk6_1(k1_tarski(k1_xboole_0)),k1_xboole_0),k1_xboole_0)|~v3_ordinal1(esk6_1(k1_tarski(k1_xboole_0))))).#trainneg
cnf(i_0_447, plain, (r1_ordinal1(X1,k10_ordinal2(X2,esk5_2(X2,X1)))|~r1_ordinal1(X2,X1)|~v3_ordinal1(X2)|~v3_ordinal1(esk5_2(X2,X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_6909, plain, (r2_hidden(k1_xboole_0,esk13_0)|esk13_0=esk15_0|r2_hidden(k1_xboole_0,esk15_0))).#trainneg
cnf(i_0_1548, negated_conjecture, (r2_hidden(X1,k10_ordinal2(X1,esk1_0))|esk2_0=esk1_0|r2_hidden(esk1_0,esk2_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_448, plain, (r1_ordinal1(X1,k10_ordinal2(X2,X3))|r2_hidden(X3,esk5_2(X2,X1))|~r1_ordinal1(X2,X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X2,X1)))).#trainneg
cnf(i_0_449, plain, (r1_ordinal1(k10_ordinal2(X1,k10_ordinal2(X2,esk5_2(X3,esk5_2(X1,X4)))),X4)|~r1_ordinal1(X1,X4)|~v3_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(X2,esk5_2(X3,esk5_2(X1,X4))))|~v3_ordinal1(X4)|~r1_ordinal1(X2,X3)|~r1_ordinal1(X3,esk5_2(X1,X4))|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(esk5_2(X1,X4)))).#trainneg
cnf(i_0_13353, negated_conjecture, (~r2_hidden(X2,X1)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(k10_ordinal2(X1,esk1_0),X2))|~r1_ordinal1(k10_ordinal2(X1,esk1_0),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_14120, plain, (r1_ordinal1(X3,k10_ordinal2(X2,X1))|~r2_hidden(esk5_2(X2,X3),X1)|~r1_ordinal1(X2,X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_13281, negated_conjecture, (~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~r1_ordinal1(k10_ordinal2(k10_ordinal2(X1,esk1_0),esk1_0),X1)|~v3_ordinal1(k10_ordinal2(k10_ordinal2(X1,esk1_0),esk1_0))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_13003, negated_conjecture, (~r2_hidden(esk1_0,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r1_ordinal1(X1,X1)|~r2_hidden(X2,k1_xboole_0)|~r2_hidden(X1,X1))).#trainneg
cnf(i_0_13109, negated_conjecture, (v4_ordinal1(k10_ordinal2(X1,esk1_0))|~r2_hidden(k2_xboole_0(esk6_1(k10_ordinal2(X1,esk1_0)),k1_tarski(esk6_1(k10_ordinal2(X1,esk1_0)))),X1)|~v3_ordinal1(k2_xboole_0(esk6_1(k10_ordinal2(X1,esk1_0)),k1_tarski(esk6_1(k10_ordinal2(X1,esk1_0)))))|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_12729, plain, (r2_hidden(X1,k10_ordinal2(X2,k2_xboole_0(X3,k1_tarski(X3))))|r2_hidden(X3,esk5_2(X2,X1))|r2_hidden(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_12986, negated_conjecture, (~r2_hidden(esk1_0,X2)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X2)|~r1_ordinal1(k10_ordinal2(k10_ordinal2(X1,esk1_0),X2),X1)|~v3_ordinal1(k10_ordinal2(k10_ordinal2(X1,esk1_0),X2))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_12576, plain, (esk6_1(esk7_3(esk4_2(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2)))),X3,X2))=k1_xboole_0|v4_ordinal1(esk7_3(esk4_2(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2)))),X3,X2))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|X3!=k10_xtuple_0(esk4_2(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2)))))|~r2_hidden(X2,X3)|~v3_ordinal1(esk6_1(esk7_3(esk4_2(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2)))),X3,X2)))|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_13781, plain, (r1_ordinal1(X1,k1_tarski(k1_xboole_0))|~r1_ordinal1(X1,k1_xboole_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_450, plain, (r2_hidden(k10_ordinal2(X1,esk5_2(X2,X3)),X4)|~r2_hidden(X3,X4)|~v3_ordinal1(X4)|~v3_ordinal1(X3)|~v3_ordinal1(k10_ordinal2(X1,esk5_2(X2,X3)))|~r1_ordinal1(X1,X2)|~r1_ordinal1(X2,X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_13784, negated_conjecture, (~r2_hidden(k10_ordinal2(esk6_1(k1_tarski(k1_xboole_0)),esk1_0),k1_xboole_0)|~v3_ordinal1(esk6_1(k1_tarski(k1_xboole_0))))).#trainneg
cnf(i_0_1462, plain, (r2_hidden(k1_xboole_0,k10_ordinal2(X1,X2))|r2_hidden(k10_ordinal2(X1,X2),X3)|v4_ordinal1(X3)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_2714, negated_conjecture, (r1_ordinal1(esk2_0,k10_ordinal2(X1,k1_xboole_0))|r1_ordinal1(k10_ordinal2(X1,k1_xboole_0),esk2_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_452, plain, (r1_ordinal1(k10_ordinal2(X1,k10_ordinal2(X2,k2_xboole_0(X3,k1_tarski(X3)))),X4)|~r1_ordinal1(X1,X4)|~v3_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(X2,k2_xboole_0(X3,k1_tarski(X3))))|~v3_ordinal1(X4)|~r2_hidden(k10_ordinal2(X2,X3),esk5_2(X1,X4))|~v3_ordinal1(esk5_2(X1,X4))|~v3_ordinal1(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_453, plain, (r2_hidden(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))),X3)|~r2_hidden(X4,X3)|~v3_ordinal1(X3)|~v3_ordinal1(X4)|~v3_ordinal1(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))))|~r2_hidden(k10_ordinal2(X1,X2),X4)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_14176, negated_conjecture, (~r1_ordinal1(k10_ordinal2(X1,esk1_0),X2)|~r2_hidden(X2,X1)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_14679, negated_conjecture, (r2_hidden(X2,k10_ordinal2(X1,esk1_0))|~r2_hidden(X2,X1)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_14669, negated_conjecture, (~r2_hidden(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)|~r1_ordinal1(k10_ordinal2(X1,esk1_0),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_13301, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(X4,X1)|~r2_hidden(esk1_0,esk7_3(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),X3,X4))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(esk7_3(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),X3,X4))|~v3_ordinal1(X1)|X3!=k10_xtuple_0(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))|~r2_hidden(X4,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_12544, negated_conjecture, (v1_xboole_0(X1)|r1_ordinal1(esk2_0,k10_ordinal2(k10_ordinal2(esk2_0,esk1_0),X1))|~v3_ordinal1(k10_ordinal2(esk2_0,esk1_0))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_12996, negated_conjecture, (v4_ordinal1(esk5_2(X2,X1))|~r2_hidden(esk1_0,esk6_1(esk5_2(X2,X1)))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(esk5_2(X2,X1)))|~r1_ordinal1(X1,X2)|~r1_ordinal1(X2,X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_486, plain, (r2_hidden(X1,X2)|~r1_ordinal1(k10_ordinal2(X1,k1_tarski(k1_xboole_0)),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_603, plain, (v3_ordinal1(k10_ordinal2(k2_xboole_0(X1,k1_tarski(X1)),X2))|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_455, plain, (r1_ordinal1(k10_ordinal2(X1,k10_ordinal2(X2,k2_xboole_0(X3,k1_tarski(X3)))),X4)|~r2_hidden(k10_ordinal2(X1,k10_ordinal2(X2,X3)),X4)|~v3_ordinal1(X4)|~v3_ordinal1(k10_ordinal2(X2,X3))|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_14511, negated_conjecture, (r2_hidden(k1_xboole_0,k10_ordinal2(X1,X2))|r2_hidden(k10_ordinal2(X1,X2),esk2_0)|v4_ordinal1(esk2_0)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_456, plain, (r1_ordinal1(k10_ordinal2(X1,k1_tarski(k1_xboole_0)),X2)|~r2_hidden(k10_ordinal2(X1,k1_xboole_0),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_457, plain, (r1_ordinal1(k2_xboole_0(X2,k1_tarski(X2)),X3)|~r2_hidden(k10_ordinal2(X1,esk5_2(X1,X2)),X3)|~v3_ordinal1(X3)|~v3_ordinal1(esk5_2(X1,X2))|~v3_ordinal1(X1)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_14705, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k10_ordinal2(X2,esk1_0),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(k10_ordinal2(X2,esk1_0))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_14699, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(k10_ordinal2(X2,esk1_0),X1)|~r2_hidden(X1,k10_ordinal2(X2,esk1_0))|~v3_ordinal1(k10_ordinal2(X2,esk1_0))|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_14746, negated_conjecture, (~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~r2_hidden(k2_xboole_0(k2_xboole_0(k10_ordinal2(X1,esk1_0),k1_tarski(k10_ordinal2(X1,esk1_0))),k1_tarski(k2_xboole_0(k10_ordinal2(X1,esk1_0),k1_tarski(k10_ordinal2(X1,esk1_0))))),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_15184, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v1_xboole_0(X2)|~r2_hidden(k10_ordinal2(X2,esk1_0),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r1_tarski(k2_xboole_0(X1,k1_tarski(X1)),X2)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_15148, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(k10_ordinal2(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),esk1_0),X1)|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_14756, negated_conjecture, (r2_hidden(X1,k10_ordinal2(X2,esk1_0))|~v3_ordinal1(k10_ordinal2(X2,esk1_0))|~v3_ordinal1(X1)|~r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_13302, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(X2,X1)|~r2_hidden(esk1_0,esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))|~v3_ordinal1(X1)|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_12413, negated_conjecture, (r2_hidden(X2,k10_ordinal2(X1,esk1_0))|~r1_ordinal1(k10_ordinal2(X2,esk1_0),X1)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(k10_ordinal2(X2,esk1_0))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_14702, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(k3_ordinal2(k10_ordinal2(X2,esk1_0)),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(k10_ordinal2(X2,esk1_0))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_2799, plain, (esk13_0=esk15_0|r2_hidden(esk15_0,esk13_0)|~v1_xboole_0(esk15_0))).#trainneg
cnf(i_0_458, plain, (r2_hidden(k10_ordinal2(X1,esk7_3(esk4_2(esk5_2(X1,X2),X3),X4,X5)),X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|X4!=k10_xtuple_0(esk4_2(esk5_2(X1,X2),X3))|~r2_hidden(X5,X4)|~v3_ordinal1(X3)|~v3_ordinal1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_15023, negated_conjecture, (r2_hidden(k10_ordinal2(X1,esk2_0),esk2_0)|r2_hidden(k1_xboole_0,k10_ordinal2(X1,esk2_0))|v4_ordinal1(esk2_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_15498, negated_conjecture, (r2_hidden(k1_xboole_0,k10_ordinal2(esk2_0,esk2_0))|r2_hidden(k10_ordinal2(esk2_0,esk2_0),esk2_0)|v4_ordinal1(esk2_0))).#trainneg
cnf(i_0_460, plain, (r2_hidden(k10_ordinal2(X1,esk3_3(esk5_2(X1,X2),X3,X4)),X2)|k10_ordinal2(X3,esk5_2(X1,X2))=k4_ordinal2(X4)|esk5_2(X1,X2)=k1_xboole_0|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|k9_xtuple_0(X4)!=esk5_2(X1,X2)|~v1_ordinal2(X4)|~v1_funct_1(X4)|~v1_relat_1(X4)|~v5_ordinal1(X4)|~v4_ordinal1(esk5_2(X1,X2))|~v3_ordinal1(X3)|~v3_ordinal1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_461, plain, (r2_hidden(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))),X3)|~r1_ordinal1(X1,X3)|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~r2_hidden(X2,esk5_2(X1,X3))|~v4_ordinal1(esk5_2(X1,X3))|~v3_ordinal1(X2)|~v3_ordinal1(esk5_2(X1,X3)))).#trainneg
cnf(i_0_15368, negated_conjecture, (r2_hidden(X1,k10_ordinal2(X2,esk1_0))|r2_hidden(X2,k10_ordinal2(X1,esk1_0))|~v3_ordinal1(k10_ordinal2(X2,esk1_0))|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_15255, negated_conjecture, (v1_xboole_0(X1)|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(k10_ordinal2(X1,esk1_0),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2))))).#trainneg
cnf(i_0_15189, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_ordinal1(k10_ordinal2(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),esk1_0),X1)|~r2_hidden(X1,k10_ordinal2(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),esk1_0))|~v3_ordinal1(k10_ordinal2(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),esk1_0))|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_15225, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v1_xboole_0(X2)|~r1_ordinal1(k10_ordinal2(X2,esk1_0),X1)|~r2_hidden(X1,k10_ordinal2(X2,esk1_0))|~v3_ordinal1(k10_ordinal2(X2,esk1_0))|~v3_ordinal1(X2)|~r1_tarski(k2_xboole_0(X1,k1_tarski(X1)),X2)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_14703, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(k10_ordinal2(X2,esk1_0),k1_tarski(k10_ordinal2(X2,esk1_0))))|~r2_hidden(esk6_1(k2_xboole_0(k10_ordinal2(X2,esk1_0),k1_tarski(k10_ordinal2(X2,esk1_0)))),X1)|~v3_ordinal1(k10_ordinal2(X2,esk1_0))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_13123, negated_conjecture, (~r1_ordinal1(X1,X1)|~r2_hidden(esk1_0,X2)|~r2_hidden(X2,k1_xboole_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_12974, negated_conjecture, (~r2_hidden(X2,X1)|~r2_hidden(esk1_0,esk5_2(X1,X2))|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_12511, negated_conjecture, (~r1_ordinal1(k2_xboole_0(X2,k1_tarski(X2)),X1)|~r2_hidden(esk1_0,k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))))|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_15091, plain, (r2_hidden(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(k10_ordinal2(X1,k1_xboole_0),X2))).#trainneg
cnf(i_0_4061, plain, (v3_ordinal1(k10_ordinal2(k10_ordinal2(X1,k1_tarski(k1_xboole_0)),X2))|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3058, plain, (esk12_0=esk15_0|r2_hidden(esk15_0,esk12_0)|~v1_xboole_0(esk15_0))).#trainneg
cnf(i_0_15451, plain, (~v1_xboole_0(X1)|X4!=k10_xtuple_0(esk4_2(esk5_2(X2,X1),X3))|~r1_ordinal1(X2,X1)|~r2_hidden(X5,X4)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_470, plain, (k10_ordinal2(X1,k10_ordinal2(esk5_2(X1,X2),k1_tarski(k1_xboole_0)))=k2_xboole_0(X2,k1_tarski(X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_8904, plain, (esk13_0=esk12_0|r2_hidden(esk13_0,esk12_0)|r1_ordinal1(esk12_0,esk13_0))).#trainneg
cnf(i_0_15499, negated_conjecture, (r2_hidden(k1_xboole_0,k10_ordinal2(esk1_0,esk2_0))|r2_hidden(k10_ordinal2(esk1_0,esk2_0),esk2_0)|v4_ordinal1(esk2_0))).#trainneg
cnf(i_0_15529, negated_conjecture, (r2_hidden(k1_xboole_0,k10_ordinal2(k1_xboole_0,esk2_0))|r2_hidden(k10_ordinal2(k1_xboole_0,esk2_0),esk2_0)|v4_ordinal1(esk2_0))).#trainneg
cnf(i_0_462, plain, (r2_hidden(k10_ordinal2(X1,k10_ordinal2(X2,k2_xboole_0(X3,k1_tarski(X3)))),X4)|~r1_ordinal1(X1,X4)|~v3_ordinal1(X1)|~v3_ordinal1(X4)|~r2_hidden(k10_ordinal2(X2,X3),esk5_2(X1,X4))|~v4_ordinal1(esk5_2(X1,X4))|~v3_ordinal1(esk5_2(X1,X4))|~v3_ordinal1(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_471, plain, (k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2)))=k10_ordinal2(k10_ordinal2(X1,X2),k1_tarski(k1_xboole_0))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(X1,X2)))).#trainneg
cnf(i_0_15653, negated_conjecture, (v1_xboole_0(X1)|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(k10_ordinal2(X1,esk1_0),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_15644, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k10_ordinal2(k2_xboole_0(X1,k1_tarski(X1)),esk1_0),X1)|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_15635, negated_conjecture, (v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(k10_ordinal2(k2_xboole_0(X1,k1_tarski(X1)),esk1_0),X2)|~v3_ordinal1(X2)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_15723, negated_conjecture, (v4_ordinal1(k2_xboole_0(k10_ordinal2(X1,esk1_0),k1_tarski(k10_ordinal2(X1,esk1_0))))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|v1_xboole_0(X2)|~r2_hidden(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_tarski(k2_xboole_0(k10_ordinal2(X1,esk1_0),k1_tarski(k10_ordinal2(X1,esk1_0))),X2)|~v3_ordinal1(k2_xboole_0(k10_ordinal2(X1,esk1_0),k1_tarski(k10_ordinal2(X1,esk1_0)))))).#trainneg
cnf(i_0_12583, negated_conjecture, (v4_ordinal1(k10_ordinal2(X1,esk1_0))|~r1_ordinal1(k10_ordinal2(esk6_1(k10_ordinal2(X1,esk1_0)),k1_tarski(k1_xboole_0)),X1)|~v3_ordinal1(k10_ordinal2(esk6_1(k10_ordinal2(X1,esk1_0)),k1_tarski(k1_xboole_0)))|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,esk1_0))))).#trainneg
cnf(i_0_12728, plain, (r2_hidden(X1,k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))))|r2_hidden(X2,esk5_2(X1,X1))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_13325, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(esk7_3(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),X3,X1))|~r2_hidden(esk1_0,esk6_1(esk7_3(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),X3,X1)))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(esk6_1(esk7_3(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),X3,X1)))|~v3_ordinal1(X1)|X3!=k10_xtuple_0(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))|~r2_hidden(X1,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_14683, negated_conjecture, (r1_ordinal1(esk1_0,k10_ordinal2(X1,esk1_0))|~r2_hidden(esk1_0,X1)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_14837, negated_conjecture, (~r1_ordinal1(k10_ordinal2(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))),esk1_0),X1)|~v3_ordinal1(k10_ordinal2(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))),esk1_0))|~v3_ordinal1(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))))|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_13686, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(esk5_2(k2_xboole_0(X2,k1_tarski(X2)),X1))|~r1_ordinal1(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v4_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(esk6_1(esk5_2(k2_xboole_0(X2,k1_tarski(X2)),X1)))|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_15152, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k10_ordinal2(k3_ordinal2(X2),esk1_0),X1)|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))).#trainneg
cnf(i_0_16502, plain, (v3_ordinal1(k10_ordinal2(k10_ordinal2(X1,X2),k1_tarski(k1_xboole_0)))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_3648, negated_conjecture, (esk2_0=esk15_0|r2_hidden(esk15_0,esk2_0)|~v1_xboole_0(esk15_0))).#trainneg
cnf(i_0_15883, plain, (X1!=k10_xtuple_0(esk4_2(k1_xboole_0,X4))|~v1_xboole_0(X3)|~r1_ordinal1(X2,X3)|~r2_hidden(X5,X1)|~v3_ordinal1(X2)|~v3_ordinal1(X4)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_472, plain, (k10_ordinal2(k1_xboole_0,k1_tarski(k1_xboole_0))=k1_tarski(k1_xboole_0))).#trainneg
cnf(i_0_16569, negated_conjecture, (~r2_hidden(k10_ordinal2(k1_tarski(k1_xboole_0),esk1_0),k1_xboole_0))).#trainneg
cnf(i_0_15024, negated_conjecture, (r2_hidden(k10_ordinal2(X1,esk1_0),esk2_0)|r2_hidden(k1_xboole_0,k10_ordinal2(X1,esk1_0))|v4_ordinal1(esk2_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_476, plain, (r2_hidden(X1,k10_ordinal2(X2,k1_tarski(k1_xboole_0)))|r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_478, plain, (r1_ordinal1(k10_ordinal2(X1,k10_ordinal2(X2,k1_tarski(k1_xboole_0))),X3)|~r2_hidden(k10_ordinal2(X1,X2),X3)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_16799, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k10_ordinal2(k3_ordinal2(k2_xboole_0(X1,k1_tarski(X1))),esk1_0),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_16751, negated_conjecture, (~v3_ordinal1(k10_ordinal2(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))),esk1_0))|~v3_ordinal1(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))))|~v3_ordinal1(X1)|~r1_ordinal1(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))),X1)|~r2_hidden(k10_ordinal2(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))),esk1_0),X1)|~r2_hidden(X1,k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1))))))).#trainneg
cnf(i_0_16764, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(k10_ordinal2(k3_ordinal2(esk6_1(k2_xboole_0(X2,k1_tarski(X2)))),esk1_0),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_16933, plain, (r1_ordinal1(k1_xboole_0,k1_tarski(k1_xboole_0)))).#trainneg
cnf(i_0_16572, negated_conjecture, (v4_ordinal1(k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(k10_ordinal2(k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))),esk1_0),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_13141, negated_conjecture, (~r2_hidden(esk1_0,X3)|~r2_hidden(X3,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r1_ordinal1(X1,X1)|~r2_hidden(k10_ordinal2(X1,X2),X1)|~r2_hidden(X1,X1))).#trainneg
cnf(i_0_15629, negated_conjecture, (r2_hidden(X1,k10_ordinal2(X2,esk1_0))|r2_hidden(X2,k10_ordinal2(X1,esk1_0))|~v3_ordinal1(k10_ordinal2(X2,esk1_0))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_14689, negated_conjecture, (~r2_hidden(X2,X1)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)|~r1_ordinal1(X1,X2)|~r2_hidden(k10_ordinal2(X1,esk1_0),X2))).#trainneg
cnf(i_0_17434, negated_conjecture, (r2_hidden(X1,k10_ordinal2(X2,esk1_0))|r2_hidden(X2,k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_17458, negated_conjecture, (~r1_ordinal1(X1,k10_ordinal2(X2,esk1_0))|~r2_hidden(k10_ordinal2(X2,esk1_0),X1)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)|~r2_hidden(k10_ordinal2(X1,esk1_0),X2)|~v3_ordinal1(k10_ordinal2(X2,esk1_0))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_15449, plain, (~v1_xboole_0(X1)|~r1_tarski(X3,X1)|X5!=k10_xtuple_0(esk4_2(esk5_2(X2,X3),X4))|~r1_ordinal1(X2,X3)|~r2_hidden(X6,X5)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X4))).#trainneg
cnf(i_0_479, plain, (r2_hidden(k10_ordinal2(X1,k10_ordinal2(X2,k1_tarski(k1_xboole_0))),X3)|~r2_hidden(k10_ordinal2(X1,X2),X3)|~v4_ordinal1(X3)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_8957, negated_conjecture, (esk2_0=esk15_0|r2_hidden(k1_xboole_0,esk2_0)|r2_hidden(k1_xboole_0,esk15_0))).#trainneg
cnf(i_0_17009, negated_conjecture, (r2_hidden(k1_xboole_0,k10_ordinal2(esk2_0,esk1_0))|r2_hidden(k10_ordinal2(esk2_0,esk1_0),esk2_0)|v4_ordinal1(esk2_0))).#trainneg
cnf(i_0_17010, negated_conjecture, (r2_hidden(k1_xboole_0,k10_ordinal2(esk1_0,esk1_0))|r2_hidden(k10_ordinal2(esk1_0,esk1_0),esk2_0)|v4_ordinal1(esk2_0))).#trainneg
cnf(i_0_480, plain, (r2_hidden(k10_ordinal2(X1,X2),X3)|~r1_ordinal1(k10_ordinal2(X1,k10_ordinal2(X2,k1_tarski(k1_xboole_0))),X3)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_484, plain, (r1_ordinal1(k10_ordinal2(X1,k1_tarski(k1_xboole_0)),X2)|~r2_hidden(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_12314, negated_conjecture, (r2_hidden(X2,esk1_0)|~r1_ordinal1(k10_ordinal2(k10_ordinal2(X1,esk1_0),X2),X1)|~v3_ordinal1(k10_ordinal2(k10_ordinal2(X1,esk1_0),X2))|~v3_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_7058, plain, (v4_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(X1,esk6_1(k10_ordinal2(X1,X2)))|~r2_hidden(k1_xboole_0,X2)|~r2_hidden(esk6_1(k10_ordinal2(X1,X2)),X1)|~v4_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_11578, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k2_xboole_0(k10_ordinal2(X2,X3),k1_tarski(k10_ordinal2(X2,X3))),X1)|~v4_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(X2,X3))|~r1_ordinal1(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_12976, negated_conjecture, (~r2_hidden(k2_xboole_0(X2,k1_tarski(X2)),X1)|~r2_hidden(esk1_0,k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))))|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(esk5_2(X1,X2),k1_tarski(esk5_2(X1,X2))))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_17563, plain, (X1!=k10_xtuple_0(esk4_2(k1_xboole_0,X4))|~v1_xboole_0(X5)|~r1_ordinal1(X2,X3)|~r2_hidden(X6,X1)|~r1_tarski(X3,X5)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X4)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_485, plain, (r2_hidden(k10_ordinal2(X1,k1_tarski(k1_xboole_0)),X2)|~r2_hidden(X1,X2)|~v4_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_8959, negated_conjecture, (esk13_0=esk2_0|r2_hidden(k1_xboole_0,esk2_0)|r2_hidden(k1_xboole_0,esk13_0))).#trainneg
cnf(i_0_17042, negated_conjecture, (r2_hidden(k1_xboole_0,k10_ordinal2(k1_xboole_0,esk1_0))|r2_hidden(k10_ordinal2(k1_xboole_0,esk1_0),esk2_0)|v4_ordinal1(esk2_0))).#trainneg
cnf(i_0_14564, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(X2,k10_ordinal2(esk2_0,X1))|r1_ordinal1(esk2_0,k10_ordinal2(X2,k1_xboole_0))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_487, plain, (r1_ordinal1(esk15_0,k1_funct_1(esk4_2(X1,X2),X3))|r1_ordinal1(k1_funct_1(esk4_2(X1,X2),X3),esk15_0)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_488, negated_conjecture, (r1_ordinal1(esk2_0,k1_funct_1(esk4_2(X1,X2),X3))|r1_ordinal1(k1_funct_1(esk4_2(X1,X2),X3),esk2_0)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_17795, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(k2_xboole_0(X2,k1_tarski(X2)),X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_17788, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(k2_xboole_0(X5,k1_tarski(X5)),X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(esk7_3(esk4_2(X3,X2),X4,X5))|X4!=k10_xtuple_0(esk4_2(X3,X2))|~r2_hidden(X5,X4)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_16948, plain, (v3_ordinal1(k10_ordinal2(k1_tarski(k1_xboole_0),k1_tarski(k1_xboole_0))))).#trainneg
cnf(i_0_6160, negated_conjecture, (esk13_0=esk2_0|r2_hidden(esk2_0,esk13_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_489, negated_conjecture, (r1_ordinal1(esk1_0,k1_funct_1(esk4_2(X1,X2),X3))|r1_ordinal1(k1_funct_1(esk4_2(X1,X2),X3),esk1_0)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_17968, negated_conjecture, (v1_xboole_0(X1)|r1_ordinal1(esk2_0,k10_ordinal2(esk2_0,k1_xboole_0))|r2_hidden(esk2_0,k10_ordinal2(esk2_0,X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_17969, negated_conjecture, (v1_xboole_0(X1)|r1_ordinal1(esk2_0,k10_ordinal2(esk1_0,k1_xboole_0))|r2_hidden(esk1_0,k10_ordinal2(esk2_0,X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_490, plain, (r1_ordinal1(X1,k1_funct_1(esk4_2(X2,X3),X4))|r1_ordinal1(k1_funct_1(esk4_2(X2,X3),X4),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X4)|~v3_ordinal1(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_18049, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),esk1_0)|~r2_hidden(k2_xboole_0(esk1_0,k1_tarski(esk1_0)),X1)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))).#trainneg
cnf(i_0_18048, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),esk2_0)|~r2_hidden(k2_xboole_0(esk2_0,k1_tarski(esk2_0)),X1)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))).#trainneg
cnf(i_0_17807, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(k10_ordinal2(X2,k2_xboole_0(X3,k1_tarski(X3))),X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_18069, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))|~r2_hidden(k2_xboole_0(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),k1_tarski(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))),X1)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))|~r1_ordinal1(X2,k1_xboole_0)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_18158, negated_conjecture, (r1_ordinal1(esk1_0,k1_funct_1(esk4_2(X1,X2),X3))|~r2_hidden(esk1_0,k1_funct_1(esk4_2(X1,X2),X3))|~v3_ordinal1(k1_funct_1(esk4_2(X1,X2),X3))|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_16978, plain, (v7_ordinal1(k1_tarski(k1_xboole_0))|~v7_ordinal1(k10_ordinal2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_16975, plain, (m1_subset_1(k10_ordinal2(k1_xboole_0,k1_xboole_0),k1_tarski(k1_xboole_0)))).#trainneg
cnf(i_0_6540, negated_conjecture, (esk12_0=esk2_0|r2_hidden(esk2_0,esk12_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_494, plain, (esk10_1(k1_zfmisc_1(X1))=X1|r2_hidden(esk10_1(k1_zfmisc_1(X1)),X1)|~v1_ordinal1(esk10_1(k1_zfmisc_1(X1)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_18001, negated_conjecture, (v1_xboole_0(X1)|r1_ordinal1(esk2_0,k10_ordinal2(k1_xboole_0,k1_xboole_0))|r2_hidden(k1_xboole_0,k10_ordinal2(esk2_0,X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_18187, negated_conjecture, (v1_xboole_0(esk2_0)|r1_ordinal1(esk2_0,k10_ordinal2(esk2_0,k1_xboole_0))|r2_hidden(esk2_0,k10_ordinal2(esk2_0,esk2_0)))).#trainneg
cnf(i_0_495, plain, (r2_hidden(esk10_1(k1_zfmisc_1(X1)),X2)|~v1_ordinal1(esk10_1(k1_zfmisc_1(X1)))|~r2_hidden(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_498, plain, (r1_tarski(esk7_3(esk4_2(k1_zfmisc_1(X1),X2),X3,X4),X1)|X3!=k10_xtuple_0(esk4_2(k1_zfmisc_1(X1),X2))|~r2_hidden(X4,X3)|~v3_ordinal1(X2)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_11379, plain, (v1_xboole_0(esk5_2(X1,X2))|v1_xboole_0(X1)|v4_ordinal1(k2_xboole_0(k2_xboole_0(X2,k1_tarski(X2)),k1_tarski(k2_xboole_0(X2,k1_tarski(X2)))))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_ordinal1(X1,X2)|~r1_tarski(k2_xboole_0(k2_xboole_0(X2,k1_tarski(X2)),k1_tarski(k2_xboole_0(X2,k1_tarski(X2)))),X1)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_5942, plain, (v1_xboole_0(esk5_2(X1,X2))|v4_ordinal1(X2)|~r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_9473, plain, (v4_ordinal1(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2))))|v4_ordinal1(esk5_2(X3,k10_ordinal2(X1,X2)))|~r2_hidden(k10_ordinal2(X1,X2),X3)|~v3_ordinal1(X3)|~r1_ordinal1(X1,X3)|~r2_hidden(X3,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_18454, plain, (v3_ordinal1(k10_ordinal2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_16985, plain, (r2_hidden(k10_ordinal2(k1_xboole_0,k1_xboole_0),k1_tarski(k1_xboole_0)))).#trainneg
cnf(i_0_6908, plain, (esk13_0=esk15_0|r2_hidden(k1_xboole_0,esk15_0)|~v1_xboole_0(esk13_0))).#trainneg
cnf(i_0_499, plain, (r1_tarski(esk6_1(k1_zfmisc_1(X1)),X1)|v4_ordinal1(k1_zfmisc_1(X1))|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_18232, negated_conjecture, (v1_xboole_0(esk2_0)|r1_ordinal1(esk2_0,k10_ordinal2(esk1_0,k1_xboole_0))|r2_hidden(esk1_0,k10_ordinal2(esk2_0,esk2_0)))).#trainneg
cnf(i_0_18468, negated_conjecture, (v1_xboole_0(esk2_0)|r1_ordinal1(esk2_0,k10_ordinal2(k1_xboole_0,k1_xboole_0))|r2_hidden(k1_xboole_0,k10_ordinal2(esk2_0,esk2_0)))).#trainneg
cnf(i_0_500, plain, (r1_tarski(esk3_3(k1_zfmisc_1(X1),X2,X3),X1)|k10_ordinal2(X2,k1_zfmisc_1(X1))=k4_ordinal2(X3)|k1_zfmisc_1(X1)=k1_xboole_0|k9_xtuple_0(X3)!=k1_zfmisc_1(X1)|~v1_ordinal2(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v5_ordinal1(X3)|~v4_ordinal1(k1_zfmisc_1(X1))|~v3_ordinal1(X2)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_501, plain, (r1_tarski(k2_xboole_0(X1,k1_tarski(X1)),X2)|~r2_hidden(X1,k1_zfmisc_1(X2))|~v4_ordinal1(k1_zfmisc_1(X2))|~v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_18385, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(esk7_3(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),X3,X4),k1_xboole_0)|~r2_hidden(k2_xboole_0(X4,k1_tarski(X4)),X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X4)|~v4_ordinal1(X1)|~v3_ordinal1(X4)|~v3_ordinal1(X1)|~v3_ordinal1(esk7_3(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),X3,X4))|X3!=k10_xtuple_0(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))|~r2_hidden(X4,X3)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_18428, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(X2,k1_xboole_0)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))|~v4_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),X1))).#trainneg
cnf(i_0_18329, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(k2_xboole_0(X3,k1_tarski(X3)),X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(esk5_2(X2,X3))|~r1_ordinal1(X2,X3)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_9329, plain, (v1_xboole_0(X1)|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),k1_tarski(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_tarski(k2_xboole_0(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),k1_tarski(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))),X1)|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_14753, negated_conjecture, (~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~r2_hidden(k2_xboole_0(k10_ordinal2(X1,esk1_0),k1_tarski(k10_ordinal2(X1,esk1_0))),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_18688, plain, (v3_ordinal1(k10_ordinal2(X1,k10_ordinal2(k1_xboole_0,k1_xboole_0)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_502, plain, (r1_tarski(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))),X3)|~r2_hidden(k10_ordinal2(X1,X2),k1_zfmisc_1(X3))|~v4_ordinal1(k1_zfmisc_1(X3))|~v3_ordinal1(k1_zfmisc_1(X3))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_10807, plain, (esk13_0=k1_xboole_0|r2_hidden(esk13_0,k1_xboole_0)|r1_ordinal1(k1_xboole_0,esk13_0))).#trainneg
cnf(i_0_18686, plain, (~r2_hidden(k1_tarski(k1_xboole_0),k10_ordinal2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_18631, plain, (r2_hidden(k1_xboole_0,k10_ordinal2(k1_xboole_0,k1_xboole_0))|r2_hidden(k10_ordinal2(k1_xboole_0,k1_xboole_0),X1)|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_18650, plain, (r1_ordinal1(X1,k10_ordinal2(k1_xboole_0,k1_xboole_0))|r1_ordinal1(k10_ordinal2(k1_xboole_0,k1_xboole_0),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_507, plain, (r1_ordinal1(esk14_0,esk10_1(X1))|r1_ordinal1(esk10_1(X1),esk14_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_508, plain, (r1_ordinal1(esk14_0,esk6_1(X1))|r1_ordinal1(esk6_1(X1),esk14_0)|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_18804, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(k3_ordinal2(X2),k1_xboole_0)|~r2_hidden(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k3_ordinal2(X2)),X1)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))).#trainneg
cnf(i_0_6030, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v1_xboole_0(X3)|~r1_ordinal1(k3_ordinal2(k10_ordinal2(X2,X3)),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_11741, negated_conjecture, (r2_hidden(X1,k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_5729, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(k2_xboole_0(k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))),k1_tarski(k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))))),X1)|~v3_ordinal1(k2_xboole_0(k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))),k1_tarski(k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))))))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))).#trainneg
cnf(i_0_18316, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k2_xboole_0(esk1_0,k1_tarski(esk1_0)),X1)|~v4_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~r1_tarski(k2_xboole_0(X1,k1_tarski(X1)),esk1_0)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_9038, plain, (v4_ordinal1(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))))|v4_ordinal1(k2_xboole_0(X3,k1_tarski(X3)))|~r2_hidden(esk6_1(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2)))),X3)|~r2_hidden(esk6_1(k2_xboole_0(X3,k1_tarski(X3))),k10_ordinal2(X1,X2))|~v3_ordinal1(X3)|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_18988, plain, (~r2_hidden(k1_xboole_0,k10_ordinal2(k1_xboole_0,k1_xboole_0))|~v4_ordinal1(k10_ordinal2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_509, plain, (r1_ordinal1(esk14_0,esk5_2(X1,X2))|r1_ordinal1(esk5_2(X1,X2),esk14_0)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_11496, plain, (esk13_0=esk14_0|r2_hidden(esk14_0,esk13_0)|r2_hidden(esk13_0,esk14_0))).#trainneg
cnf(i_0_18657, plain, (X1=k10_ordinal2(k1_xboole_0,k1_xboole_0)|r2_hidden(k1_xboole_0,X1)|r2_hidden(k1_xboole_0,k10_ordinal2(k1_xboole_0,k1_xboole_0))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_18729, negated_conjecture, (v1_xboole_0(esk2_0)|r1_ordinal1(esk2_0,esk1_0)|r2_hidden(esk1_0,k10_ordinal2(esk2_0,esk2_0)))).#trainneg
cnf(i_0_510, plain, (r1_ordinal1(esk14_0,esk3_3(X1,X2,X3))|r1_ordinal1(esk3_3(X1,X2,X3),esk14_0)|k10_ordinal2(X2,X1)=k4_ordinal2(X3)|X1=k1_xboole_0|k9_xtuple_0(X3)!=X1|~v1_ordinal2(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v5_ordinal1(X3)|~v4_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_513, plain, (r1_ordinal1(esk14_0,k3_ordinal2(X1))|r1_ordinal1(k3_ordinal2(X1),esk14_0))).#trainneg
cnf(i_0_19197, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),esk1_0),X1)|~v3_ordinal1(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),esk1_0))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))).#trainneg
cnf(i_0_9076, plain, (v1_xboole_0(esk7_3(esk4_2(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2)))),X3,X4))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(X4,X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2)|~v3_ordinal1(esk7_3(esk4_2(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2)))),X3,X4))|X3!=k10_xtuple_0(esk4_2(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2)))))|~r2_hidden(X4,X3)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_19088, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|r1_ordinal1(esk14_0,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(k2_xboole_0(esk14_0,k1_tarski(esk14_0)),X1)|~v4_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_6477, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),k1_tarski(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))))|~r2_hidden(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(esk6_1(k2_xboole_0(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),k1_tarski(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_18325, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(k10_ordinal2(X2,k1_tarski(k1_xboole_0)),X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_16662, negated_conjecture, (v4_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(k10_ordinal2(esk6_1(k10_ordinal2(X1,esk1_0)),k1_tarski(k1_xboole_0)))|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)|~r1_ordinal1(esk6_1(k10_ordinal2(X1,esk1_0)),X1)|~r2_hidden(k10_ordinal2(esk6_1(k10_ordinal2(X1,esk1_0)),k1_tarski(k1_xboole_0)),X1)|~r2_hidden(X1,esk6_1(k10_ordinal2(X1,esk1_0)))|~v3_ordinal1(esk6_1(k10_ordinal2(X1,esk1_0))))).#trainneg
cnf(i_0_11106, plain, (v7_ordinal1(X1)|r2_hidden(X2,X1)|~v7_ordinal1(X2)|~r1_ordinal1(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_19704, plain, (v7_ordinal1(esk15_0)|r2_hidden(esk12_0,esk15_0)|r1_ordinal1(esk15_0,esk12_0))).#trainneg
cnf(i_0_19317, negated_conjecture, (~r2_hidden(k2_xboole_0(esk1_0,k1_tarski(esk1_0)),k1_xboole_0)|~r1_tarski(k1_tarski(k1_xboole_0),esk1_0)|~v3_ordinal1(esk6_1(k1_tarski(k1_xboole_0))))).#trainneg
cnf(i_0_514, plain, (r1_ordinal1(esk14_0,k2_xboole_0(X1,k1_tarski(X1)))|r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),esk14_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_12101, plain, (esk12_0=esk14_0|r2_hidden(esk14_0,esk12_0)|r2_hidden(esk12_0,esk14_0))).#trainneg
cnf(i_0_515, plain, (r1_ordinal1(esk14_0,k2_xboole_0(X1,X2))|r1_ordinal1(k2_xboole_0(X1,X2),esk14_0)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_516, plain, (r1_ordinal1(esk14_0,k4_ordinal2(X1))|r1_ordinal1(k4_ordinal2(X1),esk14_0)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v5_ordinal1(X1))).#trainneg
cnf(i_0_19575, plain, (v4_ordinal1(k2_xboole_0(X4,k1_tarski(X4)))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(esk6_1(k2_xboole_0(X4,k1_tarski(X4))),X1)|~r2_hidden(X4,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(esk7_3(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),X3,X4))|X3!=k10_xtuple_0(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))|~r2_hidden(X4,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_19538, plain, (esk7_3(esk4_2(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2)))),X3,X4)=k1_xboole_0|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|X3!=k10_xtuple_0(esk4_2(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2)))))|~r2_hidden(X4,X2)|~r2_hidden(X4,X3)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_19646, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))|~r2_hidden(k10_ordinal2(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),k1_tarski(k1_xboole_0)),X1)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))|~r1_ordinal1(X2,k1_xboole_0)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_19616, plain, (v4_ordinal1(k2_xboole_0(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),k1_tarski(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v1_xboole_0(X1)|~r2_hidden(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r1_tarski(k2_xboole_0(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),k1_tarski(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))),X1)|~v3_ordinal1(k2_xboole_0(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),k1_tarski(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)))))).#trainneg
cnf(i_0_19830, plain, (v7_ordinal1(X1)|r2_hidden(X2,X1)|r2_hidden(X1,X2)|~v7_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_19783, plain, (v7_ordinal1(esk6_1(k10_ordinal2(X1,X2)))|r2_hidden(k10_ordinal2(X1,X2),esk6_1(k10_ordinal2(X1,X2)))|v4_ordinal1(k10_ordinal2(X1,X2))|~v7_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(esk6_1(k10_ordinal2(X1,X2)))|~r1_ordinal1(X1,esk6_1(k10_ordinal2(X1,X2)))|~v4_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_19034, plain, (r1_ordinal1(k10_ordinal2(k1_xboole_0,k1_xboole_0),k10_ordinal2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_20124, plain, (r1_ordinal1(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_15423, plain, (~v1_xboole_0(X1)|X4!=k10_xtuple_0(esk4_2(esk5_2(X2,esk10_1(k1_zfmisc_1(X1))),X3))|~r1_ordinal1(X2,esk10_1(k1_zfmisc_1(X1)))|~r2_hidden(X5,X4)|~v3_ordinal1(X2)|~v3_ordinal1(esk10_1(k1_zfmisc_1(X1)))|~v3_ordinal1(X3))).#trainneg
cnf(i_0_517, plain, (r1_ordinal1(esk14_0,k1_funct_1(esk4_2(X1,X2),X3))|r1_ordinal1(k1_funct_1(esk4_2(X1,X2),X3),esk14_0)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_19835, negated_conjecture, (~r2_hidden(k2_xboole_0(esk1_0,k1_tarski(esk1_0)),k1_xboole_0)|~r1_tarski(k1_tarski(k1_xboole_0),esk1_0))).#trainneg
cnf(i_0_19819, plain, (v7_ordinal1(X1)|r2_hidden(k1_xboole_0,X1)|~v7_ordinal1(k1_xboole_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_518, plain, (r1_ordinal1(esk14_0,k10_ordinal2(X1,X2))|r1_ordinal1(k10_ordinal2(X1,X2),esk14_0)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_523, plain, (r2_hidden(k10_ordinal2(X1,esk10_1(esk5_2(X1,X2))),X2)|v1_xboole_0(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_20104, plain, (v7_ordinal1(esk6_1(k10_ordinal2(X1,X2)))|r2_hidden(k10_ordinal2(X1,X2),esk6_1(k10_ordinal2(X1,X2)))|v4_ordinal1(k10_ordinal2(X1,X2))|r2_hidden(esk6_1(k10_ordinal2(X1,X2)),X1)|~v7_ordinal1(k10_ordinal2(X1,X2))|~v4_ordinal1(X2)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,X2)))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_20045, plain, (v1_xboole_0(X1)|v4_ordinal1(k2_xboole_0(X4,k1_tarski(X4)))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(X4,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r1_tarski(k2_xboole_0(X4,k1_tarski(X4)),X1)|~v3_ordinal1(k2_xboole_0(X4,k1_tarski(X4)))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(esk7_3(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),X3,X4))|X3!=k10_xtuple_0(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))|~r2_hidden(X4,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_20046, plain, (v1_xboole_0(X1)|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r1_tarski(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_19562, plain, (r2_hidden(esk14_0,X1)|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1)|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~r2_hidden(k2_xboole_0(esk14_0,k1_tarski(esk14_0)),X2)|~v4_ordinal1(X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_19553, plain, (esk6_1(k2_xboole_0(X1,k1_tarski(X1)))=esk14_0|r2_hidden(esk14_0,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(k2_xboole_0(esk14_0,k1_tarski(esk14_0)),X1)|~v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_13022, negated_conjecture, (~r2_hidden(esk1_0,k10_ordinal2(X1,k1_tarski(k1_xboole_0)))|~r2_hidden(k10_ordinal2(X2,X1),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_6675, plain, (r2_hidden(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(X3,X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X1,X3))|~v3_ordinal1(X2)|~r1_ordinal1(X1,X3)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_14720, negated_conjecture, (v1_xboole_0(X1)|v4_ordinal1(k10_ordinal2(k10_ordinal2(X2,esk1_0),X1))|~v3_ordinal1(k10_ordinal2(X2,esk1_0))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k10_ordinal2(k10_ordinal2(X2,esk1_0),X1)),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_15908, plain, (v4_ordinal1(X1)|~r2_hidden(k10_ordinal2(X2,k10_ordinal2(esk5_2(X2,esk6_1(X1)),k1_tarski(k1_xboole_0))),X1)|~v3_ordinal1(X1)|~r1_ordinal1(X2,esk6_1(X1))|~v3_ordinal1(X2)|~v3_ordinal1(esk6_1(X1)))).#trainneg
cnf(i_0_19705, plain, (v7_ordinal1(esk14_0)|r2_hidden(esk12_0,esk14_0)|r1_ordinal1(esk14_0,esk12_0))).#trainneg
cnf(i_0_20380, negated_conjecture, (~r2_hidden(k10_ordinal2(X1,k2_xboole_0(esk1_0,k1_tarski(esk1_0))),X1)|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(esk1_0,k1_tarski(esk1_0)))|~v3_ordinal1(k10_ordinal2(k2_xboole_0(esk1_0,k1_tarski(esk1_0)),k1_tarski(k1_xboole_0))))).#trainneg
cnf(i_0_528, plain, (m1_subset_1(esk10_1(X1),X2)|v1_xboole_0(X1)|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_20170, negated_conjecture, (~r2_hidden(k2_xboole_0(esk1_0,k1_tarski(esk1_0)),k1_xboole_0)|~r1_ordinal1(k1_tarski(k1_xboole_0),esk1_0))).#trainneg
cnf(i_0_20171, plain, (v7_ordinal1(X1)|r2_hidden(k1_xboole_0,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_19738, negated_conjecture, (v7_ordinal1(esk1_0)|r2_hidden(esk2_0,esk1_0)|r1_ordinal1(esk1_0,esk2_0)|~v7_ordinal1(esk2_0))).#trainneg
cnf(i_0_534, plain, (v1_xboole_0(k1_zfmisc_1(X1))|r2_hidden(k1_xboole_0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_20381, negated_conjecture, (~r2_hidden(k10_ordinal2(X2,X1),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(esk1_0,X1))).#trainneg
cnf(i_0_20315, plain, (v1_xboole_0(X1)|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~r2_hidden(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r1_tarski(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_20392, plain, (r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1)|~r2_hidden(X1,X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_20399, negated_conjecture, (v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(k2_xboole_0(esk5_2(k10_ordinal2(X1,esk1_0),X2),k1_tarski(esk5_2(k10_ordinal2(X1,esk1_0),X2))))|~v3_ordinal1(X1)|~r1_ordinal1(k10_ordinal2(X1,esk1_0),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_20383, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(k10_ordinal2(X2,X1),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_20320, plain, (r2_hidden(esk14_0,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k2_xboole_0(esk14_0,k1_tarski(esk14_0)),X2)|~v4_ordinal1(X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X2)|~r2_hidden(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_20565, plain, (v1_xboole_0(X1)|v4_ordinal1(k2_xboole_0(k2_xboole_0(X2,k1_tarski(X2)),k1_tarski(k2_xboole_0(X2,k1_tarski(X2)))))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k2_xboole_0(X2,k1_tarski(X2)),esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r1_tarski(k2_xboole_0(k2_xboole_0(X2,k1_tarski(X2)),k1_tarski(k2_xboole_0(X2,k1_tarski(X2)))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(k2_xboole_0(k2_xboole_0(X2,k1_tarski(X2)),k1_tarski(k2_xboole_0(X2,k1_tarski(X2)))))|~v3_ordinal1(X1)|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_20523, negated_conjecture, (v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_ordinal1(esk1_0,X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X1)|~r2_hidden(k10_ordinal2(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_19709, negated_conjecture, (v7_ordinal1(esk2_0)|r2_hidden(esk12_0,esk2_0)|r1_ordinal1(esk2_0,esk12_0))).#trainneg
cnf(i_0_535, plain, (v1_xboole_0(k1_zfmisc_1(X1))|r2_hidden(X2,k1_zfmisc_1(X1))|~r1_ordinal1(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_20509, negated_conjecture, (v7_ordinal1(esk1_0)|r1_ordinal1(esk1_0,esk2_0)|r2_hidden(esk2_0,esk1_0)|r2_hidden(k1_xboole_0,esk2_0))).#trainneg
cnf(i_0_20508, negated_conjecture, (v7_ordinal1(esk1_0)|r1_ordinal1(esk1_0,esk2_0)|r2_hidden(esk2_0,esk1_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_547, plain, (m1_subset_1(X1,X2)|v4_ordinal1(X3)|~r2_hidden(X1,esk6_1(X3))|~r1_tarski(X3,k1_zfmisc_1(X2))|~v3_ordinal1(X3))).#trainneg
cnf(i_0_20580, plain, (v1_xboole_0(X1)|v4_ordinal1(k2_xboole_0(k2_xboole_0(X2,k1_tarski(X2)),k1_tarski(k2_xboole_0(X2,k1_tarski(X2)))))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k2_xboole_0(X2,k1_tarski(X2)),esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r1_tarski(k2_xboole_0(k2_xboole_0(X2,k1_tarski(X2)),k1_tarski(k2_xboole_0(X2,k1_tarski(X2)))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(k2_xboole_0(k2_xboole_0(X2,k1_tarski(X2)),k1_tarski(k2_xboole_0(X2,k1_tarski(X2)))))|~v3_ordinal1(X1)|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~r2_hidden(X2,X1)|~v4_ordinal1(X1))).#trainneg
cnf(i_0_14675, negated_conjecture, (~r2_hidden(k10_ordinal2(X2,k1_xboole_0),X1)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)|~r1_ordinal1(k10_ordinal2(X1,esk1_0),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_10734, plain, (v1_xboole_0(esk5_2(k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))),X2))|v1_xboole_0(X1)|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))),X2)|~r1_tarski(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))))|~v3_ordinal1(X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))).#trainneg
cnf(i_0_20302, plain, (v1_xboole_0(X1)|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r2_hidden(X2,k10_xtuple_0(esk4_2(X3,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))))|~r1_tarski(k2_xboole_0(X2,k1_tarski(X2)),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(X1)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_13142, negated_conjecture, (~r2_hidden(esk1_0,X3)|~r2_hidden(X3,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(k10_ordinal2(X1,X2),X1))).#trainneg
cnf(i_0_8832, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k2_xboole_0(k2_xboole_0(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),k1_tarski(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))),k1_tarski(k2_xboole_0(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),k1_tarski(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(k2_xboole_0(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),k1_tarski(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))),k1_tarski(k2_xboole_0(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),k1_tarski(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))))))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_1275, plain, (v3_ordinal1(X3)|~v1_xboole_0(esk7_3(esk4_2(X1,esk12_0),X2,X3))|X2!=k10_xtuple_0(esk4_2(X1,esk12_0))|~r2_hidden(X3,X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_19710, negated_conjecture, (v7_ordinal1(esk1_0)|r2_hidden(esk12_0,esk1_0)|r1_ordinal1(esk1_0,esk12_0))).#trainneg
cnf(i_0_551, plain, (k10_ordinal2(X1,X2)=k10_ordinal2(X1,X3)|r2_hidden(k10_ordinal2(X1,X2),k10_ordinal2(X1,X3))|~v3_ordinal1(k10_ordinal2(X1,X3))|~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(X2,X3)|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_20498, negated_conjecture, (~r1_ordinal1(k1_tarski(k1_xboole_0),esk1_0)|~r2_hidden(k10_ordinal2(esk1_0,k1_tarski(k1_xboole_0)),k1_xboole_0))).#trainneg
cnf(i_0_16938, plain, (~r2_hidden(k10_ordinal2(k1_xboole_0,k2_xboole_0(k1_tarski(k1_xboole_0),k1_tarski(k1_tarski(k1_xboole_0)))),k1_tarski(k1_xboole_0)))).#trainneg
cnf(i_0_19744, negated_conjecture, (v7_ordinal1(k1_xboole_0)|r2_hidden(esk2_0,k1_xboole_0)|r1_ordinal1(k1_xboole_0,esk2_0)|~v7_ordinal1(esk2_0))).#trainneg
cnf(i_0_20506, plain, (v3_ordinal1(k10_ordinal2(X1,X2))|r2_hidden(k1_xboole_0,X2)|~v7_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_552, plain, (k10_ordinal2(X1,X2)=k10_ordinal2(X3,X2)|r2_hidden(k10_ordinal2(X1,X2),k10_ordinal2(X3,X2))|~v3_ordinal1(k10_ordinal2(X3,X2))|~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(X1,X3)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_553, plain, (k10_ordinal2(X1,esk5_2(X2,X3))=X3|r2_hidden(k10_ordinal2(X1,esk5_2(X2,X3)),X3)|~v3_ordinal1(X3)|~v3_ordinal1(k10_ordinal2(X1,esk5_2(X2,X3)))|~r1_ordinal1(X1,X2)|~r1_ordinal1(X2,X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_21038, plain, (k2_xboole_0(X1,k1_tarski(X1))=X1|~v3_ordinal1(X1)|~r2_hidden(X1,X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_21697, plain, (r2_hidden(X1,X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X1,X1))).#trainneg
cnf(i_0_21811, negated_conjecture, (~r1_ordinal1(k10_ordinal2(k2_xboole_0(X1,k1_tarski(X1)),esk1_0),X1)|~v3_ordinal1(k10_ordinal2(k2_xboole_0(X1,k1_tarski(X1)),esk1_0))|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1)|~r2_hidden(X1,X1))).#trainneg
cnf(i_0_21282, negated_conjecture, (k10_ordinal2(X1,esk1_0)=k10_ordinal2(X1,X2)|~r1_ordinal1(k10_ordinal2(X1,X2),X1)|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X1)|~r1_ordinal1(esk1_0,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_21687, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),X2)|~r2_hidden(k10_ordinal2(X2,esk1_0),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(X2,X2))).#trainneg
cnf(i_0_21665, plain, (r2_hidden(X1,esk6_1(X2))|v4_ordinal1(X2)|~r2_hidden(X1,X2)|~v3_ordinal1(esk6_1(X2))|~v3_ordinal1(X2)|~r2_hidden(X2,X2))).#trainneg
cnf(i_0_21254, plain, (v3_ordinal1(X1)|v4_ordinal1(X1)|X2!=k10_xtuple_0(esk4_2(X3,esk12_0))|~r2_hidden(X1,X2)|~v3_ordinal1(X3)|~r2_hidden(esk6_1(X1),esk12_0))).#trainneg
cnf(i_0_13026, negated_conjecture, (~r2_hidden(k10_ordinal2(X1,k3_ordinal2(k2_xboole_0(esk1_0,k1_tarski(esk1_0)))),X1)|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(k3_ordinal2(k2_xboole_0(esk1_0,k1_tarski(esk1_0))),k1_tarski(k3_ordinal2(k2_xboole_0(esk1_0,k1_tarski(esk1_0)))))))).#trainneg
cnf(i_0_554, plain, (k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2)))=X3|r2_hidden(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))),X3)|~v3_ordinal1(X3)|~v3_ordinal1(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))))|~r2_hidden(k10_ordinal2(X1,X2),X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_16305, plain, (esk13_0=esk12_0|r2_hidden(esk12_0,esk13_0)|r2_hidden(esk13_0,esk12_0))).#trainneg
cnf(i_0_21589, plain, (~r2_hidden(k1_tarski(k1_xboole_0),k1_tarski(k1_xboole_0)))).#trainneg
cnf(i_0_21334, negated_conjecture, (v7_ordinal1(k1_xboole_0)|r1_ordinal1(k1_xboole_0,esk2_0)|r2_hidden(esk2_0,k1_xboole_0)|r2_hidden(k1_xboole_0,esk2_0))).#trainneg
cnf(i_0_20504, plain, (v3_ordinal1(k10_ordinal2(X1,X2))|r2_hidden(k1_xboole_0,X1)|~v1_xboole_0(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_557, plain, (k10_ordinal2(X1,X2)=X1|r2_hidden(X1,k10_ordinal2(X1,X2))|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_559, plain, (m1_subset_1(k2_xboole_0(esk10_1(X1),k1_tarski(esk10_1(X1))),X2)|v1_xboole_0(X1)|~r1_tarski(X1,X2)|~v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_22019, plain, (r2_hidden(X1,k10_ordinal2(X1,X2))|~r2_hidden(X1,X1)|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_22417, plain, (k10_ordinal2(X1,X2)=X1|~r2_hidden(k10_ordinal2(X1,X2),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_22424, plain, (k10_ordinal2(X1,X2)=X1|~r1_ordinal1(X2,k1_xboole_0)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_22407, plain, (k10_ordinal2(k10_ordinal2(X1,X2),X3)=k10_ordinal2(X1,X2)|~r1_ordinal1(k10_ordinal2(X1,X2),X1)|~r2_hidden(X3,X2)|~v3_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_22527, plain, (~r2_hidden(k10_ordinal2(X1,X2),X1)|~r2_hidden(X1,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_22413, plain, (r2_hidden(X1,k3_ordinal2(k10_ordinal2(X2,X3)))|r2_hidden(X2,X1)|k10_ordinal2(X2,X3)=X2|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_22621, plain, (k10_ordinal2(X1,X2)=X1|r2_hidden(k1_xboole_0,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_22411, plain, (r2_hidden(X1,X2)|r2_hidden(X2,k10_ordinal2(X1,X3))|k10_ordinal2(X1,X3)=X1|~v3_ordinal1(k10_ordinal2(X1,X3))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_22144, plain, (v4_ordinal1(X1)|v3_ordinal1(X1)|~r2_hidden(esk6_1(X1),esk12_0)|~r2_hidden(X1,k10_xtuple_0(esk4_2(X2,esk12_0)))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_18629, plain, (r2_hidden(k10_ordinal2(k1_xboole_0,k1_xboole_0),k10_ordinal2(k10_ordinal2(k1_xboole_0,k1_xboole_0),X1))|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_560, plain, (m1_subset_1(k2_xboole_0(esk7_3(X1,X2,X3),k1_tarski(esk7_3(X1,X2,X3))),X4)|~r1_tarski(k9_xtuple_0(X1),X4)|~v4_ordinal1(k9_xtuple_0(X1))|~v3_ordinal1(k9_xtuple_0(X1))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_21350, plain, (r2_hidden(k1_xboole_0,X1)|v3_ordinal1(k10_ordinal2(X2,X1))|~v3_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_20507, plain, (v7_ordinal1(X1)|r2_hidden(X1,X2)|r2_hidden(X2,X1)|r2_hidden(k1_xboole_0,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_561, plain, (m1_subset_1(k2_xboole_0(esk7_3(esk4_2(X1,X2),X3,X4),k1_tarski(esk7_3(esk4_2(X1,X2),X3,X4))),X5)|~r1_tarski(X1,X5)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|X3!=k10_xtuple_0(esk4_2(X1,X2))|~r2_hidden(X4,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_563, plain, (m1_subset_1(k2_xboole_0(esk3_3(X1,X2,X3),k1_tarski(esk3_3(X1,X2,X3))),X4)|k10_ordinal2(X2,X1)=k4_ordinal2(X3)|X1=k1_xboole_0|~r1_tarski(X1,X4)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|k9_xtuple_0(X3)!=X1|~v1_ordinal2(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v5_ordinal1(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_22593, plain, (k10_ordinal2(X1,X2)=X1|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(X2,k1_xboole_0))).#trainneg
cnf(i_0_23828, plain, (r2_hidden(k10_ordinal2(X1,X2),X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(k2_xboole_0(X2,k1_tarski(X2)),k1_xboole_0)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2))))).#trainneg
cnf(i_0_23824, plain, (~r2_hidden(X1,k10_ordinal2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(k2_xboole_0(X2,k1_tarski(X2)),k1_xboole_0)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2))))).#trainneg
cnf(i_0_23329, negated_conjecture, (r2_hidden(X1,k10_ordinal2(k10_ordinal2(X2,X3),esk1_0))|k10_ordinal2(X2,X3)=X2|r2_hidden(X2,k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(k10_ordinal2(k10_ordinal2(X2,X3),esk1_0))|~v3_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(X2,X3))|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_22521, plain, (r2_hidden(X1,X2)|r2_hidden(X2,k10_ordinal2(X1,X3))|~v3_ordinal1(k10_ordinal2(X1,X3))|~v3_ordinal1(X2)|~r2_hidden(X1,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_21847, plain, (r1_ordinal1(k10_ordinal2(X1,X2),X3)|~r1_ordinal1(X1,X3)|~r2_hidden(k10_ordinal2(X1,X2),X3)|~v4_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~r2_hidden(X3,X3))).#trainneg
cnf(i_0_18701, plain, (r2_hidden(k10_ordinal2(k1_xboole_0,X1),k1_tarski(k1_xboole_0))|~r1_ordinal1(X1,k1_xboole_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_564, plain, (m1_subset_1(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))),X2)|~r1_tarski(X3,X2)|~v4_ordinal1(X3)|~v3_ordinal1(X3)|~r2_hidden(X1,X3)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_565, plain, (m1_subset_1(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2))),X3)|~r1_tarski(k10_ordinal2(X1,X4),X3)|~v4_ordinal1(k10_ordinal2(X1,X4))|~v3_ordinal1(k10_ordinal2(X1,X4))|~r2_hidden(X2,X4)|~v3_ordinal1(X1)|~v3_ordinal1(X4)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_566, plain, (m1_subset_1(k2_xboole_0(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))),k1_tarski(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))))),X3)|~r1_tarski(X4,X3)|~v4_ordinal1(X4)|~v3_ordinal1(X4)|~r2_hidden(k10_ordinal2(X1,X2),X4)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_24031, plain, (~r2_hidden(X2,k10_ordinal2(X2,X1))|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_21831, negated_conjecture, (v4_ordinal1(X1)|~r1_ordinal1(k10_ordinal2(esk6_1(X1),esk1_0),X1)|~v3_ordinal1(k10_ordinal2(esk6_1(X1),esk1_0))|~v3_ordinal1(esk6_1(X1))|~v3_ordinal1(X1)|~r2_hidden(X1,X1))).#trainneg
cnf(i_0_21839, plain, (r2_hidden(esk6_1(X1),esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|v4_ordinal1(X1)|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_ordinal1(esk6_1(X1),X1)|~r2_hidden(X1,X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(esk6_1(X1))|~v3_ordinal1(X2)|~r2_hidden(X1,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_22342, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)=esk6_1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_22201, plain, (v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|k10_ordinal2(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),k2_xboole_0(X1,k1_tarski(X1)))=X2|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k10_ordinal2(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1),X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_24483, plain, (X4=esk6_1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(X4,X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(esk7_3(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),X3,X4))|X3!=k10_xtuple_0(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))|~r2_hidden(X4,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_24135, plain, (v3_ordinal1(k10_ordinal2(k1_xboole_0,X1))|~r1_ordinal1(X1,k1_xboole_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_18702, plain, (r2_hidden(k10_ordinal2(X1,k1_xboole_0),k1_tarski(k1_xboole_0))|~r1_ordinal1(X1,k1_xboole_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_24414, plain, (~r2_hidden(X3,k1_xboole_0)|~v3_ordinal1(X3)|~v3_ordinal1(k1_funct_1(X1,X2))|k10_ordinal2(k1_funct_1(X1,X2),X3)!=k10_xtuple_0(X1)|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_567, plain, (m1_subset_1(k2_xboole_0(X1,k1_tarski(X1)),X2)|~r1_tarski(k3_ordinal2(X3),X2)|~v4_ordinal1(k3_ordinal2(X3))|~r2_hidden(X1,X3)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_18971, plain, (esk13_0=k1_xboole_0|r2_hidden(k1_xboole_0,esk13_0)|r2_hidden(esk13_0,k1_xboole_0))).#trainneg
cnf(i_0_23588, negated_conjecture, (v7_ordinal1(esk2_0)|r2_hidden(k1_xboole_0,X1)|r2_hidden(X1,esk2_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_23589, negated_conjecture, (v7_ordinal1(esk1_0)|r2_hidden(k1_xboole_0,X1)|r2_hidden(X1,esk1_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_568, plain, (m1_subset_1(k2_xboole_0(X1,k1_tarski(X1)),X2)|r2_hidden(X3,X1)|~r1_tarski(k2_xboole_0(X3,k1_tarski(X3)),X2)|~v4_ordinal1(k2_xboole_0(X3,k1_tarski(X3)))|~v3_ordinal1(k2_xboole_0(X3,k1_tarski(X3)))|~v3_ordinal1(X1)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_570, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k10_xtuple_0(X3))|~r1_tarski(k4_ordinal2(X3),X2)|~v3_ordinal1(X1)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v5_ordinal1(X3))).#trainneg
cnf(i_0_23253, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|k10_ordinal2(X2,X3)=X2|r2_hidden(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r1_ordinal1(k10_ordinal2(X2,X3),X1)|~v3_ordinal1(k10_ordinal2(X2,X3))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_24859, plain, (k10_ordinal2(X1,X2)=X1|r2_hidden(X1,esk6_1(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2)))))|v4_ordinal1(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2))))|~v3_ordinal1(esk6_1(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2)))))|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_24583, plain, (X1=esk6_1(k2_xboole_0(X2,k1_tarski(X2)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(X1,X2)|~r2_hidden(X1,k10_xtuple_0(esk4_2(X3,esk6_1(k2_xboole_0(X2,k1_tarski(X2))))))|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_21532, plain, (v1_xboole_0(X1)|v4_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))).#trainneg
cnf(i_0_22200, plain, (k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2)))=k10_ordinal2(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(k10_ordinal2(X1,X2),k10_ordinal2(X1,X2))|~v3_ordinal1(k10_ordinal2(X1,X2)))).#trainneg
cnf(i_0_24664, plain, (r2_hidden(X1,k1_tarski(k1_xboole_0))|~v3_ordinal1(X1)|~r1_ordinal1(X1,k1_xboole_0))).#trainneg
cnf(i_0_24753, negated_conjecture, (v7_ordinal1(esk2_0)|r2_hidden(k1_xboole_0,esk2_0))).#trainneg
cnf(i_0_571, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X3)|~v3_ordinal1(X1)|~r1_ordinal1(k3_ordinal2(X3),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_3672, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk1_0,esk2_0)|r2_hidden(esk15_0,esk1_0)|r2_hidden(esk2_0,esk15_0))).#trainneg
cnf(i_0_25146, negated_conjecture, (~r1_ordinal1(esk1_0,k1_xboole_0))).#trainneg
cnf(i_0_24764, negated_conjecture, (v7_ordinal1(esk1_0)|r2_hidden(esk2_0,esk1_0)|r2_hidden(k1_xboole_0,esk2_0))).#trainneg
cnf(i_0_573, plain, (r2_hidden(esk10_1(X1),X2)|r2_hidden(X2,X1)|v1_xboole_0(X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_574, plain, (r2_hidden(esk7_3(X1,X2,X3),X4)|r2_hidden(X4,k9_xtuple_0(X1))|~v3_ordinal1(k9_xtuple_0(X1))|~v3_ordinal1(X4)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_25078, plain, (k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2)))=k10_ordinal2(X1,X2)|v1_xboole_0(X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(k10_ordinal2(X1,X2),X1))).#trainneg
cnf(i_0_25048, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v1_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~r1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~r1_tarski(k2_xboole_0(X1,k1_tarski(X1)),esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_25113, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k10_ordinal2(k1_tarski(k1_xboole_0),esk1_0),X1)|~v3_ordinal1(X1)|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_xboole_0)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))).#trainneg
cnf(i_0_25090, plain, (k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2)))=k10_ordinal2(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(X2,k1_xboole_0)|~r2_hidden(X1,k10_ordinal2(X1,X2))|~v3_ordinal1(k10_ordinal2(X1,X2)))).#trainneg
cnf(i_0_25477, plain, (v1_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~r1_ordinal1(k2_xboole_0(X1,k1_tarski(X1)),esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_25347, plain, (r2_hidden(X1,k10_ordinal2(X1,X2))|v1_xboole_0(X3)|r2_hidden(X2,X3)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_25092, plain, (k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2)))=k10_ordinal2(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X2,X2))).#trainneg
cnf(i_0_25214, plain, (~r1_ordinal1(k10_ordinal2(k1_tarski(k1_xboole_0),k1_tarski(k1_xboole_0)),k1_xboole_0))).#trainneg
cnf(i_0_24699, plain, (esk13_0=k1_xboole_0|r2_hidden(k1_xboole_0,esk13_0))).#trainneg
cnf(i_0_24678, plain, (k10_ordinal2(k10_ordinal2(k1_funct_1(X1,X2),X3),k1_tarski(k1_xboole_0))!=k10_xtuple_0(X1)|~r2_hidden(X2,k9_xtuple_0(X1))|~r2_hidden(k2_xboole_0(X3,k1_tarski(X3)),k1_xboole_0)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_ordinal1(k1_funct_1(X1,X2))|~v3_ordinal1(k2_xboole_0(X3,k1_tarski(X3)))|~v3_ordinal1(X3))).#trainneg
cnf(i_0_575, plain, (r2_hidden(esk7_3(esk4_2(X1,X2),X3,X4),X5)|r2_hidden(X5,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X5)|X3!=k10_xtuple_0(esk4_2(X1,X2))|~r2_hidden(X4,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_3673, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk1_0,esk2_0)|r2_hidden(esk14_0,esk1_0)|r2_hidden(esk2_0,esk14_0))).#trainneg
cnf(i_0_24592, negated_conjecture, (v3_ordinal1(k10_ordinal2(k1_xboole_0,esk2_0))|r1_ordinal1(k1_xboole_0,esk2_0))).#trainneg
cnf(i_0_23626, plain, (v7_ordinal1(k1_xboole_0)|r2_hidden(k1_xboole_0,X1)|r2_hidden(X1,k1_xboole_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_577, plain, (r2_hidden(esk3_3(X1,X2,X3),X4)|r2_hidden(X4,X1)|k10_ordinal2(X2,X1)=k4_ordinal2(X3)|X1=k1_xboole_0|~v3_ordinal1(X1)|~v3_ordinal1(X4)|k9_xtuple_0(X3)!=X1|~v1_ordinal2(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v5_ordinal1(X3)|~v4_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_579, plain, (r2_hidden(k10_ordinal2(X1,X2),X3)|r2_hidden(X3,k10_ordinal2(X1,X4))|~v3_ordinal1(k10_ordinal2(X1,X4))|~v3_ordinal1(X3)|~r2_hidden(X2,X4)|~v3_ordinal1(X1)|~v3_ordinal1(X4)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_26003, plain, (v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_25077, plain, (k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2)))=k10_ordinal2(X1,X2)|v1_xboole_0(X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(k10_ordinal2(X1,X2),X1)|~v3_ordinal1(k10_ordinal2(X1,X2)))).#trainneg
cnf(i_0_24161, plain, (~r1_ordinal1(k2_xboole_0(k1_tarski(k1_xboole_0),k1_tarski(k1_tarski(k1_xboole_0))),k1_xboole_0)|~v3_ordinal1(k2_xboole_0(k1_tarski(k1_xboole_0),k1_tarski(k1_tarski(k1_xboole_0)))))).#trainneg
cnf(i_0_25047, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1))|~r1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X1))).#trainneg
cnf(i_0_580, plain, (r2_hidden(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))),X3)|r2_hidden(X3,X4)|~v3_ordinal1(X4)|~v3_ordinal1(X3)|~r2_hidden(k10_ordinal2(X1,X2),X4)|~v4_ordinal1(X4)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_3674, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk1_0,esk2_0)|r2_hidden(esk13_0,esk1_0)|r2_hidden(esk2_0,esk13_0))).#trainneg
cnf(i_0_582, plain, (r2_hidden(X1,X2)|r2_hidden(X2,k2_xboole_0(X3,k1_tarski(X3)))|r2_hidden(X3,X1)|~v3_ordinal1(k2_xboole_0(X3,k1_tarski(X3)))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_587, plain, (r1_tarski(esk7_3(X1,X2,X3),X4)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r1_tarski(k9_xtuple_0(X1),k1_zfmisc_1(X4)))).#trainneg
cnf(i_0_24462, negated_conjecture, (~r2_hidden(k10_ordinal2(X1,esk1_0),k10_ordinal2(X1,esk1_0))|~v3_ordinal1(k10_ordinal2(esk6_1(k10_ordinal2(X1,esk1_0)),esk1_0))|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~r1_ordinal1(esk6_1(k10_ordinal2(X1,esk1_0)),X1)|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,esk1_0))))).#trainneg
cnf(i_0_21705, plain, (r2_hidden(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_ordinal1(X1,X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X1,X1))).#trainneg
cnf(i_0_24457, negated_conjecture, (r2_hidden(X1,k10_ordinal2(esk6_1(X1),esk1_0))|~r2_hidden(X1,X1)|~v3_ordinal1(k10_ordinal2(esk6_1(X1),esk1_0))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_23323, negated_conjecture, (k10_ordinal2(X1,X2)=X1|r2_hidden(X1,k2_xboole_0(X3,k1_tarski(X3)))|~r1_ordinal1(k10_ordinal2(k10_ordinal2(X1,X2),esk1_0),X3)|~v3_ordinal1(k10_ordinal2(k10_ordinal2(X1,X2),esk1_0))|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X3)|~v3_ordinal1(k2_xboole_0(X3,k1_tarski(X3)))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_21671, plain, (r2_hidden(X1,esk6_1(X2))|v4_ordinal1(X2)|v4_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(esk6_1(X2))|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(X2,X2))).#trainneg
cnf(i_0_22036, negated_conjecture, (~r2_hidden(k10_ordinal2(X1,esk1_0),k10_ordinal2(X1,esk1_0))|~v3_ordinal1(k10_ordinal2(k2_xboole_0(k10_ordinal2(X1,esk1_0),k1_tarski(k10_ordinal2(X1,esk1_0))),esk1_0))|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~r1_ordinal1(k2_xboole_0(k10_ordinal2(X1,esk1_0),k1_tarski(k10_ordinal2(X1,esk1_0))),X1)|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(k10_ordinal2(X1,esk1_0),k1_tarski(k10_ordinal2(X1,esk1_0)))))).#trainneg
cnf(i_0_21035, plain, (v1_xboole_0(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|k2_xboole_0(X2,k1_tarski(X2))=X1|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(X2,X1)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2))))).#trainneg
cnf(i_0_21541, negated_conjecture, (v4_ordinal1(k10_ordinal2(X1,esk1_0))|~r1_ordinal1(esk6_1(k10_ordinal2(X1,esk1_0)),X1)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,esk1_0)))|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k10_ordinal2(X1,esk1_0)),esk6_1(k10_ordinal2(X1,esk1_0))))).#trainneg
cnf(i_0_21772, plain, (v4_ordinal1(X1)|~r2_hidden(k2_xboole_0(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))),k1_tarski(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))))),X1)|~v3_ordinal1(esk6_1(X1))|~v3_ordinal1(X1)|~r2_hidden(X1,X1))).#trainneg
cnf(i_0_21670, plain, (v1_xboole_0(X1)|v4_ordinal1(X2)|~r1_ordinal1(k10_ordinal2(X2,X1),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X2,X2))).#trainneg
cnf(i_0_25242, plain, (m1_subset_1(X1,X2)|~r1_tarski(k1_tarski(k1_xboole_0),X2)|~r1_ordinal1(X1,k1_xboole_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_24591, plain, (v3_ordinal1(k10_ordinal2(k1_xboole_0,esk10_1(k1_zfmisc_1(k1_xboole_0))))|~v3_ordinal1(esk10_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_277, plain, (v7_ordinal1(k1_tarski(k1_xboole_0)))).#trainneg
cnf(i_0_608, plain, (v3_ordinal1(k10_ordinal2(k1_tarski(k1_xboole_0),esk12_0)))).#trainneg
cnf(i_0_588, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk7_3(X3,X4,X5))|X4!=k10_xtuple_0(X3)|~r2_hidden(X5,X4)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r1_tarski(k9_xtuple_0(X3),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_3675, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk1_0,esk2_0)|r2_hidden(esk12_0,esk1_0)|r2_hidden(esk2_0,esk12_0))).#trainneg
cnf(i_0_24736, negated_conjecture, (v7_ordinal1(esk2_0)|r2_hidden(k10_ordinal2(k1_xboole_0,k1_xboole_0),esk2_0)|r2_hidden(k1_xboole_0,k10_ordinal2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_24785, negated_conjecture, (v7_ordinal1(esk1_0)|r2_hidden(k10_ordinal2(k1_xboole_0,k1_xboole_0),esk1_0)|r2_hidden(k1_xboole_0,k10_ordinal2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_589, plain, (v1_xboole_0(X1)|r2_hidden(esk7_3(X2,X3,X4),X1)|X3!=k10_xtuple_0(X2)|~r2_hidden(X4,X3)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r1_tarski(k9_xtuple_0(X2),X1))).#trainneg
cnf(i_0_590, plain, (v3_ordinal1(esk7_3(X1,X2,X3))|~v3_ordinal1(X4)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r1_tarski(k9_xtuple_0(X1),X4))).#trainneg
cnf(i_0_26670, negated_conjecture, (v4_ordinal1(k10_ordinal2(X1,esk1_0))|v1_xboole_0(esk6_1(k10_ordinal2(X1,esk1_0)))|~r1_ordinal1(esk6_1(k10_ordinal2(X1,esk1_0)),X1)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~v3_ordinal1(X1)|~r1_tarski(k10_ordinal2(X1,esk1_0),esk6_1(k10_ordinal2(X1,esk1_0))))).#trainneg
cnf(i_0_26664, plain, (k2_xboole_0(X1,k1_tarski(X1))=X2|v1_xboole_0(esk5_2(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1)|~r2_hidden(X1,X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_26744, plain, (~r2_hidden(X1,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_22414, plain, (r2_hidden(X1,k10_ordinal2(X1,k10_ordinal2(X2,X3)))|k10_ordinal2(X2,X3)=X2|~v3_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(X2,X3))|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_23387, plain, (r2_hidden(X1,k10_ordinal2(X2,X3))|k10_ordinal2(X2,X3)=X2|r2_hidden(X2,k10_ordinal2(X1,X4))|~v3_ordinal1(k10_ordinal2(X2,X3))|~v3_ordinal1(X1)|~v3_ordinal1(X4)|~v3_ordinal1(k10_ordinal2(X1,X4))|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_21250, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(k2_xboole_0(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),k1_tarski(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))),X1)|~v4_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),k1_tarski(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)))))).#trainneg
cnf(i_0_26757, plain, (m1_subset_1(X1,k1_tarski(k1_xboole_0))|~r1_ordinal1(X1,k1_xboole_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_17724, plain, (r2_hidden(k10_ordinal2(X1,k1_xboole_0),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_609, plain, (v3_ordinal1(k10_ordinal2(k1_tarski(k1_xboole_0),X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_591, plain, (v4_ordinal1(k3_ordinal2(X1))|~r2_hidden(k10_ordinal2(esk6_1(k3_ordinal2(X1)),k1_tarski(k1_xboole_0)),X1)|~v3_ordinal1(k10_ordinal2(esk6_1(k3_ordinal2(X1)),k1_tarski(k1_xboole_0))))).#trainneg
cnf(i_0_25304, negated_conjecture, (r2_hidden(k1_xboole_0,esk15_0)|esk2_0=esk1_0|r2_hidden(esk15_0,esk1_0)|r2_hidden(esk1_0,esk2_0))).#trainneg
cnf(i_0_24616, plain, (v3_ordinal1(k10_ordinal2(k1_xboole_0,X1))|r2_hidden(k1_xboole_0,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_25354, plain, (r2_hidden(k1_xboole_0,X1)|v1_xboole_0(X2)|r2_hidden(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_592, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|r2_hidden(X1,k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(k1_xboole_0)))|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_tarski(k1_xboole_0)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_593, plain, (r2_hidden(X1,k10_xtuple_0(esk4_2(X2,X3)))|X1!=k10_ordinal2(X3,X4)|~r2_hidden(X4,X2)|~v3_ordinal1(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_24138, plain, (v3_ordinal1(k10_ordinal2(X1,k10_ordinal2(k1_xboole_0,X2)))|~v3_ordinal1(X1)|~r1_ordinal1(X2,k1_xboole_0)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_24705, negated_conjecture, (v7_ordinal1(esk2_0)|r2_hidden(esk15_0,esk2_0)|r2_hidden(k1_xboole_0,esk15_0))).#trainneg
cnf(i_0_594, plain, (r1_ordinal1(esk15_0,k1_tarski(k1_xboole_0))|r1_ordinal1(k1_tarski(k1_xboole_0),esk15_0))).#trainneg
cnf(i_0_24707, negated_conjecture, (v7_ordinal1(esk2_0)|r2_hidden(esk13_0,esk2_0)|r2_hidden(k1_xboole_0,esk13_0))).#trainneg
cnf(i_0_27606, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(k1_xboole_0,esk2_0)|r2_hidden(esk2_0,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_595, plain, (r1_ordinal1(esk14_0,k1_tarski(k1_xboole_0))|r1_ordinal1(k1_tarski(k1_xboole_0),esk14_0))).#trainneg
cnf(i_0_596, negated_conjecture, (r1_ordinal1(esk2_0,k1_tarski(k1_xboole_0))|r1_ordinal1(k1_tarski(k1_xboole_0),esk2_0))).#trainneg
cnf(i_0_27749, plain, (r2_hidden(X1,k10_xtuple_0(esk4_2(X2,X3)))|~r2_hidden(esk5_2(X3,X1),X2)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~r1_ordinal1(X3,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_27167, plain, (k10_ordinal2(X3,X4)=X3|~r1_ordinal1(k10_ordinal2(X1,X2),X1)|~r2_hidden(k10_ordinal2(X3,X4),X2)|~v3_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X4))).#trainneg
cnf(i_0_27050, plain, (v1_xboole_0(esk6_1(X1))|v4_ordinal1(X1)|~v3_ordinal1(esk6_1(X1))|~r1_tarski(X1,esk6_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_27746, plain, (r2_hidden(k10_ordinal2(X1,X2),k10_xtuple_0(esk4_2(X3,X1)))|~r2_hidden(X2,X3)|~v3_ordinal1(X1)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_27080, negated_conjecture, (~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~r1_ordinal1(k10_ordinal2(X1,esk1_0),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_27178, plain, (k10_ordinal2(X2,X3)=X2|~r2_hidden(k10_ordinal2(X1,k10_ordinal2(X2,X3)),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_27344, plain, (v3_ordinal1(k10_ordinal2(X1,k1_xboole_0))|~v3_ordinal1(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_24629, plain, (v3_ordinal1(k10_ordinal2(X1,k10_ordinal2(X2,k1_xboole_0)))|~v3_ordinal1(X1)|~r1_ordinal1(X2,k1_xboole_0)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_24754, negated_conjecture, (v7_ordinal1(esk1_0)|r2_hidden(esk15_0,esk1_0)|r2_hidden(k1_xboole_0,esk15_0))).#trainneg
cnf(i_0_597, negated_conjecture, (r1_ordinal1(esk1_0,k1_tarski(k1_xboole_0))|r1_ordinal1(k1_tarski(k1_xboole_0),esk1_0))).#trainneg
cnf(i_0_24756, negated_conjecture, (v7_ordinal1(esk1_0)|r2_hidden(esk13_0,esk1_0)|r2_hidden(k1_xboole_0,esk13_0))).#trainneg
cnf(i_0_28322, negated_conjecture, (v3_ordinal1(k10_ordinal2(esk2_0,k1_xboole_0))|esk2_0=esk1_0|r2_hidden(esk1_0,esk2_0))).#trainneg
cnf(i_0_598, plain, (r1_ordinal1(X1,k1_tarski(k1_xboole_0))|r1_ordinal1(k1_tarski(k1_xboole_0),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_604, plain, (v3_ordinal1(k10_ordinal2(k10_ordinal2(X1,X2),X3))|~v1_xboole_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_28164, plain, (v1_xboole_0(esk6_1(X1))|v4_ordinal1(X1)|~v3_ordinal1(X1)|~r1_ordinal1(X1,esk6_1(X1))|~v3_ordinal1(esk6_1(X1)))).#trainneg
cnf(i_0_28281, plain, (k10_ordinal2(X1,k1_tarski(k1_xboole_0))=X1|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(k10_ordinal2(X2,X1),X2)|~v4_ordinal1(X2))).#trainneg
cnf(i_0_28752, plain, (v1_xboole_0(esk6_1(k10_ordinal2(X1,X2)))|v4_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(X1,esk6_1(k10_ordinal2(X1,X2)))|~v4_ordinal1(X2)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,X2)))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_28739, plain, (v1_xboole_0(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(X1,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_28749, plain, (v1_xboole_0(esk6_1(k10_ordinal2(X1,X2)))|v4_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(X2,k1_xboole_0)|~r1_ordinal1(X1,esk6_1(k10_ordinal2(X1,X2)))|~r2_hidden(esk6_1(k10_ordinal2(X1,X2)),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_28310, plain, (v3_ordinal1(k10_ordinal2(esk6_1(X1),k1_xboole_0))|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_27880, negated_conjecture, (r2_hidden(k1_xboole_0,esk2_0)|r2_hidden(esk2_0,esk14_0))).#trainneg
cnf(i_0_607, plain, (r2_hidden(esk15_0,X1)|r1_ordinal1(esk14_0,esk15_0)|~r2_hidden(esk14_0,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_27898, negated_conjecture, (X1=k1_xboole_0|r2_hidden(k1_xboole_0,esk2_0)|r2_hidden(esk2_0,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_28334, plain, (v3_ordinal1(k10_ordinal2(k10_ordinal2(k1_xboole_0,k1_xboole_0),k1_xboole_0)))).#trainneg
cnf(i_0_612, plain, (r1_ordinal1(esk15_0,k10_ordinal2(esk12_0,esk12_0))|r1_ordinal1(k10_ordinal2(esk12_0,esk12_0),esk15_0))).#trainneg
cnf(i_0_613, plain, (r1_ordinal1(esk14_0,k10_ordinal2(esk12_0,esk12_0))|r1_ordinal1(k10_ordinal2(esk12_0,esk12_0),esk14_0))).#trainneg
cnf(i_0_28869, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(esk5_2(X1,X1))|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r1_ordinal1(X1,X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_28821, negated_conjecture, (v1_xboole_0(esk6_1(k10_ordinal2(esk2_0,X1)))|v4_ordinal1(k10_ordinal2(esk2_0,X1))|~v4_ordinal1(X1)|~v3_ordinal1(esk6_1(k10_ordinal2(esk2_0,X1)))|~v3_ordinal1(X1)|~r2_hidden(esk2_0,esk6_1(k10_ordinal2(esk2_0,X1)))|~v3_ordinal1(k10_ordinal2(esk2_0,X1)))).#trainneg
cnf(i_0_28836, plain, (v1_xboole_0(esk6_1(k10_ordinal2(X1,X2)))|v4_ordinal1(k10_ordinal2(X1,X2))|r2_hidden(esk6_1(k10_ordinal2(X1,X2)),X1)|~v4_ordinal1(X2)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,X2)))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_28837, plain, (v4_ordinal1(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))))|~r2_hidden(k10_ordinal2(X1,X2),esk6_1(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2)))))|~v3_ordinal1(k10_ordinal2(X1,k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_28852, plain, (v4_ordinal1(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))))|~v3_ordinal1(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))))|~r2_hidden(X1,esk6_1(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1))))))|~v4_ordinal1(esk6_1(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1))))))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))))))).#trainneg
cnf(i_0_27089, plain, (~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(X1,k10_ordinal2(X1,X2))|~r2_hidden(k1_xboole_0,X2)|~r2_hidden(k10_ordinal2(X1,X2),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_27755, plain, (r2_hidden(X1,k10_xtuple_0(esk4_2(X2,X1)))|~r2_hidden(k1_xboole_0,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_28611, plain, (r1_ordinal1(k1_tarski(k1_xboole_0),k1_tarski(k1_xboole_0)))).#trainneg
cnf(i_0_28866, plain, (esk6_1(k1_tarski(k1_xboole_0))=k1_xboole_0|~v3_ordinal1(esk6_1(k1_tarski(k1_xboole_0))))).#trainneg
cnf(i_0_614, negated_conjecture, (r1_ordinal1(esk2_0,k10_ordinal2(esk12_0,esk12_0))|r1_ordinal1(k10_ordinal2(esk12_0,esk12_0),esk2_0))).#trainneg
cnf(i_0_29029, negated_conjecture, (r1_ordinal1(esk14_0,esk15_0)|r2_hidden(esk15_0,esk1_0)|esk14_0=esk1_0|r2_hidden(esk1_0,esk14_0))).#trainneg
cnf(i_0_28838, plain, (~r2_hidden(k1_xboole_0,esk6_1(k1_tarski(k1_xboole_0))))).#trainneg
cnf(i_0_29067, negated_conjecture, (v3_ordinal1(k10_ordinal2(esk2_0,k1_xboole_0))|X1=k1_xboole_0|r2_hidden(k1_xboole_0,esk2_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_615, negated_conjecture, (r1_ordinal1(esk1_0,k10_ordinal2(esk12_0,esk12_0))|r1_ordinal1(k10_ordinal2(esk12_0,esk12_0),esk1_0))).#trainneg
cnf(i_0_616, plain, (r1_ordinal1(X1,k10_ordinal2(esk12_0,esk12_0))|r1_ordinal1(k10_ordinal2(esk12_0,esk12_0),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_29221, plain, (v1_xboole_0(esk6_1(k10_ordinal2(X1,X2)))|r2_hidden(esk6_1(k10_ordinal2(X1,X2)),X1)|v4_ordinal1(k10_ordinal2(X1,X2))|~v4_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(k10_ordinal2(X1,X2)))).#trainneg
cnf(i_0_29198, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(esk5_2(X1,X1))|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1))))).#trainneg
cnf(i_0_29468, plain, (esk6_1(k1_tarski(k1_xboole_0))=k1_xboole_0)).#trainneg
cnf(i_0_29269, plain, (v4_ordinal1(k2_xboole_0(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))),k1_tarski(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))))))|r2_hidden(esk6_1(k2_xboole_0(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))),k1_tarski(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1)))))),X1)|v4_ordinal1(X1)|~v4_ordinal1(esk6_1(k2_xboole_0(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))),k1_tarski(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1)))))))|~v3_ordinal1(k2_xboole_0(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))),k1_tarski(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))))))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(k2_xboole_0(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))),k1_tarski(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))))))))).#trainneg
cnf(i_0_29333, plain, (v4_ordinal1(esk7_3(esk4_2(X2,X1),X3,X1))|~r2_hidden(k1_xboole_0,esk6_1(esk7_3(esk4_2(X2,X1),X3,X1)))|~v3_ordinal1(X1)|~v3_ordinal1(esk6_1(esk7_3(esk4_2(X2,X1),X3,X1)))|X3!=k10_xtuple_0(esk4_2(X2,X1))|~r2_hidden(X1,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_29304, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k1_xboole_0,X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X2)|~r2_hidden(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_28865, plain, (v4_ordinal1(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2))))|~v3_ordinal1(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2))))|~r1_ordinal1(X2,k1_xboole_0)|~r2_hidden(X1,esk6_1(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2)))))|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(esk6_1(k2_xboole_0(k10_ordinal2(X1,X2),k1_tarski(k10_ordinal2(X1,X2)))))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_27086, plain, (v1_xboole_0(X2)|~v3_ordinal1(k10_ordinal2(X1,X2))|~r2_hidden(k10_ordinal2(X1,X2),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_27038, plain, (k2_xboole_0(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),k1_tarski(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)))=X1|v1_xboole_0(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k10_ordinal2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_29836, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(X1,esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_29614, plain, (r1_ordinal1(k10_ordinal2(esk12_0,esk12_0),k10_ordinal2(esk12_0,esk12_0)))).#trainneg
cnf(i_0_10766, plain, (X1!=k10_xtuple_0(esk4_2(X2,k1_xboole_0))|~r2_hidden(X3,k1_xboole_0)|~r2_hidden(X3,X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_618, negated_conjecture, (r2_hidden(esk2_0,X1)|r1_ordinal1(esk1_0,esk2_0)|~r2_hidden(esk1_0,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_29502, negated_conjecture, (esk15_0=esk14_0|r2_hidden(esk14_0,esk15_0)|esk14_0=esk1_0|r2_hidden(esk1_0,esk14_0)|r2_hidden(esk15_0,esk1_0))).#trainneg
cnf(i_0_29564, negated_conjecture, (k10_ordinal2(X1,k1_xboole_0)=k1_xboole_0|r2_hidden(k1_xboole_0,esk2_0)|v3_ordinal1(k10_ordinal2(esk2_0,k1_xboole_0))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_619, plain, (r1_ordinal1(esk15_0,k9_xtuple_0(esk4_2(X1,X2)))|r1_ordinal1(k9_xtuple_0(esk4_2(X1,X2)),esk15_0)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_620, plain, (r1_ordinal1(esk14_0,k9_xtuple_0(esk4_2(X1,X2)))|r1_ordinal1(k9_xtuple_0(esk4_2(X1,X2)),esk14_0)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_29889, plain, (~r2_hidden(k10_ordinal2(k10_ordinal2(X2,X1),k1_tarski(k1_xboole_0)),X2)|~v3_ordinal1(X2)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_29778, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k2_xboole_0(X2,k1_tarski(X2)),X1)|~r2_hidden(k1_xboole_0,k2_xboole_0(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),k1_tarski(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(k2_xboole_0(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2),k1_tarski(esk5_2(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))))|~v3_ordinal1(X1)|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_9039, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),k1_xboole_0)|~r2_hidden(k1_xboole_0,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_29930, plain, (~v3_ordinal1(X1)|~r2_hidden(k10_ordinal2(X1,k1_xboole_0),X1)|~v4_ordinal1(X1))).#trainneg
cnf(i_0_29775, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(X4,X1)|~r2_hidden(k1_xboole_0,esk7_3(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),X3,X4))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(esk7_3(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),X3,X4))|~v3_ordinal1(X1)|X3!=k10_xtuple_0(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))|~r2_hidden(X4,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_28335, plain, (v3_ordinal1(k10_ordinal2(k10_ordinal2(k1_xboole_0,X1),k1_xboole_0))|~r1_ordinal1(X1,k1_xboole_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_27882, negated_conjecture, (r2_hidden(k1_xboole_0,esk2_0)|r2_hidden(esk2_0,esk12_0))).#trainneg
cnf(i_0_621, negated_conjecture, (r1_ordinal1(esk2_0,k9_xtuple_0(esk4_2(X1,X2)))|r1_ordinal1(k9_xtuple_0(esk4_2(X1,X2)),esk2_0)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_29550, negated_conjecture, (k10_ordinal2(k1_xboole_0,k1_xboole_0)=k1_xboole_0|r2_hidden(k1_xboole_0,esk2_0)|v3_ordinal1(k10_ordinal2(esk2_0,k1_xboole_0)))).#trainneg
cnf(i_0_30127, negated_conjecture, (k10_ordinal2(esk2_0,k1_xboole_0)=k1_xboole_0|r2_hidden(k1_xboole_0,esk2_0)|v3_ordinal1(k10_ordinal2(esk2_0,k1_xboole_0)))).#trainneg
cnf(i_0_622, negated_conjecture, (r1_ordinal1(esk1_0,k9_xtuple_0(esk4_2(X1,X2)))|r1_ordinal1(k9_xtuple_0(esk4_2(X1,X2)),esk1_0)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_623, plain, (r1_ordinal1(X1,k9_xtuple_0(esk4_2(X2,X3)))|r1_ordinal1(k9_xtuple_0(esk4_2(X2,X3)),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_29814, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(esk7_3(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),X3,X1))|~r2_hidden(k1_xboole_0,esk6_1(esk7_3(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),X3,X1)))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(esk6_1(esk7_3(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),X3,X1)))|~v3_ordinal1(X1)|X3!=k10_xtuple_0(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))|~r2_hidden(X1,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_29813, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(esk5_2(X2,X1))|~r2_hidden(k1_xboole_0,esk6_1(esk5_2(X2,X1)))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(esk6_1(esk5_2(X2,X1)))|~v3_ordinal1(X1)|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~r1_ordinal1(X2,X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_29884, plain, (v1_xboole_0(esk7_3(esk4_2(X1,X2),X3,X4))|~r2_hidden(X4,X2)|~v3_ordinal1(X2)|~v3_ordinal1(esk7_3(esk4_2(X1,X2),X3,X4))|X3!=k10_xtuple_0(esk4_2(X1,X2))|~r2_hidden(X4,X3)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_29830, plain, (v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_ordinal1(esk5_2(X1,X2),k1_xboole_0)|~r2_hidden(X1,esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2)|~v3_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_29885, plain, (v1_xboole_0(esk5_2(X1,X2))|~r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_28336, plain, (v3_ordinal1(k10_ordinal2(k10_ordinal2(X1,k1_xboole_0),k1_xboole_0))|~r1_ordinal1(X1,k1_xboole_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_13784, negated_conjecture, (~r2_hidden(k10_ordinal2(k1_xboole_0,esk1_0),k1_xboole_0))).#trainneg
cnf(i_0_625, plain, (m1_subset_1(k10_ordinal2(X1,X2),X3)|~r2_hidden(X2,esk7_3(esk4_2(X4,X1),X5,X6))|~r1_tarski(X6,X3)|~v3_ordinal1(X1)|~v3_ordinal1(esk7_3(esk4_2(X4,X1),X5,X6))|X5!=k10_xtuple_0(esk4_2(X4,X1))|~r2_hidden(X6,X5)|~v3_ordinal1(X4))).#trainneg
cnf(i_0_626, plain, (m1_subset_1(k10_ordinal2(X1,X2),X3)|~r2_hidden(X2,esk5_2(X1,X4))|~r1_tarski(X4,X3)|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X1,X4))|~r1_ordinal1(X1,X4)|~v3_ordinal1(X4))).#trainneg
cnf(i_0_627, plain, (m1_subset_1(k10_ordinal2(X1,X2),X3)|~r2_hidden(X2,k2_xboole_0(esk5_2(X1,X4),k1_tarski(esk5_2(X1,X4))))|~r1_tarski(k2_xboole_0(X4,k1_tarski(X4)),X3)|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(esk5_2(X1,X4),k1_tarski(esk5_2(X1,X4))))|~r1_ordinal1(X1,X4)|~v3_ordinal1(X4))).#trainneg
cnf(i_0_30682, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(esk5_2(k2_xboole_0(X2,k1_tarski(X2)),X1))|~r1_ordinal1(k2_xboole_0(X2,k1_tarski(X2)),X1)|~r2_hidden(k1_xboole_0,esk6_1(esk5_2(k2_xboole_0(X2,k1_tarski(X2)),X1)))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(esk6_1(esk5_2(k2_xboole_0(X2,k1_tarski(X2)),X1)))|~v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_24394, plain, (~r2_hidden(X1,X2)|~r2_hidden(esk5_2(X1,X2),k1_xboole_0)|~v3_ordinal1(esk5_2(X1,X2))|~v3_ordinal1(X1)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_27069, plain, (~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(X2,k1_xboole_0)|~r2_hidden(X1,k10_ordinal2(X1,X2))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_29268, plain, (v4_ordinal1(k2_xboole_0(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))),k1_tarski(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))))))|v1_xboole_0(esk6_1(k2_xboole_0(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))),k1_tarski(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1)))))))|v4_ordinal1(X1)|~v4_ordinal1(esk6_1(k2_xboole_0(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))),k1_tarski(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1)))))))|~v3_ordinal1(k2_xboole_0(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))),k1_tarski(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))))))|~r1_tarski(X1,esk6_1(k2_xboole_0(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1))),k1_tarski(k2_xboole_0(esk6_1(X1),k1_tarski(esk6_1(X1)))))))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_21178, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~r2_hidden(k2_xboole_0(k2_xboole_0(X4,k1_tarski(X4)),k1_tarski(k2_xboole_0(X4,k1_tarski(X4)))),X1)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~v3_ordinal1(esk7_3(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))),X3,X4))|X3!=k10_xtuple_0(esk4_2(X2,esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))|~r2_hidden(X4,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_21312, plain, (X2=k10_ordinal2(X1,X3)|r2_hidden(X2,k10_ordinal2(X1,X3))|~r1_ordinal1(esk5_2(X1,X2),X3)|~v3_ordinal1(X1)|~v3_ordinal1(X3)|~v3_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_31147, plain, (X1=k10_ordinal2(X2,X3)|r2_hidden(X1,k10_ordinal2(X2,X3))|r2_hidden(X3,esk5_2(X2,X1))|~r1_ordinal1(X2,X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3)|~v3_ordinal1(X1)|~v3_ordinal1(esk5_2(X2,X1)))).#trainneg
cnf(i_0_28366, plain, (v3_ordinal1(k10_ordinal2(k10_ordinal2(X1,X2),k1_xboole_0))|~v3_ordinal1(X1)|~r2_hidden(X2,k1_xboole_0)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_27887, negated_conjecture, (r2_hidden(k1_xboole_0,esk2_0)|r2_hidden(esk2_0,esk1_0))).#trainneg
cnf(i_0_7244, plain, (~r2_hidden(k10_ordinal2(k1_xboole_0,k1_xboole_0),k1_xboole_0))).#trainneg
cnf(i_0_628, plain, (m1_subset_1(k10_ordinal2(X1,X2),X3)|~r2_hidden(X2,k1_xboole_0)|~r1_tarski(X1,X3)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_30128, negated_conjecture, (k10_ordinal2(esk1_0,k1_xboole_0)=k1_xboole_0|r2_hidden(k1_xboole_0,esk2_0)|v3_ordinal1(k10_ordinal2(esk2_0,k1_xboole_0)))).#trainneg
cnf(i_0_629, plain, (m1_subset_1(k10_ordinal2(X1,X2),X3)|~r2_hidden(X2,X4)|~v3_ordinal1(X1)|~v3_ordinal1(X4)|~r1_ordinal1(k10_ordinal2(X1,X4),X3)|~v3_ordinal1(X3)|~v3_ordinal1(k10_ordinal2(X1,X4)))).#trainneg
cnf(i_0_632, plain, (r2_hidden(k1_xboole_0,k9_xtuple_0(X1))|~v3_ordinal1(k9_xtuple_0(X1))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_30989, plain, (~r1_ordinal1(esk5_2(X1,X2),k1_xboole_0)|~r2_hidden(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_29922, plain, (v1_xboole_0(esk6_1(esk5_2(X1,X2)))|v4_ordinal1(esk5_2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(esk6_1(esk5_2(X1,X2)))|~r1_ordinal1(X2,X1)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_26183, plain, (r2_hidden(X1,k10_ordinal2(X2,k10_ordinal2(X3,X4)))|r2_hidden(k10_ordinal2(X2,X3),X1)|k1_xboole_0=X4|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(k10_ordinal2(X3,X4))|~v3_ordinal1(X3)|~v3_ordinal1(X4))).#trainneg
cnf(i_0_26218, plain, (k10_ordinal2(k10_ordinal2(X1,X2),k2_xboole_0(X2,k1_tarski(X2)))=k10_ordinal2(k10_ordinal2(X1,X2),X2)|v1_xboole_0(X2)|~v3_ordinal1(X2)|~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(k10_ordinal2(X1,X2),X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_21057, plain, (v4_ordinal1(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r1_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1)|~r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~r2_hidden(X1,X2)|~v4_ordinal1(X2)|~v3_ordinal1(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))))|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2)|~r1_ordinal1(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X1)),k1_tarski(k2_xboole_0(X1,k1_tarski(X1)))),X2))).#trainneg
cnf(i_0_633, plain, (r2_hidden(k1_xboole_0,X1)|~v3_ordinal1(X1)|X3!=k10_xtuple_0(esk4_2(X1,X2))|~r2_hidden(X4,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_29069, negated_conjecture, (v3_ordinal1(k10_ordinal2(X1,esk2_0))|X2=k1_xboole_0|r2_hidden(k1_xboole_0,esk2_0)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_637, plain, (r2_hidden(k1_xboole_0,k10_ordinal2(X1,X2))|~v3_ordinal1(k10_ordinal2(X1,X2))|~r2_hidden(X3,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_29294, plain, (~r2_hidden(X4,X1)|~r2_hidden(k1_xboole_0,esk7_3(esk4_2(X2,X1),X3,X4))|~v3_ordinal1(X1)|~v3_ordinal1(esk7_3(esk4_2(X2,X1),X3,X4))|X3!=k10_xtuple_0(esk4_2(X2,X1))|~r2_hidden(X4,X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_27085, plain, (v1_xboole_0(X2)|~v3_ordinal1(k10_ordinal2(X1,X2))|~r1_ordinal1(k10_ordinal2(X1,X2),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_22421, plain, (r2_hidden(k1_xboole_0,k10_ordinal2(X1,X2))|k10_ordinal2(X1,X2)=X1|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_32131, plain, (r2_hidden(k1_xboole_0,k10_ordinal2(X1,k1_tarski(k1_xboole_0)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_32104, plain, (r2_hidden(k1_xboole_0,k10_ordinal2(X1,X2))|v4_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_639, plain, (r2_hidden(k1_xboole_0,k3_ordinal2(X1))|~r2_hidden(X2,X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_640, plain, (r2_hidden(k1_xboole_0,k2_xboole_0(X1,k1_tarski(X1)))|r2_hidden(X1,X2)|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_32656, plain, (v4_ordinal1(X2)|~r2_hidden(X1,X2)|~r2_hidden(esk6_1(X2),X1)|~v4_ordinal1(X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_32879, plain, (v4_ordinal1(k3_ordinal2(X1))|~r2_hidden(k3_ordinal2(X1),X1))).#trainneg
cnf(i_0_32455, plain, (r2_hidden(k1_xboole_0,k3_ordinal2(k1_tarski(k1_xboole_0))))).#trainneg
cnf(i_0_664, plain, (r1_ordinal1(X1,X2)|r2_hidden(k1_xboole_0,esk7_3(esk4_2(X4,X2),X3,X1))|X3!=k10_xtuple_0(esk4_2(X4,X2))|~r2_hidden(X1,X3)|~v3_ordinal1(esk7_3(esk4_2(X4,X2),X3,X1))|~v3_ordinal1(X2)|~v3_ordinal1(X4))).#trainneg
cnf(i_0_32346, negated_conjecture, (r2_hidden(k1_xboole_0,k10_ordinal2(esk2_0,X1))|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_32347, negated_conjecture, (r2_hidden(k1_xboole_0,k10_ordinal2(esk1_0,X1))|v4_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_666, plain, (r1_tarski(esk7_3(esk4_2(X1,X2),X3,X4),X5)|X3!=k10_xtuple_0(esk4_2(X1,X2))|~r2_hidden(X4,X3)|~r1_tarski(X1,k1_zfmisc_1(X5))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_667, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk7_3(esk4_2(X3,X4),X5,X6))|X5!=k10_xtuple_0(esk4_2(X3,X4))|~r2_hidden(X6,X5)|~r1_tarski(X3,k1_zfmisc_1(X2))|~v3_ordinal1(X4)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_28822, plain, (v1_xboole_0(esk6_1(k10_ordinal2(k2_xboole_0(X1,k1_tarski(X1)),X2)))|v4_ordinal1(k10_ordinal2(k2_xboole_0(X1,k1_tarski(X1)),X2))|~v4_ordinal1(X2)|~v3_ordinal1(esk6_1(k10_ordinal2(k2_xboole_0(X1,k1_tarski(X1)),X2)))|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X2)|~r2_hidden(X1,esk6_1(k10_ordinal2(k2_xboole_0(X1,k1_tarski(X1)),X2)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_23360, negated_conjecture, (k10_ordinal2(X1,X2)=X1|r2_hidden(X1,k10_ordinal2(X3,esk1_0))|~r1_ordinal1(k10_ordinal2(X1,X2),X3)|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X3)|~v3_ordinal1(k10_ordinal2(X3,esk1_0))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_33054, plain, (r1_ordinal1(X4,X3)|X1!=k10_xtuple_0(esk4_2(X2,X3))|~r2_hidden(X4,X3)|~r2_hidden(X4,X1)|~v3_ordinal1(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_22150, negated_conjecture, (~r2_hidden(k10_ordinal2(X1,k3_ordinal2(k2_xboole_0(esk1_0,k1_tarski(esk1_0)))),X1)|~v3_ordinal1(k10_ordinal2(k3_ordinal2(k2_xboole_0(esk1_0,k1_tarski(esk1_0))),k1_tarski(k1_xboole_0)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_668, plain, (v1_xboole_0(X1)|r2_hidden(esk7_3(esk4_2(X2,X3),X4,X5),X1)|X4!=k10_xtuple_0(esk4_2(X2,X3))|~r2_hidden(X5,X4)|~r1_tarski(X2,X1)|~v3_ordinal1(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_25566, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(esk2_0,k10_ordinal2(esk2_0,X2))|r2_hidden(X2,X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_670, plain, (r1_ordinal1(esk15_0,k10_ordinal2(X1,k1_tarski(k1_xboole_0)))|r1_ordinal1(k10_ordinal2(X1,k1_tarski(k1_xboole_0)),esk15_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_32912, plain, (v4_ordinal1(X1)|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1)|~v4_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(X1),X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X2))).#trainneg
cnf(i_0_33472, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k10_ordinal2(X2,X3))|v1_xboole_0(X3)|~r2_hidden(esk6_1(k10_ordinal2(X2,X3)),X1)|~v4_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(k10_ordinal2(X2,X3))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_33434, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(esk6_1(esk6_1(k2_xboole_0(X2,k1_tarski(X2)))),X1)|~v4_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_33441, plain, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k3_ordinal2(X2))|~r2_hidden(esk6_1(k3_ordinal2(X2)),X1)|~v4_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1)))))).#trainneg
cnf(i_0_33462, negated_conjecture, (v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k10_ordinal2(X2,esk1_0))|~r2_hidden(esk6_1(k10_ordinal2(X2,esk1_0)),X1)|~v4_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(k10_ordinal2(X2,esk1_0))|~v3_ordinal1(X1)|~r2_hidden(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_33512, plain, (m1_subset_1(esk6_1(k10_ordinal2(X1,X2)),X3)|v4_ordinal1(k10_ordinal2(X1,X2))|~r1_tarski(X1,X3)|~v4_ordinal1(X2)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,X2)))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_33533, plain, (v4_ordinal1(k10_ordinal2(X1,X2))|~r2_hidden(X1,k10_ordinal2(X1,X2))|~v4_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(X1,X2))|~v4_ordinal1(X2)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,X2)))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_33524, plain, (v4_ordinal1(k10_ordinal2(X1,X2))|~r2_hidden(X1,esk6_1(k10_ordinal2(X1,X2)))|~v4_ordinal1(X2)|~v3_ordinal1(esk6_1(k10_ordinal2(X1,X2)))|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_32697, plain, (v7_ordinal1(X1)|v4_ordinal1(X1)|~v7_ordinal1(esk6_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_671, plain, (r1_ordinal1(esk14_0,k10_ordinal2(X1,k1_tarski(k1_xboole_0)))|r1_ordinal1(k10_ordinal2(X1,k1_tarski(k1_xboole_0)),esk14_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_33334, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(esk2_0,k10_ordinal2(esk2_0,esk2_0))|r2_hidden(esk2_0,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_33335, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(esk2_0,k10_ordinal2(esk2_0,esk1_0))|r2_hidden(esk1_0,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_672, negated_conjecture, (r1_ordinal1(esk2_0,k10_ordinal2(X1,k1_tarski(k1_xboole_0)))|r1_ordinal1(k10_ordinal2(X1,k1_tarski(k1_xboole_0)),esk2_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_673, negated_conjecture, (r1_ordinal1(esk1_0,k10_ordinal2(X1,k1_tarski(k1_xboole_0)))|r1_ordinal1(k10_ordinal2(X1,k1_tarski(k1_xboole_0)),esk1_0)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_34018, negated_conjecture, (r2_hidden(esk2_0,k10_ordinal2(esk2_0,esk1_0)))).#trainneg
cnf(i_0_33848, plain, (v4_ordinal1(k10_ordinal2(X1,X2))|v1_xboole_0(X2)|~v4_ordinal1(X1)|~v4_ordinal1(X2)|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_33850, plain, (v4_ordinal1(k10_ordinal2(X1,X2))|k10_ordinal2(X1,X2)=X1|~v4_ordinal1(X1)|~v4_ordinal1(X2)|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_33849, plain, (v4_ordinal1(k10_ordinal2(X1,X2))|k1_xboole_0=X2|~v4_ordinal1(X1)|~v4_ordinal1(X2)|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_33862, plain, (v4_ordinal1(X1)|~r2_hidden(X1,esk6_1(X1))|~v3_ordinal1(esk6_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_33788, plain, (v4_ordinal1(X2)|~r2_hidden(X1,X2)|~v4_ordinal1(X1)|~v4_ordinal1(esk5_2(X1,X2))|~v3_ordinal1(X2)|~v3_ordinal1(esk5_2(X1,X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_33853, negated_conjecture, (v4_ordinal1(k10_ordinal2(X1,X2))|r2_hidden(X2,esk1_0)|~v4_ordinal1(X1)|~v4_ordinal1(X2)|~v3_ordinal1(k10_ordinal2(X1,X2))|~v3_ordinal1(X2)|~r1_ordinal1(X1,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_33895, plain, (v7_ordinal1(X1)|v4_ordinal1(X1)|~v3_ordinal1(X1)|~v1_xboole_0(esk6_1(X1)))).#trainneg
cnf(i_0_33024, plain, (r2_hidden(k1_xboole_0,k3_ordinal2(k3_ordinal2(k1_tarski(k1_xboole_0)))))).#trainneg
cnf(i_0_9986, negated_conjecture, (esk2_0=esk15_0|r2_hidden(k1_xboole_0,esk15_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_33313, plain, (~v1_xboole_0(X1)|~r1_tarski(X6,X1)|X4!=k10_xtuple_0(esk4_2(X2,X3))|~r2_hidden(X5,X4)|~r1_tarski(X2,X6)|~v3_ordinal1(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_674, plain, (r1_ordinal1(X1,k10_ordinal2(X2,k1_tarski(k1_xboole_0)))|r1_ordinal1(k10_ordinal2(X2,k1_tarski(k1_xboole_0)),X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_34180, negated_conjecture, (~v4_ordinal1(esk2_0))).#trainneg
cnf(i_0_676, plain, (r2_hidden(esk13_0,X1)|r1_ordinal1(esk15_0,esk13_0)|~r2_hidden(esk15_0,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_679, plain, (r2_hidden(esk12_0,X1)|r1_ordinal1(esk15_0,esk12_0)|~r2_hidden(esk15_0,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_33658, plain, (v4_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|v4_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|v4_ordinal1(k2_xboole_0(X2,k1_tarski(X2)))|v1_xboole_0(X2)|~r2_hidden(esk6_1(k2_xboole_0(X2,k1_tarski(X2))),X1)|~v4_ordinal1(esk6_1(k2_xboole_0(X2,k1_tarski(X2))))|~v3_ordinal1(esk6_1(k2_xboole_0(X1,k1_tarski(X1))))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_tarski(esk6_1(k2_xboole_0(X1,k1_tarski(X1))),X2))).#trainneg
cnf(i_0_33834, negated_conjecture, (v4_ordinal1(k10_ordinal2(X1,esk1_0))|~v4_ordinal1(X1)|~v3_ordinal1(k10_ordinal2(X1,esk1_0))|~r1_ordinal1(X1,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_33535, plain, (v1_xboole_0(X1)|v4_ordinal1(k10_ordinal2(X2,X1))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v4_ordinal1(X1)|~v3_ordinal1(esk6_1(k10_ordinal2(X2,X1))))).#trainneg
cnf(i_0_32791, plain, (v3_ordinal1(k10_ordinal2(X1,esk12_0))|v4_ordinal1(X1)|~v7_ordinal1(esk6_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_33251, negated_conjecture, (~r2_hidden(k10_ordinal2(X1,k3_ordinal2(k10_ordinal2(esk1_0,k1_tarski(k1_xboole_0)))),X1)|~v3_ordinal1(k10_ordinal2(k3_ordinal2(k10_ordinal2(esk1_0,k1_tarski(k1_xboole_0))),k1_tarski(k1_xboole_0)))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_681, negated_conjecture, (r2_hidden(esk2_0,X1)|r1_ordinal1(esk15_0,esk2_0)|~r2_hidden(esk15_0,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_34458, negated_conjecture, (r1_ordinal1(esk15_0,esk13_0)|r2_hidden(esk13_0,esk1_0)|esk14_0=esk1_0|esk15_0=esk14_0|r2_hidden(esk1_0,esk14_0)|r2_hidden(esk14_0,esk15_0))).#trainneg
cnf(i_0_34391, negated_conjecture, (~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_25567, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(esk1_0,k10_ordinal2(esk1_0,X2))|r2_hidden(X2,X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_25604, plain, (v1_xboole_0(X1)|r2_hidden(k1_xboole_0,k10_ordinal2(k1_xboole_0,X2))|r2_hidden(X2,X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_683, negated_conjecture, (r2_hidden(esk15_0,X1)|r1_ordinal1(esk1_0,esk15_0)|~r2_hidden(esk1_0,X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_685, negated_conjecture, (r2_hidden(esk2_0,X1)|r1_ordinal1(k1_xboole_0,esk2_0)|~r2_hidden(k1_xboole_0,X1)|~v3_ordinal1(X1))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 67
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 102
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 100
# Processed clauses                    : 2185
# ...of these trivial                  : 13
# ...subsumed                          : 785
# ...remaining for further processing  : 1387
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 132
# Backward-rewritten                   : 40
# Generated clauses                    : 34837
# ...of the previous two non-trivial   : 33257
# Contextual simplify-reflections      : 521
# Paramodulations                      : 34789
# Factorizations                       : 2
# Equation resolutions                 : 33
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1201
#    Positive orientable unit clauses  : 62
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 31
#    Non-unit-clauses                  : 1107
# Current number of unprocessed clauses: 31005
# ...number of literals in the above   : 216749
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1573
# Proof object given clauses           : 101
# Proof search given clauses           : 1387
# Clause-clause subsumption calls (NU) : 330818
# Rec. Clause-clause subsumption calls : 22544
# Non-unit clause-clause subsumptions  : 1270
# Unit Clause-clause subsumption calls : 8452
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 18
# BW rewrite match successes           : 14
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1040629

# -------------------------------------------------
# User time                : 49.832 s
# System time              : 1.024 s
# Total time               : 50.856 s
# Maximum resident set size: 464524 pages
49.83user 1.11system 0:52.80elapsed 96%CPU (0avgtext+0avgdata 464524maxresident)k
0inputs+648outputs (0major+347851minor)pagefaults 0swaps
