2020-09-01 00:16:04.051025: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_49
2020-09-01 00:16:04.165019: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-09-01 00:16:04.372835: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-09-01 00:16:05.102278: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-09-01 00:16:08.640425: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 4589440 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_49' loaded (query=128; context=512; weigths=real; len_mult=0.000000).
# Preprocessing time       : 5.953 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t2_matrix12, conjecture, ![X1]:(v7_ordinal1(X1)=>![X2]:(v7_ordinal1(X2)=>![X3]:(v7_ordinal1(X3)=>![X4]:(v7_ordinal1(X4)=>![X5]:((((((((((((~(v2_struct_0(X5))&~(v6_struct_0(X5)))&v13_algstr_0(X5))&v33_algstr_0(X5))&v3_group_1(X5))&v5_group_1(X5))&v2_rlvect_1(X5))&v3_rlvect_1(X5))&v4_rlvect_1(X5))&v4_vectsp_1(X5))&v5_vectsp_1(X5))&l6_algstr_0(X5))=>![X6]:(m1_matrix_1(X6,u1_struct_0(X5),X2,X1)=>![X7]:(m1_matrix_1(X7,u1_struct_0(X5),X2,X1)=>![X8]:(v7_ordinal1(X8)=>((((r2_hidden(X3,k4_finseq_1(X6))&r2_hidden(X4,k4_finseq_1(X6)))&r2_hidden(X8,k4_finseq_1(X6)))&X7=k1_matrix12(X2,X1,X5,X6,X3,X4))=>(((X8=X3=>k8_matrix_1(u1_struct_0(X5),X7,X8)=k8_matrix_1(u1_struct_0(X5),X6,X4))&(X8=X4=>k8_matrix_1(u1_struct_0(X5),X7,X8)=k8_matrix_1(u1_struct_0(X5),X6,X3)))&~(((X8!=X3&X8!=X4)&k8_matrix_1(u1_struct_0(X5),X7,X8)!=k8_matrix_1(u1_struct_0(X5),X6,X8))))))))))))), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', t2_matrix12)).
fof(dt_m1_matrix_1, axiom, ![X1, X2, X3]:(((~(v1_xboole_0(X1))&v7_ordinal1(X2))&v7_ordinal1(X3))=>![X4]:(m1_matrix_1(X4,X1,X2,X3)=>(v1_matrix_1(X4)&m2_finseq_1(X4,k3_finseq_2(X1))))), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', dt_m1_matrix_1)).
fof(d7_matrix_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_matrix_1(X2)&m2_finseq_1(X2,k3_finseq_2(X1)))=>![X3]:(v7_ordinal1(X3)=>![X4]:(m2_finseq_1(X4,X1)=>(X4=k6_matrix_1(X1,X2,X3)<=>(k3_finseq_1(X4)=k1_matrix_1(X2)&![X5]:(v7_ordinal1(X5)=>(r2_hidden(X5,k2_finseq_1(k1_matrix_1(X2)))=>k1_funct_1(X4,X5)=k3_matrix_1(X1,X2,X3,X5))))))))), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', d7_matrix_1)).
fof(dt_l2_struct_0, axiom, ![X1]:(l2_struct_0(X1)=>l1_struct_0(X1)), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', dt_l2_struct_0)).
fof(dt_l2_algstr_0, axiom, ![X1]:(l2_algstr_0(X1)=>(l2_struct_0(X1)&l1_algstr_0(X1))), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', dt_l2_algstr_0)).
fof(dt_l6_algstr_0, axiom, ![X1]:(l6_algstr_0(X1)=>(l2_algstr_0(X1)&l5_algstr_0(X1))), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', dt_l6_algstr_0)).
fof(cc8_ordinal1, axiom, ![X1]:(m1_subset_1(X1,k4_ordinal1)=>v7_ordinal1(X1)), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', cc8_ordinal1)).
fof(dt_k6_matrix_1, axiom, ![X1, X2, X3]:((((~(v1_xboole_0(X1))&v1_matrix_1(X2))&m1_finseq_1(X2,k3_finseq_2(X1)))&v7_ordinal1(X3))=>m2_finseq_1(k6_matrix_1(X1,X2,X3),X1)), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', dt_k6_matrix_1)).
fof(redefinition_m2_finseq_1, axiom, ![X1, X2]:(m2_finseq_1(X2,X1)<=>m1_finseq_1(X2,X1)), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', redefinition_m2_finseq_1)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', fc2_struct_0)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', redefinition_k5_numbers)).
fof(dt_k1_matrix_1, axiom, ![X1]:((((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))&v1_matrix_1(X1))=>m1_subset_1(k1_matrix_1(X1),k5_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', dt_k1_matrix_1)).
fof(dt_m1_finseq_1, axiom, ![X1, X2]:(m1_finseq_1(X2,X1)=>((v1_relat_1(X2)&v1_funct_1(X2))&v1_finseq_1(X2))), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', dt_m1_finseq_1)).
fof(t1_matrix12, axiom, ![X1]:(v7_ordinal1(X1)=>![X2]:(v7_ordinal1(X2)=>![X3]:((((((((((((~(v2_struct_0(X3))&~(v6_struct_0(X3)))&v13_algstr_0(X3))&v33_algstr_0(X3))&v3_group_1(X3))&v5_group_1(X3))&v2_rlvect_1(X3))&v3_rlvect_1(X3))&v4_rlvect_1(X3))&v4_vectsp_1(X3))&v5_vectsp_1(X3))&l6_algstr_0(X3))=>![X4]:(m1_matrix_1(X4,u1_struct_0(X3),X1,X2)=>![X5]:(m1_matrix_1(X5,u1_struct_0(X3),X1,X2)=>k1_matrix_1(X4)=k1_matrix_1(X5)))))), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', t1_matrix12)).
fof(dt_m2_finseq_1, axiom, ![X1, X2]:(m2_finseq_1(X2,X1)=>((v1_funct_1(X2)&v1_finseq_1(X2))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', dt_m2_finseq_1)).
fof(redefinition_k8_matrix_1, axiom, ![X1, X2, X3]:((((~(v1_xboole_0(X1))&v1_matrix_1(X2))&m1_finseq_1(X2,k3_finseq_2(X1)))&v7_ordinal1(X3))=>k8_matrix_1(X1,X2,X3)=k6_matrix_1(X1,X2,X3)), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', redefinition_k8_matrix_1)).
fof(fraenkel_a_1_0_finseq_1, axiom, ![X1, X2]:(v7_ordinal1(X2)=>(r2_hidden(X1,a_1_0_finseq_1(X2))<=>?[X3]:(((m2_subset_1(X3,k1_numbers,k5_numbers)&X1=X3)&r1_xxreal_0(np__1,X3))&r1_xxreal_0(X3,X2)))), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', fraenkel_a_1_0_finseq_1)).
fof(redefinition_k2_finseq_1, axiom, ![X1]:(v7_ordinal1(X1)=>k2_finseq_1(X1)=k1_finseq_1(X1)), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', redefinition_k2_finseq_1)).
fof(d1_finseq_1, axiom, ![X1]:(v7_ordinal1(X1)=>k1_finseq_1(X1)=a_1_0_finseq_1(X1)), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', d1_finseq_1)).
fof(t14_finseq_1, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))=>![X2]:(((v1_relat_1(X2)&v1_funct_1(X2))&v1_finseq_1(X2))=>((k3_finseq_1(X1)=k3_finseq_1(X2)&![X3]:(v7_ordinal1(X3)=>((r1_xxreal_0(np__1,X3)&r1_xxreal_0(X3,k3_finseq_1(X1)))=>k1_funct_1(X1,X3)=k1_funct_1(X2,X3))))=>X1=X2))), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', t14_finseq_1)).
fof(t5_subset, axiom, ![X1, X2, X3]:~(((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))&v1_xboole_0(X3))), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', t5_subset)).
fof(fc6_ordinal1, axiom, (~(v1_xboole_0(k4_ordinal1))&v3_ordinal1(k4_ordinal1)), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', fc6_ordinal1)).
fof(dt_k5_numbers, axiom, m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', dt_k5_numbers)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', t2_subset)).
fof(redefinition_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', redefinition_m2_subset_1)).
fof(existence_m1_subset_1, axiom, ![X1]:?[X2]:m1_subset_1(X2,X1), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', existence_m1_subset_1)).
fof(d12_ordinal1, axiom, ![X1]:(v7_ordinal1(X1)<=>r2_hidden(X1,k4_ordinal1)), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', d12_ordinal1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', t1_subset)).
fof(dt_l5_algstr_0, axiom, ![X1]:(l5_algstr_0(X1)=>(l4_algstr_0(X1)&l4_struct_0(X1))), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', dt_l5_algstr_0)).
fof(cc3_struct_0, axiom, ![X1]:(l4_struct_0(X1)=>(~(v6_struct_0(X1))=>~(v7_struct_0(X1)))), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', cc3_struct_0)).
fof(cc2_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>(~(v7_struct_0(X1))=>~(v2_struct_0(X1)))), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', cc2_struct_0)).
fof(dt_m2_finseq_2, axiom, ![X1, X2]:(m1_finseq_2(X2,X1)=>![X3]:(m2_finseq_2(X3,X1,X2)=>m2_finseq_1(X3,X1))), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', dt_m2_finseq_2)).
fof(dt_k8_matrix_1, axiom, ![X1, X2, X3]:((((~(v1_xboole_0(X1))&v1_matrix_1(X2))&m1_finseq_1(X2,k3_finseq_2(X1)))&v7_ordinal1(X3))=>m2_finseq_2(k8_matrix_1(X1,X2,X3),X1,k4_finseq_2(k1_matrix_1(X2),X1))), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', dt_k8_matrix_1)).
fof(dt_k1_matrix12, axiom, ![X1, X2, X3, X4, X5, X6]:(((((((((((((((((v7_ordinal1(X1)&v7_ordinal1(X2))&~(v2_struct_0(X3)))&~(v6_struct_0(X3)))&v13_algstr_0(X3))&v33_algstr_0(X3))&v3_group_1(X3))&v5_group_1(X3))&v2_rlvect_1(X3))&v3_rlvect_1(X3))&v4_rlvect_1(X3))&v4_vectsp_1(X3))&v5_vectsp_1(X3))&l6_algstr_0(X3))&m1_matrix_1(X4,u1_struct_0(X3),X1,X2))&v7_ordinal1(X5))&v7_ordinal1(X6))=>m1_matrix_1(k1_matrix12(X1,X2,X3,X4,X5,X6),u1_struct_0(X3),X1,X2)), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', dt_k1_matrix12)).
fof(d1_matrix12, axiom, ![X1]:(v7_ordinal1(X1)=>![X2]:(v7_ordinal1(X2)=>![X3]:((((((((((((~(v2_struct_0(X3))&~(v6_struct_0(X3)))&v13_algstr_0(X3))&v33_algstr_0(X3))&v3_group_1(X3))&v5_group_1(X3))&v2_rlvect_1(X3))&v3_rlvect_1(X3))&v4_rlvect_1(X3))&v4_vectsp_1(X3))&v5_vectsp_1(X3))&l6_algstr_0(X3))=>![X4]:(m1_matrix_1(X4,u1_struct_0(X3),X1,X2)=>![X5]:(v7_ordinal1(X5)=>![X6]:(v7_ordinal1(X6)=>![X7]:(m1_matrix_1(X7,u1_struct_0(X3),X1,X2)=>(X7=k1_matrix12(X1,X2,X3,X4,X5,X6)<=>(k3_finseq_1(X7)=k3_finseq_1(X4)&![X8]:(v7_ordinal1(X8)=>![X9]:(v7_ordinal1(X9)=>((r2_hidden(X8,k4_finseq_1(X4))&r2_hidden(X9,k2_finseq_1(k1_matrix_1(X4))))=>(((X8=X5=>k3_matrix_1(u1_struct_0(X3),X7,X8,X9)=k3_matrix_1(u1_struct_0(X3),X4,X6,X9))&(X8=X6=>k3_matrix_1(u1_struct_0(X3),X7,X8,X9)=k3_matrix_1(u1_struct_0(X3),X4,X5,X9)))&~(((X8!=X5&X8!=X6)&k3_matrix_1(u1_struct_0(X3),X7,X8,X9)!=k3_matrix_1(u1_struct_0(X3),X4,X8,X9)))))))))))))))), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', d1_matrix12)).
fof(dt_k4_finseq_2, axiom, ![X1, X2]:(v7_ordinal1(X1)=>m1_finseq_2(k4_finseq_2(X1,X2),X2)), file('/scratch/mptp/hardpred/preds__0.005//t2_matrix12', dt_k4_finseq_2)).
fof(c_0_36, negated_conjecture, ~(![X1]:(v7_ordinal1(X1)=>![X2]:(v7_ordinal1(X2)=>![X3]:(v7_ordinal1(X3)=>![X4]:(v7_ordinal1(X4)=>![X5]:((((((((((((~(v2_struct_0(X5))&~(v6_struct_0(X5)))&v13_algstr_0(X5))&v33_algstr_0(X5))&v3_group_1(X5))&v5_group_1(X5))&v2_rlvect_1(X5))&v3_rlvect_1(X5))&v4_rlvect_1(X5))&v4_vectsp_1(X5))&v5_vectsp_1(X5))&l6_algstr_0(X5))=>![X6]:(m1_matrix_1(X6,u1_struct_0(X5),X2,X1)=>![X7]:(m1_matrix_1(X7,u1_struct_0(X5),X2,X1)=>![X8]:(v7_ordinal1(X8)=>((((r2_hidden(X3,k4_finseq_1(X6))&r2_hidden(X4,k4_finseq_1(X6)))&r2_hidden(X8,k4_finseq_1(X6)))&X7=k1_matrix12(X2,X1,X5,X6,X3,X4))=>(((X8=X3=>k8_matrix_1(u1_struct_0(X5),X7,X8)=k8_matrix_1(u1_struct_0(X5),X6,X4))&(X8=X4=>k8_matrix_1(u1_struct_0(X5),X7,X8)=k8_matrix_1(u1_struct_0(X5),X6,X3)))&~(((X8!=X3&X8!=X4)&k8_matrix_1(u1_struct_0(X5),X7,X8)!=k8_matrix_1(u1_struct_0(X5),X6,X8)))))))))))))), inference(assume_negation,[status(cth)],[t2_matrix12])).
fof(c_0_37, plain, ![X28, X29, X30, X31]:((v1_matrix_1(X31)|~m1_matrix_1(X31,X28,X29,X30)|(v1_xboole_0(X28)|~v7_ordinal1(X29)|~v7_ordinal1(X30)))&(m2_finseq_1(X31,k3_finseq_2(X28))|~m1_matrix_1(X31,X28,X29,X30)|(v1_xboole_0(X28)|~v7_ordinal1(X29)|~v7_ordinal1(X30)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m1_matrix_1])])])])])).
fof(c_0_38, negated_conjecture, (v7_ordinal1(esk1_0)&(v7_ordinal1(esk2_0)&(v7_ordinal1(esk3_0)&(v7_ordinal1(esk4_0)&((((((((((((~v2_struct_0(esk5_0)&~v6_struct_0(esk5_0))&v13_algstr_0(esk5_0))&v33_algstr_0(esk5_0))&v3_group_1(esk5_0))&v5_group_1(esk5_0))&v2_rlvect_1(esk5_0))&v3_rlvect_1(esk5_0))&v4_rlvect_1(esk5_0))&v4_vectsp_1(esk5_0))&v5_vectsp_1(esk5_0))&l6_algstr_0(esk5_0))&(m1_matrix_1(esk6_0,u1_struct_0(esk5_0),esk2_0,esk1_0)&(m1_matrix_1(esk7_0,u1_struct_0(esk5_0),esk2_0,esk1_0)&(v7_ordinal1(esk8_0)&((((r2_hidden(esk3_0,k4_finseq_1(esk6_0))&r2_hidden(esk4_0,k4_finseq_1(esk6_0)))&r2_hidden(esk8_0,k4_finseq_1(esk6_0)))&esk7_0=k1_matrix12(esk2_0,esk1_0,esk5_0,esk6_0,esk3_0,esk4_0))&(((((esk8_0!=esk3_0|(esk8_0=esk4_0|esk8_0=esk3_0))&(esk8_0!=esk4_0|(esk8_0=esk4_0|esk8_0=esk3_0)))&(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|(esk8_0=esk4_0|esk8_0=esk3_0)))&(((esk8_0!=esk3_0|(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk8_0=esk3_0))&(esk8_0!=esk4_0|(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk8_0=esk3_0)))&(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk8_0=esk3_0))))&((((esk8_0!=esk3_0|(esk8_0=esk4_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)))&(esk8_0!=esk4_0|(esk8_0=esk4_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))))&(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|(esk8_0=esk4_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))))&(((esk8_0!=esk3_0|(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)))&(esk8_0!=esk4_0|(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))))&(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))))))))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_36])])])])])).
fof(c_0_39, plain, ![X20, X21, X22, X23, X24]:(((k3_finseq_1(X23)=k1_matrix_1(X21)|X23!=k6_matrix_1(X20,X21,X22)|~m2_finseq_1(X23,X20)|~v7_ordinal1(X22)|(~v1_matrix_1(X21)|~m2_finseq_1(X21,k3_finseq_2(X20)))|v1_xboole_0(X20))&(~v7_ordinal1(X24)|(~r2_hidden(X24,k2_finseq_1(k1_matrix_1(X21)))|k1_funct_1(X23,X24)=k3_matrix_1(X20,X21,X22,X24))|X23!=k6_matrix_1(X20,X21,X22)|~m2_finseq_1(X23,X20)|~v7_ordinal1(X22)|(~v1_matrix_1(X21)|~m2_finseq_1(X21,k3_finseq_2(X20)))|v1_xboole_0(X20)))&((v7_ordinal1(esk9_4(X20,X21,X22,X23))|k3_finseq_1(X23)!=k1_matrix_1(X21)|X23=k6_matrix_1(X20,X21,X22)|~m2_finseq_1(X23,X20)|~v7_ordinal1(X22)|(~v1_matrix_1(X21)|~m2_finseq_1(X21,k3_finseq_2(X20)))|v1_xboole_0(X20))&((r2_hidden(esk9_4(X20,X21,X22,X23),k2_finseq_1(k1_matrix_1(X21)))|k3_finseq_1(X23)!=k1_matrix_1(X21)|X23=k6_matrix_1(X20,X21,X22)|~m2_finseq_1(X23,X20)|~v7_ordinal1(X22)|(~v1_matrix_1(X21)|~m2_finseq_1(X21,k3_finseq_2(X20)))|v1_xboole_0(X20))&(k1_funct_1(X23,esk9_4(X20,X21,X22,X23))!=k3_matrix_1(X20,X21,X22,esk9_4(X20,X21,X22,X23))|k3_finseq_1(X23)!=k1_matrix_1(X21)|X23=k6_matrix_1(X20,X21,X22)|~m2_finseq_1(X23,X20)|~v7_ordinal1(X22)|(~v1_matrix_1(X21)|~m2_finseq_1(X21,k3_finseq_2(X20)))|v1_xboole_0(X20))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_matrix_1])])])])])])).
cnf(c_0_40, plain, (m2_finseq_1(X1,k3_finseq_2(X2))|v1_xboole_0(X2)|~m1_matrix_1(X1,X2,X3,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X4)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_41, negated_conjecture, (m1_matrix_1(esk6_0,u1_struct_0(esk5_0),esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_42, negated_conjecture, (v7_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_43, negated_conjecture, (v7_ordinal1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
fof(c_0_44, plain, ![X57]:(~l2_struct_0(X57)|l1_struct_0(X57)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l2_struct_0])])).
fof(c_0_45, plain, ![X42]:((l2_struct_0(X42)|~l2_algstr_0(X42))&(l1_algstr_0(X42)|~l2_algstr_0(X42))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l2_algstr_0])])])).
cnf(c_0_46, plain, (k3_finseq_1(X1)=k1_matrix_1(X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,X2,X4)|~m2_finseq_1(X1,X3)|~v7_ordinal1(X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X3))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_47, plain, (v1_xboole_0(X2)|m2_finseq_1(X1,k3_finseq_2(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4)), inference(evalgc,[status(thm)],[c_0_40])).
cnf(c_0_48, negated_conjecture, (m1_matrix_1(esk6_0,u1_struct_0(esk5_0),esk2_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_41])).
cnf(c_0_49, negated_conjecture, (v7_ordinal1(esk1_0)), inference(evalgc,[status(thm)],[c_0_42])).
cnf(c_0_50, negated_conjecture, (v7_ordinal1(esk2_0)), inference(evalgc,[status(thm)],[c_0_43])).
cnf(c_0_51, plain, (v1_matrix_1(X1)|v1_xboole_0(X2)|~m1_matrix_1(X1,X2,X3,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X4)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_52, plain, (l1_struct_0(X1)|~l2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_53, plain, (l2_struct_0(X1)|~l2_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_54, plain, ![X26]:((l2_algstr_0(X26)|~l6_algstr_0(X26))&(l5_algstr_0(X26)|~l6_algstr_0(X26))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l6_algstr_0])])])).
fof(c_0_55, plain, ![X80]:(~m1_subset_1(X80,k4_ordinal1)|v7_ordinal1(X80)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_ordinal1])])).
cnf(c_0_56, plain, (k1_matrix_1(X2)=k3_finseq_1(X1)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,X2,X4)|~v7_ordinal1(X4)|~v1_matrix_1(X2)|~m2_finseq_1(X1,X3)|~m2_finseq_1(X2,k3_finseq_2(X3))), inference(evalgc,[status(thm)],[c_0_46])).
cnf(c_0_57, negated_conjecture, (m2_finseq_1(esk6_0,k3_finseq_2(u1_struct_0(esk5_0)))|v1_xboole_0(u1_struct_0(esk5_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49]), c_0_50])])])).
cnf(c_0_58, plain, (v1_xboole_0(X2)|v1_matrix_1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4)), inference(evalgc,[status(thm)],[c_0_51])).
fof(c_0_59, plain, ![X76, X77, X78]:(v1_xboole_0(X76)|~v1_matrix_1(X77)|~m1_finseq_1(X77,k3_finseq_2(X76))|~v7_ordinal1(X78)|m2_finseq_1(k6_matrix_1(X76,X77,X78),X76)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_matrix_1])])])).
fof(c_0_60, plain, ![X55, X56]:((~m2_finseq_1(X56,X55)|m1_finseq_1(X56,X55))&(~m1_finseq_1(X56,X55)|m2_finseq_1(X56,X55))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_m2_finseq_1])])).
fof(c_0_61, plain, ![X71]:(v2_struct_0(X71)|~l1_struct_0(X71)|~v1_xboole_0(u1_struct_0(X71))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_62, plain, (l1_struct_0(X1)|~l2_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_52])).
cnf(c_0_63, plain, (l2_struct_0(X1)|~l2_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_53])).
cnf(c_0_64, plain, (l2_algstr_0(X1)|~l6_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_65, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_66, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
fof(c_0_67, plain, ![X79]:(~v1_relat_1(X79)|~v1_funct_1(X79)|~v1_finseq_1(X79)|~v1_matrix_1(X79)|m1_subset_1(k1_matrix_1(X79),k5_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_matrix_1])])).
fof(c_0_68, plain, ![X63, X64]:(((v1_relat_1(X64)|~m1_finseq_1(X64,X63))&(v1_funct_1(X64)|~m1_finseq_1(X64,X63)))&(v1_finseq_1(X64)|~m1_finseq_1(X64,X63))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_finseq_1])])])).
fof(c_0_69, plain, ![X36, X37, X38, X39, X40]:(~v7_ordinal1(X36)|(~v7_ordinal1(X37)|(v2_struct_0(X38)|v6_struct_0(X38)|~v13_algstr_0(X38)|~v33_algstr_0(X38)|~v3_group_1(X38)|~v5_group_1(X38)|~v2_rlvect_1(X38)|~v3_rlvect_1(X38)|~v4_rlvect_1(X38)|~v4_vectsp_1(X38)|~v5_vectsp_1(X38)|~l6_algstr_0(X38)|(~m1_matrix_1(X39,u1_struct_0(X38),X36,X37)|(~m1_matrix_1(X40,u1_struct_0(X38),X36,X37)|k1_matrix_1(X39)=k1_matrix_1(X40)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t1_matrix12])])])])).
cnf(c_0_70, negated_conjecture, (k1_matrix_1(esk6_0)=k3_finseq_1(X1)|v1_xboole_0(u1_struct_0(esk5_0))|X1!=k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v1_matrix_1(esk6_0)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57])])).
cnf(c_0_71, negated_conjecture, (v1_matrix_1(esk6_0)|v1_xboole_0(u1_struct_0(esk5_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_48]), c_0_49]), c_0_50])])])).
cnf(c_0_72, plain, (v1_xboole_0(X1)|m2_finseq_1(k6_matrix_1(X1,X2,X3),X1)|~v1_matrix_1(X2)|~m1_finseq_1(X2,k3_finseq_2(X1))|~v7_ordinal1(X3)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_73, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_74, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_75, plain, (l1_struct_0(X1)|~l2_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63])])).
cnf(c_0_76, plain, (l2_algstr_0(X1)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_64])).
cnf(c_0_77, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_65])).
cnf(c_0_78, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_66])).
cnf(c_0_79, plain, (m1_subset_1(k1_matrix_1(X1),k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)|~v1_matrix_1(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_80, negated_conjecture, (m1_matrix_1(esk7_0,u1_struct_0(esk5_0),esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_81, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
fof(c_0_82, plain, ![X61, X62]:(((v1_funct_1(X62)|~m2_finseq_1(X62,X61))&(v1_finseq_1(X62)|~m2_finseq_1(X62,X61)))&(m1_subset_1(X62,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X61)))|~m2_finseq_1(X62,X61))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m2_finseq_1])])])).
cnf(c_0_83, plain, (v2_struct_0(X3)|v6_struct_0(X3)|k1_matrix_1(X4)=k1_matrix_1(X5)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X5,u1_struct_0(X3),X1,X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_84, negated_conjecture, (l6_algstr_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_85, negated_conjecture, (v5_vectsp_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_86, negated_conjecture, (v4_vectsp_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_87, negated_conjecture, (v4_rlvect_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_88, negated_conjecture, (v3_rlvect_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_89, negated_conjecture, (v2_rlvect_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_90, negated_conjecture, (v5_group_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_91, negated_conjecture, (v3_group_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_92, negated_conjecture, (v33_algstr_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_93, negated_conjecture, (v13_algstr_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_94, negated_conjecture, (~v6_struct_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_95, negated_conjecture, (~v2_struct_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_96, negated_conjecture, (k1_matrix_1(esk6_0)=k3_finseq_1(X1)|v1_xboole_0(u1_struct_0(esk5_0))|X1!=k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(X2)), inference(csr,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_97, plain, (v1_xboole_0(X1)|m2_finseq_1(k6_matrix_1(X1,X2,X3),X1)|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m1_finseq_1(X2,k3_finseq_2(X1))), inference(evalgc,[status(thm)],[c_0_72])).
cnf(c_0_98, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_73])).
cnf(c_0_99, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_74])).
cnf(c_0_100, plain, (l1_struct_0(X1)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76])])).
cnf(c_0_101, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)), inference(rw,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_102, plain, (m1_subset_1(k1_matrix_1(X1),k5_numbers)|~v1_matrix_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_79])).
cnf(c_0_103, negated_conjecture, (m1_matrix_1(esk7_0,u1_struct_0(esk5_0),esk2_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_80])).
cnf(c_0_104, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_81])).
cnf(c_0_105, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_106, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_107, plain, (k1_matrix_1(X4)=k1_matrix_1(X5)|v2_struct_0(X3)|v6_struct_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~m1_matrix_1(X5,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)), inference(evalgc,[status(thm)],[c_0_83])).
cnf(c_0_108, negated_conjecture, (l6_algstr_0(esk5_0)), inference(evalgc,[status(thm)],[c_0_84])).
cnf(c_0_109, negated_conjecture, (v5_vectsp_1(esk5_0)), inference(evalgc,[status(thm)],[c_0_85])).
cnf(c_0_110, negated_conjecture, (v4_vectsp_1(esk5_0)), inference(evalgc,[status(thm)],[c_0_86])).
cnf(c_0_111, negated_conjecture, (v4_rlvect_1(esk5_0)), inference(evalgc,[status(thm)],[c_0_87])).
cnf(c_0_112, negated_conjecture, (v3_rlvect_1(esk5_0)), inference(evalgc,[status(thm)],[c_0_88])).
cnf(c_0_113, negated_conjecture, (v2_rlvect_1(esk5_0)), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_114, negated_conjecture, (v5_group_1(esk5_0)), inference(evalgc,[status(thm)],[c_0_90])).
cnf(c_0_115, negated_conjecture, (v3_group_1(esk5_0)), inference(evalgc,[status(thm)],[c_0_91])).
cnf(c_0_116, negated_conjecture, (v33_algstr_0(esk5_0)), inference(evalgc,[status(thm)],[c_0_92])).
cnf(c_0_117, negated_conjecture, (v13_algstr_0(esk5_0)), inference(evalgc,[status(thm)],[c_0_93])).
cnf(c_0_118, negated_conjecture, (~v6_struct_0(esk5_0)), inference(evalgc,[status(thm)],[c_0_94])).
cnf(c_0_119, negated_conjecture, (~v2_struct_0(esk5_0)), inference(evalgc,[status(thm)],[c_0_95])).
cnf(c_0_120, negated_conjecture, (k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))=k1_matrix_1(esk6_0)|v1_xboole_0(u1_struct_0(esk5_0))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_96])])).
cnf(c_0_121, plain, (m2_finseq_1(k6_matrix_1(X1,X2,X3),X1)|v1_xboole_0(X1)|~v1_matrix_1(X2)|~v7_ordinal1(X3)|~m2_finseq_1(X2,k3_finseq_2(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98])])).
cnf(c_0_122, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100])])).
fof(c_0_123, plain, ![X91, X92, X93]:(v1_xboole_0(X91)|~v1_matrix_1(X92)|~m1_finseq_1(X92,k3_finseq_2(X91))|~v7_ordinal1(X93)|k8_matrix_1(X91,X92,X93)=k6_matrix_1(X91,X92,X93)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k8_matrix_1])])])).
fof(c_0_124, plain, ![X66, X67, X69]:(((((m2_subset_1(esk11_2(X66,X67),k1_numbers,k5_numbers)|~r2_hidden(X66,a_1_0_finseq_1(X67))|~v7_ordinal1(X67))&(X66=esk11_2(X66,X67)|~r2_hidden(X66,a_1_0_finseq_1(X67))|~v7_ordinal1(X67)))&(r1_xxreal_0(np__1,esk11_2(X66,X67))|~r2_hidden(X66,a_1_0_finseq_1(X67))|~v7_ordinal1(X67)))&(r1_xxreal_0(esk11_2(X66,X67),X67)|~r2_hidden(X66,a_1_0_finseq_1(X67))|~v7_ordinal1(X67)))&(~m2_subset_1(X69,k1_numbers,k5_numbers)|X66!=X69|~r1_xxreal_0(np__1,X69)|~r1_xxreal_0(X69,X67)|r2_hidden(X66,a_1_0_finseq_1(X67))|~v7_ordinal1(X67))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fraenkel_a_1_0_finseq_1])])])])])).
fof(c_0_125, plain, ![X95]:(~v7_ordinal1(X95)|k2_finseq_1(X95)=k1_finseq_1(X95)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_finseq_1])])).
fof(c_0_126, plain, ![X120]:(~v7_ordinal1(X120)|k1_finseq_1(X120)=a_1_0_finseq_1(X120)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_finseq_1])])).
cnf(c_0_127, plain, (v7_ordinal1(k1_matrix_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102])])).
cnf(c_0_128, negated_conjecture, (v1_matrix_1(esk7_0)|v1_xboole_0(u1_struct_0(esk5_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_103]), c_0_49]), c_0_50])])])).
cnf(c_0_129, plain, (v1_relat_1(X1)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_98])])).
cnf(c_0_130, negated_conjecture, (m2_finseq_1(esk7_0,k3_finseq_2(u1_struct_0(esk5_0)))|v1_xboole_0(u1_struct_0(esk5_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_103]), c_0_49]), c_0_50])])])).
cnf(c_0_131, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_132, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_106])).
cnf(c_0_133, negated_conjecture, (k1_matrix_1(X1)=k1_matrix_1(esk6_0)|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_48]), c_0_108]), c_0_109]), c_0_110]), c_0_111]), c_0_112]), c_0_113]), c_0_114]), c_0_115]), c_0_116]), c_0_117]), c_0_49]), c_0_50])]), c_0_118]), c_0_119])])).
cnf(c_0_134, negated_conjecture, (k1_matrix_1(X1)=k1_matrix_1(esk7_0)|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_103]), c_0_108]), c_0_109]), c_0_110]), c_0_111]), c_0_112]), c_0_113]), c_0_114]), c_0_115]), c_0_116]), c_0_117]), c_0_49]), c_0_50])]), c_0_118]), c_0_119])])).
cnf(c_0_135, negated_conjecture, (k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))=k1_matrix_1(esk6_0)|v1_xboole_0(u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~m2_finseq_1(esk6_0,k3_finseq_2(u1_struct_0(esk5_0)))|~v1_matrix_1(esk6_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121])])).
cnf(c_0_136, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk5_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_108]), c_0_119])])).
cnf(c_0_137, plain, (v1_xboole_0(X1)|k8_matrix_1(X1,X2,X3)=k6_matrix_1(X1,X2,X3)|~v1_matrix_1(X2)|~m1_finseq_1(X2,k3_finseq_2(X1))|~v7_ordinal1(X3)), inference(split_conjunct,[status(thm)],[c_0_123])).
cnf(c_0_138, plain, (r2_hidden(X2,a_1_0_finseq_1(X3))|~m2_subset_1(X1,k1_numbers,k5_numbers)|X2!=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X3)), inference(split_conjunct,[status(thm)],[c_0_124])).
cnf(c_0_139, plain, (k2_finseq_1(X1)=k1_finseq_1(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_125])).
cnf(c_0_140, plain, (k1_finseq_1(X1)=a_1_0_finseq_1(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_126])).
cnf(c_0_141, negated_conjecture, (v7_ordinal1(k1_matrix_1(esk7_0))|v1_xboole_0(u1_struct_0(esk5_0))|~v1_finseq_1(esk7_0)|~v1_funct_1(esk7_0)|~v1_relat_1(esk7_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128])])).
cnf(c_0_142, negated_conjecture, (v1_relat_1(esk7_0)|v1_xboole_0(u1_struct_0(esk5_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130])])).
cnf(c_0_143, negated_conjecture, (v1_funct_1(esk7_0)|v1_xboole_0(u1_struct_0(esk5_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_130])])).
cnf(c_0_144, negated_conjecture, (v1_finseq_1(esk7_0)|v1_xboole_0(u1_struct_0(esk5_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_130])])).
cnf(c_0_145, negated_conjecture, (k1_matrix_1(esk7_0)=k1_matrix_1(esk6_0)|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_134])])).
cnf(c_0_146, negated_conjecture, (k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))=k1_matrix_1(esk6_0)|~m2_finseq_1(esk6_0,k3_finseq_2(u1_struct_0(esk5_0)))|~v1_matrix_1(esk6_0)|~v7_ordinal1(X1)), inference(sr,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_147, negated_conjecture, (v1_matrix_1(esk6_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[c_0_71, c_0_136])])).
cnf(c_0_148, plain, (k6_matrix_1(X1,X2,X3)=k8_matrix_1(X1,X2,X3)|v1_xboole_0(X1)|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m1_finseq_1(X2,k3_finseq_2(X1))), inference(evalgc,[status(thm)],[c_0_137])).
cnf(c_0_149, plain, (k1_funct_1(X3,X1)=k3_matrix_1(X4,X2,X5,X1)|v1_xboole_0(X4)|~v7_ordinal1(X1)|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2)))|X3!=k6_matrix_1(X4,X2,X5)|~m2_finseq_1(X3,X4)|~v7_ordinal1(X5)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X4))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_150, plain, (r2_hidden(X2,a_1_0_finseq_1(X3))|X1!=X2|~v7_ordinal1(X3)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_138])).
cnf(c_0_151, plain, (k1_finseq_1(X1)=k2_finseq_1(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_139])).
cnf(c_0_152, plain, (k1_finseq_1(X1)=a_1_0_finseq_1(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_140])).
cnf(c_0_153, negated_conjecture, (v1_xboole_0(u1_struct_0(esk5_0))|v7_ordinal1(k1_matrix_1(esk7_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_141, c_0_142]), c_0_143]), c_0_144])).
cnf(c_0_154, negated_conjecture, (k1_matrix_1(esk7_0)=k1_matrix_1(esk6_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_103])])).
cnf(c_0_155, negated_conjecture, (k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))=k1_matrix_1(esk6_0)|~m2_finseq_1(esk6_0,k3_finseq_2(u1_struct_0(esk5_0)))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_146, c_0_147])])])).
cnf(c_0_156, negated_conjecture, (m2_finseq_1(esk6_0,k3_finseq_2(u1_struct_0(esk5_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[c_0_57, c_0_136])])).
cnf(c_0_157, plain, (k6_matrix_1(X1,X2,X3)=k8_matrix_1(X1,X2,X3)|v1_xboole_0(X1)|~v1_matrix_1(X2)|~v7_ordinal1(X3)|~m2_finseq_1(X2,k3_finseq_2(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_98])])).
cnf(c_0_158, plain, (k1_funct_1(X3,X1)=k3_matrix_1(X4,X2,X5,X1)|v1_xboole_0(X4)|X3!=k6_matrix_1(X4,X2,X5)|~v7_ordinal1(X5)|~v7_ordinal1(X1)|~v1_matrix_1(X2)|~m2_finseq_1(X3,X4)|~m2_finseq_1(X2,k3_finseq_2(X4))|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2)))), inference(evalgc,[status(thm)],[c_0_149])).
cnf(c_0_159, plain, (r2_hidden(X1,a_1_0_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_150])])).
cnf(c_0_160, plain, (a_1_0_finseq_1(X1)=k2_finseq_1(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152])])).
cnf(c_0_161, negated_conjecture, (v1_xboole_0(u1_struct_0(esk5_0))|v7_ordinal1(k1_matrix_1(esk6_0))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_153, c_0_154])])).
cnf(c_0_162, negated_conjecture, (k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))=k1_matrix_1(esk6_0)|~v7_ordinal1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_155, c_0_156])])).
cnf(c_0_163, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_156]), c_0_147])]), c_0_136])])).
cnf(c_0_164, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,X4,X5,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,X4,X5)|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v1_matrix_1(X4)|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(esk6_0)))|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_133])])).
cnf(c_0_165, plain, (r2_hidden(X1,k2_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_160])])).
cnf(c_0_166, negated_conjecture, (v7_ordinal1(k1_matrix_1(esk6_0))), inference(sr,[status(thm)],[c_0_161, c_0_136])).
cnf(c_0_167, negated_conjecture, (k1_matrix_1(X1)=k1_matrix_1(X2)|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_133])])).
cnf(c_0_168, negated_conjecture, (k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))=k1_matrix_1(esk6_0)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_163])])).
fof(c_0_169, plain, ![X50, X51]:((v7_ordinal1(esk10_2(X50,X51))|k3_finseq_1(X50)!=k3_finseq_1(X51)|X50=X51|(~v1_relat_1(X51)|~v1_funct_1(X51)|~v1_finseq_1(X51))|(~v1_relat_1(X50)|~v1_funct_1(X50)|~v1_finseq_1(X50)))&(((r1_xxreal_0(np__1,esk10_2(X50,X51))|k3_finseq_1(X50)!=k3_finseq_1(X51)|X50=X51|(~v1_relat_1(X51)|~v1_funct_1(X51)|~v1_finseq_1(X51))|(~v1_relat_1(X50)|~v1_funct_1(X50)|~v1_finseq_1(X50)))&(r1_xxreal_0(esk10_2(X50,X51),k3_finseq_1(X50))|k3_finseq_1(X50)!=k3_finseq_1(X51)|X50=X51|(~v1_relat_1(X51)|~v1_funct_1(X51)|~v1_finseq_1(X51))|(~v1_relat_1(X50)|~v1_funct_1(X50)|~v1_finseq_1(X50))))&(k1_funct_1(X50,esk10_2(X50,X51))!=k1_funct_1(X51,esk10_2(X50,X51))|k3_finseq_1(X50)!=k3_finseq_1(X51)|X50=X51|(~v1_relat_1(X51)|~v1_funct_1(X51)|~v1_finseq_1(X51))|(~v1_relat_1(X50)|~v1_funct_1(X50)|~v1_finseq_1(X50))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t14_finseq_1])])])])])).
fof(c_0_170, plain, ![X1, X2]:(epred1_2(X2,X1)<=>![X3]:((((((((((((~(v2_struct_0(X3))&~(v6_struct_0(X3)))&v13_algstr_0(X3))&v33_algstr_0(X3))&v3_group_1(X3))&v5_group_1(X3))&v2_rlvect_1(X3))&v3_rlvect_1(X3))&v4_rlvect_1(X3))&v4_vectsp_1(X3))&v5_vectsp_1(X3))&l6_algstr_0(X3))=>![X4]:(m1_matrix_1(X4,u1_struct_0(X3),X1,X2)=>![X5]:(v7_ordinal1(X5)=>![X6]:(v7_ordinal1(X6)=>![X7]:(m1_matrix_1(X7,u1_struct_0(X3),X1,X2)=>(X7=k1_matrix12(X1,X2,X3,X4,X5,X6)<=>(k3_finseq_1(X7)=k3_finseq_1(X4)&![X8]:(v7_ordinal1(X8)=>![X9]:(v7_ordinal1(X9)=>((r2_hidden(X8,k4_finseq_1(X4))&r2_hidden(X9,k2_finseq_1(k1_matrix_1(X4))))=>(((X8=X5=>k3_matrix_1(u1_struct_0(X3),X7,X8,X9)=k3_matrix_1(u1_struct_0(X3),X4,X6,X9))&(X8=X6=>k3_matrix_1(u1_struct_0(X3),X7,X8,X9)=k3_matrix_1(u1_struct_0(X3),X4,X5,X9)))&~(((X8!=X5&X8!=X6)&k3_matrix_1(u1_struct_0(X3),X7,X8,X9)!=k3_matrix_1(u1_struct_0(X3),X4,X8,X9))))))))))))))), introduced(definition)).
cnf(c_0_171, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,X4,X5,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,X4,X5)|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v1_matrix_1(X4)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,k1_matrix_1(esk6_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_165]), c_0_166])])])).
cnf(c_0_172, negated_conjecture, (k1_matrix_1(X1)=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_168]), c_0_48])])])).
cnf(c_0_173, plain, (r1_xxreal_0(esk10_2(X1,X2),k3_finseq_1(X1))|X1=X2|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_169])).
cnf(c_0_174, negated_conjecture, (k1_matrix_1(esk7_0)=k3_finseq_1(X1)|v1_xboole_0(u1_struct_0(esk5_0))|X1!=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v1_matrix_1(esk7_0)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_130])])).
fof(c_0_175, plain, ![X1, X2]:(epred1_2(X2,X1)=>![X3]:((((((((((((~(v2_struct_0(X3))&~(v6_struct_0(X3)))&v13_algstr_0(X3))&v33_algstr_0(X3))&v3_group_1(X3))&v5_group_1(X3))&v2_rlvect_1(X3))&v3_rlvect_1(X3))&v4_rlvect_1(X3))&v4_vectsp_1(X3))&v5_vectsp_1(X3))&l6_algstr_0(X3))=>![X4]:(m1_matrix_1(X4,u1_struct_0(X3),X1,X2)=>![X5]:(v7_ordinal1(X5)=>![X6]:(v7_ordinal1(X6)=>![X7]:(m1_matrix_1(X7,u1_struct_0(X3),X1,X2)=>(X7=k1_matrix12(X1,X2,X3,X4,X5,X6)<=>(k3_finseq_1(X7)=k3_finseq_1(X4)&![X8]:(v7_ordinal1(X8)=>![X9]:(v7_ordinal1(X9)=>((r2_hidden(X8,k4_finseq_1(X4))&r2_hidden(X9,k2_finseq_1(k1_matrix_1(X4))))=>(((X8=X5=>k3_matrix_1(u1_struct_0(X3),X7,X8,X9)=k3_matrix_1(u1_struct_0(X3),X4,X6,X9))&(X8=X6=>k3_matrix_1(u1_struct_0(X3),X7,X8,X9)=k3_matrix_1(u1_struct_0(X3),X4,X5,X9)))&~(((X8!=X5&X8!=X6)&k3_matrix_1(u1_struct_0(X3),X7,X8,X9)!=k3_matrix_1(u1_struct_0(X3),X4,X8,X9))))))))))))))), inference(split_equiv,[status(thm)],[c_0_170])).
cnf(c_0_176, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,X4,X5,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,X4,X5)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6)))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v1_matrix_1(X4)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~v7_ordinal1(X6)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_172]), c_0_48])])])).
cnf(c_0_177, plain, (X1=X2|r1_xxreal_0(esk10_2(X1,X2),k3_finseq_1(X1))|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_173])).
cnf(c_0_178, plain, (v7_ordinal1(esk10_2(X1,X2))|X1=X2|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_169])).
cnf(c_0_179, plain, (v1_relat_1(k6_matrix_1(X1,X2,X3))|v1_xboole_0(X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_121])])).
cnf(c_0_180, plain, (v1_funct_1(k6_matrix_1(X1,X2,X3))|v1_xboole_0(X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_121])])).
cnf(c_0_181, plain, (v1_finseq_1(k6_matrix_1(X1,X2,X3))|v1_xboole_0(X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_121])])).
cnf(c_0_182, negated_conjecture, (k1_matrix_1(esk7_0)=k3_finseq_1(X1)|v1_xboole_0(u1_struct_0(esk5_0))|X1!=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(X2)), inference(csr,[status(thm)],[c_0_174, c_0_128])).
fof(c_0_183, plain, ![X159, X160, X161, X162, X163, X164, X165, X166, X167]:(((k3_finseq_1(X165)=k3_finseq_1(X162)|X165!=k1_matrix12(X159,X160,X161,X162,X163,X164)|~m1_matrix_1(X165,u1_struct_0(X161),X159,X160)|~v7_ordinal1(X164)|~v7_ordinal1(X163)|~m1_matrix_1(X162,u1_struct_0(X161),X159,X160)|(v2_struct_0(X161)|v6_struct_0(X161)|~v13_algstr_0(X161)|~v33_algstr_0(X161)|~v3_group_1(X161)|~v5_group_1(X161)|~v2_rlvect_1(X161)|~v3_rlvect_1(X161)|~v4_rlvect_1(X161)|~v4_vectsp_1(X161)|~v5_vectsp_1(X161)|~l6_algstr_0(X161))|~epred1_2(X160,X159))&(((X166!=X163|k3_matrix_1(u1_struct_0(X161),X165,X166,X167)=k3_matrix_1(u1_struct_0(X161),X162,X164,X167)|(~r2_hidden(X166,k4_finseq_1(X162))|~r2_hidden(X167,k2_finseq_1(k1_matrix_1(X162))))|~v7_ordinal1(X167)|~v7_ordinal1(X166)|X165!=k1_matrix12(X159,X160,X161,X162,X163,X164)|~m1_matrix_1(X165,u1_struct_0(X161),X159,X160)|~v7_ordinal1(X164)|~v7_ordinal1(X163)|~m1_matrix_1(X162,u1_struct_0(X161),X159,X160)|(v2_struct_0(X161)|v6_struct_0(X161)|~v13_algstr_0(X161)|~v33_algstr_0(X161)|~v3_group_1(X161)|~v5_group_1(X161)|~v2_rlvect_1(X161)|~v3_rlvect_1(X161)|~v4_rlvect_1(X161)|~v4_vectsp_1(X161)|~v5_vectsp_1(X161)|~l6_algstr_0(X161))|~epred1_2(X160,X159))&(X166!=X164|k3_matrix_1(u1_struct_0(X161),X165,X166,X167)=k3_matrix_1(u1_struct_0(X161),X162,X163,X167)|(~r2_hidden(X166,k4_finseq_1(X162))|~r2_hidden(X167,k2_finseq_1(k1_matrix_1(X162))))|~v7_ordinal1(X167)|~v7_ordinal1(X166)|X165!=k1_matrix12(X159,X160,X161,X162,X163,X164)|~m1_matrix_1(X165,u1_struct_0(X161),X159,X160)|~v7_ordinal1(X164)|~v7_ordinal1(X163)|~m1_matrix_1(X162,u1_struct_0(X161),X159,X160)|(v2_struct_0(X161)|v6_struct_0(X161)|~v13_algstr_0(X161)|~v33_algstr_0(X161)|~v3_group_1(X161)|~v5_group_1(X161)|~v2_rlvect_1(X161)|~v3_rlvect_1(X161)|~v4_rlvect_1(X161)|~v4_vectsp_1(X161)|~v5_vectsp_1(X161)|~l6_algstr_0(X161))|~epred1_2(X160,X159)))&(X166=X163|X166=X164|k3_matrix_1(u1_struct_0(X161),X165,X166,X167)=k3_matrix_1(u1_struct_0(X161),X162,X166,X167)|(~r2_hidden(X166,k4_finseq_1(X162))|~r2_hidden(X167,k2_finseq_1(k1_matrix_1(X162))))|~v7_ordinal1(X167)|~v7_ordinal1(X166)|X165!=k1_matrix12(X159,X160,X161,X162,X163,X164)|~m1_matrix_1(X165,u1_struct_0(X161),X159,X160)|~v7_ordinal1(X164)|~v7_ordinal1(X163)|~m1_matrix_1(X162,u1_struct_0(X161),X159,X160)|(v2_struct_0(X161)|v6_struct_0(X161)|~v13_algstr_0(X161)|~v33_algstr_0(X161)|~v3_group_1(X161)|~v5_group_1(X161)|~v2_rlvect_1(X161)|~v3_rlvect_1(X161)|~v4_rlvect_1(X161)|~v4_vectsp_1(X161)|~v5_vectsp_1(X161)|~l6_algstr_0(X161))|~epred1_2(X160,X159))))&((v7_ordinal1(esk15_7(X159,X160,X161,X162,X163,X164,X165))|k3_finseq_1(X165)!=k3_finseq_1(X162)|X165=k1_matrix12(X159,X160,X161,X162,X163,X164)|~m1_matrix_1(X165,u1_struct_0(X161),X159,X160)|~v7_ordinal1(X164)|~v7_ordinal1(X163)|~m1_matrix_1(X162,u1_struct_0(X161),X159,X160)|(v2_struct_0(X161)|v6_struct_0(X161)|~v13_algstr_0(X161)|~v33_algstr_0(X161)|~v3_group_1(X161)|~v5_group_1(X161)|~v2_rlvect_1(X161)|~v3_rlvect_1(X161)|~v4_rlvect_1(X161)|~v4_vectsp_1(X161)|~v5_vectsp_1(X161)|~l6_algstr_0(X161))|~epred1_2(X160,X159))&((v7_ordinal1(esk16_7(X159,X160,X161,X162,X163,X164,X165))|k3_finseq_1(X165)!=k3_finseq_1(X162)|X165=k1_matrix12(X159,X160,X161,X162,X163,X164)|~m1_matrix_1(X165,u1_struct_0(X161),X159,X160)|~v7_ordinal1(X164)|~v7_ordinal1(X163)|~m1_matrix_1(X162,u1_struct_0(X161),X159,X160)|(v2_struct_0(X161)|v6_struct_0(X161)|~v13_algstr_0(X161)|~v33_algstr_0(X161)|~v3_group_1(X161)|~v5_group_1(X161)|~v2_rlvect_1(X161)|~v3_rlvect_1(X161)|~v4_rlvect_1(X161)|~v4_vectsp_1(X161)|~v5_vectsp_1(X161)|~l6_algstr_0(X161))|~epred1_2(X160,X159))&(((r2_hidden(esk15_7(X159,X160,X161,X162,X163,X164,X165),k4_finseq_1(X162))|k3_finseq_1(X165)!=k3_finseq_1(X162)|X165=k1_matrix12(X159,X160,X161,X162,X163,X164)|~m1_matrix_1(X165,u1_struct_0(X161),X159,X160)|~v7_ordinal1(X164)|~v7_ordinal1(X163)|~m1_matrix_1(X162,u1_struct_0(X161),X159,X160)|(v2_struct_0(X161)|v6_struct_0(X161)|~v13_algstr_0(X161)|~v33_algstr_0(X161)|~v3_group_1(X161)|~v5_group_1(X161)|~v2_rlvect_1(X161)|~v3_rlvect_1(X161)|~v4_rlvect_1(X161)|~v4_vectsp_1(X161)|~v5_vectsp_1(X161)|~l6_algstr_0(X161))|~epred1_2(X160,X159))&(r2_hidden(esk16_7(X159,X160,X161,X162,X163,X164,X165),k2_finseq_1(k1_matrix_1(X162)))|k3_finseq_1(X165)!=k3_finseq_1(X162)|X165=k1_matrix12(X159,X160,X161,X162,X163,X164)|~m1_matrix_1(X165,u1_struct_0(X161),X159,X160)|~v7_ordinal1(X164)|~v7_ordinal1(X163)|~m1_matrix_1(X162,u1_struct_0(X161),X159,X160)|(v2_struct_0(X161)|v6_struct_0(X161)|~v13_algstr_0(X161)|~v33_algstr_0(X161)|~v3_group_1(X161)|~v5_group_1(X161)|~v2_rlvect_1(X161)|~v3_rlvect_1(X161)|~v4_rlvect_1(X161)|~v4_vectsp_1(X161)|~v5_vectsp_1(X161)|~l6_algstr_0(X161))|~epred1_2(X160,X159)))&(((((esk15_7(X159,X160,X161,X162,X163,X164,X165)!=X163|(esk15_7(X159,X160,X161,X162,X163,X164,X165)=X164|esk15_7(X159,X160,X161,X162,X163,X164,X165)=X163)|k3_finseq_1(X165)!=k3_finseq_1(X162)|X165=k1_matrix12(X159,X160,X161,X162,X163,X164)|~m1_matrix_1(X165,u1_struct_0(X161),X159,X160)|~v7_ordinal1(X164)|~v7_ordinal1(X163)|~m1_matrix_1(X162,u1_struct_0(X161),X159,X160)|(v2_struct_0(X161)|v6_struct_0(X161)|~v13_algstr_0(X161)|~v33_algstr_0(X161)|~v3_group_1(X161)|~v5_group_1(X161)|~v2_rlvect_1(X161)|~v3_rlvect_1(X161)|~v4_rlvect_1(X161)|~v4_vectsp_1(X161)|~v5_vectsp_1(X161)|~l6_algstr_0(X161))|~epred1_2(X160,X159))&(esk15_7(X159,X160,X161,X162,X163,X164,X165)!=X164|(esk15_7(X159,X160,X161,X162,X163,X164,X165)=X164|esk15_7(X159,X160,X161,X162,X163,X164,X165)=X163)|k3_finseq_1(X165)!=k3_finseq_1(X162)|X165=k1_matrix12(X159,X160,X161,X162,X163,X164)|~m1_matrix_1(X165,u1_struct_0(X161),X159,X160)|~v7_ordinal1(X164)|~v7_ordinal1(X163)|~m1_matrix_1(X162,u1_struct_0(X161),X159,X160)|(v2_struct_0(X161)|v6_struct_0(X161)|~v13_algstr_0(X161)|~v33_algstr_0(X161)|~v3_group_1(X161)|~v5_group_1(X161)|~v2_rlvect_1(X161)|~v3_rlvect_1(X161)|~v4_rlvect_1(X161)|~v4_vectsp_1(X161)|~v5_vectsp_1(X161)|~l6_algstr_0(X161))|~epred1_2(X160,X159)))&(k3_matrix_1(u1_struct_0(X161),X165,esk15_7(X159,X160,X161,X162,X163,X164,X165),esk16_7(X159,X160,X161,X162,X163,X164,X165))!=k3_matrix_1(u1_struct_0(X161),X162,esk15_7(X159,X160,X161,X162,X163,X164,X165),esk16_7(X159,X160,X161,X162,X163,X164,X165))|(esk15_7(X159,X160,X161,X162,X163,X164,X165)=X164|esk15_7(X159,X160,X161,X162,X163,X164,X165)=X163)|k3_finseq_1(X165)!=k3_finseq_1(X162)|X165=k1_matrix12(X159,X160,X161,X162,X163,X164)|~m1_matrix_1(X165,u1_struct_0(X161),X159,X160)|~v7_ordinal1(X164)|~v7_ordinal1(X163)|~m1_matrix_1(X162,u1_struct_0(X161),X159,X160)|(v2_struct_0(X161)|v6_struct_0(X161)|~v13_algstr_0(X161)|~v33_algstr_0(X161)|~v3_group_1(X161)|~v5_group_1(X161)|~v2_rlvect_1(X161)|~v3_rlvect_1(X161)|~v4_rlvect_1(X161)|~v4_vectsp_1(X161)|~v5_vectsp_1(X161)|~l6_algstr_0(X161))|~epred1_2(X160,X159)))&(((esk15_7(X159,X160,X161,X162,X163,X164,X165)!=X163|(k3_matrix_1(u1_struct_0(X161),X165,esk15_7(X159,X160,X161,X162,X163,X164,X165),esk16_7(X159,X160,X161,X162,X163,X164,X165))!=k3_matrix_1(u1_struct_0(X161),X162,X163,esk16_7(X159,X160,X161,X162,X163,X164,X165))|esk15_7(X159,X160,X161,X162,X163,X164,X165)=X163)|k3_finseq_1(X165)!=k3_finseq_1(X162)|X165=k1_matrix12(X159,X160,X161,X162,X163,X164)|~m1_matrix_1(X165,u1_struct_0(X161),X159,X160)|~v7_ordinal1(X164)|~v7_ordinal1(X163)|~m1_matrix_1(X162,u1_struct_0(X161),X159,X160)|(v2_struct_0(X161)|v6_struct_0(X161)|~v13_algstr_0(X161)|~v33_algstr_0(X161)|~v3_group_1(X161)|~v5_group_1(X161)|~v2_rlvect_1(X161)|~v3_rlvect_1(X161)|~v4_rlvect_1(X161)|~v4_vectsp_1(X161)|~v5_vectsp_1(X161)|~l6_algstr_0(X161))|~epred1_2(X160,X159))&(esk15_7(X159,X160,X161,X162,X163,X164,X165)!=X164|(k3_matrix_1(u1_struct_0(X161),X165,esk15_7(X159,X160,X161,X162,X163,X164,X165),esk16_7(X159,X160,X161,X162,X163,X164,X165))!=k3_matrix_1(u1_struct_0(X161),X162,X163,esk16_7(X159,X160,X161,X162,X163,X164,X165))|esk15_7(X159,X160,X161,X162,X163,X164,X165)=X163)|k3_finseq_1(X165)!=k3_finseq_1(X162)|X165=k1_matrix12(X159,X160,X161,X162,X163,X164)|~m1_matrix_1(X165,u1_struct_0(X161),X159,X160)|~v7_ordinal1(X164)|~v7_ordinal1(X163)|~m1_matrix_1(X162,u1_struct_0(X161),X159,X160)|(v2_struct_0(X161)|v6_struct_0(X161)|~v13_algstr_0(X161)|~v33_algstr_0(X161)|~v3_group_1(X161)|~v5_group_1(X161)|~v2_rlvect_1(X161)|~v3_rlvect_1(X161)|~v4_rlvect_1(X161)|~v4_vectsp_1(X161)|~v5_vectsp_1(X161)|~l6_algstr_0(X161))|~epred1_2(X160,X159)))&(k3_matrix_1(u1_struct_0(X161),X165,esk15_7(X159,X160,X161,X162,X163,X164,X165),esk16_7(X159,X160,X161,X162,X163,X164,X165))!=k3_matrix_1(u1_struct_0(X161),X162,esk15_7(X159,X160,X161,X162,X163,X164,X165),esk16_7(X159,X160,X161,X162,X163,X164,X165))|(k3_matrix_1(u1_struct_0(X161),X165,esk15_7(X159,X160,X161,X162,X163,X164,X165),esk16_7(X159,X160,X161,X162,X163,X164,X165))!=k3_matrix_1(u1_struct_0(X161),X162,X163,esk16_7(X159,X160,X161,X162,X163,X164,X165))|esk15_7(X159,X160,X161,X162,X163,X164,X165)=X163)|k3_finseq_1(X165)!=k3_finseq_1(X162)|X165=k1_matrix12(X159,X160,X161,X162,X163,X164)|~m1_matrix_1(X165,u1_struct_0(X161),X159,X160)|~v7_ordinal1(X164)|~v7_ordinal1(X163)|~m1_matrix_1(X162,u1_struct_0(X161),X159,X160)|(v2_struct_0(X161)|v6_struct_0(X161)|~v13_algstr_0(X161)|~v33_algstr_0(X161)|~v3_group_1(X161)|~v5_group_1(X161)|~v2_rlvect_1(X161)|~v3_rlvect_1(X161)|~v4_rlvect_1(X161)|~v4_vectsp_1(X161)|~v5_vectsp_1(X161)|~l6_algstr_0(X161))|~epred1_2(X160,X159))))&((((esk15_7(X159,X160,X161,X162,X163,X164,X165)!=X163|(esk15_7(X159,X160,X161,X162,X163,X164,X165)=X164|k3_matrix_1(u1_struct_0(X161),X165,esk15_7(X159,X160,X161,X162,X163,X164,X165),esk16_7(X159,X160,X161,X162,X163,X164,X165))!=k3_matrix_1(u1_struct_0(X161),X162,X164,esk16_7(X159,X160,X161,X162,X163,X164,X165)))|k3_finseq_1(X165)!=k3_finseq_1(X162)|X165=k1_matrix12(X159,X160,X161,X162,X163,X164)|~m1_matrix_1(X165,u1_struct_0(X161),X159,X160)|~v7_ordinal1(X164)|~v7_ordinal1(X163)|~m1_matrix_1(X162,u1_struct_0(X161),X159,X160)|(v2_struct_0(X161)|v6_struct_0(X161)|~v13_algstr_0(X161)|~v33_algstr_0(X161)|~v3_group_1(X161)|~v5_group_1(X161)|~v2_rlvect_1(X161)|~v3_rlvect_1(X161)|~v4_rlvect_1(X161)|~v4_vectsp_1(X161)|~v5_vectsp_1(X161)|~l6_algstr_0(X161))|~epred1_2(X160,X159))&(esk15_7(X159,X160,X161,X162,X163,X164,X165)!=X164|(esk15_7(X159,X160,X161,X162,X163,X164,X165)=X164|k3_matrix_1(u1_struct_0(X161),X165,esk15_7(X159,X160,X161,X162,X163,X164,X165),esk16_7(X159,X160,X161,X162,X163,X164,X165))!=k3_matrix_1(u1_struct_0(X161),X162,X164,esk16_7(X159,X160,X161,X162,X163,X164,X165)))|k3_finseq_1(X165)!=k3_finseq_1(X162)|X165=k1_matrix12(X159,X160,X161,X162,X163,X164)|~m1_matrix_1(X165,u1_struct_0(X161),X159,X160)|~v7_ordinal1(X164)|~v7_ordinal1(X163)|~m1_matrix_1(X162,u1_struct_0(X161),X159,X160)|(v2_struct_0(X161)|v6_struct_0(X161)|~v13_algstr_0(X161)|~v33_algstr_0(X161)|~v3_group_1(X161)|~v5_group_1(X161)|~v2_rlvect_1(X161)|~v3_rlvect_1(X161)|~v4_rlvect_1(X161)|~v4_vectsp_1(X161)|~v5_vectsp_1(X161)|~l6_algstr_0(X161))|~epred1_2(X160,X159)))&(k3_matrix_1(u1_struct_0(X161),X165,esk15_7(X159,X160,X161,X162,X163,X164,X165),esk16_7(X159,X160,X161,X162,X163,X164,X165))!=k3_matrix_1(u1_struct_0(X161),X162,esk15_7(X159,X160,X161,X162,X163,X164,X165),esk16_7(X159,X160,X161,X162,X163,X164,X165))|(esk15_7(X159,X160,X161,X162,X163,X164,X165)=X164|k3_matrix_1(u1_struct_0(X161),X165,esk15_7(X159,X160,X161,X162,X163,X164,X165),esk16_7(X159,X160,X161,X162,X163,X164,X165))!=k3_matrix_1(u1_struct_0(X161),X162,X164,esk16_7(X159,X160,X161,X162,X163,X164,X165)))|k3_finseq_1(X165)!=k3_finseq_1(X162)|X165=k1_matrix12(X159,X160,X161,X162,X163,X164)|~m1_matrix_1(X165,u1_struct_0(X161),X159,X160)|~v7_ordinal1(X164)|~v7_ordinal1(X163)|~m1_matrix_1(X162,u1_struct_0(X161),X159,X160)|(v2_struct_0(X161)|v6_struct_0(X161)|~v13_algstr_0(X161)|~v33_algstr_0(X161)|~v3_group_1(X161)|~v5_group_1(X161)|~v2_rlvect_1(X161)|~v3_rlvect_1(X161)|~v4_rlvect_1(X161)|~v4_vectsp_1(X161)|~v5_vectsp_1(X161)|~l6_algstr_0(X161))|~epred1_2(X160,X159)))&(((esk15_7(X159,X160,X161,X162,X163,X164,X165)!=X163|(k3_matrix_1(u1_struct_0(X161),X165,esk15_7(X159,X160,X161,X162,X163,X164,X165),esk16_7(X159,X160,X161,X162,X163,X164,X165))!=k3_matrix_1(u1_struct_0(X161),X162,X163,esk16_7(X159,X160,X161,X162,X163,X164,X165))|k3_matrix_1(u1_struct_0(X161),X165,esk15_7(X159,X160,X161,X162,X163,X164,X165),esk16_7(X159,X160,X161,X162,X163,X164,X165))!=k3_matrix_1(u1_struct_0(X161),X162,X164,esk16_7(X159,X160,X161,X162,X163,X164,X165)))|k3_finseq_1(X165)!=k3_finseq_1(X162)|X165=k1_matrix12(X159,X160,X161,X162,X163,X164)|~m1_matrix_1(X165,u1_struct_0(X161),X159,X160)|~v7_ordinal1(X164)|~v7_ordinal1(X163)|~m1_matrix_1(X162,u1_struct_0(X161),X159,X160)|(v2_struct_0(X161)|v6_struct_0(X161)|~v13_algstr_0(X161)|~v33_algstr_0(X161)|~v3_group_1(X161)|~v5_group_1(X161)|~v2_rlvect_1(X161)|~v3_rlvect_1(X161)|~v4_rlvect_1(X161)|~v4_vectsp_1(X161)|~v5_vectsp_1(X161)|~l6_algstr_0(X161))|~epred1_2(X160,X159))&(esk15_7(X159,X160,X161,X162,X163,X164,X165)!=X164|(k3_matrix_1(u1_struct_0(X161),X165,esk15_7(X159,X160,X161,X162,X163,X164,X165),esk16_7(X159,X160,X161,X162,X163,X164,X165))!=k3_matrix_1(u1_struct_0(X161),X162,X163,esk16_7(X159,X160,X161,X162,X163,X164,X165))|k3_matrix_1(u1_struct_0(X161),X165,esk15_7(X159,X160,X161,X162,X163,X164,X165),esk16_7(X159,X160,X161,X162,X163,X164,X165))!=k3_matrix_1(u1_struct_0(X161),X162,X164,esk16_7(X159,X160,X161,X162,X163,X164,X165)))|k3_finseq_1(X165)!=k3_finseq_1(X162)|X165=k1_matrix12(X159,X160,X161,X162,X163,X164)|~m1_matrix_1(X165,u1_struct_0(X161),X159,X160)|~v7_ordinal1(X164)|~v7_ordinal1(X163)|~m1_matrix_1(X162,u1_struct_0(X161),X159,X160)|(v2_struct_0(X161)|v6_struct_0(X161)|~v13_algstr_0(X161)|~v33_algstr_0(X161)|~v3_group_1(X161)|~v5_group_1(X161)|~v2_rlvect_1(X161)|~v3_rlvect_1(X161)|~v4_rlvect_1(X161)|~v4_vectsp_1(X161)|~v5_vectsp_1(X161)|~l6_algstr_0(X161))|~epred1_2(X160,X159)))&(k3_matrix_1(u1_struct_0(X161),X165,esk15_7(X159,X160,X161,X162,X163,X164,X165),esk16_7(X159,X160,X161,X162,X163,X164,X165))!=k3_matrix_1(u1_struct_0(X161),X162,esk15_7(X159,X160,X161,X162,X163,X164,X165),esk16_7(X159,X160,X161,X162,X163,X164,X165))|(k3_matrix_1(u1_struct_0(X161),X165,esk15_7(X159,X160,X161,X162,X163,X164,X165),esk16_7(X159,X160,X161,X162,X163,X164,X165))!=k3_matrix_1(u1_struct_0(X161),X162,X163,esk16_7(X159,X160,X161,X162,X163,X164,X165))|k3_matrix_1(u1_struct_0(X161),X165,esk15_7(X159,X160,X161,X162,X163,X164,X165),esk16_7(X159,X160,X161,X162,X163,X164,X165))!=k3_matrix_1(u1_struct_0(X161),X162,X164,esk16_7(X159,X160,X161,X162,X163,X164,X165)))|k3_finseq_1(X165)!=k3_finseq_1(X162)|X165=k1_matrix12(X159,X160,X161,X162,X163,X164)|~m1_matrix_1(X165,u1_struct_0(X161),X159,X160)|~v7_ordinal1(X164)|~v7_ordinal1(X163)|~m1_matrix_1(X162,u1_struct_0(X161),X159,X160)|(v2_struct_0(X161)|v6_struct_0(X161)|~v13_algstr_0(X161)|~v33_algstr_0(X161)|~v3_group_1(X161)|~v5_group_1(X161)|~v2_rlvect_1(X161)|~v3_rlvect_1(X161)|~v4_rlvect_1(X161)|~v4_vectsp_1(X161)|~v5_vectsp_1(X161)|~l6_algstr_0(X161))|~epred1_2(X160,X159))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_175])])])])])])).
cnf(c_0_184, negated_conjecture, (k1_funct_1(X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(X4,X5,X6,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|v1_xboole_0(X4)|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|X1!=k6_matrix_1(X4,X5,X6)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~m2_finseq_1(X5,k3_finseq_2(X4))|~m2_finseq_1(X1,X4)|~v1_matrix_1(X5)|~m1_matrix_1(X5,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X6)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v7_ordinal1(X2)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(X3)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(X3)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_177])])).
cnf(c_0_185, plain, (X1=X2|v7_ordinal1(esk10_2(X1,X2))|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_178])).
cnf(c_0_186, negated_conjecture, (v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_163]), c_0_156]), c_0_147])]), c_0_136])])).
cnf(c_0_187, negated_conjecture, (v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_163]), c_0_156]), c_0_147])]), c_0_136])])).
cnf(c_0_188, negated_conjecture, (v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_163]), c_0_156]), c_0_147])]), c_0_136])])).
cnf(c_0_189, negated_conjecture, (k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))=k1_matrix_1(esk7_0)|v1_xboole_0(u1_struct_0(esk5_0))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_182])])).
cnf(c_0_190, plain, (X1=X2|X1=X3|k3_matrix_1(u1_struct_0(X4),X5,X1,X6)=k3_matrix_1(u1_struct_0(X4),X7,X1,X6)|v2_struct_0(X4)|v6_struct_0(X4)|~r2_hidden(X1,k4_finseq_1(X7))|~r2_hidden(X6,k2_finseq_1(k1_matrix_1(X7)))|~v7_ordinal1(X6)|~v7_ordinal1(X1)|X5!=k1_matrix12(X8,X9,X4,X7,X2,X3)|~m1_matrix_1(X5,u1_struct_0(X4),X8,X9)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X7,u1_struct_0(X4),X8,X9)|~v13_algstr_0(X4)|~v33_algstr_0(X4)|~v3_group_1(X4)|~v5_group_1(X4)|~v2_rlvect_1(X4)|~v3_rlvect_1(X4)|~v4_rlvect_1(X4)|~v4_vectsp_1(X4)|~v5_vectsp_1(X4)|~l6_algstr_0(X4)|~epred1_2(X9,X8)), inference(split_conjunct,[status(thm)],[c_0_183])).
cnf(c_0_191, negated_conjecture, (k1_funct_1(X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(X4,X5,X6,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|v1_xboole_0(X4)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))!=k3_finseq_1(X3)|X1!=k6_matrix_1(X4,X5,X6)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m2_finseq_1(X5,k3_finseq_2(X4))|~m2_finseq_1(X1,X4)|~v1_matrix_1(X5)|~m1_matrix_1(X5,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X6)|~v7_ordinal1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_184, c_0_185]), c_0_186]), c_0_187]), c_0_188])).
cnf(c_0_192, negated_conjecture, (k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))=k1_matrix_1(esk6_0)|v1_xboole_0(u1_struct_0(esk5_0))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_189, c_0_154])])).
cnf(c_0_193, plain, (X1=X3|X1=X2|k3_matrix_1(u1_struct_0(X4),X5,X1,X6)=k3_matrix_1(u1_struct_0(X4),X7,X1,X6)|v2_struct_0(X4)|v6_struct_0(X4)|X5!=k1_matrix12(X8,X9,X4,X7,X2,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v13_algstr_0(X4)|~v33_algstr_0(X4)|~v3_group_1(X4)|~v5_group_1(X4)|~v2_rlvect_1(X4)|~v3_rlvect_1(X4)|~v4_rlvect_1(X4)|~v4_vectsp_1(X4)|~v5_vectsp_1(X4)|~l6_algstr_0(X4)|~epred1_2(X9,X8)|~r2_hidden(X1,k4_finseq_1(X7))|~r2_hidden(X6,k2_finseq_1(k1_matrix_1(X7)))|~m1_matrix_1(X7,u1_struct_0(X4),X8,X9)|~m1_matrix_1(X5,u1_struct_0(X4),X8,X9)), inference(evalgc,[status(thm)],[c_0_190])).
cnf(c_0_194, negated_conjecture, (esk7_0=k1_matrix12(esk2_0,esk1_0,esk5_0,esk6_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_195, negated_conjecture, (v7_ordinal1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_196, negated_conjecture, (v7_ordinal1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_197, plain, (X1=X2|k1_funct_1(X1,esk10_2(X1,X2))!=k1_funct_1(X2,esk10_2(X1,X2))|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_169])).
cnf(c_0_198, negated_conjecture, (k1_funct_1(k6_matrix_1(X1,X2,X3),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4),X5))=k3_matrix_1(X1,X2,X3,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4),X5))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4)=X5|v1_xboole_0(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4))!=k3_finseq_1(X5)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4),X5),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4),X5))|~v1_finseq_1(X5)|~v1_funct_1(X5)|~v1_relat_1(X5)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m2_finseq_1(k6_matrix_1(X1,X2,X3),X1)|~v1_matrix_1(X2)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X3)|~v7_ordinal1(X4)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_191])])).
cnf(c_0_199, negated_conjecture, (k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))=k1_matrix_1(esk6_0)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)), inference(sr,[status(thm)],[c_0_192, c_0_136])).
cnf(c_0_200, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8)=k3_matrix_1(u1_struct_0(X1),X4,X7,X8)|X7=X5|X7=X6|v6_struct_0(X1)|v2_struct_0(X1)|~epred1_2(X3,X2)|~r2_hidden(X8,k2_finseq_1(k1_matrix_1(X4)))|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~m1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6),u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X8)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_193])])).
cnf(c_0_201, negated_conjecture, (k1_matrix12(esk2_0,esk1_0,esk5_0,esk6_0,esk3_0,esk4_0)=esk7_0), inference(evalgc,[status(thm)],[c_0_194])).
cnf(c_0_202, negated_conjecture, (v7_ordinal1(esk4_0)), inference(evalgc,[status(thm)],[c_0_195])).
cnf(c_0_203, negated_conjecture, (v7_ordinal1(esk3_0)), inference(evalgc,[status(thm)],[c_0_196])).
cnf(c_0_204, plain, (X1=X2|k3_finseq_1(X1)!=k3_finseq_1(X2)|k1_funct_1(X1,esk10_2(X1,X2))!=k1_funct_1(X2,esk10_2(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_197])).
cnf(c_0_205, negated_conjecture, (k1_funct_1(k6_matrix_1(X1,X2,X3),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4),X5))=k3_matrix_1(X1,X2,X3,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4),X5))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4)=X5|v1_xboole_0(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4))!=k3_finseq_1(X5)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4),X5),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4),X5))|~v1_finseq_1(X5)|~v1_funct_1(X5)|~v1_relat_1(X5)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X3)|~v7_ordinal1(X4)), inference(csr,[status(thm)],[c_0_198, c_0_121])).
cnf(c_0_206, negated_conjecture, (k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))=k1_matrix_1(esk6_0)|~v7_ordinal1(X1)|~m2_finseq_1(esk7_0,k3_finseq_2(u1_struct_0(esk5_0)))|~v1_matrix_1(esk7_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_121]), c_0_136])])).
cnf(c_0_207, negated_conjecture, (m2_finseq_1(esk7_0,k3_finseq_2(u1_struct_0(esk5_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[c_0_130, c_0_136])])).
cnf(c_0_208, negated_conjecture, (v1_matrix_1(esk7_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[c_0_128, c_0_136])])).
fof(c_0_209, plain, ![X102, X103, X104]:(~r2_hidden(X102,X103)|~m1_subset_1(X103,k1_zfmisc_1(X104))|~v1_xboole_0(X104)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_210, plain, (~v1_xboole_0(k4_ordinal1)&v3_ordinal1(k4_ordinal1)), inference(fof_simplification,[status(thm)],[fc6_ordinal1])).
cnf(c_0_211, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|X1=esk4_0|X1=esk3_0|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(esk6_0)))|~r2_hidden(X1,k4_finseq_1(esk6_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_201]), c_0_103]), c_0_48]), c_0_108]), c_0_109]), c_0_110]), c_0_111]), c_0_112]), c_0_113]), c_0_114]), c_0_115]), c_0_116]), c_0_117]), c_0_202]), c_0_203])]), c_0_118]), c_0_119])])).
cnf(c_0_212, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(X2,X3,X4)|v1_xboole_0(X2)|k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)))!=k3_matrix_1(X2,X3,X4,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(X2,X3,X4))|~v1_finseq_1(k6_matrix_1(X2,X3,X4))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k6_matrix_1(X2,X3,X4))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k6_matrix_1(X2,X3,X4))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3)|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X4)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_205])])).
cnf(c_0_213, negated_conjecture, (k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))=k1_matrix_1(esk6_0)|~v7_ordinal1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_206, c_0_207]), c_0_208])])).
cnf(c_0_214, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_207]), c_0_208])]), c_0_136])])).
cnf(c_0_215, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_209])).
cnf(c_0_216, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[dt_k5_numbers])).
fof(c_0_217, plain, ![X32, X33]:(~m1_subset_1(X32,X33)|(v1_xboole_0(X33)|r2_hidden(X32,X33))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_218, plain, (~v1_xboole_0(k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_210])).
cnf(c_0_219, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|X1=esk3_0|X1=esk4_0|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k4_finseq_1(esk6_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,k1_matrix_1(esk6_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_165]), c_0_166])])])).
cnf(c_0_220, negated_conjecture, (r2_hidden(esk8_0,k4_finseq_1(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_221, negated_conjecture, (v7_ordinal1(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_222, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(X2,X3,X4)|v1_xboole_0(X2)|k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)))!=k3_matrix_1(X2,X3,X4,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(X2,X3,X4))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3)|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X4)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_212, c_0_179]), c_0_186]), c_0_180]), c_0_187]), c_0_181]), c_0_188])).
cnf(c_0_223, negated_conjecture, (k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))=k1_matrix_1(esk6_0)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_214])])).
fof(c_0_224, plain, ![X107, X108, X109]:((~m2_subset_1(X109,X107,X108)|m1_subset_1(X109,X108)|(v1_xboole_0(X107)|v1_xboole_0(X108)|~m1_subset_1(X108,k1_zfmisc_1(X107))))&(~m1_subset_1(X109,X108)|m2_subset_1(X109,X107,X108)|(v1_xboole_0(X107)|v1_xboole_0(X108)|~m1_subset_1(X108,k1_zfmisc_1(X107))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m2_subset_1])])])])])).
cnf(c_0_225, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_215])).
cnf(c_0_226, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_216])).
cnf(c_0_227, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_217])).
cnf(c_0_228, plain, (~v1_xboole_0(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_218])).
fof(c_0_229, plain, ![X121]:m1_subset_1(esk12_1(X121),X121), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[existence_m1_subset_1])])).
cnf(c_0_230, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|X1=esk4_0|X1=esk3_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(X3))|~r1_xxreal_0(np__1,X2)|~r2_hidden(X1,k4_finseq_1(esk6_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_167]), c_0_48])])])).
cnf(c_0_231, negated_conjecture, (r2_hidden(esk8_0,k4_finseq_1(esk6_0))), inference(evalgc,[status(thm)],[c_0_220])).
cnf(c_0_232, negated_conjecture, (v7_ordinal1(esk8_0)), inference(evalgc,[status(thm)],[c_0_221])).
cnf(c_0_233, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_214]), c_0_207]), c_0_208]), c_0_103])]), c_0_136])])).
cnf(c_0_234, negated_conjecture, (k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_168])])).
cnf(c_0_235, plain, (m2_subset_1(X1,X3,X2)|v1_xboole_0(X3)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_224])).
cnf(c_0_236, plain, (~v1_xboole_0(k1_numbers)|~r2_hidden(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_226])])).
cnf(c_0_237, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_227])).
cnf(c_0_238, plain, (~v1_xboole_0(k5_numbers)), inference(rw,[status(thm)],[c_0_228, c_0_78])).
cnf(c_0_239, plain, (m1_subset_1(esk12_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_229])).
cnf(c_0_240, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X1)|esk3_0=esk8_0|esk4_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_matrix_1(X2))|~r1_xxreal_0(np__1,X1)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_231]), c_0_232])])])).
cnf(c_0_241, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_233, c_0_234])).
cnf(c_0_242, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))!=k3_finseq_1(X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_163]), c_0_156]), c_0_147]), c_0_48])]), c_0_136])])).
cnf(c_0_243, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_235])).
cnf(c_0_244, plain, (~v1_xboole_0(k1_numbers)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_236, c_0_237]), c_0_238])])).
cnf(c_0_245, plain, (m1_subset_1(esk12_1(X1),X1)), inference(evalgc,[status(thm)],[c_0_239])).
cnf(c_0_246, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X1)|esk4_0=esk8_0|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_168]), c_0_48])])])).
cnf(c_0_247, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_242])])).
cnf(c_0_248, negated_conjecture, (v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_214]), c_0_207]), c_0_208])]), c_0_136])])).
cnf(c_0_249, negated_conjecture, (v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_214]), c_0_207]), c_0_208])]), c_0_136])])).
cnf(c_0_250, negated_conjecture, (v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_214]), c_0_207]), c_0_208])]), c_0_136])])).
cnf(c_0_251, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|v1_xboole_0(k1_numbers)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_226]), c_0_238])])).
cnf(c_0_252, plain, (~v1_xboole_0(k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_244, c_0_245])])).
cnf(c_0_253, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|esk3_0=esk8_0|esk4_0=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_177])])).
cnf(c_0_254, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_247, c_0_248]), c_0_249]), c_0_250]), c_0_234])).
cnf(c_0_255, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(sr,[status(thm)],[c_0_251, c_0_252])).
cnf(c_0_256, plain, (r1_xxreal_0(np__1,esk10_2(X1,X2))|X1=X2|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_169])).
cnf(c_0_257, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_253, c_0_185]), c_0_186]), c_0_187]), c_0_188])).
cnf(c_0_258, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_255])])).
cnf(c_0_259, plain, (X1=X2|r1_xxreal_0(np__1,esk10_2(X1,X2))|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_256])).
cnf(c_0_260, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_255])])).
cnf(c_0_261, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_259])])).
cnf(c_0_262, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|esk3_0=esk8_0|esk4_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers)|~v7_ordinal1(X1)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_259])])).
fof(c_0_263, plain, ![X27]:((~v7_ordinal1(X27)|r2_hidden(X27,k4_ordinal1))&(~r2_hidden(X27,k4_ordinal1)|v7_ordinal1(X27))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_ordinal1])])).
cnf(c_0_264, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_261, c_0_186]), c_0_248]), c_0_187]), c_0_249]), c_0_188]), c_0_250]), c_0_234])).
cnf(c_0_265, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_262, c_0_186]), c_0_187]), c_0_188])).
fof(c_0_266, plain, ![X34, X35]:(~r2_hidden(X34,X35)|m1_subset_1(X34,X35)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_267, plain, (r2_hidden(X1,k4_ordinal1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_263])).
cnf(c_0_268, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk4_0=esk8_0|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_265]), c_0_232])])])).
cnf(c_0_269, negated_conjecture, (esk8_0=esk4_0|esk8_0=esk3_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_270, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_266])).
cnf(c_0_271, plain, (r2_hidden(X1,k4_ordinal1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_267])).
cnf(c_0_272, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|esk4_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_268]), c_0_232])])])).
cnf(c_0_273, negated_conjecture, (esk3_0=esk8_0|esk4_0=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)), inference(evalgc,[status(thm)],[c_0_269])).
cnf(c_0_274, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_270])).
cnf(c_0_275, plain, (r2_hidden(X1,k5_numbers)|~v7_ordinal1(X1)), inference(rw,[status(thm)],[c_0_271, c_0_78])).
cnf(c_0_276, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,esk7_0,X4,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,esk7_0,X4)|~m2_finseq_1(esk7_0,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v1_matrix_1(esk7_0)|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(esk6_0)))|~v7_ordinal1(X4)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_154])])).
fof(c_0_277, plain, ![X41]:((l4_algstr_0(X41)|~l5_algstr_0(X41))&(l4_struct_0(X41)|~l5_algstr_0(X41))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l5_algstr_0])])])).
cnf(c_0_278, negated_conjecture, (esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)), inference(csr,[status(thm)],[c_0_272, c_0_273])).
cnf(c_0_279, plain, (m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_274, c_0_275])])).
cnf(c_0_280, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,esk7_0,X4,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,esk7_0,X4)|~m2_finseq_1(esk7_0,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(esk6_0)))|~v7_ordinal1(X4)|~v7_ordinal1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_276, c_0_208])])).
fof(c_0_281, plain, ![X65]:(~l4_struct_0(X65)|(v6_struct_0(X65)|~v7_struct_0(X65))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc3_struct_0])])])).
cnf(c_0_282, plain, (l4_struct_0(X1)|~l5_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_277])).
cnf(c_0_283, plain, (l5_algstr_0(X1)|~l6_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_284, negated_conjecture, (esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_279])])).
cnf(c_0_285, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,esk7_0,X4,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,esk7_0,X4)|~m2_finseq_1(esk7_0,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,k1_matrix_1(esk6_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_165]), c_0_166])])])).
cnf(c_0_286, plain, (k3_matrix_1(u1_struct_0(X3),X4,X1,X5)=k3_matrix_1(u1_struct_0(X3),X6,X7,X5)|v2_struct_0(X3)|v6_struct_0(X3)|X1!=X2|~r2_hidden(X1,k4_finseq_1(X6))|~r2_hidden(X5,k2_finseq_1(k1_matrix_1(X6)))|~v7_ordinal1(X5)|~v7_ordinal1(X1)|X4!=k1_matrix12(X8,X9,X3,X6,X7,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X8,X9)|~v7_ordinal1(X2)|~v7_ordinal1(X7)|~m1_matrix_1(X6,u1_struct_0(X3),X8,X9)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~epred1_2(X9,X8)), inference(split_conjunct,[status(thm)],[c_0_183])).
cnf(c_0_287, plain, (v6_struct_0(X1)|~l4_struct_0(X1)|~v7_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_281])).
cnf(c_0_288, plain, (l4_struct_0(X1)|~l5_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_282])).
cnf(c_0_289, plain, (l5_algstr_0(X1)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_283])).
fof(c_0_290, plain, ![X88]:(~l1_struct_0(X88)|(v7_struct_0(X88)|~v2_struct_0(X88))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_struct_0])])])).
cnf(c_0_291, negated_conjecture, (esk3_0=esk8_0|esk4_0=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_284, c_0_185])])).
cnf(c_0_292, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,esk7_0,X4,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,esk7_0,X4)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(X5))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(esk7_0,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~m1_matrix_1(X5,u1_struct_0(esk5_0),esk2_0,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_285, c_0_167]), c_0_48])])])).
fof(c_0_293, plain, ![X73, X74, X75]:(~m1_finseq_2(X74,X73)|(~m2_finseq_2(X75,X73,X74)|m2_finseq_1(X75,X73))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m2_finseq_2])])])).
fof(c_0_294, plain, ![X84, X85, X86]:(v1_xboole_0(X84)|~v1_matrix_1(X85)|~m1_finseq_1(X85,k3_finseq_2(X84))|~v7_ordinal1(X86)|m2_finseq_2(k8_matrix_1(X84,X85,X86),X84,k4_finseq_2(k1_matrix_1(X85),X84))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k8_matrix_1])])])).
cnf(c_0_295, plain, (k3_matrix_1(u1_struct_0(X3),X4,X1,X5)=k3_matrix_1(u1_struct_0(X3),X6,X7,X5)|v2_struct_0(X3)|v6_struct_0(X3)|X1!=X2|X4!=k1_matrix12(X8,X9,X3,X6,X7,X2)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~epred1_2(X9,X8)|~r2_hidden(X1,k4_finseq_1(X6))|~r2_hidden(X5,k2_finseq_1(k1_matrix_1(X6)))|~m1_matrix_1(X6,u1_struct_0(X3),X8,X9)|~m1_matrix_1(X4,u1_struct_0(X3),X8,X9)), inference(evalgc,[status(thm)],[c_0_286])).
fof(c_0_296, plain, ![X43, X44, X45, X46, X47, X48]:(~v7_ordinal1(X43)|~v7_ordinal1(X44)|v2_struct_0(X45)|v6_struct_0(X45)|~v13_algstr_0(X45)|~v33_algstr_0(X45)|~v3_group_1(X45)|~v5_group_1(X45)|~v2_rlvect_1(X45)|~v3_rlvect_1(X45)|~v4_rlvect_1(X45)|~v4_vectsp_1(X45)|~v5_vectsp_1(X45)|~l6_algstr_0(X45)|~m1_matrix_1(X46,u1_struct_0(X45),X43,X44)|~v7_ordinal1(X47)|~v7_ordinal1(X48)|m1_matrix_1(k1_matrix12(X43,X44,X45,X46,X47,X48),u1_struct_0(X45),X43,X44)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k1_matrix12])])])).
fof(c_0_297, axiom, ![X1]:(v7_ordinal1(X1)=>![X2]:(v7_ordinal1(X2)=>epred1_2(X2,X1))), inference(apply_def,[status(thm)],[d1_matrix12, c_0_170])).
cnf(c_0_298, plain, (v6_struct_0(X1)|~l4_struct_0(X1)|~v7_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_287])).
cnf(c_0_299, plain, (l4_struct_0(X1)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_289])])).
cnf(c_0_300, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_290])).
cnf(c_0_301, negated_conjecture, (esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))), inference(csr,[status(thm)],[c_0_291, c_0_273])).
cnf(c_0_302, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X3,X2)|X1!=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(X4))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_207]), c_0_136])])).
cnf(c_0_303, plain, (m2_finseq_1(X3,X2)|~m1_finseq_2(X1,X2)|~m2_finseq_2(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_293])).
cnf(c_0_304, plain, (v1_xboole_0(X1)|m2_finseq_2(k8_matrix_1(X1,X2,X3),X1,k4_finseq_2(k1_matrix_1(X2),X1))|~v1_matrix_1(X2)|~m1_finseq_1(X2,k3_finseq_2(X1))|~v7_ordinal1(X3)), inference(split_conjunct,[status(thm)],[c_0_294])).
fof(c_0_305, plain, ![X81, X82]:(~v7_ordinal1(X81)|m1_finseq_2(k4_finseq_2(X81,X82),X82)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_finseq_2])])).
cnf(c_0_306, plain, (k3_matrix_1(u1_struct_0(X1),X2,X3,X4)=k3_matrix_1(u1_struct_0(X1),X5,X6,X4)|v6_struct_0(X1)|v2_struct_0(X1)|X2!=k1_matrix12(X7,X8,X1,X5,X6,X3)|~epred1_2(X8,X7)|~r2_hidden(X4,k2_finseq_1(k1_matrix_1(X5)))|~r2_hidden(X3,k4_finseq_1(X5))|~m1_matrix_1(X5,u1_struct_0(X1),X7,X8)|~m1_matrix_1(X2,u1_struct_0(X1),X7,X8)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_295])])).
cnf(c_0_307, plain, (v2_struct_0(X3)|v6_struct_0(X3)|m1_matrix_1(k1_matrix12(X1,X2,X3,X4,X5,X6),u1_struct_0(X3),X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X6)), inference(split_conjunct,[status(thm)],[c_0_296])).
fof(c_0_308, plain, ![X18, X19]:(~v7_ordinal1(X18)|(~v7_ordinal1(X19)|epred1_2(X19,X18))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_297])])])).
cnf(c_0_309, plain, (v6_struct_0(X1)|~v7_struct_0(X1)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_299])])).
cnf(c_0_310, plain, (v7_struct_0(X1)|~v2_struct_0(X1)|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_300])).
cnf(c_0_311, plain, (m1_subset_1(X1,X3)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_224])).
cnf(c_0_312, negated_conjecture, (esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_301, c_0_250]), c_0_232])])])).
cnf(c_0_313, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(X3))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_302])])).
cnf(c_0_314, plain, (m2_finseq_1(X3,X2)|~m1_finseq_2(X1,X2)|~m2_finseq_2(X3,X2,X1)), inference(evalgc,[status(thm)],[c_0_303])).
cnf(c_0_315, plain, (v1_xboole_0(X1)|m2_finseq_2(k8_matrix_1(X1,X2,X3),X1,k4_finseq_2(k1_matrix_1(X2),X1))|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m1_finseq_1(X2,k3_finseq_2(X1))), inference(evalgc,[status(thm)],[c_0_304])).
cnf(c_0_316, plain, (m1_finseq_2(k4_finseq_2(X1,X2),X2)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_305])).
cnf(c_0_317, plain, (k1_funct_1(X1,X2)=k3_matrix_1(X3,X4,X5,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,X4,X5)|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v1_matrix_1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,k1_matrix_1(X4))|~v7_ordinal1(k1_matrix_1(X4))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_165])])).
cnf(c_0_318, negated_conjecture, (k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))=k1_matrix_1(X1)|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_223]), c_0_48])])])).
cnf(c_0_319, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|v2_struct_0(X1)|~epred1_2(X3,X2)|~r2_hidden(X7,k2_finseq_1(k1_matrix_1(X4)))|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~m1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6),u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X6)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_306])])).
cnf(c_0_320, plain, (v2_struct_0(X3)|v6_struct_0(X3)|m1_matrix_1(k1_matrix12(X1,X2,X3,X4,X5,X6),u1_struct_0(X3),X1,X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)), inference(evalgc,[status(thm)],[c_0_307])).
cnf(c_0_321, plain, (epred1_2(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_308])).
cnf(c_0_322, plain, (v6_struct_0(X1)|~l6_algstr_0(X1)|~l1_struct_0(X1)|~v2_struct_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_310])])).
cnf(c_0_323, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_311])).
cnf(c_0_324, negated_conjecture, (esk3_0=esk8_0|esk4_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_312, c_0_188]), c_0_232])])])).
cnf(c_0_325, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_313, c_0_223]), c_0_48])])])).
cnf(c_0_326, plain, (m2_finseq_1(k8_matrix_1(X1,X2,X3),X1)|v1_xboole_0(X1)|~m1_finseq_2(k4_finseq_2(k1_matrix_1(X2),X1),X1)|~m1_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_314, c_0_315])])).
cnf(c_0_327, plain, (m1_finseq_2(k4_finseq_2(X1,X2),X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_316])).
cnf(c_0_328, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,esk6_0,X4,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,esk6_0,X4)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5)))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(esk6_0,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5)))|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X5)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_317, c_0_223]), c_0_147])])])).
cnf(c_0_329, negated_conjecture, (v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_318]), c_0_48])])])).
cnf(c_0_330, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|v2_struct_0(X1)|~epred1_2(X3,X2)|~r2_hidden(X7,k2_finseq_1(k1_matrix_1(X4)))|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_319, c_0_320])])).
cnf(c_0_331, plain, (epred1_2(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_321])).
cnf(c_0_332, plain, (v6_struct_0(X1)|~l6_algstr_0(X1)|~v2_struct_0(X1)), inference(csr,[status(thm)],[c_0_322, c_0_100])).
cnf(c_0_333, plain, (m1_subset_1(X1,k5_numbers)|v1_xboole_0(k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_226]), c_0_238])])).
cnf(c_0_334, plain, (m2_subset_1(esk11_2(X1,X2),k1_numbers,k5_numbers)|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_124])).
cnf(c_0_335, negated_conjecture, (esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_324, c_0_249]), c_0_232])])])).
cnf(c_0_336, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v7_ordinal1(X2)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_funct_1(X3)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_relat_1(X3)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_177])])).
cnf(c_0_337, plain, (m2_finseq_1(k8_matrix_1(X1,X2,X3),X1)|v1_xboole_0(X1)|~m1_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(k1_matrix_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_326, c_0_327])])).
cnf(c_0_338, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,esk6_0,X4,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,esk6_0,X4)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5)))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(esk6_0,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X5)), inference(csr,[status(thm)],[c_0_328, c_0_329])).
cnf(c_0_339, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|~r2_hidden(X7,k2_finseq_1(k1_matrix_1(X4)))|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_330, c_0_331]), c_0_332])).
cnf(c_0_340, plain, (r1_xxreal_0(esk11_2(X1,X2),X2)|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_124])).
cnf(c_0_341, plain, (m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(sr,[status(thm)],[c_0_333, c_0_252])).
cnf(c_0_342, plain, (m2_subset_1(esk11_2(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2))), inference(evalgc,[status(thm)],[c_0_334])).
cnf(c_0_343, negated_conjecture, (esk3_0=esk8_0|esk4_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_335, c_0_187]), c_0_232])])])).
cnf(c_0_344, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))!=k3_finseq_1(X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_336, c_0_185]), c_0_248]), c_0_249]), c_0_250])).
cnf(c_0_345, plain, (m2_finseq_1(k8_matrix_1(X1,X2,X3),X1)|v1_xboole_0(X1)|~v1_matrix_1(X2)|~v7_ordinal1(k1_matrix_1(X2))|~v7_ordinal1(X3)|~m2_finseq_1(X2,k3_finseq_2(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_337, c_0_98])])).
cnf(c_0_346, negated_conjecture, (k1_funct_1(X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(X4,esk6_0,X5,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|v1_xboole_0(X4)|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|X1!=k6_matrix_1(X4,esk6_0,X5)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~m2_finseq_1(esk6_0,k3_finseq_2(X4))|~m2_finseq_1(X1,X4)|~v7_ordinal1(X5)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v7_ordinal1(X2)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_funct_1(X3)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_relat_1(X3)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_338, c_0_177])])).
cnf(c_0_347, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(X7,k1_matrix_1(X4))|~v7_ordinal1(k1_matrix_1(X4))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_339, c_0_165])])).
cnf(c_0_348, plain, (r1_xxreal_0(esk11_2(X1,X2),X2)|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2))), inference(evalgc,[status(thm)],[c_0_340])).
cnf(c_0_349, plain, (m1_subset_1(esk11_2(X1,X2),k5_numbers)|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_341, c_0_342])])).
cnf(c_0_350, plain, (r1_xxreal_0(np__1,esk11_2(X1,X2))|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_124])).
cnf(c_0_351, plain, (X1=esk11_2(X1,X2)|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_124])).
cnf(c_0_352, negated_conjecture, (esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_343, c_0_248]), c_0_232])])])).
cnf(c_0_353, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))!=k3_finseq_1(X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_214])])).
cnf(c_0_354, negated_conjecture, (m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_345, c_0_207]), c_0_208]), c_0_154]), c_0_166])]), c_0_136])])).
cnf(c_0_355, negated_conjecture, (k1_funct_1(X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(X4,esk6_0,X5,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|v1_xboole_0(X4)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))!=k3_finseq_1(X3)|X1!=k6_matrix_1(X4,esk6_0,X5)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m2_finseq_1(esk6_0,k3_finseq_2(X4))|~m2_finseq_1(X1,X4)|~v7_ordinal1(X5)|~v7_ordinal1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_346, c_0_185]), c_0_248]), c_0_249]), c_0_250])).
cnf(c_0_356, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,esk11_2(X7,k1_matrix_1(X4)))=k3_matrix_1(u1_struct_0(X1),X4,X5,esk11_2(X7,k1_matrix_1(X4)))|v6_struct_0(X1)|~m2_subset_1(esk11_2(X7,k1_matrix_1(X4)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk11_2(X7,k1_matrix_1(X4)))|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X4))|~v7_ordinal1(X5)|~v7_ordinal1(esk11_2(X7,k1_matrix_1(X4)))|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X7,a_1_0_finseq_1(k1_matrix_1(X4)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_347, c_0_348])])).
cnf(c_0_357, plain, (v7_ordinal1(esk11_2(X1,X2))|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_349])])).
cnf(c_0_358, plain, (r1_xxreal_0(np__1,esk11_2(X1,X2))|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2))), inference(evalgc,[status(thm)],[c_0_350])).
cnf(c_0_359, plain, (esk11_2(X1,X2)=X1|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2))), inference(evalgc,[status(thm)],[c_0_351])).
cnf(c_0_360, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|r1_xxreal_0(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_matrix_1(X3))|k1_matrix_1(X3)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_318])])).
cnf(c_0_361, negated_conjecture, (esk3_0=esk8_0|esk4_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_352, c_0_186]), c_0_232])])])).
cnf(c_0_362, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))!=k3_finseq_1(X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(csr,[status(thm)],[c_0_353, c_0_354])).
cnf(c_0_363, negated_conjecture, (k1_funct_1(X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))!=k3_finseq_1(X3)|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(X4)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_355, c_0_163]), c_0_156])]), c_0_136])])).
cnf(c_0_364, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,esk11_2(X7,k1_matrix_1(X4)))=k3_matrix_1(u1_struct_0(X1),X4,X5,esk11_2(X7,k1_matrix_1(X4)))|v6_struct_0(X1)|~r2_hidden(X7,a_1_0_finseq_1(k1_matrix_1(X4)))|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X4))|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_356, c_0_357]), c_0_358]), c_0_342])).
cnf(c_0_365, plain, (esk11_2(X1,X2)=X1|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_359, c_0_159])])).
cnf(c_0_366, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|r1_xxreal_0(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_matrix_1(X3))|k1_matrix_1(X3)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(X2)|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_360, c_0_249]), c_0_250])).
cnf(c_0_367, negated_conjecture, (esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_331]), c_0_49]), c_0_50])])])).
cnf(c_0_368, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))!=k1_funct_1(X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_362])])).
cnf(c_0_369, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))!=k3_finseq_1(X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_363])])).
cnf(c_0_370, negated_conjecture, (m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_345, c_0_156]), c_0_147]), c_0_166])]), c_0_136])])).
cnf(c_0_371, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|~r2_hidden(X7,a_1_0_finseq_1(k1_matrix_1(X4)))|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X4))|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(X7,k1_matrix_1(X4))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_364, c_0_365])])).
cnf(c_0_372, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|r1_xxreal_0(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_matrix_1(esk6_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_154]), c_0_103])])])).
cnf(c_0_373, negated_conjecture, (esk3_0=esk8_0|esk4_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_234]), c_0_232])])])).
cnf(c_0_374, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))!=k1_funct_1(X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(X2)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_368, c_0_248]), c_0_249]), c_0_250])).
cnf(c_0_375, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))!=k3_finseq_1(X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(csr,[status(thm)],[c_0_369, c_0_370])).
cnf(c_0_376, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k1_matrix_1(X4))|~r1_xxreal_0(np__1,X7)|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X4))|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)), inference(csr,[status(thm)],[c_0_371, c_0_159])).
cnf(c_0_377, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|r1_xxreal_0(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_matrix_1(esk6_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_372, c_0_248])).
cnf(c_0_378, negated_conjecture, (esk4_0=esk8_0|esk3_0=esk8_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_373]), c_0_232])])])).
cnf(c_0_379, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_375])])).
cnf(c_0_380, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X5,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6),X7))=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6),X7))|v6_struct_0(X1)|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6)=X7|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6),X7),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6),X7))|~r2_hidden(X5,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|k1_matrix_1(esk6_0)!=k3_finseq_1(X7)|~v1_finseq_1(X7)|~v1_funct_1(X7)|~v1_relat_1(X7)|~v7_ordinal1(X6)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_376, c_0_377]), c_0_166])])])).
cnf(c_0_381, negated_conjecture, (k1_matrix12(esk2_0,esk1_0,esk5_0,esk6_0,esk3_0,esk8_0)=esk7_0|esk3_0=esk8_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_378])])).
cnf(c_0_382, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_379, c_0_186]), c_0_187]), c_0_188]), c_0_234])).
cnf(c_0_383, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_380, c_0_381]), c_0_231]), c_0_48]), c_0_108]), c_0_109]), c_0_110]), c_0_111]), c_0_112]), c_0_113]), c_0_114]), c_0_115]), c_0_116]), c_0_117]), c_0_203]), c_0_232]), c_0_49]), c_0_50])]), c_0_118])])).
cnf(c_0_384, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_382, c_0_255])])).
cnf(c_0_385, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_255])])).
cnf(c_0_386, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k5_numbers)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_384, c_0_259])])).
cnf(c_0_387, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k5_numbers)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_385, c_0_259])])).
cnf(c_0_388, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k5_numbers)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_386, c_0_248]), c_0_186]), c_0_249]), c_0_187]), c_0_250]), c_0_188]), c_0_234])).
cnf(c_0_389, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(X2)|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k5_numbers)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_387, c_0_248]), c_0_249]), c_0_250])).
cnf(c_0_390, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)),k5_numbers)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_388, c_0_389]), c_0_232])])])).
cnf(c_0_391, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)),k5_numbers)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_390, c_0_186]), c_0_187]), c_0_188]), c_0_168])).
cnf(c_0_392, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)),k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_391]), c_0_203])])])).
cnf(c_0_393, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_392, c_0_279])])).
cnf(c_0_394, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_185])])).
cnf(c_0_395, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_394, c_0_250]), c_0_232])])])).
cnf(c_0_396, plain, (k3_matrix_1(u1_struct_0(X3),X4,X1,X5)=k3_matrix_1(u1_struct_0(X3),X6,X7,X5)|v2_struct_0(X3)|v6_struct_0(X3)|X1!=X2|~r2_hidden(X1,k4_finseq_1(X6))|~r2_hidden(X5,k2_finseq_1(k1_matrix_1(X6)))|~v7_ordinal1(X5)|~v7_ordinal1(X1)|X4!=k1_matrix12(X8,X9,X3,X6,X2,X7)|~m1_matrix_1(X4,u1_struct_0(X3),X8,X9)|~v7_ordinal1(X7)|~v7_ordinal1(X2)|~m1_matrix_1(X6,u1_struct_0(X3),X8,X9)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~epred1_2(X9,X8)), inference(split_conjunct,[status(thm)],[c_0_183])).
cnf(c_0_397, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_188]), c_0_203])])])).
cnf(c_0_398, plain, (k3_matrix_1(u1_struct_0(X3),X4,X1,X5)=k3_matrix_1(u1_struct_0(X3),X6,X7,X5)|v2_struct_0(X3)|v6_struct_0(X3)|X1!=X2|X4!=k1_matrix12(X8,X9,X3,X6,X2,X7)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~epred1_2(X9,X8)|~r2_hidden(X1,k4_finseq_1(X6))|~r2_hidden(X5,k2_finseq_1(k1_matrix_1(X6)))|~m1_matrix_1(X6,u1_struct_0(X3),X8,X9)|~m1_matrix_1(X4,u1_struct_0(X3),X8,X9)), inference(evalgc,[status(thm)],[c_0_396])).
cnf(c_0_399, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_397, c_0_249]), c_0_232])])])).
cnf(c_0_400, plain, (k3_matrix_1(u1_struct_0(X1),X2,X3,X4)=k3_matrix_1(u1_struct_0(X1),X5,X6,X4)|v6_struct_0(X1)|v2_struct_0(X1)|X2!=k1_matrix12(X7,X8,X1,X5,X3,X6)|~epred1_2(X8,X7)|~r2_hidden(X4,k2_finseq_1(k1_matrix_1(X5)))|~r2_hidden(X3,k4_finseq_1(X5))|~m1_matrix_1(X5,u1_struct_0(X1),X7,X8)|~m1_matrix_1(X2,u1_struct_0(X1),X7,X8)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_398])])).
cnf(c_0_401, negated_conjecture, (r2_hidden(esk3_0,k4_finseq_1(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_402, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_399, c_0_187]), c_0_203])])])).
cnf(c_0_403, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,X7)=k3_matrix_1(u1_struct_0(X1),X4,X6,X7)|v6_struct_0(X1)|v2_struct_0(X1)|~epred1_2(X3,X2)|~r2_hidden(X7,k2_finseq_1(k1_matrix_1(X4)))|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~m1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6),u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X7)|~v7_ordinal1(X5)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_400])])).
cnf(c_0_404, negated_conjecture, (r2_hidden(esk3_0,k4_finseq_1(esk6_0))), inference(evalgc,[status(thm)],[c_0_401])).
cnf(c_0_405, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_248]), c_0_232])])])).
cnf(c_0_406, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,X1)|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(esk6_0)))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_403, c_0_201]), c_0_404]), c_0_103]), c_0_48]), c_0_108]), c_0_109]), c_0_110]), c_0_111]), c_0_112]), c_0_113]), c_0_114]), c_0_115]), c_0_116]), c_0_117]), c_0_202]), c_0_203])]), c_0_118]), c_0_119])])).
cnf(c_0_407, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_405, c_0_186]), c_0_203])])])).
cnf(c_0_408, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,X1)|~epred1_2(esk1_0,esk2_0)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,k1_matrix_1(esk6_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_406, c_0_165]), c_0_166])])])).
cnf(c_0_409, negated_conjecture, (esk8_0=esk3_0|esk8_0!=esk4_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_410, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_407, c_0_234]), c_0_232])])])).
cnf(c_0_411, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,X1)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_408, c_0_172]), c_0_48])])])).
cnf(c_0_412, negated_conjecture, (esk3_0=esk8_0|esk4_0!=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)), inference(evalgc,[status(thm)],[c_0_409])).
cnf(c_0_413, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_410]), c_0_203])])])).
cnf(c_0_414, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_411, c_0_177])])).
cnf(c_0_415, negated_conjecture, (esk3_0=esk8_0|esk4_0!=esk8_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_412, c_0_413])])).
cnf(c_0_416, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_414, c_0_185]), c_0_186]), c_0_187]), c_0_188])).
cnf(c_0_417, negated_conjecture, (esk3_0=esk8_0), inference(csr,[status(thm)],[c_0_415, c_0_378])).
cnf(c_0_418, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_416, c_0_417])])).
cnf(c_0_419, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_418, c_0_255])])).
cnf(c_0_420, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers)|~v7_ordinal1(X1)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_419, c_0_259])])).
cnf(c_0_421, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_420, c_0_186]), c_0_187]), c_0_188])).
cnf(c_0_422, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_421]), c_0_232])])])).
cnf(c_0_423, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_422]), c_0_202])])])).
cnf(c_0_424, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_423, c_0_279])])).
cnf(c_0_425, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_424, c_0_185])])).
cnf(c_0_426, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_425, c_0_250]), c_0_232])])])).
cnf(c_0_427, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~epred1_2(esk1_0,esk2_0)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_188]), c_0_202])])])).
cnf(c_0_428, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~epred1_2(esk1_0,esk2_0)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_427, c_0_249]), c_0_232])])])).
cnf(c_0_429, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~epred1_2(esk1_0,esk2_0)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_428, c_0_187]), c_0_202])])])).
cnf(c_0_430, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~epred1_2(esk1_0,esk2_0)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_429, c_0_248]), c_0_232])])])).
cnf(c_0_431, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~epred1_2(esk1_0,esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_430, c_0_186]), c_0_202])])])).
cnf(c_0_432, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_433, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_431, c_0_331]), c_0_49]), c_0_50])])])).
cnf(c_0_434, negated_conjecture, (esk8_0=esk4_0|esk8_0!=esk3_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_435, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)), inference(evalgc,[status(thm)],[c_0_432])).
cnf(c_0_436, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_433, c_0_234]), c_0_232])])])).
cnf(c_0_437, negated_conjecture, (esk4_0=esk8_0|esk3_0!=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)), inference(evalgc,[status(thm)],[c_0_434])).
cnf(c_0_438, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_435, c_0_417])])])).
cnf(c_0_439, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_436]), c_0_202])])])).
cnf(c_0_440, negated_conjecture, (esk4_0=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_437, c_0_417])])])).
cnf(c_0_441, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_438, c_0_439]), c_0_439])])])).
cnf(c_0_442, negated_conjecture, (esk4_0=esk8_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_440, c_0_439])])])).
cnf(c_0_443, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_441, c_0_442])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 444
# Proof object clause steps            : 369
# Proof object formula steps           : 75
# Proof object conjectures             : 225
# Proof object clause conjectures      : 222
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 76
# Proof object initial formulas used   : 36
# Proof object generating inferences   : 158
# Proof object simplifying inferences  : 378
# Training examples: 249 positive, 1713 negative
# Training: Positive examples begin
cnf(c_0_78, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_49, negated_conjecture, (v7_ordinal1(esk1_0))).# trainpos
cnf(c_0_50, negated_conjecture, (v7_ordinal1(esk2_0))).# trainpos
cnf(c_0_203, negated_conjecture, (v7_ordinal1(esk3_0))).# trainpos
cnf(c_0_202, negated_conjecture, (v7_ordinal1(esk4_0))).# trainpos
cnf(c_0_232, negated_conjecture, (v7_ordinal1(esk8_0))).# trainpos
cnf(c_0_117, negated_conjecture, (v13_algstr_0(esk5_0))).# trainpos
cnf(c_0_116, negated_conjecture, (v33_algstr_0(esk5_0))).# trainpos
cnf(c_0_115, negated_conjecture, (v3_group_1(esk5_0))).# trainpos
cnf(c_0_114, negated_conjecture, (v5_group_1(esk5_0))).# trainpos
cnf(c_0_113, negated_conjecture, (v2_rlvect_1(esk5_0))).# trainpos
cnf(c_0_112, negated_conjecture, (v3_rlvect_1(esk5_0))).# trainpos
cnf(c_0_119, negated_conjecture, (~v2_struct_0(esk5_0))).# trainpos
cnf(c_0_118, negated_conjecture, (~v6_struct_0(esk5_0))).# trainpos
cnf(c_0_111, negated_conjecture, (v4_rlvect_1(esk5_0))).# trainpos
cnf(c_0_110, negated_conjecture, (v4_vectsp_1(esk5_0))).# trainpos
cnf(c_0_228, plain, (~v1_xboole_0(k4_ordinal1))).# trainpos
cnf(c_0_109, negated_conjecture, (v5_vectsp_1(esk5_0))).# trainpos
cnf(c_0_108, negated_conjecture, (l6_algstr_0(esk5_0))).# trainpos
cnf(c_0_76, plain, (l2_algstr_0(X1)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_404, negated_conjecture, (r2_hidden(esk3_0,k4_finseq_1(esk6_0)))).# trainpos
cnf(c_0_231, negated_conjecture, (r2_hidden(esk8_0,k4_finseq_1(esk6_0)))).# trainpos
cnf(c_0_48, negated_conjecture, (m1_matrix_1(esk6_0,u1_struct_0(esk5_0),esk2_0,esk1_0))).# trainpos
cnf(c_0_103, negated_conjecture, (m1_matrix_1(esk7_0,u1_struct_0(esk5_0),esk2_0,esk1_0))).# trainpos
cnf(c_0_226, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_201, negated_conjecture, (k1_matrix12(esk2_0,esk1_0,esk5_0,esk6_0,esk3_0,esk4_0)=esk7_0)).# trainpos
cnf(c_0_412, negated_conjecture, (esk3_0=esk8_0|esk4_0!=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))).# trainpos
cnf(c_0_437, negated_conjecture, (esk4_0=esk8_0|esk3_0!=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))).# trainpos
cnf(c_0_289, plain, (l5_algstr_0(X1)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_288, plain, (l4_struct_0(X1)|~l5_algstr_0(X1))).# trainpos
cnf(c_0_245, plain, (m1_subset_1(esk12_1(X1),X1))).# trainpos
cnf(c_0_63, plain, (l2_struct_0(X1)|~l2_algstr_0(X1))).# trainpos
cnf(c_0_62, plain, (l1_struct_0(X1)|~l2_struct_0(X1))).# trainpos
cnf(c_0_273, negated_conjecture, (esk3_0=esk8_0|esk4_0=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))).# trainpos
cnf(c_0_435, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))).# trainpos
cnf(c_0_151, plain, (k1_finseq_1(X1)=k2_finseq_1(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_77, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).# trainpos
cnf(c_0_310, plain, (v7_struct_0(X1)|~v2_struct_0(X1)|~l1_struct_0(X1))).# trainpos
cnf(c_0_152, plain, (k1_finseq_1(X1)=a_1_0_finseq_1(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_271, plain, (r2_hidden(X1,k4_ordinal1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_104, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2))).# trainpos
cnf(c_0_131, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_132, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_274, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_98, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_298, plain, (v6_struct_0(X1)|~l4_struct_0(X1)|~v7_struct_0(X1))).# trainpos
cnf(c_0_204, plain, (X1=X2|k3_finseq_1(X1)!=k3_finseq_1(X2)|k1_funct_1(X1,esk10_2(X1,X2))!=k1_funct_1(X2,esk10_2(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_99, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).# trainpos
cnf(c_0_225, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_327, plain, (m1_finseq_2(k4_finseq_2(X1,X2),X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_331, plain, (epred1_2(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_237, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_359, plain, (esk11_2(X1,X2)=X1|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2)))).# trainpos
cnf(c_0_314, plain, (m2_finseq_1(X3,X2)|~m1_finseq_2(X1,X2)|~m2_finseq_2(X3,X2,X1))).# trainpos
cnf(c_0_102, plain, (m1_subset_1(k1_matrix_1(X1),k5_numbers)|~v1_matrix_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_358, plain, (r1_xxreal_0(np__1,esk11_2(X1,X2))|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2)))).# trainpos
cnf(c_0_348, plain, (r1_xxreal_0(esk11_2(X1,X2),X2)|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2)))).# trainpos
cnf(c_0_342, plain, (m2_subset_1(esk11_2(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2)))).# trainpos
cnf(c_0_58, plain, (v1_xboole_0(X2)|v1_matrix_1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4))).# trainpos
cnf(c_0_243, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_323, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_47, plain, (v1_xboole_0(X2)|m2_finseq_1(X1,k3_finseq_2(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4))).# trainpos
cnf(c_0_97, plain, (v1_xboole_0(X1)|m2_finseq_1(k6_matrix_1(X1,X2,X3),X1)|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m1_finseq_1(X2,k3_finseq_2(X1)))).# trainpos
cnf(c_0_148, plain, (k6_matrix_1(X1,X2,X3)=k8_matrix_1(X1,X2,X3)|v1_xboole_0(X1)|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m1_finseq_1(X2,k3_finseq_2(X1)))).# trainpos
cnf(c_0_150, plain, (r2_hidden(X2,a_1_0_finseq_1(X3))|X1!=X2|~v7_ordinal1(X3)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_185, plain, (X1=X2|v7_ordinal1(esk10_2(X1,X2))|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_56, plain, (k1_matrix_1(X2)=k3_finseq_1(X1)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,X2,X4)|~v7_ordinal1(X4)|~v1_matrix_1(X2)|~m2_finseq_1(X1,X3)|~m2_finseq_1(X2,k3_finseq_2(X3)))).# trainpos
cnf(c_0_259, plain, (X1=X2|r1_xxreal_0(np__1,esk10_2(X1,X2))|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_177, plain, (X1=X2|r1_xxreal_0(esk10_2(X1,X2),k3_finseq_1(X1))|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_315, plain, (v1_xboole_0(X1)|m2_finseq_2(k8_matrix_1(X1,X2,X3),X1,k4_finseq_2(k1_matrix_1(X2),X1))|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m1_finseq_1(X2,k3_finseq_2(X1)))).# trainpos
cnf(c_0_107, plain, (k1_matrix_1(X4)=k1_matrix_1(X5)|v2_struct_0(X3)|v6_struct_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~m1_matrix_1(X5,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2))).# trainpos
cnf(c_0_158, plain, (k1_funct_1(X3,X1)=k3_matrix_1(X4,X2,X5,X1)|v1_xboole_0(X4)|X3!=k6_matrix_1(X4,X2,X5)|~v7_ordinal1(X5)|~v7_ordinal1(X1)|~v1_matrix_1(X2)|~m2_finseq_1(X3,X4)|~m2_finseq_1(X2,k3_finseq_2(X4))|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2))))).# trainpos
cnf(c_0_320, plain, (v2_struct_0(X3)|v6_struct_0(X3)|m1_matrix_1(k1_matrix12(X1,X2,X3,X4,X5,X6),u1_struct_0(X3),X1,X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2))).# trainpos
cnf(c_0_295, plain, (k3_matrix_1(u1_struct_0(X3),X4,X1,X5)=k3_matrix_1(u1_struct_0(X3),X6,X7,X5)|v2_struct_0(X3)|v6_struct_0(X3)|X1!=X2|X4!=k1_matrix12(X8,X9,X3,X6,X7,X2)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~epred1_2(X9,X8)|~r2_hidden(X1,k4_finseq_1(X6))|~r2_hidden(X5,k2_finseq_1(k1_matrix_1(X6)))|~m1_matrix_1(X6,u1_struct_0(X3),X8,X9)|~m1_matrix_1(X4,u1_struct_0(X3),X8,X9))).# trainpos
cnf(c_0_398, plain, (k3_matrix_1(u1_struct_0(X3),X4,X1,X5)=k3_matrix_1(u1_struct_0(X3),X6,X7,X5)|v2_struct_0(X3)|v6_struct_0(X3)|X1!=X2|X4!=k1_matrix12(X8,X9,X3,X6,X2,X7)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~epred1_2(X9,X8)|~r2_hidden(X1,k4_finseq_1(X6))|~r2_hidden(X5,k2_finseq_1(k1_matrix_1(X6)))|~m1_matrix_1(X6,u1_struct_0(X3),X8,X9)|~m1_matrix_1(X4,u1_struct_0(X3),X8,X9))).# trainpos
cnf(c_0_193, plain, (X1=X3|X1=X2|k3_matrix_1(u1_struct_0(X4),X5,X1,X6)=k3_matrix_1(u1_struct_0(X4),X7,X1,X6)|v2_struct_0(X4)|v6_struct_0(X4)|X5!=k1_matrix12(X8,X9,X4,X7,X2,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v13_algstr_0(X4)|~v33_algstr_0(X4)|~v3_group_1(X4)|~v5_group_1(X4)|~v2_rlvect_1(X4)|~v3_rlvect_1(X4)|~v4_rlvect_1(X4)|~v4_vectsp_1(X4)|~v5_vectsp_1(X4)|~l6_algstr_0(X4)|~epred1_2(X9,X8)|~r2_hidden(X1,k4_finseq_1(X7))|~r2_hidden(X6,k2_finseq_1(k1_matrix_1(X7)))|~m1_matrix_1(X7,u1_struct_0(X4),X8,X9)|~m1_matrix_1(X5,u1_struct_0(X4),X8,X9))).# trainpos
cnf(c_0_306, plain, (k3_matrix_1(u1_struct_0(X1),X2,X3,X4)=k3_matrix_1(u1_struct_0(X1),X5,X6,X4)|v6_struct_0(X1)|v2_struct_0(X1)|X2!=k1_matrix12(X7,X8,X1,X5,X6,X3)|~epred1_2(X8,X7)|~r2_hidden(X4,k2_finseq_1(k1_matrix_1(X5)))|~r2_hidden(X3,k4_finseq_1(X5))|~m1_matrix_1(X5,u1_struct_0(X1),X7,X8)|~m1_matrix_1(X2,u1_struct_0(X1),X7,X8)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_299, plain, (l4_struct_0(X1)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_75, plain, (l1_struct_0(X1)|~l2_algstr_0(X1))).# trainpos
cnf(c_0_157, plain, (k6_matrix_1(X1,X2,X3)=k8_matrix_1(X1,X2,X3)|v1_xboole_0(X1)|~v1_matrix_1(X2)|~v7_ordinal1(X3)|~m2_finseq_1(X2,k3_finseq_2(X1)))).# trainpos
cnf(c_0_129, plain, (v1_relat_1(X1)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_400, plain, (k3_matrix_1(u1_struct_0(X1),X2,X3,X4)=k3_matrix_1(u1_struct_0(X1),X5,X6,X4)|v6_struct_0(X1)|v2_struct_0(X1)|X2!=k1_matrix12(X7,X8,X1,X5,X3,X6)|~epred1_2(X8,X7)|~r2_hidden(X4,k2_finseq_1(k1_matrix_1(X5)))|~r2_hidden(X3,k4_finseq_1(X5))|~m1_matrix_1(X5,u1_struct_0(X1),X7,X8)|~m1_matrix_1(X2,u1_struct_0(X1),X7,X8)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_200, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8)=k3_matrix_1(u1_struct_0(X1),X4,X7,X8)|X7=X5|X7=X6|v6_struct_0(X1)|v2_struct_0(X1)|~epred1_2(X3,X2)|~r2_hidden(X8,k2_finseq_1(k1_matrix_1(X4)))|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~m1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6),u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X8)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7))).# trainpos
cnf(c_0_133, negated_conjecture, (k1_matrix_1(X1)=k1_matrix_1(esk6_0)|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).# trainpos
cnf(c_0_134, negated_conjecture, (k1_matrix_1(X1)=k1_matrix_1(esk7_0)|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).# trainpos
cnf(c_0_159, plain, (r2_hidden(X1,a_1_0_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_128, negated_conjecture, (v1_matrix_1(esk7_0)|v1_xboole_0(u1_struct_0(esk5_0)))).# trainpos
cnf(c_0_71, negated_conjecture, (v1_matrix_1(esk6_0)|v1_xboole_0(u1_struct_0(esk5_0)))).# trainpos
cnf(c_0_121, plain, (m2_finseq_1(k6_matrix_1(X1,X2,X3),X1)|v1_xboole_0(X1)|~v1_matrix_1(X2)|~v7_ordinal1(X3)|~m2_finseq_1(X2,k3_finseq_2(X1)))).# trainpos
cnf(c_0_160, plain, (a_1_0_finseq_1(X1)=k2_finseq_1(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_279, plain, (m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_236, plain, (~v1_xboole_0(k1_numbers)|~r2_hidden(X1,k5_numbers))).# trainpos
cnf(c_0_130, negated_conjecture, (m2_finseq_1(esk7_0,k3_finseq_2(u1_struct_0(esk5_0)))|v1_xboole_0(u1_struct_0(esk5_0)))).# trainpos
cnf(c_0_57, negated_conjecture, (m2_finseq_1(esk6_0,k3_finseq_2(u1_struct_0(esk5_0)))|v1_xboole_0(u1_struct_0(esk5_0)))).# trainpos
cnf(c_0_326, plain, (m2_finseq_1(k8_matrix_1(X1,X2,X3),X1)|v1_xboole_0(X1)|~m1_finseq_2(k4_finseq_2(k1_matrix_1(X2),X1),X1)|~m1_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_127, plain, (v7_ordinal1(k1_matrix_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1))).# trainpos
cnf(c_0_143, negated_conjecture, (v1_funct_1(esk7_0)|v1_xboole_0(u1_struct_0(esk5_0)))).# trainpos
cnf(c_0_144, negated_conjecture, (v1_finseq_1(esk7_0)|v1_xboole_0(u1_struct_0(esk5_0)))).# trainpos
cnf(c_0_100, plain, (l1_struct_0(X1)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_403, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,X7)=k3_matrix_1(u1_struct_0(X1),X4,X6,X7)|v6_struct_0(X1)|v2_struct_0(X1)|~epred1_2(X3,X2)|~r2_hidden(X7,k2_finseq_1(k1_matrix_1(X4)))|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~m1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6),u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X7)|~v7_ordinal1(X5))).# trainpos
cnf(c_0_319, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|v2_struct_0(X1)|~epred1_2(X3,X2)|~r2_hidden(X7,k2_finseq_1(k1_matrix_1(X4)))|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~m1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6),u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X6))).# trainpos
cnf(c_0_309, plain, (v6_struct_0(X1)|~v7_struct_0(X1)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_142, negated_conjecture, (v1_relat_1(esk7_0)|v1_xboole_0(u1_struct_0(esk5_0)))).# trainpos
cnf(c_0_145, negated_conjecture, (k1_matrix_1(esk7_0)=k1_matrix_1(esk6_0)|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).# trainpos
cnf(c_0_167, negated_conjecture, (k1_matrix_1(X1)=k1_matrix_1(X2)|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0))).# trainpos
cnf(c_0_164, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,X4,X5,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,X4,X5)|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v1_matrix_1(X4)|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(esk6_0)))|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0))).# trainpos
cnf(c_0_70, negated_conjecture, (k1_matrix_1(esk6_0)=k3_finseq_1(X1)|v1_xboole_0(u1_struct_0(esk5_0))|X1!=k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v1_matrix_1(esk6_0)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_174, negated_conjecture, (k1_matrix_1(esk7_0)=k3_finseq_1(X1)|v1_xboole_0(u1_struct_0(esk5_0))|X1!=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v1_matrix_1(esk7_0)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_141, negated_conjecture, (v7_ordinal1(k1_matrix_1(esk7_0))|v1_xboole_0(u1_struct_0(esk5_0))|~v1_finseq_1(esk7_0)|~v1_funct_1(esk7_0)|~v1_relat_1(esk7_0))).# trainpos
cnf(c_0_122, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l6_algstr_0(X1))).# trainpos
cnf(c_0_322, plain, (v6_struct_0(X1)|~l6_algstr_0(X1)|~l1_struct_0(X1)|~v2_struct_0(X1))).# trainpos
cnf(c_0_330, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|v2_struct_0(X1)|~epred1_2(X3,X2)|~r2_hidden(X7,k2_finseq_1(k1_matrix_1(X4)))|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_189, negated_conjecture, (k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))=k1_matrix_1(esk7_0)|v1_xboole_0(u1_struct_0(esk5_0))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_154, negated_conjecture, (k1_matrix_1(esk7_0)=k1_matrix_1(esk6_0))).# trainpos
cnf(c_0_120, negated_conjecture, (k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))=k1_matrix_1(esk6_0)|v1_xboole_0(u1_struct_0(esk5_0))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_406, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,X1)|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(esk6_0)))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_244, plain, (~v1_xboole_0(k1_numbers)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_136, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk5_0)))).# trainpos
cnf(c_0_192, negated_conjecture, (k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))=k1_matrix_1(esk6_0)|v1_xboole_0(u1_struct_0(esk5_0))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_252, plain, (~v1_xboole_0(k1_numbers))).# trainpos
cnf(c_0_135, negated_conjecture, (k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))=k1_matrix_1(esk6_0)|v1_xboole_0(u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~m2_finseq_1(esk6_0,k3_finseq_2(u1_struct_0(esk5_0)))|~v1_matrix_1(esk6_0))).# trainpos
cnf(c_0_161, negated_conjecture, (v1_xboole_0(u1_struct_0(esk5_0))|v7_ordinal1(k1_matrix_1(esk6_0)))).# trainpos
cnf(c_0_147, negated_conjecture, (v1_matrix_1(esk6_0))).# trainpos
cnf(c_0_208, negated_conjecture, (v1_matrix_1(esk7_0))).# trainpos
cnf(c_0_156, negated_conjecture, (m2_finseq_1(esk6_0,k3_finseq_2(u1_struct_0(esk5_0))))).# trainpos
cnf(c_0_207, negated_conjecture, (m2_finseq_1(esk7_0,k3_finseq_2(u1_struct_0(esk5_0))))).# trainpos
cnf(c_0_206, negated_conjecture, (k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))=k1_matrix_1(esk6_0)|~v7_ordinal1(X1)|~m2_finseq_1(esk7_0,k3_finseq_2(u1_struct_0(esk5_0)))|~v1_matrix_1(esk7_0))).# trainpos
cnf(c_0_214, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_163, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_155, negated_conjecture, (k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))=k1_matrix_1(esk6_0)|~m2_finseq_1(esk6_0,k3_finseq_2(u1_struct_0(esk5_0)))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_365, plain, (esk11_2(X1,X2)=X1|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_223, negated_conjecture, (k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))=k1_matrix_1(esk6_0)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_168, negated_conjecture, (k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))=k1_matrix_1(esk6_0)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_180, plain, (v1_funct_1(k6_matrix_1(X1,X2,X3))|v1_xboole_0(X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_181, plain, (v1_finseq_1(k6_matrix_1(X1,X2,X3))|v1_xboole_0(X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_179, plain, (v1_relat_1(k6_matrix_1(X1,X2,X3))|v1_xboole_0(X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_165, plain, (r2_hidden(X1,k2_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_318, negated_conjecture, (k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))=k1_matrix_1(X1)|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_234, negated_conjecture, (k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_172, negated_conjecture, (k1_matrix_1(X1)=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_250, negated_conjecture, (v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_249, negated_conjecture, (v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_337, plain, (m2_finseq_1(k8_matrix_1(X1,X2,X3),X1)|v1_xboole_0(X1)|~m1_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(k1_matrix_1(X2)))).# trainpos
cnf(c_0_188, negated_conjecture, (v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_360, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|r1_xxreal_0(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_matrix_1(X3))|k1_matrix_1(X3)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_345, plain, (m2_finseq_1(k8_matrix_1(X1,X2,X3),X1)|v1_xboole_0(X1)|~v1_matrix_1(X2)|~v7_ordinal1(k1_matrix_1(X2))|~v7_ordinal1(X3)|~m2_finseq_1(X2,k3_finseq_2(X1)))).# trainpos
cnf(c_0_248, negated_conjecture, (v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_354, negated_conjecture, (m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_370, negated_conjecture, (m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_187, negated_conjecture, (v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_372, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|r1_xxreal_0(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_matrix_1(esk6_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_347, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(X7,k1_matrix_1(X4))|~v7_ordinal1(k1_matrix_1(X4)))).# trainpos
cnf(c_0_186, negated_conjecture, (v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_329, negated_conjecture, (v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_171, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,X4,X5,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,X4,X5)|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v1_matrix_1(X4)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,k1_matrix_1(esk6_0)))).# trainpos
cnf(c_0_176, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,X4,X5,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,X4,X5)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6)))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v1_matrix_1(X4)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~v7_ordinal1(X6))).# trainpos
cnf(c_0_408, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,X1)|~epred1_2(esk1_0,esk2_0)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,k1_matrix_1(esk6_0)))).# trainpos
cnf(c_0_411, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,X1)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_184, negated_conjecture, (k1_funct_1(X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(X4,X5,X6,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|v1_xboole_0(X4)|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|X1!=k6_matrix_1(X4,X5,X6)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~m2_finseq_1(X5,k3_finseq_2(X4))|~m2_finseq_1(X1,X4)|~v1_matrix_1(X5)|~m1_matrix_1(X5,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X6)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v7_ordinal1(X2)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(X3)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(X3)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))).# trainpos
cnf(c_0_198, negated_conjecture, (k1_funct_1(k6_matrix_1(X1,X2,X3),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4),X5))=k3_matrix_1(X1,X2,X3,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4),X5))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4)=X5|v1_xboole_0(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4))!=k3_finseq_1(X5)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4),X5),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4),X5))|~v1_finseq_1(X5)|~v1_funct_1(X5)|~v1_relat_1(X5)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m2_finseq_1(k6_matrix_1(X1,X2,X3),X1)|~v1_matrix_1(X2)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X3)|~v7_ordinal1(X4))).# trainpos
cnf(c_0_414, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))).# trainpos
cnf(c_0_276, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,esk7_0,X4,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,esk7_0,X4)|~m2_finseq_1(esk7_0,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v1_matrix_1(esk7_0)|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(esk6_0)))|~v7_ordinal1(X4)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_211, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|X1=esk4_0|X1=esk3_0|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(esk6_0)))|~r2_hidden(X1,k4_finseq_1(esk6_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_242, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))!=k3_finseq_1(X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_219, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|X1=esk3_0|X1=esk4_0|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k4_finseq_1(esk6_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,k1_matrix_1(esk6_0)))).# trainpos
cnf(c_0_212, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(X2,X3,X4)|v1_xboole_0(X2)|k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)))!=k3_matrix_1(X2,X3,X4,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(X2,X3,X4))|~v1_finseq_1(k6_matrix_1(X2,X3,X4))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k6_matrix_1(X2,X3,X4))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k6_matrix_1(X2,X3,X4))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3)|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X4)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_233, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_247, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))).# trainpos
cnf(c_0_285, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,esk7_0,X4,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,esk7_0,X4)|~m2_finseq_1(esk7_0,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,k1_matrix_1(esk6_0)))).# trainpos
cnf(c_0_230, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|X1=esk4_0|X1=esk3_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(X3))|~r1_xxreal_0(np__1,X2)|~r2_hidden(X1,k4_finseq_1(esk6_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0))).# trainpos
cnf(c_0_240, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X1)|esk3_0=esk8_0|esk4_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_matrix_1(X2))|~r1_xxreal_0(np__1,X1)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_246, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X1)|esk4_0=esk8_0|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_253, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|esk3_0=esk8_0|esk4_0=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))).# trainpos
cnf(c_0_292, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,esk7_0,X4,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,esk7_0,X4)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(X5))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(esk7_0,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~m1_matrix_1(X5,u1_struct_0(esk5_0),esk2_0,esk1_0))).# trainpos
cnf(c_0_302, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X3,X2)|X1!=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(X4))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_313, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(X3))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_325, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_317, plain, (k1_funct_1(X1,X2)=k3_matrix_1(X3,X4,X5,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,X4,X5)|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v1_matrix_1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,k1_matrix_1(X4))|~v7_ordinal1(k1_matrix_1(X4)))).# trainpos
cnf(c_0_336, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v7_ordinal1(X2)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_funct_1(X3)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_relat_1(X3)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))).# trainpos
cnf(c_0_353, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))!=k3_finseq_1(X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_368, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))!=k1_funct_1(X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_251, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|v1_xboole_0(k1_numbers)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_258, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k5_numbers))).# trainpos
cnf(c_0_333, plain, (m1_subset_1(X1,k5_numbers)|v1_xboole_0(k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_261, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))).# trainpos
cnf(c_0_260, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers))).# trainpos
cnf(c_0_262, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|esk3_0=esk8_0|esk4_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers)|~v7_ordinal1(X1)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))).# trainpos
cnf(c_0_268, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk4_0=esk8_0|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).# trainpos
cnf(c_0_328, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,esk6_0,X4,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,esk6_0,X4)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5)))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(esk6_0,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5)))|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X5))).# trainpos
cnf(c_0_272, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|esk4_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers))).# trainpos
cnf(c_0_284, negated_conjecture, (esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))))).# trainpos
cnf(c_0_291, negated_conjecture, (esk3_0=esk8_0|esk4_0=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).# trainpos
cnf(c_0_312, negated_conjecture, (esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).# trainpos
cnf(c_0_324, negated_conjecture, (esk3_0=esk8_0|esk4_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).# trainpos
cnf(c_0_346, negated_conjecture, (k1_funct_1(X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(X4,esk6_0,X5,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|v1_xboole_0(X4)|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|X1!=k6_matrix_1(X4,esk6_0,X5)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~m2_finseq_1(esk6_0,k3_finseq_2(X4))|~m2_finseq_1(X1,X4)|~v7_ordinal1(X5)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v7_ordinal1(X2)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_funct_1(X3)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_relat_1(X3)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))).# trainpos
cnf(c_0_335, negated_conjecture, (esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).# trainpos
cnf(c_0_363, negated_conjecture, (k1_funct_1(X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))!=k3_finseq_1(X3)|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(X4)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_343, negated_conjecture, (esk3_0=esk8_0|esk4_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).# trainpos
cnf(c_0_369, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))!=k3_finseq_1(X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_379, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_352, negated_conjecture, (esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).# trainpos
cnf(c_0_361, negated_conjecture, (esk3_0=esk8_0|esk4_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0))).# trainpos
cnf(c_0_367, negated_conjecture, (esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).# trainpos
cnf(c_0_384, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k5_numbers))).# trainpos
cnf(c_0_373, negated_conjecture, (esk3_0=esk8_0|esk4_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_378, negated_conjecture, (esk4_0=esk8_0|esk3_0=esk8_0)).# trainpos
cnf(c_0_386, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k5_numbers)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))).# trainpos
cnf(c_0_349, plain, (m1_subset_1(esk11_2(X1,X2),k5_numbers)|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2))).# trainpos
cnf(c_0_381, negated_conjecture, (k1_matrix12(esk2_0,esk1_0,esk5_0,esk6_0,esk3_0,esk8_0)=esk7_0|esk3_0=esk8_0)).# trainpos
cnf(c_0_357, plain, (v7_ordinal1(esk11_2(X1,X2))|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2))).# trainpos
cnf(c_0_356, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,esk11_2(X7,k1_matrix_1(X4)))=k3_matrix_1(u1_struct_0(X1),X4,X5,esk11_2(X7,k1_matrix_1(X4)))|v6_struct_0(X1)|~m2_subset_1(esk11_2(X7,k1_matrix_1(X4)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk11_2(X7,k1_matrix_1(X4)))|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X4))|~v7_ordinal1(X5)|~v7_ordinal1(esk11_2(X7,k1_matrix_1(X4)))|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X7,a_1_0_finseq_1(k1_matrix_1(X4))))).# trainpos
cnf(c_0_371, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|~r2_hidden(X7,a_1_0_finseq_1(k1_matrix_1(X4)))|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X4))|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(X7,k1_matrix_1(X4)))).# trainpos
cnf(c_0_380, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X5,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6),X7))=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6),X7))|v6_struct_0(X1)|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6)=X7|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6),X7),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6),X7))|~r2_hidden(X5,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|k1_matrix_1(esk6_0)!=k3_finseq_1(X7)|~v1_finseq_1(X7)|~v1_funct_1(X7)|~v1_relat_1(X7)|~v7_ordinal1(X6))).# trainpos
cnf(c_0_383, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_385, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k5_numbers))).# trainpos
cnf(c_0_387, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k5_numbers)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))).# trainpos
cnf(c_0_390, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)),k5_numbers)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))).# trainpos
cnf(c_0_392, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)),k5_numbers))).# trainpos
cnf(c_0_393, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))))).# trainpos
cnf(c_0_394, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).# trainpos
cnf(c_0_395, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).# trainpos
cnf(c_0_397, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).# trainpos
cnf(c_0_399, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).# trainpos
cnf(c_0_402, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).# trainpos
cnf(c_0_405, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)))).# trainpos
cnf(c_0_407, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)))).# trainpos
cnf(c_0_410, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_413, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0)).# trainpos
cnf(c_0_415, negated_conjecture, (esk3_0=esk8_0|esk4_0!=esk8_0)).# trainpos
cnf(c_0_438, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))).# trainpos
cnf(c_0_440, negated_conjecture, (esk4_0=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))).# trainpos
cnf(c_0_418, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_419, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers))).# trainpos
cnf(c_0_420, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers)|~v7_ordinal1(X1)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))).# trainpos
cnf(c_0_422, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).# trainpos
cnf(c_0_423, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers))).# trainpos
cnf(c_0_424, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))))).# trainpos
cnf(c_0_425, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)))).# trainpos
cnf(c_0_426, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)))).# trainpos
cnf(c_0_427, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~epred1_2(esk1_0,esk2_0)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)))).# trainpos
cnf(c_0_428, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~epred1_2(esk1_0,esk2_0)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)))).# trainpos
cnf(c_0_429, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~epred1_2(esk1_0,esk2_0)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)))).# trainpos
cnf(c_0_430, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~epred1_2(esk1_0,esk2_0)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)))).# trainpos
cnf(c_0_431, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~epred1_2(esk1_0,esk2_0))).# trainpos
cnf(c_0_433, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)))).# trainpos
cnf(c_0_436, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_439, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))).# trainpos
cnf(c_0_442, negated_conjecture, (esk4_0=esk8_0)).# trainpos
cnf(c_0_441, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_3, negated_conjecture, (esk3_0!=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))).#trainneg
cnf(i_0_153, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_91, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_159, plain, (v1_xboole_0(esk14_0))).#trainneg
cnf(i_0_108, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_154, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_125, plain, (v1_card_1(k1_card_1(X1)))).#trainneg
cnf(i_0_152, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_151, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_116, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_123, plain, (r1_xxreal_0(np__1,np__1))).#trainneg
cnf(i_0_140, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_99, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_15, negated_conjecture, (r2_hidden(esk4_0,k4_finseq_1(esk6_0)))).#trainneg
cnf(i_0_2, negated_conjecture, (esk4_0!=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))).#trainneg
cnf(i_0_112, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_105, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_52, plain, (l4_algstr_0(X1)|~l5_algstr_0(X1))).#trainneg
cnf(i_0_147, plain, (m1_finseq_2(k3_finseq_2(X1),X1))).#trainneg
cnf(i_0_102, plain, (l2_struct_0(X1)|~l4_struct_0(X1))).#trainneg
cnf(i_0_53, plain, (l1_algstr_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_81, plain, (l1_struct_0(X1)|~l1_algstr_0(X1))).#trainneg
cnf(i_0_122, plain, (l1_struct_0(X1)|~l3_struct_0(X1))).#trainneg
cnf(i_0_111, plain, (l1_struct_0(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_90, plain, (l3_struct_0(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_101, plain, (l3_struct_0(X1)|~l4_struct_0(X1))).#trainneg
cnf(i_0_89, plain, (l3_algstr_0(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_44, plain, (v7_ordinal1(X1)|~r2_hidden(X1,k4_ordinal1))).#trainneg
cnf(i_0_7, negated_conjecture, (esk3_0=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))).#trainneg
cnf(i_0_4, negated_conjecture, (esk4_0=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))).#trainneg
cnf(i_0_61, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_144, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_71, plain, (v1_funct_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_70, plain, (v1_finseq_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_63, plain, (m2_finseq_1(X1,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_162, plain, (~v1_xboole_0(X1)|~v1_xxreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_145, plain, (~r2_hidden(X3,X2)|~r2_hidden(X1,X2)|~r2_hidden(X1,esk13_2(X3,X2)))).#trainneg
cnf(i_0_126, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_96, plain, (m1_subset_1(k2_finseq_1(X1),k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_160, plain, (~v1_xxreal_0(X1)|~v2_xxreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_113, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v1_zfmisc_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_156, plain, (v2_struct_0(X1)|v3_vectsp_1(X1)|~v4_vectsp_1(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_155, plain, (v2_struct_0(X1)|v6_vectsp_1(X1)|~v4_vectsp_1(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_104, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_114, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_141, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_142, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_146, plain, (r2_hidden(esk13_2(X1,X2),X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_115, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_110, plain, (v1_xboole_0(X2)|~v7_ordinal1(X1)|~v1_xboole_0(k4_finseq_2(X1,X2)))).#trainneg
cnf(i_0_67, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_56, plain, (k9_xtuple_0(X1)=k4_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_119, plain, (k1_card_1(X1)=k3_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_149, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_66, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_158, plain, (v3_ordinal1(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_ordinal1(X1))).#trainneg
cnf(i_0_150, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_148, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_79, plain, (m1_subset_1(k4_finseq_1(X1),k1_zfmisc_1(k5_numbers))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_87, plain, (m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_120, plain, (m2_finseq_2(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_finseq_2(X2,X3))).#trainneg
cnf(i_0_121, plain, (m1_subset_1(X1,X3)|~m1_finseq_2(X3,X2)|~m2_finseq_2(X1,X2,X3))).#trainneg
cnf(i_0_124, plain, (v1_xboole_0(X1)|v3_card_1(X3,X2)|~v7_ordinal1(X2)|~m1_subset_1(X3,k4_finseq_2(X2,X1)))).#trainneg
cnf(i_0_143, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_100, plain, (m1_subset_1(k3_matrix_1(X2,X1,X3,X4),X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X1)|~m1_finseq_1(X1,k3_finseq_2(X2)))).#trainneg
cnf(i_0_37, plain, (X1=k6_matrix_1(X2,X3,X4)|v1_xboole_0(X2)|k1_matrix_1(X3)!=k3_finseq_1(X1)|k1_funct_1(X1,esk9_4(X2,X3,X4,X1))!=k3_matrix_1(X2,X3,X4,esk9_4(X2,X3,X4,X1))|~v7_ordinal1(X4)|~v1_matrix_1(X3)|~m2_finseq_1(X1,X2)|~m2_finseq_1(X3,k3_finseq_2(X2)))).#trainneg
cnf(i_0_182, plain, (k3_finseq_1(X1)=k3_finseq_1(X2)|v2_struct_0(X5)|v6_struct_0(X5)|X1!=k1_matrix12(X3,X4,X5,X2,X6,X7)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v13_algstr_0(X5)|~v33_algstr_0(X5)|~v3_group_1(X5)|~v5_group_1(X5)|~v2_rlvect_1(X5)|~v3_rlvect_1(X5)|~v4_rlvect_1(X5)|~v4_vectsp_1(X5)|~v5_vectsp_1(X5)|~l6_algstr_0(X5)|~epred1_2(X4,X3)|~m1_matrix_1(X2,u1_struct_0(X5),X3,X4)|~m1_matrix_1(X1,u1_struct_0(X5),X3,X4))).#trainneg
cnf(i_0_39, plain, (X4=k6_matrix_1(X1,X2,X3)|v1_xboole_0(X1)|v7_ordinal1(esk9_4(X1,X2,X3,X4))|k1_matrix_1(X2)!=k3_finseq_1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X4,X1)|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_38, plain, (X4=k6_matrix_1(X1,X2,X3)|v1_xboole_0(X1)|r2_hidden(esk9_4(X1,X2,X3,X4),k2_finseq_1(k1_matrix_1(X2)))|k1_matrix_1(X2)!=k3_finseq_1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X4,X1)|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_178, plain, (X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|v2_struct_0(X3)|v6_struct_0(X3)|v7_ordinal1(esk15_7(X1,X2,X3,X4,X5,X6,X7))|k3_finseq_1(X4)!=k3_finseq_1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2))).#trainneg
cnf(i_0_177, plain, (X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|v2_struct_0(X3)|v6_struct_0(X3)|v7_ordinal1(esk16_7(X1,X2,X3,X4,X5,X6,X7))|k3_finseq_1(X4)!=k3_finseq_1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2))).#trainneg
cnf(i_0_176, plain, (X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|v2_struct_0(X3)|v6_struct_0(X3)|r2_hidden(esk15_7(X1,X2,X3,X4,X5,X6,X7),k4_finseq_1(X4))|k3_finseq_1(X4)!=k3_finseq_1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2))).#trainneg
cnf(i_0_175, plain, (X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|v2_struct_0(X3)|v6_struct_0(X3)|r2_hidden(esk16_7(X1,X2,X3,X4,X5,X6,X7),k2_finseq_1(k1_matrix_1(X4)))|k3_finseq_1(X4)!=k3_finseq_1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2))).#trainneg
cnf(i_0_168, plain, (X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|esk15_7(X1,X2,X3,X4,X5,X6,X7)=X6|v2_struct_0(X3)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(X7)|esk15_7(X1,X2,X3,X4,X5,X6,X7)!=X5|k3_matrix_1(u1_struct_0(X3),X7,esk15_7(X1,X2,X3,X4,X5,X6,X7),esk16_7(X1,X2,X3,X4,X5,X6,X7))!=k3_matrix_1(u1_struct_0(X3),X4,X6,esk16_7(X1,X2,X3,X4,X5,X6,X7))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2))).#trainneg
cnf(i_0_170, plain, (X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|esk15_7(X1,X2,X3,X4,X5,X6,X7)=X5|v2_struct_0(X3)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(X7)|esk15_7(X1,X2,X3,X4,X5,X6,X7)!=X6|k3_matrix_1(u1_struct_0(X3),X7,esk15_7(X1,X2,X3,X4,X5,X6,X7),esk16_7(X1,X2,X3,X4,X5,X6,X7))!=k3_matrix_1(u1_struct_0(X3),X4,X5,esk16_7(X1,X2,X3,X4,X5,X6,X7))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2))).#trainneg
cnf(i_0_172, plain, (X2=k1_matrix12(X3,X4,X1,X5,X6,X7)|esk15_7(X3,X4,X1,X5,X6,X7,X2)=X7|esk15_7(X3,X4,X1,X5,X6,X7,X2)=X6|v2_struct_0(X1)|v6_struct_0(X1)|k3_finseq_1(X2)!=k3_finseq_1(X5)|k3_matrix_1(u1_struct_0(X1),X2,esk15_7(X3,X4,X1,X5,X6,X7,X2),esk16_7(X3,X4,X1,X5,X6,X7,X2))!=k3_matrix_1(u1_struct_0(X1),X5,esk15_7(X3,X4,X1,X5,X6,X7,X2),esk16_7(X3,X4,X1,X5,X6,X7,X2))|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~epred1_2(X4,X3)|~m1_matrix_1(X5,u1_struct_0(X1),X3,X4)|~m1_matrix_1(X2,u1_struct_0(X1),X3,X4))).#trainneg
cnf(i_0_164, plain, (X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|v2_struct_0(X3)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(X7)|esk15_7(X1,X2,X3,X4,X5,X6,X7)!=X6|k3_matrix_1(u1_struct_0(X3),X7,esk15_7(X1,X2,X3,X4,X5,X6,X7),esk16_7(X1,X2,X3,X4,X5,X6,X7))!=k3_matrix_1(u1_struct_0(X3),X4,X6,esk16_7(X1,X2,X3,X4,X5,X6,X7))|k3_matrix_1(u1_struct_0(X3),X7,esk15_7(X1,X2,X3,X4,X5,X6,X7),esk16_7(X1,X2,X3,X4,X5,X6,X7))!=k3_matrix_1(u1_struct_0(X3),X4,X5,esk16_7(X1,X2,X3,X4,X5,X6,X7))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2))).#trainneg
cnf(i_0_165, plain, (X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|v2_struct_0(X3)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(X7)|esk15_7(X1,X2,X3,X4,X5,X6,X7)!=X5|k3_matrix_1(u1_struct_0(X3),X7,esk15_7(X1,X2,X3,X4,X5,X6,X7),esk16_7(X1,X2,X3,X4,X5,X6,X7))!=k3_matrix_1(u1_struct_0(X3),X4,X6,esk16_7(X1,X2,X3,X4,X5,X6,X7))|k3_matrix_1(u1_struct_0(X3),X7,esk15_7(X1,X2,X3,X4,X5,X6,X7),esk16_7(X1,X2,X3,X4,X5,X6,X7))!=k3_matrix_1(u1_struct_0(X3),X4,X5,esk16_7(X1,X2,X3,X4,X5,X6,X7))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2))).#trainneg
cnf(i_0_166, plain, (X2=k1_matrix12(X3,X4,X1,X5,X6,X7)|esk15_7(X3,X4,X1,X5,X6,X7,X2)=X7|v2_struct_0(X1)|v6_struct_0(X1)|k3_finseq_1(X2)!=k3_finseq_1(X5)|k3_matrix_1(u1_struct_0(X1),X2,esk15_7(X3,X4,X1,X5,X6,X7,X2),esk16_7(X3,X4,X1,X5,X6,X7,X2))!=k3_matrix_1(u1_struct_0(X1),X5,X7,esk16_7(X3,X4,X1,X5,X6,X7,X2))|k3_matrix_1(u1_struct_0(X1),X2,esk15_7(X3,X4,X1,X5,X6,X7,X2),esk16_7(X3,X4,X1,X5,X6,X7,X2))!=k3_matrix_1(u1_struct_0(X1),X5,esk15_7(X3,X4,X1,X5,X6,X7,X2),esk16_7(X3,X4,X1,X5,X6,X7,X2))|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~epred1_2(X4,X3)|~m1_matrix_1(X5,u1_struct_0(X1),X3,X4)|~m1_matrix_1(X2,u1_struct_0(X1),X3,X4))).#trainneg
cnf(i_0_169, plain, (X2=k1_matrix12(X3,X4,X1,X5,X6,X7)|esk15_7(X3,X4,X1,X5,X6,X7,X2)=X6|v2_struct_0(X1)|v6_struct_0(X1)|k3_finseq_1(X2)!=k3_finseq_1(X5)|k3_matrix_1(u1_struct_0(X1),X2,esk15_7(X3,X4,X1,X5,X6,X7,X2),esk16_7(X3,X4,X1,X5,X6,X7,X2))!=k3_matrix_1(u1_struct_0(X1),X5,X6,esk16_7(X3,X4,X1,X5,X6,X7,X2))|k3_matrix_1(u1_struct_0(X1),X2,esk15_7(X3,X4,X1,X5,X6,X7,X2),esk16_7(X3,X4,X1,X5,X6,X7,X2))!=k3_matrix_1(u1_struct_0(X1),X5,esk15_7(X3,X4,X1,X5,X6,X7,X2),esk16_7(X3,X4,X1,X5,X6,X7,X2))|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~epred1_2(X4,X3)|~m1_matrix_1(X5,u1_struct_0(X1),X3,X4)|~m1_matrix_1(X2,u1_struct_0(X1),X3,X4))).#trainneg
cnf(i_0_163, plain, (X2=k1_matrix12(X3,X4,X1,X5,X6,X7)|v2_struct_0(X1)|v6_struct_0(X1)|k3_finseq_1(X2)!=k3_finseq_1(X5)|k3_matrix_1(u1_struct_0(X1),X2,esk15_7(X3,X4,X1,X5,X6,X7,X2),esk16_7(X3,X4,X1,X5,X6,X7,X2))!=k3_matrix_1(u1_struct_0(X1),X5,X7,esk16_7(X3,X4,X1,X5,X6,X7,X2))|k3_matrix_1(u1_struct_0(X1),X2,esk15_7(X3,X4,X1,X5,X6,X7,X2),esk16_7(X3,X4,X1,X5,X6,X7,X2))!=k3_matrix_1(u1_struct_0(X1),X5,X6,esk16_7(X3,X4,X1,X5,X6,X7,X2))|k3_matrix_1(u1_struct_0(X1),X2,esk15_7(X3,X4,X1,X5,X6,X7,X2),esk16_7(X3,X4,X1,X5,X6,X7,X2))!=k3_matrix_1(u1_struct_0(X1),X5,esk15_7(X3,X4,X1,X5,X6,X7,X2),esk16_7(X3,X4,X1,X5,X6,X7,X2))|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~epred1_2(X4,X3)|~m1_matrix_1(X5,u1_struct_0(X1),X3,X4)|~m1_matrix_1(X2,u1_struct_0(X1),X3,X4))).#trainneg
cnf(i_0_183, plain, (esk14_0=k1_xboole_0)).#trainneg
cnf(i_0_185, plain, (v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_197, negated_conjecture, (~v1_xboole_0(k4_finseq_1(esk6_0)))).#trainneg
cnf(i_0_186, plain, (l4_algstr_0(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_189, plain, (l1_struct_0(X1)|~l4_struct_0(X1))).#trainneg
cnf(i_0_293, plain, (k3_finseq_1(k1_matrix12(X1,X2,X3,X4,X5,X6))=k3_finseq_1(X4)|v6_struct_0(X3)|v2_struct_0(X3)|~epred1_2(X2,X1)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(k1_matrix12(X1,X2,X3,X4,X5,X6),u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_316, plain, (k1_matrix12(X2,X3,X4,X1,X5,X6)=X1|v6_struct_0(X4)|v2_struct_0(X4)|k3_matrix_1(u1_struct_0(X4),X1,esk15_7(X2,X3,X4,X1,X5,X6,X1),esk16_7(X2,X3,X4,X1,X5,X6,X1))!=k3_matrix_1(u1_struct_0(X4),X1,X6,esk16_7(X2,X3,X4,X1,X5,X6,X1))|k3_matrix_1(u1_struct_0(X4),X1,esk15_7(X2,X3,X4,X1,X5,X6,X1),esk16_7(X2,X3,X4,X1,X5,X6,X1))!=k3_matrix_1(u1_struct_0(X4),X1,X5,esk16_7(X2,X3,X4,X1,X5,X6,X1))|~epred1_2(X3,X2)|~m1_matrix_1(X1,u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_196, plain, (v7_ordinal1(np__1))).#trainneg
cnf(i_0_215, plain, (v3_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_218, plain, (v3_ordinal1(np__1))).#trainneg
cnf(i_0_213, plain, (~v1_xxreal_0(np__1)|~v1_xboole_0(np__1))).#trainneg
cnf(i_0_191, plain, (l1_struct_0(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_315, plain, (esk15_7(X1,X2,X3,X4,X5,X6,X4)=X5|k1_matrix12(X1,X2,X3,X4,X5,X6)=X4|v6_struct_0(X3)|v2_struct_0(X3)|k3_matrix_1(u1_struct_0(X3),X4,esk15_7(X1,X2,X3,X4,X5,X6,X4),esk16_7(X1,X2,X3,X4,X5,X6,X4))!=k3_matrix_1(u1_struct_0(X3),X4,X5,esk16_7(X1,X2,X3,X4,X5,X6,X4))|~epred1_2(X2,X1)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_314, plain, (esk15_7(X1,X2,X3,X4,X5,X6,X4)=X6|k1_matrix12(X1,X2,X3,X4,X5,X6)=X4|v6_struct_0(X3)|v2_struct_0(X3)|k3_matrix_1(u1_struct_0(X3),X4,esk15_7(X1,X2,X3,X4,X5,X6,X4),esk16_7(X1,X2,X3,X4,X5,X6,X4))!=k3_matrix_1(u1_struct_0(X3),X4,X6,esk16_7(X1,X2,X3,X4,X5,X6,X4))|~epred1_2(X2,X1)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_313, plain, (esk15_7(X1,X2,X3,X4,X5,X6,X4)=X5|esk15_7(X1,X2,X3,X4,X5,X6,X4)=X6|k1_matrix12(X1,X2,X3,X4,X5,X6)=X4|v6_struct_0(X3)|v2_struct_0(X3)|~epred1_2(X2,X1)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_255, plain, (m2_finseq_1(X1,X2)|~m1_finseq_2(X3,X2)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_308, plain, (k1_funct_1(X1,esk16_7(X2,X3,X4,X5,X6,X7,X8))=k3_matrix_1(X9,X5,X10,esk16_7(X2,X3,X4,X5,X6,X7,X8))|v1_xboole_0(X9)|X8=k1_matrix12(X2,X3,X4,X5,X6,X7)|v6_struct_0(X4)|v2_struct_0(X4)|X1!=k6_matrix_1(X9,X5,X10)|~m2_finseq_1(X5,k3_finseq_2(X9))|~m2_finseq_1(X1,X9)|~v1_matrix_1(X5)|~v7_ordinal1(X10)|~v7_ordinal1(esk16_7(X2,X3,X4,X5,X6,X7,X8))|k3_finseq_1(X5)!=k3_finseq_1(X8)|~epred1_2(X3,X2)|~m1_matrix_1(X8,u1_struct_0(X4),X2,X3)|~m1_matrix_1(X5,u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X7)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_306, plain, (m1_subset_1(esk16_7(X1,X2,X3,X4,X5,X6,X7),k2_finseq_1(k1_matrix_1(X4)))|X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|v2_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(X7)|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_304, plain, (m1_subset_1(esk15_7(X1,X2,X3,X4,X5,X6,X7),k4_finseq_1(X4))|X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|v2_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(X7)|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_312, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),X1,X2,X3)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,X3)|X2=esk3_0|X2=esk4_0|X1!=esk7_0|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X3,k2_finseq_1(k1_matrix_1(esk6_0)))|~r2_hidden(X2,k4_finseq_1(esk6_0))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_206, negated_conjecture, (m1_subset_1(esk4_0,k4_finseq_1(esk6_0)))).#trainneg
cnf(i_0_194, plain, (v7_ordinal1(esk12_1(k5_numbers)))).#trainneg
cnf(i_0_245, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_205, negated_conjecture, (m1_subset_1(esk8_0,k4_finseq_1(esk6_0)))).#trainneg
cnf(i_0_207, negated_conjecture, (m1_subset_1(esk3_0,k4_finseq_1(esk6_0)))).#trainneg
cnf(i_0_224, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_251, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_214, plain, (v3_ordinal1(esk12_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_219, plain, (v3_ordinal1(k2_finseq_1(X1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_290, plain, (m2_finseq_1(k1_matrix12(X1,X2,X3,X4,X5,X6),k3_finseq_2(u1_struct_0(X3)))|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|v2_struct_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_243, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_240, plain, (esk11_2(X1,X2)=X1|v1_xboole_0(a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,a_1_0_finseq_1(X2)))).#trainneg
cnf(i_0_294, negated_conjecture, (k3_finseq_1(X1)=k3_finseq_1(esk6_0)|X1!=esk7_0|~epred1_2(esk1_0,esk2_0)|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_305, plain, (X7=k1_matrix12(X2,X3,X4,X1,X5,X6)|v6_struct_0(X4)|v2_struct_0(X4)|~v1_xboole_0(k4_finseq_1(X1))|k3_finseq_1(X1)!=k3_finseq_1(X7)|~epred1_2(X3,X2)|~m1_matrix_1(X7,u1_struct_0(X4),X2,X3)|~m1_matrix_1(X1,u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_307, plain, (X7=k1_matrix12(X2,X3,X4,X1,X5,X6)|v6_struct_0(X4)|v2_struct_0(X4)|~v1_xboole_0(k2_finseq_1(k1_matrix_1(X1)))|k3_finseq_1(X1)!=k3_finseq_1(X7)|~epred1_2(X3,X2)|~m1_matrix_1(X7,u1_struct_0(X4),X2,X3)|~m1_matrix_1(X1,u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_228, plain, (v7_ordinal1(esk13_2(X1,k5_numbers))|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_225, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_229, plain, (m1_subset_1(esk13_2(X1,X2),X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_359, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_387, negated_conjecture, (v3_ordinal1(esk4_0)|~v3_ordinal1(k4_finseq_1(esk6_0)))).#trainneg
cnf(i_0_231, plain, (v1_relat_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_236, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_419, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)|v1_xboole_0(u1_struct_0(esk5_0))|~v1_matrix_1(esk6_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_327, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8)=k3_matrix_1(u1_struct_0(X1),X4,X7,X8)|X7=X6|X7=X5|v6_struct_0(X1)|v2_struct_0(X1)|~epred1_2(X3,X2)|~r2_hidden(X8,k2_finseq_1(k1_matrix_1(X4)))|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X8)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_412, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)|v1_xboole_0(u1_struct_0(esk5_0))|~v1_matrix_1(esk7_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_353, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,X4,X5,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,X4,X5)|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v1_matrix_1(X4)|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(esk7_0)))|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_388, plain, (v3_ordinal1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_239, plain, (esk11_2(esk13_2(X1,a_1_0_finseq_1(X2)),X2)=esk13_2(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2)))).#trainneg
cnf(i_0_415, negated_conjecture, (v1_funct_1(esk6_0)|v1_xboole_0(u1_struct_0(esk5_0)))).#trainneg
cnf(i_0_396, negated_conjecture, (v3_ordinal1(esk8_0)|~v3_ordinal1(k4_finseq_1(esk6_0)))).#trainneg
cnf(i_0_397, negated_conjecture, (v3_ordinal1(esk3_0)|~v3_ordinal1(k4_finseq_1(esk6_0)))).#trainneg
cnf(i_0_416, negated_conjecture, (v1_finseq_1(esk6_0)|v1_xboole_0(u1_struct_0(esk5_0)))).#trainneg
cnf(i_0_417, negated_conjecture, (v1_relat_1(esk6_0)|v1_xboole_0(u1_struct_0(esk5_0)))).#trainneg
cnf(i_0_241, plain, (v1_card_1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_242, plain, (v1_xboole_0(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_370, plain, (esk11_2(X1,X2)=X1|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_325, plain, (k3_finseq_1(k1_matrix12(X1,X2,X3,X4,X5,X6))=k3_finseq_1(X4)|v6_struct_0(X3)|v2_struct_0(X3)|~epred1_2(X2,X1)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_428, negated_conjecture, (k3_finseq_1(X1)=k3_finseq_1(esk6_0)|X1!=esk7_0|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_324, negated_conjecture, (k3_finseq_1(esk7_0)=k3_finseq_1(esk6_0)|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_367, negated_conjecture, (esk15_7(esk2_0,esk1_0,esk5_0,esk7_0,X1,X2,esk7_0)=X2|esk15_7(esk2_0,esk1_0,esk5_0,esk7_0,X1,X2,esk7_0)=X1|k1_matrix12(esk2_0,esk1_0,esk5_0,esk7_0,X1,X2)=esk7_0|~epred1_2(esk1_0,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_442, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_395, plain, (m1_subset_1(X1,k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_221, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,esk12_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_244, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_422, negated_conjecture, (v1_xboole_0(esk7_0)|v1_xboole_0(u1_struct_0(esk5_0))|~v1_xboole_0(k3_finseq_2(u1_struct_0(esk5_0))))).#trainneg
cnf(i_0_423, negated_conjecture, (v1_xboole_0(esk6_0)|v1_xboole_0(u1_struct_0(esk5_0))|~v1_xboole_0(k3_finseq_2(u1_struct_0(esk5_0))))).#trainneg
cnf(i_0_398, plain, (r1_xxreal_0(X1,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_488, negated_conjecture, (v3_ordinal1(esk4_0)|~v7_ordinal1(k4_finseq_1(esk6_0)))).#trainneg
cnf(i_0_246, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_247, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|~r2_hidden(X1,X3)|~m2_finseq_1(X3,X2))).#trainneg
cnf(i_0_466, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,X7)=k3_matrix_1(u1_struct_0(X1),X4,X6,X7)|v6_struct_0(X1)|v2_struct_0(X1)|~epred1_2(X3,X2)|~r2_hidden(X7,k2_finseq_1(k1_matrix_1(X4)))|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_548, negated_conjecture, (k3_finseq_1(esk7_0)=k3_finseq_1(esk6_0))).#trainneg
cnf(i_0_463, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_496, negated_conjecture, (v3_ordinal1(esk8_0)|~v7_ordinal1(k4_finseq_1(esk6_0)))).#trainneg
cnf(i_0_576, plain, (r1_xxreal_0(X1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_497, negated_conjecture, (v3_ordinal1(esk3_0)|~v7_ordinal1(k4_finseq_1(esk6_0)))).#trainneg
cnf(i_0_248, plain, (v3_ordinal1(k4_finseq_1(X1))|~v5_ordinal1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_441, negated_conjecture, (v7_ordinal1(k1_matrix_1(esk6_0))|v1_xboole_0(u1_struct_0(esk5_0))|~v1_finseq_1(esk6_0)|~v1_funct_1(esk6_0)|~v1_relat_1(esk6_0))).#trainneg
cnf(i_0_399, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_249, plain, (v1_xboole_0(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_252, plain, (v3_ordinal1(k4_finseq_1(X1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_473, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8)=k3_matrix_1(u1_struct_0(X1),X4,X7,X8)|X7=X5|X7=X6|v6_struct_0(X1)|v2_struct_0(X1)|~r2_hidden(X8,k2_finseq_1(k1_matrix_1(esk6_0)))|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X8)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_537, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,X4,X5,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,X4,X5)|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v1_matrix_1(X4)|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(X6)))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~m1_matrix_1(X6,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_564, plain, (v3_ordinal1(X1)|~v7_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_443, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_392, plain, (~v1_xboole_0(k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_254, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k4_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_393, plain, (m1_subset_1(esk13_2(X1,k5_numbers),k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_237, plain, (~v1_xboole_0(k2_zfmisc_1(k5_numbers,X2))|~r2_hidden(X3,X1)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_449, plain, (m2_subset_1(X1,X2,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_447, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~m2_subset_1(X1,X2,X2))).#trainneg
cnf(i_0_257, plain, (v3_ordinal1(k1_matrix_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1))).#trainneg
cnf(i_0_259, plain, (v3_card_1(esk12_1(k4_finseq_2(X1,X2)),X1)|v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_467, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,X1)|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(esk6_0)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_600, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,X7)=k3_matrix_1(u1_struct_0(X1),X4,X6,X7)|v6_struct_0(X1)|~r2_hidden(X7,k2_finseq_1(k1_matrix_1(esk6_0)))|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_634, negated_conjecture, (k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))=k1_matrix_1(esk6_0)|v1_xboole_0(u1_struct_0(esk5_0))|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_610, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|~r2_hidden(X7,k2_finseq_1(k1_matrix_1(esk6_0)))|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_622, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X4,X6)=k3_matrix_1(u1_struct_0(X1),esk7_0,X5,X6)|v6_struct_0(X1)|~r2_hidden(X6,k2_finseq_1(k1_matrix_1(esk6_0)))|~r2_hidden(X4,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_678, plain, (v3_ordinal1(X1)|~v7_ordinal1(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_413, negated_conjecture, (k1_matrix_1(esk6_0)=k3_finseq_1(X1)|v1_xboole_0(u1_struct_0(esk5_0))|X1!=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_262, plain, (m2_subset_1(X1,X2,esk12_1(k1_zfmisc_1(X2)))|v1_xboole_0(esk12_1(k1_zfmisc_1(X2)))|v1_xboole_0(X2)|~m1_subset_1(X1,esk12_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_692, negated_conjecture, (m1_subset_1(esk8_0,k5_numbers)|~v1_finseq_1(esk6_0)|~v1_funct_1(esk6_0)|~v1_relat_1(esk6_0))).#trainneg
cnf(i_0_646, negated_conjecture, (r1_xxreal_0(X1,esk8_0)|r1_xxreal_0(esk8_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_647, negated_conjecture, (r1_xxreal_0(X1,esk4_0)|r1_xxreal_0(esk4_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_266, plain, (m2_subset_1(X1,k2_zfmisc_1(k5_numbers,X2),X3)|v1_xboole_0(X3)|v1_xboole_0(k2_zfmisc_1(k5_numbers,X2))|~m1_subset_1(X1,X3)|~m2_finseq_1(X3,X2))).#trainneg
cnf(i_0_267, plain, (m1_subset_1(X1,esk12_1(k1_zfmisc_1(X2)))|v1_xboole_0(X2)|v1_xboole_0(esk12_1(k1_zfmisc_1(X2)))|~m2_subset_1(X1,X2,esk12_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_598, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,X7)=k3_matrix_1(u1_struct_0(X1),X4,X6,X7)|v6_struct_0(X1)|~r2_hidden(X7,k2_finseq_1(k1_matrix_1(X8)))|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X8,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_608, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|~r2_hidden(X7,k2_finseq_1(k1_matrix_1(X8)))|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X8,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_624, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X5,X6)=k3_matrix_1(u1_struct_0(X1),esk7_0,X4,X6)|v6_struct_0(X1)|~r2_hidden(X6,k2_finseq_1(k1_matrix_1(esk6_0)))|~r2_hidden(X5,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_623, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X6,X7)=k3_matrix_1(u1_struct_0(X1),esk7_0,X6,X7)|X6=X4|X6=X5|v6_struct_0(X1)|v2_struct_0(X1)|~r2_hidden(X7,k2_finseq_1(k1_matrix_1(esk6_0)))|~r2_hidden(X6,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_660, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8)=k3_matrix_1(u1_struct_0(X1),X4,X7,X8)|X7=X6|X7=X5|v6_struct_0(X1)|~r2_hidden(X8,k2_finseq_1(k1_matrix_1(X9)))|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X8)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X9,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_679, negated_conjecture, (v3_ordinal1(esk4_0)|~v7_ordinal1(k1_zfmisc_1(k4_finseq_1(esk6_0))))).#trainneg
cnf(i_0_444, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_271, plain, (m1_subset_1(X1,X2)|v1_xboole_0(k2_zfmisc_1(k5_numbers,X3))|v1_xboole_0(X2)|~m2_subset_1(X1,k2_zfmisc_1(k5_numbers,X3),X2)|~m2_finseq_1(X2,X3))).#trainneg
cnf(i_0_693, negated_conjecture, (m1_subset_1(esk4_0,k5_numbers)|~v1_finseq_1(esk6_0)|~v1_funct_1(esk6_0)|~v1_relat_1(esk6_0))).#trainneg
cnf(i_0_680, negated_conjecture, (v3_ordinal1(esk8_0)|~v7_ordinal1(k1_zfmisc_1(k4_finseq_1(esk6_0))))).#trainneg
cnf(i_0_681, negated_conjecture, (v3_ordinal1(esk3_0)|~v7_ordinal1(k1_zfmisc_1(k4_finseq_1(esk6_0))))).#trainneg
cnf(i_0_648, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|r1_xxreal_0(esk3_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_649, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|r1_xxreal_0(esk2_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_272, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|v1_xboole_0(esk12_1(k1_zfmisc_1(X2)))|~m2_subset_1(X1,X2,esk12_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_276, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(X3)|~m2_subset_1(X1,k2_zfmisc_1(k5_numbers,X2),X3)|~m2_finseq_1(X3,X2))).#trainneg
cnf(i_0_717, negated_conjecture, (k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))=k1_matrix_1(esk6_0)|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_409, negated_conjecture, (v1_finseq_1(esk7_0))).#trainneg
cnf(i_0_408, negated_conjecture, (v1_funct_1(esk7_0))).#trainneg
cnf(i_0_410, negated_conjecture, (v1_relat_1(esk7_0))).#trainneg
cnf(i_0_748, negated_conjecture, (r1_xxreal_0(esk8_0,esk8_0))).#trainneg
cnf(i_0_759, negated_conjecture, (r1_xxreal_0(esk4_0,esk4_0))).#trainneg
cnf(i_0_415, negated_conjecture, (v1_funct_1(esk6_0))).#trainneg
cnf(i_0_281, plain, (m1_subset_1(k3_matrix_1(X1,X2,X3,X4),X1)|~v1_matrix_1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_830, negated_conjecture, (r1_xxreal_0(esk3_0,esk3_0))).#trainneg
cnf(i_0_416, negated_conjecture, (v1_finseq_1(esk6_0))).#trainneg
cnf(i_0_417, negated_conjecture, (v1_relat_1(esk6_0))).#trainneg
cnf(i_0_842, negated_conjecture, (r1_xxreal_0(esk2_0,esk2_0))).#trainneg
cnf(i_0_749, negated_conjecture, (r1_xxreal_0(esk8_0,esk4_0)|r1_xxreal_0(esk4_0,esk8_0))).#trainneg
cnf(i_0_283, plain, (m1_subset_1(k8_matrix_1(X1,X2,X3),k4_finseq_2(k1_matrix_1(X2),X1))|v1_xboole_0(X1)|~m1_finseq_2(k4_finseq_2(k1_matrix_1(X2),X1),X1)|~m1_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_287, plain, (k1_funct_1(X1,esk13_2(X2,k2_finseq_1(k1_matrix_1(X3))))=k3_matrix_1(X4,X3,X5,esk13_2(X2,k2_finseq_1(k1_matrix_1(X3))))|v1_xboole_0(X4)|X1!=k6_matrix_1(X4,X3,X5)|~m2_finseq_1(X3,k3_finseq_2(X4))|~m2_finseq_1(X1,X4)|~v1_matrix_1(X3)|~v7_ordinal1(X5)|~v7_ordinal1(esk13_2(X2,k2_finseq_1(k1_matrix_1(X3))))|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(X3))))).#trainneg
cnf(i_0_814, plain, (~m2_finseq_1(X1,X2)|~r2_hidden(X3,X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_744, negated_conjecture, (k1_matrix_1(esk6_0)=k3_finseq_1(X1)|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_597, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X4,X6)=k3_matrix_1(u1_struct_0(X1),esk6_0,X5,X6)|v6_struct_0(X1)|~r2_hidden(X6,k2_finseq_1(k1_matrix_1(X7)))|~r2_hidden(X4,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X7,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_812, plain, (v7_ordinal1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_445, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_750, negated_conjecture, (r1_xxreal_0(esk8_0,esk3_0)|r1_xxreal_0(esk3_0,esk8_0))).#trainneg
cnf(i_0_288, plain, (k1_funct_1(X1,X2)=k3_matrix_1(X3,X4,X5,X2)|v1_xboole_0(X3)|v1_xboole_0(k2_finseq_1(k1_matrix_1(X4)))|X1!=k6_matrix_1(X3,X4,X5)|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v1_matrix_1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~m1_subset_1(X2,k2_finseq_1(k1_matrix_1(X4))))).#trainneg
cnf(i_0_751, negated_conjecture, (r1_xxreal_0(esk8_0,esk2_0)|r1_xxreal_0(esk2_0,esk8_0))).#trainneg
cnf(i_0_636, negated_conjecture, (v3_ordinal1(esk4_0)|~v5_ordinal1(esk6_0)|~v1_finseq_1(esk6_0)|~v1_funct_1(esk6_0)|~v1_relat_1(esk6_0))).#trainneg
cnf(i_0_637, negated_conjecture, (v3_ordinal1(esk8_0)|~v5_ordinal1(esk6_0)|~v1_finseq_1(esk6_0)|~v1_funct_1(esk6_0)|~v1_relat_1(esk6_0))).#trainneg
cnf(i_0_752, negated_conjecture, (r1_xxreal_0(esk8_0,esk1_0)|r1_xxreal_0(esk1_0,esk8_0))).#trainneg
cnf(i_0_289, plain, (v1_matrix_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|v2_struct_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_291, plain, (k1_matrix_1(X1)=k1_matrix_1(k1_matrix12(X2,X3,X4,X5,X6,X7))|v6_struct_0(X4)|v2_struct_0(X4)|~m1_matrix_1(X1,u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X5,u1_struct_0(X4),X2,X3)|~v7_ordinal1(X7)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_536, plain, (esk11_2(X1,X2)=X1|v1_xboole_0(k2_finseq_1(X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k2_finseq_1(X2)))).#trainneg
cnf(i_0_553, negated_conjecture, (k1_matrix_1(esk6_0)=k3_finseq_1(X1)|v1_xboole_0(u1_struct_0(esk5_0))|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_607, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X5,X6)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X6)|v6_struct_0(X1)|~r2_hidden(X6,k2_finseq_1(k1_matrix_1(X7)))|~r2_hidden(X5,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X7,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_563, plain, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X3)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,X3)|esk15_7(esk2_0,esk1_0,esk5_0,esk7_0,X1,X2,esk7_0)=X1|esk15_7(esk2_0,esk1_0,esk5_0,esk7_0,X1,X2,esk7_0)=X2|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X3,k2_finseq_1(k1_matrix_1(esk7_0)))|~r2_hidden(X1,k4_finseq_1(esk7_0))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_487, plain, (v3_ordinal1(X1)|~m2_finseq_1(X1,X2)|~v7_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_462, plain, (v3_ordinal1(esk13_2(X1,X2))|~v3_ordinal1(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_753, negated_conjecture, (r1_xxreal_0(esk8_0,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,esk8_0))).#trainneg
cnf(i_0_456, plain, (~v1_xboole_0(X2)|~r2_hidden(X3,esk13_2(X1,k1_zfmisc_1(X2)))|~r2_hidden(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_295, plain, (m1_subset_1(esk9_4(X1,X2,X3,X4),k2_finseq_1(k1_matrix_1(X2)))|X4=k6_matrix_1(X1,X2,X3)|v1_xboole_0(X1)|k1_matrix_1(X2)!=k3_finseq_1(X4)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m2_finseq_1(X4,X1)|~v1_matrix_1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_754, negated_conjecture, (r1_xxreal_0(esk8_0,np__1)|r1_xxreal_0(np__1,esk8_0))).#trainneg
cnf(i_0_638, negated_conjecture, (v3_ordinal1(esk3_0)|~v5_ordinal1(esk6_0)|~v1_finseq_1(esk6_0)|~v1_funct_1(esk6_0)|~v1_relat_1(esk6_0))).#trainneg
cnf(i_0_657, negated_conjecture, (v3_ordinal1(esk4_0)|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(esk6_0)|~v1_funct_1(esk6_0)|~v1_relat_1(esk6_0))).#trainneg
cnf(i_0_760, negated_conjecture, (r1_xxreal_0(esk4_0,esk3_0)|r1_xxreal_0(esk3_0,esk4_0))).#trainneg
cnf(i_0_761, negated_conjecture, (r1_xxreal_0(esk4_0,esk2_0)|r1_xxreal_0(esk2_0,esk4_0))).#trainneg
cnf(i_0_297, plain, (k1_funct_1(X1,esk9_4(X2,X3,X4,X5))=k3_matrix_1(X6,X3,X7,esk9_4(X2,X3,X4,X5))|v1_xboole_0(X6)|X5=k6_matrix_1(X2,X3,X4)|v1_xboole_0(X2)|X1!=k6_matrix_1(X6,X3,X7)|~m2_finseq_1(X3,k3_finseq_2(X6))|~m2_finseq_1(X1,X6)|~v1_matrix_1(X3)|~v7_ordinal1(X7)|~v7_ordinal1(esk9_4(X2,X3,X4,X5))|k1_matrix_1(X3)!=k3_finseq_1(X5)|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_300, plain, (k1_matrix12(X1,X2,X3,X4,X5,X6)=k1_matrix12(X1,X2,X3,X7,X8,X9)|v6_struct_0(X3)|v2_struct_0(X3)|v7_ordinal1(esk15_7(X1,X2,X3,X7,X8,X9,k1_matrix12(X1,X2,X3,X4,X5,X6)))|k3_finseq_1(X7)!=k3_finseq_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_605, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,X7)=k3_matrix_1(u1_struct_0(X1),X4,X6,X7)|v6_struct_0(X1)|v1_xboole_0(k2_finseq_1(k1_matrix_1(X4)))|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X7,k2_finseq_1(k1_matrix_1(X4))))).#trainneg
cnf(i_0_909, plain, (v7_ordinal1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_499, plain, (v7_ordinal1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_958, negated_conjecture, (v3_ordinal1(esk4_0)|~v7_ordinal1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_693, negated_conjecture, (m1_subset_1(esk4_0,k5_numbers)|~v1_finseq_1(esk6_0)|~v1_relat_1(esk6_0))).#trainneg
cnf(i_0_303, plain, (k1_matrix12(X1,X2,X3,X4,X5,X6)=k1_matrix12(X1,X2,X3,X7,X8,X9)|v6_struct_0(X3)|v2_struct_0(X3)|v7_ordinal1(esk16_7(X1,X2,X3,X7,X8,X9,k1_matrix12(X1,X2,X3,X4,X5,X6)))|k3_finseq_1(X7)!=k3_finseq_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_762, negated_conjecture, (r1_xxreal_0(esk4_0,esk1_0)|r1_xxreal_0(esk1_0,esk4_0))).#trainneg
cnf(i_0_858, negated_conjecture, (v1_xboole_0(esk6_0)|~v1_xboole_0(k3_finseq_2(u1_struct_0(esk5_0))))).#trainneg
cnf(i_0_867, negated_conjecture, (v1_xboole_0(esk7_0)|~v1_xboole_0(k3_finseq_2(u1_struct_0(esk5_0))))).#trainneg
cnf(i_0_356, plain, (m1_subset_1(X1,a_1_0_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_861, negated_conjecture, (k1_matrix_1(esk6_0)=k3_finseq_1(X1)|X1!=k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_923, negated_conjecture, (k1_matrix_1(k1_matrix12(esk2_0,esk1_0,esk5_0,X1,X2,X3))=k1_matrix_1(esk6_0)|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_562, plain, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,X3)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X3)|esk15_7(esk2_0,esk1_0,esk5_0,esk7_0,X1,X2,esk7_0)=X1|esk15_7(esk2_0,esk1_0,esk5_0,esk7_0,X1,X2,esk7_0)=X2|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X3,k2_finseq_1(k1_matrix_1(esk7_0)))|~r2_hidden(X2,k4_finseq_1(esk7_0))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_373, plain, (m2_finseq_1(X1,X2)|~m1_subset_1(X1,k3_finseq_2(X2)))).#trainneg
cnf(i_0_655, plain, (v7_ordinal1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_763, negated_conjecture, (r1_xxreal_0(esk4_0,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,esk4_0))).#trainneg
cnf(i_0_692, negated_conjecture, (m1_subset_1(esk8_0,k5_numbers)|~v1_finseq_1(esk6_0)|~v1_relat_1(esk6_0))).#trainneg
cnf(i_0_361, plain, (X1=k6_matrix_1(X2,k6_matrix_1(k3_finseq_2(X2),X3,X4),X5)|v1_xboole_0(X2)|v7_ordinal1(esk9_4(X2,k6_matrix_1(k3_finseq_2(X2),X3,X4),X5,X1))|v1_xboole_0(k3_finseq_2(X2))|k1_matrix_1(k6_matrix_1(k3_finseq_2(X2),X3,X4))!=k3_finseq_1(X1)|~m2_finseq_1(X1,X2)|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X2),X3,X4))|~v7_ordinal1(X5)|~m2_finseq_1(X3,k3_finseq_2(k3_finseq_2(X2)))|~v1_matrix_1(X3)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_764, negated_conjecture, (r1_xxreal_0(esk4_0,np__1)|r1_xxreal_0(np__1,esk4_0))).#trainneg
cnf(i_0_957, negated_conjecture, (v3_ordinal1(esk4_0)|~v7_ordinal1(k1_zfmisc_1(k1_zfmisc_1(k5_numbers))))).#trainneg
cnf(i_0_658, negated_conjecture, (v3_ordinal1(esk8_0)|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(esk6_0)|~v1_funct_1(esk6_0)|~v1_relat_1(esk6_0))).#trainneg
cnf(i_0_831, negated_conjecture, (r1_xxreal_0(esk3_0,esk2_0)|r1_xxreal_0(esk2_0,esk3_0))).#trainneg
cnf(i_0_832, negated_conjecture, (r1_xxreal_0(esk3_0,esk1_0)|r1_xxreal_0(esk1_0,esk3_0))).#trainneg
cnf(i_0_362, plain, (k6_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4)=k8_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4)|v1_xboole_0(X1)|v1_xboole_0(k3_finseq_2(X1))|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))|~v7_ordinal1(X4)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(X1)))|~v1_matrix_1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_363, plain, (k1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))=k3_finseq_1(X4)|v1_xboole_0(X1)|v1_xboole_0(k3_finseq_2(X1))|X4!=k6_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X5)|~m2_finseq_1(X4,X1)|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))|~v7_ordinal1(X5)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(X1)))|~v1_matrix_1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_615, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|v1_xboole_0(k2_finseq_1(k1_matrix_1(X4)))|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X7,k2_finseq_1(k1_matrix_1(X4))))).#trainneg
cnf(i_0_709, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,X1)|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2)))|~v7_ordinal1(X1)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_514, negated_conjecture, (m2_finseq_2(k8_matrix_1(X1,X2,X3),X1,k4_finseq_2(k1_matrix_1(X4),X1))|v1_xboole_0(X1)|~m1_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~v7_ordinal1(X3)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_682, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,X1)|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2)))|~v7_ordinal1(X1)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_677, plain, (v3_ordinal1(X1)|~m2_finseq_1(X1,X2)|~v7_ordinal1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))))).#trainneg
cnf(i_0_1054, plain, (m2_finseq_1(esk12_1(k3_finseq_2(X1)),X1))).#trainneg
cnf(i_0_1002, negated_conjecture, (v3_ordinal1(esk4_0))).#trainneg
cnf(i_0_833, negated_conjecture, (r1_xxreal_0(esk3_0,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,esk3_0))).#trainneg
cnf(i_0_735, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X4,X6)=k3_matrix_1(u1_struct_0(X1),esk7_0,X5,X6)|v6_struct_0(X1)|~r2_hidden(X6,k2_finseq_1(k1_matrix_1(X7)))|~r2_hidden(X4,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X7,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_659, negated_conjecture, (v3_ordinal1(esk3_0)|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(esk6_0)|~v1_funct_1(esk6_0)|~v1_relat_1(esk6_0))).#trainneg
cnf(i_0_696, plain, (v3_ordinal1(esk13_2(X1,k5_numbers))|~v3_ordinal1(k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_1060, negated_conjecture, (v3_ordinal1(esk8_0))).#trainneg
cnf(i_0_979, negated_conjecture, (k1_funct_1(X1,esk13_2(X2,k2_finseq_1(k1_matrix_1(esk6_0))))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X3,esk13_2(X2,k2_finseq_1(k1_matrix_1(esk6_0))))|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(esk6_0)))|~v7_ordinal1(esk13_2(X2,k2_finseq_1(k1_matrix_1(esk6_0))))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_369, plain, (esk15_7(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8,k1_matrix12(X1,X2,X3,X4,X5,X6))=X8|esk15_7(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8,k1_matrix12(X1,X2,X3,X4,X5,X6))=X7|k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|v2_struct_0(X3)|~epred1_2(X2,X1)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_834, negated_conjecture, (r1_xxreal_0(esk3_0,np__1)|r1_xxreal_0(np__1,esk3_0))).#trainneg
cnf(i_0_1242, negated_conjecture, (v3_ordinal1(esk3_0)|~v7_ordinal1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1241, negated_conjecture, (v3_ordinal1(esk3_0)|~v7_ordinal1(k1_zfmisc_1(k1_zfmisc_1(k5_numbers))))).#trainneg
cnf(i_0_843, negated_conjecture, (r1_xxreal_0(esk2_0,esk1_0)|r1_xxreal_0(esk1_0,esk2_0))).#trainneg
cnf(i_0_372, plain, (m2_finseq_1(X1,X2)|~m1_subset_1(X1,k4_finseq_2(X3,X2))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1155, negated_conjecture, (k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1215, plain, (v1_funct_1(esk12_1(k3_finseq_2(X1))))).#trainneg
cnf(i_0_1216, plain, (v1_finseq_1(esk12_1(k3_finseq_2(X1))))).#trainneg
cnf(i_0_844, negated_conjecture, (r1_xxreal_0(esk2_0,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,esk2_0))).#trainneg
cnf(i_0_375, plain, (v3_ordinal1(esk16_7(X1,X2,X3,X4,X5,X6,X7))|X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|v2_struct_0(X3)|~v3_ordinal1(k2_finseq_1(k1_matrix_1(X4)))|k3_finseq_1(X4)!=k3_finseq_1(X7)|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_845, negated_conjecture, (r1_xxreal_0(esk2_0,np__1)|r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_632, negated_conjecture, (v1_card_1(k3_finseq_1(esk6_0))|~v1_finseq_1(esk7_0)|~v1_funct_1(esk7_0)|~v1_relat_1(esk7_0))).#trainneg
cnf(i_0_694, negated_conjecture, (m1_subset_1(esk3_0,k5_numbers)|~v1_finseq_1(esk6_0)|~v1_funct_1(esk6_0)|~v1_relat_1(esk6_0))).#trainneg
cnf(i_0_650, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|r1_xxreal_0(esk1_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_374, plain, (k1_funct_1(X1,esk16_7(X2,X3,X4,k6_matrix_1(k3_finseq_2(X5),X6,X7),X8,X9,X10))=k3_matrix_1(X5,k6_matrix_1(k3_finseq_2(X5),X6,X7),X11,esk16_7(X2,X3,X4,k6_matrix_1(k3_finseq_2(X5),X6,X7),X8,X9,X10))|X10=k1_matrix12(X2,X3,X4,k6_matrix_1(k3_finseq_2(X5),X6,X7),X8,X9)|v1_xboole_0(X5)|v6_struct_0(X4)|v2_struct_0(X4)|v1_xboole_0(k3_finseq_2(X5))|X1!=k6_matrix_1(X5,k6_matrix_1(k3_finseq_2(X5),X6,X7),X11)|k3_finseq_1(k6_matrix_1(k3_finseq_2(X5),X6,X7))!=k3_finseq_1(X10)|~epred1_2(X3,X2)|~m2_finseq_1(X1,X5)|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X5),X6,X7))|~m1_matrix_1(X10,u1_struct_0(X4),X2,X3)|~m1_matrix_1(k6_matrix_1(k3_finseq_2(X5),X6,X7),u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X11)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~m2_finseq_1(X6,k3_finseq_2(k3_finseq_2(X5)))|~v1_matrix_1(X6)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_380, plain, (v3_ordinal1(esk15_7(X1,X2,X3,X4,X5,X6,X7))|X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|v2_struct_0(X3)|~v3_ordinal1(k4_finseq_1(X4))|k3_finseq_1(X4)!=k3_finseq_1(X7)|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1217, plain, (v1_relat_1(esk12_1(k3_finseq_2(X1))))).#trainneg
cnf(i_0_1277, plain, (m2_finseq_1(esk12_1(k4_finseq_2(X1,X2)),X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_746, negated_conjecture, (r1_xxreal_0(esk8_0,esk12_1(k5_numbers))|r1_xxreal_0(esk12_1(k5_numbers),esk8_0))).#trainneg
cnf(i_0_1219, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,esk12_1(k3_finseq_2(X1))))).#trainneg
cnf(i_0_400, plain, (v1_funct_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|v2_struct_0(X3)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_651, plain, (r1_xxreal_0(X1,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_704, negated_conjecture, (v3_ordinal1(k1_matrix_1(esk6_0))|~v1_finseq_1(esk7_0)|~v1_funct_1(esk7_0)|~v1_relat_1(esk7_0)|~v1_matrix_1(esk7_0))).#trainneg
cnf(i_0_619, negated_conjecture, (m1_subset_1(k1_matrix_1(esk6_0),k5_numbers)|~v1_finseq_1(esk7_0)|~v1_funct_1(esk7_0)|~v1_relat_1(esk7_0)|~v1_matrix_1(esk7_0))).#trainneg
cnf(i_0_401, plain, (v1_finseq_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|v2_struct_0(X3)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_402, plain, (v1_relat_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|v2_struct_0(X3)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_403, plain, (X1=k6_matrix_1(u1_struct_0(X2),k1_matrix12(X3,X4,X2,X5,X6,X7),X8)|v1_xboole_0(u1_struct_0(X2))|v7_ordinal1(esk9_4(u1_struct_0(X2),k1_matrix12(X3,X4,X2,X5,X6,X7),X8,X1))|v6_struct_0(X2)|v2_struct_0(X2)|k1_matrix_1(k1_matrix12(X3,X4,X2,X5,X6,X7))!=k3_finseq_1(X1)|~m2_finseq_1(X1,u1_struct_0(X2))|~v1_matrix_1(k1_matrix12(X3,X4,X2,X5,X6,X7))|~v7_ordinal1(X8)|~m1_matrix_1(X5,u1_struct_0(X2),X3,X4)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X7)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_404, plain, (k6_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7)=k8_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6))|~v7_ordinal1(X7)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1376, negated_conjecture, (k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1319, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,X7)=k3_matrix_1(u1_struct_0(X1),X4,X6,X7)|v6_struct_0(X1)|~r2_hidden(X7,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X8))))|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_1366, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,X7)=k3_matrix_1(u1_struct_0(X1),X4,X6,X7)|v6_struct_0(X1)|~r2_hidden(X7,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X8))))|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_1558, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,X7)=k3_matrix_1(u1_struct_0(X1),X4,X6,X7)|v6_struct_0(X1)|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X8)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(X7,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X8)))|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X8))))).#trainneg
cnf(i_0_1254, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,X7)=k3_matrix_1(u1_struct_0(X1),X4,X6,X7)|v6_struct_0(X1)|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(X7,k1_matrix_1(esk6_0)))).#trainneg
cnf(i_0_1441, negated_conjecture, (r1_xxreal_0(esk1_0,esk1_0))).#trainneg
cnf(i_0_1482, plain, (r1_xxreal_0(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_1426, negated_conjecture, (v3_ordinal1(esk3_0))).#trainneg
cnf(i_0_405, plain, (k1_matrix_1(k1_matrix12(X1,X2,X3,X4,X5,X6))=k3_finseq_1(X7)|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|v2_struct_0(X3)|X7!=k6_matrix_1(u1_struct_0(X3),k1_matrix12(X1,X2,X3,X4,X5,X6),X8)|~m2_finseq_1(X7,u1_struct_0(X3))|~v1_matrix_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~v7_ordinal1(X8)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_406, plain, (k1_funct_1(X1,esk16_7(X2,X3,X4,k1_matrix12(X5,X6,X7,X8,X9,X10),X11,X12,X13))=k3_matrix_1(u1_struct_0(X7),k1_matrix12(X5,X6,X7,X8,X9,X10),X14,esk16_7(X2,X3,X4,k1_matrix12(X5,X6,X7,X8,X9,X10),X11,X12,X13))|X13=k1_matrix12(X2,X3,X4,k1_matrix12(X5,X6,X7,X8,X9,X10),X11,X12)|v1_xboole_0(u1_struct_0(X7))|v6_struct_0(X4)|v2_struct_0(X4)|v6_struct_0(X7)|v2_struct_0(X7)|X1!=k6_matrix_1(u1_struct_0(X7),k1_matrix12(X5,X6,X7,X8,X9,X10),X14)|k3_finseq_1(k1_matrix12(X5,X6,X7,X8,X9,X10))!=k3_finseq_1(X13)|~epred1_2(X3,X2)|~m2_finseq_1(X1,u1_struct_0(X7))|~v1_matrix_1(k1_matrix12(X5,X6,X7,X8,X9,X10))|~m1_matrix_1(X13,u1_struct_0(X4),X2,X3)|~m1_matrix_1(k1_matrix12(X5,X6,X7,X8,X9,X10),u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X14)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~m1_matrix_1(X8,u1_struct_0(X7),X5,X6)|~l6_algstr_0(X7)|~v5_vectsp_1(X7)|~v4_vectsp_1(X7)|~v4_rlvect_1(X7)|~v3_rlvect_1(X7)|~v2_rlvect_1(X7)|~v5_group_1(X7)|~v3_group_1(X7)|~v33_algstr_0(X7)|~v13_algstr_0(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X10)|~v7_ordinal1(X9))).#trainneg
cnf(i_0_425, plain, (v1_xboole_0(k1_matrix12(X1,X2,X3,X4,X5,X6))|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|v2_struct_0(X3)|~v1_xboole_0(k3_finseq_2(u1_struct_0(X3)))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1550, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,X7)=k3_matrix_1(u1_struct_0(X1),X4,X6,X7)|v6_struct_0(X1)|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X8)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(X7,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X8)))|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X8))))).#trainneg
cnf(i_0_1126, negated_conjecture, (m2_finseq_2(k8_matrix_1(X1,X2,X3),X1,k4_finseq_2(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X4)),X1))|v1_xboole_0(X1)|~m1_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X3)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_1262, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,X7)=k3_matrix_1(u1_struct_0(X1),X4,X6,X7)|v6_struct_0(X1)|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(X7,k1_matrix_1(X4))|~v7_ordinal1(k1_matrix_1(X4)))).#trainneg
cnf(i_0_1218, plain, (v1_xboole_0(esk12_1(k3_finseq_2(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1460, plain, (v1_funct_1(esk12_1(k4_finseq_2(X1,X2)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1442, negated_conjecture, (r1_xxreal_0(esk1_0,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,esk1_0))).#trainneg
cnf(i_0_1445, negated_conjecture, (r1_xxreal_0(esk1_0,np__1)|r1_xxreal_0(np__1,esk1_0))).#trainneg
cnf(i_0_811, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_908, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_427, plain, (esk11_2(esk12_1(a_1_0_finseq_1(X1)),X1)=esk12_1(a_1_0_finseq_1(X1))|v1_xboole_0(a_1_0_finseq_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1148, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X4,X6)=k3_matrix_1(u1_struct_0(X1),esk7_0,X5,X6)|v6_struct_0(X1)|~r2_hidden(X6,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7))))|~r2_hidden(X4,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_1293, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|~r2_hidden(X7,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X8))))|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_458, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk13_2(X3,k1_zfmisc_1(X2)))|~r2_hidden(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1461, plain, (v1_finseq_1(esk12_1(k4_finseq_2(X1,X2)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_451, plain, (v3_card_1(esk13_2(X1,k4_finseq_2(X2,X3)),X2)|v1_xboole_0(X3)|~v7_ordinal1(X2)|~r2_hidden(X1,k4_finseq_2(X2,X3)))).#trainneg
cnf(i_0_695, plain, (m1_subset_1(X1,k5_numbers)|v1_xboole_0(k4_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(X1,k4_finseq_1(X2)))).#trainneg
cnf(i_0_879, plain, (m1_subset_1(k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6))|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_453, plain, (v1_relat_1(esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1396, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|~r2_hidden(X7,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X8))))|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_1359, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|r1_xxreal_0(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_matrix_1(X3))|k1_matrix_1(X3)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1570, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,X7)=k3_matrix_1(u1_struct_0(X1),X4,X6,X7)|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X8)))|~r1_xxreal_0(np__1,X7)|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_1567, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,X7)=k3_matrix_1(u1_struct_0(X1),X4,X6,X7)|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X8)))|~r1_xxreal_0(np__1,X7)|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_1681, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|r1_xxreal_0(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_matrix_1(esk6_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1613, plain, (v7_ordinal1(esk12_1(k3_finseq_2(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1462, plain, (v1_relat_1(esk12_1(k4_finseq_2(X1,X2)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1485, plain, (r1_xxreal_0(k1_xboole_0,np__1)|r1_xxreal_0(np__1,k1_xboole_0))).#trainneg
cnf(i_0_454, plain, (v1_xboole_0(esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))|~v1_xboole_0(X3)|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_652, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_630, negated_conjecture, (m2_subset_1(k3_finseq_1(esk6_0),k1_numbers,k5_numbers)|~v1_finseq_1(esk7_0)|~v1_funct_1(esk7_0)|~v1_relat_1(esk7_0))).#trainneg
cnf(i_0_756, negated_conjecture, (r1_xxreal_0(esk4_0,esk12_1(k5_numbers))|r1_xxreal_0(esk12_1(k5_numbers),esk4_0))).#trainneg
cnf(i_0_547, plain, (v1_card_1(k3_finseq_1(X4))|v6_struct_0(X3)|v2_struct_0(X3)|~v1_finseq_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~v1_funct_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~v1_relat_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1612, plain, (X1=esk12_1(k3_finseq_2(X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_455, plain, (v1_xboole_0(esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))|~v1_xboole_0(X2)|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_457, plain, (m1_subset_1(X1,esk13_2(X2,k1_zfmisc_1(X3)))|v1_xboole_0(X3)|v1_xboole_0(esk13_2(X2,k1_zfmisc_1(X3)))|~m2_subset_1(X1,X3,esk13_2(X2,k1_zfmisc_1(X3)))|~r2_hidden(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_1172, negated_conjecture, (m2_finseq_2(k8_matrix_1(X1,X2,X3),X1,k4_finseq_2(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4)),X1))|v1_xboole_0(X1)|~m1_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X3)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_1291, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,X4,X5,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,X4,X5)|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v1_matrix_1(X4)|~r2_hidden(X2,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6))))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_1365, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,X4,X5,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,X4,X5)|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v1_matrix_1(X4)|~r2_hidden(X2,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6))))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_1463, plain, (v1_xboole_0(esk12_1(k4_finseq_2(X1,X2)))|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1053, plain, (m2_finseq_1(esk13_2(X1,k3_finseq_2(X2)),X2)|~r2_hidden(X1,k3_finseq_2(X2)))).#trainneg
cnf(i_0_815, negated_conjecture, (r1_xxreal_0(esk4_0,k1_matrix_1(esk6_0))|r1_xxreal_0(k1_matrix_1(esk6_0),esk4_0))).#trainneg
cnf(i_0_459, plain, (m2_subset_1(X1,X2,esk13_2(X3,k1_zfmisc_1(X2)))|v1_xboole_0(esk13_2(X3,k1_zfmisc_1(X2)))|v1_xboole_0(X2)|~m1_subset_1(X1,esk13_2(X3,k1_zfmisc_1(X2)))|~r2_hidden(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1274, plain, (m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_1757, plain, (v1_card_1(k3_finseq_1(X1))|v6_struct_0(X2)|v1_xboole_0(u1_struct_0(X2))|~v1_funct_1(k1_matrix12(X3,X4,X2,X1,X5,X6))|~v1_relat_1(k1_matrix12(X3,X4,X2,X1,X5,X6))|~m1_matrix_1(X1,u1_struct_0(X2),X3,X4)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_907, plain, (r1_xxreal_0(X1,k2_zfmisc_1(X2,X3))|r1_xxreal_0(k2_zfmisc_1(X2,X3),X1)|~v7_ordinal1(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_460, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|v1_xboole_0(esk13_2(X3,k1_zfmisc_1(X2)))|~m2_subset_1(X1,X2,esk13_2(X3,k1_zfmisc_1(X2)))|~r2_hidden(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_474, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,esk16_7(X8,X9,X10,X4,X11,X12,X13))=k3_matrix_1(u1_struct_0(X1),X4,X7,esk16_7(X8,X9,X10,X4,X11,X12,X13))|X7=X5|X7=X6|v6_struct_0(X1)|v2_struct_0(X1)|X13=k1_matrix12(X8,X9,X10,X4,X11,X12)|v6_struct_0(X10)|v2_struct_0(X10)|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk16_7(X8,X9,X10,X4,X11,X12,X13))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|k3_finseq_1(X4)!=k3_finseq_1(X13)|~epred1_2(X9,X8)|~m1_matrix_1(X13,u1_struct_0(X10),X8,X9)|~m1_matrix_1(X4,u1_struct_0(X10),X8,X9)|~l6_algstr_0(X10)|~v5_vectsp_1(X10)|~v4_vectsp_1(X10)|~v4_rlvect_1(X10)|~v3_rlvect_1(X10)|~v2_rlvect_1(X10)|~v5_group_1(X10)|~v3_group_1(X10)|~v33_algstr_0(X10)|~v13_algstr_0(X10)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_1715, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|r1_xxreal_0(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1718, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|r1_xxreal_0(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1796, plain, (v7_ordinal1(esk12_1(k4_finseq_2(X1,X2)))|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1808, plain, (v1_funct_1(esk13_2(X1,k3_finseq_2(X2)))|~r2_hidden(X1,k3_finseq_2(X2)))).#trainneg
cnf(i_0_475, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,esk13_2(X8,k2_finseq_1(k1_matrix_1(X4))))=k3_matrix_1(u1_struct_0(X1),X4,X7,esk13_2(X8,k2_finseq_1(k1_matrix_1(X4))))|X7=X5|X7=X6|v6_struct_0(X1)|v2_struct_0(X1)|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk13_2(X8,k2_finseq_1(k1_matrix_1(X4))))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X8,k2_finseq_1(k1_matrix_1(X4))))).#trainneg
cnf(i_0_996, plain, (r1_xxreal_0(X1,k2_zfmisc_1(X2,X3))|r1_xxreal_0(k2_zfmisc_1(X2,X3),X1)|~v7_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1795, plain, (X1=esk12_1(k4_finseq_2(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_476, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,esk9_4(X8,X4,X9,X10))=k3_matrix_1(u1_struct_0(X1),X4,X7,esk9_4(X8,X4,X9,X10))|X7=X5|X7=X6|v6_struct_0(X1)|v2_struct_0(X1)|X10=k6_matrix_1(X8,X4,X9)|v1_xboole_0(X8)|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk9_4(X8,X4,X9,X10))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|k1_matrix_1(X4)!=k3_finseq_1(X10)|~m2_finseq_1(X4,k3_finseq_2(X8))|~m2_finseq_1(X10,X8)|~v1_matrix_1(X4)|~v7_ordinal1(X9))).#trainneg
cnf(i_0_477, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8)=k3_matrix_1(u1_struct_0(X1),X4,X7,X8)|X7=X5|X7=X6|v6_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(k2_finseq_1(k1_matrix_1(X4)))|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X8)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X8,k2_finseq_1(k1_matrix_1(X4))))).#trainneg
cnf(i_0_1402, negated_conjecture, (m2_finseq_1(k8_matrix_1(X1,X2,X3),X1)|v1_xboole_0(X1)|~m1_finseq_2(k4_finseq_2(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4)),X1),X1)|~m1_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~v7_ordinal1(X3)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_1327, negated_conjecture, (m2_finseq_1(k8_matrix_1(X1,X2,X3),X1)|v1_xboole_0(X1)|~m1_finseq_2(k4_finseq_2(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X4)),X1),X1)|~m1_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~v7_ordinal1(X3)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_1945, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|r1_xxreal_0(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_matrix_1(X3))|k1_matrix_1(X3)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_1941, negated_conjecture, (m2_finseq_1(k8_matrix_1(X1,X2,X3),X1)|v1_xboole_0(X1)|~m1_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X4))))).#trainneg
cnf(i_0_1959, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X7)|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_1809, plain, (v1_finseq_1(esk13_2(X1,k3_finseq_2(X2)))|~r2_hidden(X1,k3_finseq_2(X2)))).#trainneg
cnf(i_0_1810, plain, (v1_relat_1(esk13_2(X1,k3_finseq_2(X2)))|~r2_hidden(X1,k3_finseq_2(X2)))).#trainneg
cnf(i_0_816, negated_conjecture, (r1_xxreal_0(esk8_0,k1_matrix_1(esk6_0))|r1_xxreal_0(k1_matrix_1(esk6_0),esk8_0))).#trainneg
cnf(i_0_492, plain, (r1_xxreal_0(np__1,esk13_2(X1,a_1_0_finseq_1(X2)))|~r2_hidden(esk13_2(X1,a_1_0_finseq_1(X2)),a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2)))).#trainneg
cnf(i_0_826, negated_conjecture, (r1_xxreal_0(esk3_0,esk12_1(k5_numbers))|r1_xxreal_0(esk12_1(k5_numbers),esk3_0))).#trainneg
cnf(i_0_835, negated_conjecture, (r1_xxreal_0(esk3_0,k1_matrix_1(esk6_0))|r1_xxreal_0(k1_matrix_1(esk6_0),esk3_0))).#trainneg
cnf(i_0_837, negated_conjecture, (r1_xxreal_0(esk2_0,esk12_1(k5_numbers))|r1_xxreal_0(esk12_1(k5_numbers),esk2_0))).#trainneg
cnf(i_0_1830, plain, (v1_card_1(k3_finseq_1(k1_matrix12(X1,X2,X3,X4,X5,X6)))|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|v2_struct_0(X3)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X7)|~v7_ordinal1(X8)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1696, plain, (m2_finseq_1(k8_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6))|~v7_ordinal1(k1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6)))|~v7_ordinal1(X7)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_493, plain, (r1_xxreal_0(esk13_2(X1,a_1_0_finseq_1(X2)),X2)|~r2_hidden(esk13_2(X1,a_1_0_finseq_1(X2)),a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2)))).#trainneg
cnf(i_0_494, plain, (m2_subset_1(esk13_2(X1,a_1_0_finseq_1(X2)),k1_numbers,k5_numbers)|~r2_hidden(esk13_2(X1,a_1_0_finseq_1(X2)),a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2)))).#trainneg
cnf(i_0_1935, negated_conjecture, (m2_finseq_1(k8_matrix_1(X1,X2,X3),X1)|v1_xboole_0(X1)|~m1_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4))))).#trainneg
cnf(i_0_1952, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X5,X6)=k3_matrix_1(u1_struct_0(X1),esk7_0,X4,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X6)|~r2_hidden(X5,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1194, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X4,X6)=k3_matrix_1(u1_struct_0(X1),esk7_0,X5,X6)|v6_struct_0(X1)|~r2_hidden(X6,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X7))))|~r2_hidden(X4,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_1947, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|r1_xxreal_0(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1012, plain, (v3_ordinal1(X1)|~v3_ordinal1(a_1_0_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_587, plain, (m1_subset_1(esk13_2(X1,k2_finseq_1(X2)),k5_numbers)|~v7_ordinal1(X2)|~r2_hidden(X1,k2_finseq_1(X2)))).#trainneg
cnf(i_0_846, negated_conjecture, (r1_xxreal_0(esk2_0,k1_matrix_1(esk6_0))|r1_xxreal_0(k1_matrix_1(esk6_0),esk2_0))).#trainneg
cnf(i_0_495, plain, (esk11_2(esk13_2(X1,k2_finseq_1(X2)),X2)=esk13_2(X1,k2_finseq_1(X2))|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1434, negated_conjecture, (r1_xxreal_0(esk1_0,esk12_1(k5_numbers))|r1_xxreal_0(esk12_1(k5_numbers),esk1_0))).#trainneg
cnf(i_0_1446, negated_conjecture, (r1_xxreal_0(esk1_0,k1_matrix_1(esk6_0))|r1_xxreal_0(k1_matrix_1(esk6_0),esk1_0))).#trainneg
cnf(i_0_1474, plain, (r1_xxreal_0(k1_xboole_0,esk12_1(k5_numbers))|r1_xxreal_0(esk12_1(k5_numbers),k1_xboole_0))).#trainneg
cnf(i_0_2023, plain, (v1_card_1(k3_finseq_1(k1_matrix12(X1,X2,X3,X4,X5,X6)))|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X7)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_2039, plain, (v1_funct_1(k8_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6)))|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_498, plain, (X1=esk12_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_500, plain, (esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=k1_xboole_0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1986, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X8)))|~r1_xxreal_0(np__1,X7)|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_1989, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X8)))|~r1_xxreal_0(np__1,X7)|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_2063, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X5,X6)=k3_matrix_1(u1_struct_0(X1),esk7_0,X4,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7)))|~r1_xxreal_0(np__1,X6)|~r2_hidden(X5,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_2060, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X5,X6)=k3_matrix_1(u1_struct_0(X1),esk7_0,X4,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X7)))|~r1_xxreal_0(np__1,X6)|~r2_hidden(X5,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_2093, plain, (v7_ordinal1(esk13_2(X1,k2_finseq_1(X2)))|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2092, plain, (v3_ordinal1(esk13_2(X1,k2_finseq_1(X2)))|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1486, negated_conjecture, (r1_xxreal_0(k1_xboole_0,k1_matrix_1(esk6_0))|r1_xxreal_0(k1_matrix_1(esk6_0),k1_xboole_0))).#trainneg
cnf(i_0_533, plain, (esk11_2(esk16_7(X1,X2,X3,X4,X5,X6,X7),k1_matrix_1(X4))=esk16_7(X1,X2,X3,X4,X5,X6,X7)|X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|v2_struct_0(X3)|~v7_ordinal1(k1_matrix_1(X4))|k3_finseq_1(X4)!=k3_finseq_1(X7)|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1740, plain, (r1_xxreal_0(np__1,esk12_1(k5_numbers))|r1_xxreal_0(esk12_1(k5_numbers),np__1))).#trainneg
cnf(i_0_1754, negated_conjecture, (r1_xxreal_0(np__1,k1_matrix_1(esk6_0))|r1_xxreal_0(k1_matrix_1(esk6_0),np__1))).#trainneg
cnf(i_0_2040, plain, (v1_finseq_1(k8_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6)))|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_2154, plain, (v1_relat_1(k1_xboole_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_535, plain, (esk11_2(esk9_4(X1,X2,X3,X4),k1_matrix_1(X2))=esk9_4(X1,X2,X3,X4)|X4=k6_matrix_1(X1,X2,X3)|v1_xboole_0(X1)|~v7_ordinal1(k1_matrix_1(X2))|k1_matrix_1(X2)!=k3_finseq_1(X4)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m2_finseq_1(X4,X1)|~v1_matrix_1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1985, negated_conjecture, (m2_finseq_1(k8_matrix_1(X1,X2,X3),X1)|v1_xboole_0(X1)|~m1_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X3)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_1987, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k1_matrix_1(X8))|~r1_xxreal_0(np__1,X7)|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X8,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_2082, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X4,X6)=k3_matrix_1(u1_struct_0(X1),esk7_0,X5,X6)|v6_struct_0(X1)|~r2_hidden(X4,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X6,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X7)))|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X7))))).#trainneg
cnf(i_0_2061, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X5,X6)=k3_matrix_1(u1_struct_0(X1),esk7_0,X4,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k1_matrix_1(X7))|~r1_xxreal_0(np__1,X6)|~r2_hidden(X5,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X7,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_2091, plain, (v3_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(a_1_0_finseq_1(X2)))).#trainneg
cnf(i_0_2146, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_545, plain, (m2_subset_1(k3_finseq_1(X4),k1_numbers,k5_numbers)|v6_struct_0(X3)|v2_struct_0(X3)|~v1_finseq_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~v1_funct_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~v1_relat_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_813, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2118, plain, (v1_card_1(k3_finseq_1(k1_matrix12(X1,X2,X3,X4,X5,X6)))|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v1_xboole_0(X7))).#trainneg
cnf(i_0_546, plain, (k1_matrix12(X1,X2,X3,X4,X5,X6)=X7|r1_xxreal_0(esk10_2(k1_matrix12(X1,X2,X3,X4,X5,X6),X7),k3_finseq_1(X4))|v6_struct_0(X3)|v2_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(X7)|~v1_finseq_1(X7)|~v1_finseq_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X7)|~v1_funct_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~v1_relat_1(X7)|~v1_relat_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_573, plain, (m1_subset_1(esk13_2(X1,esk12_1(k1_zfmisc_1(X2))),X2)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_2262, negated_conjecture, (m2_finseq_1(k8_matrix_1(X1,X2,X3),X1)|v1_xboole_0(X1)|~v1_matrix_1(X2)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_1973, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|r1_xxreal_0(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2246, plain, (v1_relat_1(k1_xboole_0))).#trainneg
cnf(i_0_2375, plain, (v7_ordinal1(esk13_2(X1,esk12_1(k1_zfmisc_1(k5_numbers))))|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k5_numbers))))).#trainneg
cnf(i_0_574, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk12_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,esk12_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_586, plain, (m1_subset_1(esk16_7(X1,X2,X3,X4,X5,X6,X7),k5_numbers)|X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|v2_struct_0(X3)|~v7_ordinal1(k1_matrix_1(X4))|k3_finseq_1(X4)!=k3_finseq_1(X7)|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_588, plain, (m1_subset_1(esk9_4(X1,X2,X3,X4),k5_numbers)|X4=k6_matrix_1(X1,X2,X3)|v1_xboole_0(X1)|~v7_ordinal1(k1_matrix_1(X2))|k1_matrix_1(X2)!=k3_finseq_1(X4)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m2_finseq_1(X4,X1)|~v1_matrix_1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1141, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8)=k3_matrix_1(u1_struct_0(X1),X4,X7,X8)|X7=X5|X7=X6|v6_struct_0(X1)|~r2_hidden(X8,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X9))))|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X8)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X9))).#trainneg
cnf(i_0_2319, plain, (v7_ordinal1(esk12_1(k1_zfmisc_1(k1_xboole_0)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2149, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X1,k1_xboole_0)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_984, negated_conjecture, (k1_funct_1(X1,esk13_2(X2,k2_finseq_1(k1_matrix_1(esk6_0))))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,esk13_2(X2,k2_finseq_1(k1_matrix_1(esk6_0))))|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(esk6_0)))|~v7_ordinal1(esk13_2(X2,k2_finseq_1(k1_matrix_1(esk6_0))))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_589, plain, (m1_subset_1(X1,k5_numbers)|v1_xboole_0(k2_finseq_1(X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k2_finseq_1(X2)))).#trainneg
cnf(i_0_2320, plain, (esk12_1(k1_zfmisc_1(k1_xboole_0))=k1_xboole_0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_889, negated_conjecture, (~v1_xboole_0(k3_finseq_2(u1_struct_0(esk5_0)))|~r2_hidden(X1,esk7_0))).#trainneg
cnf(i_0_890, negated_conjecture, (~v1_xboole_0(k3_finseq_2(u1_struct_0(esk5_0)))|~r2_hidden(X1,esk6_0))).#trainneg
cnf(i_0_592, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(X2)|~r2_hidden(X1,k6_matrix_1(X2,X3,X4))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_593, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k3_finseq_2(u1_struct_0(X2))))|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|~r2_hidden(X1,k1_matrix12(X3,X4,X2,X5,X6,X7))|~m1_matrix_1(X5,u1_struct_0(X2),X3,X4)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X7)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_1150, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X5,X6)=k3_matrix_1(u1_struct_0(X1),esk7_0,X4,X6)|v6_struct_0(X1)|~r2_hidden(X6,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7))))|~r2_hidden(X5,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_1600, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X4,X6)=k3_matrix_1(u1_struct_0(X1),esk7_0,X5,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X6)|~r2_hidden(X4,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1127, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X4,X6)=k3_matrix_1(u1_struct_0(X1),esk6_0,X5,X6)|v6_struct_0(X1)|~r2_hidden(X6,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7))))|~r2_hidden(X4,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_2326, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_2308, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(k1_xboole_0))|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_603, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,esk13_2(X7,k2_finseq_1(k1_matrix_1(X4))))=k3_matrix_1(u1_struct_0(X1),X4,X6,esk13_2(X7,k2_finseq_1(k1_matrix_1(X4))))|v6_struct_0(X1)|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(esk13_2(X7,k2_finseq_1(k1_matrix_1(X4))))|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X7,k2_finseq_1(k1_matrix_1(X4))))).#trainneg
cnf(i_0_2316, plain, (v1_relat_1(esk12_1(k1_zfmisc_1(k1_xboole_0)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2435, plain, (v7_ordinal1(esk12_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_2521, plain, (v1_relat_1(esk12_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_602, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,esk16_7(X7,X8,X9,X4,X10,X11,X12))=k3_matrix_1(u1_struct_0(X1),X4,X6,esk16_7(X7,X8,X9,X4,X10,X11,X12))|v6_struct_0(X1)|X12=k1_matrix12(X7,X8,X9,X4,X10,X11)|v6_struct_0(X9)|v2_struct_0(X9)|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(esk16_7(X7,X8,X9,X4,X10,X11,X12))|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|k3_finseq_1(X4)!=k3_finseq_1(X12)|~epred1_2(X8,X7)|~m1_matrix_1(X12,u1_struct_0(X9),X7,X8)|~m1_matrix_1(X4,u1_struct_0(X9),X7,X8)|~l6_algstr_0(X9)|~v5_vectsp_1(X9)|~v4_vectsp_1(X9)|~v4_rlvect_1(X9)|~v3_rlvect_1(X9)|~v2_rlvect_1(X9)|~v5_group_1(X9)|~v3_group_1(X9)|~v33_algstr_0(X9)|~v13_algstr_0(X9)|~v7_ordinal1(X11)|~v7_ordinal1(X10))).#trainneg
cnf(i_0_604, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,esk9_4(X7,X4,X8,X9))=k3_matrix_1(u1_struct_0(X1),X4,X6,esk9_4(X7,X4,X8,X9))|v6_struct_0(X1)|X9=k6_matrix_1(X7,X4,X8)|v1_xboole_0(X7)|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(esk9_4(X7,X4,X8,X9))|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|k1_matrix_1(X4)!=k3_finseq_1(X9)|~m2_finseq_1(X4,k3_finseq_2(X7))|~m2_finseq_1(X9,X7)|~v1_matrix_1(X4)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_2476, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X4,X6)=k3_matrix_1(u1_struct_0(X1),esk7_0,X5,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7)))|~r1_xxreal_0(np__1,X6)|~r2_hidden(X4,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_2473, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X4,X6)=k3_matrix_1(u1_struct_0(X1),esk7_0,X5,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X7)))|~r1_xxreal_0(np__1,X6)|~r2_hidden(X4,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_1395, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8)=k3_matrix_1(u1_struct_0(X1),X4,X7,X8)|X7=X5|X7=X6|v6_struct_0(X1)|v2_struct_0(X1)|~r2_hidden(X8,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X9))))|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X8)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X9))).#trainneg
cnf(i_0_2327, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_2452, plain, (esk12_1(k1_zfmisc_1(k1_xboole_0))=k1_xboole_0)).#trainneg
cnf(i_0_394, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_613, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,esk13_2(X7,k2_finseq_1(k1_matrix_1(X4))))=k3_matrix_1(u1_struct_0(X1),X4,X5,esk13_2(X7,k2_finseq_1(k1_matrix_1(X4))))|v6_struct_0(X1)|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(esk13_2(X7,k2_finseq_1(k1_matrix_1(X4))))|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X7,k2_finseq_1(k1_matrix_1(X4))))).#trainneg
cnf(i_0_612, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,esk16_7(X7,X8,X9,X4,X10,X11,X12))=k3_matrix_1(u1_struct_0(X1),X4,X5,esk16_7(X7,X8,X9,X4,X10,X11,X12))|v6_struct_0(X1)|X12=k1_matrix12(X7,X8,X9,X4,X10,X11)|v6_struct_0(X9)|v2_struct_0(X9)|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(esk16_7(X7,X8,X9,X4,X10,X11,X12))|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|k3_finseq_1(X4)!=k3_finseq_1(X12)|~epred1_2(X8,X7)|~m1_matrix_1(X12,u1_struct_0(X9),X7,X8)|~m1_matrix_1(X4,u1_struct_0(X9),X7,X8)|~l6_algstr_0(X9)|~v5_vectsp_1(X9)|~v4_vectsp_1(X9)|~v4_rlvect_1(X9)|~v3_rlvect_1(X9)|~v2_rlvect_1(X9)|~v5_group_1(X9)|~v3_group_1(X9)|~v33_algstr_0(X9)|~v13_algstr_0(X9)|~v7_ordinal1(X11)|~v7_ordinal1(X10))).#trainneg
cnf(i_0_614, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,esk9_4(X7,X4,X8,X9))=k3_matrix_1(u1_struct_0(X1),X4,X5,esk9_4(X7,X4,X8,X9))|v6_struct_0(X1)|X9=k6_matrix_1(X7,X4,X8)|v1_xboole_0(X7)|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(esk9_4(X7,X4,X8,X9))|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|k1_matrix_1(X4)!=k3_finseq_1(X9)|~m2_finseq_1(X4,k3_finseq_2(X7))|~m2_finseq_1(X9,X7)|~v1_matrix_1(X4)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_1196, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X5,X6)=k3_matrix_1(u1_struct_0(X1),esk7_0,X4,X6)|v6_struct_0(X1)|~r2_hidden(X6,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X7))))|~r2_hidden(X5,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_1149, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X6,X7)=k3_matrix_1(u1_struct_0(X1),esk7_0,X6,X7)|X6=X5|X6=X4|v6_struct_0(X1)|~r2_hidden(X7,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X8))))|~r2_hidden(X6,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_2311, plain, (v3_ordinal1(X1)|~m2_finseq_1(X1,X2)|~v7_ordinal1(k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2008, plain, (r1_xxreal_0(np__1,esk13_2(X1,k2_finseq_1(X2)))|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1614, plain, (esk12_1(k3_finseq_2(X1))=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_643, plain, (r1_xxreal_0(X1,esk13_2(X2,k5_numbers))|r1_xxreal_0(esk13_2(X2,k5_numbers),X1)|~v7_ordinal1(X1)|~r2_hidden(X2,k5_numbers))).#trainneg
cnf(i_0_644, plain, (r1_xxreal_0(X1,esk12_1(k5_numbers))|r1_xxreal_0(esk12_1(k5_numbers),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_645, plain, (r1_xxreal_0(X1,esk10_2(X2,X3))|r1_xxreal_0(esk10_2(X2,X3),X1)|X2=X3|~v7_ordinal1(X1)|k3_finseq_1(X2)!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_finseq_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1568, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,X7)=k3_matrix_1(u1_struct_0(X1),X4,X6,X7)|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k1_matrix_1(X8))|~r1_xxreal_0(np__1,X7)|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X8,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_1173, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X4,X6)=k3_matrix_1(u1_struct_0(X1),esk6_0,X5,X6)|v6_struct_0(X1)|~r2_hidden(X6,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X7))))|~r2_hidden(X4,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_2632, plain, (v1_funct_1(k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2650, plain, (r1_xxreal_0(esk12_1(k5_numbers),esk12_1(k5_numbers)))).#trainneg
cnf(i_0_2577, plain, (~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_654, plain, (X1=esk12_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2633, plain, (v1_finseq_1(k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2623, negated_conjecture, (v3_ordinal1(esk7_0)|~v1_xboole_0(k3_finseq_2(u1_struct_0(esk5_0)))|~v7_ordinal1(k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_2624, negated_conjecture, (v3_ordinal1(esk6_0)|~v1_xboole_0(k3_finseq_2(u1_struct_0(esk5_0)))|~v7_ordinal1(k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_2306, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_656, plain, (esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_690, plain, (m1_subset_1(esk15_7(X1,X2,X3,X4,X5,X6,X7),k5_numbers)|X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|v2_struct_0(X3)|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4)|k3_finseq_1(X4)!=k3_finseq_1(X7)|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1273, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8)=k3_matrix_1(u1_struct_0(X1),X4,X7,X8)|X7=X6|X7=X5|v6_struct_0(X1)|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X8)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_subset_1(X8,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X8)|~r1_xxreal_0(X8,k1_matrix_1(esk6_0)))).#trainneg
cnf(i_0_1955, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X5,X6)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7)))|~r1_xxreal_0(np__1,X6)|~r2_hidden(X5,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7)))|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_1954, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X5,X6)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X7)))|~r1_xxreal_0(np__1,X6)|~r2_hidden(X5,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X7)))|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_2638, plain, (m2_finseq_1(k1_xboole_0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2706, plain, (v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_691, plain, (m1_subset_1(esk13_2(X1,k4_finseq_1(X2)),k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k4_finseq_1(X2)))).#trainneg
cnf(i_0_2726, plain, (v1_finseq_1(k1_xboole_0))).#trainneg
cnf(i_0_2732, negated_conjecture, (m1_subset_1(esk7_0,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(k3_finseq_2(u1_struct_0(esk5_0))))).#trainneg
cnf(i_0_2733, negated_conjecture, (m1_subset_1(esk6_0,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(k3_finseq_2(u1_struct_0(esk5_0))))).#trainneg
cnf(i_0_745, negated_conjecture, (r1_xxreal_0(esk8_0,esk13_2(X1,k5_numbers))|r1_xxreal_0(esk13_2(X1,k5_numbers),esk8_0)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_747, negated_conjecture, (r1_xxreal_0(esk8_0,esk10_2(X1,X2))|r1_xxreal_0(esk10_2(X1,X2),esk8_0)|X1=X2|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2760, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8)=k3_matrix_1(u1_struct_0(X1),X4,X7,X8)|X7=X5|X7=X6|v6_struct_0(X1)|~m2_subset_1(X8,k1_numbers,k5_numbers)|~r1_xxreal_0(X8,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X9)))|~r1_xxreal_0(np__1,X8)|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X8)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X9))).#trainneg
cnf(i_0_2757, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8)=k3_matrix_1(u1_struct_0(X1),X4,X7,X8)|X7=X5|X7=X6|v6_struct_0(X1)|~m2_subset_1(X8,k1_numbers,k5_numbers)|~r1_xxreal_0(X8,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X9)))|~r1_xxreal_0(np__1,X8)|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X8)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X9))).#trainneg
cnf(i_0_1195, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X6,X7)=k3_matrix_1(u1_struct_0(X1),esk7_0,X6,X7)|X6=X5|X6=X4|v6_struct_0(X1)|~r2_hidden(X7,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X8))))|~r2_hidden(X6,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_2474, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X4,X6)=k3_matrix_1(u1_struct_0(X1),esk7_0,X5,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k1_matrix_1(X7))|~r1_xxreal_0(np__1,X6)|~r2_hidden(X4,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X7,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_755, negated_conjecture, (r1_xxreal_0(esk4_0,esk13_2(X1,k5_numbers))|r1_xxreal_0(esk13_2(X1,k5_numbers),esk4_0)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_2041, plain, (v1_relat_1(k8_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6)))|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1580, plain, (k1_matrix_1(k1_matrix12(X1,X2,X3,X4,X5,X6))=k3_finseq_1(X7)|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|X7!=k8_matrix_1(u1_struct_0(X3),k1_matrix12(X1,X2,X3,X4,X5,X6),X8)|~m2_finseq_1(X7,u1_struct_0(X3))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X8)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_757, negated_conjecture, (r1_xxreal_0(esk4_0,esk10_2(X1,X2))|r1_xxreal_0(esk10_2(X1,X2),esk4_0)|X1=X2|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_825, negated_conjecture, (r1_xxreal_0(esk3_0,esk13_2(X1,k5_numbers))|r1_xxreal_0(esk13_2(X1,k5_numbers),esk3_0)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_1122, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,esk6_0,X4,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,esk6_0,X4)|~m2_finseq_1(esk6_0,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~r2_hidden(X2,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5))))|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1129, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X5,X6)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X6)|v6_struct_0(X1)|~r2_hidden(X6,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7))))|~r2_hidden(X5,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_1128, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X6,X7)=k3_matrix_1(u1_struct_0(X1),esk6_0,X6,X7)|X6=X4|X6=X5|v6_struct_0(X1)|v2_struct_0(X1)|~r2_hidden(X7,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X8))))|~r2_hidden(X6,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_2630, plain, (m1_subset_1(k1_xboole_0,k3_finseq_2(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_827, negated_conjecture, (r1_xxreal_0(esk3_0,esk10_2(X1,X2))|r1_xxreal_0(esk10_2(X1,X2),esk3_0)|X1=X2|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2664, negated_conjecture, (r1_xxreal_0(esk12_1(k5_numbers),k1_matrix_1(esk6_0))|r1_xxreal_0(k1_matrix_1(esk6_0),esk12_1(k5_numbers)))).#trainneg
cnf(i_0_1246, plain, (esk15_7(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8),X9,X10,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8))=X9|esk15_7(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8),X9,X10,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8))=X10|k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8),X9,X10)=k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8)|v6_struct_0(X3)|v2_struct_0(X3)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X10)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_2120, plain, (v1_card_1(k3_finseq_1(k1_matrix12(X1,X2,X3,X4,X5,X6)))|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_836, negated_conjecture, (r1_xxreal_0(esk2_0,esk13_2(X1,k5_numbers))|r1_xxreal_0(esk13_2(X1,k5_numbers),esk2_0)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_838, negated_conjecture, (r1_xxreal_0(esk2_0,esk10_2(X1,X2))|r1_xxreal_0(esk10_2(X1,X2),esk2_0)|X1=X2|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1260, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8)=k3_matrix_1(u1_struct_0(X1),X4,X7,X8)|X7=X5|X7=X6|v6_struct_0(X1)|v2_struct_0(X1)|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X8)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_subset_1(X8,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X8)|~r1_xxreal_0(X8,k1_matrix_1(X4))|~v7_ordinal1(k1_matrix_1(X4)))).#trainneg
cnf(i_0_2495, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X4,X6)=k3_matrix_1(u1_struct_0(X1),esk6_0,X5,X6)|v6_struct_0(X1)|~r2_hidden(X4,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X6,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7)))|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7))))).#trainneg
cnf(i_0_2904, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X6,X7)=k3_matrix_1(u1_struct_0(X1),esk7_0,X6,X7)|X6=X5|X6=X4|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X7)|~r2_hidden(X6,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1967, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,k1_matrix_1(X4))=k3_matrix_1(u1_struct_0(X1),X4,X5,k1_matrix_1(X4))|v6_struct_0(X1)|~m2_subset_1(k1_matrix_1(X4),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(X4))|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X4))|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2731, plain, (m1_subset_1(esk12_1(k3_finseq_2(X1)),k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_910, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1826, negated_conjecture, (m1_subset_1(k3_finseq_1(esk6_0),k5_numbers)|~r1_xxreal_0(np__1,k3_finseq_1(esk6_0))|~r1_xxreal_0(k3_finseq_1(esk6_0),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2852, plain, (k1_matrix_1(k1_matrix12(X1,X2,X3,X4,X5,X6))=k3_finseq_1(k8_matrix_1(u1_struct_0(X3),k1_matrix12(X1,X2,X3,X4,X5,X6),X7))|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|~m2_finseq_1(k8_matrix_1(u1_struct_0(X3),k1_matrix12(X1,X2,X3,X4,X5,X6),X7),u1_struct_0(X3))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X7)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1953, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X5,X6)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k1_matrix_1(X7))|~r1_xxreal_0(np__1,X6)|~r2_hidden(X5,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X7))|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X7,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_1611, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,k1_matrix_1(X4))=k3_matrix_1(u1_struct_0(X1),X4,X6,k1_matrix_1(X4))|v6_struct_0(X1)|~m2_subset_1(k1_matrix_1(X4),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(X4))|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X4))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1168, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,esk6_0,X4,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,esk6_0,X4)|~m2_finseq_1(esk6_0,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~r2_hidden(X2,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5))))|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1175, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X5,X6)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X6)|v6_struct_0(X1)|~r2_hidden(X6,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X7))))|~r2_hidden(X5,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_2312, plain, (v3_ordinal1(X1)|~m2_finseq_1(X1,X2)|~v7_ordinal1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2739, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2584, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_851, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(X3)|~m2_finseq_1(X3,X2)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_2900, plain, (esk15_7(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8),X9,X10),X11,X12,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8),X9,X10))=X12|esk15_7(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8),X9,X10),X11,X12,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8),X9,X10))=X11|k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8),X9,X10),X11,X12)=k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8),X9,X10)|v6_struct_0(X3)|v2_struct_0(X3)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v7_ordinal1(X10)|~v7_ordinal1(X9)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_2812, plain, (v7_ordinal1(esk13_2(X1,k4_finseq_1(X2)))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k4_finseq_1(X2)))).#trainneg
cnf(i_0_878, plain, (m1_subset_1(k3_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4,X5),X1)|v1_xboole_0(k3_finseq_2(X1))|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(X1)))|~v1_matrix_1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_885, plain, (m1_subset_1(k8_matrix_1(X1,X2,X3),k4_finseq_2(k1_matrix_1(X2),X1))|v1_xboole_0(X1)|~m1_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(k1_matrix_1(X2)))).#trainneg
cnf(i_0_2951, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X6,X7)=k3_matrix_1(u1_struct_0(X1),esk7_0,X6,X7)|X6=X4|X6=X5|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X8)))|~r1_xxreal_0(np__1,X7)|~r2_hidden(X6,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_2948, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X6,X7)=k3_matrix_1(u1_struct_0(X1),esk7_0,X6,X7)|X6=X4|X6=X5|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X8)))|~r1_xxreal_0(np__1,X7)|~r2_hidden(X6,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_3056, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,X4,X5,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,X4,X5)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6)))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v1_matrix_1(X4)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_886, plain, (k1_funct_1(k6_matrix_1(X1,X2,X3),esk13_2(X4,k2_finseq_1(k1_matrix_1(X2))))=k3_matrix_1(X1,X2,X3,esk13_2(X4,k2_finseq_1(k1_matrix_1(X2))))|v1_xboole_0(X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m2_finseq_1(k6_matrix_1(X1,X2,X3),X1)|~v1_matrix_1(X2)|~r2_hidden(X4,k2_finseq_1(k1_matrix_1(X2)))|~v7_ordinal1(esk13_2(X4,k2_finseq_1(k1_matrix_1(X2))))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2975, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,k1_matrix_1(esk6_0))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,k1_matrix_1(esk6_0))|~m2_subset_1(k1_matrix_1(esk6_0),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(esk6_0)))).#trainneg
cnf(i_0_3035, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,k1_matrix_1(esk6_0))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,k1_matrix_1(esk6_0))|~m2_subset_1(k1_matrix_1(esk6_0),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(esk6_0)))).#trainneg
cnf(i_0_3104, plain, (r1_xxreal_0(X1,esk13_2(X2,k4_finseq_1(X3)))|r1_xxreal_0(esk13_2(X2,k4_finseq_1(X3)),X1)|~v7_ordinal1(X1)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(X2,k4_finseq_1(X3)))).#trainneg
cnf(i_0_903, negated_conjecture, (r1_xxreal_0(esk2_0,k2_zfmisc_1(X1,X2))|r1_xxreal_0(k2_zfmisc_1(X1,X2),esk2_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_904, negated_conjecture, (r1_xxreal_0(esk3_0,k2_zfmisc_1(X1,X2))|r1_xxreal_0(k2_zfmisc_1(X1,X2),esk3_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3023, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X5,X6)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k1_matrix_1(X7))|~r1_xxreal_0(np__1,X6)|~m1_matrix_1(X7,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X7))|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X5,k4_finseq_1(esk6_0)))).#trainneg
cnf(i_0_3054, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,X4,X5,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,X4,X5)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(X6))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v1_matrix_1(X4)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~m1_matrix_1(X6,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_2949, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X6,X7)=k3_matrix_1(u1_struct_0(X1),esk7_0,X6,X7)|X6=X4|X6=X5|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k1_matrix_1(X8))|~r1_xxreal_0(np__1,X7)|~r2_hidden(X6,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X8,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_3031, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,k1_matrix_1(X7))=k3_matrix_1(u1_struct_0(X1),X4,X6,k1_matrix_1(X7))|v6_struct_0(X1)|~m2_subset_1(k1_matrix_1(X7),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(X7))|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X7))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X7,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_905, negated_conjecture, (r1_xxreal_0(esk4_0,k2_zfmisc_1(X1,X2))|r1_xxreal_0(k2_zfmisc_1(X1,X2),esk4_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_906, negated_conjecture, (r1_xxreal_0(esk8_0,k2_zfmisc_1(X1,X2))|r1_xxreal_0(k2_zfmisc_1(X1,X2),esk8_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3026, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X4,k1_matrix_1(X6))=k3_matrix_1(u1_struct_0(X1),esk6_0,X5,k1_matrix_1(X6))|v6_struct_0(X1)|~m2_subset_1(k1_matrix_1(X6),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(X6))|~r2_hidden(X4,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X6))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X6,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_3022, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk3_0),esk3_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k1_matrix_1(X6))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(X6,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X6))|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3021, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk4_0),esk4_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k1_matrix_1(X6))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(X6,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X6))|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3202, plain, (r1_xxreal_0(esk13_2(X1,k4_finseq_1(X2)),X3)|r1_xxreal_0(X3,esk13_2(X1,k4_finseq_1(X2)))|v1_xboole_0(k4_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X1,k4_finseq_1(X2)))).#trainneg
cnf(i_0_2811, plain, (v3_ordinal1(esk13_2(X1,k4_finseq_1(X2)))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k4_finseq_1(X2)))).#trainneg
cnf(i_0_919, plain, (k1_funct_1(X1,esk12_1(k2_finseq_1(k1_matrix_1(X2))))=k3_matrix_1(X3,X2,X4,esk12_1(k2_finseq_1(k1_matrix_1(X2))))|v1_xboole_0(k2_finseq_1(k1_matrix_1(X2)))|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,X2,X4)|~m2_finseq_1(X2,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v1_matrix_1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(esk12_1(k2_finseq_1(k1_matrix_1(X2)))))).#trainneg
cnf(i_0_920, plain, (v7_ordinal1(k1_matrix_1(k1_matrix12(X1,X2,X3,X4,X5,X6)))|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|~v1_finseq_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~v1_funct_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~v1_relat_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1174, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X6,X7)=k3_matrix_1(u1_struct_0(X1),esk6_0,X6,X7)|X6=X4|X6=X5|v6_struct_0(X1)|v2_struct_0(X1)|~r2_hidden(X7,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X8))))|~r2_hidden(X6,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_2699, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X4,X6)=k3_matrix_1(u1_struct_0(X1),esk6_0,X5,X6)|v6_struct_0(X1)|~r2_hidden(X4,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X6,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X7)))|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X7))))).#trainneg
cnf(i_0_3020, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk8_0),esk8_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k1_matrix_1(X6))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(X6,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X6))|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2878, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X6,X7)=k3_matrix_1(u1_struct_0(X1),esk6_0,X6,X7)|X6=X5|X6=X4|v6_struct_0(X1)|~r2_hidden(X6,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X8)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(X7,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X8)))|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X8))))).#trainneg
cnf(i_0_2622, plain, (v3_ordinal1(esk12_1(k3_finseq_2(X1)))|~v1_xboole_0(X1)|~v7_ordinal1(k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_925, plain, (k1_matrix_1(k1_matrix12(X1,X2,X3,X4,X5,X6))=k1_matrix_1(k1_matrix12(X1,X2,X3,X7,X8,X9))|v6_struct_0(X3)|v2_struct_0(X3)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_3349, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk8_0),esk8_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3264, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk3_0),esk3_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3282, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk4_0),esk4_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3336, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X6,X7)=k3_matrix_1(u1_struct_0(X1),esk6_0,X6,X7)|X6=X5|X6=X4|v6_struct_0(X1)|~r2_hidden(X6,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X8)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(X7,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X8)))|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X8))))).#trainneg
cnf(i_0_3352, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk8_0),esk8_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6)))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6)))|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_3285, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk4_0),esk4_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6)))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6)))|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_3267, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk3_0),esk3_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6)))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6)))|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_1137, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,X1)|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2730, plain, (m1_subset_1(esk12_1(k4_finseq_2(X1,X2)),k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3081, plain, (v3_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(k1_numbers))).#trainneg
cnf(i_0_928, plain, (esk11_2(esk12_1(k2_finseq_1(X1)),X1)=esk12_1(k2_finseq_1(X1))|v1_xboole_0(k2_finseq_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3387, plain, (v3_ordinal1(k1_matrix_1(k1_matrix12(X1,X2,X3,X7,X8,X9)))|v6_struct_0(X3)|~v1_finseq_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~v1_funct_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~v1_relat_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~v1_matrix_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X9)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_949, plain, (v3_ordinal1(esk9_4(X1,X2,X3,X4))|X4=k6_matrix_1(X1,X2,X3)|v1_xboole_0(X1)|~v3_ordinal1(k2_finseq_1(k1_matrix_1(X2)))|k1_matrix_1(X2)!=k3_finseq_1(X4)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m2_finseq_1(X4,X1)|~v1_matrix_1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1368, negated_conjecture, (v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1136, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,X1)|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2971, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,k1_matrix_1(X7))=k3_matrix_1(u1_struct_0(X1),X4,X5,k1_matrix_1(X7))|v6_struct_0(X1)|~m2_subset_1(k1_matrix_1(X7),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(X7))|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X7))|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X7,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_2052, plain, (r1_xxreal_0(esk13_2(X1,k2_finseq_1(X2)),X2)|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3080, plain, (v3_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_874, negated_conjecture, (v1_card_1(k1_matrix_1(esk6_0))|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_990, negated_conjecture, (v1_card_1(k1_matrix_1(esk6_0))|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3626, plain, (v3_ordinal1(k1_matrix_1(k1_matrix12(X1,X2,X3,X4,X5,X6)))|v6_struct_0(X3)|v1_xboole_0(u1_struct_0(X3))|~v1_finseq_1(k1_matrix12(X1,X2,X3,X7,X8,X9))|~v1_funct_1(k1_matrix12(X1,X2,X3,X7,X8,X9))|~v1_relat_1(k1_matrix12(X1,X2,X3,X7,X8,X9))|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_959, plain, (k1_funct_1(k6_matrix_1(X1,X2,X3),esk9_4(X4,X2,X5,X6))=k3_matrix_1(X1,X2,X3,esk9_4(X4,X2,X5,X6))|X6=k6_matrix_1(X4,X2,X5)|v1_xboole_0(X4)|v1_xboole_0(X1)|k1_matrix_1(X2)!=k3_finseq_1(X6)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X4))|~m2_finseq_1(k6_matrix_1(X1,X2,X3),X1)|~m2_finseq_1(X6,X4)|~v1_matrix_1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_962, negated_conjecture, (r1_xxreal_0(esk2_0,esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)))|r1_xxreal_0(esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)),esk2_0)|k1_matrix12(X1,X2,X3,X7,X8,X9)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(k1_matrix12(X1,X2,X3,X7,X8,X9))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3033, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7)))=k3_matrix_1(u1_struct_0(X1),X4,X6,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7)))|v6_struct_0(X1)|~m2_subset_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7)))|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7)))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_3351, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk8_0),esk8_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6)))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6)))|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_2921, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,k1_matrix_1(X4))=k3_matrix_1(u1_struct_0(X1),X4,X7,k1_matrix_1(X4))|X7=X6|X7=X5|v6_struct_0(X1)|~m2_subset_1(k1_matrix_1(X4),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(X4))|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X4))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3750, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|~m2_subset_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3664, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,X1)|~epred1_2(esk1_0,esk2_0)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))))).#trainneg
cnf(i_0_3751, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|esk15_7(esk2_0,esk1_0,esk5_0,esk7_0,X1,X2,esk7_0)=X1|esk15_7(esk2_0,esk1_0,esk5_0,esk7_0,X1,X2,esk7_0)=X2|~m2_subset_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~r2_hidden(X1,k4_finseq_1(esk7_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_3266, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk3_0),esk3_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6)))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6)))|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_3609, plain, (m1_subset_1(X1,k1_xboole_0)|~r2_hidden(X1,esk12_1(k4_finseq_2(X2,X3)))|~v1_xboole_0(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1811, plain, (v1_xboole_0(esk13_2(X1,k3_finseq_2(X2)))|~v1_xboole_0(X2)|~r2_hidden(X1,k3_finseq_2(X2)))).#trainneg
cnf(i_0_3618, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers)|~v7_ordinal1(k1_numbers))).#trainneg
cnf(i_0_963, negated_conjecture, (r1_xxreal_0(esk3_0,esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)))|r1_xxreal_0(esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)),esk3_0)|k1_matrix12(X1,X2,X3,X7,X8,X9)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(k1_matrix12(X1,X2,X3,X7,X8,X9))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2329, plain, (v1_relat_1(esk13_2(X1,k1_zfmisc_1(k1_xboole_0)))|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_3731, negated_conjecture, (v1_card_1(k1_matrix_1(esk6_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3733, negated_conjecture, (v1_card_1(k1_matrix_1(esk6_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_964, negated_conjecture, (r1_xxreal_0(esk4_0,esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)))|r1_xxreal_0(esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)),esk4_0)|k1_matrix12(X1,X2,X3,X7,X8,X9)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(k1_matrix12(X1,X2,X3,X7,X8,X9))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_965, negated_conjecture, (r1_xxreal_0(esk8_0,esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)))|r1_xxreal_0(esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)),esk8_0)|k1_matrix12(X1,X2,X3,X7,X8,X9)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(k1_matrix12(X1,X2,X3,X7,X8,X9))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3284, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk4_0),esk4_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6)))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6)))|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_3030, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X7)))=k3_matrix_1(u1_struct_0(X1),X4,X6,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X7)))|v6_struct_0(X1)|~m2_subset_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X7)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X7)))|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X7)))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_2977, plain, (v3_ordinal1(esk12_1(k3_finseq_2(X1)))|~v3_ordinal1(k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2621, plain, (v3_ordinal1(esk12_1(k4_finseq_2(X1,X2)))|~v1_xboole_0(X2)|~v7_ordinal1(k1_zfmisc_1(k1_xboole_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1812, plain, (~v1_xboole_0(X2)|~r2_hidden(X3,esk13_2(X1,k3_finseq_2(X2)))|~r2_hidden(X1,k3_finseq_2(X2)))).#trainneg
cnf(i_0_966, plain, (r1_xxreal_0(X1,esk15_7(X2,X3,X4,X5,X6,X7,k1_matrix12(X2,X3,X4,X8,X9,X10)))|r1_xxreal_0(esk15_7(X2,X3,X4,X5,X6,X7,k1_matrix12(X2,X3,X4,X8,X9,X10)),X1)|k1_matrix12(X2,X3,X4,X8,X9,X10)=k1_matrix12(X2,X3,X4,X5,X6,X7)|v6_struct_0(X4)|~v7_ordinal1(X1)|k3_finseq_1(X5)!=k3_finseq_1(k1_matrix12(X2,X3,X4,X8,X9,X10))|~m1_matrix_1(X5,u1_struct_0(X4),X2,X3)|~m1_matrix_1(X8,u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X10)|~v7_ordinal1(X9)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2330, plain, (v1_xboole_0(esk13_2(X1,k1_zfmisc_1(k1_xboole_0)))|~v1_xboole_0(X3)|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_3881, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk12_1(k5_numbers))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk12_1(k5_numbers))|r1_xxreal_0(k1_matrix_1(esk6_0),esk12_1(k5_numbers))|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk12_1(k5_numbers),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk12_1(k5_numbers)))).#trainneg
cnf(i_0_3320, plain, (r1_xxreal_0(X1,k1_matrix_1(k1_matrix12(X2,X3,X4,X5,X6,X7)))|r1_xxreal_0(k1_matrix_1(k1_matrix12(X2,X3,X4,X5,X6,X7)),X1)|v1_xboole_0(u1_struct_0(X4))|v6_struct_0(X4)|~v7_ordinal1(X1)|~m1_matrix_1(X5,u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_3872, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,X1)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_matrix_1(X2))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_3821, plain, (v7_ordinal1(esk13_2(X1,k3_finseq_2(X2)))|~v1_xboole_0(X2)|~r2_hidden(X1,k3_finseq_2(X2)))).#trainneg
cnf(i_0_3847, plain, (v1_relat_1(esk13_2(X1,k1_zfmisc_1(k1_xboole_0)))|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_976, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,esk12_1(k2_finseq_1(k1_matrix_1(X4))))=k3_matrix_1(u1_struct_0(X1),X4,X6,esk12_1(k2_finseq_1(k1_matrix_1(X4))))|v1_xboole_0(k2_finseq_1(k1_matrix_1(X4)))|v6_struct_0(X1)|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(esk12_1(k2_finseq_1(k1_matrix_1(X4))))|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3851, negated_conjecture, (v1_card_1(k1_matrix_1(esk6_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3853, negated_conjecture, (v1_card_1(k1_matrix_1(esk6_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_992, negated_conjecture, (r1_xxreal_0(esk2_0,k2_zfmisc_1(X1,X2))|r1_xxreal_0(k2_zfmisc_1(X1,X2),esk2_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2966, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X5,k1_matrix_1(X6))=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,k1_matrix_1(X6))|v6_struct_0(X1)|~m2_subset_1(k1_matrix_1(X6),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(X6))|~r2_hidden(X5,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X6))|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X6,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_2293, plain, (v3_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3726, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers)|~v7_ordinal1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1464, plain, (~v1_xboole_0(X2)|~r2_hidden(X3,esk12_1(k4_finseq_2(X1,X2)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_993, negated_conjecture, (r1_xxreal_0(esk3_0,k2_zfmisc_1(X1,X2))|r1_xxreal_0(k2_zfmisc_1(X1,X2),esk3_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1161, negated_conjecture, (v1_card_1(k1_matrix_1(esk6_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4042, negated_conjecture, (v1_card_1(k1_matrix_1(esk6_0)))).#trainneg
cnf(i_0_994, negated_conjecture, (r1_xxreal_0(esk4_0,k2_zfmisc_1(X1,X2))|r1_xxreal_0(k2_zfmisc_1(X1,X2),esk4_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_995, negated_conjecture, (r1_xxreal_0(esk8_0,k2_zfmisc_1(X1,X2))|r1_xxreal_0(k2_zfmisc_1(X1,X2),esk8_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4016, negated_conjecture, (k1_funct_1(X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X4,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))!=k3_finseq_1(X3)|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X4)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(X4)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4017, negated_conjecture, (k1_funct_1(X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))!=k3_finseq_1(X3)|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(X4)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3923, plain, (v1_xboole_0(esk13_2(X1,k1_zfmisc_1(k1_xboole_0)))|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_997, negated_conjecture, (r1_xxreal_0(esk2_0,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|r1_xxreal_0(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),esk2_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3993, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk12_1(k2_finseq_1(k1_matrix_1(esk6_0))))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk12_1(k2_finseq_1(k1_matrix_1(esk6_0))))|v1_xboole_0(k2_finseq_1(k1_matrix_1(esk6_0)))|~v7_ordinal1(esk12_1(k2_finseq_1(k1_matrix_1(esk6_0)))))).#trainneg
cnf(i_0_3936, plain, (r1_xxreal_0(k1_matrix_1(k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8)),X9)|r1_xxreal_0(X9,k1_matrix_1(k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8)))|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|v2_struct_0(X3)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X9)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_998, negated_conjecture, (r1_xxreal_0(esk3_0,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|r1_xxreal_0(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),esk3_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_999, negated_conjecture, (r1_xxreal_0(esk4_0,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|r1_xxreal_0(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),esk4_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3946, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(X3)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1270, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,X1)|~epred1_2(esk1_0,esk2_0)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,k1_matrix_1(esk6_0)))).#trainneg
cnf(i_0_2089, plain, (v3_ordinal1(X1)|~v3_ordinal1(k2_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2054, plain, (m2_subset_1(esk13_2(X1,k2_finseq_1(X2)),k1_numbers,k5_numbers)|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1000, negated_conjecture, (r1_xxreal_0(esk8_0,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|r1_xxreal_0(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),esk8_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_872, negated_conjecture, (m2_subset_1(k1_matrix_1(esk6_0),k1_numbers,k5_numbers)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3384, plain, (m1_subset_1(k1_matrix_1(k1_matrix12(X1,X2,X3,X7,X8,X9)),k5_numbers)|v6_struct_0(X3)|~v1_finseq_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~v1_funct_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~v1_relat_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~v1_matrix_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X9)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_1001, plain, (r1_xxreal_0(X1,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3))))|r1_xxreal_0(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3))),X1)|~v7_ordinal1(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1005, negated_conjecture, (r1_xxreal_0(esk2_0,esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)))|r1_xxreal_0(esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)),esk2_0)|k1_matrix12(X1,X2,X3,X7,X8,X9)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(k1_matrix12(X1,X2,X3,X7,X8,X9))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3967, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1358, negated_conjecture, (m2_subset_1(k1_matrix_1(X2),k1_numbers,k5_numbers)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4081, plain, (v7_ordinal1(esk13_2(X1,k1_zfmisc_1(k1_xboole_0)))|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_1617, plain, (esk12_1(k3_finseq_2(X1))=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4136, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk12_1(k5_numbers))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk12_1(k5_numbers))|r1_xxreal_0(k1_matrix_1(esk6_0),esk12_1(k5_numbers))|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk12_1(k5_numbers),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk12_1(k5_numbers)))).#trainneg
cnf(i_0_1006, negated_conjecture, (r1_xxreal_0(esk3_0,esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)))|r1_xxreal_0(esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)),esk3_0)|k1_matrix12(X1,X2,X3,X7,X8,X9)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(k1_matrix12(X1,X2,X3,X7,X8,X9))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1007, negated_conjecture, (r1_xxreal_0(esk4_0,esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)))|r1_xxreal_0(esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)),esk4_0)|k1_matrix12(X1,X2,X3,X7,X8,X9)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(k1_matrix12(X1,X2,X3,X7,X8,X9))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1008, negated_conjecture, (r1_xxreal_0(esk8_0,esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)))|r1_xxreal_0(esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)),esk8_0)|k1_matrix12(X1,X2,X3,X7,X8,X9)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(k1_matrix12(X1,X2,X3,X7,X8,X9))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4075, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))!=k3_finseq_1(X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4129, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,X1)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4126, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,X1)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4279, plain, (m2_finseq_1(k2_zfmisc_1(X2,X3),X1)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4275, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k2_zfmisc_1(X3,X4))|~v1_xboole_0(X4))).#trainneg
cnf(i_0_1009, plain, (r1_xxreal_0(X1,esk16_7(X2,X3,X4,X5,X6,X7,k1_matrix12(X2,X3,X4,X8,X9,X10)))|r1_xxreal_0(esk16_7(X2,X3,X4,X5,X6,X7,k1_matrix12(X2,X3,X4,X8,X9,X10)),X1)|k1_matrix12(X2,X3,X4,X8,X9,X10)=k1_matrix12(X2,X3,X4,X5,X6,X7)|v6_struct_0(X4)|~v7_ordinal1(X1)|k3_finseq_1(X5)!=k3_finseq_1(k1_matrix12(X2,X3,X4,X8,X9,X10))|~m1_matrix_1(X5,u1_struct_0(X4),X2,X3)|~m1_matrix_1(X8,u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X10)|~v7_ordinal1(X9)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4271, plain, (v1_funct_1(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1014, plain, (m1_subset_1(X1,k2_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4297, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))!=k1_funct_1(X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4343, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4299, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|X1=esk4_0|X1=esk3_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~r1_xxreal_0(np__1,X2)|~r2_hidden(X1,k4_finseq_1(esk6_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_4302, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|X1=esk4_0|X1=esk3_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~r1_xxreal_0(np__1,X2)|~r2_hidden(X1,k4_finseq_1(esk6_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_4340, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))).#trainneg
cnf(i_0_4327, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))).#trainneg
cnf(i_0_4260, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk12_1(k3_finseq_2(X4))))|~v1_xboole_0(X3)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_1619, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X4)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_818, negated_conjecture, (v7_ordinal1(k1_matrix_1(X1))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_1055, negated_conjecture, (r1_xxreal_0(esk2_0,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|r1_xxreal_0(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),esk2_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1056, negated_conjecture, (r1_xxreal_0(esk3_0,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|r1_xxreal_0(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),esk3_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1057, negated_conjecture, (r1_xxreal_0(esk4_0,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|r1_xxreal_0(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),esk4_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4413, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4411, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(X2,X3,X4)|v1_xboole_0(X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)))!=k3_matrix_1(X2,X3,X4,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(X2,X3,X4))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)))|~v1_finseq_1(k6_matrix_1(X2,X3,X4))|~v1_funct_1(k6_matrix_1(X2,X3,X4))|~v1_relat_1(k6_matrix_1(X2,X3,X4))|~v7_ordinal1(X1)|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3)|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_4401, plain, (v1_funct_1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1058, negated_conjecture, (r1_xxreal_0(esk8_0,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|r1_xxreal_0(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),esk8_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4272, plain, (v1_finseq_1(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1059, plain, (r1_xxreal_0(X1,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3))))|r1_xxreal_0(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3))),X1)|~v7_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1063, negated_conjecture, (r1_xxreal_0(esk2_0,esk9_4(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4,X5))|r1_xxreal_0(esk9_4(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4,X5),esk2_0)|X5=k6_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4)|v1_xboole_0(k3_finseq_2(X1))|v1_xboole_0(X1)|k1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))!=k3_finseq_1(X5)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(X1)))|~m2_finseq_1(X5,X1)|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))|~v1_matrix_1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_4426, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|X1=esk3_0|X1=esk4_0|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~r2_hidden(X1,k4_finseq_1(esk6_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(X3)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(X3)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))).#trainneg
cnf(i_0_4338, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(X3)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4439, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|X1=esk3_0|X1=esk4_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~r2_hidden(X1,k4_finseq_1(esk6_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_funct_1(X3)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_relat_1(X3)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))).#trainneg
cnf(i_0_4127, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,X1)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_matrix_1(X2))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_4325, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(X3)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4407, plain, (k1_funct_1(X1,X2)=k3_matrix_1(X3,X4,X5,X2)|v1_xboole_0(k2_finseq_1(k1_matrix_1(X4)))|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,X4,X5)|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v1_matrix_1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,k1_matrix_1(X4))|~v7_ordinal1(k1_matrix_1(X4)))).#trainneg
cnf(i_0_4631, plain, (v1_finseq_1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1064, negated_conjecture, (r1_xxreal_0(esk3_0,esk9_4(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4,X5))|r1_xxreal_0(esk9_4(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4,X5),esk3_0)|X5=k6_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4)|v1_xboole_0(k3_finseq_2(X1))|v1_xboole_0(X1)|k1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))!=k3_finseq_1(X5)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(X1)))|~m2_finseq_1(X5,X1)|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))|~v1_matrix_1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1065, negated_conjecture, (r1_xxreal_0(esk4_0,esk9_4(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4,X5))|r1_xxreal_0(esk9_4(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4,X5),esk4_0)|X5=k6_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4)|v1_xboole_0(k3_finseq_2(X1))|v1_xboole_0(X1)|k1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))!=k3_finseq_1(X5)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(X1)))|~m2_finseq_1(X5,X1)|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))|~v1_matrix_1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1066, negated_conjecture, (r1_xxreal_0(esk8_0,esk9_4(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4,X5))|r1_xxreal_0(esk9_4(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4,X5),esk8_0)|X5=k6_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4)|v1_xboole_0(k3_finseq_2(X1))|v1_xboole_0(X1)|k1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))!=k3_finseq_1(X5)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(X1)))|~m2_finseq_1(X5,X1)|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))|~v1_matrix_1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_4589, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,esk7_0,X4,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,esk7_0,X4)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5)))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(esk7_0,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_4586, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,esk7_0,X4,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,esk7_0,X4)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5)))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(esk7_0,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_4139, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4312, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|X1=esk4_0|X1=esk3_0|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~r2_hidden(X1,k4_finseq_1(esk6_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v7_ordinal1(X1)|k1_matrix_1(esk6_0)!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(X3)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4261, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk12_1(k3_finseq_2(X4))))|~v1_xboole_0(X3)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_1620, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X4)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4702, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk12_1(k5_numbers))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk12_1(k5_numbers))|esk4_0=esk8_0|esk3_0=esk8_0|r1_xxreal_0(k1_matrix_1(esk6_0),esk12_1(k5_numbers))|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk12_1(k5_numbers),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk12_1(k5_numbers)))).#trainneg
cnf(i_0_1067, plain, (r1_xxreal_0(X1,esk9_4(X2,k6_matrix_1(k3_finseq_2(X2),X3,X4),X5,X6))|r1_xxreal_0(esk9_4(X2,k6_matrix_1(k3_finseq_2(X2),X3,X4),X5,X6),X1)|X6=k6_matrix_1(X2,k6_matrix_1(k3_finseq_2(X2),X3,X4),X5)|v1_xboole_0(k3_finseq_2(X2))|v1_xboole_0(X2)|~v7_ordinal1(X1)|k1_matrix_1(k6_matrix_1(k3_finseq_2(X2),X3,X4))!=k3_finseq_1(X6)|~m2_finseq_1(X3,k3_finseq_2(k3_finseq_2(X2)))|~m2_finseq_1(X6,X2)|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X2),X3,X4))|~v1_matrix_1(X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_1070, plain, (X1=k6_matrix_1(X2,k8_matrix_1(k3_finseq_2(X2),k6_matrix_1(k3_finseq_2(k3_finseq_2(X2)),X3,X4),X5),X6)|v1_xboole_0(k3_finseq_2(X2))|v1_xboole_0(X2)|v7_ordinal1(esk9_4(X2,k8_matrix_1(k3_finseq_2(X2),k6_matrix_1(k3_finseq_2(k3_finseq_2(X2)),X3,X4),X5),X6,X1))|v1_xboole_0(k3_finseq_2(k3_finseq_2(X2)))|k1_matrix_1(k8_matrix_1(k3_finseq_2(X2),k6_matrix_1(k3_finseq_2(k3_finseq_2(X2)),X3,X4),X5))!=k3_finseq_1(X1)|~m2_finseq_1(k6_matrix_1(k3_finseq_2(k3_finseq_2(X2)),X3,X4),k3_finseq_2(k3_finseq_2(X2)))|~m2_finseq_1(X1,X2)|~v1_matrix_1(k8_matrix_1(k3_finseq_2(X2),k6_matrix_1(k3_finseq_2(k3_finseq_2(X2)),X3,X4),X5))|~v1_matrix_1(k6_matrix_1(k3_finseq_2(k3_finseq_2(X2)),X3,X4))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m2_finseq_1(X3,k3_finseq_2(k3_finseq_2(k3_finseq_2(X2))))|~v1_matrix_1(X3)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_1071, plain, (k6_matrix_1(X1,k8_matrix_1(k3_finseq_2(X1),k6_matrix_1(k3_finseq_2(k3_finseq_2(X1)),X2,X3),X4),X5)=k8_matrix_1(X1,k8_matrix_1(k3_finseq_2(X1),k6_matrix_1(k3_finseq_2(k3_finseq_2(X1)),X2,X3),X4),X5)|v1_xboole_0(k3_finseq_2(X1))|v1_xboole_0(X1)|v1_xboole_0(k3_finseq_2(k3_finseq_2(X1)))|~m2_finseq_1(k6_matrix_1(k3_finseq_2(k3_finseq_2(X1)),X2,X3),k3_finseq_2(k3_finseq_2(X1)))|~v1_matrix_1(k8_matrix_1(k3_finseq_2(X1),k6_matrix_1(k3_finseq_2(k3_finseq_2(X1)),X2,X3),X4))|~v1_matrix_1(k6_matrix_1(k3_finseq_2(k3_finseq_2(X1)),X2,X3))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(k3_finseq_2(X1))))|~v1_matrix_1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_4688, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X1)|esk4_0=esk8_0|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4691, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X1)|esk4_0=esk8_0|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1623, plain, (esk12_1(k3_finseq_2(X1))=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1072, plain, (m2_finseq_1(k8_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4),X1)|v1_xboole_0(X1)|v1_xboole_0(k3_finseq_2(X1))|~m2_finseq_1(k6_matrix_1(k3_finseq_2(X1),X2,X3),k3_finseq_2(X1))|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))|~v7_ordinal1(X4)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(X1)))|~v1_matrix_1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_4925, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,k1_matrix_1(esk6_0))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,k1_matrix_1(esk6_0))|esk3_0=esk8_0|esk4_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(k1_matrix_1(esk6_0),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(esk6_0)))).#trainneg
cnf(i_0_1073, plain, (k1_funct_1(X1,esk13_2(X2,k2_finseq_1(k1_matrix_1(k6_matrix_1(k3_finseq_2(X3),X4,X5)))))=k3_matrix_1(X3,k6_matrix_1(k3_finseq_2(X3),X4,X5),X6,esk13_2(X2,k2_finseq_1(k1_matrix_1(k6_matrix_1(k3_finseq_2(X3),X4,X5)))))|v1_xboole_0(X3)|v1_xboole_0(k3_finseq_2(X3))|X1!=k8_matrix_1(X3,k6_matrix_1(k3_finseq_2(X3),X4,X5),X6)|~m2_finseq_1(k6_matrix_1(k3_finseq_2(X3),X4,X5),k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X3),X4,X5))|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(k6_matrix_1(k3_finseq_2(X3),X4,X5))))|~v7_ordinal1(esk13_2(X2,k2_finseq_1(k1_matrix_1(k6_matrix_1(k3_finseq_2(X3),X4,X5)))))|~v7_ordinal1(X6)|~m2_finseq_1(X4,k3_finseq_2(k3_finseq_2(X3)))|~v1_matrix_1(X4)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1074, plain, (k1_funct_1(X1,esk9_4(X2,k6_matrix_1(k3_finseq_2(X3),X4,X5),X6,X7))=k3_matrix_1(X3,k6_matrix_1(k3_finseq_2(X3),X4,X5),X8,esk9_4(X2,k6_matrix_1(k3_finseq_2(X3),X4,X5),X6,X7))|X7=k6_matrix_1(X2,k6_matrix_1(k3_finseq_2(X3),X4,X5),X6)|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(k3_finseq_2(X3))|X1!=k8_matrix_1(X3,k6_matrix_1(k3_finseq_2(X3),X4,X5),X8)|k1_matrix_1(k6_matrix_1(k3_finseq_2(X3),X4,X5))!=k3_finseq_1(X7)|~m2_finseq_1(k6_matrix_1(k3_finseq_2(X3),X4,X5),k3_finseq_2(X3))|~m2_finseq_1(k6_matrix_1(k3_finseq_2(X3),X4,X5),k3_finseq_2(X2))|~m2_finseq_1(X1,X3)|~m2_finseq_1(X7,X2)|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X3),X4,X5))|~v7_ordinal1(X8)|~v7_ordinal1(X6)|~m2_finseq_1(X4,k3_finseq_2(k3_finseq_2(X3)))|~v1_matrix_1(X4)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_4937, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|esk3_0=esk8_0|esk4_0=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))).#trainneg
cnf(i_0_5022, plain, (m2_finseq_1(k2_zfmisc_1(X2,X3),X1)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1075, plain, (k3_finseq_1(k6_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4))=k1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))|v1_xboole_0(k3_finseq_2(X1))|v1_xboole_0(X1)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(X1)))|~m2_finseq_1(k6_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4),X1)|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))|~v1_matrix_1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1076, plain, (k1_matrix_1(k8_matrix_1(k3_finseq_2(X1),k6_matrix_1(k3_finseq_2(k3_finseq_2(X1)),X2,X3),X4))=k3_finseq_1(X5)|v1_xboole_0(k3_finseq_2(X1))|v1_xboole_0(X1)|v1_xboole_0(k3_finseq_2(k3_finseq_2(X1)))|X5!=k6_matrix_1(X1,k8_matrix_1(k3_finseq_2(X1),k6_matrix_1(k3_finseq_2(k3_finseq_2(X1)),X2,X3),X4),X6)|~m2_finseq_1(k6_matrix_1(k3_finseq_2(k3_finseq_2(X1)),X2,X3),k3_finseq_2(k3_finseq_2(X1)))|~m2_finseq_1(X5,X1)|~v1_matrix_1(k8_matrix_1(k3_finseq_2(X1),k6_matrix_1(k3_finseq_2(k3_finseq_2(X1)),X2,X3),X4))|~v1_matrix_1(k6_matrix_1(k3_finseq_2(k3_finseq_2(X1)),X2,X3))|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(k3_finseq_2(X1))))|~v1_matrix_1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1077, plain, (k1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))=k3_finseq_1(X4)|v1_xboole_0(k3_finseq_2(X1))|v1_xboole_0(X1)|X4!=k8_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X5)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(X1)))|~m2_finseq_1(X4,X1)|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))|~v1_matrix_1(X2)|~v7_ordinal1(X5)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_4337, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(X3)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4138, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4551, negated_conjecture, (m2_subset_1(k1_matrix_1(X1),k1_numbers,k5_numbers)|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_4324, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(X3)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4116, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,esk7_0,X4,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,esk7_0,X4)|~m2_finseq_1(esk7_0,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~r2_hidden(X2,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5))))|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1488, negated_conjecture, (v3_ordinal1(k1_matrix_1(X1))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_1086, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,esk12_1(k2_finseq_1(k1_matrix_1(X4))))=k3_matrix_1(u1_struct_0(X1),X4,X5,esk12_1(k2_finseq_1(k1_matrix_1(X4))))|v1_xboole_0(k2_finseq_1(k1_matrix_1(X4)))|v6_struct_0(X1)|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(esk12_1(k2_finseq_1(k1_matrix_1(X4))))|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4948, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|esk3_0=esk8_0|esk4_0=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(X3)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4311, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|X1=esk4_0|X1=esk3_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~r2_hidden(X1,k4_finseq_1(esk6_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v7_ordinal1(X1)|k1_matrix_1(esk6_0)!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4606, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|X1=esk3_0|X1=esk4_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(X3))|~r1_xxreal_0(np__1,X2)|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k4_finseq_1(esk6_0)))).#trainneg
cnf(i_0_4192, negated_conjecture, (m2_subset_1(k1_matrix_1(X1),k1_numbers,k5_numbers)|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(X2,k4_finseq_1(X3)))).#trainneg
cnf(i_0_4935, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|esk3_0=esk8_0|esk4_0=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(X3)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3945, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(X3)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3182, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,k1_matrix_1(X1))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,k1_matrix_1(X1))|~m2_subset_1(k1_matrix_1(X1),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(X1))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_1626, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5130, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk12_1(k2_finseq_1(k1_matrix_1(esk6_0))))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk12_1(k2_finseq_1(k1_matrix_1(esk6_0))))|v1_xboole_0(k2_finseq_1(k1_matrix_1(esk6_0)))|~v7_ordinal1(esk12_1(k2_finseq_1(k1_matrix_1(esk6_0)))))).#trainneg
cnf(i_0_1209, plain, (X1=k6_matrix_1(X2,esk12_1(k3_finseq_2(k3_finseq_2(X2))),X3)|v1_xboole_0(X2)|v7_ordinal1(esk9_4(X2,esk12_1(k3_finseq_2(k3_finseq_2(X2))),X3,X1))|k1_matrix_1(esk12_1(k3_finseq_2(k3_finseq_2(X2))))!=k3_finseq_1(X1)|~m2_finseq_1(X1,X2)|~v1_matrix_1(esk12_1(k3_finseq_2(k3_finseq_2(X2))))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1210, plain, (k6_matrix_1(X1,esk12_1(k3_finseq_2(k3_finseq_2(X1))),X2)=k8_matrix_1(X1,esk12_1(k3_finseq_2(k3_finseq_2(X1))),X2)|v1_xboole_0(X1)|~v1_matrix_1(esk12_1(k3_finseq_2(k3_finseq_2(X1))))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5047, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|esk4_0=esk8_0|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5045, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,k1_matrix_1(X1))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,k1_matrix_1(X1))|esk4_0=esk8_0|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(k1_matrix_1(X1),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(X1))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_5094, negated_conjecture, (m2_subset_1(k1_matrix_1(X1),k1_numbers,k5_numbers)|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_5044, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))|esk4_0=esk8_0|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5146, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|~m2_subset_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1761, plain, (esk12_1(k3_finseq_2(X1))=esk12_1(k3_finseq_2(X2))|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5357, negated_conjecture, (m2_subset_1(k1_matrix_1(esk6_0),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5276, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X3,X2)|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(X4))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1211, plain, (m1_subset_1(k3_matrix_1(X1,esk12_1(k3_finseq_2(k3_finseq_2(X1))),X2,X3),X1)|~v1_matrix_1(esk12_1(k3_finseq_2(k3_finseq_2(X1))))|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5339, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,k1_matrix_1(esk6_0))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,k1_matrix_1(esk6_0))|esk3_0=esk8_0|esk4_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(k1_matrix_1(esk6_0),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(esk6_0)))).#trainneg
cnf(i_0_1212, plain, (k1_matrix_1(esk12_1(k3_finseq_2(k3_finseq_2(X1))))=k3_finseq_1(X2)|v1_xboole_0(X1)|X2!=k6_matrix_1(X1,esk12_1(k3_finseq_2(k3_finseq_2(X1))),X3)|~m2_finseq_1(X2,X1)|~v1_matrix_1(esk12_1(k3_finseq_2(k3_finseq_2(X1))))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1213, plain, (k1_funct_1(X1,esk16_7(X2,X3,X4,esk12_1(k3_finseq_2(k3_finseq_2(X5))),X6,X7,X8))=k3_matrix_1(X5,esk12_1(k3_finseq_2(k3_finseq_2(X5))),X9,esk16_7(X2,X3,X4,esk12_1(k3_finseq_2(k3_finseq_2(X5))),X6,X7,X8))|X8=k1_matrix12(X2,X3,X4,esk12_1(k3_finseq_2(k3_finseq_2(X5))),X6,X7)|v1_xboole_0(X5)|v6_struct_0(X4)|v2_struct_0(X4)|X1!=k6_matrix_1(X5,esk12_1(k3_finseq_2(k3_finseq_2(X5))),X9)|k3_finseq_1(esk12_1(k3_finseq_2(k3_finseq_2(X5))))!=k3_finseq_1(X8)|~epred1_2(X3,X2)|~m2_finseq_1(X1,X5)|~v1_matrix_1(esk12_1(k3_finseq_2(k3_finseq_2(X5))))|~m1_matrix_1(X8,u1_struct_0(X4),X2,X3)|~m1_matrix_1(esk12_1(k3_finseq_2(k3_finseq_2(X5))),u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X9)|~v7_ordinal1(X7)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_5356, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,k1_matrix_1(X1))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,k1_matrix_1(X1))|esk3_0=esk8_0|esk4_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~r1_xxreal_0(np__1,k1_matrix_1(X1))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_5355, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,k1_matrix_1(X1))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,k1_matrix_1(X1))|~r1_xxreal_0(np__1,k1_matrix_1(X1))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_5338, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|esk3_0=esk8_0|esk4_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|~v7_ordinal1(X1)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))).#trainneg
cnf(i_0_5360, negated_conjecture, (m2_subset_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)),k1_numbers,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5374, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))|esk3_0=esk8_0|esk4_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))|~v7_ordinal1(X1)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))).#trainneg
cnf(i_0_5380, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|~v7_ordinal1(X1)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))).#trainneg
cnf(i_0_3028, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X4,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6)))=k3_matrix_1(u1_struct_0(X1),esk6_0,X5,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6)))|v6_struct_0(X1)|~m2_subset_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6)))|~r2_hidden(X4,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6)))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_5397, plain, (m2_finseq_1(esk12_1(k3_finseq_2(X2)),X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3035, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,k1_matrix_1(esk6_0))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,k1_matrix_1(esk6_0)))).#trainneg
cnf(i_0_2975, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,k1_matrix_1(esk6_0))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,k1_matrix_1(esk6_0)))).#trainneg
cnf(i_0_1214, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|~r2_hidden(X1,esk12_1(k3_finseq_2(X2))))).#trainneg
cnf(i_0_4759, plain, (v1_funct_1(k2_zfmisc_1(X3,X4))|~v1_xboole_0(X2)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1222, plain, (v1_funct_1(k8_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4))|v1_xboole_0(X1)|v1_xboole_0(k3_finseq_2(X1))|~m2_finseq_1(k6_matrix_1(k3_finseq_2(X1),X2,X3),k3_finseq_2(X1))|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))|~v7_ordinal1(X4)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(X1)))|~v1_matrix_1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1225, plain, (v1_finseq_1(k8_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4))|v1_xboole_0(X1)|v1_xboole_0(k3_finseq_2(X1))|~m2_finseq_1(k6_matrix_1(k3_finseq_2(X1),X2,X3),k3_finseq_2(X1))|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))|~v7_ordinal1(X4)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(X1)))|~v1_matrix_1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_5462, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5464, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_5434, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(X3))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1228, plain, (v1_relat_1(k8_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4))|v1_xboole_0(X1)|v1_xboole_0(k3_finseq_2(X1))|~m2_finseq_1(k6_matrix_1(k3_finseq_2(X1),X2,X3),k3_finseq_2(X1))|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))|~v7_ordinal1(X4)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(X1)))|~v1_matrix_1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_4760, plain, (v1_finseq_1(k2_zfmisc_1(X3,X4))|~v1_xboole_0(X2)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_5631, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5633, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_5634, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_5616, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v7_ordinal1(X2)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(X3)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(X3)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))).#trainneg
cnf(i_0_5590, plain, (v1_funct_1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5663, plain, (v1_finseq_1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1276, plain, (m2_finseq_1(esk13_2(X1,k4_finseq_2(X2,X3)),X3)|~v7_ordinal1(X2)|~r2_hidden(X1,k4_finseq_2(X2,X3)))).#trainneg
cnf(i_0_1423, plain, (X1=k1_matrix12(X2,X3,X4,X5,X6,X7)|v3_ordinal1(esk16_7(X2,X3,X4,X5,X6,X7,X1))|v6_struct_0(X4)|k3_finseq_1(X5)!=k3_finseq_1(X1)|~epred1_2(X3,X2)|~m1_matrix_1(X1,u1_struct_0(X4),X2,X3)|~m1_matrix_1(X5,u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(k1_zfmisc_1(k2_finseq_1(k1_matrix_1(X5)))))).#trainneg
cnf(i_0_5602, negated_conjecture, (X1=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k1_funct_1(X1,esk10_2(X1,k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(X1,k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|k3_finseq_1(X1)!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_finseq_1(X1)|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_funct_1(X1)|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_relat_1(X1)|~m2_subset_1(esk10_2(X1,k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(esk10_2(X1,k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,esk10_2(X1,k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~v7_ordinal1(esk10_2(X1,k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))))).#trainneg
cnf(i_0_5600, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v7_ordinal1(X1)|~r1_xxreal_0(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_matrix_1(esk6_0))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))))).#trainneg
cnf(i_0_5686, plain, (k1_funct_1(X1,k1_matrix_1(X2))=k3_matrix_1(X3,X2,X4,k1_matrix_1(X2))|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,X2,X4)|~m2_subset_1(k1_matrix_1(X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(X2))|~m2_finseq_1(X2,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v1_matrix_1(X2)|~v7_ordinal1(k1_matrix_1(X2))|~v7_ordinal1(X4))).#trainneg
cnf(i_0_5650, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|k1_matrix_1(X4)!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(X3)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5718, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v7_ordinal1(X1)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5717, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5800, negated_conjecture, (k1_funct_1(X1,k1_matrix_1(esk6_0))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,k1_matrix_1(esk6_0))|X1!=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|~r1_xxreal_0(np__1,k1_matrix_1(esk6_0))|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1424, plain, (X1=k1_matrix12(X2,X3,X4,X5,X6,X7)|v3_ordinal1(esk16_7(X2,X3,X4,X5,X6,X7,X1))|v6_struct_0(X4)|k3_finseq_1(X5)!=k3_finseq_1(X1)|~epred1_2(X3,X2)|~m1_matrix_1(X1,u1_struct_0(X4),X2,X3)|~m1_matrix_1(X5,u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(k2_finseq_1(k1_matrix_1(X5))))).#trainneg
cnf(i_0_1425, plain, (X1=k1_matrix12(X2,X3,X4,X5,X6,X7)|v3_ordinal1(esk16_7(X2,X3,X4,X5,X6,X7,X1))|v6_struct_0(X4)|k3_finseq_1(X5)!=k3_finseq_1(X1)|~epred1_2(X3,X2)|~m1_matrix_1(X1,u1_struct_0(X4),X2,X3)|~m1_matrix_1(X5,u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(k1_matrix_1(X5)))).#trainneg
cnf(i_0_1429, plain, (r1_xxreal_0(esk1_0,esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)))|r1_xxreal_0(esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)),esk1_0)|k1_matrix12(X1,X2,X3,X7,X8,X9)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(k1_matrix12(X1,X2,X3,X7,X8,X9))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5720, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5779, negated_conjecture, (X1=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k1_funct_1(X1,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|k3_finseq_1(X1)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~m2_subset_1(esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5688, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))!=k1_funct_1(X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(X1)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2)))).#trainneg
cnf(i_0_5774, negated_conjecture, (X1=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k1_funct_1(X1,esk10_2(X1,k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(X1,k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|k3_finseq_1(X1)!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~m2_subset_1(esk10_2(X1,k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(esk10_2(X1,k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~r1_xxreal_0(np__1,esk10_2(X1,k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_5801, negated_conjecture, (k1_funct_1(X1,k1_matrix_1(esk6_0))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,k1_matrix_1(esk6_0))|X1!=k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|~r1_xxreal_0(np__1,k1_matrix_1(esk6_0))|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1430, plain, (r1_xxreal_0(esk1_0,esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)))|r1_xxreal_0(esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)),esk1_0)|k1_matrix12(X1,X2,X3,X7,X8,X9)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(k1_matrix12(X1,X2,X3,X7,X8,X9))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1431, negated_conjecture, (r1_xxreal_0(esk1_0,esk13_2(X1,k5_numbers))|r1_xxreal_0(esk13_2(X1,k5_numbers),esk1_0)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_1432, negated_conjecture, (r1_xxreal_0(esk1_0,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|r1_xxreal_0(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),esk1_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5771, negated_conjecture, (X1=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k1_funct_1(X1,esk10_2(X1,k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(X1,k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|k3_finseq_1(X1)!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~m2_subset_1(esk10_2(X1,k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(esk10_2(X1,k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~r1_xxreal_0(np__1,esk10_2(X1,k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_5821, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|k1_matrix_1(esk6_0)!=k3_finseq_1(X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5772, negated_conjecture, (X1=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k1_funct_1(X1,esk10_2(X1,k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(X1,k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|k3_finseq_1(X1)!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~m2_subset_1(esk10_2(X1,k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(esk10_2(X1,k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_matrix_1(X3))|~r1_xxreal_0(np__1,esk10_2(X1,k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_5699, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v7_ordinal1(X1)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v7_ordinal1(X4)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4))!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5840, negated_conjecture, (k1_funct_1(X1,k1_matrix_1(esk6_0))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,k1_matrix_1(esk6_0))|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|~r1_xxreal_0(np__1,k1_matrix_1(esk6_0))|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1433, negated_conjecture, (r1_xxreal_0(esk1_0,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|r1_xxreal_0(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),esk1_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1435, negated_conjecture, (r1_xxreal_0(esk1_0,esk10_2(X1,X2))|r1_xxreal_0(esk10_2(X1,X2),esk1_0)|X1=X2|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1436, negated_conjecture, (r1_xxreal_0(esk1_0,esk9_4(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4,X5))|r1_xxreal_0(esk9_4(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4,X5),esk1_0)|X5=k6_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4)|v1_xboole_0(k3_finseq_2(X1))|v1_xboole_0(X1)|k1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))!=k3_finseq_1(X5)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(X1)))|~m2_finseq_1(X5,X1)|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))|~v1_matrix_1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_5861, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5864, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))).#trainneg
cnf(i_0_5918, negated_conjecture, (X1=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k1_funct_1(X1,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|k3_finseq_1(X1)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~m2_subset_1(esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~r1_xxreal_0(np__1,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_5892, negated_conjecture, (X1=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k1_funct_1(X1,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|k3_finseq_1(X1)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~m2_subset_1(esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~r1_xxreal_0(np__1,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_4268, plain, (m1_subset_1(k2_zfmisc_1(X2,X3),k3_finseq_2(X1))|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5898, negated_conjecture, (k1_funct_1(X1,k1_matrix_1(esk6_0))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,k1_matrix_1(esk6_0))|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|~r1_xxreal_0(np__1,k1_matrix_1(esk6_0))|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1443, negated_conjecture, (r1_xxreal_0(esk1_0,k2_zfmisc_1(X1,X2))|r1_xxreal_0(k2_zfmisc_1(X1,X2),esk1_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1444, negated_conjecture, (r1_xxreal_0(esk1_0,k2_zfmisc_1(X1,X2))|r1_xxreal_0(k2_zfmisc_1(X1,X2),esk1_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1447, plain, (k1_funct_1(k6_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4),esk16_7(X5,X6,X7,k6_matrix_1(k3_finseq_2(X1),X2,X3),X8,X9,X10))=k3_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4,esk16_7(X5,X6,X7,k6_matrix_1(k3_finseq_2(X1),X2,X3),X8,X9,X10))|X10=k1_matrix12(X5,X6,X7,k6_matrix_1(k3_finseq_2(X1),X2,X3),X8,X9)|v1_xboole_0(k3_finseq_2(X1))|v1_xboole_0(X1)|v6_struct_0(X7)|k3_finseq_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))!=k3_finseq_1(X10)|~epred1_2(X6,X5)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(X1)))|~m2_finseq_1(k6_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4),X1)|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))|~v1_matrix_1(X2)|~m1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3),u1_struct_0(X7),X5,X6)|~m1_matrix_1(X10,u1_struct_0(X7),X5,X6)|~l6_algstr_0(X7)|~v5_vectsp_1(X7)|~v4_vectsp_1(X7)|~v4_rlvect_1(X7)|~v3_rlvect_1(X7)|~v2_rlvect_1(X7)|~v5_group_1(X7)|~v3_group_1(X7)|~v33_algstr_0(X7)|~v13_algstr_0(X7)|~v7_ordinal1(X4)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_5942, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5937, negated_conjecture, (X1=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k1_funct_1(X1,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|k3_finseq_1(X1)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~m2_subset_1(esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_matrix_1(X3))|~r1_xxreal_0(np__1,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),u1_struct_0(esk5_0))|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5009, plain, (m1_subset_1(k2_zfmisc_1(X2,X3),k3_finseq_2(X1))|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4061, negated_conjecture, (v1_card_1(k1_matrix_1(X1))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_1449, plain, (k1_funct_1(X1,esk16_7(X2,X3,X4,k6_matrix_1(k3_finseq_2(X5),X6,X7),X8,X9,X10))=k3_matrix_1(X5,k6_matrix_1(k3_finseq_2(X5),X6,X7),X11,esk16_7(X2,X3,X4,k6_matrix_1(k3_finseq_2(X5),X6,X7),X8,X9,X10))|X10=k1_matrix12(X2,X3,X4,k6_matrix_1(k3_finseq_2(X5),X6,X7),X8,X9)|v1_xboole_0(k3_finseq_2(X5))|v1_xboole_0(X5)|v6_struct_0(X4)|X1!=k8_matrix_1(X5,k6_matrix_1(k3_finseq_2(X5),X6,X7),X11)|k3_finseq_1(k6_matrix_1(k3_finseq_2(X5),X6,X7))!=k3_finseq_1(X10)|~epred1_2(X3,X2)|~m2_finseq_1(X6,k3_finseq_2(k3_finseq_2(X5)))|~m2_finseq_1(X1,X5)|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X5),X6,X7))|~v1_matrix_1(X6)|~m1_matrix_1(k6_matrix_1(k3_finseq_2(X5),X6,X7),u1_struct_0(X4),X2,X3)|~m1_matrix_1(X10,u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X11)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_1448, plain, (k1_funct_1(X1,esk16_7(X2,X3,X4,k8_matrix_1(k3_finseq_2(X5),k6_matrix_1(k3_finseq_2(k3_finseq_2(X5)),X6,X7),X8),X9,X10,X11))=k3_matrix_1(X5,k8_matrix_1(k3_finseq_2(X5),k6_matrix_1(k3_finseq_2(k3_finseq_2(X5)),X6,X7),X8),X12,esk16_7(X2,X3,X4,k8_matrix_1(k3_finseq_2(X5),k6_matrix_1(k3_finseq_2(k3_finseq_2(X5)),X6,X7),X8),X9,X10,X11))|X11=k1_matrix12(X2,X3,X4,k8_matrix_1(k3_finseq_2(X5),k6_matrix_1(k3_finseq_2(k3_finseq_2(X5)),X6,X7),X8),X9,X10)|v1_xboole_0(k3_finseq_2(X5))|v1_xboole_0(X5)|v6_struct_0(X4)|v1_xboole_0(k3_finseq_2(k3_finseq_2(X5)))|X1!=k6_matrix_1(X5,k8_matrix_1(k3_finseq_2(X5),k6_matrix_1(k3_finseq_2(k3_finseq_2(X5)),X6,X7),X8),X12)|k3_finseq_1(k8_matrix_1(k3_finseq_2(X5),k6_matrix_1(k3_finseq_2(k3_finseq_2(X5)),X6,X7),X8))!=k3_finseq_1(X11)|~epred1_2(X3,X2)|~m2_finseq_1(k6_matrix_1(k3_finseq_2(k3_finseq_2(X5)),X6,X7),k3_finseq_2(k3_finseq_2(X5)))|~m2_finseq_1(X1,X5)|~v1_matrix_1(k8_matrix_1(k3_finseq_2(X5),k6_matrix_1(k3_finseq_2(k3_finseq_2(X5)),X6,X7),X8))|~v1_matrix_1(k6_matrix_1(k3_finseq_2(k3_finseq_2(X5)),X6,X7))|~m1_matrix_1(k8_matrix_1(k3_finseq_2(X5),k6_matrix_1(k3_finseq_2(k3_finseq_2(X5)),X6,X7),X8),u1_struct_0(X4),X2,X3)|~m1_matrix_1(X11,u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X12)|~v7_ordinal1(X10)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~m2_finseq_1(X6,k3_finseq_2(k3_finseq_2(k3_finseq_2(X5))))|~v1_matrix_1(X6)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_1450, plain, (X1=k1_matrix12(X2,X3,X4,X5,X6,X7)|v3_ordinal1(esk15_7(X2,X3,X4,X5,X6,X7,X1))|v6_struct_0(X4)|k3_finseq_1(X5)!=k3_finseq_1(X1)|~epred1_2(X3,X2)|~m1_matrix_1(X1,u1_struct_0(X4),X2,X3)|~m1_matrix_1(X5,u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(k1_zfmisc_1(k4_finseq_1(X5))))).#trainneg
cnf(i_0_5983, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1451, plain, (X1=k1_matrix12(X2,X3,X4,X5,X6,X7)|v3_ordinal1(esk15_7(X2,X3,X4,X5,X6,X7,X1))|v6_struct_0(X4)|k3_finseq_1(X5)!=k3_finseq_1(X1)|~epred1_2(X3,X2)|~m1_matrix_1(X1,u1_struct_0(X4),X2,X3)|~m1_matrix_1(X5,u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(k4_finseq_1(X5)))).#trainneg
cnf(i_0_1452, plain, (X1=k1_matrix12(X2,X3,X4,X5,X6,X7)|v3_ordinal1(esk15_7(X2,X3,X4,X5,X6,X7,X1))|v6_struct_0(X4)|k3_finseq_1(X5)!=k3_finseq_1(X1)|~epred1_2(X3,X2)|~m1_matrix_1(X1,u1_struct_0(X4),X2,X3)|~m1_matrix_1(X5,u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v5_ordinal1(X5)|~v1_finseq_1(X5)|~v1_funct_1(X5)|~v1_relat_1(X5))).#trainneg
cnf(i_0_1453, plain, (X1=k1_matrix12(X2,X3,X4,X5,X6,X7)|v3_ordinal1(esk15_7(X2,X3,X4,X5,X6,X7,X1))|v6_struct_0(X4)|k3_finseq_1(X5)!=k3_finseq_1(X1)|~epred1_2(X3,X2)|~m1_matrix_1(X1,u1_struct_0(X4),X2,X3)|~m1_matrix_1(X5,u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X5)|~v1_funct_1(X5)|~v1_relat_1(X5))).#trainneg
cnf(i_0_5823, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4))!=k3_finseq_1(X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_5712, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v7_ordinal1(X1)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v7_ordinal1(X4)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X4))!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1454, plain, (X1=k6_matrix_1(X2,esk12_1(k4_finseq_2(X3,k3_finseq_2(X2))),X4)|v1_xboole_0(X2)|v7_ordinal1(esk9_4(X2,esk12_1(k4_finseq_2(X3,k3_finseq_2(X2))),X4,X1))|k1_matrix_1(esk12_1(k4_finseq_2(X3,k3_finseq_2(X2))))!=k3_finseq_1(X1)|~m2_finseq_1(X1,X2)|~v1_matrix_1(esk12_1(k4_finseq_2(X3,k3_finseq_2(X2))))|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1455, plain, (k6_matrix_1(X1,esk12_1(k4_finseq_2(X2,k3_finseq_2(X1))),X3)=k8_matrix_1(X1,esk12_1(k4_finseq_2(X2,k3_finseq_2(X1))),X3)|v1_xboole_0(X1)|~v1_matrix_1(esk12_1(k4_finseq_2(X2,k3_finseq_2(X1))))|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1456, plain, (m1_subset_1(k3_matrix_1(X1,esk12_1(k4_finseq_2(X2,k3_finseq_2(X1))),X3,X4),X1)|~v1_matrix_1(esk12_1(k4_finseq_2(X2,k3_finseq_2(X1))))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5824, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X4))!=k3_finseq_1(X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_5381, plain, (m1_subset_1(esk12_1(k3_finseq_2(X2)),k3_finseq_2(X1))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1457, plain, (k1_matrix_1(esk12_1(k4_finseq_2(X1,k3_finseq_2(X2))))=k3_finseq_1(X3)|v1_xboole_0(X2)|X3!=k6_matrix_1(X2,esk12_1(k4_finseq_2(X1,k3_finseq_2(X2))),X4)|~m2_finseq_1(X3,X2)|~v1_matrix_1(esk12_1(k4_finseq_2(X1,k3_finseq_2(X2))))|~v7_ordinal1(X4)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1458, plain, (k1_funct_1(X1,esk16_7(X2,X3,X4,esk12_1(k4_finseq_2(X5,k3_finseq_2(X6))),X7,X8,X9))=k3_matrix_1(X6,esk12_1(k4_finseq_2(X5,k3_finseq_2(X6))),X10,esk16_7(X2,X3,X4,esk12_1(k4_finseq_2(X5,k3_finseq_2(X6))),X7,X8,X9))|X9=k1_matrix12(X2,X3,X4,esk12_1(k4_finseq_2(X5,k3_finseq_2(X6))),X7,X8)|v1_xboole_0(X6)|v6_struct_0(X4)|v2_struct_0(X4)|X1!=k6_matrix_1(X6,esk12_1(k4_finseq_2(X5,k3_finseq_2(X6))),X10)|k3_finseq_1(esk12_1(k4_finseq_2(X5,k3_finseq_2(X6))))!=k3_finseq_1(X9)|~epred1_2(X3,X2)|~m2_finseq_1(X1,X6)|~v1_matrix_1(esk12_1(k4_finseq_2(X5,k3_finseq_2(X6))))|~m1_matrix_1(X9,u1_struct_0(X4),X2,X3)|~m1_matrix_1(esk12_1(k4_finseq_2(X5,k3_finseq_2(X6))),u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X10)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1459, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|~r2_hidden(X1,esk12_1(k4_finseq_2(X3,X2)))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_6014, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k1_matrix_1(esk6_0))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,k1_matrix_1(esk6_0))|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6098, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_6219, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6184, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1469, plain, (r1_xxreal_0(k1_xboole_0,esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)))|r1_xxreal_0(esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)),k1_xboole_0)|k1_matrix12(X1,X2,X3,X7,X8,X9)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(k1_matrix12(X1,X2,X3,X7,X8,X9))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5018, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k2_zfmisc_1(X3,X4))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1470, plain, (r1_xxreal_0(k1_xboole_0,esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)))|r1_xxreal_0(esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)),k1_xboole_0)|k1_matrix12(X1,X2,X3,X7,X8,X9)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(k1_matrix12(X1,X2,X3,X7,X8,X9))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1471, plain, (r1_xxreal_0(k1_xboole_0,esk13_2(X1,k5_numbers))|r1_xxreal_0(esk13_2(X1,k5_numbers),k1_xboole_0)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_6117, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6216, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6217, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k1_matrix_1(X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,k1_matrix_1(X2))|~r1_xxreal_0(np__1,k1_matrix_1(X2))|~v7_ordinal1(X1)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_1472, plain, (r1_xxreal_0(k1_xboole_0,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|r1_xxreal_0(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1473, plain, (r1_xxreal_0(k1_xboole_0,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|r1_xxreal_0(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),k1_xboole_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1475, plain, (r1_xxreal_0(k1_xboole_0,esk10_2(X1,X2))|r1_xxreal_0(esk10_2(X1,X2),k1_xboole_0)|X1=X2|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5965, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k1_matrix_1(esk6_0))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,k1_matrix_1(esk6_0))|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5752, plain, (v1_funct_1(esk13_2(X1,k4_finseq_2(X2,X3)))|~r2_hidden(X1,k4_finseq_2(X2,X3))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1476, plain, (r1_xxreal_0(k1_xboole_0,esk9_4(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4,X5))|r1_xxreal_0(esk9_4(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4,X5),k1_xboole_0)|X5=k6_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4)|v1_xboole_0(k3_finseq_2(X1))|v1_xboole_0(X1)|k1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))!=k3_finseq_1(X5)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(X1)))|~m2_finseq_1(X5,X1)|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))|~v1_matrix_1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_3834, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X1)),k5_numbers)|~v7_ordinal1(k1_numbers))).#trainneg
cnf(i_0_4024, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X1)),k5_numbers)|~v7_ordinal1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1483, plain, (r1_xxreal_0(k1_xboole_0,k2_zfmisc_1(X1,X2))|r1_xxreal_0(k2_zfmisc_1(X1,X2),k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1484, plain, (r1_xxreal_0(k1_xboole_0,k2_zfmisc_1(X1,X2))|r1_xxreal_0(k2_zfmisc_1(X1,X2),k1_xboole_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_5753, plain, (v1_finseq_1(esk13_2(X1,k4_finseq_2(X2,X3)))|~r2_hidden(X1,k4_finseq_2(X2,X3))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1510, plain, (r1_xxreal_0(k1_xboole_0,esk9_4(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8))|r1_xxreal_0(esk9_4(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8),k1_xboole_0)|X8=k6_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6))!=k3_finseq_1(X8)|~m2_finseq_1(X8,u1_struct_0(X1))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1499, negated_conjecture, (m1_subset_1(k1_matrix_1(X1),k5_numbers)|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_1511, negated_conjecture, (r1_xxreal_0(esk1_0,esk9_4(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8))|r1_xxreal_0(esk9_4(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8),esk1_0)|X8=k6_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6))!=k3_finseq_1(X8)|~m2_finseq_1(X8,u1_struct_0(X1))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1512, negated_conjecture, (r1_xxreal_0(esk2_0,esk9_4(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8))|r1_xxreal_0(esk9_4(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8),esk2_0)|X8=k6_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6))!=k3_finseq_1(X8)|~m2_finseq_1(X8,u1_struct_0(X1))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_5614, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v7_ordinal1(X4)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4))!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(X3)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5755, plain, (v1_relat_1(esk13_2(X1,k4_finseq_2(X2,X3)))|~r2_hidden(X1,k4_finseq_2(X2,X3))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1513, negated_conjecture, (r1_xxreal_0(esk3_0,esk9_4(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8))|r1_xxreal_0(esk9_4(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8),esk3_0)|X8=k6_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6))!=k3_finseq_1(X8)|~m2_finseq_1(X8,u1_struct_0(X1))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1514, negated_conjecture, (r1_xxreal_0(esk4_0,esk9_4(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8))|r1_xxreal_0(esk9_4(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8),esk4_0)|X8=k6_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6))!=k3_finseq_1(X8)|~m2_finseq_1(X8,u1_struct_0(X1))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1515, negated_conjecture, (r1_xxreal_0(esk8_0,esk9_4(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8))|r1_xxreal_0(esk9_4(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8),esk8_0)|X8=k6_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6))!=k3_finseq_1(X8)|~m2_finseq_1(X8,u1_struct_0(X1))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_6295, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6292, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6348, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6352, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4))!=k1_matrix_1(X5)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~m1_matrix_1(X5,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_3606, plain, (v3_ordinal1(esk12_1(k4_finseq_2(X1,X2)))|~v3_ordinal1(k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1516, plain, (r1_xxreal_0(X1,esk9_4(u1_struct_0(X2),k1_matrix12(X3,X4,X2,X5,X6,X7),X8,X9))|r1_xxreal_0(esk9_4(u1_struct_0(X2),k1_matrix12(X3,X4,X2,X5,X6,X7),X8,X9),X1)|X9=k6_matrix_1(u1_struct_0(X2),k1_matrix12(X3,X4,X2,X5,X6,X7),X8)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|~v7_ordinal1(X1)|k1_matrix_1(k1_matrix12(X3,X4,X2,X5,X6,X7))!=k3_finseq_1(X9)|~m2_finseq_1(X9,u1_struct_0(X2))|~m1_matrix_1(X5,u1_struct_0(X2),X3,X4)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X8)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X7)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_1519, plain, (v1_funct_1(k8_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~m2_finseq_1(k1_matrix12(X2,X3,X1,X4,X5,X6),k3_finseq_2(u1_struct_0(X1)))|~v1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6))|~v7_ordinal1(X7)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1520, plain, (v1_finseq_1(k8_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~m2_finseq_1(k1_matrix12(X2,X3,X1,X4,X5,X6),k3_finseq_2(u1_struct_0(X1)))|~v1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6))|~v7_ordinal1(X7)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_6354, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_4084, plain, (esk13_2(X1,k1_zfmisc_1(k1_xboole_0))=k1_xboole_0|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_1521, plain, (v1_relat_1(k8_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~m2_finseq_1(k1_matrix12(X2,X3,X1,X4,X5,X6),k3_finseq_2(u1_struct_0(X1)))|~v1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6))|~v7_ordinal1(X7)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1522, plain, (m2_finseq_1(k8_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~m2_finseq_1(k1_matrix12(X2,X3,X1,X4,X5,X6),k3_finseq_2(u1_struct_0(X1)))|~v1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6))|~v7_ordinal1(X7)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1523, plain, (k1_funct_1(X1,esk13_2(X2,k2_finseq_1(k1_matrix_1(k1_matrix12(X3,X4,X5,X6,X7,X8)))))=k3_matrix_1(u1_struct_0(X5),k1_matrix12(X3,X4,X5,X6,X7,X8),X9,esk13_2(X2,k2_finseq_1(k1_matrix_1(k1_matrix12(X3,X4,X5,X6,X7,X8)))))|v1_xboole_0(u1_struct_0(X5))|v6_struct_0(X5)|X1!=k8_matrix_1(u1_struct_0(X5),k1_matrix12(X3,X4,X5,X6,X7,X8),X9)|~m2_finseq_1(k1_matrix12(X3,X4,X5,X6,X7,X8),k3_finseq_2(u1_struct_0(X5)))|~m2_finseq_1(X1,u1_struct_0(X5))|~v1_matrix_1(k1_matrix12(X3,X4,X5,X6,X7,X8))|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(k1_matrix12(X3,X4,X5,X6,X7,X8))))|~v7_ordinal1(esk13_2(X2,k2_finseq_1(k1_matrix_1(k1_matrix12(X3,X4,X5,X6,X7,X8)))))|~v7_ordinal1(X9)|~m1_matrix_1(X6,u1_struct_0(X5),X3,X4)|~l6_algstr_0(X5)|~v5_vectsp_1(X5)|~v4_vectsp_1(X5)|~v4_rlvect_1(X5)|~v3_rlvect_1(X5)|~v2_rlvect_1(X5)|~v5_group_1(X5)|~v3_group_1(X5)|~v33_algstr_0(X5)|~v13_algstr_0(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X8)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_5627, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v7_ordinal1(X4)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X4))!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(X3)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6465, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k5_numbers))).#trainneg
cnf(i_0_6462, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k5_numbers))).#trainneg
cnf(i_0_6449, plain, (r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_1798, plain, (esk12_1(k4_finseq_2(X1,X2))=k1_xboole_0|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1577, plain, (k3_finseq_1(k6_matrix_1(u1_struct_0(X3),k1_matrix12(X1,X2,X3,X4,X5,X6),X7))=k1_matrix_1(k1_matrix12(X1,X2,X3,X4,X5,X6))|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|~m2_finseq_1(k6_matrix_1(u1_struct_0(X3),k1_matrix12(X1,X2,X3,X4,X5,X6),X7),u1_struct_0(X3))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X7)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_6412, plain, (v1_funct_1(k8_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_6415, plain, (v1_finseq_1(k8_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1524, plain, (k1_funct_1(X1,esk9_4(X2,k1_matrix12(X3,X4,X5,X6,X7,X8),X9,X10))=k3_matrix_1(u1_struct_0(X5),k1_matrix12(X3,X4,X5,X6,X7,X8),X11,esk9_4(X2,k1_matrix12(X3,X4,X5,X6,X7,X8),X9,X10))|X10=k6_matrix_1(X2,k1_matrix12(X3,X4,X5,X6,X7,X8),X9)|v1_xboole_0(X2)|v1_xboole_0(u1_struct_0(X5))|v6_struct_0(X5)|X1!=k8_matrix_1(u1_struct_0(X5),k1_matrix12(X3,X4,X5,X6,X7,X8),X11)|k1_matrix_1(k1_matrix12(X3,X4,X5,X6,X7,X8))!=k3_finseq_1(X10)|~m2_finseq_1(k1_matrix12(X3,X4,X5,X6,X7,X8),k3_finseq_2(u1_struct_0(X5)))|~m2_finseq_1(k1_matrix12(X3,X4,X5,X6,X7,X8),k3_finseq_2(X2))|~m2_finseq_1(X1,u1_struct_0(X5))|~m2_finseq_1(X10,X2)|~v1_matrix_1(k1_matrix12(X3,X4,X5,X6,X7,X8))|~v7_ordinal1(X11)|~v7_ordinal1(X9)|~m1_matrix_1(X6,u1_struct_0(X5),X3,X4)|~l6_algstr_0(X5)|~v5_vectsp_1(X5)|~v4_vectsp_1(X5)|~v4_rlvect_1(X5)|~v3_rlvect_1(X5)|~v2_rlvect_1(X5)|~v5_group_1(X5)|~v3_group_1(X5)|~v33_algstr_0(X5)|~v13_algstr_0(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X8)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_1581, plain, (k1_funct_1(k6_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7),esk16_7(X8,X9,X10,k1_matrix12(X2,X3,X1,X4,X5,X6),X11,X12,X13))=k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,esk16_7(X8,X9,X10,k1_matrix12(X2,X3,X1,X4,X5,X6),X11,X12,X13))|X13=k1_matrix12(X8,X9,X10,k1_matrix12(X2,X3,X1,X4,X5,X6),X11,X12)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v6_struct_0(X10)|k3_finseq_1(k1_matrix12(X2,X3,X1,X4,X5,X6))!=k3_finseq_1(X13)|~epred1_2(X9,X8)|~m2_finseq_1(k6_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7),u1_struct_0(X1))|~m1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6),u1_struct_0(X10),X8,X9)|~m1_matrix_1(X13,u1_struct_0(X10),X8,X9)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X10)|~l6_algstr_0(X1)|~v5_vectsp_1(X10)|~v5_vectsp_1(X1)|~v4_vectsp_1(X10)|~v4_vectsp_1(X1)|~v4_rlvect_1(X10)|~v4_rlvect_1(X1)|~v3_rlvect_1(X10)|~v3_rlvect_1(X1)|~v2_rlvect_1(X10)|~v2_rlvect_1(X1)|~v5_group_1(X10)|~v5_group_1(X1)|~v3_group_1(X10)|~v3_group_1(X1)|~v33_algstr_0(X10)|~v33_algstr_0(X1)|~v13_algstr_0(X10)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_6466, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k5_numbers))).#trainneg
cnf(i_0_6463, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k5_numbers))).#trainneg
cnf(i_0_6499, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6461, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(X2,X3,X4)|v1_xboole_0(X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)))!=k3_matrix_1(X2,X3,X4,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(X2,X3,X4))|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3)|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)),k5_numbers))).#trainneg
cnf(i_0_4229, plain, (v3_ordinal1(X1)|~m2_finseq_1(X1,X2)|~v7_ordinal1(k1_zfmisc_1(esk12_1(k3_finseq_2(X3))))|~v1_xboole_0(X2)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1610, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,esk11_2(X7,k1_matrix_1(X4)))=k3_matrix_1(u1_struct_0(X1),X4,X6,esk11_2(X7,k1_matrix_1(X4)))|v6_struct_0(X1)|~m2_subset_1(esk11_2(X7,k1_matrix_1(X4)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk11_2(X7,k1_matrix_1(X4)))|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X4))|~v7_ordinal1(X6)|~v7_ordinal1(esk11_2(X7,k1_matrix_1(X4)))|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X7,a_1_0_finseq_1(k1_matrix_1(X4))))).#trainneg
cnf(i_0_273, plain, (m1_subset_1(X1,k1_numbers)|v1_xboole_0(k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6492, plain, (v1_relat_1(k8_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_6495, plain, (m2_finseq_1(k8_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1584, plain, (k1_funct_1(X1,esk16_7(X2,X3,X4,k1_matrix12(X5,X6,X7,X8,X9,X10),X11,X12,X13))=k3_matrix_1(u1_struct_0(X7),k1_matrix12(X5,X6,X7,X8,X9,X10),X14,esk16_7(X2,X3,X4,k1_matrix12(X5,X6,X7,X8,X9,X10),X11,X12,X13))|X13=k1_matrix12(X2,X3,X4,k1_matrix12(X5,X6,X7,X8,X9,X10),X11,X12)|v1_xboole_0(u1_struct_0(X7))|v6_struct_0(X7)|v6_struct_0(X4)|X1!=k8_matrix_1(u1_struct_0(X7),k1_matrix12(X5,X6,X7,X8,X9,X10),X14)|k3_finseq_1(k1_matrix12(X5,X6,X7,X8,X9,X10))!=k3_finseq_1(X13)|~epred1_2(X3,X2)|~m2_finseq_1(X1,u1_struct_0(X7))|~m1_matrix_1(k1_matrix12(X5,X6,X7,X8,X9,X10),u1_struct_0(X4),X2,X3)|~m1_matrix_1(X13,u1_struct_0(X4),X2,X3)|~m1_matrix_1(X8,u1_struct_0(X7),X5,X6)|~l6_algstr_0(X4)|~l6_algstr_0(X7)|~v5_vectsp_1(X4)|~v5_vectsp_1(X7)|~v4_vectsp_1(X4)|~v4_vectsp_1(X7)|~v4_rlvect_1(X4)|~v4_rlvect_1(X7)|~v3_rlvect_1(X4)|~v3_rlvect_1(X7)|~v2_rlvect_1(X4)|~v2_rlvect_1(X7)|~v5_group_1(X4)|~v5_group_1(X7)|~v3_group_1(X4)|~v3_group_1(X7)|~v33_algstr_0(X4)|~v33_algstr_0(X7)|~v13_algstr_0(X4)|~v13_algstr_0(X7)|~v7_ordinal1(X14)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X10)|~v7_ordinal1(X9))).#trainneg
cnf(i_0_1615, plain, (esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X4)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_6518, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k5_numbers)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))).#trainneg
cnf(i_0_6560, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))).#trainneg
cnf(i_0_6521, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k5_numbers)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))).#trainneg
cnf(i_0_6562, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k5_numbers)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))).#trainneg
cnf(i_0_6524, plain, (m1_subset_1(k1_xboole_0,k4_finseq_2(X1,X2))|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6605, negated_conjecture, (m1_subset_1(k3_finseq_1(esk6_0),k5_numbers))).#trainneg
cnf(i_0_1616, plain, (esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X4)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6613, negated_conjecture, (m1_subset_1(k1_matrix_1(esk6_0),k1_numbers))).#trainneg
cnf(i_0_1621, plain, (esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1622, plain, (esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6571, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(X2,X3,X4)|v1_xboole_0(X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)))!=k3_matrix_1(X2,X3,X4,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(X2,X3,X4))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,X3,X4)),k5_numbers)|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3)|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v1_finseq_1(k6_matrix_1(X2,X3,X4))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k6_matrix_1(X2,X3,X4))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k6_matrix_1(X2,X3,X4))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))).#trainneg
cnf(i_0_4082, plain, (esk13_2(X1,k1_zfmisc_1(k1_xboole_0))=esk12_1(k3_finseq_2(X2))|~v1_xboole_0(X2)|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_6760, negated_conjecture, (v3_ordinal1(k3_finseq_1(esk6_0)))).#trainneg
cnf(i_0_1627, plain, (r1_xxreal_0(np__1,esk12_1(a_1_0_finseq_1(X1)))|v1_xboole_0(a_1_0_finseq_1(X1))|~r2_hidden(esk12_1(a_1_0_finseq_1(X1)),a_1_0_finseq_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6761, negated_conjecture, (v7_ordinal1(k3_finseq_1(esk6_0)))).#trainneg
cnf(i_0_6615, negated_conjecture, (m1_subset_1(k3_finseq_1(esk6_0),k1_numbers))).#trainneg
cnf(i_0_1628, plain, (r1_xxreal_0(esk12_1(a_1_0_finseq_1(X1)),X1)|v1_xboole_0(a_1_0_finseq_1(X1))|~r2_hidden(esk12_1(a_1_0_finseq_1(X1)),a_1_0_finseq_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1629, plain, (m2_subset_1(esk12_1(a_1_0_finseq_1(X1)),k1_numbers,k5_numbers)|v1_xboole_0(a_1_0_finseq_1(X1))|~r2_hidden(esk12_1(a_1_0_finseq_1(X1)),a_1_0_finseq_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6512, plain, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k1_matrix12(X3,X4,X5,X6,X7,X8)))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k1_matrix12(X3,X4,X5,X6,X7,X8)))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=k1_matrix12(X3,X4,X5,X6,X7,X8)|v6_struct_0(X5)|v2_struct_0(X5)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X9))!=k3_finseq_1(X6)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k1_matrix12(X3,X4,X5,X6,X7,X8)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k1_matrix12(X3,X4,X5,X6,X7,X8)))|~v1_finseq_1(k1_matrix12(X3,X4,X5,X6,X7,X8))|~v1_funct_1(k1_matrix12(X3,X4,X5,X6,X7,X8))|~v1_relat_1(k1_matrix12(X3,X4,X5,X6,X7,X8))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k1_matrix12(X3,X4,X5,X6,X7,X8)))|~v7_ordinal1(X1)|~v7_ordinal1(X9)|~v7_ordinal1(X2)|~m1_matrix_1(X6,u1_struct_0(X5),X3,X4)|~l6_algstr_0(X5)|~v5_vectsp_1(X5)|~v4_vectsp_1(X5)|~v4_rlvect_1(X5)|~v3_rlvect_1(X5)|~v2_rlvect_1(X5)|~v5_group_1(X5)|~v3_group_1(X5)|~v33_algstr_0(X5)|~v13_algstr_0(X5)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_6472, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k5_numbers))).#trainneg
cnf(i_0_4086, plain, (esk13_2(X1,k1_zfmisc_1(k1_xboole_0))=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X3)|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_7200, negated_conjecture, (r1_xxreal_0(k1_xboole_0,k3_finseq_1(esk6_0))|r1_xxreal_0(k3_finseq_1(esk6_0),k1_xboole_0))).#trainneg
cnf(i_0_7201, negated_conjecture, (r1_xxreal_0(np__1,k3_finseq_1(esk6_0))|r1_xxreal_0(k3_finseq_1(esk6_0),np__1))).#trainneg
cnf(i_0_7202, negated_conjecture, (r1_xxreal_0(esk1_0,k3_finseq_1(esk6_0))|r1_xxreal_0(k3_finseq_1(esk6_0),esk1_0))).#trainneg
cnf(i_0_1648, plain, (m1_subset_1(esk13_2(X1,esk13_2(X2,k1_zfmisc_1(X3))),X3)|~r2_hidden(X2,k1_zfmisc_1(X3))|~r2_hidden(X1,esk13_2(X2,k1_zfmisc_1(X3))))).#trainneg
cnf(i_0_1649, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk13_2(X3,k1_zfmisc_1(X2)))|~r2_hidden(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,esk13_2(X3,k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_7237, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|esk3_0=esk8_0|esk4_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k5_numbers)|~v7_ordinal1(X1)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))).#trainneg
cnf(i_0_6501, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X4))!=k1_matrix_1(esk6_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_6500, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X4))!=k1_matrix_1(esk6_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_6607, plain, (m1_subset_1(k3_finseq_1(X1),k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4087, plain, (esk13_2(X1,k1_zfmisc_1(k1_xboole_0))=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X2)|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_7457, plain, (m1_subset_1(k3_finseq_1(k1_xboole_0),k5_numbers))).#trainneg
cnf(i_0_7377, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk4_0=esk8_0|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_7203, negated_conjecture, (r1_xxreal_0(esk2_0,k3_finseq_1(esk6_0))|r1_xxreal_0(k3_finseq_1(esk6_0),esk2_0))).#trainneg
cnf(i_0_7204, negated_conjecture, (r1_xxreal_0(esk3_0,k3_finseq_1(esk6_0))|r1_xxreal_0(k3_finseq_1(esk6_0),esk3_0))).#trainneg
cnf(i_0_1652, plain, (m1_subset_1(esk12_1(k4_finseq_1(X1)),k5_numbers)|v1_xboole_0(k4_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7394, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_7428, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_7382, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk4_0=esk8_0|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_7376, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk4_0=esk8_0|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_7380, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk4_0=esk8_0|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|~v7_ordinal1(X1)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_7379, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)|esk4_0=esk8_0|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)),k5_numbers)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))).#trainneg
cnf(i_0_7378, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk4_0=esk8_0|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)),k5_numbers)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))).#trainneg
cnf(i_0_5601, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))!=k1_funct_1(X2,esk10_2(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(X2)|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~m2_subset_1(esk10_2(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(esk10_2(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,esk10_2(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(esk10_2(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2)))).#trainneg
cnf(i_0_5647, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v7_ordinal1(X1)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k1_matrix_1(esk6_0)!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2729, plain, (m1_subset_1(esk13_2(X1,k3_finseq_2(X2)),k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2)|~r2_hidden(X1,k3_finseq_2(X2)))).#trainneg
cnf(i_0_7205, negated_conjecture, (r1_xxreal_0(esk4_0,k3_finseq_1(esk6_0))|r1_xxreal_0(k3_finseq_1(esk6_0),esk4_0))).#trainneg
cnf(i_0_1656, plain, (v3_ordinal1(k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~v3_ordinal1(u1_struct_0(X1))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_7206, negated_conjecture, (r1_xxreal_0(esk8_0,k3_finseq_1(esk6_0))|r1_xxreal_0(k3_finseq_1(esk6_0),esk8_0))).#trainneg
cnf(i_0_7207, negated_conjecture, (r1_xxreal_0(esk12_1(k5_numbers),k3_finseq_1(esk6_0))|r1_xxreal_0(k3_finseq_1(esk6_0),esk12_1(k5_numbers)))).#trainneg
cnf(i_0_1691, plain, (m2_finseq_1(k8_matrix_1(X1,esk12_1(k4_finseq_2(X2,k3_finseq_2(X1))),X3),X1)|v1_xboole_0(X1)|~v1_matrix_1(esk12_1(k4_finseq_2(X2,k3_finseq_2(X1))))|~v7_ordinal1(k1_matrix_1(esk12_1(k4_finseq_2(X2,k3_finseq_2(X1)))))|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1692, plain, (m2_finseq_1(k8_matrix_1(X1,esk12_1(k3_finseq_2(k3_finseq_2(X1))),X2),X1)|v1_xboole_0(X1)|~v1_matrix_1(esk12_1(k3_finseq_2(k3_finseq_2(X1))))|~v7_ordinal1(k1_matrix_1(esk12_1(k3_finseq_2(k3_finseq_2(X1)))))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6488, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|~v7_ordinal1(X1)|~m1_subset_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)),k5_numbers))).#trainneg
cnf(i_0_5649, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k1_matrix_1(X4)!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1797, plain, (esk12_1(k4_finseq_2(X1,X2))=esk12_1(k3_finseq_2(X3))|~v1_xboole_0(X3)|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7567, plain, (v3_ordinal1(k3_finseq_1(k1_xboole_0)))).#trainneg
cnf(i_0_7583, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|esk4_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),u1_struct_0(esk5_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7568, plain, (v7_ordinal1(k3_finseq_1(k1_xboole_0)))).#trainneg
cnf(i_0_7580, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|esk4_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),u1_struct_0(esk5_0)))).#trainneg
cnf(i_0_1723, plain, (r1_xxreal_0(k1_xboole_0,esk12_1(k3_finseq_2(X1)))|r1_xxreal_0(esk12_1(k3_finseq_2(X1)),k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1724, negated_conjecture, (r1_xxreal_0(esk1_0,esk12_1(k3_finseq_2(X1)))|r1_xxreal_0(esk12_1(k3_finseq_2(X1)),esk1_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7585, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|esk4_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers))).#trainneg
cnf(i_0_7593, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))!=k1_funct_1(X2,esk10_2(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(X2)|~m2_subset_1(esk10_2(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(esk10_2(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~r1_xxreal_0(np__1,esk10_2(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_7590, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))!=k1_funct_1(X2,esk10_2(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(X2)|~m2_subset_1(esk10_2(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(esk10_2(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~r1_xxreal_0(np__1,esk10_2(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_7581, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|esk4_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),u1_struct_0(esk5_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1799, plain, (esk12_1(k4_finseq_2(X1,X2))=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X4)|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7774, plain, (r1_xxreal_0(k1_xboole_0,k3_finseq_1(k1_xboole_0))|r1_xxreal_0(k3_finseq_1(k1_xboole_0),k1_xboole_0))).#trainneg
cnf(i_0_1725, negated_conjecture, (r1_xxreal_0(esk2_0,esk12_1(k3_finseq_2(X1)))|r1_xxreal_0(esk12_1(k3_finseq_2(X1)),esk2_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7775, plain, (r1_xxreal_0(np__1,k3_finseq_1(k1_xboole_0))|r1_xxreal_0(k3_finseq_1(k1_xboole_0),np__1))).#trainneg
cnf(i_0_7776, negated_conjecture, (r1_xxreal_0(esk1_0,k3_finseq_1(k1_xboole_0))|r1_xxreal_0(k3_finseq_1(k1_xboole_0),esk1_0))).#trainneg
cnf(i_0_1726, negated_conjecture, (r1_xxreal_0(esk3_0,esk12_1(k3_finseq_2(X1)))|r1_xxreal_0(esk12_1(k3_finseq_2(X1)),esk3_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1727, negated_conjecture, (r1_xxreal_0(esk4_0,esk12_1(k3_finseq_2(X1)))|r1_xxreal_0(esk12_1(k3_finseq_2(X1)),esk4_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7783, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),u1_struct_0(esk5_0)))).#trainneg
cnf(i_0_7784, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))))).#trainneg
cnf(i_0_7800, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))))).#trainneg
cnf(i_0_7591, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))!=k1_funct_1(X2,esk10_2(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(X2)|~m2_subset_1(esk10_2(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(esk10_2(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_matrix_1(X3))|~r1_xxreal_0(np__1,esk10_2(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_1800, plain, (esk12_1(k4_finseq_2(X1,X2))=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X3)|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7777, negated_conjecture, (r1_xxreal_0(esk2_0,k3_finseq_1(k1_xboole_0))|r1_xxreal_0(k3_finseq_1(k1_xboole_0),esk2_0))).#trainneg
cnf(i_0_7943, negated_conjecture, (esk3_0=esk8_0|esk4_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))))).#trainneg
cnf(i_0_1728, negated_conjecture, (r1_xxreal_0(esk8_0,esk12_1(k3_finseq_2(X1)))|r1_xxreal_0(esk12_1(k3_finseq_2(X1)),esk8_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7778, negated_conjecture, (r1_xxreal_0(esk3_0,k3_finseq_1(k1_xboole_0))|r1_xxreal_0(k3_finseq_1(k1_xboole_0),esk3_0))).#trainneg
cnf(i_0_7779, negated_conjecture, (r1_xxreal_0(esk4_0,k3_finseq_1(k1_xboole_0))|r1_xxreal_0(k3_finseq_1(k1_xboole_0),esk4_0))).#trainneg
cnf(i_0_1729, plain, (r1_xxreal_0(X1,esk12_1(k3_finseq_2(X2)))|r1_xxreal_0(esk12_1(k3_finseq_2(X2)),X1)|~v7_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1730, plain, (X1=esk13_2(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_xboole_0(X1)|~v1_xboole_0(X4)|~r2_hidden(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_7947, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|esk4_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).#trainneg
cnf(i_0_7752, plain, (m1_subset_1(esk12_1(k3_finseq_2(X3)),k4_finseq_2(X1,X2))|~v1_xboole_0(X3)|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7780, negated_conjecture, (r1_xxreal_0(esk8_0,k3_finseq_1(k1_xboole_0))|r1_xxreal_0(k3_finseq_1(k1_xboole_0),esk8_0))).#trainneg
cnf(i_0_1731, plain, (v7_ordinal1(esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))|~v1_xboole_0(X3)|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_7781, plain, (r1_xxreal_0(esk12_1(k5_numbers),k3_finseq_1(k1_xboole_0))|r1_xxreal_0(k3_finseq_1(k1_xboole_0),esk12_1(k5_numbers)))).#trainneg
cnf(i_0_7584, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|esk4_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),u1_struct_0(esk5_0))|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_1732, plain, (esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=k1_xboole_0|~v1_xboole_0(X3)|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1733, plain, (esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=k2_zfmisc_1(X4,X5)|~v1_xboole_0(X5)|~v1_xboole_0(X3)|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_8102, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).#trainneg
cnf(i_0_8140, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))))).#trainneg
cnf(i_0_7898, plain, (m1_subset_1(k2_zfmisc_1(X3,X4),k4_finseq_2(X1,X2))|~v1_xboole_0(X4)|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1734, plain, (esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=k2_zfmisc_1(X4,X5)|~v1_xboole_0(X4)|~v1_xboole_0(X3)|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_8300, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|esk4_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).#trainneg
cnf(i_0_1735, plain, (r1_xxreal_0(np__1,esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)))|r1_xxreal_0(esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)),np__1)|k1_matrix12(X1,X2,X3,X7,X8,X9)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(k1_matrix12(X1,X2,X3,X7,X8,X9))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1736, plain, (r1_xxreal_0(np__1,esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)))|r1_xxreal_0(esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)),np__1)|k1_matrix12(X1,X2,X3,X7,X8,X9)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(k1_matrix12(X1,X2,X3,X7,X8,X9))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8301, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|esk4_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),u1_struct_0(esk5_0))|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_7574, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_7578, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_8046, plain, (m1_subset_1(k2_zfmisc_1(X3,X4),k4_finseq_2(X1,X2))|~v1_xboole_0(X3)|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1737, plain, (r1_xxreal_0(np__1,esk13_2(X1,k5_numbers))|r1_xxreal_0(esk13_2(X1,k5_numbers),np__1)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_6523, plain, (r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))|v1_xboole_0(k1_zfmisc_1(k1_xboole_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_7171, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,esk13_2(X3,k1_zfmisc_1(k1_xboole_0)))|~r2_hidden(X3,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_1738, plain, (r1_xxreal_0(np__1,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|r1_xxreal_0(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),np__1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1739, plain, (r1_xxreal_0(np__1,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|r1_xxreal_0(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),np__1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_6293, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k1_matrix_1(X2))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,k1_matrix_1(X2))|~r1_xxreal_0(np__1,k1_matrix_1(X2))|~v7_ordinal1(X1)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_5668, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,esk6_0,X4,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,esk6_0,X4)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5)))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(esk6_0,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5)))|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_7686, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,esk6_0,X4,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,esk6_0,X4)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(esk6_0,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_4452, plain, (v3_ordinal1(X1)|~m2_finseq_1(X1,X2)|~v7_ordinal1(k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_xboole_0(X2)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_7177, plain, (m2_finseq_1(esk13_2(X2,k1_zfmisc_1(k1_xboole_0)),X1)|~v1_xboole_0(X1)|~r2_hidden(X2,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_8492, plain, (v1_xboole_0(k1_zfmisc_1(k1_xboole_0))|r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_8455, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).#trainneg
cnf(i_0_1741, plain, (r1_xxreal_0(np__1,esk12_1(k3_finseq_2(X1)))|r1_xxreal_0(esk12_1(k3_finseq_2(X1)),np__1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_8460, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_1743, plain, (r1_xxreal_0(np__1,esk9_4(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8))|r1_xxreal_0(esk9_4(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8),np__1)|X8=k6_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6))!=k3_finseq_1(X8)|~m2_finseq_1(X8,u1_struct_0(X1))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_8577, negated_conjecture, (k1_funct_1(X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(X4,esk6_0,X5,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|v1_xboole_0(X4)|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|X1!=k6_matrix_1(X4,esk6_0,X5)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~m2_finseq_1(esk6_0,k3_finseq_2(X4))|~m2_finseq_1(X1,X4)|~v7_ordinal1(X5)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v7_ordinal1(X2)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(X3)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(X3)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))).#trainneg
cnf(i_0_8461, negated_conjecture, (k1_funct_1(k6_matrix_1(X1,esk6_0,X2),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3),X4))=k3_matrix_1(X1,esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3),X4))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)=X4|v1_xboole_0(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))!=k3_finseq_1(X4)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3),X4),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3),X4))|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4)|~m2_finseq_1(esk6_0,k3_finseq_2(X1))|~m2_finseq_1(k6_matrix_1(X1,esk6_0,X2),X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_8575, negated_conjecture, (k1_funct_1(X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(X4,esk6_0,X5,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|v1_xboole_0(X4)|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|X1!=k6_matrix_1(X4,esk6_0,X5)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~m2_finseq_1(esk6_0,k3_finseq_2(X4))|~m2_finseq_1(X1,X4)|~v7_ordinal1(X5)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v7_ordinal1(X6)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6))!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(X3)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_8574, negated_conjecture, (k1_funct_1(X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(X4,esk6_0,X5,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|v1_xboole_0(X4)|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|X1!=k6_matrix_1(X4,esk6_0,X5)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~m2_finseq_1(esk6_0,k3_finseq_2(X4))|~m2_finseq_1(X1,X4)|~v7_ordinal1(X5)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v7_ordinal1(X6)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6))!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2384, plain, (v3_ordinal1(esk13_2(X1,esk12_1(k1_zfmisc_1(X2))))|~v3_ordinal1(X2)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_7454, plain, (m1_subset_1(k3_finseq_1(esk12_1(k3_finseq_2(X1))),k5_numbers))).#trainneg
cnf(i_0_1744, plain, (r1_xxreal_0(np__1,esk9_4(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4,X5))|r1_xxreal_0(esk9_4(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4,X5),np__1)|X5=k6_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4)|v1_xboole_0(k3_finseq_2(X1))|v1_xboole_0(X1)|k1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))!=k3_finseq_1(X5)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(X1)))|~m2_finseq_1(X5,X1)|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))|~v1_matrix_1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1751, plain, (r1_xxreal_0(np__1,k2_zfmisc_1(X1,X2))|r1_xxreal_0(k2_zfmisc_1(X1,X2),np__1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1752, plain, (r1_xxreal_0(np__1,k2_zfmisc_1(X1,X2))|r1_xxreal_0(k2_zfmisc_1(X1,X2),np__1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_8640, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|esk4_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).#trainneg
cnf(i_0_8641, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|esk4_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).#trainneg
cnf(i_0_8650, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|esk4_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_8657, negated_conjecture, (k1_funct_1(k6_matrix_1(X1,esk6_0,X2),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3),X4))=k3_matrix_1(X1,esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3),X4))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)=X4|v1_xboole_0(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))!=k3_finseq_1(X4)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3),X4),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3),X4))|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4)|~m2_finseq_1(esk6_0,k3_finseq_2(X1))|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_8654, negated_conjecture, (k1_funct_1(k6_matrix_1(X1,esk6_0,X2),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3),X4))=k3_matrix_1(X1,esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3),X4))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)=X4|v1_xboole_0(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))!=k3_finseq_1(X4)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3),X4),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3),X4))|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4)|~m2_finseq_1(esk6_0,k3_finseq_2(X1))|~m2_finseq_1(k6_matrix_1(X1,esk6_0,X2),X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1758, plain, (esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=esk12_1(k3_finseq_2(X4))|~v1_xboole_0(X4)|~v1_xboole_0(X3)|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1759, plain, (esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=esk12_1(k3_finseq_2(X3))|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1760, plain, (esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=esk12_1(k3_finseq_2(X3))|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_8759, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8757, negated_conjecture, (esk4_0=esk8_0|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).#trainneg
cnf(i_0_8761, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk4_0=esk8_0|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).#trainneg
cnf(i_0_8762, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk4_0=esk8_0|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_8767, negated_conjecture, (k1_funct_1(k6_matrix_1(X1,esk6_0,X2),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3),X4))=k3_matrix_1(X1,esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3),X4))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)=X4|v1_xboole_0(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))!=k3_finseq_1(X4)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3),X4),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3),X4))|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4)|~m2_finseq_1(esk6_0,k3_finseq_2(X1))|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_8764, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k6_matrix_1(X2,esk6_0,X3)|v1_xboole_0(X2)|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,esk6_0,X3)))!=k3_matrix_1(X2,esk6_0,X3,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,esk6_0,X3)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k6_matrix_1(X2,esk6_0,X3))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,esk6_0,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,esk6_0,X3)))|~v1_finseq_1(k6_matrix_1(X2,esk6_0,X3))|~v1_funct_1(k6_matrix_1(X2,esk6_0,X3))|~v1_relat_1(k6_matrix_1(X2,esk6_0,X3))|~v7_ordinal1(X1)|~m2_finseq_1(esk6_0,k3_finseq_2(X2))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_8763, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k6_matrix_1(X2,esk6_0,X3)|v1_xboole_0(X2)|k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,esk6_0,X3)))!=k3_matrix_1(X2,esk6_0,X3,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,esk6_0,X3)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k6_matrix_1(X2,esk6_0,X3))|~v1_finseq_1(k6_matrix_1(X2,esk6_0,X3))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k6_matrix_1(X2,esk6_0,X3))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k6_matrix_1(X2,esk6_0,X3))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,esk6_0,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,esk6_0,X3)))|~m2_finseq_1(esk6_0,k3_finseq_2(X2))|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8878, plain, (m2_finseq_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3))),X1)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1765, plain, (X1=esk13_2(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_xboole_0(X1)|~v1_xboole_0(X3)|~r2_hidden(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_9053, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|esk4_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).#trainneg
cnf(i_0_9054, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|esk4_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_1766, plain, (v7_ordinal1(esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))|~v1_xboole_0(X2)|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1767, plain, (esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=esk12_1(k3_finseq_2(X4))|~v1_xboole_0(X4)|~v1_xboole_0(X2)|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_9122, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk4_0=esk8_0|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_9123, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk4_0=esk8_0|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_8977, plain, (m2_finseq_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3))),X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_9295, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|esk4_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_9242, negated_conjecture, (esk3_0=esk8_0|esk4_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1768, plain, (esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=k1_xboole_0|~v1_xboole_0(X2)|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_9302, negated_conjecture, (v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|esk3_0=esk8_0|esk4_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_9303, negated_conjecture, (v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|esk3_0=esk8_0|esk4_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_1769, plain, (esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=k2_zfmisc_1(X4,X5)|~v1_xboole_0(X5)|~v1_xboole_0(X2)|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1770, plain, (esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=k2_zfmisc_1(X4,X5)|~v1_xboole_0(X4)|~v1_xboole_0(X2)|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_9393, negated_conjecture, (esk4_0=esk8_0|esk3_0=esk8_0|k1_matrix_1(X1)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_7747, plain, (m2_finseq_1(esk12_1(k4_finseq_2(X2,X3)),X1)|~v1_xboole_0(X1)|~v1_xboole_0(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_9304, negated_conjecture, (v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|esk3_0=esk8_0|esk4_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_9404, negated_conjecture, (esk4_0=esk8_0|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1801, plain, (X1=k6_matrix_1(X2,esk13_2(X3,k3_finseq_2(k3_finseq_2(X2))),X4)|v1_xboole_0(X2)|v7_ordinal1(esk9_4(X2,esk13_2(X3,k3_finseq_2(k3_finseq_2(X2))),X4,X1))|k1_matrix_1(esk13_2(X3,k3_finseq_2(k3_finseq_2(X2))))!=k3_finseq_1(X1)|~m2_finseq_1(X1,X2)|~v1_matrix_1(esk13_2(X3,k3_finseq_2(k3_finseq_2(X2))))|~v7_ordinal1(X4)|~r2_hidden(X3,k3_finseq_2(k3_finseq_2(X2))))).#trainneg
cnf(i_0_9305, negated_conjecture, (m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),u1_struct_0(esk5_0))|esk3_0=esk8_0|esk4_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_9445, negated_conjecture, (m1_subset_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|esk4_0=esk8_0|esk3_0=esk8_0|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_1802, plain, (k6_matrix_1(X1,esk13_2(X2,k3_finseq_2(k3_finseq_2(X1))),X3)=k8_matrix_1(X1,esk13_2(X2,k3_finseq_2(k3_finseq_2(X1))),X3)|v1_xboole_0(X1)|~v1_matrix_1(esk13_2(X2,k3_finseq_2(k3_finseq_2(X1))))|~v7_ordinal1(X3)|~r2_hidden(X2,k3_finseq_2(k3_finseq_2(X1))))).#trainneg
cnf(i_0_1803, plain, (m1_subset_1(k3_matrix_1(X1,esk13_2(X2,k3_finseq_2(k3_finseq_2(X1))),X3,X4),X1)|~v1_matrix_1(esk13_2(X2,k3_finseq_2(k3_finseq_2(X1))))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X2,k3_finseq_2(k3_finseq_2(X1))))).#trainneg
cnf(i_0_9401, negated_conjecture, (esk4_0=esk8_0|esk3_0=esk8_0|k1_matrix_1(X2)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v7_ordinal1(X1)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_9056, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(X2,esk6_0,X3)|v1_xboole_0(X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,esk6_0,X3)))!=k3_matrix_1(X2,esk6_0,X3,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,esk6_0,X3)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(X2,esk6_0,X3))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,esk6_0,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,esk6_0,X3)))|~v1_finseq_1(k6_matrix_1(X2,esk6_0,X3))|~v1_funct_1(k6_matrix_1(X2,esk6_0,X3))|~v1_relat_1(k6_matrix_1(X2,esk6_0,X3))|~v7_ordinal1(X1)|~m2_finseq_1(esk6_0,k3_finseq_2(X2))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_9055, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(X2,esk6_0,X3)|v1_xboole_0(X2)|k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,esk6_0,X3)))!=k3_matrix_1(X2,esk6_0,X3,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,esk6_0,X3)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(X2,esk6_0,X3))|~v1_finseq_1(k6_matrix_1(X2,esk6_0,X3))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k6_matrix_1(X2,esk6_0,X3))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k6_matrix_1(X2,esk6_0,X3))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,esk6_0,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,esk6_0,X3)))|~m2_finseq_1(esk6_0,k3_finseq_2(X2))|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9059, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k6_matrix_1(X2,esk6_0,X3)|v1_xboole_0(X2)|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,esk6_0,X3)))!=k3_matrix_1(X2,esk6_0,X3,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,esk6_0,X3)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k6_matrix_1(X2,esk6_0,X3))|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,esk6_0,X3)))|~v1_finseq_1(k6_matrix_1(X2,esk6_0,X3))|~v1_funct_1(k6_matrix_1(X2,esk6_0,X3))|~v1_relat_1(k6_matrix_1(X2,esk6_0,X3))|~m2_finseq_1(esk6_0,k3_finseq_2(X2))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,esk6_0,X3)),k5_numbers))).#trainneg
cnf(i_0_9881, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(X2,esk6_0,X3)|v1_xboole_0(X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,esk6_0,X3)))!=k3_matrix_1(X2,esk6_0,X3,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,esk6_0,X3)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(X2,esk6_0,X3))|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,esk6_0,X3)))|~v1_finseq_1(k6_matrix_1(X2,esk6_0,X3))|~v1_funct_1(k6_matrix_1(X2,esk6_0,X3))|~v1_relat_1(k6_matrix_1(X2,esk6_0,X3))|~m2_finseq_1(esk6_0,k3_finseq_2(X2))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,esk6_0,X3)),k5_numbers))).#trainneg
cnf(i_0_8703, plain, (v3_ordinal1(k3_finseq_1(esk12_1(k3_finseq_2(X1)))))).#trainneg
cnf(i_0_1804, plain, (m2_finseq_1(k8_matrix_1(X1,esk13_2(X2,k3_finseq_2(k3_finseq_2(X1))),X3),X1)|v1_xboole_0(X1)|~v1_matrix_1(esk13_2(X2,k3_finseq_2(k3_finseq_2(X1))))|~v7_ordinal1(k1_matrix_1(esk13_2(X2,k3_finseq_2(k3_finseq_2(X1)))))|~v7_ordinal1(X3)|~r2_hidden(X2,k3_finseq_2(k3_finseq_2(X1))))).#trainneg
cnf(i_0_1805, plain, (k1_matrix_1(esk13_2(X1,k3_finseq_2(k3_finseq_2(X2))))=k3_finseq_1(X3)|v1_xboole_0(X2)|X3!=k6_matrix_1(X2,esk13_2(X1,k3_finseq_2(k3_finseq_2(X2))),X4)|~m2_finseq_1(X3,X2)|~v1_matrix_1(esk13_2(X1,k3_finseq_2(k3_finseq_2(X2))))|~v7_ordinal1(X4)|~r2_hidden(X1,k3_finseq_2(k3_finseq_2(X2))))).#trainneg
cnf(i_0_1806, plain, (k1_funct_1(X1,esk16_7(X2,X3,X4,esk13_2(X5,k3_finseq_2(k3_finseq_2(X6))),X7,X8,X9))=k3_matrix_1(X6,esk13_2(X5,k3_finseq_2(k3_finseq_2(X6))),X10,esk16_7(X2,X3,X4,esk13_2(X5,k3_finseq_2(k3_finseq_2(X6))),X7,X8,X9))|X9=k1_matrix12(X2,X3,X4,esk13_2(X5,k3_finseq_2(k3_finseq_2(X6))),X7,X8)|v1_xboole_0(X6)|v6_struct_0(X4)|v2_struct_0(X4)|X1!=k6_matrix_1(X6,esk13_2(X5,k3_finseq_2(k3_finseq_2(X6))),X10)|k3_finseq_1(esk13_2(X5,k3_finseq_2(k3_finseq_2(X6))))!=k3_finseq_1(X9)|~epred1_2(X3,X2)|~m2_finseq_1(X1,X6)|~v1_matrix_1(esk13_2(X5,k3_finseq_2(k3_finseq_2(X6))))|~m1_matrix_1(X9,u1_struct_0(X4),X2,X3)|~m1_matrix_1(esk13_2(X5,k3_finseq_2(k3_finseq_2(X6))),u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X10)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~r2_hidden(X5,k3_finseq_2(k3_finseq_2(X6))))).#trainneg
cnf(i_0_4424, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|X1=esk3_0|X1=esk4_0|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~r2_hidden(X1,k4_finseq_1(esk6_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v7_ordinal1(X1)|~v7_ordinal1(X4)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4))!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(X3)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_8704, plain, (v7_ordinal1(k3_finseq_1(esk12_1(k3_finseq_2(X1)))))).#trainneg
cnf(i_0_1807, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|~r2_hidden(X1,esk13_2(X3,k3_finseq_2(X2)))|~r2_hidden(X3,k3_finseq_2(X2)))).#trainneg
cnf(i_0_7329, plain, (~r2_hidden(X1,k2_zfmisc_1(X3,X4))|~r2_hidden(X2,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X4))).#trainneg
cnf(i_0_1831, plain, (r1_xxreal_0(k2_zfmisc_1(X1,esk13_2(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))),X5)|r1_xxreal_0(X5,k2_zfmisc_1(X1,esk13_2(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))))|~v7_ordinal1(X5)|~v1_xboole_0(X4)|~r2_hidden(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_9915, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k6_matrix_1(X2,esk6_0,X3)|v1_xboole_0(X2)|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,esk6_0,X3)))!=k3_matrix_1(X2,esk6_0,X3,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,esk6_0,X3)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k6_matrix_1(X2,esk6_0,X3))|~v1_finseq_1(k6_matrix_1(X2,esk6_0,X3))|~v1_funct_1(k6_matrix_1(X2,esk6_0,X3))|~v1_relat_1(k6_matrix_1(X2,esk6_0,X3))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,esk6_0,X3)),k5_numbers)|~m2_finseq_1(esk6_0,k3_finseq_2(X2))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))).#trainneg
cnf(i_0_9917, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(X2,esk6_0,X3)|v1_xboole_0(X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,esk6_0,X3)))!=k3_matrix_1(X2,esk6_0,X3,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,esk6_0,X3)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(X2,esk6_0,X3))|~v1_finseq_1(k6_matrix_1(X2,esk6_0,X3))|~v1_funct_1(k6_matrix_1(X2,esk6_0,X3))|~v1_relat_1(k6_matrix_1(X2,esk6_0,X3))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(X2,esk6_0,X3)),k5_numbers)|~m2_finseq_1(esk6_0,k3_finseq_2(X2))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))).#trainneg
cnf(i_0_2620, plain, (v3_ordinal1(esk13_2(X1,k3_finseq_2(X2)))|~v1_xboole_0(X2)|~v7_ordinal1(k1_zfmisc_1(k1_xboole_0))|~r2_hidden(X1,k3_finseq_2(X2)))).#trainneg
cnf(i_0_1832, plain, (r1_xxreal_0(k2_zfmisc_1(X1,esk13_2(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))),X5)|r1_xxreal_0(X5,k2_zfmisc_1(X1,esk13_2(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))))|~v7_ordinal1(X5)|~v1_xboole_0(X3)|~r2_hidden(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_1833, plain, (r1_xxreal_0(k2_zfmisc_1(X1,esk12_1(k4_finseq_2(X2,X3))),X4)|r1_xxreal_0(X4,k2_zfmisc_1(X1,esk12_1(k4_finseq_2(X2,X3))))|~v7_ordinal1(X4)|~v1_xboole_0(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1834, plain, (r1_xxreal_0(k2_zfmisc_1(X1,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)))),X4)|r1_xxreal_0(X4,k2_zfmisc_1(X1,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~v7_ordinal1(X4)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_4755, plain, (v3_ordinal1(X1)|~m2_finseq_1(X1,X2)|~v7_ordinal1(k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_xboole_0(X2)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_5580, plain, (m1_subset_1(esk13_2(X1,esk12_1(k3_finseq_2(X2))),k2_zfmisc_1(k5_numbers,X2))|~r2_hidden(X1,esk12_1(k3_finseq_2(X2))))).#trainneg
cnf(i_0_1835, plain, (r1_xxreal_0(k2_zfmisc_1(X1,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)))),X4)|r1_xxreal_0(X4,k2_zfmisc_1(X1,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~v7_ordinal1(X4)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1836, plain, (r1_xxreal_0(k2_zfmisc_1(X1,esk12_1(k3_finseq_2(X2))),X3)|r1_xxreal_0(X3,k2_zfmisc_1(X1,esk12_1(k3_finseq_2(X2))))|~v7_ordinal1(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1837, plain, (r1_xxreal_0(k2_zfmisc_1(X1,k1_xboole_0),X2)|r1_xxreal_0(X2,k2_zfmisc_1(X1,k1_xboole_0))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10074, plain, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,X1)|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(esk6_0)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10306, plain, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,X1)|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,k1_matrix_1(esk6_0)))).#trainneg
cnf(i_0_10295, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,X1)|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2)))|~v7_ordinal1(X1)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_10087, plain, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X1)|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(esk6_0)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10046, plain, (esk7_0=X1|r1_xxreal_0(esk10_2(esk7_0,X1),k3_finseq_1(esk6_0))|esk3_0=esk8_0|k3_finseq_1(esk6_0)!=k3_finseq_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10297, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,X1)|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1838, plain, (r1_xxreal_0(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)),X4)|r1_xxreal_0(X4,k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))|~v7_ordinal1(X4)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1839, plain, (r1_xxreal_0(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)),X4)|r1_xxreal_0(X4,k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))|~v7_ordinal1(X4)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_10094, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|esk3_0=esk8_0|~m2_subset_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10294, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,X1)|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10083, plain, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk9_4(X1,esk6_0,X2,X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk9_4(X1,esk6_0,X2,X3))|X3=k6_matrix_1(X1,esk6_0,X2)|v1_xboole_0(X1)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(X3)|~m2_finseq_1(esk6_0,k3_finseq_2(X1))|~m2_finseq_1(X3,X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10069, plain, (k3_matrix_1(u1_struct_0(esk5_0),X1,esk8_0,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,X2)|esk3_0=esk8_0|X1!=esk7_0|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(esk6_0)))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10096, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))|esk3_0=esk8_0|~m2_subset_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10102, plain, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk9_4(X2,esk6_0,X3,X4))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk9_4(X2,esk6_0,X3,X4))|X4=k6_matrix_1(X2,esk6_0,X3)|X1=esk8_0|X1=esk3_0|v1_xboole_0(X2)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(X4)|~m2_finseq_1(esk6_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~r2_hidden(X1,k4_finseq_1(esk6_0))|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_10086, plain, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk9_4(X1,esk6_0,X2,X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk9_4(X1,esk6_0,X2,X3))|X3=k6_matrix_1(X1,esk6_0,X2)|v1_xboole_0(X1)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(X3)|~m2_finseq_1(esk6_0,k3_finseq_2(X1))|~m2_finseq_1(X3,X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10075, plain, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk16_7(X2,X3,X4,esk7_0,X5,X6,X7))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk16_7(X2,X3,X4,esk7_0,X5,X6,X7))|X7=k1_matrix12(X2,X3,X4,esk7_0,X5,X6)|v6_struct_0(X4)|esk3_0=esk8_0|k3_finseq_1(esk6_0)!=k3_finseq_1(X7)|~epred1_2(X3,X2)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~m1_matrix_1(esk7_0,u1_struct_0(X4),X2,X3)|~m1_matrix_1(X7,u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_10268, plain, (r1_xxreal_0(esk12_1(k5_numbers),k2_zfmisc_1(X1,k1_xboole_0))|r1_xxreal_0(k2_zfmisc_1(X1,k1_xboole_0),esk12_1(k5_numbers)))).#trainneg
cnf(i_0_10090, plain, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,k1_matrix_1(esk6_0))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,k1_matrix_1(esk6_0))|esk3_0=esk8_0|~r1_xxreal_0(np__1,k1_matrix_1(esk6_0)))).#trainneg
cnf(i_0_1842, plain, (esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10))=X6|k1_matrix12(X7,X8,X3,X4,X9,X10)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|v2_struct_0(X3)|esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10))=X10|esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10))=X9|k3_matrix_1(u1_struct_0(X3),X4,esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10)),esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10)))!=k3_matrix_1(u1_struct_0(X3),X4,X6,esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10)))|esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10))!=X5|k3_finseq_1(k1_matrix12(X7,X8,X3,X4,X9,X10))!=k3_finseq_1(X4)|~epred1_2(X2,X1)|~m1_matrix_1(k1_matrix12(X7,X8,X3,X4,X9,X10),u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~r2_hidden(esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10)),k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X3),X7,X8)|~v7_ordinal1(X10)|~v7_ordinal1(X9)|~v7_ordinal1(esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10)))|~v7_ordinal1(X8)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_1843, plain, (esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10))=X5|k1_matrix12(X7,X8,X3,X4,X9,X10)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|v2_struct_0(X3)|esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10))=X10|esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10))=X9|k3_matrix_1(u1_struct_0(X3),X4,esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10)),esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10)))!=k3_matrix_1(u1_struct_0(X3),X4,X5,esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10)))|esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10))!=X6|k3_finseq_1(k1_matrix12(X7,X8,X3,X4,X9,X10))!=k3_finseq_1(X4)|~epred1_2(X2,X1)|~m1_matrix_1(k1_matrix12(X7,X8,X3,X4,X9,X10),u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~r2_hidden(esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10)),k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X3),X7,X8)|~v7_ordinal1(X10)|~v7_ordinal1(X9)|~v7_ordinal1(esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10)))|~v7_ordinal1(X8)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_10310, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,X1)|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10307, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,X1)|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10353, plain, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X1)|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,k1_matrix_1(esk6_0)))).#trainneg
cnf(i_0_10409, plain, (k3_matrix_1(u1_struct_0(esk5_0),X1,esk8_0,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,X2)|esk3_0=esk8_0|X1!=esk7_0|~epred1_2(esk1_0,esk2_0)|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,k1_matrix_1(esk6_0)))).#trainneg
cnf(i_0_10308, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,X1)|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_matrix_1(X2))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_10416, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))|esk3_0=esk8_0|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))|~v7_ordinal1(X1)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))).#trainneg
cnf(i_0_10091, plain, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,k1_matrix_1(esk6_0))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,k1_matrix_1(esk6_0))|esk3_0=esk8_0|~r1_xxreal_0(np__1,k1_matrix_1(esk6_0)))).#trainneg
cnf(i_0_1876, plain, (r1_xxreal_0(k1_xboole_0,esk12_1(k4_finseq_2(X1,X2)))|r1_xxreal_0(esk12_1(k4_finseq_2(X1,X2)),k1_xboole_0)|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10317, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk12_1(k5_numbers))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk12_1(k5_numbers))|esk3_0=esk8_0|r1_xxreal_0(k1_matrix_1(esk6_0),esk12_1(k5_numbers))|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk12_1(k5_numbers),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk12_1(k5_numbers)))).#trainneg
cnf(i_0_1844, plain, (esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10))=X5|esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10))=X6|k1_matrix12(X7,X8,X3,X4,X9,X10)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|v2_struct_0(X3)|esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10))=X10|esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10))=X9|k3_finseq_1(k1_matrix12(X7,X8,X3,X4,X9,X10))!=k3_finseq_1(X4)|~epred1_2(X2,X1)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(k1_matrix12(X7,X8,X3,X4,X9,X10),u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~r2_hidden(esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10)),k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X3),X7,X8)|~v7_ordinal1(X10)|~v7_ordinal1(X9)|~v7_ordinal1(esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10)))|~v7_ordinal1(X8)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_1845, plain, (esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10))=X6|k1_matrix12(X7,X8,X3,X4,X9,X10)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|v2_struct_0(X3)|esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10))=X10|esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10))=X9|k3_matrix_1(u1_struct_0(X3),X4,esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10)),esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10)))!=k3_matrix_1(u1_struct_0(X3),X4,X6,esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10)))|k3_finseq_1(k1_matrix12(X7,X8,X3,X4,X9,X10))!=k3_finseq_1(X4)|~epred1_2(X2,X1)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(k1_matrix12(X7,X8,X3,X4,X9,X10),u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~r2_hidden(esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10)),k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X3),X7,X8)|~v7_ordinal1(X10)|~v7_ordinal1(X9)|~v7_ordinal1(esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10)))|~v7_ordinal1(X8)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_10591, plain, (esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X4,X7,X8))=X7|esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X4,X7,X8))=X8|esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X4,X7,X8))=X6|esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X4,X7,X8))=X5|k1_matrix12(X1,X2,X3,X4,X7,X8)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|v2_struct_0(X3)|~epred1_2(X2,X1)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10588, negated_conjecture, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk3_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk4_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_10511, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X1)|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10508, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X1)|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10463, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,k1_matrix_1(X1))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,k1_matrix_1(X1))|esk3_0=esk8_0|~r1_xxreal_0(np__1,k1_matrix_1(X1))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_10719, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|esk3_0=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))).#trainneg
cnf(i_0_1877, plain, (r1_xxreal_0(np__1,esk12_1(k4_finseq_2(X1,X2)))|r1_xxreal_0(esk12_1(k4_finseq_2(X1,X2)),np__1)|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1846, plain, (esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10))=X5|k1_matrix12(X7,X8,X3,X4,X9,X10)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|v2_struct_0(X3)|esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10))=X10|esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10))=X9|k3_matrix_1(u1_struct_0(X3),X4,esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10)),esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10)))!=k3_matrix_1(u1_struct_0(X3),X4,X5,esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10)))|k3_finseq_1(k1_matrix12(X7,X8,X3,X4,X9,X10))!=k3_finseq_1(X4)|~epred1_2(X2,X1)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(k1_matrix12(X7,X8,X3,X4,X9,X10),u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~r2_hidden(esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10)),k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X3),X7,X8)|~v7_ordinal1(X10)|~v7_ordinal1(X9)|~v7_ordinal1(esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X7,X8,X3,X4,X9,X10)))|~v7_ordinal1(X8)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_1847, plain, (k1_matrix12(X1,X2,X3,X4,X5,X6)=k1_matrix12(X7,X8,X3,X4,X9,X10)|v6_struct_0(X3)|v2_struct_0(X3)|esk15_7(X7,X8,X3,X4,X9,X10,k1_matrix12(X1,X2,X3,X4,X5,X6))=X6|esk15_7(X7,X8,X3,X4,X9,X10,k1_matrix12(X1,X2,X3,X4,X5,X6))=X5|k3_matrix_1(u1_struct_0(X3),X4,esk15_7(X7,X8,X3,X4,X9,X10,k1_matrix12(X1,X2,X3,X4,X5,X6)),esk16_7(X7,X8,X3,X4,X9,X10,k1_matrix12(X1,X2,X3,X4,X5,X6)))!=k3_matrix_1(u1_struct_0(X3),X4,X10,esk16_7(X7,X8,X3,X4,X9,X10,k1_matrix12(X1,X2,X3,X4,X5,X6)))|k3_matrix_1(u1_struct_0(X3),X4,esk15_7(X7,X8,X3,X4,X9,X10,k1_matrix12(X1,X2,X3,X4,X5,X6)),esk16_7(X7,X8,X3,X4,X9,X10,k1_matrix12(X1,X2,X3,X4,X5,X6)))!=k3_matrix_1(u1_struct_0(X3),X4,X9,esk16_7(X7,X8,X3,X4,X9,X10,k1_matrix12(X1,X2,X3,X4,X5,X6)))|esk15_7(X7,X8,X3,X4,X9,X10,k1_matrix12(X1,X2,X3,X4,X5,X6))!=X10|k3_finseq_1(k1_matrix12(X1,X2,X3,X4,X5,X6))!=k3_finseq_1(X4)|~epred1_2(X8,X7)|~m1_matrix_1(k1_matrix12(X1,X2,X3,X4,X5,X6),u1_struct_0(X3),X7,X8)|~m1_matrix_1(X4,u1_struct_0(X3),X7,X8)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X10)|~v7_ordinal1(X9)|~r2_hidden(esk15_7(X7,X8,X3,X4,X9,X10,k1_matrix12(X1,X2,X3,X4,X5,X6)),k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(esk15_7(X7,X8,X3,X4,X9,X10,k1_matrix12(X1,X2,X3,X4,X5,X6)))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10593, negated_conjecture, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk3_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk8_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_10592, negated_conjecture, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk3_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk4_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_10667, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X3,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5)))|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk4_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk3_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|~m2_subset_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5)))|~r2_hidden(X3,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X5)|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2))).#trainneg
cnf(i_0_10681, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X3,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5)))|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk4_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk3_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5)))|~r2_hidden(X3,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X5)|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2))).#trainneg
cnf(i_0_10706, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|esk3_0=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))).#trainneg
cnf(i_0_1878, negated_conjecture, (r1_xxreal_0(esk1_0,esk12_1(k4_finseq_2(X1,X2)))|r1_xxreal_0(esk12_1(k4_finseq_2(X1,X2)),esk1_0)|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10518, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk12_1(k5_numbers))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk12_1(k5_numbers))|esk3_0=esk8_0|r1_xxreal_0(k1_matrix_1(esk6_0),esk12_1(k5_numbers))|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk12_1(k5_numbers),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk12_1(k5_numbers)))).#trainneg
cnf(i_0_10594, negated_conjecture, (esk15_7(X1,X2,esk5_0,esk7_0,X3,X4,esk7_0)=X5|esk15_7(X1,X2,esk5_0,esk7_0,X3,X4,esk7_0)=X6|esk15_7(X1,X2,esk5_0,esk7_0,X3,X4,esk7_0)=X4|k1_matrix12(X1,X2,esk5_0,esk7_0,X3,X4)=esk7_0|esk15_7(esk2_0,esk1_0,esk5_0,esk7_0,X5,X6,esk7_0)=X5|esk15_7(esk2_0,esk1_0,esk5_0,esk7_0,X5,X6,esk7_0)=X6|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk15_7(X1,X2,esk5_0,esk7_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk7_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X4,esk16_7(X1,X2,esk5_0,esk7_0,X3,X4,esk7_0))|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_10678, plain, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X7|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X8|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=esk4_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=esk3_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=X7|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=X8|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X5,X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~epred1_2(X6,X5)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X5,X6))).#trainneg
cnf(i_0_10569, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,k1_matrix_1(X1))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,k1_matrix_1(X1))|esk3_0=esk8_0|~r1_xxreal_0(np__1,k1_matrix_1(X1))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_1879, negated_conjecture, (r1_xxreal_0(esk2_0,esk12_1(k4_finseq_2(X1,X2)))|r1_xxreal_0(esk12_1(k4_finseq_2(X1,X2)),esk2_0)|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9896, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk12_1(k2_finseq_1(k1_matrix_1(esk6_0))))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk12_1(k2_finseq_1(k1_matrix_1(esk6_0))))|v1_xboole_0(k2_finseq_1(k1_matrix_1(esk6_0)))|esk3_0=esk8_0|~v7_ordinal1(esk12_1(k2_finseq_1(k1_matrix_1(esk6_0)))))).#trainneg
cnf(i_0_10098, plain, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk12_1(k2_finseq_1(k1_matrix_1(esk6_0))))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk12_1(k2_finseq_1(k1_matrix_1(esk6_0))))|v1_xboole_0(k2_finseq_1(k1_matrix_1(esk6_0)))|esk3_0=esk8_0|~v7_ordinal1(esk12_1(k2_finseq_1(k1_matrix_1(esk6_0)))))).#trainneg
cnf(i_0_1850, plain, (m1_subset_1(k3_matrix_1(u1_struct_0(X1),X4,X7,esk16_7(X8,X9,X10,X4,X11,X12,X13)),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|X13=k1_matrix12(X8,X9,X10,X4,X11,X12)|X7=X6|X7=X5|v6_struct_0(X10)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk16_7(X8,X9,X10,X4,X11,X12,X13))|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|k3_finseq_1(X4)!=k3_finseq_1(X13)|~epred1_2(X9,X8)|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X13,u1_struct_0(X10),X8,X9)|~m1_matrix_1(X4,u1_struct_0(X10),X8,X9)|~l6_algstr_0(X10)|~v5_vectsp_1(X10)|~v4_vectsp_1(X10)|~v4_rlvect_1(X10)|~v3_rlvect_1(X10)|~v2_rlvect_1(X10)|~v5_group_1(X10)|~v3_group_1(X10)|~v33_algstr_0(X10)|~v13_algstr_0(X10)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_1880, negated_conjecture, (r1_xxreal_0(esk3_0,esk12_1(k4_finseq_2(X1,X2)))|r1_xxreal_0(esk12_1(k4_finseq_2(X1,X2)),esk3_0)|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10589, negated_conjecture, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk3_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk8_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|esk3_0=esk8_0|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_10855, plain, (k3_finseq_1(X1)=k3_finseq_1(esk6_0)|esk15_7(X2,X3,esk5_0,esk6_0,X4,X5,esk7_0)=esk8_0|esk15_7(X2,X3,esk5_0,esk6_0,X4,X5,esk7_0)=esk3_0|esk15_7(X2,X3,esk5_0,esk6_0,X4,X5,esk7_0)=X4|esk15_7(X2,X3,esk5_0,esk6_0,X4,X5,esk7_0)=X5|esk3_0=esk8_0|X1!=esk7_0|~epred1_2(X3,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X2,X3)|~m1_matrix_1(X1,u1_struct_0(esk5_0),X2,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X2,X3))).#trainneg
cnf(i_0_10747, negated_conjecture, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk3_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk8_0|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_10886, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X3,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5)))|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk8_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk3_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|esk3_0=esk8_0|~m2_subset_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5)))|~r2_hidden(X3,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X5)|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2))).#trainneg
cnf(i_0_10884, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X3,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5)))|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk8_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk3_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|esk3_0=esk8_0|~m2_subset_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5)))|~r2_hidden(X3,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X5)|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2))).#trainneg
cnf(i_0_1881, negated_conjecture, (r1_xxreal_0(esk4_0,esk12_1(k4_finseq_2(X1,X2)))|r1_xxreal_0(esk12_1(k4_finseq_2(X1,X2)),esk4_0)|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1882, negated_conjecture, (r1_xxreal_0(esk8_0,esk12_1(k4_finseq_2(X1,X2)))|r1_xxreal_0(esk12_1(k4_finseq_2(X1,X2)),esk8_0)|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1883, plain, (r1_xxreal_0(X1,esk12_1(k4_finseq_2(X2,X3)))|r1_xxreal_0(esk12_1(k4_finseq_2(X2,X3)),X1)|~v7_ordinal1(X1)|~v1_xboole_0(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10739, negated_conjecture, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk3_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk8_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_10760, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|esk15_7(X4,X5,esk5_0,esk6_0,X1,X3,esk7_0)=esk3_0|esk15_7(X4,X5,esk5_0,esk6_0,X1,X3,esk7_0)=esk4_0|esk15_7(X4,X5,esk5_0,esk6_0,X1,X3,esk7_0)=X3|esk15_7(X4,X5,esk5_0,esk6_0,X1,X3,esk7_0)=X1|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~r2_hidden(X1,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X4,X5)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X4,X5)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))).#trainneg
cnf(i_0_10746, negated_conjecture, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk3_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk4_0|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_10782, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))|esk3_0=esk8_0|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10738, negated_conjecture, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk3_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk4_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1884, plain, (m1_subset_1(k3_matrix_1(u1_struct_0(X1),X4,X7,esk13_2(X8,k2_finseq_1(k1_matrix_1(X4)))),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|X7=X6|X7=X5|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk13_2(X8,k2_finseq_1(k1_matrix_1(X4))))|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~r2_hidden(X8,k2_finseq_1(k1_matrix_1(X4)))|~r2_hidden(X7,k4_finseq_1(X4)))).#trainneg
cnf(i_0_10804, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk12_1(k2_finseq_1(k1_matrix_1(esk6_0))))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk12_1(k2_finseq_1(k1_matrix_1(esk6_0))))|v1_xboole_0(k2_finseq_1(k1_matrix_1(esk6_0)))|esk3_0=esk8_0|~v7_ordinal1(esk12_1(k2_finseq_1(k1_matrix_1(esk6_0)))))).#trainneg
cnf(i_0_1896, plain, (r1_xxreal_0(k2_zfmisc_1(esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))),X4),X5)|r1_xxreal_0(X5,k2_zfmisc_1(esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))),X4))|~v7_ordinal1(X5)|~v1_xboole_0(X3)|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1897, plain, (r1_xxreal_0(k2_zfmisc_1(esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))),X4),X5)|r1_xxreal_0(X5,k2_zfmisc_1(esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))),X4))|~v7_ordinal1(X5)|~v1_xboole_0(X2)|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_10783, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|esk3_0=esk8_0|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10716, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|esk3_0=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(X3)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10899, plain, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X7|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X8|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=esk8_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=esk3_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=X7|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=X8|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X5,X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~epred1_2(X6,X5)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X5,X6))).#trainneg
cnf(i_0_10703, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|esk3_0=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(X3)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2377, plain, (m2_finseq_1(esk13_2(X1,esk12_1(k1_zfmisc_1(k3_finseq_2(X2)))),X2)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k3_finseq_2(X2)))))).#trainneg
cnf(i_0_1898, plain, (r1_xxreal_0(k2_zfmisc_1(esk12_1(k4_finseq_2(X1,X2)),X3),X4)|r1_xxreal_0(X4,k2_zfmisc_1(esk12_1(k4_finseq_2(X1,X2)),X3))|~v7_ordinal1(X4)|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1899, plain, (r1_xxreal_0(k2_zfmisc_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X3),X4)|r1_xxreal_0(X4,k2_zfmisc_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X3))|~v7_ordinal1(X4)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1900, plain, (r1_xxreal_0(k2_zfmisc_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X3),X4)|r1_xxreal_0(X4,k2_zfmisc_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X3))|~v7_ordinal1(X4)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_10923, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|esk15_7(X4,X5,esk5_0,esk6_0,X1,X3,esk7_0)=esk3_0|esk15_7(X4,X5,esk5_0,esk6_0,X1,X3,esk7_0)=esk8_0|esk15_7(X4,X5,esk5_0,esk6_0,X1,X3,esk7_0)=X3|esk15_7(X4,X5,esk5_0,esk6_0,X1,X3,esk7_0)=X1|esk3_0=esk8_0|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~r2_hidden(X1,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X4,X5)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X4,X5)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))).#trainneg
cnf(i_0_10636, plain, (k3_finseq_1(X1)=k3_finseq_1(esk6_0)|esk15_7(X2,X3,esk5_0,esk6_0,X4,X5,esk7_0)=esk4_0|esk15_7(X2,X3,esk5_0,esk6_0,X4,X5,esk7_0)=esk3_0|esk15_7(X2,X3,esk5_0,esk6_0,X4,X5,esk7_0)=X4|esk15_7(X2,X3,esk5_0,esk6_0,X4,X5,esk7_0)=X5|X1!=esk7_0|~epred1_2(X3,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X2,X3)|~m1_matrix_1(X1,u1_struct_0(esk5_0),X2,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X2,X3))).#trainneg
cnf(i_0_10679, plain, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X7|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X8|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=esk4_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=esk3_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=X7|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=X8|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X5,X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~epred1_2(X6,X5)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X5,X6))).#trainneg
cnf(i_0_10896, plain, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X7|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X8|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=esk8_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=esk3_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=X7|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=X8|esk3_0=esk8_0|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X5,X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~epred1_2(X6,X5)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X5,X6))).#trainneg
cnf(i_0_10509, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X1)|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_matrix_1(X2))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_11035, plain, (v1_funct_1(esk13_2(X1,esk12_1(k1_zfmisc_1(k3_finseq_2(X2)))))|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k3_finseq_2(X2)))))).#trainneg
cnf(i_0_1901, plain, (r1_xxreal_0(k2_zfmisc_1(esk12_1(k3_finseq_2(X1)),X2),X3)|r1_xxreal_0(X3,k2_zfmisc_1(esk12_1(k3_finseq_2(X1)),X2))|~v7_ordinal1(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1902, plain, (r1_xxreal_0(k2_zfmisc_1(k1_xboole_0,X1),X2)|r1_xxreal_0(X2,k2_zfmisc_1(k1_xboole_0,X1))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1903, plain, (r1_xxreal_0(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)|r1_xxreal_0(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))|~v7_ordinal1(X4)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_10748, negated_conjecture, (esk15_7(X1,X2,esk5_0,esk7_0,X3,X4,esk7_0)=X5|esk15_7(X1,X2,esk5_0,esk7_0,X3,X4,esk7_0)=X6|k1_matrix12(X1,X2,esk5_0,esk7_0,X3,X4)=esk7_0|esk15_7(esk2_0,esk1_0,esk5_0,esk7_0,X5,X6,esk7_0)=X5|esk15_7(esk2_0,esk1_0,esk5_0,esk7_0,X5,X6,esk7_0)=X6|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk15_7(X1,X2,esk5_0,esk7_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk7_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X4,esk16_7(X1,X2,esk5_0,esk7_0,X3,X4,esk7_0))|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk15_7(X1,X2,esk5_0,esk7_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk7_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X3,esk16_7(X1,X2,esk5_0,esk7_0,X3,X4,esk7_0))|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_10897, plain, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X7|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X8|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=esk8_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=esk3_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=X7|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=X8|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X5,X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~epred1_2(X6,X5)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X5,X6))).#trainneg
cnf(i_0_10749, negated_conjecture, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X7|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X8|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=esk4_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=esk3_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=X7|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=X8|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X5,X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~epred1_2(X6,X5)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X5,X6))).#trainneg
cnf(i_0_10880, plain, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X3,k1_matrix_1(esk6_0))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,k1_matrix_1(esk6_0))|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk8_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk3_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|esk3_0=esk8_0|~r1_xxreal_0(np__1,k1_matrix_1(esk6_0))|~r2_hidden(X3,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2))).#trainneg
cnf(i_0_10664, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X4,k1_matrix_1(X5))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,k1_matrix_1(X5))|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk4_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk3_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|~m2_subset_1(k1_matrix_1(X5),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(X5))|~r2_hidden(X4,k4_finseq_1(esk6_0))|~m1_matrix_1(X5,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(k1_matrix_1(X5))|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2))).#trainneg
cnf(i_0_11128, plain, (r1_xxreal_0(esk12_1(k5_numbers),k2_zfmisc_1(k1_xboole_0,X1))|r1_xxreal_0(k2_zfmisc_1(k1_xboole_0,X1),esk12_1(k5_numbers)))).#trainneg
cnf(i_0_1904, plain, (r1_xxreal_0(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)|r1_xxreal_0(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))|~v7_ordinal1(X4)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1905, plain, (esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=esk12_1(k4_finseq_2(X4,X5))|~v1_xboole_0(X5)|~v7_ordinal1(X4)|~v1_xboole_0(X3)|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1906, plain, (esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=esk12_1(k4_finseq_2(X4,X5))|~v1_xboole_0(X5)|~v7_ordinal1(X4)|~v1_xboole_0(X2)|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_11100, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|esk3_0=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11099, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|esk3_0=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11102, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|esk3_0=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k1_matrix_1(X3)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11036, plain, (v1_finseq_1(esk13_2(X1,esk12_1(k1_zfmisc_1(k3_finseq_2(X2)))))|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k3_finseq_2(X2)))))).#trainneg
cnf(i_0_1907, plain, (esk12_1(k4_finseq_2(X1,X2))=esk12_1(k4_finseq_2(X3,X4))|~v1_xboole_0(X4)|~v7_ordinal1(X3)|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1908, plain, (esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=esk12_1(k4_finseq_2(X3,X4))|~v1_xboole_0(X4)|~v7_ordinal1(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1909, plain, (esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=esk12_1(k4_finseq_2(X3,X4))|~v1_xboole_0(X4)|~v7_ordinal1(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_11403, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(X1)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k5_numbers))).#trainneg
cnf(i_0_11402, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(X1)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers))).#trainneg
cnf(i_0_11404, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|esk3_0=esk8_0|k1_matrix_1(X3)!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(X1)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers))).#trainneg
cnf(i_0_10654, plain, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X3,esk9_4(X5,esk6_0,X6,X7))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,esk9_4(X5,esk6_0,X6,X7))|X7=k6_matrix_1(X5,esk6_0,X6)|v1_xboole_0(X5)|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk4_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk3_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|k1_matrix_1(esk6_0)!=k3_finseq_1(X7)|~m2_finseq_1(esk6_0,k3_finseq_2(X5))|~m2_finseq_1(X7,X5)|~r2_hidden(X3,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2))).#trainneg
cnf(i_0_10666, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X3,k1_matrix_1(X5))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,k1_matrix_1(X5))|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk4_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk3_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|~m2_subset_1(k1_matrix_1(X5),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(X5))|~r2_hidden(X3,k4_finseq_1(esk6_0))|~m1_matrix_1(X5,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(k1_matrix_1(X5))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2))).#trainneg
cnf(i_0_331, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,esk6_0,X4,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,esk6_0,X4)|~m2_finseq_1(esk6_0,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~v1_matrix_1(esk6_0)|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(X5)))|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~m1_matrix_1(X5,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_11038, plain, (v1_relat_1(esk13_2(X1,esk12_1(k1_zfmisc_1(k3_finseq_2(X2)))))|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k3_finseq_2(X2)))))).#trainneg
cnf(i_0_10541, plain, (k3_matrix_1(u1_struct_0(esk5_0),X1,esk8_0,k1_matrix_1(esk6_0))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,k1_matrix_1(esk6_0))|esk3_0=esk8_0|X1!=esk7_0|~epred1_2(esk1_0,esk2_0)|~r1_xxreal_0(np__1,k1_matrix_1(esk6_0))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_7559, plain, (~r2_hidden(X1,k2_zfmisc_1(X3,X4))|~r2_hidden(X2,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_8870, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X3,X4))))|~v1_xboole_0(X4))).#trainneg
cnf(i_0_3165, negated_conjecture, (k1_funct_1(X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(X4,X5,X6,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|v1_xboole_0(X4)|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|X1!=k6_matrix_1(X4,X5,X6)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~m2_finseq_1(X5,k3_finseq_2(X4))|~m2_finseq_1(X1,X4)|~v1_matrix_1(X5)|~m1_matrix_1(X5,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X6)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v7_ordinal1(X2)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_funct_1(X3)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_relat_1(X3)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))).#trainneg
cnf(i_0_4176, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|X1=esk3_0|X1=esk4_0|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X2,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))))|~r2_hidden(X1,k4_finseq_1(esk6_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_5613, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v7_ordinal1(X4)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4))!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_11779, negated_conjecture, (k1_funct_1(X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X4,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))!=k3_finseq_1(X3)|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X4)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(X4)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_11778, negated_conjecture, (k1_funct_1(k6_matrix_1(X1,X2,X3),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X4),X5))=k3_matrix_1(X1,X2,X3,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X4),X5))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X4)=X5|v1_xboole_0(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X4))!=k3_finseq_1(X5)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X4),X5),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X4),X5))|~v1_finseq_1(X5)|~v1_funct_1(X5)|~v1_relat_1(X5)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m2_finseq_1(k6_matrix_1(X1,X2,X3),X1)|~v1_matrix_1(X2)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X3)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_11795, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6711, plain, (m1_subset_1(k2_zfmisc_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X4)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1914, plain, (m1_subset_1(k3_matrix_1(u1_struct_0(X1),X4,X7,esk9_4(X8,X4,X9,X10)),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|X10=k6_matrix_1(X8,X4,X9)|X7=X6|X7=X5|v1_xboole_0(X8)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk9_4(X8,X4,X9,X10))|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|k1_matrix_1(X4)!=k3_finseq_1(X10)|~m2_finseq_1(X4,k3_finseq_2(X8))|~m2_finseq_1(X10,X8)|~v1_matrix_1(X4)|~r2_hidden(X7,k4_finseq_1(X4))|~v7_ordinal1(X9))).#trainneg
cnf(i_0_11681, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|esk3_0=esk8_0|k1_matrix_1(X3)!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers)|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))).#trainneg
cnf(i_0_11680, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))).#trainneg
cnf(i_0_11679, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k5_numbers)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))).#trainneg
cnf(i_0_11715, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),X1,esk8_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|esk3_0=esk8_0|X1!=esk7_0|~epred1_2(esk1_0,esk2_0)|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2905, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X6,X7)=k3_matrix_1(u1_struct_0(X1),esk6_0,X6,X7)|X6=X5|X6=X4|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k1_matrix_1(X8))|~r1_xxreal_0(np__1,X7)|~r2_hidden(X6,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X8))|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X8,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_11713, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),X1,esk8_0,k1_matrix_1(X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,k1_matrix_1(X2))|esk3_0=esk8_0|X1!=esk7_0|~epred1_2(esk1_0,esk2_0)|~r1_xxreal_0(np__1,k1_matrix_1(X2))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_6840, plain, (m1_subset_1(k2_zfmisc_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X4)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_11911, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)),k5_numbers)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))).#trainneg
cnf(i_0_1928, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,esk12_1(k2_finseq_1(k1_matrix_1(X4))))=k3_matrix_1(u1_struct_0(X1),X4,X7,esk12_1(k2_finseq_1(k1_matrix_1(X4))))|X7=X6|X7=X5|v1_xboole_0(k2_finseq_1(k1_matrix_1(X4)))|v6_struct_0(X1)|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk12_1(k2_finseq_1(k1_matrix_1(X4))))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_11897, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_11898, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_11816, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k6_matrix_1(X2,X3,X4)|v1_xboole_0(X2)|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,X3,X4)))!=k3_matrix_1(X2,X3,X4,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,X3,X4)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k6_matrix_1(X2,X3,X4))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,X3,X4)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,X3,X4)))|~v1_finseq_1(k6_matrix_1(X2,X3,X4))|~v1_funct_1(k6_matrix_1(X2,X3,X4))|~v1_relat_1(k6_matrix_1(X2,X3,X4))|~v7_ordinal1(X1)|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3)|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_11926, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X6,X7)=k3_matrix_1(u1_struct_0(X1),esk6_0,X6,X7)|X6=X4|X6=X5|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k1_matrix_1(X8))|~r1_xxreal_0(np__1,X7)|~m1_matrix_1(X8,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X6,k4_finseq_1(esk6_0)))).#trainneg
cnf(i_0_11815, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k6_matrix_1(X2,X3,X4)|v1_xboole_0(X2)|k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,X3,X4)))!=k3_matrix_1(X2,X3,X4,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,X3,X4)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k6_matrix_1(X2,X3,X4))|~v1_finseq_1(k6_matrix_1(X2,X3,X4))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k6_matrix_1(X2,X3,X4))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k6_matrix_1(X2,X3,X4))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,X3,X4)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,X3,X4)))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3)|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X4)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6969, plain, (m1_subset_1(k2_zfmisc_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_10535, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),X1,esk8_0,esk12_1(k5_numbers))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk12_1(k5_numbers))|esk3_0=esk8_0|r1_xxreal_0(k1_matrix_1(esk6_0),esk12_1(k5_numbers))|X1!=esk7_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk12_1(k5_numbers),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk12_1(k5_numbers))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_11913, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)),k5_numbers)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))).#trainneg
cnf(i_0_2048, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,esk13_2(X7,a_1_0_finseq_1(k1_matrix_1(X4))))=k3_matrix_1(u1_struct_0(X1),X4,X6,esk13_2(X7,a_1_0_finseq_1(k1_matrix_1(X4))))|v6_struct_0(X1)|~m2_subset_1(esk13_2(X7,a_1_0_finseq_1(k1_matrix_1(X4))),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk13_2(X7,a_1_0_finseq_1(k1_matrix_1(X4))))|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X4))|~v7_ordinal1(X6)|~v7_ordinal1(esk13_2(X7,a_1_0_finseq_1(k1_matrix_1(X4))))|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X7,a_1_0_finseq_1(k1_matrix_1(X4))))).#trainneg
cnf(i_0_8969, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X3,X4))))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_2049, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,esk13_2(X7,a_1_0_finseq_1(k1_matrix_1(X4))))=k3_matrix_1(u1_struct_0(X1),X4,X5,esk13_2(X7,a_1_0_finseq_1(k1_matrix_1(X4))))|v6_struct_0(X1)|~m2_subset_1(esk13_2(X7,a_1_0_finseq_1(k1_matrix_1(X4))),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk13_2(X7,a_1_0_finseq_1(k1_matrix_1(X4))))|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X4))|~v7_ordinal1(X5)|~v7_ordinal1(esk13_2(X7,a_1_0_finseq_1(k1_matrix_1(X4))))|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X7,a_1_0_finseq_1(k1_matrix_1(X4))))).#trainneg
cnf(i_0_2090, plain, (v3_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(k1_zfmisc_1(a_1_0_finseq_1(X2))))).#trainneg
cnf(i_0_12082, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k6_matrix_1(X2,X3,X4)|v1_xboole_0(X2)|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,X3,X4)))!=k3_matrix_1(X2,X3,X4,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,X3,X4)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k6_matrix_1(X2,X3,X4))|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,X3,X4)))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3)|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,X3,X4)),k5_numbers))).#trainneg
cnf(i_0_11910, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)),k5_numbers)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k1_matrix_1(esk6_0)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))).#trainneg
cnf(i_0_7093, plain, (m1_subset_1(k2_zfmisc_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_12270, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k1_matrix_1(esk6_0)|~epred1_2(esk1_0,esk2_0)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)),k5_numbers)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0),u1_struct_0(esk5_0))|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))).#trainneg
cnf(i_0_12269, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~epred1_2(esk1_0,esk2_0)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)),k5_numbers)|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0),u1_struct_0(esk5_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_12248, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X7)|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_12239, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X5,X6)=k3_matrix_1(u1_struct_0(X1),esk7_0,X4,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X6)|~r2_hidden(X5,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_12268, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k6_matrix_1(X2,X3,X4)|v1_xboole_0(X2)|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,X3,X4)))!=k3_matrix_1(X2,X3,X4,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,X3,X4)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k6_matrix_1(X2,X3,X4))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(X2,X3,X4)),k5_numbers)|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3)|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v1_finseq_1(k6_matrix_1(X2,X3,X4))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k6_matrix_1(X2,X3,X4))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k6_matrix_1(X2,X3,X4))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))).#trainneg
cnf(i_0_12249, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X8)))|~r1_xxreal_0(np__1,X7)|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X8)))|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_12246, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X8)))|~r1_xxreal_0(np__1,X7)|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X8)))|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_12242, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X5,X6)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7)))|~r1_xxreal_0(np__1,X6)|~r2_hidden(X5,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7)))|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_12241, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X5,X6)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X7)))|~r1_xxreal_0(np__1,X6)|~r2_hidden(X5,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X7)))|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_12247, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,X7)=k3_matrix_1(u1_struct_0(X1),X4,X5,X7)|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k1_matrix_1(X8))|~r1_xxreal_0(np__1,X7)|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X8))|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X8,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_8848, plain, (m1_subset_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3))),k3_finseq_2(X1))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_12324, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k1_matrix_1(esk6_0)|~epred1_2(esk1_0,esk2_0)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))))).#trainneg
cnf(i_0_11912, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)),k5_numbers)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k1_matrix_1(esk6_0)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))).#trainneg
cnf(i_0_2137, plain, (esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=esk12_1(k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~v1_xboole_0(X5)|~v1_xboole_0(X3)|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2138, plain, (esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=esk12_1(k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~v1_xboole_0(X5)|~v1_xboole_0(X2)|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_12485, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k1_matrix_1(esk6_0)|~epred1_2(esk1_0,esk2_0)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0),u1_struct_0(esk5_0))|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))).#trainneg
cnf(i_0_12163, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),X1,esk8_0,esk12_1(k5_numbers))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk12_1(k5_numbers))|esk3_0=esk8_0|r1_xxreal_0(k1_matrix_1(esk6_0),esk12_1(k5_numbers))|X1!=esk7_0|~epred1_2(esk1_0,esk2_0)|~r1_xxreal_0(np__1,esk12_1(k5_numbers))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_12368, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)),k5_numbers)|~v7_ordinal1(X1)|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k1_matrix_1(esk6_0)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))).#trainneg
cnf(i_0_2140, plain, (esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=esk12_1(k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_xboole_0(X4)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_12722, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k1_matrix_1(esk6_0)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))).#trainneg
cnf(i_0_12724, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k1_matrix_1(esk6_0)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)),k5_numbers))).#trainneg
cnf(i_0_2141, plain, (esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=esk12_1(k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_xboole_0(X4)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2200, plain, (r1_xxreal_0(k1_xboole_0,esk13_2(X1,k2_finseq_1(X2)))|r1_xxreal_0(esk13_2(X1,k2_finseq_1(X2)),k1_xboole_0)|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_12865, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))).#trainneg
cnf(i_0_12345, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X5,X6)=k3_matrix_1(u1_struct_0(X1),esk7_0,X4,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7)))|~r1_xxreal_0(np__1,X6)|~r2_hidden(X5,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_12342, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X5,X6)=k3_matrix_1(u1_struct_0(X1),esk7_0,X4,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X7)))|~r1_xxreal_0(np__1,X6)|~r2_hidden(X5,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_12866, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k1_matrix_1(esk6_0)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))))).#trainneg
cnf(i_0_8890, plain, (m1_subset_1(esk12_1(k3_finseq_2(X3)),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_10277, negated_conjecture, (r1_xxreal_0(esk8_0,k2_zfmisc_1(X1,k1_xboole_0))|r1_xxreal_0(k2_zfmisc_1(X1,k1_xboole_0),esk8_0))).#trainneg
cnf(i_0_2202, negated_conjecture, (r1_xxreal_0(esk1_0,esk13_2(X1,k2_finseq_1(X2)))|r1_xxreal_0(esk13_2(X1,k2_finseq_1(X2)),esk1_0)|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2203, negated_conjecture, (r1_xxreal_0(esk2_0,esk13_2(X1,k2_finseq_1(X2)))|r1_xxreal_0(esk13_2(X1,k2_finseq_1(X2)),esk2_0)|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_11907, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)),k5_numbers)|~v7_ordinal1(X1)|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k1_matrix_1(esk6_0)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))).#trainneg
cnf(i_0_11929, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),X1,esk8_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|esk3_0=esk8_0|X1!=esk7_0|~epred1_2(esk1_0,esk2_0)|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_12343, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X5,X6)=k3_matrix_1(u1_struct_0(X1),esk7_0,X4,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k1_matrix_1(X7))|~r1_xxreal_0(np__1,X6)|~r2_hidden(X5,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X7,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_8946, plain, (m1_subset_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3))),k3_finseq_2(X1))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_11908, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)),k5_numbers)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))).#trainneg
cnf(i_0_2204, negated_conjecture, (r1_xxreal_0(esk3_0,esk13_2(X1,k2_finseq_1(X2)))|r1_xxreal_0(esk13_2(X1,k2_finseq_1(X2)),esk3_0)|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13011, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~epred1_2(esk1_0,esk2_0)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))).#trainneg
cnf(i_0_13103, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k1_matrix_1(esk6_0)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)),k5_numbers))).#trainneg
cnf(i_0_2205, negated_conjecture, (r1_xxreal_0(esk4_0,esk13_2(X1,k2_finseq_1(X2)))|r1_xxreal_0(esk13_2(X1,k2_finseq_1(X2)),esk4_0)|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2206, negated_conjecture, (r1_xxreal_0(esk8_0,esk13_2(X1,k2_finseq_1(X2)))|r1_xxreal_0(esk13_2(X1,k2_finseq_1(X2)),esk8_0)|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_12240, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X5,X6)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k1_matrix_1(X7))|~r1_xxreal_0(np__1,X6)|~r2_hidden(X5,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X7))|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X7,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_12256, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X5,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6),X7))=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6),X7))|v6_struct_0(X1)|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6)=X7|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6),X7),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6),X7))|~r2_hidden(X5,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|k1_matrix_1(esk6_0)!=k3_finseq_1(X7)|~v1_finseq_1(X7)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6))|~v1_funct_1(X7)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6))|~v1_relat_1(X7)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_10070, plain, (k3_matrix_1(u1_struct_0(esk5_0),X1,esk3_0,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X2)|esk3_0=esk8_0|X1!=esk7_0|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(esk6_0)))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10898, plain, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X7|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X8|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=esk8_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=esk3_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=X7|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=X8|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X5,X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~epred1_2(X6,X5)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X5,X6))).#trainneg
cnf(i_0_10278, negated_conjecture, (r1_xxreal_0(esk4_0,k2_zfmisc_1(X1,k1_xboole_0))|r1_xxreal_0(k2_zfmisc_1(X1,k1_xboole_0),esk4_0))).#trainneg
cnf(i_0_2207, plain, (r1_xxreal_0(X1,esk13_2(X2,k2_finseq_1(X3)))|r1_xxreal_0(esk13_2(X2,k2_finseq_1(X3)),X1)|~v7_ordinal1(X1)|~r2_hidden(X2,k2_finseq_1(X3))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2217, plain, (r1_xxreal_0(np__1,esk16_7(X1,X2,X3,X4,X5,X6,X7))|X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|~r2_hidden(esk16_7(X1,X2,X3,X4,X5,X6,X7),a_1_0_finseq_1(k1_matrix_1(X4)))|~v7_ordinal1(k1_matrix_1(X4))|k3_finseq_1(X4)!=k3_finseq_1(X7)|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_2218, plain, (r1_xxreal_0(esk16_7(X1,X2,X3,X4,X5,X6,X7),k1_matrix_1(X4))|X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|~r2_hidden(esk16_7(X1,X2,X3,X4,X5,X6,X7),a_1_0_finseq_1(k1_matrix_1(X4)))|~v7_ordinal1(k1_matrix_1(X4))|k3_finseq_1(X4)!=k3_finseq_1(X7)|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_4114, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,esk7_0,X4,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,esk7_0,X4)|~m2_finseq_1(esk7_0,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(X5)))|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~m1_matrix_1(X5,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_1378, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,X1)|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_8989, plain, (m1_subset_1(esk12_1(k3_finseq_2(X3)),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_13157, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k1_matrix_1(esk6_0)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))))).#trainneg
cnf(i_0_2219, plain, (m2_subset_1(esk16_7(X1,X2,X3,X4,X5,X6,X7),k1_numbers,k5_numbers)|X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|~r2_hidden(esk16_7(X1,X2,X3,X4,X5,X6,X7),a_1_0_finseq_1(k1_matrix_1(X4)))|~v7_ordinal1(k1_matrix_1(X4))|k3_finseq_1(X4)!=k3_finseq_1(X7)|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_13155, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~epred1_2(esk1_0,esk2_0)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))).#trainneg
cnf(i_0_13154, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)),k5_numbers))).#trainneg
cnf(i_0_2249, plain, (r1_xxreal_0(np__1,esk9_4(X1,X2,X3,X4))|X4=k6_matrix_1(X1,X2,X3)|v1_xboole_0(X1)|~r2_hidden(esk9_4(X1,X2,X3,X4),a_1_0_finseq_1(k1_matrix_1(X2)))|~v7_ordinal1(k1_matrix_1(X2))|k1_matrix_1(X2)!=k3_finseq_1(X4)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m2_finseq_1(X4,X1)|~v1_matrix_1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2250, plain, (r1_xxreal_0(esk9_4(X1,X2,X3,X4),k1_matrix_1(X2))|X4=k6_matrix_1(X1,X2,X3)|v1_xboole_0(X1)|~r2_hidden(esk9_4(X1,X2,X3,X4),a_1_0_finseq_1(k1_matrix_1(X2)))|~v7_ordinal1(k1_matrix_1(X2))|k1_matrix_1(X2)!=k3_finseq_1(X4)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m2_finseq_1(X4,X1)|~v1_matrix_1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_13190, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13169, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13170, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13164, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X5,X6)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k1_matrix_1(X7))|~r1_xxreal_0(np__1,X6)|~m1_matrix_1(X7,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X5,k4_finseq_1(esk6_0)))).#trainneg
cnf(i_0_13185, plain, (k3_matrix_1(u1_struct_0(esk5_0),X1,esk3_0,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X2)|esk3_0=esk8_0|X1!=esk7_0|~epred1_2(esk1_0,esk2_0)|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,k1_matrix_1(esk6_0)))).#trainneg
cnf(i_0_13171, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X4,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5),X6))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5),X6))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5)=X6|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk4_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk3_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|k1_matrix_1(esk6_0)!=k3_finseq_1(X6)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5),X6),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5),X6))|~v1_finseq_1(X6)|~v1_funct_1(X6)|~v1_relat_1(X6)|~r2_hidden(X4,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X5)|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2))).#trainneg
cnf(i_0_13192, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X4,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5),X6))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5),X6))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5)=X6|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk4_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk3_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|k1_matrix_1(esk6_0)!=k3_finseq_1(X6)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5),X6),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5),X6))|~v1_finseq_1(X6)|~v1_funct_1(X6)|~v1_relat_1(X6)|~r2_hidden(X4,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X5)|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2))).#trainneg
cnf(i_0_13163, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk3_0),esk3_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k1_matrix_1(X6))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(X6,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13162, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk4_0),esk4_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k1_matrix_1(X6))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(X6,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6696, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X4,X5)))))|~v1_xboole_0(X3)|~v1_xboole_0(X5))).#trainneg
cnf(i_0_7153, plain, (m1_subset_1(esk13_2(X2,k1_zfmisc_1(k1_xboole_0)),k3_finseq_2(X1))|~v1_xboole_0(X1)|~r2_hidden(X2,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_13336, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_13319, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))))).#trainneg
cnf(i_0_2251, plain, (m2_subset_1(esk9_4(X1,X2,X3,X4),k1_numbers,k5_numbers)|X4=k6_matrix_1(X1,X2,X3)|v1_xboole_0(X1)|~r2_hidden(esk9_4(X1,X2,X3,X4),a_1_0_finseq_1(k1_matrix_1(X2)))|~v7_ordinal1(k1_matrix_1(X2))|k1_matrix_1(X2)!=k3_finseq_1(X4)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m2_finseq_1(X4,X1)|~v1_matrix_1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2305, plain, (m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|v6_struct_0(X2)|v1_xboole_0(u1_struct_0(X2))|~v1_funct_1(k1_matrix12(X3,X4,X2,X1,X5,X6))|~v1_relat_1(k1_matrix12(X3,X4,X2,X1,X5,X6))|~m1_matrix_1(X1,u1_struct_0(X2),X3,X4)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_13320, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k1_matrix_1(esk6_0)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).#trainneg
cnf(i_0_13365, plain, (X1=k6_matrix_1(X2,X3,X4)|r1_xxreal_0(esk9_4(X2,X3,X4,X1),k1_matrix_1(X3))|v1_xboole_0(X2)|k1_matrix_1(X3)!=k3_finseq_1(X1)|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X1,X2)|~v1_matrix_1(X3)|~r2_hidden(esk9_4(X2,X3,X4,X1),k2_finseq_1(k1_matrix_1(X3)))|~v7_ordinal1(k1_matrix_1(X3))|~v7_ordinal1(X4))).#trainneg
cnf(i_0_13402, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk3_0),esk3_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6)))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_13565, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k1_matrix_1(esk6_0)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).#trainneg
cnf(i_0_13401, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk3_0),esk3_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6)))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_13423, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk4_0),esk4_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6)))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_13368, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k5_numbers))).#trainneg
cnf(i_0_13422, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk4_0),esk4_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6)))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_13161, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk8_0),esk8_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k1_matrix_1(X6))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(X6,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13500, negated_conjecture, (v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_2364, plain, (v3_card_1(esk13_2(X1,esk12_1(k1_zfmisc_1(k4_finseq_2(X2,X3)))),X2)|v1_xboole_0(X3)|~v7_ordinal1(X2)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k4_finseq_2(X2,X3)))))).#trainneg
cnf(i_0_13501, negated_conjecture, (v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_13502, negated_conjecture, (v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_2361, plain, (k1_matrix12(X1,X2,X3,k1_matrix12(X4,X5,X6,X7,X8,X9),X10,X11)=X12|r1_xxreal_0(esk10_2(k1_matrix12(X1,X2,X3,k1_matrix12(X4,X5,X6,X7,X8,X9),X10,X11),X12),k3_finseq_1(X7))|v6_struct_0(X3)|v6_struct_0(X6)|v2_struct_0(X6)|k3_finseq_1(X7)!=k3_finseq_1(X12)|~v1_finseq_1(k1_matrix12(X1,X2,X3,k1_matrix12(X4,X5,X6,X7,X8,X9),X10,X11))|~v1_finseq_1(X12)|~v1_funct_1(k1_matrix12(X1,X2,X3,k1_matrix12(X4,X5,X6,X7,X8,X9),X10,X11))|~v1_funct_1(X12)|~v1_relat_1(k1_matrix12(X1,X2,X3,k1_matrix12(X4,X5,X6,X7,X8,X9),X10,X11))|~v1_relat_1(X12)|~m1_matrix_1(k1_matrix12(X4,X5,X6,X7,X8,X9),u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X11)|~v7_ordinal1(X10)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_matrix_1(X7,u1_struct_0(X6),X4,X5)|~l6_algstr_0(X6)|~v5_vectsp_1(X6)|~v4_vectsp_1(X6)|~v4_rlvect_1(X6)|~v3_rlvect_1(X6)|~v2_rlvect_1(X6)|~v5_group_1(X6)|~v3_group_1(X6)|~v33_algstr_0(X6)|~v13_algstr_0(X6)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X5)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_2365, plain, (m2_finseq_1(esk13_2(X1,esk12_1(k1_zfmisc_1(k4_finseq_2(X2,X3)))),X3)|~v7_ordinal1(X2)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k4_finseq_2(X2,X3)))))).#trainneg
cnf(i_0_13686, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk8_0),esk8_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6)))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_13630, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)),k5_numbers)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k1_matrix_1(esk6_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))).#trainneg
cnf(i_0_13687, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk8_0),esk8_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6)))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_13634, negated_conjecture, (k6_matrix_1(X1,X2,X3)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|v1_xboole_0(X1)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(X1,X2,X3)))!=k3_matrix_1(X1,X2,X3,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(X1,X2,X3)))|k3_finseq_1(k6_matrix_1(X1,X2,X3))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(X1,X2,X3)),k5_numbers)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X3)|k3_finseq_1(k6_matrix_1(X1,X2,X3))!=k1_matrix_1(esk6_0)|~v1_finseq_1(k6_matrix_1(X1,X2,X3))|~v1_funct_1(k6_matrix_1(X1,X2,X3))|~v1_relat_1(k6_matrix_1(X1,X2,X3)))).#trainneg
cnf(i_0_13684, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk8_0),esk8_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13633, negated_conjecture, (k6_matrix_1(X1,esk6_0,X2)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|v1_xboole_0(X1)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(X1,esk6_0,X2)))!=k3_matrix_1(X1,esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(X1,esk6_0,X2)))|k3_finseq_1(k6_matrix_1(X1,esk6_0,X2))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_finseq_1(k6_matrix_1(X1,esk6_0,X2))|~v1_funct_1(k6_matrix_1(X1,esk6_0,X2))|~v1_relat_1(k6_matrix_1(X1,esk6_0,X2))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(X1,esk6_0,X2)),k5_numbers)|~m2_finseq_1(esk6_0,k3_finseq_2(X1))|~v7_ordinal1(X2)|k3_finseq_1(k6_matrix_1(X1,esk6_0,X2))!=k1_matrix_1(esk6_0))).#trainneg
cnf(i_0_13631, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)),k5_numbers)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))).#trainneg
cnf(i_0_13771, negated_conjecture, (k1_matrix12(X1,X2,X3,esk7_0,X4,X5)=X6|r1_xxreal_0(esk10_2(k1_matrix12(X1,X2,X3,esk7_0,X4,X5),X6),k3_finseq_1(esk6_0))|v6_struct_0(X3)|esk3_0=esk8_0|k3_finseq_1(esk6_0)!=k3_finseq_1(X6)|~v1_finseq_1(k1_matrix12(X1,X2,X3,esk7_0,X4,X5))|~v1_finseq_1(X6)|~v1_funct_1(k1_matrix12(X1,X2,X3,esk7_0,X4,X5))|~v1_funct_1(X6)|~v1_relat_1(k1_matrix12(X1,X2,X3,esk7_0,X4,X5))|~v1_relat_1(X6)|~m1_matrix_1(esk7_0,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13643, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k1_matrix_1(esk6_0)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).#trainneg
cnf(i_0_13503, negated_conjecture, (m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0),u1_struct_0(esk5_0))|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_13636, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)),k5_numbers))).#trainneg
cnf(i_0_2366, plain, (esk11_2(esk13_2(X1,esk12_1(k1_zfmisc_1(a_1_0_finseq_1(X2)))),X2)=esk13_2(X1,esk12_1(k1_zfmisc_1(a_1_0_finseq_1(X2))))|v1_xboole_0(a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(a_1_0_finseq_1(X2)))))).#trainneg
cnf(i_0_13706, negated_conjecture, (m1_subset_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)),k5_numbers)|esk3_0=esk8_0|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_2367, plain, (v1_relat_1(esk13_2(X1,esk12_1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3))))))|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3))))))).#trainneg
cnf(i_0_2368, plain, (v1_xboole_0(esk13_2(X1,esk12_1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3))))))|~v1_xboole_0(X3)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3))))))).#trainneg
cnf(i_0_13864, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k1_matrix_1(esk6_0)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).#trainneg
cnf(i_0_13826, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)),k5_numbers)|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_2369, plain, (v1_xboole_0(esk13_2(X1,esk12_1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3))))))|~v1_xboole_0(X2)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3))))))).#trainneg
cnf(i_0_2371, plain, (m1_subset_1(X1,esk13_2(X2,esk12_1(k1_zfmisc_1(k1_zfmisc_1(X3)))))|v1_xboole_0(X3)|v1_xboole_0(esk13_2(X2,esk12_1(k1_zfmisc_1(k1_zfmisc_1(X3)))))|~m2_subset_1(X1,X3,esk13_2(X2,esk12_1(k1_zfmisc_1(k1_zfmisc_1(X3)))))|~r2_hidden(X2,esk12_1(k1_zfmisc_1(k1_zfmisc_1(X3)))))).#trainneg
cnf(i_0_2372, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk13_2(X3,esk12_1(k1_zfmisc_1(k1_zfmisc_1(X2)))))|~r2_hidden(X3,esk12_1(k1_zfmisc_1(k1_zfmisc_1(X2)))))).#trainneg
cnf(i_0_13809, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),u1_struct_0(esk5_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k1_matrix_1(esk6_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_13835, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k1_matrix_1(esk6_0)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_13824, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))!=k1_matrix_1(esk6_0)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)),k5_numbers))).#trainneg
cnf(i_0_13628, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k1_matrix_1(esk6_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_13627, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|~v7_ordinal1(X1)|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k1_matrix_1(esk6_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_7721, plain, (m1_subset_1(esk12_1(k4_finseq_2(X2,X3)),k3_finseq_2(X1))|~v1_xboole_0(X1)|~v1_xboole_0(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13882, negated_conjecture, (v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_2373, plain, (m2_subset_1(X1,X2,esk13_2(X3,esk12_1(k1_zfmisc_1(k1_zfmisc_1(X2)))))|v1_xboole_0(esk13_2(X3,esk12_1(k1_zfmisc_1(k1_zfmisc_1(X2)))))|v1_xboole_0(X2)|~m1_subset_1(X1,esk13_2(X3,esk12_1(k1_zfmisc_1(k1_zfmisc_1(X2)))))|~r2_hidden(X3,esk12_1(k1_zfmisc_1(k1_zfmisc_1(X2)))))).#trainneg
cnf(i_0_2374, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|v1_xboole_0(esk13_2(X3,esk12_1(k1_zfmisc_1(k1_zfmisc_1(X2)))))|~m2_subset_1(X1,X2,esk13_2(X3,esk12_1(k1_zfmisc_1(k1_zfmisc_1(X2)))))|~r2_hidden(X3,esk12_1(k1_zfmisc_1(k1_zfmisc_1(X2)))))).#trainneg
cnf(i_0_2378, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X5,esk13_2(X7,esk12_1(k1_zfmisc_1(k2_finseq_1(k1_matrix_1(X4))))))=k3_matrix_1(u1_struct_0(X1),X4,X6,esk13_2(X7,esk12_1(k1_zfmisc_1(k2_finseq_1(k1_matrix_1(X4))))))|v1_xboole_0(k2_finseq_1(k1_matrix_1(X4)))|v6_struct_0(X1)|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(esk13_2(X7,esk12_1(k1_zfmisc_1(k2_finseq_1(k1_matrix_1(X4))))))|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X7,esk12_1(k1_zfmisc_1(k2_finseq_1(k1_matrix_1(X4))))))).#trainneg
cnf(i_0_13963, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).#trainneg
cnf(i_0_13961, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).#trainneg
cnf(i_0_14023, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),u1_struct_0(esk5_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14026, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14022, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))!=k1_matrix_1(esk6_0)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))))).#trainneg
cnf(i_0_14019, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k1_matrix_1(esk6_0)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))))).#trainneg
cnf(i_0_14020, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k1_matrix_1(esk6_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))))).#trainneg
cnf(i_0_13881, negated_conjecture, (v3_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_13640, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_2379, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,esk13_2(X7,esk12_1(k1_zfmisc_1(k2_finseq_1(k1_matrix_1(X4))))))=k3_matrix_1(u1_struct_0(X1),X4,X5,esk13_2(X7,esk12_1(k1_zfmisc_1(k2_finseq_1(k1_matrix_1(X4))))))|v1_xboole_0(k2_finseq_1(k1_matrix_1(X4)))|v6_struct_0(X1)|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(esk13_2(X7,esk12_1(k1_zfmisc_1(k2_finseq_1(k1_matrix_1(X4))))))|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X7,esk12_1(k1_zfmisc_1(k2_finseq_1(k1_matrix_1(X4))))))).#trainneg
cnf(i_0_2380, plain, (k1_funct_1(X1,esk13_2(X2,esk12_1(k1_zfmisc_1(k2_finseq_1(k1_matrix_1(X3))))))=k3_matrix_1(X4,X3,X5,esk13_2(X2,esk12_1(k1_zfmisc_1(k2_finseq_1(k1_matrix_1(X3))))))|v1_xboole_0(k2_finseq_1(k1_matrix_1(X3)))|v1_xboole_0(X4)|X1!=k6_matrix_1(X4,X3,X5)|~m2_finseq_1(X3,k3_finseq_2(X4))|~m2_finseq_1(X1,X4)|~v1_matrix_1(X3)|~v7_ordinal1(X5)|~v7_ordinal1(esk13_2(X2,esk12_1(k1_zfmisc_1(k2_finseq_1(k1_matrix_1(X3))))))|~r2_hidden(X2,esk12_1(k1_zfmisc_1(k2_finseq_1(k1_matrix_1(X3))))))).#trainneg
cnf(i_0_2381, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,esk13_2(X8,esk12_1(k1_zfmisc_1(k2_finseq_1(k1_matrix_1(X4))))))=k3_matrix_1(u1_struct_0(X1),X4,X7,esk13_2(X8,esk12_1(k1_zfmisc_1(k2_finseq_1(k1_matrix_1(X4))))))|X7=X6|X7=X5|v1_xboole_0(k2_finseq_1(k1_matrix_1(X4)))|v6_struct_0(X1)|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk13_2(X8,esk12_1(k1_zfmisc_1(k2_finseq_1(k1_matrix_1(X4))))))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X8,esk12_1(k1_zfmisc_1(k2_finseq_1(k1_matrix_1(X4))))))).#trainneg
cnf(i_0_13851, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk8_0),esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5),X6))=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5),X6))|v6_struct_0(X1)|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5)=X6|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5),X6),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5),X6))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|k1_matrix_1(esk6_0)!=k3_finseq_1(X6)|~v1_finseq_1(X6)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5))|~v1_funct_1(X6)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5))|~v1_relat_1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_6697, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X4,X5)))))|~v1_xboole_0(X3)|~v1_xboole_0(X5))).#trainneg
cnf(i_0_5756, plain, (v1_xboole_0(esk13_2(X1,k4_finseq_2(X2,X3)))|~v1_xboole_0(X3)|~r2_hidden(X1,k4_finseq_2(X2,X3))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2382, plain, (esk11_2(esk13_2(X1,esk12_1(k1_zfmisc_1(k2_finseq_1(X2)))),X2)=esk13_2(X1,esk12_1(k1_zfmisc_1(k2_finseq_1(X2))))|v1_xboole_0(k2_finseq_1(X2))|~v7_ordinal1(X2)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k2_finseq_1(X2)))))).#trainneg
cnf(i_0_14074, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).#trainneg
cnf(i_0_2383, plain, (m1_subset_1(esk13_2(X1,esk12_1(k1_zfmisc_1(k4_finseq_1(X2)))),k5_numbers)|v1_xboole_0(k4_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k4_finseq_1(X2)))))).#trainneg
cnf(i_0_2402, plain, (r1_xxreal_0(k1_xboole_0,esk13_2(X1,esk12_1(k1_zfmisc_1(k5_numbers))))|r1_xxreal_0(esk13_2(X1,esk12_1(k1_zfmisc_1(k5_numbers))),k1_xboole_0)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k5_numbers))))).#trainneg
cnf(i_0_14079, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))!=k1_matrix_1(esk6_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_14081, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k1_matrix_1(esk6_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_14137, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4),X5))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4),X5))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4)=X5|esk15_7(X1,X2,esk5_0,esk6_0,X3,esk8_0,esk7_0)=esk4_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,esk8_0,esk7_0)=esk3_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,esk8_0,esk7_0)=X3|esk15_7(X1,X2,esk5_0,esk6_0,X3,esk8_0,esk7_0)=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(X5)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4),X5),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4),X5))|~v1_finseq_1(X5)|~v1_funct_1(X5)|~v1_relat_1(X5)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2))).#trainneg
cnf(i_0_14075, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~epred1_2(esk1_0,esk2_0)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).#trainneg
cnf(i_0_14138, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4),X5))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4),X5))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4)=X5|esk15_7(X1,X2,esk5_0,esk6_0,X3,esk8_0,esk7_0)=esk8_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,esk8_0,esk7_0)=esk3_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,esk8_0,esk7_0)=X3|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(X5)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4),X5),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4),X5))|~v1_finseq_1(X5)|~v1_funct_1(X5)|~v1_relat_1(X5)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2))).#trainneg
cnf(i_0_13371, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers))).#trainneg
cnf(i_0_7604, plain, (v3_ordinal1(esk13_2(X1,k3_finseq_2(X2)))|~v3_ordinal1(k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2)|~r2_hidden(X1,k3_finseq_2(X2)))).#trainneg
cnf(i_0_2403, plain, (r1_xxreal_0(np__1,esk13_2(X1,esk12_1(k1_zfmisc_1(k5_numbers))))|r1_xxreal_0(esk13_2(X1,esk12_1(k1_zfmisc_1(k5_numbers))),np__1)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k5_numbers))))).#trainneg
cnf(i_0_13638, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k1_matrix_1(esk6_0)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)),k5_numbers))).#trainneg
cnf(i_0_2404, negated_conjecture, (r1_xxreal_0(esk1_0,esk13_2(X1,esk12_1(k1_zfmisc_1(k5_numbers))))|r1_xxreal_0(esk13_2(X1,esk12_1(k1_zfmisc_1(k5_numbers))),esk1_0)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k5_numbers))))).#trainneg
cnf(i_0_2405, negated_conjecture, (r1_xxreal_0(esk2_0,esk13_2(X1,esk12_1(k1_zfmisc_1(k5_numbers))))|r1_xxreal_0(esk13_2(X1,esk12_1(k1_zfmisc_1(k5_numbers))),esk2_0)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k5_numbers))))).#trainneg
cnf(i_0_10657, plain, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X4,esk9_4(X5,esk6_0,X6,X7))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,esk9_4(X5,esk6_0,X6,X7))|X7=k6_matrix_1(X5,esk6_0,X6)|v1_xboole_0(X5)|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk4_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk3_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|k1_matrix_1(esk6_0)!=k3_finseq_1(X7)|~m2_finseq_1(esk6_0,k3_finseq_2(X5))|~m2_finseq_1(X7,X5)|~r2_hidden(X4,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2))).#trainneg
cnf(i_0_13648, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk3_0),esk3_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_471, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X6,X7)=k3_matrix_1(u1_struct_0(X1),esk6_0,X6,X7)|X6=X4|X6=X5|v6_struct_0(X1)|v2_struct_0(X1)|~r2_hidden(X7,k2_finseq_1(k1_matrix_1(X8)))|~r2_hidden(X6,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X8,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_4437, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|X1=esk3_0|X1=esk4_0|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~r2_hidden(X1,k4_finseq_1(esk6_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v7_ordinal1(X1)|~v7_ordinal1(X4)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X4))!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(X3)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13573, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X5,esk9_4(X6,esk6_0,X7,X8))=k3_matrix_1(u1_struct_0(X1),esk7_0,X4,esk9_4(X6,esk6_0,X7,X8))|v6_struct_0(X1)|X8=k6_matrix_1(X6,esk6_0,X7)|v1_xboole_0(X6)|~m2_subset_1(esk9_4(X6,esk6_0,X7,X8),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk9_4(X6,esk6_0,X7,X8))|~r2_hidden(X5,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|k1_matrix_1(esk6_0)!=k3_finseq_1(X8)|~m2_finseq_1(esk6_0,k3_finseq_2(X6))|~m2_finseq_1(X8,X6)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_13661, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,esk4_0),esk4_0,X5)=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,X5)|v6_struct_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X5,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X5)|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_10342, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X1)|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2)))|~v7_ordinal1(X1)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_3824, plain, (esk13_2(X1,k3_finseq_2(X2))=k1_xboole_0|~v1_xboole_0(X2)|~r2_hidden(X1,k3_finseq_2(X2)))).#trainneg
cnf(i_0_14217, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).#trainneg
cnf(i_0_2406, negated_conjecture, (r1_xxreal_0(esk3_0,esk13_2(X1,esk12_1(k1_zfmisc_1(k5_numbers))))|r1_xxreal_0(esk13_2(X1,esk12_1(k1_zfmisc_1(k5_numbers))),esk3_0)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k5_numbers))))).#trainneg
cnf(i_0_14226, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~epred1_2(esk1_0,esk2_0)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))).#trainneg
cnf(i_0_2407, negated_conjecture, (r1_xxreal_0(esk4_0,esk13_2(X1,esk12_1(k1_zfmisc_1(k5_numbers))))|r1_xxreal_0(esk13_2(X1,esk12_1(k1_zfmisc_1(k5_numbers))),esk4_0)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k5_numbers))))).#trainneg
cnf(i_0_2408, negated_conjecture, (r1_xxreal_0(esk8_0,esk13_2(X1,esk12_1(k1_zfmisc_1(k5_numbers))))|r1_xxreal_0(esk13_2(X1,esk12_1(k1_zfmisc_1(k5_numbers))),esk8_0)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k5_numbers))))).#trainneg
cnf(i_0_14228, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))).#trainneg
cnf(i_0_14223, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))!=k1_matrix_1(esk6_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_14224, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k1_matrix_1(esk6_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_14292, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X1)|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_14198, plain, (v7_ordinal1(esk13_2(X1,k4_finseq_2(X2,X3)))|~v1_xboole_0(X3)|~r2_hidden(X1,k4_finseq_2(X2,X3))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2409, plain, (r1_xxreal_0(X1,esk13_2(X2,esk12_1(k1_zfmisc_1(k5_numbers))))|r1_xxreal_0(esk13_2(X2,esk12_1(k1_zfmisc_1(k5_numbers))),X1)|~v7_ordinal1(X1)|~r2_hidden(X2,esk12_1(k1_zfmisc_1(k5_numbers))))).#trainneg
cnf(i_0_14027, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k1_matrix_1(esk6_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_2412, plain, (m1_subset_1(esk13_2(X1,esk12_1(k1_zfmisc_1(esk12_1(k1_zfmisc_1(X2))))),X2)|v1_xboole_0(esk12_1(k1_zfmisc_1(X2)))|~r2_hidden(X1,esk12_1(k1_zfmisc_1(esk12_1(k1_zfmisc_1(X2))))))).#trainneg
cnf(i_0_14291, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X1)|esk3_0=esk8_0|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_14225, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|esk15_7(X4,X5,esk5_0,esk6_0,X3,esk8_0,esk7_0)=esk8_0|esk15_7(X4,X5,esk5_0,esk6_0,X3,esk8_0,esk7_0)=esk3_0|esk15_7(X4,X5,esk5_0,esk6_0,X3,esk8_0,esk7_0)=esk4_0|esk15_7(X4,X5,esk5_0,esk6_0,X3,esk8_0,esk7_0)=X3|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X4,X5)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X4,X5)|~v7_ordinal1(X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X1)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers))).#trainneg
cnf(i_0_14231, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk9_4(X2,esk6_0,X3,X4))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X5,esk9_4(X2,esk6_0,X3,X4))|esk15_7(esk2_0,esk1_0,esk5_0,esk6_0,X5,X1,esk7_0)=esk3_0|esk15_7(esk2_0,esk1_0,esk5_0,esk6_0,X5,X1,esk7_0)=esk4_0|esk15_7(esk2_0,esk1_0,esk5_0,esk6_0,X5,X1,esk7_0)=X1|esk15_7(esk2_0,esk1_0,esk5_0,esk6_0,X5,X1,esk7_0)=X5|X4=k6_matrix_1(X2,esk6_0,X3)|v1_xboole_0(X2)|k1_matrix_1(esk6_0)!=k3_finseq_1(X4)|~m2_finseq_1(esk6_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~r2_hidden(X1,k4_finseq_1(esk6_0))|~v7_ordinal1(X5)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1183, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,X1)|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_14306, plain, (r2_hidden(k1_xboole_0,k3_finseq_2(X2))|~r2_hidden(X1,k3_finseq_2(X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_4085, plain, (esk13_2(X1,k1_zfmisc_1(k1_xboole_0))=esk12_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xboole_0(X3)|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_14327, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_14328, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_2414, plain, (m1_subset_1(esk12_1(esk12_1(k1_zfmisc_1(X1))),X1)|v1_xboole_0(esk12_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_14335, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))!=k1_matrix_1(esk6_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_2415, plain, (v3_ordinal1(esk16_7(X1,X2,X3,X4,X5,X6,X7))|X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(X7)|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(k1_matrix_1(X4))|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_2416, plain, (v7_ordinal1(esk16_7(X1,X2,X3,X4,X5,X6,X7))|X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(X7)|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(k1_matrix_1(X4))|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_14654, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_14336, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k1_matrix_1(esk6_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_14655, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))!=k1_matrix_1(esk6_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_14332, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_14580, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_14418, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|esk15_7(X4,X5,esk5_0,esk6_0,X3,esk8_0,esk7_0)=esk4_0|esk15_7(X4,X5,esk5_0,esk6_0,X3,esk8_0,esk7_0)=esk3_0|esk15_7(X4,X5,esk5_0,esk6_0,X3,esk8_0,esk7_0)=esk8_0|esk15_7(X4,X5,esk5_0,esk6_0,X3,esk8_0,esk7_0)=X3|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X4,X5)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X4,X5)|~v7_ordinal1(X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))).#trainneg
cnf(i_0_14551, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X2)|esk3_0=esk8_0|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(X3))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X2)|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_14331, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers)|~v7_ordinal1(X1)|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k1_matrix_1(esk6_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_4089, plain, (esk13_2(X1,k1_zfmisc_1(k1_xboole_0))=esk12_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xboole_0(X2)|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_2418, plain, (v3_ordinal1(esk9_4(X1,X2,X3,X4))|X4=k6_matrix_1(X1,X2,X3)|v1_xboole_0(X1)|k1_matrix_1(X2)!=k3_finseq_1(X4)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m2_finseq_1(X4,X1)|~v1_matrix_1(X2)|~v7_ordinal1(k1_matrix_1(X2))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_14620, plain, (v7_ordinal1(esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|v1_xboole_0(esk12_1(k1_zfmisc_1(k5_numbers))))).#trainneg
cnf(i_0_14539, negated_conjecture, (v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_14676, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers))).#trainneg
cnf(i_0_14674, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_14708, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k1_matrix_1(esk6_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k5_numbers))).#trainneg
cnf(i_0_14675, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_14677, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14680, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14673, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k1_matrix_1(esk6_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_14688, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X2)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_14686, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X2)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_14540, negated_conjecture, (v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_2442, plain, (m1_subset_1(esk13_2(X1,esk12_1(k1_zfmisc_1(k2_finseq_1(X2)))),k5_numbers)|v1_xboole_0(k2_finseq_1(X2))|~v7_ordinal1(X2)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k2_finseq_1(X2)))))).#trainneg
cnf(i_0_14541, negated_conjecture, (v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_14542, negated_conjecture, (m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0),u1_struct_0(esk5_0))|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_2444, plain, (m1_subset_1(esk12_1(k2_finseq_1(X1)),k5_numbers)|v1_xboole_0(k2_finseq_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14689, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X2)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_14685, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|esk15_7(X4,X5,esk5_0,esk6_0,X3,esk8_0,esk7_0)=esk8_0|esk15_7(X4,X5,esk5_0,esk6_0,X3,esk8_0,esk7_0)=esk3_0|esk15_7(X4,X5,esk5_0,esk6_0,X3,esk8_0,esk7_0)=esk4_0|esk15_7(X4,X5,esk5_0,esk6_0,X3,esk8_0,esk7_0)=X3|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X4,X5)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X4,X5)|~v7_ordinal1(X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X1)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2)))).#trainneg
cnf(i_0_14703, negated_conjecture, (k1_funct_1(X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|esk3_0=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|k1_matrix_1(esk6_0)!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(X3)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_14678, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|k1_matrix_1(X1)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0)|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_14566, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0),X1)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X1)|esk3_0=esk8_0|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X1)|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_14826, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_11479, plain, (m1_subset_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X3,X4))),k4_finseq_2(X1,X2))|~v1_xboole_0(X2)|~v1_xboole_0(X4)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2457, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|~m2_finseq_1(k1_matrix12(X3,X4,X2,X5,X6,X7),k3_finseq_2(u1_struct_0(X2)))|~v1_matrix_1(k1_matrix12(X3,X4,X2,X5,X6,X7))|~r2_hidden(X1,k8_matrix_1(u1_struct_0(X2),k1_matrix12(X3,X4,X2,X5,X6,X7),X8))|~v7_ordinal1(X8)|~m1_matrix_1(X5,u1_struct_0(X2),X3,X4)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X7)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_2458, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(X2)|v1_xboole_0(k3_finseq_2(X2))|~m2_finseq_1(k6_matrix_1(k3_finseq_2(X2),X3,X4),k3_finseq_2(X2))|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X2),X3,X4))|~r2_hidden(X1,k8_matrix_1(X2,k6_matrix_1(k3_finseq_2(X2),X3,X4),X5))|~v7_ordinal1(X5)|~m2_finseq_1(X3,k3_finseq_2(k3_finseq_2(X2)))|~v1_matrix_1(X3)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_14828, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k1_matrix_1(esk6_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))))).#trainneg
cnf(i_0_14825, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))))).#trainneg
cnf(i_0_14848, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X2)|esk3_0=esk8_0|k1_matrix_1(X3)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(X3))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_14850, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X2)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_14847, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,X2)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_14829, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_14832, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_6824, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X4,X5)))))|~v1_xboole_0(X3)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_11578, plain, (m1_subset_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X3,X4))),k4_finseq_2(X1,X2))|~v1_xboole_0(X2)|~v1_xboole_0(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14813, plain, (r1_xxreal_0(k1_xboole_0,esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|r1_xxreal_0(esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))),k1_xboole_0)|v1_xboole_0(esk12_1(k1_zfmisc_1(k5_numbers))))).#trainneg
cnf(i_0_2459, plain, (m1_subset_1(esk13_2(X1,k6_matrix_1(X2,X3,X4)),k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(X2)|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3)|~v7_ordinal1(X4)|~r2_hidden(X1,k6_matrix_1(X2,X3,X4)))).#trainneg
cnf(i_0_2460, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(X2)|v1_xboole_0(k6_matrix_1(X2,X3,X4))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3)|~v7_ordinal1(X4)|~m1_subset_1(X1,k6_matrix_1(X2,X3,X4)))).#trainneg
cnf(i_0_2463, plain, (m1_subset_1(esk13_2(X1,k1_matrix12(X2,X3,X4,X5,X6,X7)),k2_zfmisc_1(k5_numbers,k3_finseq_2(u1_struct_0(X4))))|v1_xboole_0(u1_struct_0(X4))|v6_struct_0(X4)|~m1_matrix_1(X5,u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~r2_hidden(X1,k1_matrix12(X2,X3,X4,X5,X6,X7)))).#trainneg
cnf(i_0_15009, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k1_matrix_1(esk6_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)))).#trainneg
cnf(i_0_15061, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_15062, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_15059, negated_conjecture, (k1_funct_1(X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|esk3_0=esk8_0|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v7_ordinal1(X2)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(X3)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(X3)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))).#trainneg
cnf(i_0_15230, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4083, plain, (esk13_2(X1,k1_zfmisc_1(k1_xboole_0))=esk12_1(k4_finseq_2(X2,X3))|~v1_xboole_0(X3)|~v7_ordinal1(X2)|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_15235, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_14545, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)),k5_numbers)|~epred1_2(esk1_0,esk2_0))).#trainneg
cnf(i_0_2464, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k3_finseq_2(u1_struct_0(X2))))|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v1_xboole_0(k1_matrix12(X3,X4,X2,X5,X6,X7))|~m1_matrix_1(X5,u1_struct_0(X2),X3,X4)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~m1_subset_1(X1,k1_matrix12(X3,X4,X2,X5,X6,X7)))).#trainneg
cnf(i_0_2496, plain, (v1_relat_1(esk13_2(X1,esk12_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))|~v1_xboole_0(X2)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_2502, plain, (m1_subset_1(k3_matrix_1(u1_struct_0(X1),X4,X6,esk13_2(X7,k2_finseq_1(k1_matrix_1(X4)))),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk13_2(X7,k2_finseq_1(k1_matrix_1(X4))))|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~r2_hidden(X7,k2_finseq_1(k1_matrix_1(X4)))|~r2_hidden(X5,k4_finseq_1(X4)))).#trainneg
cnf(i_0_15368, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_15260, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|X2!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|~epred1_2(esk1_0,esk2_0)|~m2_finseq_1(X2,u1_struct_0(esk5_0)))).#trainneg
cnf(i_0_15259, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|esk3_0=esk8_0|k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(X2)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|X2!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~m2_finseq_1(X2,u1_struct_0(esk5_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15255, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),u1_struct_0(esk5_0)))).#trainneg
cnf(i_0_15234, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))!=k1_matrix_1(esk6_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)))).#trainneg
cnf(i_0_15256, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),u1_struct_0(esk5_0)))).#trainneg
cnf(i_0_15253, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),u1_struct_0(esk5_0)))).#trainneg
cnf(i_0_15257, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))!=k3_finseq_1(X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|~epred1_2(esk1_0,esk2_0)|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),u1_struct_0(esk5_0)))).#trainneg
cnf(i_0_15356, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v7_ordinal1(X1)|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_15392, plain, (v1_relat_1(esk13_2(X1,esk12_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_15342, negated_conjecture, (v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_15265, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2536, plain, (m1_subset_1(k3_matrix_1(u1_struct_0(X1),X4,X6,esk9_4(X7,X4,X8,X9)),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|X9=k6_matrix_1(X7,X4,X8)|v1_xboole_0(X7)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk9_4(X7,X4,X8,X9))|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|k1_matrix_1(X4)!=k3_finseq_1(X9)|~m2_finseq_1(X4,k3_finseq_2(X7))|~m2_finseq_1(X9,X7)|~v1_matrix_1(X4)|~r2_hidden(X5,k4_finseq_1(X4))|~v7_ordinal1(X8))).#trainneg
cnf(i_0_15343, negated_conjecture, (v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_15344, negated_conjecture, (v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_2533, plain, (m1_subset_1(k3_matrix_1(u1_struct_0(X1),X4,X6,esk16_7(X7,X8,X9,X4,X10,X11,X12)),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|X12=k1_matrix12(X7,X8,X9,X4,X10,X11)|v6_struct_0(X9)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk16_7(X7,X8,X9,X4,X10,X11,X12))|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|k3_finseq_1(X4)!=k3_finseq_1(X12)|~epred1_2(X8,X7)|~r2_hidden(X5,k4_finseq_1(X4))|~m1_matrix_1(X12,u1_struct_0(X9),X7,X8)|~m1_matrix_1(X4,u1_struct_0(X9),X7,X8)|~l6_algstr_0(X9)|~v5_vectsp_1(X9)|~v4_vectsp_1(X9)|~v4_rlvect_1(X9)|~v3_rlvect_1(X9)|~v2_rlvect_1(X9)|~v5_group_1(X9)|~v3_group_1(X9)|~v33_algstr_0(X9)|~v13_algstr_0(X9)|~v7_ordinal1(X11)|~v7_ordinal1(X10))).#trainneg
cnf(i_0_2571, plain, (v1_xboole_0(esk13_2(X1,esk12_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))|~v1_xboole_0(X2)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_15341, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(X2)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_15250, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))!=k3_finseq_1(X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0)|~epred1_2(esk1_0,esk2_0)|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0)))).#trainneg
cnf(i_0_15340, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|esk3_0=esk8_0|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_15354, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk3_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v7_ordinal1(X1)|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_15236, negated_conjecture, (k6_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)|esk3_0=esk8_0|k1_matrix_1(esk6_0)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)))).#trainneg
cnf(i_0_11801, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X4))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_15457, plain, (v1_xboole_0(esk13_2(X1,esk12_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_13, negated_conjecture, (k1_matrix12(esk2_0,esk1_0,esk5_0,esk6_0,esk8_0,esk4_0)=esk7_0)).#trainneg
cnf(i_0_13169, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13190, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3871, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,X1)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6825, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X4,X5)))))|~v1_xboole_0(X3)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_15810, plain, (v7_ordinal1(esk13_2(X1,esk12_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_467, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X1)|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(esk6_0)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2629, plain, (v3_ordinal1(k1_matrix12(X1,X2,X3,X4,X5,X6))|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|v2_struct_0(X3)|~v1_xboole_0(k3_finseq_2(u1_struct_0(X3)))|~v7_ordinal1(k1_zfmisc_1(k1_xboole_0))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_2639, plain, (r1_xxreal_0(esk13_2(esk13_2(X1,k5_numbers),k5_numbers),X2)|r1_xxreal_0(X2,esk13_2(esk13_2(X1,k5_numbers),k5_numbers))|~v7_ordinal1(X2)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_2641, plain, (r1_xxreal_0(esk13_2(X1,k5_numbers),X2)|r1_xxreal_0(X2,esk13_2(X1,k5_numbers))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1269, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,X1)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15904, plain, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk8_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk4_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1270, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X1)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3664, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,X1)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4126, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X1)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_15266, plain, (m1_subset_1(esk13_2(X3,k1_zfmisc_1(k1_xboole_0)),k4_finseq_2(X1,X2))|~v1_xboole_0(X2)|~r2_hidden(X3,k1_zfmisc_1(k1_xboole_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_465, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,X1)|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(esk6_0)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2642, plain, (r1_xxreal_0(esk12_1(k5_numbers),esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)))|r1_xxreal_0(esk16_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)),esk12_1(k5_numbers))|k1_matrix12(X1,X2,X3,X7,X8,X9)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(k1_matrix12(X1,X2,X3,X7,X8,X9))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2643, plain, (r1_xxreal_0(esk12_1(k5_numbers),esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)))|r1_xxreal_0(esk15_7(X1,X2,X3,X4,X5,X6,k1_matrix12(X1,X2,X3,X7,X8,X9)),esk12_1(k5_numbers))|k1_matrix12(X1,X2,X3,X7,X8,X9)=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(k1_matrix12(X1,X2,X3,X7,X8,X9))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2644, plain, (r1_xxreal_0(esk12_1(k5_numbers),esk13_2(X1,esk12_1(k1_zfmisc_1(k5_numbers))))|r1_xxreal_0(esk13_2(X1,esk12_1(k1_zfmisc_1(k5_numbers))),esk12_1(k5_numbers))|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k5_numbers))))).#trainneg
cnf(i_0_16189, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X1)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_16187, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X1)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_matrix_1(X2))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_4340, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4185, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|X1=esk4_0|X1=esk8_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X2)|~r2_hidden(X1,k4_finseq_1(esk6_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15910, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4327, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11407, plain, (m1_subset_1(esk12_1(k4_finseq_2(X3,X4)),k4_finseq_2(X1,X2))|~v1_xboole_0(X2)|~v1_xboole_0(X4)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_709, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X1)|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2)))|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2645, plain, (r1_xxreal_0(esk12_1(k5_numbers),esk13_2(X1,k5_numbers))|r1_xxreal_0(esk13_2(X1,k5_numbers),esk12_1(k5_numbers))|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_4149, plain, (v3_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~v3_ordinal1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_14814, plain, (r1_xxreal_0(np__1,esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|r1_xxreal_0(esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))),np__1)|v1_xboole_0(esk12_1(k1_zfmisc_1(k5_numbers))))).#trainneg
cnf(i_0_2646, plain, (r1_xxreal_0(esk12_1(k5_numbers),esk13_2(X1,k2_finseq_1(X2)))|r1_xxreal_0(esk13_2(X1,k2_finseq_1(X2)),esk12_1(k5_numbers))|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2647, plain, (r1_xxreal_0(esk12_1(k5_numbers),esk12_1(k4_finseq_2(X1,X2)))|r1_xxreal_0(esk12_1(k4_finseq_2(X1,X2)),esk12_1(k5_numbers))|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16298, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers))).#trainneg
cnf(i_0_16322, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k5_numbers))).#trainneg
cnf(i_0_15928, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers))).#trainneg
cnf(i_0_4299, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|X1=esk8_0|X1=esk4_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~r1_xxreal_0(np__1,X2)|~r2_hidden(X1,k4_finseq_1(esk6_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_4302, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|X1=esk8_0|X1=esk4_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~r1_xxreal_0(np__1,X2)|~r2_hidden(X1,k4_finseq_1(esk6_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_3035, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,k1_matrix_1(esk6_0))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,k1_matrix_1(esk6_0)))).#trainneg
cnf(i_0_682, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,X1)|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2)))|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2648, plain, (r1_xxreal_0(esk12_1(k5_numbers),esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|r1_xxreal_0(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),esk12_1(k5_numbers))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_16360, plain, (v3_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_16581, negated_conjecture, (v3_ordinal1(esk2_0)|~m2_subset_1(esk2_0,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk2_0)|~v7_ordinal1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_16586, negated_conjecture, (v3_ordinal1(esk1_0)|~m2_subset_1(esk1_0,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk1_0)|~v7_ordinal1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2649, plain, (r1_xxreal_0(esk12_1(k5_numbers),esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|r1_xxreal_0(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),esk12_1(k5_numbers))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2651, plain, (r1_xxreal_0(esk12_1(k5_numbers),esk12_1(k3_finseq_2(X1)))|r1_xxreal_0(esk12_1(k3_finseq_2(X1)),esk12_1(k5_numbers))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_15905, plain, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk8_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk4_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_16152, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X3,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5)))|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk4_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk8_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|~m2_subset_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5)))|~r2_hidden(X3,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X5)|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2))).#trainneg
cnf(i_0_15929, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k5_numbers))).#trainneg
cnf(i_0_16169, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X3,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5)))|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk4_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk8_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5)))|~r2_hidden(X3,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X5)|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2))).#trainneg
cnf(i_0_15894, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))|~m2_subset_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16406, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11039, plain, (v1_xboole_0(esk13_2(X1,esk12_1(k1_zfmisc_1(k3_finseq_2(X2)))))|~v1_xboole_0(X2)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k3_finseq_2(X2)))))).#trainneg
cnf(i_0_1378, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,X1)|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2652, plain, (r1_xxreal_0(esk12_1(k5_numbers),esk10_2(X1,X2))|r1_xxreal_0(esk10_2(X1,X2),esk12_1(k5_numbers))|X1=X2|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16508, plain, (v3_ordinal1(esk12_1(k5_numbers))|~m2_subset_1(esk12_1(k5_numbers),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk12_1(k5_numbers))|~v7_ordinal1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2653, plain, (r1_xxreal_0(esk12_1(k5_numbers),esk9_4(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8))|r1_xxreal_0(esk9_4(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8),esk12_1(k5_numbers))|X8=k6_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6))!=k3_finseq_1(X8)|~m2_finseq_1(X8,u1_struct_0(X1))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_2654, plain, (r1_xxreal_0(esk12_1(k5_numbers),esk9_4(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4,X5))|r1_xxreal_0(esk9_4(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4,X5),esk12_1(k5_numbers))|X5=k6_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4)|v1_xboole_0(k3_finseq_2(X1))|v1_xboole_0(X1)|k1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))!=k3_finseq_1(X5)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(X1)))|~m2_finseq_1(X5,X1)|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))|~v1_matrix_1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_15906, plain, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk8_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk4_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_3872, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,X1)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_matrix_1(X2))|~r1_xxreal_0(np__1,X1)|~m1_matrix_1(X2,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15877, plain, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk9_4(X1,esk6_0,X2,X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk9_4(X1,esk6_0,X2,X3))|X3=k6_matrix_1(X1,esk6_0,X2)|v1_xboole_0(X1)|k1_matrix_1(esk6_0)!=k3_finseq_1(X3)|~m2_finseq_1(esk6_0,k3_finseq_2(X1))|~m2_finseq_1(X3,X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_15874, plain, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk9_4(X1,esk6_0,X2,X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk9_4(X1,esk6_0,X2,X3))|X3=k6_matrix_1(X1,esk6_0,X2)|v1_xboole_0(X1)|k1_matrix_1(esk6_0)!=k3_finseq_1(X3)|~m2_finseq_1(esk6_0,k3_finseq_2(X1))|~m2_finseq_1(X3,X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_16163, plain, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X7|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X8|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=esk4_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=esk8_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=X7|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=X8|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X5,X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~epred1_2(X6,X5)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X5,X6))).#trainneg
cnf(i_0_16162, plain, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X7|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X8|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=esk4_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=esk8_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=X7|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=X8|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X5,X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~epred1_2(X6,X5)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X5,X6))).#trainneg
cnf(i_0_16730, plain, (v7_ordinal1(esk13_2(X1,esk12_1(k1_zfmisc_1(k3_finseq_2(X2)))))|~v1_xboole_0(X2)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k3_finseq_2(X2)))))).#trainneg
cnf(i_0_1183, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X1)|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2661, plain, (r1_xxreal_0(esk12_1(k5_numbers),k2_zfmisc_1(X1,X2))|r1_xxreal_0(k2_zfmisc_1(X1,X2),esk12_1(k5_numbers))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2662, plain, (r1_xxreal_0(esk12_1(k5_numbers),k2_zfmisc_1(X1,X2))|r1_xxreal_0(k2_zfmisc_1(X1,X2),esk12_1(k5_numbers))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2665, plain, (X1=esk13_2(X2,k3_finseq_2(X3))|r1_xxreal_0(esk10_2(X1,esk13_2(X2,k3_finseq_2(X3))),X4)|r1_xxreal_0(X4,esk10_2(X1,esk13_2(X2,k3_finseq_2(X3))))|k3_finseq_1(X1)!=k3_finseq_1(esk13_2(X2,k3_finseq_2(X3)))|~v1_finseq_1(X1)|~v1_funct_1(esk13_2(X2,k3_finseq_2(X3)))|~v1_funct_1(X1)|~v1_relat_1(esk13_2(X2,k3_finseq_2(X3)))|~v1_relat_1(X1)|~v7_ordinal1(X4)|~r2_hidden(X2,k3_finseq_2(X3)))).#trainneg
cnf(i_0_16165, plain, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X7|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X8|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=esk4_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=esk8_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=X7|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=X8|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X5,X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~epred1_2(X6,X5)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X5,X6))).#trainneg
cnf(i_0_16376, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k5_numbers)|~v7_ordinal1(X1)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))).#trainneg
cnf(i_0_5355, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,k1_matrix_1(X1))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,k1_matrix_1(X1))|~r1_xxreal_0(np__1,k1_matrix_1(X1))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_15907, plain, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk8_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk4_0|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_16164, plain, (esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X7|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X8|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|k1_matrix12(X1,X2,esk5_0,esk6_0,X3,X4)=esk7_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=esk4_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=esk8_0|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=X7|esk15_7(X5,X6,esk5_0,esk6_0,X7,X8,esk7_0)=X8|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0),esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,esk16_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0))|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X5,X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~epred1_2(X6,X5)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X5,X6))).#trainneg
cnf(i_0_3822, plain, (esk13_2(X1,k3_finseq_2(X2))=esk12_1(k3_finseq_2(X3))|~v1_xboole_0(X3)|~v1_xboole_0(X2)|~r2_hidden(X1,k3_finseq_2(X2)))).#trainneg
cnf(i_0_2975, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,k1_matrix_1(esk6_0))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,k1_matrix_1(esk6_0)))).#trainneg
cnf(i_0_1137, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,X1)|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2666, plain, (X1=esk12_1(k4_finseq_2(X2,X3))|r1_xxreal_0(esk10_2(X1,esk12_1(k4_finseq_2(X2,X3))),X4)|r1_xxreal_0(X4,esk10_2(X1,esk12_1(k4_finseq_2(X2,X3))))|k3_finseq_1(X1)!=k3_finseq_1(esk12_1(k4_finseq_2(X2,X3)))|~v1_finseq_1(X1)|~v1_funct_1(esk12_1(k4_finseq_2(X2,X3)))|~v1_funct_1(X1)|~v1_relat_1(esk12_1(k4_finseq_2(X2,X3)))|~v1_relat_1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2667, plain, (X1=esk12_1(k3_finseq_2(X2))|r1_xxreal_0(esk10_2(X1,esk12_1(k3_finseq_2(X2))),X3)|r1_xxreal_0(X3,esk10_2(X1,esk12_1(k3_finseq_2(X2))))|k3_finseq_1(X1)!=k3_finseq_1(esk12_1(k3_finseq_2(X2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2670, plain, (X1=k6_matrix_1(X2,X3,X4)|r1_xxreal_0(esk10_2(X1,k6_matrix_1(X2,X3,X4)),X5)|r1_xxreal_0(X5,esk10_2(X1,k6_matrix_1(X2,X3,X4)))|v1_xboole_0(X2)|k3_finseq_1(X1)!=k3_finseq_1(k6_matrix_1(X2,X3,X4))|~v1_finseq_1(X1)|~v1_funct_1(k6_matrix_1(X2,X3,X4))|~v1_funct_1(X1)|~v1_relat_1(k6_matrix_1(X2,X3,X4))|~v1_relat_1(X1)|~v7_ordinal1(X5)|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_16407, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,k1_matrix_1(X1))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,k1_matrix_1(X1))|~r1_xxreal_0(np__1,k1_matrix_1(X1))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_16375, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k5_numbers)|~v7_ordinal1(X1)|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)))).#trainneg
cnf(i_0_15854, negated_conjecture, (k1_funct_1(X1,esk9_4(X2,esk7_0,X3,X4))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X5,esk9_4(X2,esk7_0,X3,X4))|X4=k6_matrix_1(X2,esk7_0,X3)|v1_xboole_0(X2)|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X5)|k1_matrix_1(esk6_0)!=k3_finseq_1(X4)|~m2_finseq_1(esk7_0,k3_finseq_2(X2))|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~m2_finseq_1(X4,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1136, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,X1)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,X1)|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X1,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2))))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2674, plain, (X1=k1_matrix12(X2,X3,X4,X5,X6,X7)|r1_xxreal_0(esk10_2(X1,k1_matrix12(X2,X3,X4,X5,X6,X7)),X8)|r1_xxreal_0(X8,esk10_2(X1,k1_matrix12(X2,X3,X4,X5,X6,X7)))|v1_xboole_0(u1_struct_0(X4))|v6_struct_0(X4)|k3_finseq_1(X1)!=k3_finseq_1(k1_matrix12(X2,X3,X4,X5,X6,X7))|~v1_finseq_1(X1)|~v1_funct_1(k1_matrix12(X2,X3,X4,X5,X6,X7))|~v1_funct_1(X1)|~v1_relat_1(k1_matrix12(X2,X3,X4,X5,X6,X7))|~v1_relat_1(X1)|~v7_ordinal1(X8)|~m1_matrix_1(X5,u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_2711, plain, (esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=esk12_1(k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~v1_xboole_0(X4)|~v1_xboole_0(X3)|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2712, plain, (esk13_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=esk12_1(k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~v1_xboole_0(X4)|~v1_xboole_0(X2)|~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_4300, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|X1=esk8_0|X1=esk4_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(X3))|~r1_xxreal_0(np__1,X2)|~r2_hidden(X1,k4_finseq_1(esk6_0))|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2715, plain, (esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=esk12_1(k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2738, plain, (m1_subset_1(k1_matrix12(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k1_xboole_0))|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|v2_struct_0(X3)|~v1_xboole_0(k3_finseq_2(u1_struct_0(X3)))|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_17372, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|X1=esk4_0|X1=esk8_0|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(X3))|~r1_xxreal_0(np__1,X2)|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k4_finseq_1(esk6_0)))).#trainneg
cnf(i_0_16114, plain, (k3_finseq_1(X1)=k3_finseq_1(esk6_0)|esk15_7(X2,X3,esk5_0,esk6_0,X4,X5,esk7_0)=esk4_0|esk15_7(X2,X3,esk5_0,esk6_0,X4,X5,esk7_0)=esk8_0|esk15_7(X2,X3,esk5_0,esk6_0,X4,X5,esk7_0)=X4|esk15_7(X2,X3,esk5_0,esk6_0,X4,X5,esk7_0)=X5|X1!=esk7_0|~epred1_2(X3,X2)|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X2,X3)|~m1_matrix_1(X1,u1_struct_0(esk5_0),X2,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X2,X3))).#trainneg
cnf(i_0_16878, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16172, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X4,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5),X6))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5),X6))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5)=X6|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk4_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk8_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|k1_matrix_1(esk6_0)!=k3_finseq_1(X6)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5),X6),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X5),X6))|~v1_finseq_1(X6)|~v1_funct_1(X6)|~v1_relat_1(X6)|~r2_hidden(X4,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X5)|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2))).#trainneg
cnf(i_0_3826, plain, (esk13_2(X1,k3_finseq_2(X2))=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X4)|~v1_xboole_0(X2)|~r2_hidden(X1,k3_finseq_2(X2)))).#trainneg
cnf(i_0_326, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|X1=esk8_0|X1=esk4_0|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(esk6_0)))|~r2_hidden(X1,k4_finseq_1(esk6_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2754, plain, (v3_ordinal1(esk15_7(X1,X2,X3,X4,X5,X6,X7))|X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(X7)|~epred1_2(X2,X1)|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_14815, negated_conjecture, (r1_xxreal_0(esk1_0,esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|r1_xxreal_0(esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))),esk1_0)|v1_xboole_0(esk12_1(k1_zfmisc_1(k5_numbers))))).#trainneg
cnf(i_0_14816, negated_conjecture, (r1_xxreal_0(esk2_0,esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|r1_xxreal_0(esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))),esk2_0)|v1_xboole_0(esk12_1(k1_zfmisc_1(k5_numbers))))).#trainneg
cnf(i_0_2755, plain, (v7_ordinal1(esk15_7(X1,X2,X3,X4,X5,X6,X7))|X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|k3_finseq_1(X4)!=k3_finseq_1(X7)|~epred1_2(X2,X1)|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_2796, plain, (X1=k6_matrix_1(X2,k1_xboole_0,X3)|v1_xboole_0(X2)|v7_ordinal1(esk9_4(X2,k1_xboole_0,X3,X1))|k1_matrix_1(k1_xboole_0)!=k3_finseq_1(X1)|~m2_finseq_1(X1,X2)|~v1_matrix_1(k1_xboole_0)|~v7_ordinal1(X3)|~v1_xboole_0(k3_finseq_2(X2)))).#trainneg
cnf(i_0_2910, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8)=k3_matrix_1(u1_struct_0(X1),X4,X7,X8)|X7=X6|X7=X5|v6_struct_0(X1)|~m2_subset_1(X8,k1_numbers,k5_numbers)|~r1_xxreal_0(X8,k1_matrix_1(X9))|~r1_xxreal_0(np__1,X8)|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X9))|~v7_ordinal1(X8)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X9,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_16132, plain, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X3,esk9_4(X5,esk6_0,X6,X7))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X4,esk9_4(X5,esk6_0,X6,X7))|X7=k6_matrix_1(X5,esk6_0,X6)|v1_xboole_0(X5)|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk4_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=esk8_0|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X3|esk15_7(X1,X2,esk5_0,esk6_0,X3,X4,esk7_0)=X4|k1_matrix_1(esk6_0)!=k3_finseq_1(X7)|~m2_finseq_1(esk6_0,k3_finseq_2(X5))|~m2_finseq_1(X7,X5)|~r2_hidden(X3,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(esk5_0),X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~epred1_2(X2,X1)|~m1_matrix_1(esk7_0,u1_struct_0(esk5_0),X1,X2))).#trainneg
cnf(i_0_2849, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X4,X6)=k3_matrix_1(u1_struct_0(X1),esk7_0,X5,X6)|v6_struct_0(X1)|v1_xboole_0(k4_finseq_1(esk7_0))|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k1_matrix_1(X7))|~r1_xxreal_0(np__1,X6)|~m1_matrix_1(X7,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X4,k4_finseq_1(esk7_0)))).#trainneg
cnf(i_0_3027, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X4,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6)))=k3_matrix_1(u1_struct_0(X1),esk6_0,X5,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6)))|v6_struct_0(X1)|~m2_subset_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6)))|~r2_hidden(X4,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X6)))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_3827, plain, (esk13_2(X1,k3_finseq_2(X2))=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X3)|~v1_xboole_0(X2)|~r2_hidden(X1,k3_finseq_2(X2)))).#trainneg
cnf(i_0_4176, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|X1=esk4_0|X1=esk8_0|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X2,k2_finseq_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))))|~r2_hidden(X1,k4_finseq_1(esk6_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2797, plain, (k6_matrix_1(X1,k1_xboole_0,X2)=k8_matrix_1(X1,k1_xboole_0,X2)|v1_xboole_0(X1)|~v1_matrix_1(k1_xboole_0)|~v7_ordinal1(X2)|~v1_xboole_0(k3_finseq_2(X1)))).#trainneg
cnf(i_0_14818, negated_conjecture, (r1_xxreal_0(esk4_0,esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|r1_xxreal_0(esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))),esk4_0)|v1_xboole_0(esk12_1(k1_zfmisc_1(k5_numbers))))).#trainneg
cnf(i_0_14819, negated_conjecture, (r1_xxreal_0(esk8_0,esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|r1_xxreal_0(esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))),esk8_0)|v1_xboole_0(esk12_1(k1_zfmisc_1(k5_numbers))))).#trainneg
cnf(i_0_2798, plain, (m1_subset_1(k3_matrix_1(X1,k1_xboole_0,X2,X3),X1)|~v1_matrix_1(k1_xboole_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v1_xboole_0(k3_finseq_2(X1)))).#trainneg
cnf(i_0_2799, plain, (m2_finseq_1(k8_matrix_1(X1,k1_xboole_0,X2),X1)|v1_xboole_0(X1)|~v1_matrix_1(k1_xboole_0)|~v7_ordinal1(k1_matrix_1(k1_xboole_0))|~v7_ordinal1(X2)|~v1_xboole_0(k3_finseq_2(X1)))).#trainneg
cnf(i_0_10199, plain, (v3_ordinal1(esk13_2(X1,esk12_1(k3_finseq_2(X2))))|~v3_ordinal1(k2_zfmisc_1(k5_numbers,X2))|~r2_hidden(X1,esk12_1(k3_finseq_2(X2))))).#trainneg
cnf(i_0_312, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),X1,X2,X3)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,X3)|X2=esk4_0|X2=esk8_0|X1!=esk7_0|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X3,k2_finseq_1(k1_matrix_1(esk6_0)))|~r2_hidden(X2,k4_finseq_1(esk6_0))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2801, plain, (k1_matrix_1(k1_xboole_0)=k3_finseq_1(X1)|v1_xboole_0(X2)|X1!=k6_matrix_1(X2,k1_xboole_0,X3)|~m2_finseq_1(X1,X2)|~v1_matrix_1(k1_xboole_0)|~v7_ordinal1(X3)|~v1_xboole_0(k3_finseq_2(X2)))).#trainneg
cnf(i_0_2802, plain, (k1_funct_1(X1,esk16_7(X2,X3,X4,k1_xboole_0,X5,X6,X7))=k3_matrix_1(X8,k1_xboole_0,X9,esk16_7(X2,X3,X4,k1_xboole_0,X5,X6,X7))|X7=k1_matrix12(X2,X3,X4,k1_xboole_0,X5,X6)|v1_xboole_0(X8)|v6_struct_0(X4)|v2_struct_0(X4)|X1!=k6_matrix_1(X8,k1_xboole_0,X9)|k3_finseq_1(k1_xboole_0)!=k3_finseq_1(X7)|~epred1_2(X3,X2)|~m2_finseq_1(X1,X8)|~v1_matrix_1(k1_xboole_0)|~m1_matrix_1(X7,u1_struct_0(X4),X2,X3)|~m1_matrix_1(k1_xboole_0,u1_struct_0(X4),X2,X3)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v7_ordinal1(X9)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v1_xboole_0(k3_finseq_2(X8)))).#trainneg
cnf(i_0_2814, plain, (X1=k1_xboole_0|r1_xxreal_0(esk10_2(X1,k1_xboole_0),X2)|r1_xxreal_0(X2,esk10_2(X1,k1_xboole_0))|k3_finseq_1(X1)!=k3_finseq_1(k1_xboole_0)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2973, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X6,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7)))=k3_matrix_1(u1_struct_0(X1),X4,X5,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7)))|v6_struct_0(X1)|~m2_subset_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7)))|~r2_hidden(X6,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X7)))|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_7302, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk13_2(X4,k1_zfmisc_1(k1_xboole_0))))|~v1_xboole_0(X3)|~r2_hidden(X4,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_7352, plain, (v7_ordinal1(esk13_2(X1,esk13_2(X2,k1_zfmisc_1(k5_numbers))))|~r2_hidden(X1,esk13_2(X2,k1_zfmisc_1(k5_numbers)))|~r2_hidden(X2,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2914, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,esk13_2(X8,a_1_0_finseq_1(k1_matrix_1(X4))))=k3_matrix_1(u1_struct_0(X1),X4,X7,esk13_2(X8,a_1_0_finseq_1(k1_matrix_1(X4))))|X7=X6|X7=X5|v6_struct_0(X1)|~m2_subset_1(esk13_2(X8,a_1_0_finseq_1(k1_matrix_1(X4))),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk13_2(X8,a_1_0_finseq_1(k1_matrix_1(X4))))|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X4))|~v7_ordinal1(esk13_2(X8,a_1_0_finseq_1(k1_matrix_1(X4))))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X8,a_1_0_finseq_1(k1_matrix_1(X4))))).#trainneg
cnf(i_0_14820, plain, (r1_xxreal_0(esk12_1(k5_numbers),esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|r1_xxreal_0(esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))),esk12_1(k5_numbers))|v1_xboole_0(esk12_1(k1_zfmisc_1(k5_numbers))))).#trainneg
cnf(i_0_2916, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,esk11_2(X8,k1_matrix_1(X4)))=k3_matrix_1(u1_struct_0(X1),X4,X7,esk11_2(X8,k1_matrix_1(X4)))|X7=X6|X7=X5|v6_struct_0(X1)|~m2_subset_1(esk11_2(X8,k1_matrix_1(X4)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk11_2(X8,k1_matrix_1(X4)))|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X4))|~v7_ordinal1(esk11_2(X8,k1_matrix_1(X4)))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X8,a_1_0_finseq_1(k1_matrix_1(X4))))).#trainneg
cnf(i_0_2964, plain, (m1_subset_1(k3_matrix_1(u1_struct_0(X1),X4,X5,k1_matrix_1(X4)),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X4))|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X5)|~m2_subset_1(k1_matrix_1(X4),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(X4))|~r2_hidden(X6,k4_finseq_1(X4)))).#trainneg
cnf(i_0_18156, plain, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8)=k3_matrix_1(u1_struct_0(X1),X4,X7,X8)|X7=X5|X7=X6|v6_struct_0(X1)|~r2_hidden(X8,a_1_0_finseq_1(k1_matrix_1(X4)))|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k1_matrix_1(X4))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_subset_1(X8,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X8)|~r1_xxreal_0(X8,k1_matrix_1(X4)))).#trainneg
cnf(i_0_18063, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),X1,X2,X3)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,X3)|X2=esk8_0|X2=esk4_0|X1!=esk7_0|~epred1_2(esk1_0,esk2_0)|~r2_hidden(X2,k4_finseq_1(esk6_0))|~m1_matrix_1(X1,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,k1_matrix_1(esk6_0)))).#trainneg
cnf(i_0_5626, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v7_ordinal1(X4)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X4))!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4005, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X5,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6)))=k3_matrix_1(u1_struct_0(X1),esk6_0,X4,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6)))|v6_struct_0(X1)|~m2_subset_1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6)))|~r2_hidden(X5,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X6)))|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_2926, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X4,X6)=k3_matrix_1(u1_struct_0(X1),esk6_0,X5,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X6)|~r2_hidden(X4,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_11711, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(X3,esk6_0,X4,X2)|v1_xboole_0(X3)|X1!=k6_matrix_1(X3,esk6_0,X4)|~m2_finseq_1(esk6_0,k3_finseq_2(X3))|~m2_finseq_1(X1,X3)|~m1_matrix_1(X5,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,k1_matrix_1(X5))|~v7_ordinal1(k1_matrix_1(X5)))).#trainneg
cnf(i_0_16925, plain, (m2_finseq_1(esk13_2(X2,k3_finseq_2(X3)),X1)|~v1_xboole_0(X1)|~v1_xboole_0(X3)|~r2_hidden(X2,k3_finseq_2(X3)))).#trainneg
cnf(i_0_4136, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk12_1(k5_numbers))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk12_1(k5_numbers))|r1_xxreal_0(k1_matrix_1(esk6_0),esk12_1(k5_numbers))|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk12_1(k5_numbers),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk12_1(k5_numbers)))).#trainneg
cnf(i_0_16529, plain, (v3_ordinal1(esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|r1_xxreal_0(k1_xboole_0,esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|v1_xboole_0(esk12_1(k1_zfmisc_1(k5_numbers)))|~m2_subset_1(esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|~v7_ordinal1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_17778, negated_conjecture, (v3_ordinal1(esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|r1_xxreal_0(esk1_0,esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|v1_xboole_0(esk12_1(k1_zfmisc_1(k5_numbers)))|~m2_subset_1(esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|~v7_ordinal1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_18184, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8)=k3_matrix_1(u1_struct_0(X1),X4,X7,X8)|X7=X6|X7=X5|v6_struct_0(X1)|~m2_subset_1(X8,k1_numbers,k5_numbers)|~r1_xxreal_0(X8,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X8)|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_18261, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,X2)|X1!=k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(X4))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_18175, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X6,X7)=k3_matrix_1(u1_struct_0(X1),esk7_0,X6,X7)|X6=X5|X6=X4|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X7)|~r2_hidden(X6,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_18178, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X6,X7)=k3_matrix_1(u1_struct_0(X1),esk6_0,X6,X7)|X6=X5|X6=X4|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X8)))|~r1_xxreal_0(np__1,X7)|~r2_hidden(X6,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X8)))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_18177, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X6,X7)=k3_matrix_1(u1_struct_0(X1),esk6_0,X6,X7)|X6=X5|X6=X4|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X8)))|~r1_xxreal_0(np__1,X7)|~r2_hidden(X6,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X8)))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_18215, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))=k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_18185, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8)=k3_matrix_1(u1_struct_0(X1),X4,X7,X8)|X7=X6|X7=X5|v6_struct_0(X1)|~m2_subset_1(X8,k1_numbers,k5_numbers)|~r1_xxreal_0(X8,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X9)))|~r1_xxreal_0(np__1,X8)|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X9)))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X9))).#trainneg
cnf(i_0_18243, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X4,X6)=k3_matrix_1(u1_struct_0(X1),esk6_0,X5,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X8)))|~r1_xxreal_0(np__1,X6)|~r2_hidden(X4,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_18246, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk6_0,X4,X5),X4,X6)=k3_matrix_1(u1_struct_0(X1),esk6_0,X5,X6)|v6_struct_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~r1_xxreal_0(X6,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X8)))|~r1_xxreal_0(np__1,X6)|~r2_hidden(X4,k4_finseq_1(esk6_0))|~m1_matrix_1(esk6_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X7)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_12725, plain, (m1_subset_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X3,X4))),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X2)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_3881, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk12_1(k5_numbers))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk12_1(k5_numbers))|r1_xxreal_0(k1_matrix_1(esk6_0),esk12_1(k5_numbers))|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk12_1(k5_numbers),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk12_1(k5_numbers)))).#trainneg
cnf(i_0_3082, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(esk13_2(X3,k3_finseq_2(X2)))|~m1_subset_1(X1,esk13_2(X3,k3_finseq_2(X2)))|~r2_hidden(X3,k3_finseq_2(X2)))).#trainneg
cnf(i_0_17779, negated_conjecture, (v3_ordinal1(esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|r1_xxreal_0(esk2_0,esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|v1_xboole_0(esk12_1(k1_zfmisc_1(k5_numbers)))|~m2_subset_1(esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|~v7_ordinal1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_18036, negated_conjecture, (v3_ordinal1(esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|r1_xxreal_0(esk4_0,esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|v1_xboole_0(esk12_1(k1_zfmisc_1(k5_numbers)))|~m2_subset_1(esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|~v7_ordinal1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3083, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(esk12_1(k4_finseq_2(X3,X2)))|~m1_subset_1(X1,esk12_1(k4_finseq_2(X3,X2)))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_3084, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(esk12_1(k3_finseq_2(X2)))|~m1_subset_1(X1,esk12_1(k3_finseq_2(X2))))).#trainneg
cnf(i_0_18293, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,X8)=k3_matrix_1(u1_struct_0(X1),X4,X7,X8)|X7=X5|X7=X6|v6_struct_0(X1)|~m2_subset_1(X8,k1_numbers,k5_numbers)|~r1_xxreal_0(X8,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X9)))|~r1_xxreal_0(np__1,X8)|~r2_hidden(X7,k4_finseq_1(X4))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X9))).#trainneg
cnf(i_0_18316, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X6,X7)=k3_matrix_1(u1_struct_0(X1),esk7_0,X6,X7)|X6=X4|X6=X5|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X8)))|~r1_xxreal_0(np__1,X7)|~r2_hidden(X6,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_18312, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X3,X2)|X1!=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(X4))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(esk5_0))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_18313, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,esk7_0,X4,X5),X6,X7)=k3_matrix_1(u1_struct_0(X1),esk7_0,X6,X7)|X6=X4|X6=X5|v6_struct_0(X1)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~r1_xxreal_0(X7,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X8)))|~r1_xxreal_0(np__1,X7)|~r2_hidden(X6,k4_finseq_1(esk7_0))|~m1_matrix_1(esk7_0,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_7303, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk13_2(X4,k1_zfmisc_1(k1_xboole_0))))|~v1_xboole_0(X3)|~r2_hidden(X4,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_3967, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|k1_matrix_1(esk6_0)!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_18037, negated_conjecture, (v3_ordinal1(esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|r1_xxreal_0(esk8_0,esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|v1_xboole_0(esk12_1(k1_zfmisc_1(k5_numbers)))|~m2_subset_1(esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|~v7_ordinal1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_18133, plain, (v3_ordinal1(esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|r1_xxreal_0(esk12_1(k5_numbers),esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|v1_xboole_0(esk12_1(k1_zfmisc_1(k5_numbers)))|~m2_subset_1(esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk12_1(esk12_1(k1_zfmisc_1(k5_numbers))))|~v7_ordinal1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3095, plain, (esk15_7(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8),X9,X10),X11,X12),X13,X14,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8),X9,X10),X11,X12))=X13|esk15_7(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8),X9,X10),X11,X12),X13,X14,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8),X9,X10),X11,X12))=X14|k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8),X9,X10),X11,X12),X13,X14)=k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,k1_matrix12(X1,X2,X3,X4,X5,X6),X7,X8),X9,X10),X11,X12)|v6_struct_0(X3)|v2_struct_0(X3)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X14)|~v7_ordinal1(X13)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X10)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_18311, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(X3))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),u1_struct_0(esk5_0))|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_18600, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_18601, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_18598, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_18564, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(X3))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),u1_struct_0(esk5_0))|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_18614, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k1_matrix_1(esk6_0)!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_12867, plain, (m1_subset_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X3,X4))),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X2)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_4338, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3096, plain, (r1_xxreal_0(k1_xboole_0,esk13_2(X1,k4_finseq_1(X2)))|r1_xxreal_0(esk13_2(X1,k4_finseq_1(X2)),k1_xboole_0)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k4_finseq_1(X2)))).#trainneg
cnf(i_0_3097, plain, (r1_xxreal_0(np__1,esk13_2(X1,k4_finseq_1(X2)))|r1_xxreal_0(esk13_2(X1,k4_finseq_1(X2)),np__1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k4_finseq_1(X2)))).#trainneg
cnf(i_0_3098, negated_conjecture, (r1_xxreal_0(esk1_0,esk13_2(X1,k4_finseq_1(X2)))|r1_xxreal_0(esk13_2(X1,k4_finseq_1(X2)),esk1_0)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k4_finseq_1(X2)))).#trainneg
cnf(i_0_18657, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_18658, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_18655, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2)=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_18650, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v7_ordinal1(X1)|~r1_xxreal_0(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k1_matrix_1(esk6_0))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))))).#trainneg
cnf(i_0_18678, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk7_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v7_ordinal1(X1)|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))!=k1_matrix_1(esk6_0)|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_18679, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|k1_matrix_1(esk6_0)!=k3_finseq_1(X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_18649, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v7_ordinal1(X1)|~r1_xxreal_0(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k1_matrix_1(esk6_0))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))))).#trainneg
cnf(i_0_12917, plain, (m1_subset_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X3,X4))),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X4)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4337, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3099, negated_conjecture, (r1_xxreal_0(esk2_0,esk13_2(X1,k4_finseq_1(X2)))|r1_xxreal_0(esk13_2(X1,k4_finseq_1(X2)),esk2_0)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k4_finseq_1(X2)))).#trainneg
cnf(i_0_3101, negated_conjecture, (r1_xxreal_0(esk4_0,esk13_2(X1,k4_finseq_1(X2)))|r1_xxreal_0(esk13_2(X1,k4_finseq_1(X2)),esk4_0)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k4_finseq_1(X2)))).#trainneg
cnf(i_0_3102, negated_conjecture, (r1_xxreal_0(esk8_0,esk13_2(X1,k4_finseq_1(X2)))|r1_xxreal_0(esk13_2(X1,k4_finseq_1(X2)),esk8_0)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k4_finseq_1(X2)))).#trainneg
cnf(i_0_18795, negated_conjecture, (X1=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k1_funct_1(X1,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|k3_finseq_1(X1)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_finseq_1(X1)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(X1)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(X1)|~m2_subset_1(esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k1_matrix_1(esk6_0))|~r1_xxreal_0(np__1,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))))).#trainneg
cnf(i_0_18826, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))!=k1_matrix_1(esk6_0)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))).#trainneg
cnf(i_0_18615, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(X1)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|k1_matrix_1(esk6_0)!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(X3)|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_17373, plain, (m1_subset_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X3,X4))),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_4325, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3103, plain, (r1_xxreal_0(esk12_1(k5_numbers),esk13_2(X1,k4_finseq_1(X2)))|r1_xxreal_0(esk13_2(X1,k4_finseq_1(X2)),esk12_1(k5_numbers))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k4_finseq_1(X2)))).#trainneg
cnf(i_0_5130, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk12_1(k2_finseq_1(k1_matrix_1(esk6_0))))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk12_1(k2_finseq_1(k1_matrix_1(esk6_0))))|v1_xboole_0(k2_finseq_1(k1_matrix_1(esk6_0)))|~v7_ordinal1(esk12_1(k2_finseq_1(k1_matrix_1(esk6_0)))))).#trainneg
cnf(i_0_3105, plain, (m1_subset_1(k3_matrix_1(X1,k8_matrix_1(k3_finseq_2(X1),k6_matrix_1(k3_finseq_2(k3_finseq_2(X1)),X2,X3),X4),X5,X6),X1)|v1_xboole_0(k3_finseq_2(X1))|v1_xboole_0(k3_finseq_2(k3_finseq_2(X1)))|~m2_finseq_1(k6_matrix_1(k3_finseq_2(k3_finseq_2(X1)),X2,X3),k3_finseq_2(k3_finseq_2(X1)))|~v1_matrix_1(k8_matrix_1(k3_finseq_2(X1),k6_matrix_1(k3_finseq_2(k3_finseq_2(X1)),X2,X3),X4))|~v1_matrix_1(k6_matrix_1(k3_finseq_2(k3_finseq_2(X1)),X2,X3))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(k3_finseq_2(X1))))|~v1_matrix_1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_3106, plain, (v3_ordinal1(k8_matrix_1(X1,X2,X3))|v1_xboole_0(X1)|~v3_ordinal1(k4_finseq_2(k1_matrix_1(X2),X1))|~m1_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~v7_ordinal1(k1_matrix_1(X2))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_18915, negated_conjecture, (X1=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k1_funct_1(X1,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|k3_finseq_1(X1)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~m2_subset_1(esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3)))|~r1_xxreal_0(np__1,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_18912, negated_conjecture, (X1=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k1_funct_1(X1,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|k3_finseq_1(X1)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~m2_subset_1(esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3)))|~r1_xxreal_0(np__1,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_18928, negated_conjecture, (k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)=X3|k1_matrix_1(esk6_0)!=k3_finseq_1(X3)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m2_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),X3))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3993, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk12_1(k2_finseq_1(k1_matrix_1(esk6_0))))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk12_1(k2_finseq_1(k1_matrix_1(esk6_0))))|v1_xboole_0(k2_finseq_1(k1_matrix_1(esk6_0)))|~v7_ordinal1(esk12_1(k2_finseq_1(k1_matrix_1(esk6_0)))))).#trainneg
cnf(i_0_4324, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X3))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3107, plain, (v3_card_1(k8_matrix_1(X1,X2,X3),k1_matrix_1(X2))|v1_xboole_0(X1)|~v7_ordinal1(k1_matrix_1(X2))|~m1_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_3179, plain, (k1_funct_1(k8_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7),esk13_2(X8,k2_finseq_1(k1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6)))))=k3_matrix_1(u1_struct_0(X1),k1_matrix12(X2,X3,X1,X4,X5,X6),X7,esk13_2(X8,k2_finseq_1(k1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6)))))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~m2_finseq_1(k1_matrix12(X2,X3,X1,X4,X5,X6),k3_finseq_2(u1_struct_0(X1)))|~v1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6))|~r2_hidden(X8,k2_finseq_1(k1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6))))|~v7_ordinal1(esk13_2(X8,k2_finseq_1(k1_matrix_1(k1_matrix12(X2,X3,X1,X4,X5,X6)))))|~v7_ordinal1(X7)|~m1_matrix_1(X4,u1_struct_0(X1),X2,X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_3180, plain, (k1_funct_1(k8_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4),esk13_2(X5,k2_finseq_1(k1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3)))))=k3_matrix_1(X1,k6_matrix_1(k3_finseq_2(X1),X2,X3),X4,esk13_2(X5,k2_finseq_1(k1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3)))))|v1_xboole_0(X1)|v1_xboole_0(k3_finseq_2(X1))|~m2_finseq_1(k6_matrix_1(k3_finseq_2(X1),X2,X3),k3_finseq_2(X1))|~v1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))|~r2_hidden(X5,k2_finseq_1(k1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3))))|~v7_ordinal1(esk13_2(X5,k2_finseq_1(k1_matrix_1(k6_matrix_1(k3_finseq_2(X1),X2,X3)))))|~v7_ordinal1(X4)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(X1)))|~v1_matrix_1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_18925, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k5_numbers))).#trainneg
cnf(i_0_18926, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))|k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))!=k1_matrix_1(esk6_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_18616, negated_conjecture, (k1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2)=X3|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),u1_struct_0(esk5_0))|~m1_matrix_1(X4,u1_struct_0(esk5_0),esk2_0,esk1_0)|~v7_ordinal1(X1)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X2),X3))|k1_matrix_1(X4)!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_18913, negated_conjecture, (X1=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k1_funct_1(X1,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|k3_finseq_1(X1)!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~m2_subset_1(esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k1_matrix_1(X3))|~r1_xxreal_0(np__1,esk10_2(X1,k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(X2)|~m1_matrix_1(X3,u1_struct_0(esk5_0),esk2_0,esk1_0))).#trainneg
cnf(i_0_19041, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0)=k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0)|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1))!=k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_19127, negated_conjecture, (k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)|k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))!=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,X2,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1))!=k3_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))|~m2_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2),u1_struct_0(esk5_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_finseq_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_funct_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2))|~v1_relat_1(k6_matrix_1(u1_struct_0(esk5_0),esk6_0,X2)))).#trainneg
cnf(i_0_14430, plain, (m1_subset_1(esk13_2(X3,k1_zfmisc_1(k1_xboole_0)),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X2)|~r2_hidden(X3,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_3946, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1)=X2|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X1),X2))|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3197, plain, (r1_xxreal_0(esk13_2(esk15_7(X1,X2,X3,X4,X5,X6,X7),k4_finseq_1(X4)),X8)|r1_xxreal_0(X8,esk13_2(esk15_7(X1,X2,X3,X4,X5,X6,X7),k4_finseq_1(X4)))|X7=k1_matrix12(X1,X2,X3,X4,X5,X6)|v6_struct_0(X3)|v2_struct_0(X3)|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4)|~v7_ordinal1(X8)|k3_finseq_1(X4)!=k3_finseq_1(X7)|~epred1_2(X2,X1)|~m1_matrix_1(X7,u1_struct_0(X3),X1,X2)|~m1_matrix_1(X4,u1_struct_0(X3),X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_3198, plain, (r1_xxreal_0(esk13_2(esk13_2(X1,k4_finseq_1(X2)),k4_finseq_1(X2)),X3)|r1_xxreal_0(X3,esk13_2(esk13_2(X1,k4_finseq_1(X2)),k4_finseq_1(X2)))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X3)|~r2_hidden(X1,k4_finseq_1(X2)))).#trainneg
cnf(i_0_14711, plain, (m1_subset_1(esk13_2(X3,k1_zfmisc_1(k1_xboole_0)),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X1)|~r2_hidden(X3,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_3945, negated_conjecture, (k3_matrix_1(u1_struct_0(esk5_0),esk7_0,esk8_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))=k3_matrix_1(u1_struct_0(esk5_0),esk6_0,esk4_0,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1)=X2|k3_finseq_1(k8_matrix_1(u1_struct_0(esk5_0),esk6_0,X3))!=k3_finseq_1(X2)|~epred1_2(esk1_0,esk2_0)|~m2_subset_1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(esk10_2(k8_matrix_1(u1_struct_0(esk5_0),esk7_0,X1),X2))|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3301, plain, (r1_xxreal_0(X1,esk13_2(esk13_2(X2,esk12_1(k1_zfmisc_1(k4_finseq_1(X3)))),k4_finseq_1(X3)))|r1_xxreal_0(esk13_2(esk13_2(X2,esk12_1(k1_zfmisc_1(k4_finseq_1(X3)))),k4_finseq_1(X3)),X1)|v1_xboole_0(k4_finseq_1(X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X1)|~r2_hidden(X2,esk12_1(k1_zfmisc_1(k4_finseq_1(X3)))))).#trainneg
cnf(i_0_14610, plain, (v1_relat_1(esk12_1(esk12_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))|v1_xboole_0(esk12_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_14611, plain, (v1_xboole_0(esk12_1(esk12_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))|v1_xboole_0(esk12_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3303, plain, (r1_xxreal_0(X1,esk13_2(esk12_1(k4_finseq_1(X2)),k4_finseq_1(X2)))|r1_xxreal_0(esk13_2(esk12_1(k4_finseq_1(X2)),k4_finseq_1(X2)),X1)|v1_xboole_0(k4_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 93
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 182
# Removed in clause preprocessing      : 23
# Initial clauses in saturation        : 159
# Processed clauses                    : 4155
# ...of these trivial                  : 25
# ...subsumed                          : 2168
# ...remaining for further processing  : 1961
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 262
# Backward-rewritten                   : 438
# Generated clauses                    : 19303
# ...of the previous two non-trivial   : 18922
# Contextual simplify-reflections      : 762
# Paramodulations                      : 19187
# Factorizations                       : 0
# Equation resolutions                 : 105
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1247
#    Positive orientable unit clauses  : 87
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 8
#    Non-unit-clauses                  : 1152
# Current number of unprocessed clauses: 14128
# ...number of literals in the above   : 149679
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2673
# Proof object given clauses           : 249
# Proof search given clauses           : 1962
# Clause-clause subsumption calls (NU) : 5339015
# Rec. Clause-clause subsumption calls : 160460
# Non-unit clause-clause subsumptions  : 3093
# Unit Clause-clause subsumption calls : 20250
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 41
# BW rewrite match successes           : 20
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 828053

# -------------------------------------------------
# User time                : 111.848 s
# System time              : 3.540 s
# Total time               : 115.388 s
# Maximum resident set size: 454000 pages
111.85user 3.65system 1:56.93elapsed 98%CPU (0avgtext+0avgdata 454592maxresident)k
32inputs+1856outputs (0major+864517minor)pagefaults 0swaps
