# ENIGMA: Connected to the TF server '127.0.0.1:8888'.
# Preprocessing time       : 0.033 s
#TF#SERVER: Sending data to the TF server (conj=7, context=207, query=131, total=345).
#TF#SERVER: Sent 54666 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1356 bytes expected).
#TF#SERVER: Received 1356 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=240, query=132, total=379).
#TF#SERVER: Sent 82586 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1492 bytes expected).
#TF#SERVER: Received 1492 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=281, query=138, total=426).
#TF#SERVER: Sent 135100 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1680 bytes expected).
#TF#SERVER: Received 1680 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=348, query=129, total=484).
#TF#SERVER: Sent 130588 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1912 bytes expected).
#TF#SERVER: Received 1912 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=416, query=128, total=551).
#TF#SERVER: Sent 149980 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2180 bytes expected).
#TF#SERVER: Received 2180 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=456, query=129, total=592).
#TF#SERVER: Sent 179330 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2344 bytes expected).
#TF#SERVER: Received 2344 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=492, query=151, total=650).
#TF#SERVER: Sent 212319 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2576 bytes expected).
#TF#SERVER: Received 2576 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=136, total=655).
#TF#SERVER: Sent 247434 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2596 bytes expected).
#TF#SERVER: Received 2596 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=130, total=649).
#TF#SERVER: Sent 229141 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2572 bytes expected).
#TF#SERVER: Received 2572 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=130, total=649).
#TF#SERVER: Sent 273967 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2572 bytes expected).
#TF#SERVER: Received 2572 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 337839 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=139, total=658).
#TF#SERVER: Sent 341108 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2608 bytes expected).
#TF#SERVER: Received 2608 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=135, total=654).
#TF#SERVER: Sent 337590 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2592 bytes expected).
#TF#SERVER: Received 2592 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=129, total=648).
#TF#SERVER: Sent 342947 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2568 bytes expected).
#TF#SERVER: Received 2568 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=143, total=662).
#TF#SERVER: Sent 335033 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2624 bytes expected).
#TF#SERVER: Received 2624 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=132, total=651).
#TF#SERVER: Sent 325349 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2580 bytes expected).
#TF#SERVER: Received 2580 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=134, total=653).
#TF#SERVER: Sent 358227 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2588 bytes expected).
#TF#SERVER: Received 2588 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=130, total=649).
#TF#SERVER: Sent 256015 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2572 bytes expected).
#TF#SERVER: Received 2572 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=137, total=656).
#TF#SERVER: Sent 361570 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2600 bytes expected).
#TF#SERVER: Received 2600 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=133, total=652).
#TF#SERVER: Sent 334523 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2584 bytes expected).
#TF#SERVER: Received 2584 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=143, total=662).
#TF#SERVER: Sent 370344 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2624 bytes expected).
#TF#SERVER: Received 2624 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=129, total=648).
#TF#SERVER: Sent 412024 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2568 bytes expected).
#TF#SERVER: Received 2568 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=135, total=654).
#TF#SERVER: Sent 373814 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2592 bytes expected).
#TF#SERVER: Received 2592 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 306348 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=135, total=654).
#TF#SERVER: Sent 333269 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2592 bytes expected).
#TF#SERVER: Received 2592 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=132, total=651).
#TF#SERVER: Sent 372512 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2580 bytes expected).
#TF#SERVER: Received 2580 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=137, total=656).
#TF#SERVER: Sent 384355 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2600 bytes expected).
#TF#SERVER: Received 2600 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=129, total=648).
#TF#SERVER: Sent 365124 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2568 bytes expected).
#TF#SERVER: Received 2568 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=129, total=648).
#TF#SERVER: Sent 338118 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2568 bytes expected).
#TF#SERVER: Received 2568 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=135, total=654).
#TF#SERVER: Sent 295078 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2592 bytes expected).
#TF#SERVER: Received 2592 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=129, total=648).
#TF#SERVER: Sent 331951 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2568 bytes expected).
#TF#SERVER: Received 2568 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=132, total=651).
#TF#SERVER: Sent 280896 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2580 bytes expected).
#TF#SERVER: Received 2580 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=150, total=669).
#TF#SERVER: Sent 361934 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2652 bytes expected).
#TF#SERVER: Received 2652 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 403929 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=93, total=612).
#TF#SERVER: Sent 365250 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2424 bytes expected).
#TF#SERVER: Received 2424 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=131, total=650).
#TF#SERVER: Sent 389328 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2576 bytes expected).
#TF#SERVER: Received 2576 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=129, total=648).
#TF#SERVER: Sent 431086 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2568 bytes expected).
#TF#SERVER: Received 2568 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=170, total=689).
#TF#SERVER: Sent 429083 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2732 bytes expected).
#TF#SERVER: Received 2732 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=148, total=667).
#TF#SERVER: Sent 424190 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2644 bytes expected).
#TF#SERVER: Received 2644 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 346343 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=129, total=648).
#TF#SERVER: Sent 417644 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2568 bytes expected).
#TF#SERVER: Received 2568 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=147, total=666).
#TF#SERVER: Sent 356323 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2640 bytes expected).
#TF#SERVER: Received 2640 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=165, total=684).
#TF#SERVER: Sent 421868 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2712 bytes expected).
#TF#SERVER: Received 2712 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=140, total=659).
#TF#SERVER: Sent 492361 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2612 bytes expected).
#TF#SERVER: Received 2612 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 467048 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=118, total=637).
#TF#SERVER: Sent 428556 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2524 bytes expected).
#TF#SERVER: Received 2524 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=162, total=681).
#TF#SERVER: Sent 511307 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2700 bytes expected).
#TF#SERVER: Received 2700 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=150, total=669).
#TF#SERVER: Sent 512121 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2652 bytes expected).
#TF#SERVER: Received 2652 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=156, total=675).
#TF#SERVER: Sent 522336 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2676 bytes expected).
#TF#SERVER: Received 2676 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 443852 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=91, total=610).
#TF#SERVER: Sent 362535 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2416 bytes expected).
#TF#SERVER: Received 2416 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=131, total=650).
#TF#SERVER: Sent 422619 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2576 bytes expected).
#TF#SERVER: Received 2576 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=164, total=683).
#TF#SERVER: Sent 541916 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2708 bytes expected).
#TF#SERVER: Received 2708 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=166, total=685).
#TF#SERVER: Sent 561121 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2716 bytes expected).
#TF#SERVER: Received 2716 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 315860 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 472167 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=15, total=534).
#TF#SERVER: Sent 173317 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2112 bytes expected).
#TF#SERVER: Received 2112 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=174, total=693).
#TF#SERVER: Sent 562468 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2748 bytes expected).
#TF#SERVER: Received 2748 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 305331 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 449631 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=65, total=584).
#TF#SERVER: Sent 285734 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2312 bytes expected).
#TF#SERVER: Received 2312 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=136, total=655).
#TF#SERVER: Sent 326592 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2596 bytes expected).
#TF#SERVER: Received 2596 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 418355 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 444045 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=11, total=530).
#TF#SERVER: Sent 163762 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2096 bytes expected).
#TF#SERVER: Received 2096 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=184, total=703).
#TF#SERVER: Sent 582211 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2788 bytes expected).
#TF#SERVER: Received 2788 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=145, total=664).
#TF#SERVER: Sent 453877 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2632 bytes expected).
#TF#SERVER: Received 2632 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=129, total=648).
#TF#SERVER: Sent 389247 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2568 bytes expected).
#TF#SERVER: Received 2568 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=141, total=660).
#TF#SERVER: Sent 490218 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2616 bytes expected).
#TF#SERVER: Received 2616 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=130, total=649).
#TF#SERVER: Sent 389721 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2572 bytes expected).
#TF#SERVER: Received 2572 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=175, total=694).
#TF#SERVER: Sent 554003 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2752 bytes expected).
#TF#SERVER: Received 2752 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=164, total=683).
#TF#SERVER: Sent 483704 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2708 bytes expected).
#TF#SERVER: Received 2708 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=145, total=664).
#TF#SERVER: Sent 490708 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2632 bytes expected).
#TF#SERVER: Received 2632 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=166, total=685).
#TF#SERVER: Sent 431231 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2716 bytes expected).
#TF#SERVER: Received 2716 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=148, total=667).
#TF#SERVER: Sent 513688 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2644 bytes expected).
#TF#SERVER: Received 2644 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 373299 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=90, total=609).
#TF#SERVER: Sent 358046 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2412 bytes expected).
#TF#SERVER: Received 2412 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=148, total=667).
#TF#SERVER: Sent 497955 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2644 bytes expected).
#TF#SERVER: Received 2644 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=181, total=700).
#TF#SERVER: Sent 517399 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2776 bytes expected).
#TF#SERVER: Received 2776 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=151, total=670).
#TF#SERVER: Sent 499490 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2656 bytes expected).
#TF#SERVER: Received 2656 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=180, total=699).
#TF#SERVER: Sent 531272 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2772 bytes expected).
#TF#SERVER: Received 2772 bytes from the TF server.

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(dt_k2_pre_poly, axiom, ![X1]:(v1_xboole_0(k2_pre_poly(X1))&m2_finseq_2(k2_pre_poly(X1),X1,k3_finseq_2(X1))), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', dt_k2_pre_poly)).
fof(redefinition_k2_pre_poly, axiom, ![X1]:k2_pre_poly(X1)=k6_finseq_1(X1), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', redefinition_k2_pre_poly)).
fof(fc7_finseq_5, axiom, ![X1]:((((v1_xboole_0(X1)&v1_relat_1(X1))&v1_funct_1(X1))&v1_finseq_1(X1))=>(((v1_xboole_0(k3_finseq_5(X1))&v1_relat_1(k3_finseq_5(X1)))&v1_funct_1(k3_finseq_5(X1)))&v1_finseq_1(k3_finseq_5(X1)))), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', fc7_finseq_5)).
fof(cc1_funct_1, axiom, ![X1]:(v1_xboole_0(X1)=>v1_funct_1(X1)), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', cc1_funct_1)).
fof(cc4_finseq_1, axiom, ![X1]:((v1_relat_1(X1)&v1_xboole_0(X1))=>(v1_relat_1(X1)&v1_finseq_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', cc4_finseq_1)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', t6_boole)).
fof(rc2_finseq_1, axiom, ![X1]:?[X2]:(((((((m1_finseq_1(X2,X1)&v1_relat_1(X2))&v4_relat_1(X2,k5_numbers))&v5_relat_1(X2,X1))&v1_funct_1(X2))&v1_xboole_0(X2))&v1_finset_1(X2))&v1_finseq_1(X2)), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', rc2_finseq_1)).
fof(s2_finseq_2__e3_2__polynom3, axiom, ![X1]:(((((~(v2_struct_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&l2_algstr_0(X1))=>((k4_rlvect_1(X1,k6_finseq_1(u1_struct_0(X1)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k6_finseq_1(u1_struct_0(X1))))&![X2]:(m2_finseq_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k4_rlvect_1(X1,X2)=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=>k4_rlvect_1(X1,k7_finseq_1(X2,k9_finseq_1(X3)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k7_finseq_1(X2,k9_finseq_1(X3))))))))=>![X2]:(m2_finseq_1(X2,u1_struct_0(X1))=>k4_rlvect_1(X1,X2)=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))))), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', s2_finseq_2__e3_2__polynom3)).
fof(involutiveness_k3_finseq_5, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))=>k3_finseq_5(k3_finseq_5(X1))=X1), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', involutiveness_k3_finseq_5)).
fof(involutiveness_k4_finseq_5, axiom, ![X1, X2]:(m1_finseq_1(X2,X1)=>k4_finseq_5(X1,k4_finseq_5(X1,X2))=X2), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', involutiveness_k4_finseq_5)).
fof(redefinition_k4_finseq_5, axiom, ![X1, X2]:(m1_finseq_1(X2,X1)=>k4_finseq_5(X1,X2)=k3_finseq_5(X2)), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', redefinition_k4_finseq_5)).
fof(redefinition_m2_finseq_1, axiom, ![X1, X2]:(m2_finseq_1(X2,X1)<=>m1_finseq_1(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', redefinition_m2_finseq_1)).
fof(dt_k4_finseq_5, axiom, ![X1, X2]:(m1_finseq_1(X2,X1)=>m2_finseq_1(k4_finseq_5(X1,X2),X1)), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', dt_k4_finseq_5)).
fof(dt_o_0_0_xboole_0, axiom, v1_xboole_0(o_0_0_xboole_0), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', dt_o_0_0_xboole_0)).
fof(d2_xboole_0, axiom, k1_xboole_0=o_0_0_xboole_0, file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', d2_xboole_0)).
fof(t2_polynom3, conjecture, ![X1]:(((((~(v2_struct_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&l2_algstr_0(X1))=>![X2]:(m2_finseq_1(X2,u1_struct_0(X1))=>k4_rlvect_1(X1,X2)=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2)))), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', t2_polynom3)).
fof(dt_m1_finseq_1, axiom, ![X1, X2]:(m1_finseq_1(X2,X1)=>((v1_relat_1(X2)&v1_funct_1(X2))&v1_finseq_1(X2))), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', dt_m1_finseq_1)).
fof(dt_m2_finseq_1, axiom, ![X1, X2]:(m2_finseq_1(X2,X1)=>((v1_funct_1(X2)&v1_finseq_1(X2))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', dt_m2_finseq_1)).
fof(redefinition_k3_rlvect_1, axiom, ![X1, X2, X3]:((((v2_rlvect_1(X1)&l1_algstr_0(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>k3_rlvect_1(X1,X2,X3)=k1_algstr_0(X1,X2,X3)), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', redefinition_k3_rlvect_1)).
fof(commutativity_k3_rlvect_1, axiom, ![X1, X2, X3]:((((v2_rlvect_1(X1)&l1_algstr_0(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>k3_rlvect_1(X1,X2,X3)=k3_rlvect_1(X1,X3,X2)), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', commutativity_k3_rlvect_1)).
fof(dt_k4_rlvect_1, axiom, ![X1, X2]:((((((~(v2_struct_0(X1))&l2_algstr_0(X1))&v1_relat_1(X2))&v5_relat_1(X2,u1_struct_0(X1)))&v1_funct_1(X2))&v1_finseq_1(X2))=>m1_subset_1(k4_rlvect_1(X1,X2),u1_struct_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', dt_k4_rlvect_1)).
fof(cc3_finseq_1, axiom, ![X1, X2]:(m1_finseq_1(X2,X1)=>v5_relat_1(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', cc3_finseq_1)).
fof(t41_rlvect_1, axiom, ![X1]:((((~(v2_struct_0(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&l2_algstr_0(X1))=>![X2]:(m2_finseq_1(X2,u1_struct_0(X1))=>![X3]:(m2_finseq_1(X3,u1_struct_0(X1))=>k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3))=k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,X3))))), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', t41_rlvect_1)).
fof(dt_k3_finseq_5, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))=>((v1_relat_1(k3_finseq_5(X1))&v1_funct_1(k3_finseq_5(X1)))&v1_finseq_1(k3_finseq_5(X1)))), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', dt_k3_finseq_5)).
fof(dt_l2_algstr_0, axiom, ![X1]:(l2_algstr_0(X1)=>(l2_struct_0(X1)&l1_algstr_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', dt_l2_algstr_0)).
fof(redefinition_k8_finseq_1, axiom, ![X1, X2, X3]:((m1_finseq_1(X2,X1)&m1_finseq_1(X3,X1))=>k8_finseq_1(X1,X2,X3)=k7_finseq_1(X2,X3)), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', redefinition_k8_finseq_1)).
fof(t63_finseq_5, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))=>![X2]:k3_finseq_5(k7_finseq_1(X1,k9_finseq_1(X2)))=k7_finseq_1(k9_finseq_1(X2),k3_finseq_5(X1))), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', t63_finseq_5)).
fof(dt_k8_finseq_1, axiom, ![X1, X2, X3]:((m1_finseq_1(X2,X1)&m1_finseq_1(X3,X1))=>m2_finseq_1(k8_finseq_1(X1,X2,X3),X1)), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', dt_k8_finseq_1)).
fof(redefinition_m2_finseq_2, axiom, ![X1, X2]:(m1_finseq_2(X2,X1)=>![X3]:(m2_finseq_2(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', redefinition_m2_finseq_2)).
fof(dt_k3_finseq_2, axiom, ![X1]:m1_finseq_2(k3_finseq_2(X1),X1), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', dt_k3_finseq_2)).
fof(dt_k3_pre_poly, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>m2_finseq_2(k3_pre_poly(X1,X2),X1,k3_finseq_2(X1))), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', dt_k3_pre_poly)).
fof(redefinition_k3_pre_poly, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k3_pre_poly(X1,X2)=k5_finseq_1(X2)), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', redefinition_k3_pre_poly)).
fof(redefinition_k9_finseq_1, axiom, ![X1]:k9_finseq_1(X1)=k5_finseq_1(X1), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', redefinition_k9_finseq_1)).
fof(dt_m2_finseq_2, axiom, ![X1, X2]:(m1_finseq_2(X2,X1)=>![X3]:(m2_finseq_2(X3,X1,X2)=>m2_finseq_1(X3,X1))), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', dt_m2_finseq_2)).
fof(dt_l2_struct_0, axiom, ![X1]:(l2_struct_0(X1)=>l1_struct_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', dt_l2_struct_0)).
fof(fc1_struct_0, axiom, ![X1]:((v2_struct_0(X1)&l1_struct_0(X1))=>v1_xboole_0(u1_struct_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', fc1_struct_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', fc2_struct_0)).
fof(dt_l1_algstr_0, axiom, ![X1]:(l1_algstr_0(X1)=>l1_struct_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t2_polynom3', dt_l1_algstr_0)).
fof(c_0_38, plain, ![X41]:(v1_xboole_0(k2_pre_poly(X41))&m2_finseq_2(k2_pre_poly(X41),X41,k3_finseq_2(X41))), inference(variable_rename,[status(thm)],[dt_k2_pre_poly])).
fof(c_0_39, plain, ![X31]:k2_pre_poly(X31)=k6_finseq_1(X31), inference(variable_rename,[status(thm)],[redefinition_k2_pre_poly])).
fof(c_0_40, plain, ![X92]:((((v1_xboole_0(k3_finseq_5(X92))|(~v1_xboole_0(X92)|~v1_relat_1(X92)|~v1_funct_1(X92)|~v1_finseq_1(X92)))&(v1_relat_1(k3_finseq_5(X92))|(~v1_xboole_0(X92)|~v1_relat_1(X92)|~v1_funct_1(X92)|~v1_finseq_1(X92))))&(v1_funct_1(k3_finseq_5(X92))|(~v1_xboole_0(X92)|~v1_relat_1(X92)|~v1_funct_1(X92)|~v1_finseq_1(X92))))&(v1_finseq_1(k3_finseq_5(X92))|(~v1_xboole_0(X92)|~v1_relat_1(X92)|~v1_funct_1(X92)|~v1_finseq_1(X92)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc7_finseq_5])])])).
fof(c_0_41, plain, ![X112]:(~v1_xboole_0(X112)|v1_funct_1(X112)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_funct_1])])).
fof(c_0_42, plain, ![X144]:((v1_relat_1(X144)|(~v1_relat_1(X144)|~v1_xboole_0(X144)))&(v1_finseq_1(X144)|(~v1_relat_1(X144)|~v1_xboole_0(X144)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_finseq_1])])])).
fof(c_0_43, plain, ![X49]:(~v1_xboole_0(X49)|X49=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
cnf(c_0_44, plain, (v1_xboole_0(k2_pre_poly(X1))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_45, plain, (k2_pre_poly(X1)=k6_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_46, plain, ![X119]:(((((((m1_finseq_1(esk6_1(X119),X119)&v1_relat_1(esk6_1(X119)))&v4_relat_1(esk6_1(X119),k5_numbers))&v5_relat_1(esk6_1(X119),X119))&v1_funct_1(esk6_1(X119)))&v1_xboole_0(esk6_1(X119)))&v1_finset_1(esk6_1(X119)))&v1_finseq_1(esk6_1(X119))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc2_finseq_1])])).
cnf(c_0_47, plain, (v1_xboole_0(k3_finseq_5(X1))|~v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_48, plain, (v1_funct_1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_49, plain, (v1_finseq_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_50, plain, ![X7, X10]:((m2_finseq_1(esk3_1(X7),u1_struct_0(X7))|k4_rlvect_1(X7,k6_finseq_1(u1_struct_0(X7)))!=k4_rlvect_1(X7,k4_finseq_5(u1_struct_0(X7),k6_finseq_1(u1_struct_0(X7))))|(~m2_finseq_1(X10,u1_struct_0(X7))|k4_rlvect_1(X7,X10)=k4_rlvect_1(X7,k4_finseq_5(u1_struct_0(X7),X10)))|(v2_struct_0(X7)|~v2_rlvect_1(X7)|~v3_rlvect_1(X7)|~v4_rlvect_1(X7)|~l2_algstr_0(X7)))&((m1_subset_1(esk4_1(X7),u1_struct_0(X7))|k4_rlvect_1(X7,k6_finseq_1(u1_struct_0(X7)))!=k4_rlvect_1(X7,k4_finseq_5(u1_struct_0(X7),k6_finseq_1(u1_struct_0(X7))))|(~m2_finseq_1(X10,u1_struct_0(X7))|k4_rlvect_1(X7,X10)=k4_rlvect_1(X7,k4_finseq_5(u1_struct_0(X7),X10)))|(v2_struct_0(X7)|~v2_rlvect_1(X7)|~v3_rlvect_1(X7)|~v4_rlvect_1(X7)|~l2_algstr_0(X7)))&((k4_rlvect_1(X7,esk3_1(X7))=k4_rlvect_1(X7,k4_finseq_5(u1_struct_0(X7),esk3_1(X7)))|k4_rlvect_1(X7,k6_finseq_1(u1_struct_0(X7)))!=k4_rlvect_1(X7,k4_finseq_5(u1_struct_0(X7),k6_finseq_1(u1_struct_0(X7))))|(~m2_finseq_1(X10,u1_struct_0(X7))|k4_rlvect_1(X7,X10)=k4_rlvect_1(X7,k4_finseq_5(u1_struct_0(X7),X10)))|(v2_struct_0(X7)|~v2_rlvect_1(X7)|~v3_rlvect_1(X7)|~v4_rlvect_1(X7)|~l2_algstr_0(X7)))&(k4_rlvect_1(X7,k7_finseq_1(esk3_1(X7),k9_finseq_1(esk4_1(X7))))!=k4_rlvect_1(X7,k4_finseq_5(u1_struct_0(X7),k7_finseq_1(esk3_1(X7),k9_finseq_1(esk4_1(X7)))))|k4_rlvect_1(X7,k6_finseq_1(u1_struct_0(X7)))!=k4_rlvect_1(X7,k4_finseq_5(u1_struct_0(X7),k6_finseq_1(u1_struct_0(X7))))|(~m2_finseq_1(X10,u1_struct_0(X7))|k4_rlvect_1(X7,X10)=k4_rlvect_1(X7,k4_finseq_5(u1_struct_0(X7),X10)))|(v2_struct_0(X7)|~v2_rlvect_1(X7)|~v3_rlvect_1(X7)|~v4_rlvect_1(X7)|~l2_algstr_0(X7)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[s2_finseq_2__e3_2__polynom3])])])])])])).
cnf(c_0_51, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_52, plain, (v1_xboole_0(k6_finseq_1(X1))), inference(rw,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_53, plain, (v1_xboole_0(esk6_1(X1))), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_54, plain, ![X21]:(~v1_relat_1(X21)|~v1_funct_1(X21)|~v1_finseq_1(X21)|k3_finseq_5(k3_finseq_5(X21))=X21), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k3_finseq_5])])).
cnf(c_0_55, plain, (v1_xboole_0(k3_finseq_5(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_47])).
cnf(c_0_56, plain, (v1_funct_1(X1)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_48])).
cnf(c_0_57, plain, (v1_finseq_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_49])).
fof(c_0_58, plain, ![X38, X39]:(~m1_finseq_1(X39,X38)|k4_finseq_5(X38,k4_finseq_5(X38,X39))=X39), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k4_finseq_5])])).
fof(c_0_59, plain, ![X11, X12]:(~m1_finseq_1(X12,X11)|k4_finseq_5(X11,X12)=k3_finseq_5(X12)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_finseq_5])])).
cnf(c_0_60, plain, (k4_rlvect_1(X1,esk3_1(X1))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))|k4_rlvect_1(X1,X2)=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k6_finseq_1(u1_struct_0(X1)))!=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k6_finseq_1(u1_struct_0(X1))))|~m2_finseq_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_61, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_51])).
cnf(c_0_62, plain, (v1_xboole_0(k6_finseq_1(X1))), inference(evalgc,[status(thm)],[c_0_52])).
cnf(c_0_63, plain, (m1_finseq_1(esk6_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_64, plain, (v1_xboole_0(esk6_1(X1))), inference(evalgc,[status(thm)],[c_0_53])).
cnf(c_0_65, plain, (k3_finseq_5(k3_finseq_5(X1))=X1|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_66, plain, (v1_xboole_0(k3_finseq_5(X1))|~v1_xboole_0(X1)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_55, c_0_56]), c_0_57])).
cnf(c_0_67, plain, (k4_finseq_5(X2,k4_finseq_5(X2,X1))=X1|~m1_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_68, plain, (k4_finseq_5(X2,X1)=k3_finseq_5(X1)|~m1_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
fof(c_0_69, plain, ![X18, X19]:((~m2_finseq_1(X19,X18)|m1_finseq_1(X19,X18))&(~m1_finseq_1(X19,X18)|m2_finseq_1(X19,X18))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_m2_finseq_1])])).
cnf(c_0_70, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k6_finseq_1(u1_struct_0(X1))))!=k4_rlvect_1(X1,k6_finseq_1(u1_struct_0(X1)))|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_60])).
cnf(c_0_71, plain, (k6_finseq_1(X1)=k1_xboole_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62])])).
cnf(c_0_72, plain, (m1_finseq_1(esk6_1(X1),X1)), inference(evalgc,[status(thm)],[c_0_63])).
cnf(c_0_73, plain, (esk6_1(X1)=k1_xboole_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_64])])).
fof(c_0_74, plain, ![X25, X26]:(~m1_finseq_1(X26,X25)|m2_finseq_1(k4_finseq_5(X25,X26),X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_finseq_5])])).
cnf(c_0_75, plain, (k3_finseq_5(k3_finseq_5(X1))=X1|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_65])).
cnf(c_0_76, plain, (k3_finseq_5(X1)=k1_xboole_0|~v1_xboole_0(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_66])])).
cnf(c_0_77, plain, (v1_xboole_0(o_0_0_xboole_0)), inference(split_conjunct,[status(thm)],[dt_o_0_0_xboole_0])).
cnf(c_0_78, plain, (k1_xboole_0=o_0_0_xboole_0), inference(split_conjunct,[status(thm)],[d2_xboole_0])).
cnf(c_0_79, plain, (v1_relat_1(esk6_1(X1))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_80, plain, (k4_finseq_5(X2,k4_finseq_5(X2,X1))=X1|~m1_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_67])).
cnf(c_0_81, plain, (k4_finseq_5(X2,X1)=k3_finseq_5(X1)|~m1_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_68])).
cnf(c_0_82, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_83, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k1_xboole_0))!=k4_rlvect_1(X1,k1_xboole_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_71]), c_0_71])])).
cnf(c_0_84, plain, (m1_finseq_1(k1_xboole_0,X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_73])])).
cnf(c_0_85, plain, (m2_finseq_1(k4_finseq_5(X2,X1),X2)|~m1_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_86, plain, (k3_finseq_5(k1_xboole_0)=X1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76])])).
cnf(c_0_87, plain, (v1_xboole_0(o_0_0_xboole_0)), inference(evalgc,[status(thm)],[c_0_77])).
cnf(c_0_88, plain, (o_0_0_xboole_0=k1_xboole_0), inference(evalgc,[status(thm)],[c_0_78])).
cnf(c_0_89, plain, (v1_relat_1(esk6_1(X1))), inference(evalgc,[status(thm)],[c_0_79])).
cnf(c_0_90, plain, (k4_finseq_5(X1,k3_finseq_5(X2))=X2|~m1_finseq_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81])])).
cnf(c_0_91, plain, (k3_finseq_5(k4_finseq_5(X1,X2))=X2|~m1_finseq_1(k4_finseq_5(X1,X2),X1)|~m1_finseq_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_80])])).
cnf(c_0_92, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_82])).
cnf(c_0_93, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|k4_rlvect_1(X1,k3_finseq_5(k1_xboole_0))!=k4_rlvect_1(X1,k1_xboole_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_81]), c_0_84])])])).
cnf(c_0_94, plain, (m2_finseq_1(k4_finseq_5(X2,X1),X2)|~m1_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_85])).
cnf(c_0_95, plain, (k3_finseq_5(k1_xboole_0)=X1|~v1_xboole_0(X1)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_86, c_0_56]), c_0_57])).
cnf(c_0_96, plain, (v1_xboole_0(k1_xboole_0)), inference(rw,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_97, plain, (v1_relat_1(k1_xboole_0)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_73])])).
cnf(c_0_98, plain, (k3_finseq_5(k3_finseq_5(X2))=X2|~m1_finseq_1(k3_finseq_5(X2),X1)|~m1_finseq_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_90])])).
fof(c_0_99, negated_conjecture, ~(![X1]:(((((~(v2_struct_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&l2_algstr_0(X1))=>![X2]:(m2_finseq_1(X2,u1_struct_0(X1))=>k4_rlvect_1(X1,X2)=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))))), inference(assume_negation,[status(cth)],[t2_polynom3])).
cnf(c_0_100, plain, (k3_finseq_5(k4_finseq_5(X1,X2))=X2|~m1_finseq_1(X2,X1)|~m2_finseq_1(k4_finseq_5(X1,X2),X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92])])).
cnf(c_0_101, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),X2)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k3_finseq_5(k1_xboole_0))!=k4_rlvect_1(X1,k1_xboole_0)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94])])).
cnf(c_0_102, plain, (k3_finseq_5(k1_xboole_0)=k1_xboole_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_97])])])).
cnf(c_0_103, plain, (k3_finseq_5(k3_finseq_5(X1))=X1|~m1_finseq_1(X1,X2)|~m2_finseq_1(k3_finseq_5(X1),X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_92])])).
cnf(c_0_104, plain, (m2_finseq_1(k3_finseq_5(X2),X1)|~m1_finseq_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_81])])).
fof(c_0_105, negated_conjecture, (((((~v2_struct_0(esk1_0)&v2_rlvect_1(esk1_0))&v3_rlvect_1(esk1_0))&v4_rlvect_1(esk1_0))&l2_algstr_0(esk1_0))&(m2_finseq_1(esk2_0,u1_struct_0(esk1_0))&k4_rlvect_1(esk1_0,esk2_0)!=k4_rlvect_1(esk1_0,k4_finseq_5(u1_struct_0(esk1_0),esk2_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_99])])])])).
cnf(c_0_106, plain, (k3_finseq_5(k4_finseq_5(X1,X2))=X2|~m1_finseq_1(X2,X1)), inference(csr,[status(thm)],[c_0_100, c_0_94])).
cnf(c_0_107, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),X2)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_101, c_0_102])])])).
cnf(c_0_108, plain, (k3_finseq_5(k3_finseq_5(X1))=X1|~m1_finseq_1(X1,X2)), inference(csr,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_109, negated_conjecture, (m2_finseq_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_110, plain, (k4_finseq_5(X1,X3)=k4_finseq_5(X2,X3)|~m1_finseq_1(k4_finseq_5(X2,X3),X1)|~m1_finseq_1(X3,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_106])])).
fof(c_0_111, plain, ![X36, X37]:(((v1_relat_1(X37)|~m1_finseq_1(X37,X36))&(v1_funct_1(X37)|~m1_finseq_1(X37,X36)))&(v1_finseq_1(X37)|~m1_finseq_1(X37,X36))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_finseq_1])])])).
cnf(c_0_112, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),X2)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_92])])).
cnf(c_0_113, plain, (k3_finseq_5(k3_finseq_5(X1))=X1|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_92])])).
cnf(c_0_114, negated_conjecture, (m2_finseq_1(esk2_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_109])).
cnf(c_0_115, plain, (k4_finseq_5(X1,X2)=k4_finseq_5(X3,X2)|~m1_finseq_1(X2,X3)|~m2_finseq_1(k4_finseq_5(X3,X2),X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_92])])).
cnf(c_0_116, plain, (m2_finseq_1(X1,X2)|~m1_finseq_1(k3_finseq_5(X1),X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_75])])).
cnf(c_0_117, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_111])).
fof(c_0_118, plain, ![X29, X30]:(((v1_funct_1(X30)|~m2_finseq_1(X30,X29))&(v1_finseq_1(X30)|~m2_finseq_1(X30,X29)))&(m1_subset_1(X30,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X29)))|~m2_finseq_1(X30,X29))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m2_finseq_1])])])).
cnf(c_0_119, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),k3_finseq_5(X2))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k3_finseq_5(X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_104])])).
cnf(c_0_120, negated_conjecture, (k3_finseq_5(k3_finseq_5(esk2_0))=esk2_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114])])).
cnf(c_0_121, plain, (k4_finseq_5(X1,k4_finseq_5(X2,X3))=X3|~m1_finseq_1(k4_finseq_5(X2,X3),X2)|~m2_finseq_1(X3,X1)|~m1_finseq_1(X3,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_80])])).
cnf(c_0_122, plain, (m2_finseq_1(X1,X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m2_finseq_1(k3_finseq_5(X1),X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_92])])).
cnf(c_0_123, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_117])).
cnf(c_0_124, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_125, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_118])).
fof(c_0_126, plain, ![X43, X44, X45]:(~v2_rlvect_1(X43)|~l1_algstr_0(X43)|~m1_subset_1(X44,u1_struct_0(X43))|~m1_subset_1(X45,u1_struct_0(X43))|k3_rlvect_1(X43,X44,X45)=k1_algstr_0(X43,X44,X45)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k3_rlvect_1])])).
fof(c_0_127, plain, ![X33, X34, X35]:(~v2_rlvect_1(X33)|~l1_algstr_0(X33)|~m1_subset_1(X34,u1_struct_0(X33))|~m1_subset_1(X35,u1_struct_0(X33))|k3_rlvect_1(X33,X34,X35)=k3_rlvect_1(X33,X35,X34)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k3_rlvect_1])])).
fof(c_0_128, plain, ![X68, X69]:(v2_struct_0(X68)|~l2_algstr_0(X68)|~v1_relat_1(X69)|~v5_relat_1(X69,u1_struct_0(X68))|~v1_funct_1(X69)|~v1_finseq_1(X69)|m1_subset_1(k4_rlvect_1(X68,X69),u1_struct_0(X68))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k4_rlvect_1])])])).
fof(c_0_129, plain, ![X54, X55]:(~m1_finseq_1(X55,X54)|v5_relat_1(X55,X54)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_finseq_1])])).
cnf(c_0_130, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),esk2_0)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120])])).
cnf(c_0_131, plain, (k4_finseq_5(X1,k4_finseq_5(X2,X3))=X3|~m1_finseq_1(X3,X2)|~m2_finseq_1(X3,X1)|~m2_finseq_1(k4_finseq_5(X2,X3),X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_92])])).
cnf(c_0_132, plain, (k4_finseq_5(X1,X2)=k3_finseq_5(X2)|~m1_finseq_1(k3_finseq_5(X2),X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_75])])).
cnf(c_0_133, plain, (m2_finseq_1(k3_finseq_5(X1),X2)|~v1_finseq_1(k3_finseq_5(X1))|~v1_funct_1(k3_finseq_5(X1))|~v1_relat_1(k3_finseq_5(X1))|~m2_finseq_1(X1,X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_75])])).
cnf(c_0_134, plain, (v1_relat_1(X1)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_92])])).
cnf(c_0_135, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_124])).
cnf(c_0_136, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_125])).
cnf(c_0_137, plain, (k3_rlvect_1(X1,X2,X3)=k1_algstr_0(X1,X2,X3)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_126])).
cnf(c_0_138, plain, (k3_rlvect_1(X1,X2,X3)=k3_rlvect_1(X1,X3,X2)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_139, plain, (v2_struct_0(X1)|m1_subset_1(k4_rlvect_1(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v1_relat_1(X2)|~v5_relat_1(X2,u1_struct_0(X1))|~v1_funct_1(X2)|~v1_finseq_1(X2)), inference(split_conjunct,[status(thm)],[c_0_128])).
cnf(c_0_140, plain, (v5_relat_1(X1,X2)|~m1_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_129])).
fof(c_0_141, plain, ![X15, X16, X17]:(v2_struct_0(X15)|~v3_rlvect_1(X15)|~v4_rlvect_1(X15)|~l2_algstr_0(X15)|(~m2_finseq_1(X16,u1_struct_0(X15))|(~m2_finseq_1(X17,u1_struct_0(X15))|k4_rlvect_1(X15,k8_finseq_1(u1_struct_0(X15),X16,X17))=k1_algstr_0(X15,k4_rlvect_1(X15,X16),k4_rlvect_1(X15,X17))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t41_rlvect_1])])])])).
cnf(c_0_142, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),esk2_0)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))|k4_rlvect_1(X1,k3_finseq_5(esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_81])])).
cnf(c_0_143, plain, (k4_finseq_5(X1,k4_finseq_5(X2,X3))=X3|~m1_finseq_1(X3,X2)|~m2_finseq_1(X3,X1)), inference(csr,[status(thm)],[c_0_131, c_0_94])).
cnf(c_0_144, plain, (m2_finseq_1(X1,X2)|~m1_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_145, plain, (k4_finseq_5(X1,X2)=k3_finseq_5(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m2_finseq_1(k3_finseq_5(X2),X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_92])])).
cnf(c_0_146, plain, (m2_finseq_1(k3_finseq_5(X1),X2)|~v1_finseq_1(k3_finseq_5(X1))|~v1_funct_1(k3_finseq_5(X1))|~v1_relat_1(k3_finseq_5(X1))|~m2_finseq_1(X1,X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_133, c_0_134]), c_0_135]), c_0_136])).
fof(c_0_147, plain, ![X40]:(((v1_relat_1(k3_finseq_5(X40))|(~v1_relat_1(X40)|~v1_funct_1(X40)|~v1_finseq_1(X40)))&(v1_funct_1(k3_finseq_5(X40))|(~v1_relat_1(X40)|~v1_funct_1(X40)|~v1_finseq_1(X40))))&(v1_finseq_1(k3_finseq_5(X40))|(~v1_relat_1(X40)|~v1_funct_1(X40)|~v1_finseq_1(X40)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_finseq_5])])])).
cnf(c_0_148, plain, (k3_rlvect_1(X1,X2,X3)=k1_algstr_0(X1,X2,X3)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_137])).
cnf(c_0_149, plain, (k3_rlvect_1(X1,X2,X3)=k3_rlvect_1(X1,X3,X2)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_138])).
cnf(c_0_150, plain, (v2_struct_0(X1)|m1_subset_1(k4_rlvect_1(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)|~v5_relat_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_139])).
cnf(c_0_151, plain, (v5_relat_1(X1,X2)|~m1_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_140])).
cnf(c_0_152, plain, (v1_funct_1(X1)|~m1_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_153, plain, (v1_finseq_1(X1)|~m1_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_154, plain, (v2_struct_0(X1)|k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3))=k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,X3))|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_141])).
cnf(c_0_155, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k3_finseq_5(esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(esk2_0,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_143])])).
cnf(c_0_156, plain, (m2_finseq_1(X1,X2)|~m1_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_144])).
cnf(c_0_157, plain, (k4_finseq_5(X1,X2)=k3_finseq_5(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_finseq_1(k3_finseq_5(X2))|~v1_funct_1(k3_finseq_5(X2))|~v1_relat_1(k3_finseq_5(X2))|~m2_finseq_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_146])])).
cnf(c_0_158, plain, (v1_finseq_1(k3_finseq_5(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_147])).
cnf(c_0_159, plain, (k3_rlvect_1(X1,X3,X2)=k1_algstr_0(X1,X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_149])])).
cnf(c_0_160, plain, (m1_subset_1(k4_rlvect_1(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~l2_algstr_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_151])])).
cnf(c_0_161, plain, (v1_funct_1(X1)|~m1_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_152])).
cnf(c_0_162, plain, (v1_finseq_1(X1)|~m1_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_153])).
fof(c_0_163, plain, ![X32]:((l2_struct_0(X32)|~l2_algstr_0(X32))&(l1_algstr_0(X32)|~l2_algstr_0(X32))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l2_algstr_0])])])).
cnf(c_0_164, plain, (k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,X3))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3))|v2_struct_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m2_finseq_1(X3,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_154])).
cnf(c_0_165, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k3_finseq_5(esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(csr,[status(thm)],[c_0_155, c_0_156])).
cnf(c_0_166, plain, (k4_finseq_5(X1,X2)=k3_finseq_5(X2)|~v1_finseq_1(k3_finseq_5(X2))|~v1_funct_1(k3_finseq_5(X2))|~v1_relat_1(k3_finseq_5(X2))|~m2_finseq_1(X2,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_157, c_0_134]), c_0_135]), c_0_136])).
cnf(c_0_167, plain, (v1_finseq_1(k3_finseq_5(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_158])).
cnf(c_0_168, plain, (k1_algstr_0(X1,X3,X2)=k1_algstr_0(X1,X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_159])])).
cnf(c_0_169, plain, (m1_subset_1(k4_rlvect_1(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_160, c_0_123]), c_0_161]), c_0_162])).
cnf(c_0_170, plain, (l1_algstr_0(X1)|~l2_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_163])).
cnf(c_0_171, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,esk3_1(X1)))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k3_finseq_5(esk3_1(X1))))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m2_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_165])])).
cnf(c_0_172, plain, (k4_finseq_5(X1,X2)=k3_finseq_5(X2)|~v1_funct_1(k3_finseq_5(X2))|~v1_relat_1(k3_finseq_5(X2))|~m2_finseq_1(X2,X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_167])])).
cnf(c_0_173, plain, (v1_funct_1(k3_finseq_5(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_147])).
cnf(c_0_174, plain, (m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|k4_rlvect_1(X1,X2)=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k6_finseq_1(u1_struct_0(X1)))!=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k6_finseq_1(u1_struct_0(X1))))|~m2_finseq_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_175, plain, (k1_algstr_0(X1,k4_rlvect_1(X1,X3),k4_rlvect_1(X1,X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3))|v2_struct_0(X1)|~m2_finseq_1(X3,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_rlvect_1(X1,X3),u1_struct_0(X1))|~m1_subset_1(k4_rlvect_1(X1,X2),u1_struct_0(X1))|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_168])])).
cnf(c_0_176, plain, (m1_subset_1(k4_rlvect_1(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_92])])).
cnf(c_0_177, plain, (l1_algstr_0(X1)|~l2_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_170])).
cnf(c_0_178, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,esk3_1(X1)))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k3_finseq_5(esk3_1(X1))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(csr,[status(thm)],[c_0_171, c_0_104])).
cnf(c_0_179, plain, (k4_finseq_5(X1,X2)=k3_finseq_5(X2)|~v1_funct_1(k3_finseq_5(X2))|~v1_relat_1(k3_finseq_5(X2))|~m2_finseq_1(X2,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_172, c_0_134]), c_0_135]), c_0_136])).
cnf(c_0_180, plain, (v1_funct_1(k3_finseq_5(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_173])).
cnf(c_0_181, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k6_finseq_1(u1_struct_0(X1))))!=k4_rlvect_1(X1,k6_finseq_1(u1_struct_0(X1)))|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_174])).
cnf(c_0_182, plain, (k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,X3))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X3,X2))|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_175, c_0_176]), c_0_176]), c_0_177])).
cnf(c_0_183, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k3_finseq_5(esk3_1(X1))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_rlvect_1(X1,esk3_1(X1)),u1_struct_0(X1))|~m1_subset_1(k4_rlvect_1(X1,X2),u1_struct_0(X1))|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_178])])).
cnf(c_0_184, plain, (k4_finseq_5(X1,X2)=k3_finseq_5(X2)|~v1_relat_1(k3_finseq_5(X2))|~m2_finseq_1(X2,X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_180])])).
cnf(c_0_185, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k1_xboole_0))!=k4_rlvect_1(X1,k1_xboole_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_181, c_0_71]), c_0_71])])).
cnf(c_0_186, plain, (k4_finseq_5(X1,k1_xboole_0)=k1_xboole_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_102]), c_0_84])])])).
cnf(c_0_187, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X3,X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3))|v2_struct_0(X1)|~m2_finseq_1(X3,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_182])])).
cnf(c_0_188, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k3_finseq_5(esk3_1(X1))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_183, c_0_176]), c_0_169]), c_0_177])).
cnf(c_0_189, plain, (k4_finseq_5(X1,X2)=k3_finseq_5(X2)|~v1_relat_1(k3_finseq_5(X2))|~m2_finseq_1(X2,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_184, c_0_134]), c_0_135]), c_0_136])).
cnf(c_0_190, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_185, c_0_186])])])).
cnf(c_0_191, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k3_finseq_5(esk3_1(X1)),X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_188])])).
cnf(c_0_192, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),esk2_0)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))|k4_rlvect_1(X1,k3_finseq_5(esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_relat_1(k3_finseq_5(esk3_1(X1)))|~m2_finseq_1(esk3_1(X1),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_189])])).
cnf(c_0_193, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),X2)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))|m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_190, c_0_94])])).
cnf(c_0_194, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k3_finseq_5(esk3_1(X1)),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(csr,[status(thm)],[c_0_191, c_0_104])).
cnf(c_0_195, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k3_finseq_5(esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~v1_relat_1(k3_finseq_5(esk3_1(X1)))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(esk2_0,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_143])])).
cnf(c_0_196, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_80])])).
fof(c_0_197, plain, ![X22, X23, X24]:(~m1_finseq_1(X23,X22)|~m1_finseq_1(X24,X22)|k8_finseq_1(X22,X23,X24)=k7_finseq_1(X23,X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_finseq_1])])).
cnf(c_0_198, negated_conjecture, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k3_finseq_5(esk3_1(X1)),X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),esk3_1(X1),X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_194])])).
cnf(c_0_199, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k3_finseq_5(esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~v1_relat_1(k3_finseq_5(esk3_1(X1)))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_195, c_0_156]), c_0_196])).
cnf(c_0_200, plain, (k8_finseq_1(X2,X1,X3)=k7_finseq_1(X1,X3)|~m1_finseq_1(X1,X2)|~m1_finseq_1(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_197])).
cnf(c_0_201, negated_conjecture, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k3_finseq_5(esk3_1(X1)),X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),esk3_1(X1),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(csr,[status(thm)],[c_0_198, c_0_156])).
cnf(c_0_202, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k3_finseq_5(esk3_1(X1))))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m2_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_relat_1(k3_finseq_5(esk3_1(X1)))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_199])])).
cnf(c_0_203, plain, (k8_finseq_1(X2,X1,X3)=k7_finseq_1(X1,X3)|~m1_finseq_1(X3,X2)|~m1_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_200])).
cnf(c_0_204, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),esk3_1(X1),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_194, c_0_201])])).
cnf(c_0_205, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k3_finseq_5(esk3_1(X1))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(csr,[status(thm)],[c_0_202, c_0_134])).
cnf(c_0_206, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X3,X2))=k4_rlvect_1(X1,k7_finseq_1(X2,X3))|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_203])])).
cnf(c_0_207, negated_conjecture, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k3_finseq_5(esk3_1(X1))))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),esk3_1(X1),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_205])])).
cnf(c_0_208, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3))=k4_rlvect_1(X1,k7_finseq_1(X3,X2))|v2_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~m1_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_206, c_0_156]), c_0_156])).
cnf(c_0_209, negated_conjecture, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k3_finseq_5(esk3_1(X1))))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),esk3_1(X1),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(csr,[status(thm)],[c_0_207, c_0_104])).
cnf(c_0_210, plain, (k4_rlvect_1(X1,X2)=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))))!=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1)))))|k4_rlvect_1(X1,k6_finseq_1(u1_struct_0(X1)))!=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k6_finseq_1(u1_struct_0(X1))))|~m2_finseq_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_211, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3))=k4_rlvect_1(X1,k7_finseq_1(X2,X3))|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_208])])).
cnf(c_0_212, negated_conjecture, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),esk3_1(X1),X2))=k4_rlvect_1(X1,k7_finseq_1(X2,k3_finseq_5(esk3_1(X1))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_203])])).
fof(c_0_213, plain, ![X13, X14]:(~v1_relat_1(X13)|~v1_funct_1(X13)|~v1_finseq_1(X13)|k3_finseq_5(k7_finseq_1(X13,k9_finseq_1(X14)))=k7_finseq_1(k9_finseq_1(X14),k3_finseq_5(X13))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_finseq_5])])])).
cnf(c_0_214, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k6_finseq_1(u1_struct_0(X1))))!=k4_rlvect_1(X1,k6_finseq_1(u1_struct_0(X1)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1)))))!=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))))|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_210])).
cnf(c_0_215, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3))=k4_rlvect_1(X1,k7_finseq_1(X2,X3))|v2_struct_0(X1)|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_211, c_0_156]), c_0_156])).
cnf(c_0_216, negated_conjecture, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),esk3_1(X1),X2))=k4_rlvect_1(X1,k7_finseq_1(X2,k3_finseq_5(esk3_1(X1))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(csr,[status(thm)],[c_0_212, c_0_156])).
cnf(c_0_217, plain, (k3_finseq_5(k7_finseq_1(X1,k9_finseq_1(X2)))=k7_finseq_1(k9_finseq_1(X2),k3_finseq_5(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_213])).
cnf(c_0_218, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1)))))!=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k1_xboole_0))!=k4_rlvect_1(X1,k1_xboole_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_214, c_0_71]), c_0_71])])).
cnf(c_0_219, negated_conjecture, (k4_rlvect_1(X1,k7_finseq_1(X2,k3_finseq_5(esk3_1(X1))))=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m1_finseq_1(X2,u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_216])])).
cnf(c_0_220, plain, (k7_finseq_1(k9_finseq_1(X2),k3_finseq_5(X1))=k3_finseq_5(k7_finseq_1(X1,k9_finseq_1(X2)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_217])).
cnf(c_0_221, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1)))))!=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k1_xboole_0))!=k4_rlvect_1(X1,k1_xboole_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_81])])).
cnf(c_0_222, negated_conjecture, (k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(esk3_1(X1),k9_finseq_1(X2))))=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),k9_finseq_1(X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m1_finseq_1(k9_finseq_1(X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_finseq_1(esk3_1(X1))|~v1_funct_1(esk3_1(X1))|~v1_relat_1(esk3_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_220])])).
fof(c_0_223, plain, ![X89, X90, X91]:(~m1_finseq_1(X90,X89)|~m1_finseq_1(X91,X89)|m2_finseq_1(k8_finseq_1(X89,X90,X91),X89)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_finseq_1])])).
fof(c_0_224, plain, ![X57, X58, X59]:((~m2_finseq_2(X59,X57,X58)|m1_subset_1(X59,X58)|~m1_finseq_2(X58,X57))&(~m1_subset_1(X59,X58)|m2_finseq_2(X59,X57,X58)|~m1_finseq_2(X58,X57))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_m2_finseq_2])])])])).
fof(c_0_225, plain, ![X70]:m1_finseq_2(k3_finseq_2(X70),X70), inference(variable_rename,[status(thm)],[dt_k3_finseq_2])).
fof(c_0_226, plain, ![X47, X48]:(v1_xboole_0(X47)|~m1_subset_1(X48,X47)|m2_finseq_2(k3_pre_poly(X47,X48),X47,k3_finseq_2(X47))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_pre_poly])])])).
fof(c_0_227, plain, ![X27, X28]:(v1_xboole_0(X27)|~m1_subset_1(X28,X27)|k3_pre_poly(X27,X28)=k5_finseq_1(X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_pre_poly])])])).
fof(c_0_228, plain, ![X20]:k9_finseq_1(X20)=k5_finseq_1(X20), inference(variable_rename,[status(thm)],[redefinition_k9_finseq_1])).
cnf(c_0_229, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1)))))!=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))))|~m1_finseq_1(k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_221, c_0_186])])])).
cnf(c_0_230, negated_conjecture, (k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(esk3_1(X1),k9_finseq_1(X2))))=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),k9_finseq_1(X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(k9_finseq_1(X2),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_222, c_0_123]), c_0_161]), c_0_162])).
cnf(c_0_231, plain, (m2_finseq_1(k8_finseq_1(X2,X1,X3),X2)|~m1_finseq_1(X1,X2)|~m1_finseq_1(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_223])).
fof(c_0_232, plain, ![X77, X78, X79]:(~m1_finseq_2(X78,X77)|(~m2_finseq_2(X79,X77,X78)|m2_finseq_1(X79,X77))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m2_finseq_2])])])).
cnf(c_0_233, plain, (m2_finseq_2(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_finseq_2(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_224])).
cnf(c_0_234, plain, (m1_finseq_2(k3_finseq_2(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_225])).
cnf(c_0_235, plain, (m1_subset_1(X1,X3)|~m2_finseq_2(X1,X2,X3)|~m1_finseq_2(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_224])).
cnf(c_0_236, plain, (v1_xboole_0(X1)|m2_finseq_2(k3_pre_poly(X1,X2),X1,k3_finseq_2(X1))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_226])).
cnf(c_0_237, plain, (v1_xboole_0(X1)|k3_pre_poly(X1,X2)=k5_finseq_1(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_227])).
cnf(c_0_238, plain, (k9_finseq_1(X1)=k5_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_228])).
cnf(c_0_239, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m1_finseq_1(k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(k9_finseq_1(esk4_1(X1)),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_230])])).
cnf(c_0_240, plain, (m2_finseq_1(k8_finseq_1(X2,X1,X3),X2)|~m1_finseq_1(X3,X2)|~m1_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_231])).
cnf(c_0_241, plain, (m2_finseq_1(X3,X2)|~m1_finseq_2(X1,X2)|~m2_finseq_2(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_232])).
cnf(c_0_242, plain, (m2_finseq_2(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_finseq_2(X2,X3)), inference(evalgc,[status(thm)],[c_0_233])).
cnf(c_0_243, plain, (m1_finseq_2(k3_finseq_2(X1),X1)), inference(evalgc,[status(thm)],[c_0_234])).
cnf(c_0_244, plain, (m1_subset_1(X1,X3)|~m1_finseq_2(X3,X2)|~m2_finseq_2(X1,X2,X3)), inference(evalgc,[status(thm)],[c_0_235])).
cnf(c_0_245, plain, (v1_xboole_0(X1)|m2_finseq_2(k3_pre_poly(X1,X2),X1,k3_finseq_2(X1))|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_236])).
cnf(c_0_246, plain, (k3_pre_poly(X1,X2)=k9_finseq_1(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_237, c_0_238])).
cnf(c_0_247, plain, (m1_subset_1(esk4_1(X1),u1_struct_0(X1))|k4_rlvect_1(X1,X2)=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k6_finseq_1(u1_struct_0(X1)))!=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k6_finseq_1(u1_struct_0(X1))))|~m2_finseq_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
fof(c_0_248, plain, ![X60]:(~l2_struct_0(X60)|l1_struct_0(X60)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l2_struct_0])])).
cnf(c_0_249, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k9_finseq_1(esk4_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_92])])).
cnf(c_0_250, plain, (m2_finseq_1(k7_finseq_1(X2,X3),X1)|~m1_finseq_1(X3,X1)|~m1_finseq_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_203])])).
cnf(c_0_251, plain, (m2_finseq_1(X3,X2)|~m1_finseq_2(X1,X2)|~m2_finseq_2(X3,X2,X1)), inference(evalgc,[status(thm)],[c_0_241])).
cnf(c_0_252, plain, (m2_finseq_2(X1,X2,k3_finseq_2(X2))|~m1_subset_1(X1,k3_finseq_2(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_243])])).
cnf(c_0_253, plain, (m1_subset_1(k3_pre_poly(X1,X2),k3_finseq_2(X1))|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_244, c_0_245]), c_0_243])])])).
cnf(c_0_254, plain, (k3_pre_poly(X1,X2)=k9_finseq_1(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_246])).
cnf(c_0_255, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|m1_subset_1(esk4_1(X1),u1_struct_0(X1))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k6_finseq_1(u1_struct_0(X1))))!=k4_rlvect_1(X1,k6_finseq_1(u1_struct_0(X1)))|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_247])).
fof(c_0_256, plain, ![X118]:(~v2_struct_0(X118)|~l1_struct_0(X118)|v1_xboole_0(u1_struct_0(X118))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_struct_0])])).
cnf(c_0_257, plain, (l1_struct_0(X1)|~l2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_248])).
cnf(c_0_258, plain, (l2_struct_0(X1)|~l2_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_163])).
cnf(c_0_259, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k9_finseq_1(esk4_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(csr,[status(thm)],[c_0_249, c_0_250])).
cnf(c_0_260, plain, (m2_finseq_1(X1,X2)|~m1_subset_1(X1,k3_finseq_2(X2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_252]), c_0_243])])])).
cnf(c_0_261, plain, (v1_xboole_0(X1)|m1_subset_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_254])])).
cnf(c_0_262, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|m1_subset_1(esk4_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k1_xboole_0))!=k4_rlvect_1(X1,k1_xboole_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_255, c_0_71]), c_0_71])])).
cnf(c_0_263, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_256])).
cnf(c_0_264, plain, (l1_struct_0(X1)|~l2_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_257])).
cnf(c_0_265, plain, (l2_struct_0(X1)|~l2_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_258])).
cnf(c_0_266, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(k9_finseq_1(esk4_1(X1)),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_259, c_0_92])])).
cnf(c_0_267, plain, (m2_finseq_1(k9_finseq_1(X1),X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_261])])).
cnf(c_0_268, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|m1_subset_1(esk4_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|k4_rlvect_1(X1,k3_finseq_5(k1_xboole_0))!=k4_rlvect_1(X1,k1_xboole_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_81]), c_0_84])])])).
cnf(c_0_269, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_263])).
cnf(c_0_270, plain, (l1_struct_0(X1)|~l2_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_265])])).
cnf(c_0_271, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(esk4_1(X1),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_267])])).
cnf(c_0_272, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|m1_subset_1(esk4_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_268, c_0_102])])])).
cnf(c_0_273, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l2_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_269, c_0_270])])).
fof(c_0_274, plain, ![X52]:(v2_struct_0(X52)|~l1_struct_0(X52)|~v1_xboole_0(u1_struct_0(X52))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_275, plain, ![X46]:(~l1_algstr_0(X46)|l1_struct_0(X46)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_algstr_0])])).
cnf(c_0_276, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_271, c_0_272]), c_0_273])).
cnf(c_0_277, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_274])).
cnf(c_0_278, plain, (l1_struct_0(X1)|~l1_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_275])).
cnf(c_0_279, negated_conjecture, (l2_algstr_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_280, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_276, c_0_92])])).
cnf(c_0_281, negated_conjecture, (v4_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_282, negated_conjecture, (v3_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_283, negated_conjecture, (v2_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_284, negated_conjecture, (k4_rlvect_1(esk1_0,esk2_0)!=k4_rlvect_1(esk1_0,k4_finseq_5(u1_struct_0(esk1_0),esk2_0))), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_285, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_286, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_277])).
cnf(c_0_287, plain, (l1_struct_0(X1)|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_278])).
cnf(c_0_288, negated_conjecture, (l2_algstr_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_279])).
cnf(c_0_289, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(csr,[status(thm)],[c_0_280, c_0_104])).
cnf(c_0_290, negated_conjecture, (v4_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_281])).
cnf(c_0_291, negated_conjecture, (v3_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_282])).
cnf(c_0_292, negated_conjecture, (v2_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_283])).
cnf(c_0_293, negated_conjecture, (k4_rlvect_1(esk1_0,k4_finseq_5(u1_struct_0(esk1_0),esk2_0))!=k4_rlvect_1(esk1_0,esk2_0)), inference(evalgc,[status(thm)],[c_0_284])).
cnf(c_0_294, negated_conjecture, (~v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_285])).
cnf(c_0_295, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_287])])).
cnf(c_0_296, negated_conjecture, (l1_algstr_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_288])])).
cnf(c_0_297, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(esk3_1(X1),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_289, c_0_92])])).
cnf(c_0_298, negated_conjecture, (m2_finseq_1(esk3_1(esk1_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_190, c_0_114]), c_0_288]), c_0_290]), c_0_291]), c_0_292])]), c_0_293]), c_0_294])])).
cnf(c_0_299, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_295, c_0_296]), c_0_294])])).
cnf(c_0_300, negated_conjecture, (k4_rlvect_1(esk1_0,k4_finseq_5(u1_struct_0(esk1_0),X1))=k4_rlvect_1(esk1_0,X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(esk1_0))|~m1_finseq_1(esk2_0,u1_struct_0(esk1_0))|~m2_finseq_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_297, c_0_298]), c_0_288]), c_0_290]), c_0_291]), c_0_292])]), c_0_293]), c_0_299])])).
cnf(c_0_301, negated_conjecture, (k4_rlvect_1(esk1_0,k4_finseq_5(u1_struct_0(esk1_0),X1))=k4_rlvect_1(esk1_0,X1)|~m1_finseq_1(esk2_0,u1_struct_0(esk1_0))|~m2_finseq_1(X1,u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_300, c_0_92])])).
cnf(c_0_302, negated_conjecture, (k4_rlvect_1(esk1_0,k4_finseq_5(u1_struct_0(esk1_0),X1))=k4_rlvect_1(esk1_0,X1)|~m1_finseq_1(esk2_0,u1_struct_0(esk1_0))|~m2_finseq_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[c_0_301, c_0_104])).
cnf(c_0_303, negated_conjecture, (~m1_finseq_1(esk2_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_293, c_0_302]), c_0_114])])])).
cnf(c_0_304, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_303, c_0_92]), c_0_114])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 305
# Proof object clause steps            : 230
# Proof object formula steps           : 75
# Proof object conjectures             : 60
# Proof object clause conjectures      : 57
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 56
# Proof object initial formulas used   : 38
# Proof object generating inferences   : 80
# Proof object simplifying inferences  : 101
# Training examples: 143 positive, 1080 negative
# Training: Positive examples begin
cnf(c_0_88, plain, (o_0_0_xboole_0=k1_xboole_0)).# trainpos
cnf(c_0_292, negated_conjecture, (v2_rlvect_1(esk1_0))).# trainpos
cnf(c_0_291, negated_conjecture, (v3_rlvect_1(esk1_0))).# trainpos
cnf(c_0_290, negated_conjecture, (v4_rlvect_1(esk1_0))).# trainpos
cnf(c_0_288, negated_conjecture, (l2_algstr_0(esk1_0))).# trainpos
cnf(c_0_87, plain, (v1_xboole_0(o_0_0_xboole_0))).# trainpos
cnf(c_0_294, negated_conjecture, (~v2_struct_0(esk1_0))).# trainpos
cnf(c_0_293, negated_conjecture, (k4_rlvect_1(esk1_0,k4_finseq_5(u1_struct_0(esk1_0),esk2_0))!=k4_rlvect_1(esk1_0,esk2_0))).# trainpos
cnf(c_0_265, plain, (l2_struct_0(X1)|~l2_algstr_0(X1))).# trainpos
cnf(c_0_287, plain, (l1_struct_0(X1)|~l1_algstr_0(X1))).# trainpos
cnf(c_0_114, negated_conjecture, (m2_finseq_1(esk2_0,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_89, plain, (v1_relat_1(esk6_1(X1)))).# trainpos
cnf(c_0_61, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_62, plain, (v1_xboole_0(k6_finseq_1(X1)))).# trainpos
cnf(c_0_64, plain, (v1_xboole_0(esk6_1(X1)))).# trainpos
cnf(c_0_156, plain, (m2_finseq_1(X1,X2)|~m1_finseq_1(X1,X2))).# trainpos
cnf(c_0_177, plain, (l1_algstr_0(X1)|~l2_algstr_0(X1))).# trainpos
cnf(c_0_92, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_286, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).# trainpos
cnf(c_0_56, plain, (v1_funct_1(X1)|~v1_xboole_0(X1))).# trainpos
cnf(c_0_264, plain, (l1_struct_0(X1)|~l2_struct_0(X1))).# trainpos
cnf(c_0_72, plain, (m1_finseq_1(esk6_1(X1),X1))).# trainpos
cnf(c_0_243, plain, (m1_finseq_2(k3_finseq_2(X1),X1))).# trainpos
cnf(c_0_135, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_136, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_269, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1))).# trainpos
cnf(c_0_57, plain, (v1_finseq_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).# trainpos
cnf(c_0_123, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2))).# trainpos
cnf(c_0_161, plain, (v1_funct_1(X1)|~m1_finseq_1(X1,X2))).# trainpos
cnf(c_0_162, plain, (v1_finseq_1(X1)|~m1_finseq_1(X1,X2))).# trainpos
cnf(c_0_151, plain, (v5_relat_1(X1,X2)|~m1_finseq_1(X1,X2))).# trainpos
cnf(c_0_81, plain, (k4_finseq_5(X2,X1)=k3_finseq_5(X1)|~m1_finseq_1(X1,X2))).# trainpos
cnf(c_0_180, plain, (v1_funct_1(k3_finseq_5(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_94, plain, (m2_finseq_1(k4_finseq_5(X2,X1),X2)|~m1_finseq_1(X1,X2))).# trainpos
cnf(c_0_167, plain, (v1_finseq_1(k3_finseq_5(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_80, plain, (k4_finseq_5(X2,k4_finseq_5(X2,X1))=X1|~m1_finseq_1(X1,X2))).# trainpos
cnf(c_0_251, plain, (m2_finseq_1(X3,X2)|~m1_finseq_2(X1,X2)|~m2_finseq_2(X3,X2,X1))).# trainpos
cnf(c_0_75, plain, (k3_finseq_5(k3_finseq_5(X1))=X1|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_242, plain, (m2_finseq_2(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_finseq_2(X2,X3))).# trainpos
cnf(c_0_244, plain, (m1_subset_1(X1,X3)|~m1_finseq_2(X3,X2)|~m2_finseq_2(X1,X2,X3))).# trainpos
cnf(c_0_254, plain, (k3_pre_poly(X1,X2)=k9_finseq_1(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_55, plain, (v1_xboole_0(k3_finseq_5(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)|~v1_xboole_0(X1))).# trainpos
cnf(c_0_150, plain, (v2_struct_0(X1)|m1_subset_1(k4_rlvect_1(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)|~v5_relat_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_203, plain, (k8_finseq_1(X2,X1,X3)=k7_finseq_1(X1,X3)|~m1_finseq_1(X3,X2)|~m1_finseq_1(X1,X2))).# trainpos
cnf(c_0_245, plain, (v1_xboole_0(X1)|m2_finseq_2(k3_pre_poly(X1,X2),X1,k3_finseq_2(X1))|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_240, plain, (m2_finseq_1(k8_finseq_1(X2,X1,X3),X2)|~m1_finseq_1(X3,X2)|~m1_finseq_1(X1,X2))).# trainpos
cnf(c_0_220, plain, (k7_finseq_1(k9_finseq_1(X2),k3_finseq_5(X1))=k3_finseq_5(k7_finseq_1(X1,k9_finseq_1(X2)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_148, plain, (k3_rlvect_1(X1,X2,X3)=k1_algstr_0(X1,X2,X3)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_149, plain, (k3_rlvect_1(X1,X2,X3)=k3_rlvect_1(X1,X3,X2)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_214, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k6_finseq_1(u1_struct_0(X1))))!=k4_rlvect_1(X1,k6_finseq_1(u1_struct_0(X1)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1)))))!=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))))|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_164, plain, (k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,X3))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3))|v2_struct_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m2_finseq_1(X3,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_70, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k6_finseq_1(u1_struct_0(X1))))!=k4_rlvect_1(X1,k6_finseq_1(u1_struct_0(X1)))|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_181, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k6_finseq_1(u1_struct_0(X1))))!=k4_rlvect_1(X1,k6_finseq_1(u1_struct_0(X1)))|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_255, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|m1_subset_1(esk4_1(X1),u1_struct_0(X1))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k6_finseq_1(u1_struct_0(X1))))!=k4_rlvect_1(X1,k6_finseq_1(u1_struct_0(X1)))|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_296, negated_conjecture, (l1_algstr_0(esk1_0))).# trainpos
cnf(c_0_295, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(X1))).# trainpos
cnf(c_0_270, plain, (l1_struct_0(X1)|~l2_algstr_0(X1))).# trainpos
cnf(c_0_299, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_71, plain, (k6_finseq_1(X1)=k1_xboole_0)).# trainpos
cnf(c_0_73, plain, (esk6_1(X1)=k1_xboole_0)).# trainpos
cnf(c_0_159, plain, (k3_rlvect_1(X1,X3,X2)=k1_algstr_0(X1,X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_134, plain, (v1_relat_1(X1)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_253, plain, (m1_subset_1(k3_pre_poly(X1,X2),k3_finseq_2(X1))|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_250, plain, (m2_finseq_1(k7_finseq_1(X2,X3),X1)|~m1_finseq_1(X3,X1)|~m1_finseq_1(X2,X1))).# trainpos
cnf(c_0_160, plain, (m1_subset_1(k4_rlvect_1(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~l2_algstr_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_76, plain, (k3_finseq_5(X1)=k1_xboole_0|~v1_xboole_0(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_176, plain, (m1_subset_1(k4_rlvect_1(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_273, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l2_algstr_0(X1))).# trainpos
cnf(c_0_218, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1)))))!=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k1_xboole_0))!=k4_rlvect_1(X1,k1_xboole_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_84, plain, (m1_finseq_1(k1_xboole_0,X1))).# trainpos
cnf(c_0_168, plain, (k1_algstr_0(X1,X3,X2)=k1_algstr_0(X1,X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_83, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k1_xboole_0))!=k4_rlvect_1(X1,k1_xboole_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_97, plain, (v1_relat_1(k1_xboole_0))).# trainpos
cnf(c_0_262, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|m1_subset_1(esk4_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k1_xboole_0))!=k4_rlvect_1(X1,k1_xboole_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_185, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k1_xboole_0))!=k4_rlvect_1(X1,k1_xboole_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_91, plain, (k3_finseq_5(k4_finseq_5(X1,X2))=X2|~m1_finseq_1(k4_finseq_5(X1,X2),X1)|~m1_finseq_1(X2,X1))).# trainpos
cnf(c_0_175, plain, (k1_algstr_0(X1,k4_rlvect_1(X1,X3),k4_rlvect_1(X1,X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3))|v2_struct_0(X1)|~m2_finseq_1(X3,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_rlvect_1(X1,X3),u1_struct_0(X1))|~m1_subset_1(k4_rlvect_1(X1,X2),u1_struct_0(X1))|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_221, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1)))))!=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k1_xboole_0))!=k4_rlvect_1(X1,k1_xboole_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))),u1_struct_0(X1)))).# trainpos
cnf(c_0_100, plain, (k3_finseq_5(k4_finseq_5(X1,X2))=X2|~m1_finseq_1(X2,X1)|~m2_finseq_1(k4_finseq_5(X1,X2),X1))).# trainpos
cnf(c_0_268, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|m1_subset_1(esk4_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|k4_rlvect_1(X1,k3_finseq_5(k1_xboole_0))!=k4_rlvect_1(X1,k1_xboole_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_93, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|k4_rlvect_1(X1,k3_finseq_5(k1_xboole_0))!=k4_rlvect_1(X1,k1_xboole_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_252, plain, (m2_finseq_2(X1,X2,k3_finseq_2(X2))|~m1_subset_1(X1,k3_finseq_2(X2)))).# trainpos
cnf(c_0_261, plain, (v1_xboole_0(X1)|m1_subset_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_187, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X3,X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3))|v2_struct_0(X1)|~m2_finseq_1(X3,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_260, plain, (m2_finseq_1(X1,X2)|~m1_subset_1(X1,k3_finseq_2(X2)))).# trainpos
cnf(c_0_267, plain, (m2_finseq_1(k9_finseq_1(X1),X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_101, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),X2)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k3_finseq_5(k1_xboole_0))!=k4_rlvect_1(X1,k1_xboole_0)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_86, plain, (k3_finseq_5(k1_xboole_0)=X1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).# trainpos
cnf(c_0_102, plain, (k3_finseq_5(k1_xboole_0)=k1_xboole_0)).# trainpos
cnf(c_0_206, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X3,X2))=k4_rlvect_1(X1,k7_finseq_1(X2,X3))|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_90, plain, (k4_finseq_5(X1,k3_finseq_5(X2))=X2|~m1_finseq_1(X2,X1))).# trainpos
cnf(c_0_186, plain, (k4_finseq_5(X1,k1_xboole_0)=k1_xboole_0)).# trainpos
cnf(c_0_107, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),X2)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_272, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|m1_subset_1(esk4_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_132, plain, (k4_finseq_5(X1,X2)=k3_finseq_5(X2)|~m1_finseq_1(k3_finseq_5(X2),X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).# trainpos
cnf(c_0_104, plain, (m2_finseq_1(k3_finseq_5(X2),X1)|~m1_finseq_1(X2,X1))).# trainpos
cnf(c_0_98, plain, (k3_finseq_5(k3_finseq_5(X2))=X2|~m1_finseq_1(k3_finseq_5(X2),X1)|~m1_finseq_1(X2,X1))).# trainpos
cnf(c_0_211, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3))=k4_rlvect_1(X1,k7_finseq_1(X2,X3))|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_110, plain, (k4_finseq_5(X1,X3)=k4_finseq_5(X2,X3)|~m1_finseq_1(k4_finseq_5(X2,X3),X1)|~m1_finseq_1(X3,X2))).# trainpos
cnf(c_0_116, plain, (m2_finseq_1(X1,X2)|~m1_finseq_1(k3_finseq_5(X1),X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_229, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1)))))!=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))))|~m1_finseq_1(k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_190, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_112, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),X2)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_103, plain, (k3_finseq_5(k3_finseq_5(X1))=X1|~m1_finseq_1(X1,X2)|~m2_finseq_1(k3_finseq_5(X1),X2))).# trainpos
cnf(c_0_145, plain, (k4_finseq_5(X1,X2)=k3_finseq_5(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m2_finseq_1(k3_finseq_5(X2),X1))).# trainpos
cnf(c_0_115, plain, (k4_finseq_5(X1,X2)=k4_finseq_5(X3,X2)|~m1_finseq_1(X2,X3)|~m2_finseq_1(k4_finseq_5(X3,X2),X1))).# trainpos
cnf(c_0_113, plain, (k3_finseq_5(k3_finseq_5(X1))=X1|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_122, plain, (m2_finseq_1(X1,X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m2_finseq_1(k3_finseq_5(X1),X2))).# trainpos
cnf(c_0_120, negated_conjecture, (k3_finseq_5(k3_finseq_5(esk2_0))=esk2_0)).# trainpos
cnf(c_0_133, plain, (m2_finseq_1(k3_finseq_5(X1),X2)|~v1_finseq_1(k3_finseq_5(X1))|~v1_funct_1(k3_finseq_5(X1))|~v1_relat_1(k3_finseq_5(X1))|~m2_finseq_1(X1,X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_157, plain, (k4_finseq_5(X1,X2)=k3_finseq_5(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_finseq_1(k3_finseq_5(X2))|~v1_funct_1(k3_finseq_5(X2))|~v1_relat_1(k3_finseq_5(X2))|~m2_finseq_1(X2,X1))).# trainpos
cnf(c_0_121, plain, (k4_finseq_5(X1,k4_finseq_5(X2,X3))=X3|~m1_finseq_1(k4_finseq_5(X2,X3),X2)|~m2_finseq_1(X3,X1)|~m1_finseq_1(X3,X2))).# trainpos
cnf(c_0_193, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),X2)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))|m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_196, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_172, plain, (k4_finseq_5(X1,X2)=k3_finseq_5(X2)|~v1_funct_1(k3_finseq_5(X2))|~v1_relat_1(k3_finseq_5(X2))|~m2_finseq_1(X2,X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).# trainpos
cnf(c_0_131, plain, (k4_finseq_5(X1,k4_finseq_5(X2,X3))=X3|~m1_finseq_1(X3,X2)|~m2_finseq_1(X3,X1)|~m2_finseq_1(k4_finseq_5(X2,X3),X2))).# trainpos
cnf(c_0_184, plain, (k4_finseq_5(X1,X2)=k3_finseq_5(X2)|~v1_relat_1(k3_finseq_5(X2))|~m2_finseq_1(X2,X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).# trainpos
cnf(c_0_298, negated_conjecture, (m2_finseq_1(esk3_1(esk1_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_119, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),k3_finseq_5(X2))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k3_finseq_5(X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_130, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),esk2_0)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_192, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),esk2_0)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))|k4_rlvect_1(X1,k3_finseq_5(esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_relat_1(k3_finseq_5(esk3_1(X1)))|~m2_finseq_1(esk3_1(X1),u1_struct_0(X1)))).# trainpos
cnf(c_0_142, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),esk2_0)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))|k4_rlvect_1(X1,k3_finseq_5(esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1)))).# trainpos
cnf(c_0_155, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k3_finseq_5(esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(esk2_0,u1_struct_0(X1)))).# trainpos
cnf(c_0_195, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k3_finseq_5(esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~v1_relat_1(k3_finseq_5(esk3_1(X1)))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(esk2_0,u1_struct_0(X1)))).# trainpos
cnf(c_0_171, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,esk3_1(X1)))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k3_finseq_5(esk3_1(X1))))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m2_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_183, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k3_finseq_5(esk3_1(X1))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_rlvect_1(X1,esk3_1(X1)),u1_struct_0(X1))|~m1_subset_1(k4_rlvect_1(X1,X2),u1_struct_0(X1))|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1))).# trainpos
cnf(c_0_191, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k3_finseq_5(esk3_1(X1)),X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1)))).# trainpos
cnf(c_0_198, negated_conjecture, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k3_finseq_5(esk3_1(X1)),X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),esk3_1(X1),X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_204, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),esk3_1(X1),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_202, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k3_finseq_5(esk3_1(X1))))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m2_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_relat_1(k3_finseq_5(esk3_1(X1)))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1)))).# trainpos
cnf(c_0_207, negated_conjecture, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k3_finseq_5(esk3_1(X1))))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),esk3_1(X1),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1)))).# trainpos
cnf(c_0_212, negated_conjecture, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),esk3_1(X1),X2))=k4_rlvect_1(X1,k7_finseq_1(X2,k3_finseq_5(esk3_1(X1))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_219, negated_conjecture, (k4_rlvect_1(X1,k7_finseq_1(X2,k3_finseq_5(esk3_1(X1))))=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m1_finseq_1(X2,u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1)))).# trainpos
cnf(c_0_222, negated_conjecture, (k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(esk3_1(X1),k9_finseq_1(X2))))=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),k9_finseq_1(X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m1_finseq_1(k9_finseq_1(X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_finseq_1(esk3_1(X1))|~v1_funct_1(esk3_1(X1))|~v1_relat_1(esk3_1(X1)))).# trainpos
cnf(c_0_239, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m1_finseq_1(k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(k9_finseq_1(esk4_1(X1)),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1)))).# trainpos
cnf(c_0_249, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k9_finseq_1(esk4_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))),u1_struct_0(X1)))).# trainpos
cnf(c_0_266, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(k9_finseq_1(esk4_1(X1)),u1_struct_0(X1)))).# trainpos
cnf(c_0_271, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(esk4_1(X1),u1_struct_0(X1)))).# trainpos
cnf(c_0_280, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1)))).# trainpos
cnf(c_0_297, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(esk3_1(X1),u1_struct_0(X1)))).# trainpos
cnf(c_0_300, negated_conjecture, (k4_rlvect_1(esk1_0,k4_finseq_5(u1_struct_0(esk1_0),X1))=k4_rlvect_1(esk1_0,X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(esk1_0))|~m1_finseq_1(esk2_0,u1_struct_0(esk1_0))|~m2_finseq_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_301, negated_conjecture, (k4_rlvect_1(esk1_0,k4_finseq_5(u1_struct_0(esk1_0),X1))=k4_rlvect_1(esk1_0,X1)|~m1_finseq_1(esk2_0,u1_struct_0(esk1_0))|~m2_finseq_1(X1,u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_303, negated_conjecture, (~m1_finseq_1(esk2_0,u1_struct_0(esk1_0)))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_65, plain, (k4_ordinal1=k5_numbers)).#trainneg
cnf(i_0_295, plain, (v1_relat_1(esk21_0))).#trainneg
cnf(i_0_292, plain, (v1_funct_1(esk21_0))).#trainneg
cnf(i_0_153, plain, (v1_xboole_0(esk7_0))).#trainneg
cnf(i_0_159, plain, (v1_xboole_0(esk8_0))).#trainneg
cnf(i_0_277, plain, (v7_ordinal1(esk20_0))).#trainneg
cnf(i_0_70, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_71, plain, (~v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_205, plain, (v3_ordinal1(esk13_0))).#trainneg
cnf(i_0_290, plain, (v1_funct_2(esk21_0,k5_numbers,k5_numbers))).#trainneg
cnf(i_0_137, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_206, plain, (~v1_xboole_0(esk13_0))).#trainneg
cnf(i_0_278, plain, (v3_ordinal1(esk20_0))).#trainneg
cnf(i_0_273, plain, (v1_finset_1(esk20_0))).#trainneg
cnf(i_0_110, plain, (v3_membered(k1_numbers))).#trainneg
cnf(i_0_116, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_272, plain, (v1_card_1(esk20_0))).#trainneg
cnf(i_0_127, plain, (v6_membered(k4_ordinal1))).#trainneg
cnf(i_0_152, plain, (v1_xxreal_0(esk7_0))).#trainneg
cnf(i_0_157, plain, (v1_xxreal_0(esk8_0))).#trainneg
cnf(i_0_275, plain, (v1_xxreal_0(esk20_0))).#trainneg
cnf(i_0_326, plain, (v1_xxreal_0(esk22_0))).#trainneg
cnf(i_0_156, plain, (v1_xreal_0(esk8_0))).#trainneg
cnf(i_0_274, plain, (v1_xreal_0(esk20_0))).#trainneg
cnf(i_0_158, plain, (v1_xcmplx_0(esk8_0))).#trainneg
cnf(i_0_276, plain, (v1_xcmplx_0(esk20_0))).#trainneg
cnf(i_0_281, plain, (~v1_xboole_0(esk20_0))).#trainneg
cnf(i_0_287, plain, (v3_valued_0(esk21_0))).#trainneg
cnf(i_0_296, plain, (~v1_xboole_0(esk21_0))).#trainneg
cnf(i_0_286, plain, (v4_valued_0(esk21_0))).#trainneg
cnf(i_0_115, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_128, plain, (~v1_finset_1(k4_ordinal1))).#trainneg
cnf(i_0_289, plain, (v1_valued_0(esk21_0))).#trainneg
cnf(i_0_141, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_120, plain, (v3_ordinal1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_325, plain, (v3_xxreal_0(esk22_0))).#trainneg
cnf(i_0_288, plain, (v2_valued_0(esk21_0))).#trainneg
cnf(i_0_280, plain, (v1_ordinal1(esk20_0))).#trainneg
cnf(i_0_279, plain, (v2_ordinal1(esk20_0))).#trainneg
cnf(i_0_285, plain, (v5_valued_0(esk21_0))).#trainneg
cnf(i_0_46, plain, (v1_relat_1(k9_finseq_1(X1)))).#trainneg
cnf(i_0_169, plain, (v1_relat_1(esk9_1(X1)))).#trainneg
cnf(i_0_175, plain, (v1_relat_1(esk10_1(X1)))).#trainneg
cnf(i_0_200, plain, (v1_relat_1(esk12_1(X1)))).#trainneg
cnf(i_0_248, plain, (v1_relat_1(esk16_1(X1)))).#trainneg
cnf(i_0_45, plain, (v1_funct_1(k9_finseq_1(X1)))).#trainneg
cnf(i_0_187, plain, (~v1_xboole_0(k9_finseq_1(X1)))).#trainneg
cnf(i_0_161, plain, (~v1_xboole_0(k3_finseq_2(X1)))).#trainneg
cnf(i_0_98, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_102, plain, (v1_funct_1(esk6_1(X1)))).#trainneg
cnf(i_0_114, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_113, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_282, plain, (m1_subset_1(esk20_0,k5_numbers))).#trainneg
cnf(i_0_139, plain, (v1_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_188, plain, (v1_finset_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_166, plain, (v1_funct_1(esk9_1(X1)))).#trainneg
cnf(i_0_173, plain, (v1_funct_1(esk10_1(X1)))).#trainneg
cnf(i_0_198, plain, (v1_funct_1(esk12_1(X1)))).#trainneg
cnf(i_0_245, plain, (v1_funct_1(esk16_1(X1)))).#trainneg
cnf(i_0_48, plain, (v1_finseq_1(k9_finseq_1(X1)))).#trainneg
cnf(i_0_99, plain, (v1_finseq_1(esk6_1(X1)))).#trainneg
cnf(i_0_100, plain, (v1_finset_1(esk6_1(X1)))).#trainneg
cnf(i_0_240, plain, (v3_valued_0(esk16_1(X1)))).#trainneg
cnf(i_0_155, plain, (v1_pre_poly(k6_finseq_1(X1)))).#trainneg
cnf(i_0_165, plain, (v3_funct_1(esk9_1(X1)))).#trainneg
cnf(i_0_171, plain, (v4_valued_0(esk10_1(X1)))).#trainneg
cnf(i_0_89, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_293, plain, (v5_relat_1(esk21_0,k5_numbers))).#trainneg
cnf(i_0_196, plain, (v4_valued_0(esk12_1(X1)))).#trainneg
cnf(i_0_294, plain, (v4_relat_1(esk21_0,k5_numbers))).#trainneg
cnf(i_0_291, plain, (v1_partfun1(esk21_0,k5_numbers))).#trainneg
cnf(i_0_62, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_324, plain, (v1_card_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_255, plain, (v1_card_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_239, plain, (v4_valued_0(esk16_1(X1)))).#trainneg
cnf(i_0_242, plain, (v1_valued_0(esk16_1(X1)))).#trainneg
cnf(i_0_195, plain, (v2_pre_poly(esk12_1(X1)))).#trainneg
cnf(i_0_238, plain, (v2_pre_poly(esk16_1(X1)))).#trainneg
cnf(i_0_241, plain, (v2_valued_0(esk16_1(X1)))).#trainneg
cnf(i_0_254, plain, (v3_finseq_1(k3_finseq_2(X1)))).#trainneg
cnf(i_0_271, plain, (v4_funct_1(k3_finseq_2(X1)))).#trainneg
cnf(i_0_126, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_207, plain, (m1_subset_1(esk13_0,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_246, plain, (v5_relat_1(esk16_1(X1),k5_numbers))).#trainneg
cnf(i_0_104, plain, (v4_relat_1(esk6_1(X1),k5_numbers))).#trainneg
cnf(i_0_168, plain, (v4_relat_1(esk9_1(X1),k5_numbers))).#trainneg
cnf(i_0_256, plain, (v3_card_1(k9_finseq_1(X1),np__1))).#trainneg
cnf(i_0_77, plain, (m2_finseq_1(k6_finseq_1(X1),X1))).#trainneg
cnf(i_0_250, plain, (~v7_ordinal1(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_297, plain, (m1_subset_1(esk21_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_232, plain, (~v1_xboole_0(X1)|~v3_card_1(X1,np__1))).#trainneg
cnf(i_0_218, plain, (~v3_xxreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_54, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_125, plain, (v1_int_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_140, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_163, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_92, plain, (m1_subset_1(esk5_1(X1),X1))).#trainneg
cnf(i_0_103, plain, (v5_relat_1(esk6_1(X1),X1))).#trainneg
cnf(i_0_167, plain, (v5_relat_1(esk9_1(X1),X1))).#trainneg
cnf(i_0_174, plain, (v4_relat_1(esk10_1(X1),X1))).#trainneg
cnf(i_0_199, plain, (v4_relat_1(esk12_1(X1),X1))).#trainneg
cnf(i_0_247, plain, (v4_relat_1(esk16_1(X1),X1))).#trainneg
cnf(i_0_172, plain, (v1_partfun1(esk10_1(X1),X1))).#trainneg
cnf(i_0_197, plain, (v1_partfun1(esk12_1(X1),X1))).#trainneg
cnf(i_0_244, plain, (v1_partfun1(esk16_1(X1),X1))).#trainneg
cnf(i_0_181, plain, (v1_relat_1(esk11_2(X1,X2)))).#trainneg
cnf(i_0_268, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(esk19_1(X1)))).#trainneg
cnf(i_0_284, plain, (v1_relat_1(X1)|~v4_funct_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_203, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_190, plain, (v1_xreal_0(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_194, plain, (v1_xboole_0(X1)|~v3_card_1(X1,k1_xboole_0))).#trainneg
cnf(i_0_301, plain, (~v1_xboole_0(X1)|~v2_xxreal_0(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_299, plain, (~v2_xxreal_0(X1)|~v1_xxreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_192, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_160, plain, (v1_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_177, plain, (v2_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_208, plain, (v5_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_262, plain, (v1_relat_1(esk18_2(X1,X2)))).#trainneg
cnf(i_0_259, plain, (v1_funct_1(esk18_2(X1,X2)))).#trainneg
cnf(i_0_182, plain, (v1_xboole_0(esk11_2(X1,X2)))).#trainneg
cnf(i_0_179, plain, (v5_relat_1(esk11_2(X1,X2),X2))).#trainneg
cnf(i_0_260, plain, (v5_relat_1(esk18_2(X1,X2),X2))).#trainneg
cnf(i_0_180, plain, (v4_relat_1(esk11_2(X1,X2),X1))).#trainneg
cnf(i_0_261, plain, (v4_relat_1(esk18_2(X1,X2),X1))).#trainneg
cnf(i_0_243, plain, (v1_funct_2(esk16_1(X1),X1,k5_numbers))).#trainneg
cnf(i_0_164, plain, (v1_funct_2(esk9_1(X1),k5_numbers,X1))).#trainneg
cnf(i_0_221, plain, (v1_relat_1(esk14_1(X1))|~v1_card_1(X1))).#trainneg
cnf(i_0_230, plain, (v1_relat_1(esk15_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_88, plain, (v1_funct_1(u1_algstr_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_283, plain, (v1_funct_1(X1)|~v4_funct_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_60, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_220, plain, (v1_funct_1(esk14_1(X1))|~v1_card_1(X1))).#trainneg
cnf(i_0_111, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_107, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_318, plain, (v1_int_1(X1)|~m1_subset_1(X1,k4_numbers))).#trainneg
cnf(i_0_228, plain, (v1_funct_1(esk15_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_225, plain, (v1_finseq_1(esk15_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_253, plain, (l1_struct_0(esk17_1(X1))|~v1_card_1(X1))).#trainneg
cnf(i_0_227, plain, (v1_finset_1(esk15_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_224, plain, (v2_finseq_1(esk15_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_270, plain, (v5_finset_1(k1_zfmisc_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_184, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_189, plain, (v1_card_3(X1)|~m1_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_231, plain, (v1_zfmisc_1(X1)|~v3_card_1(X1,np__1))).#trainneg
cnf(i_0_36, plain, (m2_finseq_2(k6_finseq_1(X1),X1,k3_finseq_2(X1)))).#trainneg
cnf(i_0_258, plain, (v1_funct_2(esk18_2(X1,X2),X1,X2))).#trainneg
cnf(i_0_233, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_121, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_229, plain, (v4_relat_1(esk15_1(X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_204, plain, (v1_funct_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_91, plain, (v7_ordinal1(X1)|~v3_ordinal1(X1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_222, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainneg
cnf(i_0_201, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_122, plain, (v1_finseq_1(X1)|~m1_subset_1(X1,k3_finseq_2(X2)))).#trainneg
cnf(i_0_149, plain, (v1_finset_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_267, plain, (v2_struct_0(X1)|v1_zfmisc_1(esk19_1(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_219, plain, (v3_card_1(esk14_1(X1),X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_226, plain, (v3_card_1(esk15_1(X1),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_252, plain, (v13_struct_0(esk17_1(X1),X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_53, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_265, plain, (v3_valued_0(k2_zfmisc_1(X2,X1))|~v3_membered(X1))).#trainneg
cnf(i_0_249, plain, (m1_subset_1(esk16_1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,k5_numbers))))).#trainneg
cnf(i_0_170, plain, (m1_subset_1(esk9_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_95, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_96, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_183, plain, (m1_subset_1(esk11_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_263, plain, (m1_subset_1(esk18_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_112, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_237, plain, (v1_finset_1(X2)|~v5_finset_1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_186, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_35, plain, (v1_relat_1(k3_finseq_5(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_215, plain, (v1_xboole_0(X2)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_xboole_0(k7_finseq_1(X1,X2)))).#trainneg
cnf(i_0_84, plain, (v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_143, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_185, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_43, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_209, plain, (v1_pre_poly(k9_finseq_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_55, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_178, plain, (v1_finset_1(X2)|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_138, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_298, plain, (v1_membered(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_234, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_142, plain, (v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_145, plain, (v1_int_1(k2_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_236, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_22, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_58, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_59, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_154, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_211, plain, (v1_xboole_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_xboole_0(k7_finseq_1(X1,X2)))).#trainneg
cnf(i_0_321, plain, (v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)|~v3_valued_0(X1)|~v5_valued_0(X1))).#trainneg
cnf(i_0_235, plain, (v3_valued_0(X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_257, plain, (v1_valued_0(X2)|~v1_relat_1(X1)|~v1_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_144, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_269, plain, (v2_struct_0(X1)|m1_subset_1(esk19_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_78, plain, (k1_nat_1(X1,X2)=k1_nat_1(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_63, plain, (k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_76, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_66, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_67, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_193, plain, (v1_valued_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_317, plain, (v2_valued_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_81, plain, (m2_subset_1(k1_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_87, plain, (v1_funct_2(u1_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_264, plain, (v1_partfun1(X1,X2)|~v1_funct_2(X1,X2,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2))))).#trainneg
cnf(i_0_130, plain, (v2_funct_1(k3_finseq_5(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)|~v2_funct_1(X1))).#trainneg
cnf(i_0_119, plain, (v1_relat_1(k7_finseq_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_118, plain, (v1_funct_1(k7_finseq_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_108, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_117, plain, (v1_finseq_1(k7_finseq_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_68, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_86, plain, (m1_subset_1(u1_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_123, plain, (v1_xboole_0(X1)|v1_partfun1(X2,X3)|~v1_funct_2(X2,X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_109, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_162, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_90, plain, (m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_147, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|~v1_funct_1(X1)|~v1_xboole_0(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_94, plain, (m1_subset_1(k3_rlvect_1(X1,X2,X3),u1_struct_0(X1))|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_176, plain, (k5_binop_1(u1_struct_0(X1),u1_algstr_0(X1),X2,X3)=k1_algstr_0(X1,X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_266, plain, (m1_subset_1(k5_binop_1(X2,X1,X3,X4),X2)|~v1_funct_1(X1)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v1_funct_2(X1,k2_zfmisc_1(X2,X2),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2))))).#trainneg
cnf(i_0_191, plain, (k5_binop_1(X2,X1,X3,X4)=k1_binop_1(X1,X3,X4)|~v1_funct_1(X1)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v1_funct_2(X1,k2_zfmisc_1(X2,X2),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2))))).#trainneg
cnf(i_0_424, negated_conjecture, (k4_rlvect_1(esk1_0,k3_finseq_5(esk2_0))!=k4_rlvect_1(esk1_0,esk2_0)|~m1_finseq_1(esk2_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_339, plain, (~v7_ordinal1(esk22_0))).#trainneg
cnf(i_0_330, plain, (~v7_ordinal1(k5_numbers))).#trainneg
cnf(i_0_425, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_774, plain, (~m1_subset_1(esk22_0,k5_numbers))).#trainneg
cnf(i_0_773, plain, (~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_342, plain, (v1_funct_1(esk8_0))).#trainneg
cnf(i_0_343, plain, (v1_funct_1(esk7_0))).#trainneg
cnf(i_0_327, plain, (esk8_0=k1_xboole_0)).#trainneg
cnf(i_0_328, plain, (esk7_0=k1_xboole_0)).#trainneg
cnf(i_0_345, plain, (v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_518, plain, (~v1_finset_1(k1_numbers))).#trainneg
cnf(i_0_359, negated_conjecture, (v1_funct_1(esk2_0))).#trainneg
cnf(i_0_367, negated_conjecture, (v1_finseq_1(esk2_0))).#trainneg
cnf(i_0_407, plain, (v7_ordinal1(X1)|~v1_finset_1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_526, plain, (v3_membered(k5_numbers))).#trainneg
cnf(i_0_397, plain, (v7_ordinal1(np__1))).#trainneg
cnf(i_0_409, plain, (~v1_finseq_1(k5_numbers)|~v1_funct_1(k5_numbers)|~v1_relat_1(k5_numbers))).#trainneg
cnf(i_0_398, plain, (v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_777, plain, (~m1_subset_1(k5_numbers,k5_numbers))).#trainneg
cnf(i_0_405, plain, (v7_ordinal1(esk13_0)|~v1_finset_1(esk13_0))).#trainneg
cnf(i_0_446, plain, (v3_ordinal1(np__1))).#trainneg
cnf(i_0_447, plain, (v3_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_481, plain, (v1_xreal_0(np__1))).#trainneg
cnf(i_0_522, plain, (v3_membered(esk13_0))).#trainneg
cnf(i_0_565, plain, (v5_relat_1(X1,X2)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_352, plain, (v1_relat_1(esk20_0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_353, plain, (v1_relat_1(esk13_0)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_685, plain, (m2_finseq_1(k3_pre_poly(X1,X2),X1)|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_645, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_411, plain, (v1_xboole_0(esk14_1(k1_xboole_0))|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_413, plain, (v1_xboole_0(esk15_1(k1_xboole_0))|~v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_659, plain, (v1_valued_0(X1)|~v1_membered(X2)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_580, plain, (v3_valued_0(X1)|~v3_membered(X2)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_772, plain, (k5_binop_1(u1_struct_0(X1),u1_algstr_0(X1),X2,X3)=k1_binop_1(u1_algstr_0(X1),X2,X3)|~v1_funct_1(u1_algstr_0(X1))|~m1_subset_1(u1_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_355, plain, (v1_relat_1(np__1)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_365, plain, (v1_funct_1(esk11_2(X1,X2)))).#trainneg
cnf(i_0_354, plain, (v1_relat_1(esk5_1(X1))|~v4_funct_1(X1))).#trainneg
cnf(i_0_356, plain, (v1_relat_1(np__1)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_460, plain, (~v5_finset_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_377, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_algstr_0(X1))).#trainneg
cnf(i_0_499, plain, (v1_funct_1(k4_finseq_5(X1,X2))|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_357, plain, (v1_relat_1(k1_xboole_0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_358, plain, (v1_relat_1(k5_numbers)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_667, plain, (m2_subset_1(k1_nat_1(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_786, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_668, plain, (m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_627, plain, (k1_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_789, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_769, plain, (m1_subset_1(k5_binop_1(u1_struct_0(X1),u1_algstr_0(X1),X2,X3),u1_struct_0(X1))|~v1_funct_1(u1_algstr_0(X1))|~m1_subset_1(u1_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_370, plain, (v1_funct_1(esk20_0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_408, plain, (v1_finseq_1(esk5_1(k3_finseq_2(X1))))).#trainneg
cnf(i_0_410, plain, (~v1_xboole_0(esk14_1(np__1))|~v1_card_1(np__1))).#trainneg
cnf(i_0_364, plain, (esk11_2(X1,X2)=k1_xboole_0)).#trainneg
cnf(i_0_373, plain, (v1_funct_1(np__1)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_362, plain, (~v1_xxreal_0(np__1)|~v1_xboole_0(np__1))).#trainneg
cnf(i_0_363, plain, (~v3_xxreal_0(np__1)|~v1_xxreal_0(np__1))).#trainneg
cnf(i_0_500, plain, (v1_finseq_1(k4_finseq_5(X1,X2))|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_666, plain, (v2_valued_0(X1)|~v2_membered(X2)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_371, plain, (v1_funct_1(esk13_0)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_372, plain, (v1_funct_1(esk5_1(X1))|~v4_funct_1(X1))).#trainneg
cnf(i_0_804, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1))).#trainneg
cnf(i_0_806, plain, (~v7_ordinal1(k1_numbers))).#trainneg
cnf(i_0_807, plain, (~m1_subset_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_157, plain, (v1_xxreal_0(k1_xboole_0))).#trainneg
cnf(i_0_156, plain, (v1_xreal_0(k1_xboole_0))).#trainneg
cnf(i_0_374, plain, (v1_funct_1(np__1)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_104, plain, (v4_relat_1(k1_xboole_0,k5_numbers))).#trainneg
cnf(i_0_899, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_817, negated_conjecture, (v1_relat_1(esk2_0))).#trainneg
cnf(i_0_376, plain, (v1_funct_1(k5_numbers)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_808, plain, (~v1_finseq_1(k1_numbers)|~v1_funct_1(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_859, plain, (k1_binop_1(u1_algstr_0(X1),X2,X3)=k1_algstr_0(X1,X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_100, plain, (v1_finset_1(k1_xboole_0))).#trainneg
cnf(i_0_183, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_714, plain, (v1_xboole_0(u1_algstr_0(X1))|~v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_158, plain, (v1_xcmplx_0(k1_xboole_0))).#trainneg
cnf(i_0_155, plain, (v1_pre_poly(k1_xboole_0))).#trainneg
cnf(i_0_383, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k4_numbers))).#trainneg
cnf(i_0_99, plain, (v1_finseq_1(k1_xboole_0))).#trainneg
cnf(i_0_412, plain, (~v1_xboole_0(esk15_1(np__1))|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_819, plain, (v1_relat_1(k4_finseq_5(X1,X2))|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_384, plain, (v1_xboole_0(u1_struct_0(esk17_1(X1)))|~v2_struct_0(esk17_1(X1))|~v1_card_1(X1))).#trainneg
cnf(i_0_385, plain, (v2_struct_0(esk17_1(X1))|~v1_xboole_0(u1_struct_0(esk17_1(X1)))|~v1_card_1(X1))).#trainneg
cnf(i_0_543, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_77, plain, (m2_finseq_1(k1_xboole_0,X1))).#trainneg
cnf(i_0_643, plain, (v1_xboole_0(esk9_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_885, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_793, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_845, plain, (esk14_1(k1_xboole_0)=k1_xboole_0|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_648, plain, (v1_xboole_0(esk16_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_833, plain, (v7_ordinal1(esk13_0)|~v1_finseq_1(esk13_0)|~v1_funct_1(esk13_0)|~v1_relat_1(esk13_0))).#trainneg
cnf(i_0_771, plain, (k1_binop_1(esk16_1(k2_zfmisc_1(k5_numbers,k5_numbers)),X1,X2)=k5_binop_1(k5_numbers,esk16_1(k2_zfmisc_1(k5_numbers,k5_numbers)),X1,X2)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_103, plain, (v5_relat_1(k1_xboole_0,X1))).#trainneg
cnf(i_0_848, plain, (esk15_1(k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_392, plain, (v7_ordinal1(esk21_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_457, plain, (v1_finset_1(np__1)|~v5_finset_1(k1_numbers))).#trainneg
cnf(i_0_851, negated_conjecture, (v1_valued_0(esk2_0)|~v1_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_855, negated_conjecture, (v3_valued_0(esk2_0)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_394, plain, (v7_ordinal1(esk13_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_395, plain, (v7_ordinal1(esk5_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_908, negated_conjecture, (k4_rlvect_1(esk1_0,esk2_0)!=k4_rlvect_1(esk1_0,k1_xboole_0)|~m1_finseq_1(esk2_0,u1_struct_0(esk1_0))|~v1_xboole_0(esk2_0)|~v1_relat_1(esk2_0))).#trainneg
cnf(i_0_871, plain, (m2_subset_1(k2_xcmplx_0(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_750, plain, (v1_xboole_0(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|v1_xboole_0(u1_struct_0(X1))|~v1_xboole_0(u1_algstr_0(X1))|~m1_subset_1(u1_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_534, plain, (v1_membered(k5_numbers)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_458, plain, (v1_finset_1(np__1)|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_180, plain, (v4_relat_1(k1_xboole_0,X1))).#trainneg
cnf(i_0_530, plain, (v1_membered(esk13_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_399, plain, (v7_ordinal1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_469, plain, (v1_xxreal_0(np__1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_470, plain, (v1_xxreal_0(np__1)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_903, negated_conjecture, (v2_valued_0(esk2_0)|~v2_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_401, plain, (v1_funct_1(esk13_0)|~v1_funct_1(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_402, plain, (v1_funct_1(esk5_1(k1_zfmisc_1(X1)))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_681, plain, (m1_subset_1(k4_rlvect_1(X1,esk9_1(u1_struct_0(X1))),u1_struct_0(X1))|v2_struct_0(X1)|~v1_finseq_1(esk9_1(u1_struct_0(X1)))|~l2_algstr_0(X1))).#trainneg
cnf(i_0_824, plain, (v1_finset_1(k3_pre_poly(X1,X2))|v1_xboole_0(X1)|~v5_finset_1(k3_finseq_2(X1))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_870, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_843, plain, (v1_xboole_0(k8_finseq_1(X1,X2,X3))|~v1_xboole_0(X1)|~m1_finseq_1(X3,X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_864, plain, (m1_subset_1(k1_nat_1(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_36, plain, (m2_finseq_2(k1_xboole_0,X1,k3_finseq_2(X1)))).#trainneg
cnf(i_0_403, plain, (v1_funct_1(k5_numbers)|~v1_funct_1(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1022, plain, (v3_card_1(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_1039, negated_conjecture, (v2_valued_0(esk2_0)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1025, negated_conjecture, (v1_valued_0(esk2_0)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_416, plain, (v7_ordinal1(esk16_1(X1))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k5_numbers))))).#trainneg
cnf(i_0_421, plain, (v7_ordinal1(esk9_1(X1))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_918, negated_conjecture, (m1_subset_1(k4_rlvect_1(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1036, plain, (v1_membered(k5_numbers))).#trainneg
cnf(i_0_919, plain, (m1_subset_1(k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),X2)),u1_struct_0(X1))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1020, plain, (v2_finseq_1(k1_xboole_0))).#trainneg
cnf(i_0_1037, plain, (v1_xxreal_0(np__1))).#trainneg
cnf(i_0_538, plain, (v2_membered(esk13_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1027, negated_conjecture, (k4_rlvect_1(esk1_0,esk2_0)!=k4_rlvect_1(esk1_0,k1_xboole_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_434, plain, (v7_ordinal1(esk18_2(X1,X2))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_844, plain, (v1_xboole_0(k4_finseq_5(X1,X2))|~v1_xboole_0(X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_854, plain, (v1_valued_0(k4_finseq_5(X1,X2))|~v1_membered(X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_437, plain, (v3_ordinal1(esk21_0)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_439, plain, (v3_ordinal1(esk18_2(X1,X2))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_865, plain, (m1_subset_1(k1_nat_1(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_735, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_601, plain, (v1_valued_0(k5_numbers)|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_874, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_542, plain, (v2_membered(k5_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_787, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_711, plain, (v1_relat_1(u1_algstr_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_734, plain, (m1_subset_1(k1_nat_1(X1,X2),k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_584, plain, (v2_funct_1(esk21_0)|~v1_finseq_1(esk21_0))).#trainneg
cnf(i_0_552, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,X3))|~m2_finseq_1(X2,X3))).#trainneg
cnf(i_0_440, plain, (v3_ordinal1(esk16_1(X1))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(X1,k5_numbers))))).#trainneg
cnf(i_0_493, plain, (v1_xcmplx_0(np__1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_494, plain, (v1_xcmplx_0(np__1)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_858, plain, (v3_valued_0(k4_finseq_5(X1,X2))|~v3_membered(X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_907, plain, (v2_valued_0(k4_finseq_5(X1,X2))|~v2_membered(X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_443, plain, (v3_ordinal1(esk9_1(X1))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_444, plain, (v3_ordinal1(esk5_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1174, plain, (v2_membered(k5_numbers))).#trainneg
cnf(i_0_1074, plain, (m1_subset_1(k1_xboole_0,k3_finseq_2(X1)))).#trainneg
cnf(i_0_1156, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1109, negated_conjecture, (k4_rlvect_1(esk1_0,k4_finseq_5(u1_struct_0(esk1_0),esk3_1(esk1_0)))=k4_rlvect_1(esk1_0,esk3_1(esk1_0))|k4_rlvect_1(esk1_0,k3_finseq_5(k1_xboole_0))!=k4_rlvect_1(esk1_0,k1_xboole_0))).#trainneg
cnf(i_0_1111, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2)))=k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),X2))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|k4_rlvect_1(X1,k3_finseq_5(k1_xboole_0))!=k4_rlvect_1(X1,k1_xboole_0)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_886, plain, (v1_funct_1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_363, plain, (~v3_xxreal_0(np__1))).#trainneg
cnf(i_0_449, plain, (v1_finset_1(esk21_0)|~v5_finset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_465, plain, (v1_xxreal_0(esk13_0)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1127, plain, (X1=k4_finseq_5(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X2)|~m1_finseq_1(X3,X2))).#trainneg
cnf(i_0_451, plain, (v1_finset_1(esk18_2(X1,X2))|~v5_finset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_452, plain, (v1_finset_1(esk16_1(X1))|~v5_finset_1(k1_zfmisc_1(k2_zfmisc_1(X1,k5_numbers))))).#trainneg
cnf(i_0_748, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk16_1(X1)))).#trainneg
cnf(i_0_922, plain, (m1_subset_1(k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2)),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_763, plain, (m1_subset_1(k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3)),u1_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_rlvect_1(X1,X3),u1_struct_0(X1))|~m1_subset_1(k4_rlvect_1(X1,X2),u1_struct_0(X1))|~m2_finseq_1(X3,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1))).#trainneg
cnf(i_0_677, plain, (v1_xboole_0(X1)|~v1_xboole_0(k3_finseq_5(X1))|~v1_relat_1(k3_finseq_5(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_898, plain, (m1_subset_1(k1_binop_1(u1_algstr_0(X1),X2,X3),u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_869, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1202, plain, (~r2_hidden(X1,X2)|~m2_finseq_1(X2,X3)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_949, plain, (v3_valued_0(k1_xboole_0)|~v3_membered(X1))).#trainneg
cnf(i_0_673, plain, (v1_partfun1(esk9_1(k5_numbers),k5_numbers))).#trainneg
cnf(i_0_588, plain, (v3_valued_0(esk13_0)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_936, plain, (~r2_hidden(X1,k1_xboole_0)|~v1_xboole_0(k2_zfmisc_1(X2,X3)))).#trainneg
cnf(i_0_455, plain, (v1_finset_1(esk9_1(X1))|~v5_finset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_592, plain, (v3_valued_0(k5_numbers)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_472, plain, (v1_xxreal_0(k5_numbers)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_477, plain, (v1_xreal_0(esk13_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_456, plain, (v1_finset_1(esk5_1(X1))|~v5_finset_1(X1))).#trainneg
cnf(i_0_461, plain, (v1_xxreal_0(esk21_0)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_640, plain, (v1_xboole_0(esk18_2(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1175, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1123, plain, (v1_xboole_0(u1_struct_0(X1))|m1_subset_1(k4_rlvect_1(X1,k9_finseq_1(X2)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1))).#trainneg
cnf(i_0_782, plain, (v1_finset_1(X1)|~v5_finset_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1214, plain, (k1_algstr_0(X1,k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X3,X2)),k4_rlvect_1(X1,X4))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k8_finseq_1(u1_struct_0(X1),X2,X3),X4))|v2_struct_0(X1)|~m2_finseq_1(X4,u1_struct_0(X1))|~m2_finseq_1(k8_finseq_1(u1_struct_0(X1),X2,X3),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~m2_finseq_1(X3,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_1324, plain, (v3_valued_0(k1_xboole_0))).#trainneg
cnf(i_0_1253, plain, (m2_finseq_1(esk5_1(k3_finseq_2(X1)),X1))).#trainneg
cnf(i_0_597, plain, (v1_valued_0(esk13_0)|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_722, plain, (~v1_xboole_0(k3_finseq_5(k7_finseq_1(X2,k9_finseq_1(X1))))|~v1_finseq_1(k3_finseq_5(X2))|~v1_funct_1(k3_finseq_5(X2))|~v1_relat_1(k3_finseq_5(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_463, plain, (v1_xxreal_0(esk18_2(X1,X2))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1334, plain, (v3_valued_0(k5_numbers)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k1_numbers))).#trainneg
cnf(i_0_484, plain, (v1_xreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_489, plain, (v1_xcmplx_0(esk13_0)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_464, plain, (v1_xxreal_0(esk16_1(X1))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k5_numbers))))).#trainneg
cnf(i_0_467, plain, (v1_xxreal_0(esk9_1(X1))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_893, plain, (v1_finset_1(k5_binop_1(u1_struct_0(X1),u1_algstr_0(X1),X2,X3))|~v5_finset_1(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1092, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))))!=k4_rlvect_1(X1,k1_xboole_0)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k1_xboole_0))!=k4_rlvect_1(X1,k1_xboole_0)|~m1_finseq_1(k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_xboole_0(k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))))|~v1_relat_1(k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1)))))).#trainneg
cnf(i_0_953, plain, (m1_subset_1(k4_rlvect_1(X1,k1_xboole_0),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_578, plain, (v3_valued_0(esk9_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_882, plain, (v1_xboole_0(k7_finseq_1(X1,X2))|~v1_xboole_0(X3)|~m1_finseq_1(X2,X3)|~m1_finseq_1(X1,X3))).#trainneg
cnf(i_0_630, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,esk19_1(X2))|~l1_struct_0(X2))).#trainneg
cnf(i_0_686, plain, (m2_finseq_2(k9_finseq_1(X2),X1,k3_finseq_2(X1))|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_950, plain, (v1_valued_0(k1_xboole_0)|~v1_membered(X1))).#trainneg
cnf(i_0_671, plain, (v1_partfun1(esk18_2(X1,X1),X1))).#trainneg
cnf(i_0_1159, plain, (v1_valued_0(k5_numbers)|~v1_relat_1(k1_numbers)|~v3_valued_0(k1_numbers))).#trainneg
cnf(i_0_468, plain, (v1_xxreal_0(esk5_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_951, plain, (v2_valued_0(k1_xboole_0)|~v2_membered(X1))).#trainneg
cnf(i_0_496, plain, (v1_xcmplx_0(k5_numbers)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1259, plain, (v1_xxreal_0(esk13_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1269, plain, (k4_finseq_5(X1,X2)=k4_finseq_5(X3,X4)|~v1_xboole_0(X3)|~m1_finseq_1(X4,X3)|~v1_xboole_0(X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_1099, plain, (v1_xboole_0(X2)|~v1_xboole_0(k4_finseq_5(X1,X2))|~v1_relat_1(k4_finseq_5(X1,X2))|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_473, plain, (v1_xreal_0(esk21_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_475, plain, (v1_xreal_0(esk18_2(X1,X2))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1112, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k8_finseq_1(u1_struct_0(X1),X2,X3)))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3))|v2_struct_0(X1)|k4_rlvect_1(X1,k3_finseq_5(k1_xboole_0))!=k4_rlvect_1(X1,k1_xboole_0)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1423, plain, (~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(k7_finseq_1(X1,k9_finseq_1(X2)))|~v1_relat_1(k7_finseq_1(X1,k9_finseq_1(X2))))).#trainneg
cnf(i_0_1436, negated_conjecture, (m1_subset_1(k4_rlvect_1(esk1_0,k1_xboole_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1402, plain, (k1_algstr_0(X1,k4_rlvect_1(X1,X4),k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3)))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k8_finseq_1(u1_struct_0(X1),X3,X2),X4))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3)),u1_struct_0(X1))|~m1_subset_1(k4_rlvect_1(X1,X4),u1_struct_0(X1))|~v2_rlvect_1(X1)|~m2_finseq_1(k8_finseq_1(u1_struct_0(X1),X3,X2),u1_struct_0(X1))|~m2_finseq_1(X4,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1))).#trainneg
cnf(i_0_1445, plain, (v1_valued_0(k1_xboole_0))).#trainneg
cnf(i_0_1460, plain, (v2_valued_0(k1_xboole_0))).#trainneg
cnf(i_0_476, plain, (v1_xreal_0(esk16_1(X1))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k5_numbers))))).#trainneg
cnf(i_0_1335, plain, (v1_xxreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1426, plain, (v1_xcmplx_0(esk13_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1512, plain, (v1_xboole_0(X1)|~m1_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_479, plain, (v1_xreal_0(esk9_1(X1))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_480, plain, (v1_xreal_0(esk5_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_1509, plain, (v1_xboole_0(X1)|~v1_xboole_0(k3_finseq_5(X1))|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_1087, negated_conjecture, (v3_ordinal1(k4_rlvect_1(esk1_0,esk2_0))|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_768, plain, (m1_subset_1(k5_binop_1(k5_numbers,esk16_1(k2_zfmisc_1(k5_numbers,k5_numbers)),X1,X2),k5_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_724, plain, (k3_finseq_5(k7_finseq_1(k3_finseq_5(X2),k9_finseq_1(X1)))=k7_finseq_1(k9_finseq_1(X1),X2)|~v1_finseq_1(k3_finseq_5(X2))|~v1_funct_1(k3_finseq_5(X2))|~v1_relat_1(k3_finseq_5(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1264, plain, (k4_finseq_5(X1,X2)=k1_xboole_0|~v1_xboole_0(X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_718, plain, (v5_relat_1(u1_algstr_0(X1),u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_731, plain, (m1_subset_1(k1_nat_1(X1,X2),k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1416, plain, (v1_funct_1(esk5_1(k3_finseq_2(X1))))).#trainneg
cnf(i_0_1026, negated_conjecture, (v3_valued_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_485, plain, (v1_xcmplx_0(esk21_0)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_1462, plain, (v1_xcmplx_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_511, plain, (v1_finset_1(esk21_0)|~v1_finset_1(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_487, plain, (v1_xcmplx_0(esk18_2(X1,X2))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_488, plain, (v1_xcmplx_0(esk16_1(X1))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k5_numbers))))).#trainneg
cnf(i_0_965, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_1111, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2)))=k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_1112, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k8_finseq_1(u1_struct_0(X1),X2,X3)))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_1109, negated_conjecture, (k4_rlvect_1(esk1_0,k4_finseq_5(u1_struct_0(esk1_0),esk3_1(esk1_0)))=k4_rlvect_1(esk1_0,esk3_1(esk1_0)))).#trainneg
cnf(i_0_1603, plain, (k4_rlvect_1(X1,k7_finseq_1(X2,X3))=k4_rlvect_1(X1,k7_finseq_1(X3,X2))|v2_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~m1_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_1418, plain, (v1_relat_1(esk5_1(k3_finseq_2(X1))))).#trainneg
cnf(i_0_1664, plain, (v1_finset_1(esk21_0)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1587, plain, (~v1_xboole_0(k7_finseq_1(k9_finseq_1(X2),X1))|~v1_finseq_1(k3_finseq_5(X1))|~v1_funct_1(k3_finseq_5(X1))|~v1_relat_1(k3_finseq_5(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_491, plain, (v1_xcmplx_0(esk9_1(X1))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_1666, plain, (v1_finset_1(esk21_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k5_numbers))).#trainneg
cnf(i_0_519, plain, (v3_membered(esk21_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_527, plain, (v1_membered(esk21_0)|~v1_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_492, plain, (v1_xcmplx_0(esk5_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_1667, plain, (v1_finset_1(esk21_0)|~v1_finseq_1(k2_zfmisc_1(k5_numbers,k5_numbers))|~v1_funct_1(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1240, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1645, plain, (m1_subset_1(k4_rlvect_1(X1,u1_algstr_0(X1)),u1_struct_0(X1))|v2_struct_0(X1)|~v1_finseq_1(u1_algstr_0(X1))|~v1_funct_1(u1_algstr_0(X1))|~v1_relat_1(u1_algstr_0(X1))|~l2_algstr_0(X1)|~l1_algstr_0(X1))).#trainneg
cnf(i_0_788, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_1077, negated_conjecture, (v1_valued_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_510, plain, (v1_relat_1(esk5_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_535, plain, (v2_membered(esk21_0)|~v2_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_512, plain, (v1_finset_1(esk18_2(X1,X2))|~v1_finset_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_513, plain, (v1_finset_1(esk16_1(X1))|~v1_finset_1(k2_zfmisc_1(X1,k5_numbers)))).#trainneg
cnf(i_0_8, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1)))))!=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_1768, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k3_finseq_5(X2)))=k4_rlvect_1(X1,k3_finseq_5(X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1676, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k7_finseq_1(X2,X3)))=k4_rlvect_1(X1,k7_finseq_1(X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_778, plain, (v1_relat_1(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_657, plain, (v1_valued_0(esk9_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_1336, plain, (v1_xreal_0(esk13_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1092, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))))!=k4_rlvect_1(X1,k1_xboole_0)|~v1_xboole_0(k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))))|~m1_finseq_1(k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_516, plain, (v1_finset_1(esk9_1(X1))|~v1_finset_1(k2_zfmisc_1(k5_numbers,X1)))).#trainneg
cnf(i_0_1665, plain, (v1_finset_1(esk21_0)|~v7_ordinal1(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_517, plain, (v1_finset_1(esk5_1(k1_zfmisc_1(X1)))|~v1_finset_1(X1))).#trainneg
cnf(i_0_520, plain, (v3_membered(esk18_2(X1,X2))|~v3_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1677, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)))=k4_rlvect_1(X1,k9_finseq_1(X2))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1682, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3))))=k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),X3)))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_1687, plain, (k4_rlvect_1(X1,k3_finseq_5(k3_pre_poly(u1_struct_0(X1),X2)))=k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_pre_poly(u1_struct_0(X1),X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_1770, plain, (v1_finseq_1(k3_finseq_5(X1))|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_1683, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3)),X2))=k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),X3)))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_1671, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk5_1(k3_finseq_2(u1_struct_0(X1)))))=k4_rlvect_1(X1,esk5_1(k3_finseq_2(u1_struct_0(X1))))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_1724, plain, (k4_rlvect_1(X1,k3_finseq_5(k8_finseq_1(u1_struct_0(X1),X2,X3)))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k8_finseq_1(u1_struct_0(X1),X2,X3),u1_struct_0(X1)))).#trainneg
cnf(i_0_1679, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2)),X3))=k1_algstr_0(X1,k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),X2)),k4_rlvect_1(X1,X3))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m2_finseq_1(X3,u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_779, plain, (v1_funct_1(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_664, plain, (v2_valued_0(esk9_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_1425, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_521, plain, (v3_membered(esk16_1(X1))|~v3_membered(k2_zfmisc_1(X1,k5_numbers)))).#trainneg
cnf(i_0_1819, plain, (m2_finseq_1(X1,X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_524, plain, (v3_membered(esk9_1(X1))|~v3_membered(k2_zfmisc_1(k5_numbers,X1)))).#trainneg
cnf(i_0_525, plain, (v3_membered(esk5_1(k1_zfmisc_1(X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_1845, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1842, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),esk5_1(k3_finseq_2(u1_struct_0(X1))))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk5_1(k3_finseq_2(u1_struct_0(X1)))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_1893, plain, (k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),X3)))=k4_rlvect_1(X1,k7_finseq_1(X2,k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3)),u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1961, plain, (k4_finseq_5(X1,X2)=k3_finseq_5(X2)|~m1_finseq_1(X2,X3)|~m2_finseq_1(k3_finseq_5(X2),X1))).#trainneg
cnf(i_0_528, plain, (v1_membered(esk18_2(X1,X2))|~v1_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1765, plain, (v1_membered(esk21_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1474, plain, (m2_finseq_1(k4_finseq_5(X3,X4),X1)|~m1_finseq_1(X2,X1)|~v1_xboole_0(X1)|~v1_xboole_0(X3)|~m1_finseq_1(X4,X3))).#trainneg
cnf(i_0_529, plain, (v1_membered(esk16_1(X1))|~v1_membered(k2_zfmisc_1(X1,k5_numbers)))).#trainneg
cnf(i_0_532, plain, (v1_membered(esk9_1(X1))|~v1_membered(k2_zfmisc_1(k5_numbers,X1)))).#trainneg
cnf(i_0_1901, plain, (k4_rlvect_1(X1,k3_finseq_5(k3_pre_poly(u1_struct_0(X1),X2)))=k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(k3_pre_poly(u1_struct_0(X1),X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_1852, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k3_finseq_5(X2)))=k4_rlvect_1(X1,k3_finseq_5(X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1850, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),X2))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1865, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k9_finseq_1(k3_rlvect_1(X1,X2,X3))))=k4_rlvect_1(X1,k9_finseq_1(k3_rlvect_1(X1,X2,X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2105, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_2065, plain, (k4_rlvect_1(X1,k3_finseq_5(k3_pre_poly(u1_struct_0(X1),k5_binop_1(u1_struct_0(X1),u1_algstr_0(X1),X2,X3))))=k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),k5_binop_1(u1_struct_0(X1),u1_algstr_0(X1),X2,X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_533, plain, (v1_membered(esk5_1(k1_zfmisc_1(X1)))|~v1_membered(X1))).#trainneg
cnf(i_0_1963, plain, (k4_finseq_5(X1,X2)=k4_finseq_5(X4,X5)|~m1_finseq_1(X2,X3)|~m2_finseq_1(k4_finseq_5(X4,X5),X1)|~v1_xboole_0(X4)|~v1_xboole_0(X3)|~m1_finseq_1(X5,X4))).#trainneg
cnf(i_0_536, plain, (v2_membered(esk18_2(X1,X2))|~v2_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1828, plain, (v2_membered(esk21_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1975, plain, (v3_valued_0(X1)|~v3_membered(X2)|~v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_537, plain, (v2_membered(esk16_1(X1))|~v2_membered(k2_zfmisc_1(X1,k5_numbers)))).#trainneg
cnf(i_0_540, plain, (v2_membered(esk9_1(X1))|~v2_membered(k2_zfmisc_1(k5_numbers,X1)))).#trainneg
cnf(i_0_2066, plain, (k4_rlvect_1(X1,k3_finseq_5(k3_pre_poly(u1_struct_0(X1),k3_rlvect_1(X1,X2,X3))))=k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),k3_rlvect_1(X1,X2,X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1866, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k9_finseq_1(k1_algstr_0(X1,X2,X3))))=k4_rlvect_1(X1,k9_finseq_1(k1_algstr_0(X1,X2,X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1849, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1854, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k7_finseq_1(X2,X3)))=k4_rlvect_1(X1,k7_finseq_1(X2,X3))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_1946, plain, (k1_algstr_0(X1,k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),X2)),k4_rlvect_1(X1,X3))=k4_rlvect_1(X1,k7_finseq_1(X3,k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2))))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2)),u1_struct_0(X1))|~m1_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2)),u1_struct_0(X1))|~m2_finseq_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_2196, negated_conjecture, (m2_finseq_1(esk2_0,X1)|~m2_finseq_1(k3_finseq_5(esk2_0),X1))).#trainneg
cnf(i_0_2138, negated_conjecture, (m2_finseq_1(esk2_0,X1)|~m1_finseq_1(k3_finseq_5(esk2_0),X1))).#trainneg
cnf(i_0_541, plain, (v2_membered(esk5_1(k1_zfmisc_1(X1)))|~v2_membered(X1))).#trainneg
cnf(i_0_2172, plain, (v3_valued_0(X1)|~v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1977, plain, (v1_valued_0(X1)|~v1_membered(X2)|~v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2154, plain, (k4_rlvect_1(X1,k3_finseq_5(k3_pre_poly(u1_struct_0(X1),k1_algstr_0(X1,X2,X3))))=k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),k1_algstr_0(X1,X2,X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~l1_algstr_0(X1))).#trainneg
cnf(i_0_1884, plain, (k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),X3)))=k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3)),X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~m1_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1903, plain, (v1_finseq_1(k3_finseq_5(X1))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_1945, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X3,k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2))))=k1_algstr_0(X1,k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),X2)),k4_rlvect_1(X1,X3))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2)),u1_struct_0(X1))|~m2_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2290, negated_conjecture, (m2_finseq_1(esk2_0,X1)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_2364, negated_conjecture, (v1_finseq_1(k3_finseq_5(esk2_0)))).#trainneg
cnf(i_0_2161, plain, (k4_finseq_5(X1,X2)=k3_finseq_5(X4)|~v1_xboole_0(X3)|~v1_xboole_0(X5)|~m1_finseq_1(X2,X5)|~m1_finseq_1(X4,X3)|~m2_finseq_1(k3_finseq_5(X4),X1))).#trainneg
cnf(i_0_546, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_1133, plain, (v3_ordinal1(esk21_0)|~v1_card_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_2302, plain, (v1_valued_0(X1)|~v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_547, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_548, plain, (v1_finset_1(X1)|~v5_finset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_2388, plain, (k1_algstr_0(X1,k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),X3)),k4_rlvect_1(X1,X2))=k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3)),X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~m1_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2340, plain, (k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,k9_finseq_1(X3)))=k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X3)),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X3)),u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_2291, negated_conjecture, (m2_finseq_1(esk2_0,X1)|~v1_xboole_0(k3_finseq_5(esk2_0))|~v1_relat_1(k3_finseq_5(esk2_0)))).#trainneg
cnf(i_0_1463, plain, (v1_xxreal_0(esk13_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_549, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_1338, plain, (v1_xxreal_0(esk21_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_2430, negated_conjecture, (v3_valued_0(esk2_0)|~v3_membered(X1)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_550, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_551, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_2504, plain, (k1_algstr_0(X1,k4_rlvect_1(X1,k9_finseq_1(X2)),k4_rlvect_1(X1,X3))=k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),X3))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),u1_struct_0(X1))|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_2416, plain, (k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),X2)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))|m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k4_finseq_5(u1_struct_0(X1),X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_2320, plain, (k4_rlvect_1(X1,k3_finseq_5(k9_finseq_1(k1_algstr_0(X1,X2,X3))))=k4_rlvect_1(X1,k9_finseq_1(k1_algstr_0(X1,X2,X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1)))).#trainneg
cnf(i_0_2607, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1)))))!=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_relat_1(k3_finseq_5(k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1)))))|~m2_finseq_1(k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))),u1_struct_0(X1)))).#trainneg
cnf(i_0_2584, negated_conjecture, (v3_valued_0(esk2_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_2599, negated_conjecture, (k4_rlvect_1(esk1_0,k3_finseq_5(esk2_0))!=k4_rlvect_1(esk1_0,esk2_0)|~v1_relat_1(k3_finseq_5(esk2_0)))).#trainneg
cnf(i_0_553, plain, (v1_finset_1(X1)|~v1_finset_1(k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_2432, negated_conjecture, (v1_valued_0(esk2_0)|~v1_membered(X1)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_2433, negated_conjecture, (v2_valued_0(esk2_0)|~v2_membered(X1)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_554, plain, (v3_membered(X1)|~v3_membered(k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_555, plain, (v1_membered(X1)|~v1_membered(k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_2614, plain, (k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),X2))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_relat_1(k3_finseq_5(k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),X2))))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),X2)),u1_struct_0(X1)))).#trainneg
cnf(i_0_2612, plain, (k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_relat_1(k3_finseq_5(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2))))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2)),u1_struct_0(X1)))).#trainneg
cnf(i_0_2613, plain, (k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_relat_1(k3_finseq_5(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2))))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),u1_struct_0(X1)))).#trainneg
cnf(i_0_2712, plain, (k4_rlvect_1(X1,k3_finseq_5(k9_finseq_1(k1_algstr_0(X1,X2,X3))))=k4_rlvect_1(X1,k9_finseq_1(k1_algstr_0(X1,X2,X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~l1_algstr_0(X1))).#trainneg
cnf(i_0_2647, plain, (k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),X3))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k9_finseq_1(X2),X3))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m2_finseq_1(X3,u1_struct_0(X1))|~m2_finseq_1(k9_finseq_1(X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~m1_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),u1_struct_0(X1))|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_2063, plain, (k4_rlvect_1(X1,k3_finseq_5(k3_pre_poly(u1_struct_0(X1),esk5_1(u1_struct_0(X1)))))=k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),esk5_1(u1_struct_0(X1))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_2339, plain, (k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3)),X2))=k4_rlvect_1(X1,k7_finseq_1(X2,k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3)),u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_2767, plain, (k4_rlvect_1(X1,k3_finseq_5(X2))=k4_rlvect_1(X1,X2)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~v1_relat_1(k3_finseq_5(X2))|~m1_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_2726, negated_conjecture, (v1_valued_0(esk2_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_2733, negated_conjecture, (v2_valued_0(esk2_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_1556, plain, (v1_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_556, plain, (v2_membered(X1)|~v2_membered(k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_633, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk13_0))).#trainneg
cnf(i_0_2421, negated_conjecture, (m1_subset_1(k4_rlvect_1(X1,esk2_0),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_1978, plain, (v2_valued_0(X1)|~v2_membered(X2)|~v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_564, plain, (v5_relat_1(esk5_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_2816, plain, (k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),X3))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k9_finseq_1(X2),X3))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),u1_struct_0(X1)))).#trainneg
cnf(i_0_2596, plain, (k4_rlvect_1(X1,k3_finseq_5(esk5_1(k3_finseq_2(u1_struct_0(X1)))))=k4_rlvect_1(X1,esk5_1(k3_finseq_2(u1_struct_0(X1))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_relat_1(k3_finseq_5(esk5_1(k3_finseq_2(u1_struct_0(X1))))))).#trainneg
cnf(i_0_1895, plain, (k4_finseq_5(X1,k3_finseq_5(X2))=X2|~v1_finseq_1(k3_finseq_5(X2))|~v1_funct_1(k3_finseq_5(X2))|~v1_relat_1(k3_finseq_5(X2))|~m2_finseq_1(X2,X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_2715, plain, (v1_finset_1(X1)|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2363, plain, (v1_finseq_1(k3_finseq_5(esk5_1(k3_finseq_2(X1)))))).#trainneg
cnf(i_0_1557, plain, (v1_xcmplx_0(esk13_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_571, plain, (v4_relat_1(esk5_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_637, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_1967, plain, (m1_subset_1(k4_rlvect_1(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v1_xboole_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_2061, plain, (m2_finseq_1(k4_finseq_5(X1,X2),X3)|~v1_xboole_0(X3)|~v1_xboole_0(X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_572, plain, (v4_relat_1(X1,k5_numbers)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_575, plain, (v3_valued_0(esk18_2(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_2245, plain, (k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),u1_struct_0(X1)))).#trainneg
cnf(i_0_2757, plain, (k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),k3_finseq_5(X2))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k3_finseq_5(X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~v1_relat_1(k3_finseq_5(k4_finseq_5(u1_struct_0(X1),k3_finseq_5(X2))))|~m1_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_finseq_5(X2)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_1948, plain, (k1_algstr_0(X1,k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2))),k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),X3)))=k1_algstr_0(X1,k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),X2)),k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3)),u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2892, plain, (k4_rlvect_1(X1,k3_finseq_5(esk5_1(k3_finseq_2(u1_struct_0(X1)))))=k4_rlvect_1(X1,esk5_1(k3_finseq_2(u1_struct_0(X1))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_2868, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|~v1_relat_1(k3_finseq_5(X2))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2918, negated_conjecture, (k4_finseq_5(X1,esk2_0)=k3_finseq_5(esk2_0)|~m2_finseq_1(k3_finseq_5(esk2_0),X1))).#trainneg
cnf(i_0_780, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2971, negated_conjecture, (v4_relat_1(esk2_0,k5_numbers))).#trainneg
cnf(i_0_579, plain, (v3_valued_0(esk5_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v3_membered(X2))).#trainneg
cnf(i_0_2879, plain, (v2_valued_0(X1)|~v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_586, plain, (v3_valued_0(esk18_2(X1,X2))|~v3_valued_0(k2_zfmisc_1(X1,X2))|~v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_590, plain, (v3_valued_0(esk9_1(X1))|~v3_valued_0(k2_zfmisc_1(k5_numbers,X1))|~v1_relat_1(k2_zfmisc_1(k5_numbers,X1)))).#trainneg
cnf(i_0_2740, plain, (v1_membered(X1)|~m2_finseq_1(X1,X2)|~v3_membered(k2_zfmisc_1(k5_numbers,X2)))).#trainneg
cnf(i_0_1870, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k9_finseq_1(k4_rlvect_1(X1,u1_algstr_0(X1)))))=k4_rlvect_1(X1,k9_finseq_1(k4_rlvect_1(X1,u1_algstr_0(X1))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_finseq_1(u1_algstr_0(X1)))).#trainneg
cnf(i_0_1935, plain, (k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(X2,X3)))=k4_rlvect_1(X1,k7_finseq_1(X2,X3))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(k7_finseq_1(X2,X3),u1_struct_0(X1))|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_1952, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),X3))=k1_algstr_0(X1,k4_rlvect_1(X1,k9_finseq_1(X2)),k4_rlvect_1(X1,X3))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),u1_struct_0(X1))|~m2_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_1898, plain, (k4_finseq_5(X1,X2)=k1_xboole_0|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2690, plain, (k4_rlvect_1(X1,k3_finseq_5(X2))=k4_rlvect_1(X1,X2)|m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_2988, plain, (k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),u1_struct_0(X1)))).#trainneg
cnf(i_0_1924, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),esk5_1(k3_finseq_2(u1_struct_0(X1)))),X2))=k1_algstr_0(X1,k4_rlvect_1(X1,esk5_1(k3_finseq_2(u1_struct_0(X1)))),k4_rlvect_1(X1,X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),esk5_1(k3_finseq_2(u1_struct_0(X1)))),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_781, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2970, plain, (v4_relat_1(esk5_1(k3_finseq_2(X1)),k5_numbers))).#trainneg
cnf(i_0_1661, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_591, plain, (v3_valued_0(esk5_1(k1_zfmisc_1(X1)))|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1660, plain, (v1_xcmplx_0(esk21_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_2933, plain, (v1_finset_1(k4_finseq_5(X1,X2))|~v1_xboole_0(X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_593, plain, (v3_valued_0(X1)|~v3_valued_0(k2_zfmisc_1(k5_numbers,X2))|~v1_relat_1(k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_595, plain, (v1_valued_0(esk18_2(X1,X2))|~v1_valued_0(k2_zfmisc_1(X1,X2))|~v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_3033, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),X2)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))|v2_struct_0(X1)|~v1_relat_1(k3_finseq_5(k4_finseq_5(u1_struct_0(X1),X2)))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3030, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k3_finseq_5(X2)))=k4_rlvect_1(X1,k3_finseq_5(X2))|v2_struct_0(X1)|~v1_relat_1(k3_finseq_5(k3_finseq_5(X2)))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_finseq_1(k3_finseq_5(X2))|~v1_funct_1(k3_finseq_5(X2))|~v1_relat_1(k3_finseq_5(X2))|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1846, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),k8_finseq_1(u1_struct_0(X1),X2,X3))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k8_finseq_1(u1_struct_0(X1),X2,X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2155, plain, (k4_rlvect_1(X1,k3_finseq_5(k3_pre_poly(u1_struct_0(X1),k1_binop_1(u1_algstr_0(X1),X2,X3))))=k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),k1_binop_1(u1_algstr_0(X1),X2,X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~l1_algstr_0(X1))).#trainneg
cnf(i_0_3032, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k9_finseq_1(X2)))=k4_rlvect_1(X1,k9_finseq_1(X2))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~v1_relat_1(k3_finseq_5(k9_finseq_1(X2)))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1848, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),k7_finseq_1(X2,X3))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k7_finseq_1(X2,X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_783, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_647, plain, (v1_xboole_0(esk18_2(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2451, plain, (v3_ordinal1(esk21_0)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_599, plain, (v1_valued_0(esk9_1(X1))|~v1_valued_0(k2_zfmisc_1(k5_numbers,X1))|~v1_relat_1(k2_zfmisc_1(k5_numbers,X1)))).#trainneg
cnf(i_0_2452, plain, (v3_ordinal1(esk21_0)|~v7_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_600, plain, (v1_valued_0(esk5_1(k1_zfmisc_1(X1)))|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_602, plain, (v1_valued_0(X1)|~v1_valued_0(k2_zfmisc_1(k5_numbers,X2))|~v1_relat_1(k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_2251, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k7_finseq_1(X2,X3)))=k4_rlvect_1(X1,k7_finseq_1(X2,X3))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m2_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3226, plain, (k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),X2)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_finseq_1(k4_finseq_5(u1_struct_0(X1),X2))|~v1_funct_1(k4_finseq_5(u1_struct_0(X1),X2))|~v1_relat_1(k4_finseq_5(u1_struct_0(X1),X2)))).#trainneg
cnf(i_0_3284, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k9_finseq_1(X2)))=k4_rlvect_1(X1,k9_finseq_1(X2))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_3210, plain, (v3_valued_0(X1)|~m2_finseq_1(X1,X2)|~v5_relat_1(k2_zfmisc_1(k5_numbers,X2),k1_numbers))).#trainneg
cnf(i_0_3208, plain, (v3_valued_0(X1)|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_654, plain, (v1_valued_0(esk18_2(X1,X2))|~v1_membered(X2))).#trainneg
cnf(i_0_3323, plain, (v1_valued_0(X1)|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_611, plain, (v1_relat_1(esk19_1(X1))|v2_struct_0(X1)|~v4_funct_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1874, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k9_finseq_1(k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2)))))=k4_rlvect_1(X1,k9_finseq_1(k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1918, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),X3))=k1_algstr_0(X1,k4_rlvect_1(X1,X3),k4_rlvect_1(X1,k9_finseq_1(X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),u1_struct_0(X1))|~m2_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_2871, plain, (v2_membered(X1)|~m2_finseq_1(X1,X2)|~v3_membered(k2_zfmisc_1(k5_numbers,X2)))).#trainneg
cnf(i_0_1832, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2)))=k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),X2))|m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2338, plain, (k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3)),X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3),X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(k3_pre_poly(u1_struct_0(X1),X3),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3)),u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_3341, plain, (k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),X2)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3446, plain, (k1_algstr_0(X1,k4_rlvect_1(X1,X3),k4_rlvect_1(X1,k9_finseq_1(X2)))=k4_rlvect_1(X1,k7_finseq_1(X3,k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2))))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),u1_struct_0(X1))|~m1_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),u1_struct_0(X1))|~m2_finseq_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_3359, plain, (v3_valued_0(X1)|~m2_finseq_1(X1,k1_numbers))).#trainneg
cnf(i_0_612, plain, (v1_funct_1(esk19_1(X1))|v2_struct_0(X1)|~v4_funct_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_613, plain, (v7_ordinal1(esk19_1(X1))|v2_struct_0(X1)|~v6_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_614, plain, (v3_ordinal1(esk19_1(X1))|v2_struct_0(X1)|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_3350, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k9_finseq_1(k4_rlvect_1(X1,u1_algstr_0(X1)))))=k4_rlvect_1(X1,k9_finseq_1(k4_rlvect_1(X1,u1_algstr_0(X1))))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_finseq_1(u1_algstr_0(X1)))).#trainneg
cnf(i_0_3498, plain, (k4_rlvect_1(X1,k7_finseq_1(X3,k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2))))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2),X3))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2)),u1_struct_0(X1))|~m1_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1888, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X3))))=k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,k9_finseq_1(X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X3)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_3508, plain, (k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),esk5_1(k3_finseq_2(u1_struct_0(X1))))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk5_1(k3_finseq_2(u1_struct_0(X1)))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_3514, plain, (k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),k3_finseq_5(X2))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k3_finseq_5(X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3518, plain, (k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),X2))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3513, plain, (k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),k8_finseq_1(u1_struct_0(X1),X2,X3))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k8_finseq_1(u1_struct_0(X1),X2,X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3512, plain, (k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3520, plain, (k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),X2)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_xboole_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_661, plain, (v2_valued_0(esk18_2(X1,X2))|~v2_membered(X2))).#trainneg
cnf(i_0_615, plain, (v1_finset_1(esk19_1(X1))|v2_struct_0(X1)|~v5_finset_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_3572, plain, (v3_valued_0(k5_numbers)|~v1_relat_1(k1_numbers)|~m2_finseq_1(k1_numbers,k1_numbers))).#trainneg
cnf(i_0_1762, plain, (v3_membered(esk21_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2060, plain, (v1_membered(esk21_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3370, plain, (v3_valued_0(k4_finseq_5(X1,X2))|~v1_xboole_0(X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_616, plain, (v1_xxreal_0(esk19_1(X1))|v2_struct_0(X1)|~v2_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_617, plain, (v1_xreal_0(esk19_1(X1))|v2_struct_0(X1)|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_3517, plain, (k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3516, plain, (k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),k7_finseq_1(X2,X3))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k7_finseq_1(X2,X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3351, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k9_finseq_1(k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),X2)))))=k4_rlvect_1(X1,k9_finseq_1(k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),X2))))|v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3699, plain, (k4_rlvect_1(X1,k3_finseq_5(X2))=k4_rlvect_1(X1,X2)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_3668, plain, (k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),k3_finseq_5(X2))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k3_finseq_5(X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_784, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_791, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~v1_membered(X2))).#trainneg
cnf(i_0_2169, plain, (v2_membered(esk21_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2135, negated_conjecture, (esk2_0=k1_xboole_0|~v1_xboole_0(k3_finseq_5(esk2_0))|~v1_relat_1(k3_finseq_5(esk2_0)))).#trainneg
cnf(i_0_618, plain, (v1_xcmplx_0(esk19_1(X1))|v2_struct_0(X1)|~v1_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_3322, plain, (v3_ordinal1(esk21_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)),k5_numbers))).#trainneg
cnf(i_0_2136, negated_conjecture, (v1_xboole_0(esk2_0)|~v1_xboole_0(k3_finseq_5(esk2_0))|~v1_relat_1(k3_finseq_5(esk2_0)))).#trainneg
cnf(i_0_1515, plain, (v1_xreal_0(esk21_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3383, plain, (v1_valued_0(k4_finseq_5(X1,X2))|~v1_xboole_0(X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_620, plain, (v1_finset_1(esk19_1(X1))|v2_struct_0(X1)|~v1_finset_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_621, plain, (v3_membered(esk19_1(X1))|v2_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_3785, plain, (k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),k7_finseq_1(X2,X3))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k7_finseq_1(X2,X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_1709, negated_conjecture, (m1_subset_1(esk4_1(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3345, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k9_finseq_1(k3_rlvect_1(X1,X2,X3))))=k4_rlvect_1(X1,k9_finseq_1(k3_rlvect_1(X1,X2,X3)))|v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_785, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_792, plain, (v2_valued_0(k2_zfmisc_1(X1,X2))|~v2_membered(X2))).#trainneg
cnf(i_0_2920, plain, (k4_finseq_5(X1,k7_finseq_1(k9_finseq_1(X3),X2))=k7_finseq_1(k3_finseq_5(X2),k9_finseq_1(X3))|~v1_funct_1(k7_finseq_1(k9_finseq_1(X3),X2))|~v1_relat_1(k7_finseq_1(k9_finseq_1(X3),X2))|~m2_finseq_1(k7_finseq_1(k3_finseq_5(X2),k9_finseq_1(X3)),X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_622, plain, (v1_membered(esk19_1(X1))|v2_struct_0(X1)|~v1_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_2581, plain, (v1_xxreal_0(esk21_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3196, plain, (v1_xcmplx_0(esk21_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_623, plain, (v2_membered(esk19_1(X1))|v2_struct_0(X1)|~v2_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_624, plain, (v1_funct_1(esk19_1(X1))|v2_struct_0(X1)|~v1_funct_1(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_2393, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X3))))=k1_algstr_0(X1,k4_rlvect_1(X1,k9_finseq_1(X3)),k4_rlvect_1(X1,X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X3)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_1769, plain, (v1_funct_1(k3_finseq_5(X1))|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_3325, plain, (v1_valued_0(X1)|~m2_finseq_1(X1,X2)|~v3_valued_0(k2_zfmisc_1(k5_numbers,X2)))).#trainneg
cnf(i_0_1862, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k9_finseq_1(esk5_1(u1_struct_0(X1)))))=k4_rlvect_1(X1,k9_finseq_1(esk5_1(u1_struct_0(X1))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_3029, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k3_finseq_5(X2)))=k4_rlvect_1(X1,k3_finseq_5(X2))|v2_struct_0(X1)|~v1_relat_1(k3_finseq_5(k3_finseq_5(X2)))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3193, plain, (v3_ordinal1(X1)|~v1_card_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1176, plain, (v1_funct_1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_625, plain, (v3_valued_0(esk19_1(X1))|v2_struct_0(X1)|~v3_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_3318, plain, (v7_ordinal1(esk21_0)|~v1_finset_1(esk21_0)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_626, plain, (v1_valued_0(esk19_1(X1))|v2_struct_0(X1)|~v1_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_629, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k5_numbers))|~r2_hidden(X1,esk21_0))).#trainneg
cnf(i_0_3802, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3922, plain, (v1_valued_0(X1)|~m2_finseq_1(X1,X2)|~v3_membered(X2))).#trainneg
cnf(i_0_3916, plain, (v1_funct_1(k3_finseq_5(X1))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_3839, plain, (k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),k7_finseq_1(X2,X3))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k7_finseq_1(X2,X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m2_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3953, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k3_finseq_5(X2)))=k4_rlvect_1(X1,k3_finseq_5(X2))|v2_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_finseq_1(k3_finseq_5(X2))|~v1_funct_1(k3_finseq_5(X2))|~v1_relat_1(k3_finseq_5(X2)))).#trainneg
cnf(i_0_3945, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~v1_relat_1(k3_finseq_5(esk2_0))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_3923, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),k9_finseq_1(esk5_1(u1_struct_0(X1)))),X2))=k1_algstr_0(X1,k4_rlvect_1(X1,k9_finseq_1(esk5_1(u1_struct_0(X1)))),k4_rlvect_1(X1,X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(esk5_1(u1_struct_0(X1)))),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_3950, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(X2))=k4_rlvect_1(X1,k1_xboole_0)|v2_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_xboole_0(k3_finseq_5(X2))|~v1_relat_1(k3_finseq_5(X2)))).#trainneg
cnf(i_0_3954, plain, (v3_ordinal1(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3842, negated_conjecture, (v7_ordinal1(esk4_1(esk1_0))|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_631, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X1,esk18_2(X2,X3)))).#trainneg
cnf(i_0_632, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,k5_numbers))|~r2_hidden(X1,esk16_1(X2)))).#trainneg
cnf(i_0_635, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|~r2_hidden(X1,esk9_1(X2)))).#trainneg
cnf(i_0_3479, plain, (k4_rlvect_1(X1,k3_finseq_5(k3_pre_poly(u1_struct_0(X1),X2)))=k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),X2))|v1_xboole_0(u1_struct_0(X1))|m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_pre_poly(u1_struct_0(X1),X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_3838, plain, (k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),k7_finseq_1(k1_xboole_0,X2))))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k7_finseq_1(k1_xboole_0,X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_3966, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k8_finseq_1(u1_struct_0(X1),X2,X3)))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3969, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(X2,X3)))=k4_rlvect_1(X1,k7_finseq_1(X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3313, plain, (v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3988, negated_conjecture, (v1_funct_1(k3_finseq_5(esk2_0)))).#trainneg
cnf(i_0_3037, negated_conjecture, (k4_finseq_5(X1,esk2_0)=k3_finseq_5(esk2_0)|~v1_xboole_0(k3_finseq_5(esk2_0))|~v1_relat_1(k3_finseq_5(esk2_0)))).#trainneg
cnf(i_0_636, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk5_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_4070, plain, (v3_ordinal1(esk21_0)|~v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)))))).#trainneg
cnf(i_0_3984, plain, (v1_valued_0(k4_finseq_5(X1,X2))|~v3_membered(X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_638, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|~r2_hidden(X1,X3)|~m2_finseq_1(X3,X2))).#trainneg
cnf(i_0_644, plain, (v1_xboole_0(esk5_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_4014, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k3_finseq_5(X2)))=k4_rlvect_1(X1,k3_finseq_5(X2))|v2_struct_0(X1)|~v1_relat_1(k3_finseq_5(X2))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_4022, negated_conjecture, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k4_finseq_5(u1_struct_0(X1),esk3_1(X1))))=k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,esk3_1(X1)))|v2_struct_0(X1)|k4_rlvect_1(X1,k3_finseq_5(esk2_0))=k4_rlvect_1(X1,esk2_0)|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),esk3_1(X1)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_4034, negated_conjecture, (k4_rlvect_1(X1,k3_finseq_5(esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_relat_1(k3_finseq_5(esk3_1(X1)))|~m2_finseq_1(esk3_1(X1),u1_struct_0(X1)))).#trainneg
cnf(i_0_4030, negated_conjecture, (k4_rlvect_1(X1,k3_finseq_5(esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1)))).#trainneg
cnf(i_0_4016, negated_conjecture, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k4_finseq_5(u1_struct_0(X1),esk3_1(X1))))=k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k3_finseq_5(esk2_0))=k4_rlvect_1(X1,esk2_0)|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),esk3_1(X1)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1)))).#trainneg
cnf(i_0_4107, plain, (k4_rlvect_1(X1,k3_finseq_5(k8_finseq_1(u1_struct_0(X1),X2,X3)))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_4109, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(X2,X3)))=k4_rlvect_1(X1,k7_finseq_1(X2,X3))|v2_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_3967, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k3_finseq_5(X2)))=k4_rlvect_1(X1,k3_finseq_5(X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_4023, negated_conjecture, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),esk3_1(X1)),X2))=k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,esk3_1(X1)))|v2_struct_0(X1)|k4_rlvect_1(X1,k3_finseq_5(esk2_0))=k4_rlvect_1(X1,esk2_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),esk3_1(X1)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1)))).#trainneg
cnf(i_0_4015, negated_conjecture, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),esk3_1(X1)),X2))=k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k3_finseq_5(esk2_0))=k4_rlvect_1(X1,esk2_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),esk3_1(X1)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_2691, plain, (k4_rlvect_1(X1,k3_finseq_5(X2))=k4_rlvect_1(X1,X2)|m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~m1_finseq_1(k4_finseq_5(u1_struct_0(X1),X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_3867, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3987, plain, (v1_funct_1(k3_finseq_5(esk5_1(k3_finseq_2(X1)))))).#trainneg
cnf(i_0_3843, negated_conjecture, (v3_ordinal1(esk4_1(esk1_0))|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3106, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))=k4_rlvect_1(X1,X2)|v2_struct_0(X1)|k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))))!=k4_rlvect_1(X1,k1_xboole_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_xboole_0(k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1))))|~v1_relat_1(k7_finseq_1(esk3_1(X1),k9_finseq_1(esk4_1(X1)))))).#trainneg
cnf(i_0_651, plain, (v1_xboole_0(esk5_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2952, plain, (v2_valued_0(k4_finseq_5(X1,X2))|~v2_membered(X3)|~v1_xboole_0(X3)|~v1_xboole_0(X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_658, plain, (v1_valued_0(esk5_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_membered(X2))).#trainneg
cnf(i_0_665, plain, (v2_valued_0(esk5_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v2_membered(X2))).#trainneg
cnf(i_0_3179, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k4_finseq_5(u1_struct_0(X1),esk5_1(k3_finseq_2(u1_struct_0(X1))))))=k1_algstr_0(X1,k4_rlvect_1(X1,esk5_1(k3_finseq_2(u1_struct_0(X1)))),k4_rlvect_1(X1,X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),esk5_1(k3_finseq_2(u1_struct_0(X1)))),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_4245, plain, (k4_rlvect_1(X1,k3_finseq_5(k8_finseq_1(u1_struct_0(X1),X2,X3)))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m2_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_4247, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(X2,X3)))=k4_rlvect_1(X1,k7_finseq_1(X2,X3))|v2_struct_0(X1)|~m2_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_4249, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k3_finseq_5(X2)))=k4_rlvect_1(X1,k3_finseq_5(X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3904, plain, (k1_algstr_0(X1,k4_rlvect_1(X1,k9_finseq_1(X3)),k4_rlvect_1(X1,X2))=k4_rlvect_1(X1,k7_finseq_1(X2,k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X3))))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X3)),u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X3)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2890, plain, (k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),X3))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k9_finseq_1(X2),X3))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k9_finseq_1(X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_3530, plain, (k4_rlvect_1(X1,k7_finseq_1(X2,k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X3))))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k9_finseq_1(X3)))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m2_finseq_1(k9_finseq_1(X3),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~m1_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X3)),u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_3955, plain, (v3_ordinal1(X1)|~v7_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1419, plain, (v3_valued_0(esk5_1(k3_finseq_2(X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_3844, negated_conjecture, (v1_finset_1(esk4_1(esk1_0))|~v5_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2128, negated_conjecture, (k4_finseq_5(X1,esk2_0)=k3_finseq_5(esk2_0)|~m1_finseq_1(k3_finseq_5(esk2_0),X1))).#trainneg
cnf(i_0_675, plain, (v1_funct_1(k3_finseq_5(X1))|~v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3840, negated_conjecture, (v1_relat_1(esk4_1(esk1_0))|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3841, negated_conjecture, (v1_funct_1(esk4_1(esk1_0))|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4345, plain, (v2_valued_0(k4_finseq_5(X1,X2))|~v1_xboole_0(X3)|~v1_xboole_0(X1)|~m1_finseq_1(X2,X1)|~v3_membered(X3))).#trainneg
cnf(i_0_676, plain, (X1=k3_finseq_5(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_678, plain, (v2_funct_1(X1)|~v2_funct_1(k3_finseq_5(X1))|~v1_finseq_1(k3_finseq_5(X1))|~v1_funct_1(k3_finseq_5(X1))|~v1_relat_1(k3_finseq_5(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4394, plain, (k4_rlvect_1(X1,k3_finseq_5(k8_finseq_1(u1_struct_0(X1),X2,k4_finseq_5(u1_struct_0(X1),X3))))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k4_finseq_5(u1_struct_0(X1),X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_4388, plain, (k4_rlvect_1(X1,k3_finseq_5(k8_finseq_1(u1_struct_0(X1),X2,k3_pre_poly(u1_struct_0(X1),X3))))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k3_pre_poly(u1_struct_0(X1),X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_4407, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(X2,k4_finseq_5(u1_struct_0(X1),X3))))=k4_rlvect_1(X1,k7_finseq_1(X2,k4_finseq_5(u1_struct_0(X1),X3)))|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_4494, plain, (k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),X3))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k9_finseq_1(X2),X3))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(k9_finseq_1(X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_4501, plain, (k4_rlvect_1(X1,k7_finseq_1(X2,k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X3))))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k9_finseq_1(X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X3)),u1_struct_0(X1)))).#trainneg
cnf(i_0_4420, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k3_finseq_5(k4_finseq_5(u1_struct_0(X1),X2))))=k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),X2)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_4393, plain, (k4_rlvect_1(X1,k3_finseq_5(k8_finseq_1(u1_struct_0(X1),X2,k9_finseq_1(X3))))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k9_finseq_1(X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_4409, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(X2,X3)))=k4_rlvect_1(X1,k7_finseq_1(X2,X3))|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_xboole_0(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_4403, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(X2,k3_finseq_5(X3))))=k4_rlvect_1(X1,k7_finseq_1(X2,k3_finseq_5(X3)))|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_4401, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(X2,k3_pre_poly(u1_struct_0(X1),X3))))=k4_rlvect_1(X1,k7_finseq_1(X2,k3_pre_poly(u1_struct_0(X1),X3)))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_4404, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(X2,k3_finseq_5(X3))))=k4_rlvect_1(X1,k7_finseq_1(X2,k3_finseq_5(X3)))|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_finseq_1(k3_finseq_5(X3))|~v1_funct_1(k3_finseq_5(X3))|~v1_relat_1(k3_finseq_5(X3))|~m2_finseq_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_4390, plain, (k4_rlvect_1(X1,k3_finseq_5(k8_finseq_1(u1_struct_0(X1),X2,k3_finseq_5(X3))))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k3_finseq_5(X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_938, plain, (v3_membered(k1_xboole_0)|~v3_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1420, plain, (v1_xboole_0(esk5_1(k3_finseq_2(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3845, negated_conjecture, (v1_xxreal_0(esk4_1(esk1_0))|~v2_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_679, plain, (m1_subset_1(k4_rlvect_1(X1,esk18_2(X2,u1_struct_0(X1))),u1_struct_0(X1))|v2_struct_0(X1)|~v1_finseq_1(esk18_2(X2,u1_struct_0(X1)))|~l2_algstr_0(X1))).#trainneg
cnf(i_0_4502, plain, (v3_ordinal1(esk21_0)|~v7_ordinal1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)))))).#trainneg
cnf(i_0_3846, negated_conjecture, (v1_xreal_0(esk4_1(esk1_0))|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_687, plain, (v1_funct_1(k8_finseq_1(X1,X2,X3))|~m1_finseq_1(X3,X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_688, plain, (v1_finseq_1(k8_finseq_1(X1,X2,X3))|~m1_finseq_1(X3,X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_4662, plain, (k4_rlvect_1(X1,k3_finseq_5(k8_finseq_1(u1_struct_0(X1),X2,k3_finseq_5(X3))))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k3_finseq_5(X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_4591, plain, (k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(X2,k4_finseq_5(u1_struct_0(X1),X3))))=k4_rlvect_1(X1,k7_finseq_1(X2,k4_finseq_5(u1_struct_0(X1),X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_4593, plain, (k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),X3))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k9_finseq_1(X2),X3))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_4633, plain, (k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(X2,k3_finseq_5(X3))))=k4_rlvect_1(X1,k7_finseq_1(X2,k3_finseq_5(X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_4604, plain, (k4_rlvect_1(X1,k3_finseq_5(k3_finseq_5(k4_finseq_5(u1_struct_0(X1),X2))))=k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_4528, plain, (esk14_1(k1_xboole_0)=k3_finseq_5(X1)|~v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_4693, plain, (v1_finseq_1(k7_finseq_1(X2,X3))|~m1_finseq_1(X3,X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_939, plain, (v1_membered(k1_xboole_0)|~v1_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1421, plain, (v1_valued_0(esk5_1(k3_finseq_2(X1)))|~v1_membered(X1))).#trainneg
cnf(i_0_3847, negated_conjecture, (v1_xcmplx_0(esk4_1(esk1_0))|~v1_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4791, plain, (k4_finseq_5(X1,X2)=esk14_1(k1_xboole_0)|~m1_finseq_1(X2,X3)|~m2_finseq_1(esk14_1(k1_xboole_0),X1)|~v1_card_1(k1_xboole_0)|~v1_xboole_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_690, plain, (v1_funct_2(esk21_0,k5_numbers,X1)|~m1_subset_1(esk21_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_691, plain, (v1_funct_2(esk16_1(X1),X1,X2)|~m1_subset_1(esk16_1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_692, plain, (v1_funct_2(esk12_1(X1),X1,X2)|~m1_subset_1(esk12_1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4707, plain, (k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(esk5_1(k3_finseq_2(u1_struct_0(X1))),k4_finseq_5(u1_struct_0(X1),X2))))=k4_rlvect_1(X1,k7_finseq_1(esk5_1(k3_finseq_2(u1_struct_0(X1))),k4_finseq_5(u1_struct_0(X1),X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_4717, plain, (k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(k4_finseq_5(u1_struct_0(X1),X2),k4_finseq_5(u1_struct_0(X1),X3))))=k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),X2),k4_finseq_5(u1_struct_0(X1),X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m2_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_4735, plain, (k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(esk5_1(k3_finseq_2(u1_struct_0(X1))),k3_finseq_5(X2))))=k4_rlvect_1(X1,k7_finseq_1(esk5_1(k3_finseq_2(u1_struct_0(X1))),k3_finseq_5(X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_4561, plain, (k4_rlvect_1(X1,k3_finseq_5(k8_finseq_1(u1_struct_0(X1),X2,X3)))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(X3),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_4435, plain, (k4_rlvect_1(X1,k7_finseq_1(X3,k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2))))=k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),X3))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),u1_struct_0(X1))|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_4417, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|k4_rlvect_1(X1,k3_finseq_5(k3_finseq_5(k3_finseq_5(X2))))=k4_rlvect_1(X1,k3_finseq_5(k3_finseq_5(X2)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_finseq_1(k3_finseq_5(X2))|~v1_funct_1(k3_finseq_5(X2))|~v1_relat_1(k3_finseq_5(X2))|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2052, plain, (k4_finseq_5(X1,k3_finseq_5(X2))=X2|~m1_finseq_1(k3_finseq_5(X2),X3)|~m2_finseq_1(X2,X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_940, plain, (v2_membered(k1_xboole_0)|~v2_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1422, plain, (v2_valued_0(esk5_1(k3_finseq_2(X1)))|~v2_membered(X1))).#trainneg
cnf(i_0_4789, plain, (k4_finseq_5(X1,X2)=esk14_1(k1_xboole_0)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m2_finseq_1(esk14_1(k1_xboole_0),X1)|~v1_card_1(k1_xboole_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_693, plain, (v1_funct_2(esk10_1(X1),X1,X2)|~m1_subset_1(esk10_1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2982, plain, (v4_relat_1(X1,k5_numbers)|~v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4691, plain, (v3_ordinal1(esk21_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))),k5_numbers))).#trainneg
cnf(i_0_2959, plain, (k4_finseq_5(X1,X2)=k4_finseq_5(X3,X2)|~m1_finseq_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_4951, plain, (k4_finseq_5(X1,k3_finseq_5(X2))=X2|~m2_finseq_1(X2,X1)|~m2_finseq_1(k3_finseq_5(X2),X3))).#trainneg
cnf(i_0_4893, plain, (k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(X2,k4_finseq_5(u1_struct_0(X1),X3))))=k4_rlvect_1(X1,k7_finseq_1(X2,k4_finseq_5(u1_struct_0(X1),X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(k4_finseq_5(u1_struct_0(X1),X2),u1_struct_0(X1))|~m2_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_4923, plain, (k4_rlvect_1(X1,k3_finseq_5(k8_finseq_1(u1_struct_0(X1),X2,X3)))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(X3,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(k3_finseq_5(X3),u1_struct_0(X1)))).#trainneg
cnf(i_0_4891, plain, (k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(X2,k4_finseq_5(u1_struct_0(X1),X3))))=k4_rlvect_1(X1,k7_finseq_1(X2,k4_finseq_5(u1_struct_0(X1),X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(X2),u1_struct_0(X1))|~m2_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2314, plain, (k4_rlvect_1(X1,k3_finseq_5(k3_pre_poly(u1_struct_0(X1),k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3)))))=k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,X3))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k4_rlvect_1(X1,X3),u1_struct_0(X1))|~m1_subset_1(k4_rlvect_1(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X3,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2337, plain, (k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3)),X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k3_pre_poly(u1_struct_0(X1),X3)))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m2_finseq_1(k3_pre_poly(u1_struct_0(X1),X3),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~m1_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3)),u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_2923, plain, (v1_finset_1(esk5_1(k3_finseq_2(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_696, plain, (v7_ordinal1(u1_algstr_0(X1))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_3822, plain, (v7_ordinal1(esk21_0)|~v1_finset_1(esk21_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)),k5_numbers))).#trainneg
cnf(i_0_4540, negated_conjecture, (v2_funct_1(k3_finseq_5(esk2_0))|~v2_funct_1(esk2_0)|~v1_relat_1(k3_finseq_5(esk2_0)))).#trainneg
cnf(i_0_5034, plain, (m2_finseq_1(k4_finseq_5(X3,X2),X1)|~m1_finseq_1(X2,X1)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5036, plain, (v1_funct_1(k4_finseq_5(X3,X2))|~m1_finseq_1(X2,X1)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_697, plain, (v3_ordinal1(u1_algstr_0(X1))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_698, plain, (v1_finset_1(u1_algstr_0(X1))|~v5_finset_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_5180, plain, (k4_rlvect_1(X1,k3_finseq_5(k7_finseq_1(X2,k4_finseq_5(u1_struct_0(X1),X3))))=k4_rlvect_1(X1,k7_finseq_1(X2,k4_finseq_5(u1_struct_0(X1),X3)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_5244, plain, (k4_rlvect_1(X1,k7_finseq_1(X3,k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2))))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X3,k3_pre_poly(u1_struct_0(X1),X2)))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2)),u1_struct_0(X1))|~m1_finseq_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_5246, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X3,k3_pre_poly(u1_struct_0(X1),X2)))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2),X3))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X2)),u1_struct_0(X1))|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_5216, plain, (k4_rlvect_1(X1,k3_finseq_5(k3_pre_poly(u1_struct_0(X1),k4_rlvect_1(X1,k7_finseq_1(X2,X3)))))=k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),k4_rlvect_1(X1,k7_finseq_1(X2,X3))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m2_finseq_1(X3,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_5208, plain, (k4_rlvect_1(X1,k3_finseq_5(k3_pre_poly(u1_struct_0(X1),k1_algstr_0(X1,k4_rlvect_1(X1,X3),k4_rlvect_1(X1,k9_finseq_1(X2))))))=k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),k1_algstr_0(X1,k4_rlvect_1(X1,X3),k4_rlvect_1(X1,k9_finseq_1(X2)))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m2_finseq_1(X3,u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3360, plain, (v3_valued_0(esk5_1(k3_finseq_2(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5264, negated_conjecture, (v2_funct_1(k1_xboole_0)|~v2_funct_1(esk2_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_699, plain, (v1_xxreal_0(u1_algstr_0(X1))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_4329, negated_conjecture, (v7_ordinal1(esk4_1(esk1_0))|~v1_finset_1(esk4_1(esk1_0))|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5319, plain, (v1_funct_1(k4_finseq_5(X1,X2))|~v1_xboole_0(X1)|~v1_xboole_0(X3)|~m2_finseq_1(X2,X3))).#trainneg
cnf(i_0_5037, plain, (v1_finseq_1(k4_finseq_5(X3,X2))|~m1_finseq_1(X2,X1)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_700, plain, (v1_xreal_0(u1_algstr_0(X1))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_701, plain, (v1_xcmplx_0(u1_algstr_0(X1))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_5248, plain, (k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),X3))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X3,k9_finseq_1(X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),u1_struct_0(X1))|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_5210, plain, (k4_rlvect_1(X1,k3_finseq_5(k3_pre_poly(u1_struct_0(X1),k1_algstr_0(X1,k4_rlvect_1(X1,esk5_1(k3_finseq_2(u1_struct_0(X1)))),k4_rlvect_1(X1,X2)))))=k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),k1_algstr_0(X1,k4_rlvect_1(X1,esk5_1(k3_finseq_2(u1_struct_0(X1)))),k4_rlvect_1(X1,X2))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),esk5_1(k3_finseq_2(u1_struct_0(X1)))),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_5590, plain, (k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),X3))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X3,k9_finseq_1(X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),u1_struct_0(X1)))).#trainneg
cnf(i_0_5586, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X3,k9_finseq_1(X2)))=k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(k9_finseq_1(X2)),X3))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(k9_finseq_1(X2)),u1_struct_0(X1))|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_relat_1(k3_finseq_5(k9_finseq_1(X2)))|~m2_finseq_1(k9_finseq_1(X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_5625, plain, (k4_rlvect_1(X1,k3_finseq_5(k3_pre_poly(u1_struct_0(X1),k1_algstr_0(X1,k4_rlvect_1(X1,esk5_1(k3_finseq_2(u1_struct_0(X1)))),k1_algstr_0(X1,k4_rlvect_1(X1,esk5_1(k3_finseq_2(u1_struct_0(X1)))),k4_rlvect_1(X1,X2))))))=k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),k1_algstr_0(X1,k4_rlvect_1(X1,esk5_1(k3_finseq_2(u1_struct_0(X1)))),k1_algstr_0(X1,k4_rlvect_1(X1,esk5_1(k3_finseq_2(u1_struct_0(X1)))),k4_rlvect_1(X1,X2)))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),esk5_1(k3_finseq_2(u1_struct_0(X1)))),u1_struct_0(X1))|~m2_finseq_1(k8_finseq_1(u1_struct_0(X1),X2,k4_finseq_5(u1_struct_0(X1),esk5_1(k3_finseq_2(u1_struct_0(X1))))),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_4073, plain, (v3_ordinal1(X1)|~v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_3373, plain, (v1_valued_0(esk5_1(k3_finseq_2(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_703, plain, (v1_finset_1(u1_algstr_0(X1))|~v1_finset_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_2735, plain, (v3_membered(X1)|~v3_membered(k1_xboole_0)|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_5808, negated_conjecture, (v1_funct_1(esk3_1(esk1_0)))).#trainneg
cnf(i_0_5809, negated_conjecture, (v1_finseq_1(esk3_1(esk1_0)))).#trainneg
cnf(i_0_5558, negated_conjecture, (v1_funct_1(k4_finseq_5(X1,esk2_0))|~v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_704, plain, (v3_membered(u1_algstr_0(X1))|~v3_membered(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_705, plain, (v1_membered(u1_algstr_0(X1))|~v1_membered(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_5692, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X3,k9_finseq_1(X2)))=k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(k9_finseq_1(X2)),X3))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m2_finseq_1(k3_finseq_5(k9_finseq_1(X2)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_relat_1(k3_finseq_5(k9_finseq_1(X2)))|~m2_finseq_1(k9_finseq_1(X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_5697, plain, (k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),X3))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X3,k9_finseq_1(X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k9_finseq_1(X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_5816, negated_conjecture, (k3_finseq_5(k3_finseq_5(esk3_1(esk1_0)))=esk3_1(esk1_0))).#trainneg
cnf(i_0_5217, plain, (k4_rlvect_1(X1,k3_finseq_5(k3_pre_poly(u1_struct_0(X1),k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),X3))))))=k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),X3)))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k3_pre_poly(u1_struct_0(X1),X3)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_5884, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k9_finseq_1(X3)))=k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(k9_finseq_1(X3)),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_finseq_1(k3_finseq_5(k9_finseq_1(X3)))|~v1_funct_1(k3_finseq_5(k9_finseq_1(X3)))|~v1_relat_1(k3_finseq_5(k9_finseq_1(X3)))|~m2_finseq_1(k9_finseq_1(X3),u1_struct_0(X1)))).#trainneg
cnf(i_0_5885, plain, (k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),X3))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X3,k9_finseq_1(X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(k9_finseq_1(X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_5848, plain, (v1_finset_1(u1_algstr_0(X1))|~l1_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_5883, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k9_finseq_1(X3)))=k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(k9_finseq_1(X3)),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k9_finseq_1(X3),u1_struct_0(X1)))).#trainneg
cnf(i_0_4505, plain, (v3_ordinal1(X1)|~v7_ordinal1(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_3974, plain, (v1_valued_0(esk5_1(k3_finseq_2(X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_5810, negated_conjecture, (v1_relat_1(esk3_1(esk1_0)))).#trainneg
cnf(i_0_706, plain, (v2_membered(u1_algstr_0(X1))|~v2_membered(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_2738, plain, (v1_membered(X1)|~v1_membered(k1_xboole_0)|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_5802, plain, (v3_ordinal1(esk21_0)|~v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))))).#trainneg
cnf(i_0_5812, negated_conjecture, (v3_valued_0(esk3_1(esk1_0))|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5864, negated_conjecture, (v1_funct_1(k4_finseq_5(X1,esk2_0))|~v1_xboole_0(esk2_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5569, plain, (v1_funct_1(k4_finseq_5(X1,X2))|~v1_xboole_0(X1)|~v1_xboole_0(X3)|~v1_xboole_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_708, plain, (v3_valued_0(u1_algstr_0(X1))|~v3_valued_0(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)))|~v1_relat_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_5221, plain, (k4_rlvect_1(X1,k3_finseq_5(k3_pre_poly(u1_struct_0(X1),k1_algstr_0(X1,k4_rlvect_1(X1,k9_finseq_1(X3)),k4_rlvect_1(X1,X2)))))=k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),k1_algstr_0(X1,k4_rlvect_1(X1,k9_finseq_1(X3)),k4_rlvect_1(X1,X2))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X3)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_6025, plain, (k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),X3))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X3,k9_finseq_1(X2)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_6031, plain, (v2_membered(u1_algstr_0(X1))|~v2_membered(k1_xboole_0)|~l1_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_6026, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k9_finseq_1(X3)))=k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(k9_finseq_1(X3)),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(k9_finseq_1(X3),u1_struct_0(X1)))).#trainneg
cnf(i_0_6150, plain, (k4_rlvect_1(X1,k3_finseq_5(k3_pre_poly(u1_struct_0(X1),k1_algstr_0(X1,k4_rlvect_1(X1,k9_finseq_1(X2)),k4_rlvect_1(X1,k9_finseq_1(k4_rlvect_1(X1,u1_algstr_0(X1))))))))=k4_rlvect_1(X1,k3_pre_poly(u1_struct_0(X1),k1_algstr_0(X1,k4_rlvect_1(X1,k9_finseq_1(X2)),k4_rlvect_1(X1,k9_finseq_1(k4_rlvect_1(X1,u1_algstr_0(X1)))))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(X2)),u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(k4_rlvect_1(X1,u1_algstr_0(X1)))),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_finseq_1(u1_algstr_0(X1)))).#trainneg
cnf(i_0_4821, plain, (v1_membered(k1_xboole_0)|~v3_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_5811, negated_conjecture, (v4_relat_1(esk3_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_5894, negated_conjecture, (k4_finseq_5(X1,esk3_1(esk1_0))=k3_finseq_5(esk3_1(esk1_0))|~m2_finseq_1(esk3_1(esk1_0),X2)|~m2_finseq_1(k3_finseq_5(esk3_1(esk1_0)),X1))).#trainneg
cnf(i_0_709, plain, (v1_valued_0(u1_algstr_0(X1))|~v1_valued_0(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)))|~v1_relat_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_2869, plain, (v2_membered(X1)|~v2_membered(k1_xboole_0)|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_5814, negated_conjecture, (v1_valued_0(esk3_1(esk1_0))|~v1_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5815, negated_conjecture, (v2_valued_0(esk3_1(esk1_0))|~v2_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6054, plain, (v1_funct_1(k4_finseq_5(X1,X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_5571, plain, (v1_finseq_1(k4_finseq_5(X1,X2))|~v1_xboole_0(X1)|~v1_xboole_0(X3)|~m2_finseq_1(X2,X3))).#trainneg
cnf(i_0_710, plain, (m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2)),u1_struct_0(X2)))|~r2_hidden(X1,u1_algstr_0(X2))|~l1_algstr_0(X2))).#trainneg
cnf(i_0_712, plain, (v1_xboole_0(u1_algstr_0(X1))|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_6190, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k9_finseq_1(X3)))=k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(k9_finseq_1(X3)),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_6245, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),X2)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6177, plain, (k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(k3_rlvect_1(X1,X2,X3))),X4))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X4,k9_finseq_1(k3_rlvect_1(X1,X2,X3))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m2_finseq_1(X4,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_6178, plain, (k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(k1_algstr_0(X1,X2,X3))),X4))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X4,k9_finseq_1(k1_algstr_0(X1,X2,X3))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m2_finseq_1(X4,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_4954, plain, (v2_membered(k1_xboole_0)|~v3_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2424, negated_conjecture, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_5817, negated_conjecture, (v1_finseq_1(k3_finseq_5(esk3_1(esk1_0))))).#trainneg
cnf(i_0_6273, negated_conjecture, (k4_finseq_5(X1,esk3_1(esk1_0))=k1_xboole_0|~m2_finseq_1(esk3_1(esk1_0),X2)|~v1_xboole_0(esk3_1(esk1_0)))).#trainneg
cnf(i_0_713, plain, (v4_relat_1(u1_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_6302, negated_conjecture, (v1_finseq_1(k4_finseq_5(X1,esk2_0))|~v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_715, plain, (v3_valued_0(u1_algstr_0(X1))|~v3_membered(u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_716, plain, (v1_valued_0(u1_algstr_0(X1))|~v1_membered(u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_6342, plain, (k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),X2)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m2_finseq_1(k3_finseq_5(X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_2439, negated_conjecture, (k4_finseq_5(X1,X2)=esk2_0|~v1_xboole_0(X3)|~v1_xboole_0(X4)|~m1_finseq_1(X2,X4)|~m1_finseq_1(k3_finseq_5(esk2_0),X3)|~m2_finseq_1(esk2_0,X1))).#trainneg
cnf(i_0_3036, negated_conjecture, (k4_finseq_5(X1,esk2_0)=k1_xboole_0|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_5818, negated_conjecture, (v1_funct_1(k3_finseq_5(esk3_1(esk1_0))))).#trainneg
cnf(i_0_717, plain, (v2_valued_0(u1_algstr_0(X1))|~v2_membered(u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_5822, negated_conjecture, (v1_valued_0(esk3_1(esk1_0))|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6515, negated_conjecture, (v1_finseq_1(k4_finseq_5(X1,esk2_0))|~v1_xboole_0(esk2_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6313, plain, (v1_finseq_1(k4_finseq_5(X1,X2))|~v1_xboole_0(X1)|~v1_xboole_0(X3)|~v1_xboole_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_719, plain, (v1_relat_1(k3_finseq_5(k7_finseq_1(X2,k9_finseq_1(X1))))|~v1_finseq_1(k3_finseq_5(X2))|~v1_funct_1(k3_finseq_5(X2))|~v1_relat_1(k3_finseq_5(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_720, plain, (v1_funct_1(k3_finseq_5(k7_finseq_1(X2,k9_finseq_1(X1))))|~v1_finseq_1(k3_finseq_5(X2))|~v1_funct_1(k3_finseq_5(X2))|~v1_relat_1(k3_finseq_5(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6314, plain, (m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2)),u1_struct_0(X2)))|v1_xboole_0(u1_algstr_0(X2))|~l1_algstr_0(X2)|~m1_subset_1(X1,u1_algstr_0(X2)))).#trainneg
cnf(i_0_6604, negated_conjecture, (k4_rlvect_1(X1,k3_finseq_5(k4_finseq_5(u1_struct_0(X1),esk2_0)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))|k4_rlvect_1(X1,k3_finseq_5(esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v1_relat_1(k3_finseq_5(k4_finseq_5(u1_struct_0(X1),esk2_0)))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),esk2_0),u1_struct_0(X1)))).#trainneg
cnf(i_0_6355, negated_conjecture, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),esk3_1(X1)),X2))=k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,esk3_1(X1)))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),k4_finseq_5(u1_struct_0(X1),esk2_0)))=k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))|~m2_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(k4_finseq_5(u1_struct_0(X1),esk3_1(X1)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1)))).#trainneg
cnf(i_0_6840, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,esk3_1(X1)))=k4_rlvect_1(X1,k7_finseq_1(X2,k3_finseq_5(esk3_1(X1))))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_6652, negated_conjecture, (k4_finseq_5(X1,X2)=esk2_0|~v1_xboole_0(X3)|~v1_xboole_0(X4)|~m1_finseq_1(X2,X4)|~m2_finseq_1(esk2_0,X1)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_6500, negated_conjecture, (v1_finseq_1(esk14_1(k1_xboole_0))|~v1_card_1(k1_xboole_0)|~v1_xboole_0(esk3_1(esk1_0)))).#trainneg
cnf(i_0_721, plain, (v1_finseq_1(k3_finseq_5(k7_finseq_1(X2,k9_finseq_1(X1))))|~v1_finseq_1(k3_finseq_5(X2))|~v1_funct_1(k3_finseq_5(X2))|~v1_relat_1(k3_finseq_5(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6717, plain, (v1_finseq_1(k4_finseq_5(X1,X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_726, plain, (v1_partfun1(esk18_2(X1,X2),X1)|v1_xboole_0(X2))).#trainneg
cnf(i_0_7255, negated_conjecture, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k3_finseq_5(esk3_1(X1)),X2))=k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,esk3_1(X1)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_rlvect_1(X1,esk3_1(X1)),u1_struct_0(X1))|~m1_subset_1(k4_rlvect_1(X1,X2),u1_struct_0(X1))|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1))).#trainneg
cnf(i_0_7083, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,X2))=k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(esk3_1(X1)),X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m1_finseq_1(X2,u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_7662, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,esk3_1(X1)))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),esk3_1(X1),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_7193, negated_conjecture, (k4_finseq_5(X1,X2)=esk2_0|~v1_xboole_0(esk2_0)|~v1_xboole_0(X3)|~v1_xboole_0(X4)|~m2_finseq_1(X2,X4))).#trainneg
cnf(i_0_6274, negated_conjecture, (k4_finseq_5(X1,esk3_1(esk1_0))=esk14_1(k1_xboole_0)|~m2_finseq_1(esk14_1(k1_xboole_0),X1)|~m2_finseq_1(esk3_1(esk1_0),X2)|~v1_card_1(k1_xboole_0)|~v1_xboole_0(esk3_1(esk1_0)))).#trainneg
cnf(i_0_728, plain, (v1_partfun1(esk9_1(X1),k5_numbers)|v1_xboole_0(X1))).#trainneg
cnf(i_0_5555, negated_conjecture, (v7_ordinal1(esk4_1(esk1_0))|~v3_ordinal1(u1_struct_0(esk1_0))|~v5_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5038, plain, (v1_relat_1(k4_finseq_5(X3,X2))|~m1_finseq_1(X2,X1)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_729, plain, (v1_partfun1(u1_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(u1_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_736, plain, (v1_relat_1(k1_algstr_0(X1,X2,X3))|~v4_funct_1(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_7084, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,X2))=k4_rlvect_1(X1,k7_finseq_1(X2,k3_finseq_5(esk3_1(X1))))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_6642, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,esk3_1(X1)))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k3_finseq_5(esk3_1(X1))))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m2_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v1_relat_1(k3_finseq_5(esk3_1(X1)))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_8076, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k3_finseq_5(esk3_1(X1))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_rlvect_1(X1,esk3_1(X1)),u1_struct_0(X1))|~m1_subset_1(k4_rlvect_1(X1,X2),u1_struct_0(X1))|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1))).#trainneg
cnf(i_0_7889, negated_conjecture, (k4_finseq_5(X1,esk2_0)=esk2_0|~v1_xboole_0(esk2_0)|~v1_xboole_0(X2)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1090, negated_conjecture, (v1_xreal_0(k4_rlvect_1(esk1_0,esk2_0))|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5890, negated_conjecture, (k4_finseq_5(X1,esk3_1(esk1_0))=k3_finseq_5(esk3_1(esk1_0))|~v1_relat_1(esk3_1(esk1_0))|~m2_finseq_1(k3_finseq_5(esk3_1(esk1_0)),X1))).#trainneg
cnf(i_0_737, plain, (v1_funct_1(k1_algstr_0(X1,X2,X3))|~v4_funct_1(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1084, negated_conjecture, (v1_relat_1(k4_rlvect_1(esk1_0,esk2_0))|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1085, negated_conjecture, (v1_funct_1(k4_rlvect_1(esk1_0,esk2_0))|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7906, plain, (v1_relat_1(k4_finseq_5(X1,X2))|~v1_xboole_0(X1)|~v1_xboole_0(X3)|~m2_finseq_1(X2,X3))).#trainneg
cnf(i_0_7900, negated_conjecture, (k4_finseq_5(X1,X2)=esk2_0|~v1_xboole_0(esk2_0)|~v1_xboole_0(X3)|~v1_xboole_0(X4)|~v1_xboole_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_738, plain, (v7_ordinal1(k1_algstr_0(X1,X2,X3))|~v6_membered(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_739, plain, (v3_ordinal1(k1_algstr_0(X1,X2,X3))|~v3_ordinal1(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_8233, negated_conjecture, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k3_finseq_5(esk3_1(X1)),X2))=k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m2_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m1_subset_1(k4_rlvect_1(X1,esk3_1(X1)),u1_struct_0(X1)))).#trainneg
cnf(i_0_8064, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,esk3_1(X1)))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k3_finseq_5(esk3_1(X1)),X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1)))).#trainneg
cnf(i_0_8360, negated_conjecture, (k4_finseq_5(X1,esk2_0)=esk2_0|~v1_xboole_0(esk2_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_984, plain, (m1_subset_1(esk9_1(X1),k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1091, negated_conjecture, (v1_xcmplx_0(k4_rlvect_1(esk1_0,esk2_0))|~v1_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_740, plain, (v1_finset_1(k1_algstr_0(X1,X2,X3))|~v5_finset_1(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1086, negated_conjecture, (v7_ordinal1(k4_rlvect_1(esk1_0,esk2_0))|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1088, negated_conjecture, (v1_finset_1(k4_rlvect_1(esk1_0,esk2_0))|~v5_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8375, negated_conjecture, (v1_relat_1(k4_finseq_5(X1,esk2_0))|~v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_741, plain, (v1_xxreal_0(k1_algstr_0(X1,X2,X3))|~v2_membered(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_742, plain, (v1_xreal_0(k1_algstr_0(X1,X2,X3))|~v3_membered(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_6839, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,esk3_1(X1)))=k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(esk3_1(X1)),X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m1_finseq_1(X2,u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_6987, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,esk3_1(X1)))=k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(esk3_1(X1)),k3_finseq_5(esk3_1(X1))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|k4_rlvect_1(X1,k3_finseq_5(esk2_0))=k4_rlvect_1(X1,esk2_0)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_8808, negated_conjecture, (k4_finseq_5(X1,esk2_0)=esk2_0|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_8823, plain, (v1_finset_1(esk9_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1537, negated_conjecture, (v3_ordinal1(k4_rlvect_1(esk1_0,k1_xboole_0))|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6653, negated_conjecture, (k4_finseq_5(X1,X2)=esk2_0|~v1_xboole_0(X3)|~v1_xboole_0(X4)|~m1_finseq_1(X2,X4)|~m2_finseq_1(esk2_0,X1)|~m2_finseq_1(k3_finseq_5(esk2_0),X3))).#trainneg
cnf(i_0_743, plain, (v1_xcmplx_0(k1_algstr_0(X1,X2,X3))|~v1_membered(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1089, negated_conjecture, (v1_xxreal_0(k4_rlvect_1(esk1_0,esk2_0))|~v2_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1534, negated_conjecture, (v1_relat_1(k4_rlvect_1(esk1_0,k1_xboole_0))|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8869, negated_conjecture, (v1_relat_1(k4_finseq_5(X1,esk2_0))|~v1_xboole_0(esk2_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_8396, negated_conjecture, (k4_finseq_5(X1,X2)=esk2_0|~v1_xboole_0(esk2_0)|~v1_xboole_0(X3)|~v1_xboole_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_747, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(esk18_2(X1,X2)))).#trainneg
cnf(i_0_751, plain, (v1_relat_1(k3_rlvect_1(X1,X2,X3))|~v4_funct_1(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_9109, negated_conjecture, (k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(esk3_1(X1)),k3_finseq_5(esk3_1(X1))))=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),k3_finseq_5(esk3_1(X1))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|k4_rlvect_1(X1,k3_finseq_5(esk2_0))=k4_rlvect_1(X1,esk2_0)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_9135, negated_conjecture, (~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_9271, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,esk3_1(X1)))=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),k3_finseq_5(esk3_1(X1))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k3_finseq_5(esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_9108, negated_conjecture, (k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(esk3_1(X1)),k3_finseq_5(esk3_1(X1))))=k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(esk3_1(X1)),esk3_1(X1)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|k4_rlvect_1(X1,k3_finseq_5(esk2_0))=k4_rlvect_1(X1,esk2_0)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_9176, negated_conjecture, (k4_finseq_5(X1,X2)=esk2_0|~v1_xboole_0(X3)|~v1_xboole_0(X4)|~m1_finseq_1(X2,X4)|~m2_finseq_1(esk2_0,X1)|~v1_relat_1(k3_finseq_5(esk2_0))|~m2_finseq_1(esk2_0,X3))).#trainneg
cnf(i_0_752, plain, (v1_funct_1(k3_rlvect_1(X1,X2,X3))|~v4_funct_1(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_1535, negated_conjecture, (v1_funct_1(k4_rlvect_1(esk1_0,k1_xboole_0))|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1536, negated_conjecture, (v7_ordinal1(k4_rlvect_1(esk1_0,k1_xboole_0))|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_753, plain, (v7_ordinal1(k3_rlvect_1(X1,X2,X3))|~v6_membered(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_754, plain, (v3_ordinal1(k3_rlvect_1(X1,X2,X3))|~v3_ordinal1(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_9405, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k3_finseq_5(esk3_1(X1)),X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1)))).#trainneg
cnf(i_0_9339, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,esk3_1(X1)))=k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(esk3_1(X1)),esk3_1(X1)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k3_finseq_5(esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_1540, negated_conjecture, (v1_xreal_0(k4_rlvect_1(esk1_0,k1_xboole_0))|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_755, plain, (v1_finset_1(k3_rlvect_1(X1,X2,X3))|~v5_finset_1(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_1538, negated_conjecture, (v1_finset_1(k4_rlvect_1(esk1_0,k1_xboole_0))|~v5_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1539, negated_conjecture, (v1_xxreal_0(k4_rlvect_1(esk1_0,k1_xboole_0))|~v2_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_756, plain, (v1_xxreal_0(k3_rlvect_1(X1,X2,X3))|~v2_membered(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_757, plain, (v1_xreal_0(k3_rlvect_1(X1,X2,X3))|~v3_membered(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_9288, negated_conjecture, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),esk3_1(X1),esk3_1(X1)))=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),k3_finseq_5(esk3_1(X1))))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k3_finseq_5(esk2_0))=k4_rlvect_1(X1,esk2_0)|~m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_8828, plain, (v3_valued_0(esk9_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1541, negated_conjecture, (v1_xcmplx_0(k4_rlvect_1(esk1_0,k1_xboole_0))|~v1_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_758, plain, (v1_xcmplx_0(k3_rlvect_1(X1,X2,X3))|~v1_membered(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_9207, negated_conjecture, (v1_xxreal_0(k4_rlvect_1(esk1_0,esk2_0))|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_767, plain, (m1_subset_1(k5_binop_1(X1,esk18_2(k2_zfmisc_1(X1,X1),X1),X2,X3),X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_770, plain, (k5_binop_1(X1,esk18_2(k2_zfmisc_1(X1,X1),X1),X2,X3)=k1_binop_1(esk18_2(k2_zfmisc_1(X1,X1),X1),X2,X3)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_9804, negated_conjecture, (k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),k3_finseq_5(esk3_1(X1))))=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),esk3_1(X1)))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k3_finseq_5(esk2_0))=k4_rlvect_1(X1,esk2_0)|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1)))).#trainneg
cnf(i_0_9740, negated_conjecture, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),esk3_1(X1),esk3_1(X1)))=k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(esk3_1(X1)),esk3_1(X1)))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k3_finseq_5(esk2_0))=k4_rlvect_1(X1,esk2_0)|~m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1)))).#trainneg
cnf(i_0_9994, negated_conjecture, (k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(esk3_1(X1)),esk3_1(X1)))=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),esk3_1(X1)))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k3_finseq_5(esk2_0))=k4_rlvect_1(X1,esk2_0)|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1)))).#trainneg
cnf(i_0_9905, plain, (k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(k5_binop_1(u1_struct_0(X1),esk18_2(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)),X2,X3))),X4))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k9_finseq_1(k5_binop_1(u1_struct_0(X1),esk18_2(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)),X2,X3)),X4))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m2_finseq_1(X4,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_8829, plain, (v1_valued_0(esk9_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_9759, negated_conjecture, (v1_xxreal_0(k4_rlvect_1(esk1_0,k1_xboole_0))|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5163, plain, (k4_finseq_5(X1,esk14_1(k1_xboole_0))=X2|~m2_finseq_1(esk14_1(k1_xboole_0),X3)|~m2_finseq_1(X2,X1)|~v1_card_1(k1_xboole_0)|~v1_xboole_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_818, plain, (v1_relat_1(k8_finseq_1(X1,X2,X3))|~m1_finseq_1(X3,X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_2137, negated_conjecture, (v2_funct_1(esk2_0)|~v2_funct_1(k3_finseq_5(esk2_0))|~v1_finseq_1(k3_finseq_5(esk2_0))|~v1_funct_1(k3_finseq_5(esk2_0))|~v1_relat_1(k3_finseq_5(esk2_0)))).#trainneg
cnf(i_0_5900, negated_conjecture, (v1_xboole_0(esk3_1(esk1_0))|~v1_xboole_0(k3_finseq_5(esk3_1(esk1_0)))|~v1_relat_1(k3_finseq_5(esk3_1(esk1_0))))).#trainneg
cnf(i_0_820, plain, (v1_relat_1(k3_pre_poly(X1,X2))|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_821, plain, (v1_funct_1(k3_pre_poly(X1,X2))|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_4809, plain, (k4_finseq_5(X1,X2)=esk14_1(k1_xboole_0)|~v1_xboole_0(X4)|~v1_xboole_0(X5)|~m1_finseq_1(X2,X5)|~m1_finseq_1(X3,X4)|~m2_finseq_1(esk14_1(k1_xboole_0),X1)|~v1_card_1(k1_xboole_0)|~v1_xboole_0(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_822, plain, (v7_ordinal1(k3_pre_poly(X1,X2))|v1_xboole_0(X1)|~v6_membered(k3_finseq_2(X1))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1565, negated_conjecture, (v7_ordinal1(k4_rlvect_1(esk1_0,esk2_0))|~v1_finset_1(k4_rlvect_1(esk1_0,esk2_0))|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_823, plain, (v3_ordinal1(k3_pre_poly(X1,X2))|v1_xboole_0(X1)|~v3_ordinal1(k3_finseq_2(X1))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_825, plain, (v1_xxreal_0(k3_pre_poly(X1,X2))|v1_xboole_0(X1)|~v2_membered(k3_finseq_2(X1))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_10092, plain, (v1_relat_1(k7_finseq_1(X2,X3))|~m1_finseq_1(X3,X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_10080, plain, (k4_rlvect_1(X1,k7_finseq_1(k4_finseq_5(u1_struct_0(X1),k9_finseq_1(k1_binop_1(esk18_2(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)),X2,X3))),X4))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k9_finseq_1(k1_binop_1(esk18_2(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)),X2,X3)),X4))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m2_finseq_1(X4,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_5914, negated_conjecture, (k3_finseq_5(esk14_1(k1_xboole_0))=esk3_1(esk1_0)|~v1_card_1(k1_xboole_0)|~v1_xboole_0(esk3_1(esk1_0))|~v1_relat_1(esk3_1(esk1_0)))).#trainneg
cnf(i_0_826, plain, (v1_xreal_0(k3_pre_poly(X1,X2))|v1_xboole_0(X1)|~v3_membered(k3_finseq_2(X1))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_9172, negated_conjecture, (v7_ordinal1(k4_rlvect_1(esk1_0,k1_xboole_0))|~v1_finset_1(k4_rlvect_1(esk1_0,k1_xboole_0))|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_827, plain, (v1_xcmplx_0(k3_pre_poly(X1,X2))|v1_xboole_0(X1)|~v1_membered(k3_finseq_2(X1))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_828, plain, (v1_finseq_1(k3_pre_poly(X1,X2))|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_10199, negated_conjecture, (k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(esk3_1(X1)),X2))=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1)))).#trainneg
cnf(i_0_9104, negated_conjecture, (k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(esk3_1(X1)),k3_finseq_5(esk3_1(X1))))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),esk3_1(X1),esk3_1(X1)))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|k4_rlvect_1(X1,k3_finseq_5(esk2_0))=k4_rlvect_1(X1,esk2_0)|~m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_10206, negated_conjecture, (v7_ordinal1(k4_rlvect_1(esk1_0,esk2_0))|~v1_finset_1(k4_rlvect_1(esk1_0,esk2_0))|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10299, negated_conjecture, (v7_ordinal1(k4_rlvect_1(esk1_0,k1_xboole_0))|~v1_finset_1(k4_rlvect_1(esk1_0,k1_xboole_0))|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5554, negated_conjecture, (v7_ordinal1(esk4_1(esk1_0))|~v3_ordinal1(u1_struct_0(esk1_0))|~v1_finseq_1(esk4_1(esk1_0))|~v1_funct_1(esk4_1(esk1_0))|~v1_relat_1(esk4_1(esk1_0)))).#trainneg
cnf(i_0_8386, plain, (v1_relat_1(k4_finseq_5(X1,X2))|~v1_xboole_0(X1)|~v1_xboole_0(X3)|~v1_xboole_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_9406, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,X2))=k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(esk3_1(X1)),X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m1_finseq_1(X2,u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_10501, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,esk3_1(X1)))=k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(esk3_1(X1)),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_rlvect_1(X1,esk3_1(X1)),u1_struct_0(X1))|~m1_subset_1(k4_rlvect_1(X1,X2),u1_struct_0(X1))|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1))).#trainneg
cnf(i_0_8947, negated_conjecture, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),esk3_1(X1),esk3_1(X1)))=k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(esk3_1(X1)),esk3_1(X1)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1)))).#trainneg
cnf(i_0_1357, plain, (m1_subset_1(esk16_1(X1),k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5824, negated_conjecture, (m1_subset_1(k4_rlvect_1(esk1_0,esk3_1(esk1_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4790, plain, (k4_finseq_5(X1,esk14_1(k1_xboole_0))=X2|~v1_funct_1(esk14_1(k1_xboole_0))|~v1_relat_1(esk14_1(k1_xboole_0))|~m2_finseq_1(X2,X1)|~v1_card_1(k1_xboole_0)|~v1_xboole_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_10476, negated_conjecture, (v7_ordinal1(k4_rlvect_1(esk1_0,esk2_0))|~v1_finset_1(k4_rlvect_1(esk1_0,esk2_0))|~v7_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10488, plain, (v1_relat_1(k4_finseq_5(X1,X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_847, plain, (X1=esk14_1(k1_xboole_0)|~v1_xboole_0(X1)|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_852, plain, (v1_valued_0(k3_pre_poly(X1,X2))|v1_xboole_0(X1)|~v1_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_10081, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k9_finseq_1(k5_binop_1(u1_struct_0(X1),esk18_2(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)),X2,X3)),X4))=k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(k9_finseq_1(k5_binop_1(u1_struct_0(X1),esk18_2(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)),X2,X3))),X4))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m2_finseq_1(X4,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k9_finseq_1(k5_binop_1(u1_struct_0(X1),esk18_2(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)),X2,X3)),u1_struct_0(X1)))).#trainneg
cnf(i_0_10829, negated_conjecture, (k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(esk3_1(X1)),esk3_1(X1)))=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),esk3_1(X1)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_10839, plain, (v1_finset_1(esk16_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_853, plain, (v1_valued_0(k8_finseq_1(X1,X2,X3))|~v1_membered(X1)|~m1_finseq_1(X3,X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_10478, negated_conjecture, (v7_ordinal1(k4_rlvect_1(esk1_0,k1_xboole_0))|~v1_finset_1(k4_rlvect_1(esk1_0,k1_xboole_0))|~v7_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5039, plain, (v1_xboole_0(k4_finseq_5(X3,X2))|~v1_xboole_0(X1)|~m1_finseq_1(X2,X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_5049, plain, (v1_finset_1(k4_finseq_5(X3,X2))|~v1_xboole_0(X1)|~m1_finseq_1(X2,X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_856, plain, (v3_valued_0(k3_pre_poly(X1,X2))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_857, plain, (v3_valued_0(k8_finseq_1(X1,X2,X3))|~v3_membered(X1)|~m1_finseq_1(X3,X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_9900, plain, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k9_finseq_1(k5_binop_1(u1_struct_0(X1),esk18_2(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)),X3,X4))))=k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(k9_finseq_1(k5_binop_1(u1_struct_0(X1),esk18_2(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)),X3,X4))),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_9898, plain, (k4_rlvect_1(X1,k3_finseq_5(k8_finseq_1(u1_struct_0(X1),X2,k9_finseq_1(k5_binop_1(u1_struct_0(X1),esk18_2(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)),X3,X4)))))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),X2,k9_finseq_1(k5_binop_1(u1_struct_0(X1),esk18_2(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)),X3,X4))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk3_1(X1)))=k4_rlvect_1(X1,esk3_1(X1))|v1_xboole_0(u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_11011, negated_conjecture, (k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),k3_finseq_5(esk3_1(X1))))=k4_rlvect_1(X1,k7_finseq_1(esk3_1(X1),esk3_1(X1)))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1)))).#trainneg
cnf(i_0_9175, negated_conjecture, (k4_finseq_5(X1,X2)=esk2_0|~v1_xboole_0(X3)|~v1_xboole_0(X4)|~m1_finseq_1(X2,X4)|~m2_finseq_1(esk2_0,X1)|~m1_finseq_1(esk2_0,X3))).#trainneg
cnf(i_0_2882, plain, (v5_relat_1(esk5_1(k1_zfmisc_1(k1_xboole_0)),X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_10848, negated_conjecture, (v1_relat_1(k4_rlvect_1(esk1_0,esk3_1(esk1_0)))|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10849, negated_conjecture, (v1_funct_1(k4_rlvect_1(esk1_0,esk3_1(esk1_0)))|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10850, negated_conjecture, (v7_ordinal1(k4_rlvect_1(esk1_0,esk3_1(esk1_0)))|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11022, plain, (v1_xboole_0(k4_finseq_5(X1,X2))|~v1_xboole_0(X3)|~v1_xboole_0(X1)|~m2_finseq_1(X2,X3))).#trainneg
cnf(i_0_11024, plain, (v1_finset_1(k4_finseq_5(X1,X2))|~v1_xboole_0(X3)|~v1_xboole_0(X1)|~m2_finseq_1(X2,X3))).#trainneg
cnf(i_0_878, plain, (v1_funct_1(k7_finseq_1(X1,X2))|~m1_finseq_1(X2,X3)|~m1_finseq_1(X1,X3))).#trainneg
cnf(i_0_8066, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,esk3_1(X1)))=k4_rlvect_1(X1,k7_finseq_1(X2,k3_finseq_5(esk3_1(X1))))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_11213, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,X2))=k4_rlvect_1(X1,k7_finseq_1(X2,k3_finseq_5(esk3_1(X1))))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_rlvect_1(X1,X2),u1_struct_0(X1))|~m1_subset_1(k4_rlvect_1(X1,esk3_1(X1)),u1_struct_0(X1))|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m1_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1))).#trainneg
cnf(i_0_11170, negated_conjecture, (k4_finseq_5(X1,X2)=esk2_0|~v1_xboole_0(X3)|~v1_xboole_0(X4)|~m1_finseq_1(X2,X4)|~m2_finseq_1(esk2_0,X1)|~m2_finseq_1(esk2_0,X3))).#trainneg
cnf(i_0_2937, plain, (v4_relat_1(esk5_1(k1_zfmisc_1(k1_xboole_0)),X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_10851, negated_conjecture, (v3_ordinal1(k4_rlvect_1(esk1_0,esk3_1(esk1_0)))|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5126, plain, (k4_finseq_5(X3,k3_finseq_5(X2))=X2|~v1_funct_1(k3_finseq_5(X2))|~v1_relat_1(k3_finseq_5(X2))|~m2_finseq_1(X2,X1)|~v1_xboole_0(X1)|~v1_xboole_0(X3)|~m1_finseq_1(k3_finseq_5(X2),X3))).#trainneg
cnf(i_0_10852, negated_conjecture, (v1_finset_1(k4_rlvect_1(esk1_0,esk3_1(esk1_0)))|~v5_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10853, negated_conjecture, (v1_xxreal_0(k4_rlvect_1(esk1_0,esk3_1(esk1_0)))|~v2_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_881, plain, (v3_valued_0(k7_finseq_1(X1,X2))|~v3_membered(X3)|~m1_finseq_1(X2,X3)|~m1_finseq_1(X1,X3))).#trainneg
cnf(i_0_883, plain, (v1_valued_0(k7_finseq_1(X1,X2))|~v1_membered(X3)|~m1_finseq_1(X2,X3)|~m1_finseq_1(X1,X3))).#trainneg
cnf(i_0_9609, negated_conjecture, (k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),k3_finseq_5(esk3_1(X1)),X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),esk3_1(X1),X2))|v2_struct_0(X1)|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|~m2_finseq_1(X2,u1_struct_0(X1))|~m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_10153, negated_conjecture, (k4_rlvect_1(X1,k7_finseq_1(k3_finseq_5(esk3_1(X1)),k3_finseq_5(esk3_1(X1))))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),esk3_1(X1),esk3_1(X1)))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk3_1(X1),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(esk3_1(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1)))).#trainneg
cnf(i_0_11587, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,X2),k4_rlvect_1(X1,esk3_1(X1)))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),esk3_1(X1),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_11695, negated_conjecture, (k1_algstr_0(X1,k4_rlvect_1(X1,esk3_1(X1)),k4_rlvect_1(X1,X2))=k4_rlvect_1(X1,k8_finseq_1(u1_struct_0(X1),esk3_1(X1),X2))|k4_rlvect_1(X1,k4_finseq_5(u1_struct_0(X1),esk2_0))=k4_rlvect_1(X1,esk2_0)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_rlvect_1(X1,X2),u1_struct_0(X1))|~m1_subset_1(k4_rlvect_1(X1,esk3_1(X1)),u1_struct_0(X1))|~v2_rlvect_1(X1)|~m1_finseq_1(k3_finseq_5(esk2_0),u1_struct_0(X1))|~m1_finseq_1(esk2_0,u1_struct_0(X1))|~m2_finseq_1(k3_finseq_5(esk3_1(X1)),u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 180
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 326
# Removed in clause preprocessing      : 35
# Initial clauses in saturation        : 291
# Processed clauses                    : 1927
# ...of these trivial                  : 36
# ...subsumed                          : 668
# ...remaining for further processing  : 1223
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 72
# Backward-rewritten                   : 55
# Generated clauses                    : 11706
# ...of the previous two non-trivial   : 10893
# Contextual simplify-reflections      : 235
# Paramodulations                      : 11706
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1096
#    Positive orientable unit clauses  : 164
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 26
#    Non-unit-clauses                  : 906
# Current number of unprocessed clauses: 9049
# ...number of literals in the above   : 107347
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1353
# Proof object given clauses           : 143
# Proof search given clauses           : 1223
# Clause-clause subsumption calls (NU) : 600350
# Rec. Clause-clause subsumption calls : 75970
# Non-unit clause-clause subsumptions  : 814
# Unit Clause-clause subsumption calls : 10211
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 38
# BW rewrite match successes           : 18
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 610558

# -------------------------------------------------
# User time                : 4.846 s
# System time              : 0.124 s
# Total time               : 4.970 s
# Maximum resident set size: 63876 pages
