# ENIGMA: LightGBM model 'Enigma/mizar40-all-T10/mega-VHSLCAXPh2e15/lgb-t150-d50-l900-e0.15/model.lgb' loaded. (hash_base: 32768; conj_feats: 68; version: 991; iters: 150)
# Preprocessing time       : 0.295 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(commutativity_k8_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k8_real_1(X1,X2)=k8_real_1(X2,X1)), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', commutativity_k8_real_1)).
fof(redefinition_k8_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', redefinition_k8_real_1)).
fof(t30_prob_2, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((((~(v1_xboole_0(X2))&v1_prob_1(X2,X1))&v4_prob_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))=>![X3]:(m2_prob_1(X3,X1,X2)=>![X4]:(m1_prob_1(X4,X1,X2)=>![X5]:(m1_prob_1(X5,X1,X2)=>![X6]:(m1_prob_1(X6,X1,X2)=>(~(r1_xxreal_0(k1_seq_1(X3,k5_prob_1(X1,X2,X4,X5)),k6_numbers))=>k1_seq_1(X3,k5_prob_1(X1,X2,k5_prob_1(X1,X2,X4,X5),X6))=k8_real_1(k8_real_1(k1_seq_1(X3,X4),k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5)),k1_seq_1(k3_prob_2(X1,X2,X3,k5_prob_1(X1,X2,X4,X5)),X6))))))))), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', t30_prob_2)).
fof(cc1_xreal_0, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>v1_xreal_0(X1)), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', cc1_xreal_0)).
fof(t29_prob_2, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((((~(v1_xboole_0(X2))&v1_prob_1(X2,X1))&v4_prob_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))=>![X3]:(m2_prob_1(X3,X1,X2)=>![X4]:(m1_prob_1(X4,X1,X2)=>![X5]:(m1_prob_1(X5,X1,X2)=>(~(r1_xxreal_0(k1_seq_1(X3,X4),k6_numbers))=>k1_seq_1(X3,k5_prob_1(X1,X2,X5,X4))=k8_real_1(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5),k1_seq_1(X3,X4)))))))), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', t29_prob_2)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', redefinition_k6_numbers)).
fof(redefinition_m1_prob_1, axiom, ![X1, X2]:((((~(v1_xboole_0(X2))&v1_prob_1(X2,X1))&v4_prob_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))=>![X3]:(m1_prob_1(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', redefinition_m1_prob_1)).
fof(commutativity_k5_prob_1, axiom, ![X1, X2, X3, X4]:((((((~(v1_xboole_0(X2))&v1_prob_1(X2,X1))&v4_prob_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))&m1_subset_1(X3,X2))&m1_subset_1(X4,X2))=>k5_prob_1(X1,X2,X3,X4)=k5_prob_1(X1,X2,X4,X3)), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', commutativity_k5_prob_1)).
fof(fc63_valued_0, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v3_valued_0(X1))=>v1_xreal_0(k1_funct_1(X1,X2))), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', fc63_valued_0)).
fof(redefinition_k1_seq_1, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v3_valued_0(X1))=>k1_seq_1(X1,X2)=k1_funct_1(X1,X2)), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', redefinition_k1_seq_1)).
fof(dt_m2_prob_1, axiom, ![X1, X2]:(((((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&v1_prob_1(X2,X1))&v4_prob_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))=>![X3]:(m2_prob_1(X3,X1,X2)=>((v1_funct_1(X3)&v1_funct_2(X3,X2,k1_numbers))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))))), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', dt_m2_prob_1)).
fof(dt_k3_prob_2, axiom, ![X1, X2, X3, X4]:(((((((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&v1_prob_1(X2,X1))&v4_prob_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))&m2_prob_1(X3,X1,X2))&m1_subset_1(X4,X2))=>m2_prob_1(k3_prob_2(X1,X2,X3,X4),X1,X2)), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', dt_k3_prob_2)).
fof(cc30_valued_0, axiom, ![X1]:((v1_relat_1(X1)&v5_relat_1(X1,k1_numbers))=>(v1_relat_1(X1)&v3_valued_0(X1))), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', cc30_valued_0)).
fof(t2_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k3_xcmplx_0(X1,k6_numbers)=k6_numbers), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', t2_arithm)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', t6_boole)).
fof(rc4_xreal_0, axiom, ?[X1]:(((v1_xboole_0(X1)&v1_xcmplx_0(X1))&v1_xxreal_0(X1))&v1_xreal_0(X1)), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', rc4_xreal_0)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', cc2_relset_1)).
fof(cc2_xxreal_0, axiom, ![X1]:(v7_ordinal1(X1)=>v1_xxreal_0(X1)), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', cc2_xxreal_0)).
fof(cc7_ordinal1, axiom, ![X1]:(v1_xboole_0(X1)=>v7_ordinal1(X1)), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', cc7_ordinal1)).
fof(rc3_xreal_0, axiom, ?[X1]:(((v1_xcmplx_0(X1)&v1_xxreal_0(X1))&v3_xxreal_0(X1))&v1_xreal_0(X1)), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', rc3_xreal_0)).
fof(cc3_xxreal_0, axiom, ![X1]:((v1_xxreal_0(X1)&v2_xxreal_0(X1))=>((~(v1_xboole_0(X1))&v1_xxreal_0(X1))&~(v3_xxreal_0(X1)))), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', cc3_xxreal_0)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', cc1_relset_1)).
fof(cc5_xxreal_0, axiom, ![X1]:((v1_xxreal_0(X1)&v3_xxreal_0(X1))=>((~(v1_xboole_0(X1))&v1_xxreal_0(X1))&~(v2_xxreal_0(X1)))), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', cc5_xxreal_0)).
fof(fc23_xreal_0, axiom, ![X1, X2]:((((~(v2_xxreal_0(X1))&v1_xreal_0(X1))&~(v3_xxreal_0(X2)))&v1_xreal_0(X2))=>~(v2_xxreal_0(k3_xcmplx_0(X1,X2)))), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', fc23_xreal_0)).
fof(t34_prob_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((((~(v1_xboole_0(X2))&v1_prob_1(X2,X1))&v4_prob_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))=>![X3]:(m1_prob_1(X3,X1,X2)=>![X4]:(m1_prob_1(X4,X1,X2)=>![X5]:(m2_prob_1(X5,X1,X2)=>(r1_tarski(X3,X4)=>r1_xxreal_0(k1_seq_1(X5,X3),k1_seq_1(X5,X4)))))))), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', t34_prob_1)).
fof(fc25_xreal_0, axiom, ![X1, X2]:((((~(v2_xxreal_0(X1))&v1_xreal_0(X1))&~(v2_xxreal_0(X2)))&v1_xreal_0(X2))=>~(v3_xxreal_0(k3_xcmplx_0(X1,X2)))), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', fc25_xreal_0)).
fof(dt_k1_seq_1, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v3_valued_0(X1))=>m1_subset_1(k1_seq_1(X1,X2),k1_numbers)), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', dt_k1_seq_1)).
fof(t1_real, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>((r1_xxreal_0(X1,X2)&v2_xxreal_0(X1))=>v2_xxreal_0(X2)))), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', t1_real)).
fof(t7_real, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>~(((~(r1_xxreal_0(X1,X2))&~(v2_xxreal_0(X1)))&~(v3_xxreal_0(X2)))))), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', t7_real)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', fc1_xboole_0)).
fof(cc18_valued_0, axiom, ![X1, X2]:(v3_membered(X2)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v3_valued_0(X3))), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', cc18_valued_0)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', fc3_membered)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', t3_subset)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', t17_xboole_1)).
fof(dt_k8_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>m1_subset_1(k8_real_1(X1,X2),k1_numbers)), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', dt_k8_real_1)).
fof(t3_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k3_xcmplx_0(np__1,X1)=X1), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', t3_arithm)).
fof(cc3_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xcmplx_0(X1)), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', cc3_xreal_0)).
fof(redefinition_k5_prob_1, axiom, ![X1, X2, X3, X4]:((((((~(v1_xboole_0(X2))&v1_prob_1(X2,X1))&v4_prob_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))&m1_subset_1(X3,X2))&m1_subset_1(X4,X2))=>k5_prob_1(X1,X2,X3,X4)=k3_xboole_0(X3,X4)), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', redefinition_k5_prob_1)).
fof(dt_k5_prob_1, axiom, ![X1, X2, X3, X4]:((((((~(v1_xboole_0(X2))&v1_prob_1(X2,X1))&v4_prob_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))&m1_subset_1(X3,X2))&m1_subset_1(X4,X2))=>m1_prob_1(k5_prob_1(X1,X2,X3,X4),X1,X2)), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', dt_k5_prob_1)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/home/yan/atp/benchmarks/mizar40/all/t30_prob_2', spc1_numerals)).
fof(c_0_40, plain, ![X148, X149]:(~m1_subset_1(X148,k1_numbers)|~v1_xreal_0(X149)|k8_real_1(X148,X149)=k8_real_1(X149,X148)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k8_real_1])])).
fof(c_0_41, plain, ![X316, X317]:(~m1_subset_1(X316,k1_numbers)|~v1_xreal_0(X317)|k8_real_1(X316,X317)=k3_xcmplx_0(X316,X317)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_real_1])])).
fof(c_0_42, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:((((~(v1_xboole_0(X2))&v1_prob_1(X2,X1))&v4_prob_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))=>![X3]:(m2_prob_1(X3,X1,X2)=>![X4]:(m1_prob_1(X4,X1,X2)=>![X5]:(m1_prob_1(X5,X1,X2)=>![X6]:(m1_prob_1(X6,X1,X2)=>(~(r1_xxreal_0(k1_seq_1(X3,k5_prob_1(X1,X2,X4,X5)),k6_numbers))=>k1_seq_1(X3,k5_prob_1(X1,X2,k5_prob_1(X1,X2,X4,X5),X6))=k8_real_1(k8_real_1(k1_seq_1(X3,X4),k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5)),k1_seq_1(k3_prob_2(X1,X2,X3,k5_prob_1(X1,X2,X4,X5)),X6)))))))))), inference(assume_negation,[status(cth)],[t30_prob_2])).
cnf(c_0_43, plain, (k8_real_1(X1,X2)=k8_real_1(X2,X1)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_44, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_45, plain, ![X68]:(~m1_subset_1(X68,k1_numbers)|v1_xreal_0(X68)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_xreal_0])])).
fof(c_0_46, plain, ![X337, X338, X339, X340, X341]:(v1_xboole_0(X337)|(v1_xboole_0(X338)|~v1_prob_1(X338,X337)|~v4_prob_1(X338,X337)|~m1_subset_1(X338,k1_zfmisc_1(k1_zfmisc_1(X337)))|(~m2_prob_1(X339,X337,X338)|(~m1_prob_1(X340,X337,X338)|(~m1_prob_1(X341,X337,X338)|(r1_xxreal_0(k1_seq_1(X339,X340),k6_numbers)|k1_seq_1(X339,k5_prob_1(X337,X338,X341,X340))=k8_real_1(k1_seq_1(k3_prob_2(X337,X338,X339,X340),X341),k1_seq_1(X339,X340)))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t29_prob_2])])])])])])).
fof(c_0_47, negated_conjecture, (~v1_xboole_0(esk1_0)&((((~v1_xboole_0(esk2_0)&v1_prob_1(esk2_0,esk1_0))&v4_prob_1(esk2_0,esk1_0))&m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0))))&(m2_prob_1(esk3_0,esk1_0,esk2_0)&(m1_prob_1(esk4_0,esk1_0,esk2_0)&(m1_prob_1(esk5_0,esk1_0,esk2_0)&(m1_prob_1(esk6_0,esk1_0,esk2_0)&(~r1_xxreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k6_numbers)&k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k8_real_1(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_42])])])])])])).
cnf(c_0_48, plain, (k8_real_1(X1,X2)=k8_real_1(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_43])).
cnf(c_0_49, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_44])).
cnf(c_0_50, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_51, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|r1_xxreal_0(k1_seq_1(X3,X4),k6_numbers)|k1_seq_1(X3,k5_prob_1(X1,X2,X5,X4))=k8_real_1(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5),k1_seq_1(X3,X4))|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m2_prob_1(X3,X1,X2)|~m1_prob_1(X4,X1,X2)|~m1_prob_1(X5,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_52, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
cnf(c_0_53, negated_conjecture, (k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k8_real_1(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_54, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49])])).
cnf(c_0_55, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_50])).
cnf(c_0_56, plain, (k8_real_1(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5),k1_seq_1(X3,X4))=k1_seq_1(X3,k5_prob_1(X1,X2,X5,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|r1_xxreal_0(k1_seq_1(X3,X4),k6_numbers)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m2_prob_1(X3,X1,X2)|~m1_prob_1(X5,X1,X2)|~m1_prob_1(X4,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(evalgc,[status(thm)],[c_0_51])).
cnf(c_0_57, plain, (k6_numbers=k1_xboole_0), inference(evalgc,[status(thm)],[c_0_52])).
cnf(c_0_58, negated_conjecture, (k8_real_1(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))), inference(evalgc,[status(thm)],[c_0_53])).
cnf(c_0_59, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_54, c_0_55]), c_0_55])).
cnf(c_0_60, plain, (k8_real_1(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5),k1_seq_1(X3,X4))=k1_seq_1(X3,k5_prob_1(X1,X2,X5,X4))|r1_xxreal_0(k1_seq_1(X3,X4),k1_xboole_0)|v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_prob_1(X5,X1,X2)|~m1_prob_1(X4,X1,X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)), inference(rw,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_61, negated_conjecture, (m1_prob_1(esk5_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_62, negated_conjecture, (m1_prob_1(esk4_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_63, negated_conjecture, (m2_prob_1(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_64, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_65, negated_conjecture, (v4_prob_1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_66, negated_conjecture, (v1_prob_1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_67, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_68, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_69, plain, ![X318, X319, X320, X320]:((~m1_prob_1(X320,X318,X319)|m1_subset_1(X320,X319)|(v1_xboole_0(X319)|~v1_prob_1(X319,X318)|~v4_prob_1(X319,X318)|~m1_subset_1(X319,k1_zfmisc_1(k1_zfmisc_1(X318)))))&(~m1_subset_1(X320,X319)|m1_prob_1(X320,X318,X319)|(v1_xboole_0(X319)|~v1_prob_1(X319,X318)|~v4_prob_1(X319,X318)|~m1_subset_1(X319,k1_zfmisc_1(k1_zfmisc_1(X318)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m1_prob_1])])])])])])])).
cnf(c_0_70, negated_conjecture, (k8_real_1(k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59])])).
cnf(c_0_71, plain, (k3_xcmplx_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5),k1_seq_1(X3,X4))=k1_seq_1(X3,k5_prob_1(X1,X2,X5,X4))|r1_xxreal_0(k1_seq_1(X3,X4),k1_xboole_0)|v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_prob_1(X5,X1,X2)|~m1_prob_1(X4,X1,X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_xreal_0(k1_seq_1(X3,X4))|~m1_subset_1(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_49])])).
cnf(c_0_72, negated_conjecture, (m1_prob_1(esk5_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[c_0_61])).
cnf(c_0_73, negated_conjecture, (m1_prob_1(esk4_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[c_0_62])).
cnf(c_0_74, negated_conjecture, (m2_prob_1(esk3_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[c_0_63])).
cnf(c_0_75, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))), inference(evalgc,[status(thm)],[c_0_64])).
cnf(c_0_76, negated_conjecture, (v4_prob_1(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_65])).
cnf(c_0_77, negated_conjecture, (v1_prob_1(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_66])).
cnf(c_0_78, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(evalgc,[status(thm)],[c_0_67])).
cnf(c_0_79, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_68])).
fof(c_0_80, plain, ![X144, X145, X146, X147]:(v1_xboole_0(X145)|~v1_prob_1(X145,X144)|~v4_prob_1(X145,X144)|~m1_subset_1(X145,k1_zfmisc_1(k1_zfmisc_1(X144)))|~m1_subset_1(X146,X145)|~m1_subset_1(X147,X145)|k5_prob_1(X144,X145,X146,X147)=k5_prob_1(X144,X145,X147,X146)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[commutativity_k5_prob_1])])])).
cnf(c_0_81, plain, (m1_subset_1(X1,X3)|v1_xboole_0(X3)|~m1_prob_1(X1,X2,X3)|~v1_prob_1(X3,X2)|~v4_prob_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_82, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_72]), c_0_73]), c_0_74]), c_0_75]), c_0_76]), c_0_77])]), c_0_78]), c_0_79])])).
cnf(c_0_83, plain, (v1_xboole_0(X1)|k5_prob_1(X2,X1,X3,X4)=k5_prob_1(X2,X1,X4,X3)|~v1_prob_1(X1,X2)|~v4_prob_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X3,X1)|~m1_subset_1(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_84, plain, (v1_xboole_0(X3)|m1_subset_1(X1,X3)|~v1_prob_1(X3,X2)|~v4_prob_1(X3,X2)|~m1_prob_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(evalgc,[status(thm)],[c_0_81])).
cnf(c_0_85, negated_conjecture, (~r1_xxreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k6_numbers)), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_86, plain, ![X244, X245]:(~v1_relat_1(X244)|~v1_funct_1(X244)|~v3_valued_0(X244)|v1_xreal_0(k1_funct_1(X244,X245))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc63_valued_0])])])])).
fof(c_0_87, plain, ![X310, X311]:(~v1_relat_1(X310)|~v1_funct_1(X310)|~v3_valued_0(X310)|k1_seq_1(X310,X311)=k1_funct_1(X310,X311)), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_seq_1])])])])).
cnf(c_0_88, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)), inference(csr,[status(thm)],[c_0_82, c_0_55])).
cnf(c_0_89, plain, (k5_prob_1(X2,X1,X3,X4)=k5_prob_1(X2,X1,X4,X3)|v1_xboole_0(X1)|~v1_prob_1(X1,X2)|~v4_prob_1(X1,X2)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(evalgc,[status(thm)],[c_0_83])).
cnf(c_0_90, negated_conjecture, (m1_subset_1(esk4_0,esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_73]), c_0_75]), c_0_76]), c_0_77])]), c_0_78])])).
cnf(c_0_91, negated_conjecture, (m1_subset_1(esk5_0,esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_72]), c_0_75]), c_0_76]), c_0_77])]), c_0_78])])).
cnf(c_0_92, negated_conjecture, (m1_prob_1(esk6_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_93, negated_conjecture, (~r1_xxreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k6_numbers)), inference(evalgc,[status(thm)],[c_0_85])).
cnf(c_0_94, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_95, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_87])).
fof(c_0_96, plain, ![X167, X168, X169]:(((v1_funct_1(X169)|~m2_prob_1(X169,X167,X168)|(v1_xboole_0(X167)|v1_xboole_0(X168)|~v1_prob_1(X168,X167)|~v4_prob_1(X168,X167)|~m1_subset_1(X168,k1_zfmisc_1(k1_zfmisc_1(X167)))))&(v1_funct_2(X169,X168,k1_numbers)|~m2_prob_1(X169,X167,X168)|(v1_xboole_0(X167)|v1_xboole_0(X168)|~v1_prob_1(X168,X167)|~v4_prob_1(X168,X167)|~m1_subset_1(X168,k1_zfmisc_1(k1_zfmisc_1(X167))))))&(m1_subset_1(X169,k1_zfmisc_1(k2_zfmisc_1(X168,k1_numbers)))|~m2_prob_1(X169,X167,X168)|(v1_xboole_0(X167)|v1_xboole_0(X168)|~v1_prob_1(X168,X167)|~v4_prob_1(X168,X167)|~m1_subset_1(X168,k1_zfmisc_1(k1_zfmisc_1(X167)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m2_prob_1])])])])])])])).
fof(c_0_97, plain, ![X154, X155, X156, X157]:(v1_xboole_0(X154)|v1_xboole_0(X155)|~v1_prob_1(X155,X154)|~v4_prob_1(X155,X154)|~m1_subset_1(X155,k1_zfmisc_1(k1_zfmisc_1(X154)))|~m2_prob_1(X156,X154,X155)|~m1_subset_1(X157,X155)|m2_prob_1(k3_prob_2(X154,X155,X156,X157),X154,X155)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_prob_2])])])).
cnf(c_0_98, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_75]), c_0_90]), c_0_91]), c_0_76]), c_0_77])]), c_0_78])])).
cnf(c_0_99, plain, (k8_real_1(k1_seq_1(X3,X4),k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))=k1_seq_1(X3,k5_prob_1(X1,X2,X5,X4))|r1_xxreal_0(k1_seq_1(X3,X4),k1_xboole_0)|v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_prob_1(X5,X1,X2)|~m1_prob_1(X4,X1,X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|~m1_subset_1(k1_seq_1(X3,X4),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_48])])).
cnf(c_0_100, negated_conjecture, (m1_prob_1(esk6_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[c_0_92])).
cnf(c_0_101, negated_conjecture, (~r1_xxreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_xboole_0)), inference(rw,[status(thm)],[c_0_93, c_0_57])).
cnf(c_0_102, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[c_0_94])).
cnf(c_0_103, plain, (k1_funct_1(X1,X2)=k1_seq_1(X1,X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[c_0_95])).
cnf(c_0_104, plain, (v1_funct_1(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m2_prob_1(X1,X2,X3)|~v1_prob_1(X3,X2)|~v4_prob_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_105, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m2_prob_1(k3_prob_2(X1,X2,X3,X4),X1,X2)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_106, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_100]), c_0_74]), c_0_75]), c_0_76]), c_0_77])]), c_0_101]), c_0_78]), c_0_79])])).
cnf(c_0_107, plain, (v1_xreal_0(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103])])).
cnf(c_0_108, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|v1_funct_1(X1)|~v1_prob_1(X3,X2)|~v4_prob_1(X3,X2)|~m2_prob_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(evalgc,[status(thm)],[c_0_104])).
cnf(c_0_109, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m2_prob_1(k3_prob_2(X1,X2,X3,X4),X1,X2)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X4,X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(evalgc,[status(thm)],[c_0_105])).
fof(c_0_110, plain, ![X100]:((v1_relat_1(X100)|(~v1_relat_1(X100)|~v5_relat_1(X100,k1_numbers)))&(v3_valued_0(X100)|(~v1_relat_1(X100)|~v5_relat_1(X100,k1_numbers)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc30_valued_0])])])).
fof(c_0_111, plain, ![X342]:(~v1_xcmplx_0(X342)|k3_xcmplx_0(X342,k6_numbers)=k6_numbers), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_arithm])])).
fof(c_0_112, plain, ![X368]:(~v1_xboole_0(X368)|X368=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
fof(c_0_113, plain, (((v1_xboole_0(esk35_0)&v1_xcmplx_0(esk35_0))&v1_xxreal_0(esk35_0))&v1_xreal_0(esk35_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc4_xreal_0])])).
cnf(c_0_114, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_funct_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v3_valued_0(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107])])).
cnf(c_0_115, plain, (v1_funct_1(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X4,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109])])).
cnf(c_0_116, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_110])).
fof(c_0_117, plain, ![X92, X93, X94]:((v4_relat_1(X94,X92)|~m1_subset_1(X94,k1_zfmisc_1(k2_zfmisc_1(X92,X93))))&(v5_relat_1(X94,X93)|~m1_subset_1(X94,k1_zfmisc_1(k2_zfmisc_1(X92,X93))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
cnf(c_0_118, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|v1_xboole_0(X3)|v1_xboole_0(X2)|~m2_prob_1(X1,X3,X2)|~v1_prob_1(X2,X3)|~v4_prob_1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))), inference(split_conjunct,[status(thm)],[c_0_96])).
fof(c_0_119, plain, ![X99]:(~v7_ordinal1(X99)|v1_xxreal_0(X99)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_xxreal_0])])).
fof(c_0_120, plain, ![X130]:(~v1_xboole_0(X130)|v7_ordinal1(X130)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc7_ordinal1])])).
cnf(c_0_121, plain, (k3_xcmplx_0(X1,k6_numbers)=k6_numbers|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_111])).
fof(c_0_122, plain, (((v1_xcmplx_0(esk30_0)&v1_xxreal_0(esk30_0))&v3_xxreal_0(esk30_0))&v1_xreal_0(esk30_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc3_xreal_0])])).
cnf(c_0_123, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_124, plain, (v1_xboole_0(esk35_0)), inference(split_conjunct,[status(thm)],[c_0_113])).
fof(c_0_125, plain, ![X111]:(((~v1_xboole_0(X111)|(~v1_xxreal_0(X111)|~v2_xxreal_0(X111)))&(v1_xxreal_0(X111)|(~v1_xxreal_0(X111)|~v2_xxreal_0(X111))))&(~v3_xxreal_0(X111)|(~v1_xxreal_0(X111)|~v2_xxreal_0(X111)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc3_xxreal_0])])])])).
cnf(c_0_126, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v3_valued_0(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_74]), c_0_75]), c_0_76]), c_0_77])]), c_0_78]), c_0_79])])).
cnf(c_0_127, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_116])).
cnf(c_0_128, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_117])).
cnf(c_0_129, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_prob_1(X2,X3)|~v4_prob_1(X2,X3)|~m2_prob_1(X1,X3,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))), inference(evalgc,[status(thm)],[c_0_118])).
fof(c_0_130, plain, ![X62, X63, X64]:(~m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(X62,X63)))|v1_relat_1(X64)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_131, plain, ![X126]:(((~v1_xboole_0(X126)|(~v1_xxreal_0(X126)|~v3_xxreal_0(X126)))&(v1_xxreal_0(X126)|(~v1_xxreal_0(X126)|~v3_xxreal_0(X126))))&(~v2_xxreal_0(X126)|(~v1_xxreal_0(X126)|~v3_xxreal_0(X126)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc5_xxreal_0])])])])).
cnf(c_0_132, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_133, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_120])).
fof(c_0_134, plain, ![X200, X201]:(v2_xxreal_0(X200)|~v1_xreal_0(X200)|v3_xxreal_0(X201)|~v1_xreal_0(X201)|~v2_xxreal_0(k3_xcmplx_0(X200,X201))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc23_xreal_0])])])).
cnf(c_0_135, plain, (k3_xcmplx_0(X1,k6_numbers)=k6_numbers|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_121])).
cnf(c_0_136, plain, (v1_xcmplx_0(esk30_0)), inference(split_conjunct,[status(thm)],[c_0_122])).
cnf(c_0_137, plain, (v1_xreal_0(esk35_0)), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_138, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_123])).
cnf(c_0_139, plain, (v1_xboole_0(esk35_0)), inference(evalgc,[status(thm)],[c_0_124])).
cnf(c_0_140, plain, (~v3_xxreal_0(X1)|~v1_xxreal_0(X1)|~v2_xxreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_125])).
cnf(c_0_141, plain, (v3_xxreal_0(esk30_0)), inference(split_conjunct,[status(thm)],[c_0_122])).
cnf(c_0_142, plain, (v1_xxreal_0(esk30_0)), inference(split_conjunct,[status(thm)],[c_0_122])).
cnf(c_0_143, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v5_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_127])])).
cnf(c_0_144, negated_conjecture, (m1_subset_1(esk6_0,esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_100]), c_0_75]), c_0_76]), c_0_77])]), c_0_78])])).
cnf(c_0_145, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(evalgc,[status(thm)],[c_0_128])).
cnf(c_0_146, plain, (m1_subset_1(k3_prob_2(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|v1_xboole_0(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X4,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_109])])).
cnf(c_0_147, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_130])).
fof(c_0_148, plain, ![X348, X349, X350, X351, X352]:(v1_xboole_0(X348)|(v1_xboole_0(X349)|~v1_prob_1(X349,X348)|~v4_prob_1(X349,X348)|~m1_subset_1(X349,k1_zfmisc_1(k1_zfmisc_1(X348)))|(~m1_prob_1(X350,X348,X349)|(~m1_prob_1(X351,X348,X349)|(~m2_prob_1(X352,X348,X349)|(~r1_tarski(X350,X351)|r1_xxreal_0(k1_seq_1(X352,X350),k1_seq_1(X352,X351)))))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_prob_1])])])])])])).
fof(c_0_149, plain, ![X204, X205]:(v2_xxreal_0(X204)|~v1_xreal_0(X204)|v2_xxreal_0(X205)|~v1_xreal_0(X205)|~v3_xxreal_0(k3_xcmplx_0(X204,X205))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc25_xreal_0])])])).
cnf(c_0_150, plain, (~v1_xboole_0(X1)|~v1_xxreal_0(X1)|~v3_xxreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_131])).
cnf(c_0_151, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_132])).
cnf(c_0_152, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_133])).
cnf(c_0_153, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v2_xxreal_0(k3_xcmplx_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_154, plain, (k3_xcmplx_0(X1,k1_xboole_0)=k1_xboole_0|~v1_xcmplx_0(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_135, c_0_57]), c_0_57])).
cnf(c_0_155, plain, (v1_xcmplx_0(esk30_0)), inference(evalgc,[status(thm)],[c_0_136])).
cnf(c_0_156, plain, (v1_xreal_0(esk35_0)), inference(evalgc,[status(thm)],[c_0_137])).
cnf(c_0_157, plain, (esk35_0=k1_xboole_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139])])).
cnf(c_0_158, plain, (v1_xreal_0(esk30_0)), inference(split_conjunct,[status(thm)],[c_0_122])).
cnf(c_0_159, plain, (~v1_xxreal_0(X1)|~v2_xxreal_0(X1)|~v3_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_140])).
cnf(c_0_160, plain, (v3_xxreal_0(esk30_0)), inference(evalgc,[status(thm)],[c_0_141])).
cnf(c_0_161, plain, (v1_xxreal_0(esk30_0)), inference(evalgc,[status(thm)],[c_0_142])).
cnf(c_0_162, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v5_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_89]), c_0_75]), c_0_144]), c_0_76]), c_0_77])]), c_0_78])])).
cnf(c_0_163, plain, (v5_relat_1(k3_prob_2(X1,X2,X3,X4),k1_numbers)|v1_xboole_0(X1)|v1_xboole_0(X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_146])])).
cnf(c_0_164, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_147])).
fof(c_0_165, plain, ![X152, X153]:(~v1_relat_1(X152)|~v1_funct_1(X152)|~v3_valued_0(X152)|m1_subset_1(k1_seq_1(X152,X153),k1_numbers)), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_seq_1])])])])).
fof(c_0_166, plain, ![X333, X334]:(~v1_xreal_0(X333)|(~v1_xreal_0(X334)|(~r1_xxreal_0(X333,X334)|~v2_xxreal_0(X333)|v2_xxreal_0(X334)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_real])])])])])).
cnf(c_0_167, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|r1_xxreal_0(k1_seq_1(X5,X3),k1_seq_1(X5,X4))|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_prob_1(X3,X1,X2)|~m1_prob_1(X4,X1,X2)|~m2_prob_1(X5,X1,X2)|~r1_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_148])).
fof(c_0_168, plain, ![X373, X374]:(~v1_xreal_0(X373)|(~v1_xreal_0(X374)|(r1_xxreal_0(X373,X374)|v2_xxreal_0(X373)|v3_xxreal_0(X374)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_real])])])])])])).
cnf(c_0_169, plain, (v2_xxreal_0(X1)|v2_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v3_xxreal_0(k3_xcmplx_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_149])).
cnf(c_0_170, plain, (~v1_xboole_0(X1)|~v1_xxreal_0(X1)|~v3_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_150])).
cnf(c_0_171, plain, (v1_xxreal_0(X1)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152])])).
cnf(c_0_172, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(k3_xcmplx_0(X1,X2))), inference(evalgc,[status(thm)],[c_0_153])).
cnf(c_0_173, plain, (k3_xcmplx_0(esk30_0,k1_xboole_0)=k1_xboole_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_155])])).
cnf(c_0_174, plain, (v1_xreal_0(k1_xboole_0)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_156, c_0_157])])).
cnf(c_0_175, plain, (v1_xreal_0(esk30_0)), inference(evalgc,[status(thm)],[c_0_158])).
cnf(c_0_176, plain, (~v2_xxreal_0(esk30_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_160]), c_0_161])])])).
cnf(c_0_177, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_178, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_163]), c_0_74]), c_0_75]), c_0_76]), c_0_77])]), c_0_78]), c_0_79])])).
cnf(c_0_179, plain, (v1_relat_1(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_146])])).
cnf(c_0_180, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_165])).
fof(c_0_181, plain, ![X50, X51, X52]:(~v3_membered(X51)|(~m1_subset_1(X52,k1_zfmisc_1(k2_zfmisc_1(X50,X51)))|v3_valued_0(X52))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc18_valued_0])])])])])).
cnf(c_0_182, plain, (v2_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)|~v2_xxreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_166])).
cnf(c_0_183, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|r1_xxreal_0(k1_seq_1(X5,X3),k1_seq_1(X5,X4))|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~r1_tarski(X3,X4)|~m2_prob_1(X5,X1,X2)|~m1_prob_1(X4,X1,X2)|~m1_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(evalgc,[status(thm)],[c_0_167])).
cnf(c_0_184, plain, (r1_xxreal_0(X1,X2)|v2_xxreal_0(X1)|v3_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_168])).
cnf(c_0_185, plain, (v2_xxreal_0(X2)|v2_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(k3_xcmplx_0(X1,X2))), inference(evalgc,[status(thm)],[c_0_169])).
cnf(c_0_186, plain, (~v3_xxreal_0(X1)|~v1_xboole_0(X1)), inference(csr,[status(thm)],[c_0_170, c_0_171])).
cnf(c_0_187, plain, (v3_xxreal_0(k1_xboole_0)|~v2_xxreal_0(k1_xboole_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_173]), c_0_174]), c_0_175])]), c_0_176])])).
cnf(c_0_188, plain, (v1_xboole_0(k1_xboole_0)), inference(evalgc,[status(thm)],[c_0_177])).
cnf(c_0_189, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_179]), c_0_74]), c_0_75]), c_0_76]), c_0_77])]), c_0_78]), c_0_79])])).
cnf(c_0_190, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[c_0_180])).
cnf(c_0_191, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_181])).
cnf(c_0_192, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
cnf(c_0_193, plain, (v2_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_182])).
cnf(c_0_194, negated_conjecture, (r1_xxreal_0(k1_seq_1(X1,X2),k1_seq_1(X1,esk4_0))|~r1_tarski(X2,esk4_0)|~m1_prob_1(X2,esk1_0,esk2_0)|~m2_prob_1(X1,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_73]), c_0_75]), c_0_76]), c_0_77])]), c_0_79]), c_0_78])])).
cnf(c_0_195, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X2)|r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_184])).
cnf(c_0_196, plain, (v2_xxreal_0(k1_xboole_0)|~v3_xxreal_0(k1_xboole_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_173]), c_0_174]), c_0_175])]), c_0_176])])).
cnf(c_0_197, plain, (~v2_xxreal_0(k1_xboole_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_187]), c_0_188])])])).
fof(c_0_198, plain, ![X356, X357, X356, X357]:((~m1_subset_1(X356,k1_zfmisc_1(X357))|r1_tarski(X356,X357))&(~r1_tarski(X356,X357)|m1_subset_1(X356,k1_zfmisc_1(X357)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])])])).
cnf(c_0_199, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_funct_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~v3_valued_0(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_190])])).
cnf(c_0_200, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_191])).
cnf(c_0_201, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_192])).
cnf(c_0_202, negated_conjecture, (v2_xxreal_0(k1_seq_1(X1,esk4_0))|~v2_xxreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,esk4_0))|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_tarski(X2,esk4_0)|~m1_prob_1(X2,esk1_0,esk2_0)|~m2_prob_1(X1,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_194])])).
cnf(c_0_203, negated_conjecture, (v3_xxreal_0(k1_xboole_0)|v2_xxreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_xreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_195]), c_0_174])])])).
cnf(c_0_204, plain, (~v3_xxreal_0(k1_xboole_0)), inference(sr,[status(thm)],[c_0_196, c_0_197])).
cnf(c_0_205, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_198])).
fof(c_0_206, plain, ![X331, X332]:r1_tarski(k3_xboole_0(X331,X332),X331), inference(variable_rename,[status(thm)],[t17_xboole_1])).
cnf(c_0_207, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v3_valued_0(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_115]), c_0_74]), c_0_75]), c_0_90]), c_0_76]), c_0_77])]), c_0_78]), c_0_79])])).
cnf(c_0_208, plain, (v3_valued_0(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_146]), c_0_201])])])).
cnf(c_0_209, plain, (m1_prob_1(X1,X3,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)|~v1_prob_1(X2,X3)|~v4_prob_1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))), inference(split_conjunct,[status(thm)],[c_0_69])).
fof(c_0_210, plain, ![X162, X163]:(~m1_subset_1(X162,k1_numbers)|~v1_xreal_0(X163)|m1_subset_1(k8_real_1(X162,X163),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_real_1])])).
fof(c_0_211, plain, ![X353]:(~v1_xcmplx_0(X353)|k3_xcmplx_0(np__1,X353)=X353), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_arithm])])).
fof(c_0_212, plain, ![X110]:(~v1_xreal_0(X110)|v1_xcmplx_0(X110)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_xreal_0])])).
cnf(c_0_213, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~r1_tarski(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk4_0)|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_203]), c_0_74])]), c_0_204])])).
cnf(c_0_214, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_205])).
cnf(c_0_215, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_74]), c_0_75]), c_0_76]), c_0_77])]), c_0_78]), c_0_79])])).
fof(c_0_216, plain, ![X312, X313, X314, X315]:(v1_xboole_0(X313)|~v1_prob_1(X313,X312)|~v4_prob_1(X313,X312)|~m1_subset_1(X313,k1_zfmisc_1(k1_zfmisc_1(X312)))|~m1_subset_1(X314,X313)|~m1_subset_1(X315,X313)|k5_prob_1(X312,X313,X314,X315)=k3_xboole_0(X314,X315)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k5_prob_1])])])).
cnf(c_0_217, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_198])).
cnf(c_0_218, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_206])).
fof(c_0_219, plain, ![X158, X159, X160, X161]:(v1_xboole_0(X159)|~v1_prob_1(X159,X158)|~v4_prob_1(X159,X158)|~m1_subset_1(X159,k1_zfmisc_1(k1_zfmisc_1(X158)))|~m1_subset_1(X160,X159)|~m1_subset_1(X161,X159)|m1_prob_1(k5_prob_1(X158,X159,X160,X161),X158,X159)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_prob_1])])])).
cnf(c_0_220, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_208]), c_0_74]), c_0_75]), c_0_90]), c_0_76]), c_0_77])]), c_0_78]), c_0_79])])).
cnf(c_0_221, plain, (v1_xboole_0(X2)|m1_prob_1(X1,X3,X2)|~v1_prob_1(X2,X3)|~v4_prob_1(X2,X3)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))), inference(evalgc,[status(thm)],[c_0_209])).
cnf(c_0_222, plain, (m1_subset_1(k8_real_1(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_210])).
cnf(c_0_223, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_211])).
cnf(c_0_224, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_212])).
cnf(c_0_225, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk4_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_214])])).
cnf(c_0_226, negated_conjecture, (v1_funct_1(esk3_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_74]), c_0_75]), c_0_76]), c_0_77])]), c_0_79]), c_0_78])])).
cnf(c_0_227, negated_conjecture, (v3_valued_0(esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_215]), c_0_201])])])).
cnf(c_0_228, negated_conjecture, (v1_relat_1(esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_215])])).
cnf(c_0_229, plain, (v1_xboole_0(X1)|k5_prob_1(X2,X1,X3,X4)=k3_xboole_0(X3,X4)|~v1_prob_1(X1,X2)|~v4_prob_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X3,X1)|~m1_subset_1(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_216])).
cnf(c_0_230, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_217])).
cnf(c_0_231, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(evalgc,[status(thm)],[c_0_218])).
cnf(c_0_232, plain, (v1_xboole_0(X1)|m1_prob_1(k5_prob_1(X2,X1,X3,X4),X2,X1)|~v1_prob_1(X1,X2)|~v4_prob_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X3,X1)|~m1_subset_1(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_219])).
cnf(c_0_233, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_221]), c_0_75]), c_0_76]), c_0_77])]), c_0_78])])).
cnf(c_0_234, plain, (m1_subset_1(k8_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_222])).
cnf(c_0_235, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_223])).
cnf(c_0_236, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_224])).
cnf(c_0_237, plain, (m1_subset_1(np__1,k1_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_238, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk4_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_107]), c_0_226]), c_0_227]), c_0_228])])])).
cnf(c_0_239, plain, (k5_prob_1(X2,X1,X3,X4)=k3_xboole_0(X3,X4)|v1_xboole_0(X1)|~v1_prob_1(X1,X2)|~v4_prob_1(X1,X2)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(evalgc,[status(thm)],[c_0_229])).
cnf(c_0_240, plain, (m1_subset_1(k3_xboole_0(X1,X2),k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_231])])).
cnf(c_0_241, plain, (v1_xboole_0(X1)|m1_prob_1(k5_prob_1(X2,X1,X3,X4),X2,X1)|~v1_prob_1(X1,X2)|~v4_prob_1(X1,X2)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(evalgc,[status(thm)],[c_0_232])).
cnf(c_0_242, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_190]), c_0_226]), c_0_227]), c_0_228])])])).
cnf(c_0_243, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_49])])).
cnf(c_0_244, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_236])])).
cnf(c_0_245, plain, (m1_subset_1(np__1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_237])).
cnf(c_0_246, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_prob_1(k3_xboole_0(esk4_0,esk5_0),esk1_0,esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_239]), c_0_240]), c_0_75]), c_0_91]), c_0_90]), c_0_76]), c_0_77])]), c_0_78])])).
cnf(c_0_247, plain, (m1_prob_1(k3_xboole_0(X3,X4),X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_239])])).
cnf(c_0_248, negated_conjecture, (~v2_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_242]), c_0_174])]), c_0_197])])).
cnf(c_0_249, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_244]), c_0_245])])])).
cnf(c_0_250, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_247]), c_0_75]), c_0_91]), c_0_90]), c_0_76]), c_0_77])]), c_0_78])])).
cnf(c_0_251, negated_conjecture, (~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_248, c_0_249]), c_0_250])).
cnf(c_0_252, negated_conjecture, (~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_179]), c_0_74]), c_0_75]), c_0_90]), c_0_76]), c_0_77])]), c_0_78]), c_0_79])])).
cnf(c_0_253, plain, (m1_subset_1(k5_prob_1(X1,X2,X3,X4),X2)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_241])])).
cnf(c_0_254, negated_conjecture, (~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_253]), c_0_75]), c_0_91]), c_0_90]), c_0_76]), c_0_77])]), c_0_78])])).
cnf(c_0_255, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_107]), c_0_226]), c_0_227]), c_0_228])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 256
# Proof object clause steps            : 179
# Proof object formula steps           : 77
# Proof object conjectures             : 62
# Proof object clause conjectures      : 59
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 57
# Proof object initial formulas used   : 40
# Proof object generating inferences   : 56
# Proof object simplifying inferences  : 186
# Training examples: 113 positive, 1743 negative
# Training: Positive examples begin
cnf(c_0_57, plain, (k6_numbers=k1_xboole_0)).#trainpos
cnf(c_0_79, negated_conjecture, (~v1_xboole_0(esk1_0))).#trainpos
cnf(c_0_78, negated_conjecture, (~v1_xboole_0(esk2_0))).#trainpos
cnf(c_0_77, negated_conjecture, (v1_prob_1(esk2_0,esk1_0))).#trainpos
cnf(c_0_76, negated_conjecture, (v4_prob_1(esk2_0,esk1_0))).#trainpos
cnf(c_0_93, negated_conjecture, (~r1_xxreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k6_numbers))).#trainpos
cnf(c_0_74, negated_conjecture, (m2_prob_1(esk3_0,esk1_0,esk2_0))).#trainpos
cnf(c_0_188, plain, (v1_xboole_0(k1_xboole_0))).#trainpos
cnf(c_0_58, negated_conjecture, (k8_real_1(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0)))).#trainpos
cnf(c_0_73, negated_conjecture, (m1_prob_1(esk4_0,esk1_0,esk2_0))).#trainpos
cnf(c_0_139, plain, (v1_xboole_0(esk35_0))).#trainpos
cnf(c_0_75, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainpos
cnf(c_0_103, plain, (k1_funct_1(X1,X2)=k1_seq_1(X1,X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1))).#trainpos
cnf(c_0_72, negated_conjecture, (m1_prob_1(esk5_0,esk1_0,esk2_0))).#trainpos
cnf(c_0_100, negated_conjecture, (m1_prob_1(esk6_0,esk1_0,esk2_0))).#trainpos
cnf(c_0_221, plain, (v1_xboole_0(X2)|m1_prob_1(X1,X3,X2)|~v1_prob_1(X2,X3)|~v4_prob_1(X2,X3)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3))))).#trainpos
cnf(c_0_108, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|v1_funct_1(X1)|~v1_prob_1(X3,X2)|~v4_prob_1(X3,X2)|~m2_prob_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainpos
cnf(c_0_84, plain, (v1_xboole_0(X3)|m1_subset_1(X1,X3)|~v1_prob_1(X3,X2)|~v4_prob_1(X3,X2)|~m1_prob_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainpos
cnf(c_0_129, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_prob_1(X2,X3)|~v4_prob_1(X2,X3)|~m2_prob_1(X1,X3,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3))))).#trainpos
cnf(c_0_226, negated_conjecture, (v1_funct_1(esk3_0))).#trainpos
cnf(c_0_239, plain, (k5_prob_1(X2,X1,X3,X4)=k3_xboole_0(X3,X4)|v1_xboole_0(X1)|~v1_prob_1(X1,X2)|~v4_prob_1(X1,X2)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainpos
cnf(c_0_183, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|r1_xxreal_0(k1_seq_1(X5,X3),k1_seq_1(X5,X4))|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~r1_tarski(X3,X4)|~m2_prob_1(X5,X1,X2)|~m1_prob_1(X4,X1,X2)|~m1_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainpos
cnf(c_0_89, plain, (k5_prob_1(X2,X1,X3,X4)=k5_prob_1(X2,X1,X4,X3)|v1_xboole_0(X1)|~v1_prob_1(X1,X2)|~v4_prob_1(X1,X2)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainpos
cnf(c_0_241, plain, (v1_xboole_0(X1)|m1_prob_1(k5_prob_1(X2,X1,X3,X4),X2,X1)|~v1_prob_1(X1,X2)|~v4_prob_1(X1,X2)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainpos
cnf(c_0_109, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m2_prob_1(k3_prob_2(X1,X2,X3,X4),X1,X2)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X4,X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainpos
cnf(c_0_56, plain, (k8_real_1(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5),k1_seq_1(X3,X4))=k1_seq_1(X3,k5_prob_1(X1,X2,X5,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|r1_xxreal_0(k1_seq_1(X3,X4),k6_numbers)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m2_prob_1(X3,X1,X2)|~m1_prob_1(X5,X1,X2)|~m1_prob_1(X4,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainpos
cnf(c_0_144, negated_conjecture, (m1_subset_1(esk6_0,esk2_0))).#trainpos
cnf(c_0_91, negated_conjecture, (m1_subset_1(esk5_0,esk2_0))).#trainpos
cnf(c_0_90, negated_conjecture, (m1_subset_1(esk4_0,esk2_0))).#trainpos
cnf(c_0_215, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers))))).#trainpos
cnf(c_0_247, plain, (m1_prob_1(k3_xboole_0(X3,X4),X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainpos
cnf(c_0_253, plain, (m1_subset_1(k5_prob_1(X1,X2,X3,X4),X2)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainpos
cnf(c_0_194, negated_conjecture, (r1_xxreal_0(k1_seq_1(X1,X2),k1_seq_1(X1,esk4_0))|~r1_tarski(X2,esk4_0)|~m1_prob_1(X2,esk1_0,esk2_0)|~m2_prob_1(X1,esk1_0,esk2_0))).#trainpos
cnf(c_0_146, plain, (m1_subset_1(k3_prob_2(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|v1_xboole_0(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X4,X2))).#trainpos
cnf(c_0_115, plain, (v1_funct_1(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X4,X2))).#trainpos
cnf(c_0_155, plain, (v1_xcmplx_0(esk30_0))).#trainpos
cnf(c_0_161, plain, (v1_xxreal_0(esk30_0))).#trainpos
cnf(c_0_201, plain, (v3_membered(k1_numbers))).#trainpos
cnf(c_0_175, plain, (v1_xreal_0(esk30_0))).#trainpos
cnf(c_0_156, plain, (v1_xreal_0(esk35_0))).#trainpos
cnf(c_0_160, plain, (v3_xxreal_0(esk30_0))).#trainpos
cnf(c_0_138, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainpos
cnf(c_0_157, plain, (esk35_0=k1_xboole_0)).#trainpos
cnf(c_0_174, plain, (v1_xreal_0(k1_xboole_0))).#trainpos
cnf(c_0_245, plain, (m1_subset_1(np__1,k1_numbers))).#trainpos
cnf(c_0_236, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainpos
cnf(c_0_151, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).#trainpos
cnf(c_0_152, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainpos
cnf(c_0_171, plain, (v1_xxreal_0(X1)|~v1_xboole_0(X1))).#trainpos
cnf(c_0_135, plain, (k3_xcmplx_0(X1,k6_numbers)=k6_numbers|~v1_xcmplx_0(X1))).#trainpos
cnf(c_0_173, plain, (k3_xcmplx_0(esk30_0,k1_xboole_0)=k1_xboole_0)).#trainpos
cnf(c_0_235, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1))).#trainpos
cnf(c_0_170, plain, (~v1_xboole_0(X1)|~v1_xxreal_0(X1)|~v3_xxreal_0(X1))).#trainpos
cnf(c_0_159, plain, (~v1_xxreal_0(X1)|~v2_xxreal_0(X1)|~v3_xxreal_0(X1))).#trainpos
cnf(c_0_176, plain, (~v2_xxreal_0(esk30_0))).#trainpos
cnf(c_0_55, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_244, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xreal_0(X1))).#trainpos
cnf(c_0_127, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainpos
cnf(c_0_195, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X2)|r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainpos
cnf(c_0_231, plain, (r1_tarski(k3_xboole_0(X1,X2),X1))).#trainpos
cnf(c_0_230, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainpos
cnf(c_0_214, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainpos
cnf(c_0_193, plain, (v2_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainpos
cnf(c_0_102, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1))).#trainpos
cnf(c_0_48, plain, (k8_real_1(X1,X2)=k8_real_1(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_99, plain, (k8_real_1(k1_seq_1(X3,X4),k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))=k1_seq_1(X3,k5_prob_1(X1,X2,X5,X4))|r1_xxreal_0(k1_seq_1(X3,X4),k1_xboole_0)|v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_prob_1(X5,X1,X2)|~m1_prob_1(X4,X1,X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|~m1_subset_1(k1_seq_1(X3,X4),k1_numbers))).#trainpos
cnf(c_0_49, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_71, plain, (k3_xcmplx_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5),k1_seq_1(X3,X4))=k1_seq_1(X3,k5_prob_1(X1,X2,X5,X4))|r1_xxreal_0(k1_seq_1(X3,X4),k1_xboole_0)|v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_prob_1(X5,X1,X2)|~m1_prob_1(X4,X1,X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_xreal_0(k1_seq_1(X3,X4))|~m1_subset_1(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5),k1_numbers))).#trainpos
cnf(c_0_190, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1))).#trainpos
cnf(c_0_234, plain, (m1_subset_1(k8_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_185, plain, (v2_xxreal_0(X2)|v2_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(k3_xcmplx_0(X1,X2)))).#trainpos
cnf(c_0_172, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(k3_xcmplx_0(X1,X2)))).#trainpos
cnf(c_0_164, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainpos
cnf(c_0_228, negated_conjecture, (v1_relat_1(esk3_0))).#trainpos
cnf(c_0_200, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainpos
cnf(c_0_145, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainpos
cnf(c_0_163, plain, (v5_relat_1(k3_prob_2(X1,X2,X3,X4),k1_numbers)|v1_xboole_0(X1)|v1_xboole_0(X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainpos
cnf(c_0_227, negated_conjecture, (v3_valued_0(esk3_0))).#trainpos
cnf(c_0_187, plain, (v3_xxreal_0(k1_xboole_0)|~v2_xxreal_0(k1_xboole_0))).#trainpos
cnf(c_0_197, plain, (~v2_xxreal_0(k1_xboole_0))).#trainpos
cnf(c_0_203, negated_conjecture, (v3_xxreal_0(k1_xboole_0)|v2_xxreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_xreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0))))).#trainpos
cnf(c_0_196, plain, (v2_xxreal_0(k1_xboole_0)|~v3_xxreal_0(k1_xboole_0))).#trainpos
cnf(c_0_240, plain, (m1_subset_1(k3_xboole_0(X1,X2),k1_zfmisc_1(X1)))).#trainpos
cnf(c_0_107, plain, (v1_xreal_0(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainpos
cnf(c_0_202, negated_conjecture, (v2_xxreal_0(k1_seq_1(X1,esk4_0))|~v2_xxreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,esk4_0))|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_tarski(X2,esk4_0)|~m1_prob_1(X2,esk1_0,esk2_0)|~m2_prob_1(X1,esk1_0,esk2_0))).#trainpos
cnf(c_0_213, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~r1_tarski(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk4_0)|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0))).#trainpos
cnf(c_0_225, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk4_0)))).#trainpos
cnf(c_0_238, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk4_0)))).#trainpos
cnf(c_0_246, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_prob_1(k3_xboole_0(esk4_0,esk5_0),esk1_0,esk2_0))).#trainpos
cnf(c_0_250, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0)))).#trainpos
cnf(c_0_54, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainpos
cnf(c_0_70, negated_conjecture, (k8_real_1(k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainpos
cnf(c_0_82, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0)))).#trainpos
cnf(c_0_98, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainpos
cnf(c_0_106, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers))).#trainpos
cnf(c_0_114, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_funct_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v3_valued_0(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0))))).#trainpos
cnf(c_0_126, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v3_valued_0(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0))).#trainpos
cnf(c_0_143, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v5_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers))).#trainpos
cnf(c_0_162, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v5_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainpos
cnf(c_0_178, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainpos
cnf(c_0_243, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_249, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainpos
cnf(c_0_179, plain, (v1_relat_1(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainpos
cnf(c_0_189, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainpos
cnf(c_0_199, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_funct_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~v3_valued_0(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0)))).#trainpos
cnf(c_0_207, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v3_valued_0(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainpos
cnf(c_0_208, plain, (v3_valued_0(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainpos
cnf(c_0_220, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainpos
cnf(c_0_233, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainpos
cnf(c_0_242, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainpos
cnf(c_0_248, negated_conjecture, (~v2_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainpos
cnf(c_0_252, negated_conjecture, (~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0))).#trainpos
cnf(c_0_254, negated_conjecture, (~v1_xreal_0(k1_seq_1(esk3_0,esk4_0)))).#trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_276, plain, (k4_ordinal1=k5_numbers)).#trainneg
cnf(i_0_209, plain, (~v1_xboole_0(esk11_0))).#trainneg
cnf(i_0_229, plain, (~v1_xboole_0(esk20_0))).#trainneg
cnf(i_0_233, plain, (~v1_xboole_0(esk21_0))).#trainneg
cnf(i_0_240, plain, (~v1_xboole_0(esk24_0))).#trainneg
cnf(i_0_249, plain, (~v1_xboole_0(esk27_0))).#trainneg
cnf(i_0_274, plain, (~v1_xboole_0(esk37_0))).#trainneg
cnf(i_0_225, plain, (v1_xboole_0(esk17_0))).#trainneg
cnf(i_0_179, plain, (~v1_xboole_0(k4_numbers))).#trainneg
cnf(i_0_174, plain, (~v1_xboole_0(k3_numbers))).#trainneg
cnf(i_0_113, plain, (v6_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_30, plain, (v6_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_272, plain, (v1_xboole_0(esk36_0))).#trainneg
cnf(i_0_116, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1))).#trainneg
cnf(i_0_57, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|v10_valued_0(X3)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_137, plain, (v1_xboole_0(X1)|m1_subset_1(X3,k1_zfmisc_1(X2))|~v1_prob_1(X1,X2)|~v4_prob_1(X1,X2)|~m1_prob_1(X3,X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_140, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|v1_funct_2(X1,X2,k1_numbers)|~v1_prob_1(X2,X3)|~v4_prob_1(X2,X3)|~m2_prob_1(X1,X3,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3))))).#trainneg
cnf(i_0_327, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_152, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_185, plain, (v5_membered(k4_numbers))).#trainneg
cnf(i_0_196, plain, (~v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_224, plain, (v1_relat_1(esk16_0))).#trainneg
cnf(i_0_135, plain, (m2_subset_1(k6_numbers,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_288, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_194, plain, (~v1_finset_1(k4_numbers))).#trainneg
cnf(i_0_319, negated_conjecture, (v10_valued_0(esk3_0))).#trainneg
cnf(i_0_193, plain, (v6_membered(k4_ordinal1))).#trainneg
cnf(i_0_208, plain, (v6_membered(esk11_0))).#trainneg
cnf(i_0_207, plain, (k5_prob_1(X2,X1,X3,X3)=X3|v1_xboole_0(X1)|~v1_prob_1(X1,X2)|~v4_prob_1(X1,X2)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_328, negated_conjecture, (v1_funct_2(esk3_0,esk2_0,k1_numbers))).#trainneg
cnf(i_0_228, plain, (v6_membered(esk20_0))).#trainneg
cnf(i_0_333, negated_conjecture, (k8_real_1(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~m1_subset_1(esk5_0,esk2_0)|~m1_subset_1(esk4_0,esk2_0))).#trainneg
cnf(i_0_331, negated_conjecture, (k5_prob_1(X1,k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X2,X2)=X2|v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X1)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X1))).#trainneg
cnf(i_0_354, plain, (m1_subset_1(k3_xboole_0(X1,X2),X3)|v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X4)))|~v4_prob_1(X3,X4)|~v1_prob_1(X3,X4)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_198, plain, (~v1_finset_1(k3_numbers))).#trainneg
cnf(i_0_204, plain, (~v1_finset_1(k1_numbers))).#trainneg
cnf(i_0_337, plain, (r1_xxreal_0(k1_seq_1(X1,X2),k1_seq_1(X1,X3))|v1_xboole_0(X4)|v1_xboole_0(X5)|~r1_tarski(X2,X3)|~m1_prob_1(X2,X4,X5)|~m2_prob_1(X1,X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(k1_zfmisc_1(X4)))|~v4_prob_1(X5,X4)|~v1_prob_1(X5,X4)|~m1_subset_1(X3,X5))).#trainneg
cnf(i_0_338, plain, (k5_prob_1(X1,X2,X4,X3)=k3_xboole_0(X3,X4)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_332, negated_conjecture, (~r1_xxreal_0(k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0)),k1_xboole_0)|~m1_subset_1(esk5_0,esk2_0)|~m1_subset_1(esk4_0,esk2_0))).#trainneg
cnf(i_0_334, negated_conjecture, (r1_xxreal_0(k1_seq_1(X1,X2),k1_seq_1(X1,esk6_0))|~r1_tarski(X2,esk6_0)|~m1_prob_1(X2,esk1_0,esk2_0)|~m2_prob_1(X1,esk1_0,esk2_0))).#trainneg
cnf(i_0_335, negated_conjecture, (r1_xxreal_0(k1_seq_1(X1,X2),k1_seq_1(X1,esk5_0))|~r1_tarski(X2,esk5_0)|~m1_prob_1(X2,esk1_0,esk2_0)|~m2_prob_1(X1,esk1_0,esk2_0))).#trainneg
cnf(i_0_342, plain, (r1_xxreal_0(k1_seq_1(X1,X2),k1_seq_1(X1,k5_prob_1(X3,X4,X5,X6)))|v1_xboole_0(X3)|v1_xboole_0(X4)|~r1_tarski(X2,k5_prob_1(X3,X4,X5,X6))|~m1_prob_1(X2,X3,X4)|~m2_prob_1(X1,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3)))|~v4_prob_1(X4,X3)|~v1_prob_1(X4,X3)|~m1_subset_1(X6,X4)|~m1_subset_1(X5,X4))).#trainneg
cnf(i_0_348, plain, (v1_funct_2(k3_prob_2(X1,X2,X3,X4),X2,k1_numbers)|v1_xboole_0(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_363, negated_conjecture, (m1_subset_1(k3_xboole_0(X1,X2),esk2_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_248, plain, (v6_membered(esk27_0))).#trainneg
cnf(i_0_341, plain, (m1_subset_1(k5_prob_1(X1,X2,X3,X4),k1_zfmisc_1(X1))|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_244, plain, (v1_xcmplx_0(esk25_0))).#trainneg
cnf(i_0_347, plain, (v10_valued_0(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_355, plain, (m1_subset_1(k3_xboole_0(X1,X2),k1_zfmisc_1(X3))|v1_xboole_0(X4)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3)))|~v4_prob_1(X4,X3)|~v1_prob_1(X4,X3)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X4))).#trainneg
cnf(i_0_369, negated_conjecture, (r1_xxreal_0(k1_seq_1(X1,esk6_0),k1_seq_1(X1,X2))|~r1_tarski(esk6_0,X2)|~m2_prob_1(X1,esk1_0,esk2_0)|~m1_subset_1(X2,esk2_0))).#trainneg
cnf(i_0_370, negated_conjecture, (r1_xxreal_0(k1_seq_1(X1,esk5_0),k1_seq_1(X1,X2))|~r1_tarski(esk5_0,X2)|~m2_prob_1(X1,esk1_0,esk2_0)|~m1_subset_1(X2,esk2_0))).#trainneg
cnf(i_0_371, negated_conjecture, (r1_xxreal_0(k1_seq_1(X1,esk4_0),k1_seq_1(X1,X2))|~r1_tarski(esk4_0,X2)|~m2_prob_1(X1,esk1_0,esk2_0)|~m1_subset_1(X2,esk2_0))).#trainneg
cnf(i_0_373, plain, (r1_xxreal_0(k1_seq_1(X1,k5_prob_1(X2,X3,X4,X5)),k1_seq_1(X1,X6))|v1_xboole_0(X3)|v1_xboole_0(X2)|~r1_tarski(k5_prob_1(X2,X3,X4,X5),X6)|~m2_prob_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X6,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_374, plain, (r1_xxreal_0(k1_seq_1(X1,X2),k1_seq_1(X1,X3))|v1_xboole_0(X4)|v1_xboole_0(X5)|~r1_tarski(X2,X3)|~m2_prob_1(X1,X5,X4)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X5)))|~m1_subset_1(X3,X4)|~v4_prob_1(X4,X5)|~v1_prob_1(X4,X5)|~m1_subset_1(X2,X4))).#trainneg
cnf(i_0_389, negated_conjecture, (k5_prob_1(X1,esk2_0,X2,X2)=X2|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,esk2_0)|~v4_prob_1(esk2_0,X1)|~v1_prob_1(esk2_0,X1)|~m1_subset_1(X4,esk2_0)|~m1_subset_1(X3,esk2_0))).#trainneg
cnf(i_0_154, plain, (~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_356, plain, (r1_xxreal_0(k1_seq_1(X1,X2),k1_seq_1(X1,k3_xboole_0(X3,X4)))|v1_xboole_0(X5)|v1_xboole_0(X6)|~r1_tarski(X2,k3_xboole_0(X3,X4))|~m1_prob_1(X2,X5,X6)|~m2_prob_1(X1,X5,X6)|~m1_subset_1(X6,k1_zfmisc_1(k1_zfmisc_1(X5)))|~v4_prob_1(X6,X5)|~v1_prob_1(X6,X5)|~m1_subset_1(X4,X6)|~m1_subset_1(X3,X6))).#trainneg
cnf(i_0_269, plain, (v1_xcmplx_0(esk35_0))).#trainneg
cnf(i_0_227, plain, (v1_xxreal_0(esk19_0))).#trainneg
cnf(i_0_361, negated_conjecture, (m1_prob_1(X2,X1,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X1)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X1))).#trainneg
cnf(i_0_359, negated_conjecture, (k3_xboole_0(X2,X2)=X2|v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X1)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X1))).#trainneg
cnf(i_0_360, negated_conjecture, (k5_prob_1(X1,k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X2,X2)=X2|v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X1)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X1))).#trainneg
cnf(i_0_405, plain, (r1_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5),k1_seq_1(k3_prob_2(X1,X2,X3,X4),X6))|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_tarski(X5,X6)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_406, negated_conjecture, (k5_prob_1(X1,esk2_0,X2,X2)=X2|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X3,esk2_0)|~v4_prob_1(esk2_0,X1)|~v1_prob_1(esk2_0,X1))).#trainneg
cnf(i_0_243, plain, (v1_xxreal_0(esk25_0))).#trainneg
cnf(i_0_246, plain, (v1_xxreal_0(esk26_0))).#trainneg
cnf(i_0_261, plain, (v1_xxreal_0(esk31_0))).#trainneg
cnf(i_0_268, plain, (v1_xxreal_0(esk35_0))).#trainneg
cnf(i_0_271, plain, (v1_xxreal_0(esk36_0))).#trainneg
cnf(i_0_416, plain, (r1_xxreal_0(k1_seq_1(X1,X2),k1_seq_1(X1,k3_xboole_0(X3,X4)))|v1_xboole_0(X5)|v1_xboole_0(X6)|~r1_tarski(X2,k3_xboole_0(X3,X4))|~m2_prob_1(X1,X6,X5)|~m1_subset_1(X5,k1_zfmisc_1(k1_zfmisc_1(X6)))|~m1_subset_1(X4,X5)|~m1_subset_1(X3,X5)|~v4_prob_1(X5,X6)|~v1_prob_1(X5,X6)|~m1_subset_1(X2,X5))).#trainneg
cnf(i_0_415, plain, (r1_xxreal_0(k1_seq_1(X1,k5_prob_1(X2,X3,X4,X5)),k1_seq_1(X1,k3_xboole_0(X6,X7)))|v1_xboole_0(X3)|v1_xboole_0(X2)|~r1_tarski(k5_prob_1(X2,X3,X4,X5),k3_xboole_0(X6,X7))|~m2_prob_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X7,X3)|~m1_subset_1(X6,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_213, plain, (~v1_xboole_0(esk13_1(X1)))).#trainneg
cnf(i_0_372, plain, (r1_xxreal_0(k1_seq_1(X1,k3_xboole_0(X2,X3)),k1_seq_1(X1,X4))|v1_xboole_0(X5)|v1_xboole_0(X6)|~r1_tarski(k3_xboole_0(X2,X3),X4)|~m2_prob_1(X1,X6,X5)|~m1_subset_1(X5,k1_zfmisc_1(k1_zfmisc_1(X6)))|~m1_subset_1(X4,X5)|~v4_prob_1(X5,X6)|~v1_prob_1(X5,X6)|~m1_subset_1(X3,X5)|~m1_subset_1(X2,X5))).#trainneg
cnf(i_0_226, plain, (v1_xreal_0(esk18_0))).#trainneg
cnf(i_0_241, plain, (v1_xreal_0(esk25_0))).#trainneg
cnf(i_0_392, plain, (k5_prob_1(X1,k1_zfmisc_1(X2),X3,X3)=X3|v1_xboole_0(k1_zfmisc_1(X2))|v1_xboole_0(X4)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k1_zfmisc_1(X2),X1)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X6,X4)|~m1_subset_1(X5,X4)|~v4_prob_1(X4,X2)|~v1_prob_1(X4,X2))).#trainneg
cnf(i_0_178, plain, (v4_membered(k3_numbers))).#trainneg
cnf(i_0_222, plain, (v4_valued_0(esk16_0))).#trainneg
cnf(i_0_262, plain, (v7_ordinal1(esk32_0))).#trainneg
cnf(i_0_273, plain, (v7_ordinal1(esk37_0))).#trainneg
cnf(i_0_195, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_210, plain, (v3_ordinal1(esk12_0))).#trainneg
cnf(i_0_230, plain, (v3_ordinal1(esk21_0))).#trainneg
cnf(i_0_232, plain, (v1_ordinal1(esk21_0))).#trainneg
cnf(i_0_231, plain, (v2_ordinal1(esk21_0))).#trainneg
cnf(i_0_183, plain, (v7_membered(k4_numbers))).#trainneg
cnf(i_0_182, plain, (v7_membered(k3_numbers))).#trainneg
cnf(i_0_236, plain, (~v1_xboole_0(esk22_1(X1)))).#trainneg
cnf(i_0_388, plain, (k5_prob_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)),X3,X3)=X3|v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|v1_xboole_0(X4)|v1_xboole_0(X2)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v4_prob_1(k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)),X1)|~v1_prob_1(k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)),X1)|~m2_prob_1(X5,X4,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X4)))|~m1_subset_1(X6,X2)|~v4_prob_1(X2,X4)|~v1_prob_1(X2,X4))).#trainneg
cnf(i_0_414, plain, (r1_xxreal_0(k1_seq_1(X1,k3_xboole_0(X2,X3)),k1_seq_1(X1,k3_xboole_0(X4,X5)))|v1_xboole_0(X6)|v1_xboole_0(X7)|~r1_tarski(k3_xboole_0(X2,X3),k3_xboole_0(X4,X5))|~m2_prob_1(X1,X7,X6)|~m1_subset_1(X6,k1_zfmisc_1(k1_zfmisc_1(X7)))|~m1_subset_1(X5,X6)|~m1_subset_1(X4,X6)|~v4_prob_1(X6,X7)|~v1_prob_1(X6,X7)|~m1_subset_1(X3,X6)|~m1_subset_1(X2,X6))).#trainneg
cnf(i_0_181, plain, (v7_membered(k1_numbers))).#trainneg
cnf(i_0_184, plain, (v7_membered(k4_ordinal1))).#trainneg
cnf(i_0_437, plain, (r1_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5),k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X6,X7)))|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_tarski(X5,k3_xboole_0(X6,X7))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X7,X2)|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_247, plain, (v7_membered(esk27_0))).#trainneg
cnf(i_0_292, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_242, plain, (v2_xxreal_0(esk25_0))).#trainneg
cnf(i_0_245, plain, (v2_xxreal_0(esk26_0))).#trainneg
cnf(i_0_260, plain, (v3_xxreal_0(esk31_0))).#trainneg
cnf(i_0_223, plain, (v1_funct_1(esk16_0))).#trainneg
cnf(i_0_238, plain, (v1_xboole_0(esk23_1(X1)))).#trainneg
cnf(i_0_253, plain, (v1_relat_1(esk28_1(X1)))).#trainneg
cnf(i_0_153, plain, (v2_finsub_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_212, plain, (v2_finsub_1(esk13_1(X1)))).#trainneg
cnf(i_0_462, plain, (esk36_0=k1_xboole_0)).#trainneg
cnf(i_0_271, plain, (v1_xxreal_0(k1_xboole_0))).#trainneg
cnf(i_0_269, plain, (v1_xcmplx_0(k1_xboole_0))).#trainneg
cnf(i_0_464, plain, (esk17_0=k1_xboole_0)).#trainneg
cnf(i_0_295, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_290, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_465, plain, (esk23_1(X1)=k1_xboole_0)).#trainneg
cnf(i_0_446, plain, (r1_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X5,X6)),k1_seq_1(k3_prob_2(X1,X2,X3,X4),X7))|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_tarski(k3_xboole_0(X5,X6),X7)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X7,X2)|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_250, plain, (v5_ordinal1(esk28_1(X1)))).#trainneg
cnf(i_0_251, plain, (v1_funct_1(esk28_1(X1)))).#trainneg
cnf(i_0_287, plain, (k3_xcmplx_0(np__1,np__1)=np__1)).#trainneg
cnf(i_0_286, plain, (r1_xxreal_0(np__1,np__1))).#trainneg
cnf(i_0_34, plain, (v5_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_88, plain, (v1_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_79, plain, (v2_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_85, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_69, plain, (v3_membered(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_474, plain, (v1_xxreal_0(esk37_0))).#trainneg
cnf(i_0_475, plain, (v1_xxreal_0(esk32_0))).#trainneg
cnf(i_0_476, plain, (v3_membered(k3_numbers))).#trainneg
cnf(i_0_459, plain, (r1_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X5,X6)),k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X7,X8)))|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_tarski(k3_xboole_0(X5,X6),k3_xboole_0(X7,X8))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X8,X2)|~m1_subset_1(X7,X2)|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_133, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_291, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_63, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_55, plain, (v4_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_70, plain, (v3_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_95, plain, (v3_ordinal1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_36, plain, (v1_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_35, plain, (v2_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_54, plain, (v7_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_80, plain, (v5_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_300, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_284, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_206, plain, (k3_xboole_0(X1,X1)=X1)).#trainneg
cnf(i_0_483, plain, (v1_xreal_0(esk37_0))).#trainneg
cnf(i_0_484, plain, (v1_xreal_0(esk32_0))).#trainneg
cnf(i_0_488, plain, (v3_ordinal1(esk37_0))).#trainneg
cnf(i_0_485, plain, (v3_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_489, plain, (v3_ordinal1(esk32_0))).#trainneg
cnf(i_0_509, plain, (v3_membered(k4_numbers))).#trainneg
cnf(i_0_324, negated_conjecture, (m1_subset_1(esk6_0,k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_487, plain, (v1_xreal_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_316, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_96, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_90, plain, (v2_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_108, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_114, plain, (v4_valued_0(X1)|~v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_56, plain, (v3_ordinal1(X1)|~v1_ordinal1(X1)|~v2_ordinal1(X1))).#trainneg
cnf(i_0_162, plain, (v1_prob_1(k1_zfmisc_1(X1),X1))).#trainneg
cnf(i_0_514, plain, (v3_valued_0(esk16_0))).#trainneg
cnf(i_0_519, plain, (k3_xcmplx_0(esk25_0,k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_325, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_326, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_510, plain, (v3_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_515, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_491, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0)).#trainneg
cnf(i_0_211, plain, (v1_prob_1(esk13_1(X1),X1))).#trainneg
cnf(i_0_235, plain, (v1_prob_1(esk22_1(X1),X1))).#trainneg
cnf(i_0_180, plain, (v4_prob_1(k1_zfmisc_1(X1),X1))).#trainneg
cnf(i_0_234, plain, (v4_prob_1(esk22_1(X1),X1))).#trainneg
cnf(i_0_144, plain, (m1_subset_1(esk8_1(X1),X1))).#trainneg
cnf(i_0_252, plain, (v5_relat_1(esk28_1(X1),X1))).#trainneg
cnf(i_0_86, plain, (v1_xboole_0(X1)|v2_xxreal_0(X1)|v3_xxreal_0(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_220, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk15_1(X1)))).#trainneg
cnf(i_0_78, plain, (~v1_xboole_0(X1)|~v1_xxreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_548, plain, (~v1_xboole_0(esk26_0))).#trainneg
cnf(i_0_549, plain, (~v1_xboole_0(esk25_0))).#trainneg
cnf(i_0_520, plain, (k3_xcmplx_0(k1_xboole_0,k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_544, plain, (k3_xcmplx_0(np__1,esk30_0)=esk30_0)).#trainneg
cnf(i_0_545, plain, (k3_xcmplx_0(np__1,esk25_0)=esk25_0)).#trainneg
cnf(i_0_542, plain, (k5_prob_1(X1,X2,X3,X3)=X3|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_561, plain, (~v1_xboole_0(esk31_0))).#trainneg
cnf(i_0_562, plain, (~v1_xboole_0(esk30_0))).#trainneg
cnf(i_0_564, plain, (~v2_xxreal_0(esk31_0))).#trainneg
cnf(i_0_53, plain, (v6_membered(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_546, plain, (k3_xcmplx_0(np__1,k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_456, negated_conjecture, (k5_prob_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),X2,X2)=X2|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)))|~m1_subset_1(X3,esk2_0)|~v4_prob_1(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),X1)|~v1_prob_1(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),X1))).#trainneg
cnf(i_0_521, plain, (k3_xcmplx_0(X1,k1_xboole_0)=k1_xboole_0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_254, plain, (~v1_subset_1(esk29_1(X1),X1))).#trainneg
cnf(i_0_314, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_107, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_265, plain, (v1_xboole_0(X1)|v1_subset_1(esk34_1(X1),X1))).#trainneg
cnf(i_0_239, plain, (m1_subset_1(esk23_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_255, plain, (m1_subset_1(esk29_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_43, plain, (v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_101, plain, (v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_285, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_218, plain, (v1_xboole_0(esk14_2(X1,X2)))).#trainneg
cnf(i_0_217, plain, (v1_relat_1(esk14_2(X1,X2)))).#trainneg
cnf(i_0_221, plain, (v1_xboole_0(X1)|m1_subset_1(esk15_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_266, plain, (v1_xboole_0(X1)|m1_subset_1(esk34_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_13, plain, (v5_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_582, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_578, plain, (v5_relat_1(esk16_0,k4_numbers))).#trainneg
cnf(i_0_580, plain, (v5_relat_1(esk16_0,k3_numbers))).#trainneg
cnf(i_0_593, plain, (X1=esk14_2(X2,X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_573, plain, (v1_xxreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_12, plain, (~r2_hidden(X2,X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_311, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_574, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_575, plain, (v3_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_16, plain, (v6_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_106, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_112, plain, (v4_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_73, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_83, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k3_numbers))).#trainneg
cnf(i_0_67, plain, (v4_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_18, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_20, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_22, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_24, plain, (v1_rat_1(X2)|~v4_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_26, plain, (v1_int_1(X2)|~v5_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_671, plain, (v1_xreal_0(np__1))).#trainneg
cnf(i_0_596, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_684, negated_conjecture, (v1_rat_1(esk6_0)|~v4_membered(esk2_0))).#trainneg
cnf(i_0_629, plain, (v1_xcmplx_0(np__1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_722, plain, (v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_621, negated_conjecture, (v1_xcmplx_0(esk6_0)|~v1_membered(esk2_0))).#trainneg
cnf(i_0_691, negated_conjecture, (v1_rat_1(esk2_0)|~v4_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_712, negated_conjecture, (v1_int_1(esk2_0)|~v5_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_28, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_89, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_297, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_175, plain, (v5_membered(k3_xboole_0(X2,X1))|~v5_membered(X1))).#trainneg
cnf(i_0_172, plain, (v5_membered(k3_xboole_0(X1,X2))|~v5_membered(X1))).#trainneg
cnf(i_0_199, plain, (v1_valued_0(k2_zfmisc_1(X2,X1))|~v1_membered(X1))).#trainneg
cnf(i_0_177, plain, (v6_membered(k3_xboole_0(X2,X1))|~v6_membered(X1))).#trainneg
cnf(i_0_176, plain, (v6_membered(k3_xboole_0(X1,X2))|~v6_membered(X1))).#trainneg
cnf(i_0_200, plain, (v2_valued_0(k2_zfmisc_1(X2,X1))|~v2_membered(X1))).#trainneg
cnf(i_0_165, plain, (v1_membered(k3_xboole_0(X2,X1))|~v1_membered(X1))).#trainneg
cnf(i_0_164, plain, (v1_membered(k3_xboole_0(X1,X2))|~v1_membered(X1))).#trainneg
cnf(i_0_201, plain, (v3_valued_0(k2_zfmisc_1(X2,X1))|~v3_membered(X1))).#trainneg
cnf(i_0_739, plain, (v7_ordinal1(np__1))).#trainneg
cnf(i_0_740, plain, (v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_760, plain, (v3_ordinal1(np__1))).#trainneg
cnf(i_0_579, plain, (v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_761, plain, (v3_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_800, plain, (v1_xxreal_0(np__1))).#trainneg
cnf(i_0_737, negated_conjecture, (v7_ordinal1(esk2_0)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_808, negated_conjecture, (v3_ordinal1(esk2_0)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_167, plain, (v2_membered(k3_xboole_0(X2,X1))|~v2_membered(X1))).#trainneg
cnf(i_0_166, plain, (v2_membered(k3_xboole_0(X1,X2))|~v2_membered(X1))).#trainneg
cnf(i_0_169, plain, (v3_membered(k3_xboole_0(X2,X1))|~v3_membered(X1))).#trainneg
cnf(i_0_168, plain, (v3_membered(k3_xboole_0(X1,X2))|~v3_membered(X1))).#trainneg
cnf(i_0_171, plain, (v4_membered(k3_xboole_0(X2,X1))|~v4_membered(X1))).#trainneg
cnf(i_0_170, plain, (v4_membered(k3_xboole_0(X1,X2))|~v4_membered(X1))).#trainneg
cnf(i_0_205, plain, (v4_valued_0(k2_zfmisc_1(X2,X1))|~v6_membered(X1))).#trainneg
cnf(i_0_214, plain, (m1_subset_1(esk13_1(X1),k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_237, plain, (m1_subset_1(esk22_1(X1),k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_215, plain, (v5_relat_1(esk14_2(X1,X2),X2))).#trainneg
cnf(i_0_613, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_686, negated_conjecture, (v1_rat_1(esk5_0)|~v4_membered(esk2_0))).#trainneg
cnf(i_0_447, negated_conjecture, (k5_prob_1(X1,k1_zfmisc_1(X2),X3,X3)=X3|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X4,esk2_0)|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v4_prob_1(esk2_0,X2)|~v1_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(esk2_0,X2))).#trainneg
cnf(i_0_581, plain, (v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_623, negated_conjecture, (v1_xcmplx_0(esk5_0)|~v1_membered(esk2_0))).#trainneg
cnf(i_0_625, negated_conjecture, (v1_xcmplx_0(esk4_0)|~v1_membered(esk2_0))).#trainneg
cnf(i_0_628, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_649, negated_conjecture, (v1_xxreal_0(esk2_0)|~v2_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_216, plain, (v4_relat_1(esk14_2(X1,X2),X1))).#trainneg
cnf(i_0_61, plain, (v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_302, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_117, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_120, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_42, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_51, plain, (v5_membered(X2)|~v5_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_14, plain, (v6_ordinal1(X2)|~v6_ordinal1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_52, plain, (v6_membered(X2)|~v6_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_32, plain, (v1_membered(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_888, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_902, plain, (v1_xboole_0(esk29_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_688, negated_conjecture, (v1_rat_1(esk4_0)|~v4_membered(esk2_0))).#trainneg
cnf(i_0_592, plain, (esk14_2(X1,X2)=k1_xboole_0)).#trainneg
cnf(i_0_217, plain, (v1_relat_1(k1_xboole_0))).#trainneg
cnf(i_0_642, negated_conjecture, (v1_xxreal_0(esk6_0)|~v2_membered(esk2_0))).#trainneg
cnf(i_0_984, plain, (v1_rat_1(X1)|~v4_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_985, plain, (v1_int_1(X1)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_614, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_615, plain, (v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_46, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_48, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_50, plain, (v4_membered(X2)|~v4_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_197, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_203, plain, (v5_relat_1(k2_zfmisc_1(X2,X1),k4_numbers)|~v5_membered(X1))).#trainneg
cnf(i_0_202, plain, (v5_relat_1(k2_zfmisc_1(X2,X1),k3_numbers)|~v4_membered(X1))).#trainneg
cnf(i_0_301, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(X2)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_310, plain, (v1_xboole_0(X2)|v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_1029, plain, (v3_membered(k5_numbers))).#trainneg
cnf(i_0_692, plain, (v1_rat_1(np__1)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_216, plain, (v4_relat_1(k1_xboole_0,X1))).#trainneg
cnf(i_0_616, plain, (v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_644, negated_conjecture, (v1_xxreal_0(esk5_0)|~v2_membered(esk2_0))).#trainneg
cnf(i_0_646, negated_conjecture, (v1_xxreal_0(esk4_0)|~v2_membered(esk2_0))).#trainneg
cnf(i_0_986, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1114, plain, (v3_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_617, plain, (v1_xcmplx_0(esk34_1(X1))|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_618, plain, (v1_xcmplx_0(esk29_1(X1))|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_313, plain, (v1_xboole_0(X1)|v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_15, plain, (v1_valued_0(X2)|~v1_relat_1(X1)|~v1_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_17, plain, (v2_valued_0(X2)|~v1_relat_1(X1)|~v2_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_19, plain, (v3_valued_0(X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_25, plain, (v4_valued_0(X2)|~v1_relat_1(X1)|~v4_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_148, plain, (v1_valued_0(k3_xboole_0(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_149, plain, (v2_valued_0(k3_xboole_0(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v2_valued_0(X1))).#trainneg
cnf(i_0_187, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_valued_0(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_150, plain, (v3_valued_0(k3_xboole_0(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_188, plain, (v1_xxreal_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v2_valued_0(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_157, plain, (v4_valued_0(k3_xboole_0(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_875, plain, (r1_tarski(k1_xboole_0,X1))).#trainneg
cnf(i_0_693, plain, (v1_rat_1(np__1)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_215, plain, (v5_relat_1(k1_xboole_0,X2))).#trainneg
cnf(i_0_619, plain, (v1_xcmplx_0(esk15_1(X1))|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_663, negated_conjecture, (v1_xreal_0(esk6_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_665, negated_conjecture, (v1_xreal_0(esk5_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_1108, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X1)|v1_xboole_0(X1)|~v1_xreal_0(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_620, plain, (v1_xcmplx_0(esk8_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_634, plain, (v1_xcmplx_0(k5_numbers)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_192, plain, (v7_ordinal1(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v4_valued_0(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1264, plain, (k8_real_1(k1_seq_1(X3,X4),k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))=k1_seq_1(X3,k5_prob_1(X1,X2,X5,X4))|r1_xxreal_0(k1_seq_1(X3,X4),k1_xboole_0)|v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_prob_1(X5,X1,X2)|~m1_prob_1(X4,X1,X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_xreal_0(k1_seq_1(X3,X4))|~m1_subset_1(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5),k1_numbers))).#trainneg
cnf(i_0_1265, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_1266, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_1267, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_1268, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_1272, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_694, plain, (v1_rat_1(k1_xboole_0)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_906, plain, (v1_xboole_0(esk8_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_705, negated_conjecture, (v1_int_1(esk6_0)|~v5_membered(esk2_0))).#trainneg
cnf(i_0_1259, negated_conjecture, (k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers))).#trainneg
cnf(i_0_635, plain, (v1_xcmplx_0(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_713, plain, (v1_int_1(np__1)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_667, negated_conjecture, (v1_xreal_0(esk4_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_707, negated_conjecture, (v1_int_1(esk5_0)|~v5_membered(esk2_0))).#trainneg
cnf(i_0_899, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1319, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_636, plain, (v1_xcmplx_0(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v1_membered(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_637, plain, (v1_xcmplx_0(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1276, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_1282, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_1286, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_1333, plain, (k3_xcmplx_0(X1,k5_prob_1(X2,X3,X4,X5))=k3_xcmplx_0(k5_prob_1(X2,X3,X4,X5),X1)|v1_xboole_0(X3)|~v1_xcmplx_0(X1)|~v1_membered(X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_1334, plain, (k3_xcmplx_0(k5_prob_1(X1,X2,X3,X4),k1_xboole_0)=k1_xboole_0|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1335, plain, (k3_xcmplx_0(np__1,k5_prob_1(X1,X2,X3,X4))=k5_prob_1(X1,X2,X3,X4)|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_714, plain, (v1_int_1(np__1)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_999, plain, (v5_relat_1(esk28_1(k5_numbers),k4_numbers))).#trainneg
cnf(i_0_709, negated_conjecture, (v1_int_1(esk4_0)|~v5_membered(esk2_0))).#trainneg
cnf(i_0_1262, negated_conjecture, (k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_numbers))).#trainneg
cnf(i_0_638, plain, (v1_xxreal_0(esk34_1(X1))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_715, plain, (v1_int_1(k1_xboole_0)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_725, negated_conjecture, (v1_xcmplx_0(esk6_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_730, negated_conjecture, (v7_ordinal1(esk6_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_1318, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_639, plain, (v1_xxreal_0(esk29_1(X1))|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_640, plain, (v1_xxreal_0(esk15_1(X1))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1360, plain, (k3_xcmplx_0(X1,X4)=k3_xcmplx_0(X4,X1)|v1_xboole_0(X3)|~v1_xcmplx_0(X1)|~v1_membered(X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_1368, plain, (k3_xcmplx_0(k1_xboole_0,k5_prob_1(X1,X2,X3,X4))=k1_xboole_0|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1370, plain, (k3_xcmplx_0(k3_xboole_0(X3,X4),k1_xboole_0)=k1_xboole_0|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1377, plain, (k3_xcmplx_0(np__1,k3_xboole_0(X3,X4))=k3_xboole_0(X3,X4)|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1109, plain, (v5_relat_1(esk28_1(k5_numbers),k3_numbers))).#trainneg
cnf(i_0_732, negated_conjecture, (v7_ordinal1(esk5_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_1260, negated_conjecture, (k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers))).#trainneg
cnf(i_0_641, plain, (v1_xxreal_0(esk8_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_734, negated_conjecture, (v7_ordinal1(esk4_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_751, negated_conjecture, (v3_ordinal1(esk6_0)|~v3_ordinal1(esk2_0))).#trainneg
cnf(i_0_898, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_655, plain, (v1_xxreal_0(k5_numbers)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_656, plain, (v1_xxreal_0(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1381, plain, (k3_xcmplx_0(np__1,X3)=X3|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1543, negated_conjecture, (k8_real_1(k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_1557, plain, (k3_xcmplx_0(k1_seq_1(X1,X2),k1_seq_1(k3_prob_2(X3,X4,X1,X2),X5))=k1_seq_1(X1,k5_prob_1(X3,X4,X5,X2))|r1_xxreal_0(k1_seq_1(X1,X2),k1_xboole_0)|v1_xboole_0(X4)|v1_xboole_0(X3)|~v1_xreal_0(k1_seq_1(k3_prob_2(X3,X4,X1,X2),X5))|~m1_prob_1(X5,X3,X4)|~m1_prob_1(X2,X3,X4)|~m2_prob_1(X1,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v4_prob_1(X4,X3)|~v1_prob_1(X4,X3))).#trainneg
cnf(i_0_147, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_769, plain, (v5_membered(k1_xboole_0)|~v5_membered(X1))).#trainneg
cnf(i_0_1571, plain, (v5_membered(k1_xboole_0))).#trainneg
cnf(i_0_1574, plain, (v3_membered(k1_xboole_0))).#trainneg
cnf(i_0_1263, negated_conjecture, (k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_numbers))).#trainneg
cnf(i_0_657, plain, (v1_xxreal_0(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v2_membered(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_929, plain, (v5_membered(k5_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_753, negated_conjecture, (v3_ordinal1(esk5_0)|~v3_ordinal1(esk2_0))).#trainneg
cnf(i_0_755, negated_conjecture, (v3_ordinal1(esk4_0)|~v3_ordinal1(esk2_0))).#trainneg
cnf(i_0_927, negated_conjecture, (v5_membered(esk2_0)|~v5_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1588, negated_conjecture, (v5_membered(esk2_0)|~v6_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_658, plain, (v1_xxreal_0(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v2_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_659, plain, (v1_xreal_0(esk34_1(X1))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_60, plain, (v1_xboole_0(X1)|v1_subset_1(X2,X1)|~v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_82, plain, (~v1_xboole_0(X1)|~v1_subset_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_190, plain, (v1_rat_1(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_relat_1(X1,k3_numbers))).#trainneg
cnf(i_0_191, plain, (v1_int_1(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_309, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_159, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(k3_xcmplx_0(X2,X1)))).#trainneg
cnf(i_0_161, plain, (v3_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(k3_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_1647, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_870, negated_conjecture, (v1_xcmplx_0(esk5_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_660, plain, (v1_xreal_0(esk29_1(X1))|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1661, plain, (v2_xxreal_0(k1_xboole_0)|~v3_xxreal_0(k1_xboole_0))).#trainneg
cnf(i_0_871, negated_conjecture, (v1_xcmplx_0(esk4_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_943, negated_conjecture, (v6_ordinal1(esk2_0)|~v6_ordinal1(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_661, plain, (v1_xreal_0(esk15_1(X1))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_662, plain, (v1_xreal_0(esk8_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_23, plain, (v5_relat_1(X2,k4_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_21, plain, (v5_relat_1(X2,k3_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k3_numbers)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_156, plain, (v5_relat_1(k3_xboole_0(X1,X2),k4_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_151, plain, (v5_relat_1(k3_xboole_0(X1,X2),k3_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v5_relat_1(X1,k3_numbers))).#trainneg
cnf(i_0_186, plain, (v5_relat_1(k3_xboole_0(X1,X3),X2)|~v1_relat_1(X3)|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).#trainneg
cnf(i_0_163, plain, (v4_relat_1(k3_xboole_0(X1,X3),X2)|~v1_relat_1(X3)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_293, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_39, plain, (v1_xboole_0(X1)|v2_finsub_1(X1)|~v1_prob_1(X1,X2)|~v4_prob_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_282, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_146, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m2_subset_1(esk10_2(X1,X2),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1637, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk29_1(X2)))).#trainneg
cnf(i_0_1638, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X1,esk22_1(X2)))).#trainneg
cnf(i_0_1818, negated_conjecture, (v2_finsub_1(esk2_0))).#trainneg
cnf(i_0_670, negated_conjecture, (v1_xreal_0(esk2_0)|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_1749, plain, (v5_relat_1(k5_numbers,k4_numbers)|~v5_relat_1(k1_numbers,k4_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_959, negated_conjecture, (v6_membered(esk2_0)|~v6_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_676, plain, (v1_xreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_677, plain, (v1_xreal_0(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_81, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_71, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_27, plain, (v1_valued_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_29, plain, (v2_valued_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_49, plain, (v4_valued_0(X2)|~v6_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_59, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1919, negated_conjecture, (v4_relat_1(esk3_0,esk2_0))).#trainneg
cnf(i_0_1921, plain, (v4_relat_1(k3_prob_2(X1,X2,X3,X4),X2)|v1_xboole_0(X1)|v1_xboole_0(X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_263, plain, (v1_xboole_0(X2)|v1_xboole_0(esk33_2(X1,X2))|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_1640, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X1,esk13_1(X2)))).#trainneg
cnf(i_0_1641, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk8_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_678, plain, (v1_xreal_0(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v3_membered(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1766, plain, (v5_relat_1(k5_numbers,k3_numbers)|~v5_relat_1(k1_numbers,k3_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_923, negated_conjecture, (v5_membered(esk6_0)|~v5_membered(esk1_0))).#trainneg
cnf(i_0_975, negated_conjecture, (v1_membered(esk2_0)|~v1_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1939, negated_conjecture, (v1_membered(esk2_0)|~v3_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_679, plain, (v1_xreal_0(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v3_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_680, plain, (v1_rat_1(esk34_1(X1))|v1_xboole_0(X1)|~v4_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_47, plain, (v5_relat_1(X2,k4_numbers)|~v5_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_33, plain, (v5_relat_1(X2,k3_numbers)|~v4_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_264, plain, (v1_xboole_0(X2)|m1_subset_1(esk33_2(X1,X2),X2)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_283, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_142, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_143, plain, (v1_xboole_0(X1)|m1_prob_1(esk7_2(X2,X1),X2,X1)|~v1_prob_1(X1,X2)|~v4_prob_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_145, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m2_prob_1(esk9_2(X1,X2),X1,X2)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_330, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(esk1_0))|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_398, negated_conjecture, (m1_subset_1(k3_xboole_0(X1,X2),k1_zfmisc_1(esk1_0))|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_404, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,X1),k1_seq_1(esk3_0,X2))|~r1_tarski(X1,X2)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_411, negated_conjecture, (r1_xxreal_0(k1_seq_1(X1,esk6_0),k1_seq_1(X1,k3_xboole_0(X2,X3)))|~r1_tarski(esk6_0,k3_xboole_0(X2,X3))|~m2_prob_1(X1,esk1_0,esk2_0)|~m1_subset_1(X3,esk2_0)|~m1_subset_1(X2,esk2_0))).#trainneg
cnf(i_0_1648, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_1864, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_924, negated_conjecture, (v5_membered(esk5_0)|~v5_membered(esk1_0))).#trainneg
cnf(i_0_681, plain, (v1_rat_1(esk29_1(X1))|~v4_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_945, plain, (v6_ordinal1(k5_numbers)|~v6_ordinal1(k1_numbers))).#trainneg
cnf(i_0_925, negated_conjecture, (v5_membered(esk4_0)|~v5_membered(esk1_0))).#trainneg
cnf(i_0_939, negated_conjecture, (v6_ordinal1(esk6_0)|~v6_ordinal1(esk1_0))).#trainneg
cnf(i_0_981, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_982, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_682, plain, (v1_rat_1(esk15_1(X1))|v1_xboole_0(X1)|~v4_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_683, plain, (v1_rat_1(esk8_1(X1))|~v4_membered(X1))).#trainneg
cnf(i_0_412, negated_conjecture, (r1_xxreal_0(k1_seq_1(X1,esk5_0),k1_seq_1(X1,k3_xboole_0(X2,X3)))|~r1_tarski(esk5_0,k3_xboole_0(X2,X3))|~m2_prob_1(X1,esk1_0,esk2_0)|~m1_subset_1(X3,esk2_0)|~m1_subset_1(X2,esk2_0))).#trainneg
cnf(i_0_413, negated_conjecture, (r1_xxreal_0(k1_seq_1(X1,esk4_0),k1_seq_1(X1,k3_xboole_0(X2,X3)))|~r1_tarski(esk4_0,k3_xboole_0(X2,X3))|~m2_prob_1(X1,esk1_0,esk2_0)|~m1_subset_1(X3,esk2_0)|~m1_subset_1(X2,esk2_0))).#trainneg
cnf(i_0_417, negated_conjecture, (r1_xxreal_0(k1_seq_1(X1,X2),k1_seq_1(X1,X3))|v1_xboole_0(X4)|~r1_tarski(X2,X3)|~m2_prob_1(X1,X4,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(X4)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X4)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X4)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_419, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X2)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_420, negated_conjecture, (r1_xxreal_0(k1_seq_1(X1,X2),k1_seq_1(X1,k3_xboole_0(X3,X4)))|v1_xboole_0(X5)|~r1_tarski(X2,k3_xboole_0(X3,X4))|~m2_prob_1(X1,X5,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(X5)))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X5)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X5)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_421, negated_conjecture, (r1_xxreal_0(k1_seq_1(X1,X2),k1_seq_1(X1,X3))|v1_xboole_0(X4)|~r1_tarski(X2,X3)|~m1_prob_1(X2,X4,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~m2_prob_1(X1,X4,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(X4)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X4)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X4)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_1646, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(esk1_0))|~r2_hidden(X1,esk2_0))).#trainneg
cnf(i_0_1829, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_940, negated_conjecture, (v6_ordinal1(esk5_0)|~v6_ordinal1(esk1_0))).#trainneg
cnf(i_0_1538, negated_conjecture, (k3_xcmplx_0(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_numbers))).#trainneg
cnf(i_0_695, plain, (v1_rat_1(k1_xboole_0)|~v4_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_977, plain, (v1_membered(k5_numbers)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_941, negated_conjecture, (v6_ordinal1(esk4_0)|~v6_ordinal1(esk1_0))).#trainneg
cnf(i_0_955, negated_conjecture, (v6_membered(esk6_0)|~v6_membered(esk1_0))).#trainneg
cnf(i_0_983, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_987, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_697, plain, (v1_rat_1(k5_numbers)|~v4_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_698, plain, (v1_rat_1(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v4_membered(k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_431, negated_conjecture, (k5_prob_1(X1,esk2_0,X2,X2)=X2|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,esk2_0)|~v4_prob_1(esk2_0,X1)|~v1_prob_1(esk2_0,X1))).#trainneg
cnf(i_0_436, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,X1),k1_seq_1(esk3_0,k3_xboole_0(X2,X3)))|~r1_tarski(X1,k3_xboole_0(X2,X3))|~m1_subset_1(X3,esk2_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_445, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,k3_xboole_0(X1,X2)),k1_seq_1(esk3_0,X3))|~r1_tarski(k3_xboole_0(X1,X2),X3)|~m1_subset_1(X3,esk2_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_450, negated_conjecture, (k5_prob_1(X1,k1_zfmisc_1(X2),X3,X3)=X3|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)))|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v4_prob_1(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),X2)|~v1_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),X2))).#trainneg
cnf(i_0_451, negated_conjecture, (k5_prob_1(X1,k1_zfmisc_1(X2),X3,X3)=X3|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X2)|~v1_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X2))).#trainneg
cnf(i_0_453, plain, (k5_prob_1(X1,k1_zfmisc_1(X2),X3,X3)=X3|v1_xboole_0(X5)|v1_xboole_0(X4)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X4,k1_numbers)),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X4,k1_numbers)))|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v4_prob_1(k1_zfmisc_1(k2_zfmisc_1(X4,k1_numbers)),X2)|~v1_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k1_zfmisc_1(k2_zfmisc_1(X4,k1_numbers)),X2)|~m2_prob_1(X6,X5,X4)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X5)))|~m1_subset_1(X7,X4)|~v4_prob_1(X4,X5)|~v1_prob_1(X4,X5))).#trainneg
cnf(i_0_1636, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk34_1(X2)))).#trainneg
cnf(i_0_1856, plain, (v1_xboole_0(esk29_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_956, negated_conjecture, (v6_membered(esk5_0)|~v6_membered(esk1_0))).#trainneg
cnf(i_0_1541, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers))).#trainneg
cnf(i_0_699, plain, (v1_rat_1(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v4_membered(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1012, plain, (v2_membered(k5_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_957, negated_conjecture, (v6_membered(esk4_0)|~v6_membered(esk1_0))).#trainneg
cnf(i_0_971, negated_conjecture, (v1_membered(esk6_0)|~v1_membered(esk1_0))).#trainneg
cnf(i_0_2160, plain, (v3_ordinal1(X1)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_700, plain, (v1_rat_1(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v4_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_701, plain, (v1_int_1(esk34_1(X1))|v1_xboole_0(X1)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_454, plain, (k5_prob_1(X1,k1_zfmisc_1(X2),X3,X3)=X3|v1_xboole_0(X5)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(k1_zfmisc_1(X4),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X8,k1_zfmisc_1(X4))|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v4_prob_1(k1_zfmisc_1(X4),X2)|~v1_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k1_zfmisc_1(X4),X2)|~m1_subset_1(X5,k1_zfmisc_1(k1_zfmisc_1(X4)))|~m1_subset_1(X7,X5)|~m1_subset_1(X6,X5)|~v4_prob_1(X5,X4)|~v1_prob_1(X5,X4))).#trainneg
cnf(i_0_458, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,k3_xboole_0(X1,X2)),k1_seq_1(esk3_0,k3_xboole_0(X3,X4)))|~r1_tarski(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))|~m1_subset_1(X4,esk2_0)|~m1_subset_1(X3,esk2_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_467, plain, (k5_prob_1(X1,k1_zfmisc_1(X2),X3,X3)=X3|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(k5_numbers,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X4,k5_numbers)|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v4_prob_1(k5_numbers,X2)|~v1_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k5_numbers,X2))).#trainneg
cnf(i_0_473, plain, (k5_prob_1(X1,k1_zfmisc_1(X2),X3,X3)=X3|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(k1_numbers,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X4,k1_numbers)|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v4_prob_1(k1_numbers,X2)|~v1_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k1_numbers,X2))).#trainneg
cnf(i_0_482, plain, (k5_prob_1(X1,k1_zfmisc_1(X2),X3,X3)=X3|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v4_prob_1(k1_zfmisc_1(k1_numbers),X2)|~v1_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k1_zfmisc_1(k1_numbers),X2))).#trainneg
cnf(i_0_512, negated_conjecture, (k5_prob_1(X1,k1_zfmisc_1(X2),X3,X3)=X3|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(esk1_0))|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v4_prob_1(k1_zfmisc_1(esk1_0),X2)|~v1_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k1_zfmisc_1(esk1_0),X2))).#trainneg
cnf(i_0_536, plain, (m1_subset_1(k3_xboole_0(X1,X2),k1_zfmisc_1(X3))|~m1_subset_1(k1_zfmisc_1(X3),k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_1639, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk15_1(X2)))).#trainneg
cnf(i_0_1854, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_972, negated_conjecture, (v1_membered(esk5_0)|~v1_membered(esk1_0))).#trainneg
cnf(i_0_1545, negated_conjecture, (k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_702, plain, (v1_int_1(esk29_1(X1))|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1046, plain, (v4_membered(k5_numbers)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_973, negated_conjecture, (v1_membered(esk4_0)|~v1_membered(esk1_0))).#trainneg
cnf(i_0_998, negated_conjecture, (v1_xxreal_0(esk6_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_1010, negated_conjecture, (v2_membered(esk2_0)|~v2_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_2390, negated_conjecture, (v2_membered(esk2_0)|~v3_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_703, plain, (v1_int_1(esk15_1(X1))|v1_xboole_0(X1)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_704, plain, (v1_int_1(esk8_1(X1))|~v5_membered(X1))).#trainneg
cnf(i_0_538, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(esk22_1(X2),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,esk22_1(X2)))).#trainneg
cnf(i_0_539, plain, (m1_subset_1(k3_xboole_0(X1,X2),esk22_1(X3))|~m1_subset_1(esk22_1(X3),k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X2,esk22_1(X3))|~m1_subset_1(X1,esk22_1(X3)))).#trainneg
cnf(i_0_540, plain, (m1_subset_1(k3_xboole_0(X1,X2),k1_zfmisc_1(X3))|~m1_subset_1(esk22_1(X3),k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X2,esk22_1(X3))|~m1_subset_1(X1,esk22_1(X3)))).#trainneg
cnf(i_0_543, plain, (k5_prob_1(X1,k1_zfmisc_1(X2),X3,X3)=X3|v1_xboole_0(X4)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X5,X4)|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v4_prob_1(X4,X2)|~v1_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(X4,X2))).#trainneg
cnf(i_0_567, negated_conjecture, (k5_prob_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),X2,X2)=X2|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)))|~v4_prob_1(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),X1)|~v1_prob_1(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),X1))).#trainneg
cnf(i_0_600, plain, (~r2_hidden(X1,esk29_1(X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_602, plain, (~r2_hidden(X1,esk8_1(k1_zfmisc_1(X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_606, negated_conjecture, (~r2_hidden(X1,esk3_0)|~v1_xboole_0(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_2020, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_876, plain, (r1_tarski(k3_xboole_0(X2,X1),X1))).#trainneg
cnf(i_0_1006, negated_conjecture, (v2_membered(esk6_0)|~v2_membered(esk1_0))).#trainneg
cnf(i_0_1547, negated_conjecture, (k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_716, plain, (v1_int_1(k1_xboole_0)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1149, plain, (v1_valued_0(k5_numbers)|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1007, negated_conjecture, (v2_membered(esk5_0)|~v2_membered(esk1_0))).#trainneg
cnf(i_0_1008, negated_conjecture, (v2_membered(esk4_0)|~v2_membered(esk1_0))).#trainneg
cnf(i_0_1027, negated_conjecture, (v3_membered(esk2_0)|~v3_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1044, negated_conjecture, (v4_membered(esk2_0)|~v4_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_718, plain, (v1_int_1(k5_numbers)|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_719, plain, (v1_int_1(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_608, plain, (~r2_hidden(X1,k1_xboole_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_610, plain, (v1_xboole_0(X2)|v1_xboole_0(X3)|~r2_hidden(X1,k3_prob_2(X2,X3,X4,X5))|~v1_xboole_0(k2_zfmisc_1(X3,k1_numbers))|~m2_prob_1(X4,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_611, plain, (~r2_hidden(X1,k5_prob_1(X2,k1_zfmisc_1(X3),X4,X5))|~v1_xboole_0(X3)|~m1_subset_1(k1_zfmisc_1(X3),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,k1_zfmisc_1(X3))|~m1_subset_1(X4,k1_zfmisc_1(X3))|~v4_prob_1(k1_zfmisc_1(X3),X2)|~v1_prob_1(k1_zfmisc_1(X3),X2))).#trainneg
cnf(i_0_612, plain, (v1_xboole_0(X3)|~r2_hidden(X1,k5_prob_1(X2,X3,X4,X5))|~v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_622, negated_conjecture, (v1_xcmplx_0(esk6_0)|~v1_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_624, negated_conjecture, (v1_xcmplx_0(esk5_0)|~v1_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_626, negated_conjecture, (v1_xcmplx_0(esk4_0)|~v1_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_627, negated_conjecture, (v1_xcmplx_0(esk3_0)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers))))).#trainneg
cnf(i_0_633, negated_conjecture, (v1_xcmplx_0(k3_xboole_0(X1,X2))|~v1_membered(esk2_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_643, negated_conjecture, (v1_xxreal_0(esk6_0)|~v2_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1023, negated_conjecture, (v3_membered(esk6_0)|~v3_membered(esk1_0))).#trainneg
cnf(i_0_720, plain, (v1_int_1(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v5_membered(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1166, plain, (v2_valued_0(k5_numbers)|~v2_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1024, negated_conjecture, (v3_membered(esk5_0)|~v3_membered(esk1_0))).#trainneg
cnf(i_0_1025, negated_conjecture, (v3_membered(esk4_0)|~v3_membered(esk1_0))).#trainneg
cnf(i_0_1112, plain, (v1_xxreal_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_721, plain, (v1_int_1(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v5_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_726, plain, (v7_ordinal1(esk34_1(X1))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_645, negated_conjecture, (v1_xxreal_0(esk5_0)|~v2_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_647, negated_conjecture, (v1_xxreal_0(esk4_0)|~v2_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_648, negated_conjecture, (v1_xxreal_0(esk3_0)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers))))).#trainneg
cnf(i_0_654, negated_conjecture, (v1_xxreal_0(k3_xboole_0(X1,X2))|~v2_membered(esk2_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_664, negated_conjecture, (v1_xreal_0(esk6_0)|~v3_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_666, negated_conjecture, (v1_xreal_0(esk5_0)|~v3_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_668, negated_conjecture, (v1_xreal_0(esk4_0)|~v3_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_669, negated_conjecture, (v1_xreal_0(esk3_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers))))).#trainneg
cnf(i_0_675, negated_conjecture, (v1_xreal_0(k3_xboole_0(X1,X2))|~v3_membered(esk2_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_685, negated_conjecture, (v1_rat_1(esk6_0)|~v4_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_687, negated_conjecture, (v1_rat_1(esk5_0)|~v4_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_689, negated_conjecture, (v1_rat_1(esk4_0)|~v4_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_780, plain, (v6_membered(k1_xboole_0)|~v6_membered(X1))).#trainneg
cnf(i_0_2593, plain, (v6_membered(k1_xboole_0))).#trainneg
cnf(i_0_1040, negated_conjecture, (v4_membered(esk6_0)|~v4_membered(esk1_0))).#trainneg
cnf(i_0_727, plain, (v7_ordinal1(esk29_1(X1))|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1183, plain, (v3_valued_0(k5_numbers)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1041, negated_conjecture, (v4_membered(esk5_0)|~v4_membered(esk1_0))).#trainneg
cnf(i_0_1042, negated_conjecture, (v4_membered(esk4_0)|~v4_membered(esk1_0))).#trainneg
cnf(i_0_1113, plain, (v1_xreal_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_728, plain, (v7_ordinal1(esk15_1(X1))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_729, plain, (v7_ordinal1(esk8_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_690, negated_conjecture, (v1_rat_1(esk3_0)|~v4_membered(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers))))).#trainneg
cnf(i_0_696, negated_conjecture, (v1_rat_1(k3_xboole_0(X1,X2))|~v4_membered(esk2_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_706, negated_conjecture, (v1_int_1(esk6_0)|~v5_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_708, negated_conjecture, (v1_int_1(esk5_0)|~v5_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_710, negated_conjecture, (v1_int_1(esk4_0)|~v5_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_711, negated_conjecture, (v1_int_1(esk3_0)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers))))).#trainneg
cnf(i_0_717, negated_conjecture, (v1_int_1(k3_xboole_0(X1,X2))|~v5_membered(esk2_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_731, negated_conjecture, (v7_ordinal1(esk6_0)|~v6_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_733, negated_conjecture, (v7_ordinal1(esk5_0)|~v6_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_735, negated_conjecture, (v7_ordinal1(esk4_0)|~v6_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_736, negated_conjecture, (v7_ordinal1(esk3_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers))))).#trainneg
cnf(i_0_1821, plain, (v1_relat_1(esk29_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1393, negated_conjecture, (v3_ordinal1(esk6_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_1542, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_numbers)|~v1_xreal_0(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers))).#trainneg
cnf(i_0_743, plain, (v7_ordinal1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1200, plain, (v4_valued_0(k5_numbers)|~v4_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1110, negated_conjecture, (v1_xxreal_0(esk5_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_1111, negated_conjecture, (v1_xxreal_0(esk4_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_1590, negated_conjecture, (v3_membered(esk2_0)|~v6_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_744, plain, (v7_ordinal1(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_745, plain, (v7_ordinal1(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v6_membered(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_2648, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_2649, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~m1_subset_1(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers))).#trainneg
cnf(i_0_2683, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_2686, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_2697, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_742, negated_conjecture, (v7_ordinal1(k3_xboole_0(X1,X2))|~v6_membered(esk2_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_746, plain, (v7_ordinal1(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v6_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_790, plain, (v1_membered(k1_xboole_0)|~v1_membered(X1))).#trainneg
cnf(i_0_1491, negated_conjecture, (v3_ordinal1(esk5_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_1544, negated_conjecture, (k8_real_1(k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_747, plain, (v3_ordinal1(esk34_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2655, plain, (v5_relat_1(k5_numbers,k4_numbers)|~v1_relat_1(k5_numbers)|~v4_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1242, negated_conjecture, (v1_xreal_0(esk6_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_1246, negated_conjecture, (v1_xreal_0(esk5_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_2116, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2742, plain, (v1_xcmplx_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_748, plain, (v3_ordinal1(esk29_1(X1))|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_749, plain, (v3_ordinal1(esk15_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_750, plain, (v3_ordinal1(esk8_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_752, negated_conjecture, (v3_ordinal1(esk6_0)|~v3_ordinal1(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_754, negated_conjecture, (v3_ordinal1(esk5_0)|~v3_ordinal1(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_756, negated_conjecture, (v3_ordinal1(esk4_0)|~v3_ordinal1(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_757, negated_conjecture, (v3_ordinal1(esk3_0)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers))))).#trainneg
cnf(i_0_758, negated_conjecture, (v3_ordinal1(esk2_0)|~v3_ordinal1(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_763, negated_conjecture, (v3_ordinal1(k3_xboole_0(X1,X2))|~v3_ordinal1(esk2_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_765, plain, (v3_ordinal1(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_766, plain, (v3_ordinal1(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v3_ordinal1(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_767, plain, (v3_ordinal1(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v3_ordinal1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1914, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_1498, negated_conjecture, (v3_ordinal1(esk4_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_768, plain, (v3_membered(k3_xboole_0(X1,X2))|~v5_membered(X2))).#trainneg
cnf(i_0_2734, plain, (v5_relat_1(k5_numbers,k4_numbers)|~v1_relat_1(k5_numbers)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1305, negated_conjecture, (v1_xreal_0(esk4_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_1391, negated_conjecture, (v1_xxreal_0(esk6_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_774, plain, (v3_membered(k3_xboole_0(X1,X2))|~v5_membered(X1))).#trainneg
cnf(i_0_806, negated_conjecture, (v1_xxreal_0(esk2_0)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_809, plain, (v2_membered(k1_xboole_0)|~v2_membered(X1))).#trainneg
cnf(i_0_807, negated_conjecture, (v1_xreal_0(esk2_0)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_829, plain, (v3_membered(k3_xboole_0(X1,X2))|~v4_membered(X2))).#trainneg
cnf(i_0_830, plain, (v4_membered(k1_xboole_0)|~v4_membered(X1))).#trainneg
cnf(i_0_2834, plain, (v4_membered(k1_xboole_0))).#trainneg
cnf(i_0_1924, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_1539, negated_conjecture, (k3_xcmplx_0(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_numbers))).#trainneg
cnf(i_0_2656, plain, (v5_relat_1(k5_numbers,k3_numbers)|~v1_relat_1(k5_numbers)|~v4_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1489, negated_conjecture, (v1_xxreal_0(esk5_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_2117, plain, (v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_835, plain, (v3_membered(k3_xboole_0(X1,X2))|~v4_membered(X1))).#trainneg
cnf(i_0_841, plain, (v3_valued_0(k2_zfmisc_1(X1,X2))|~v1_relat_1(k2_zfmisc_1(X1,X2))|~v6_membered(X2))).#trainneg
cnf(i_0_842, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),k4_numbers)|~v1_relat_1(k2_zfmisc_1(X1,X2))|~v6_membered(X2))).#trainneg
cnf(i_0_843, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),k3_numbers)|~v1_relat_1(k2_zfmisc_1(X1,X2))|~v6_membered(X2))).#trainneg
cnf(i_0_844, negated_conjecture, (v3_xxreal_0(k1_xboole_0)|v2_xxreal_0(k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0)))|~v1_xreal_0(k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0))))).#trainneg
cnf(i_0_846, plain, (v1_xcmplx_0(esk13_1(X1))|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_847, plain, (v1_xxreal_0(esk13_1(X1))|~v2_membered(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_848, plain, (v1_xreal_0(esk13_1(X1))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_849, plain, (v1_rat_1(esk13_1(X1))|~v4_membered(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_850, plain, (v1_int_1(esk13_1(X1))|~v5_membered(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_2789, plain, (v5_relat_1(k2_zfmisc_1(X1,k5_numbers),k4_numbers))).#trainneg
cnf(i_0_1540, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers))).#trainneg
cnf(i_0_851, plain, (v7_ordinal1(esk13_1(X1))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_2843, plain, (v5_relat_1(k5_numbers,k3_numbers)|~v1_relat_1(k5_numbers)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1496, negated_conjecture, (v1_xxreal_0(esk4_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_852, plain, (v3_ordinal1(esk13_1(X1))|~v3_ordinal1(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_855, plain, (v1_xcmplx_0(esk22_1(X1))|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_856, plain, (v1_xxreal_0(esk22_1(X1))|~v2_membered(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_857, plain, (v1_xreal_0(esk22_1(X1))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_858, plain, (v1_rat_1(esk22_1(X1))|~v4_membered(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_859, plain, (v1_int_1(esk22_1(X1))|~v5_membered(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_860, plain, (v7_ordinal1(esk22_1(X1))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_861, plain, (v3_ordinal1(esk22_1(X1))|~v3_ordinal1(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_864, negated_conjecture, (k5_prob_1(X1,k1_zfmisc_1(X2),X3,X3)=X3|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v4_prob_1(esk2_0,X2)|~v1_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(esk2_0,X2))).#trainneg
cnf(i_0_872, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_873, negated_conjecture, (v1_xxreal_0(esk2_0)|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_881, plain, (v1_xboole_0(X2)|~r2_hidden(X2,X1)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_944, plain, (v6_ordinal1(k1_xboole_0)|~v6_ordinal1(X1))).#trainneg
cnf(i_0_1869, negated_conjecture, (v1_valued_0(esk3_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_882, plain, (k3_xcmplx_0(X1,esk30_0)=k3_xcmplx_0(esk30_0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2512, negated_conjecture, (v4_membered(esk2_0)|~v5_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_883, plain, (k3_xcmplx_0(X1,esk25_0)=k3_xcmplx_0(esk25_0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_884, plain, (k3_xcmplx_0(X1,np__1)=k3_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2937, negated_conjecture, (k3_xcmplx_0(esk6_0,esk30_0)=k3_xcmplx_0(esk30_0,esk6_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_2938, negated_conjecture, (k3_xcmplx_0(esk30_0,esk5_0)=k3_xcmplx_0(esk5_0,esk30_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_2939, negated_conjecture, (k3_xcmplx_0(esk30_0,esk4_0)=k3_xcmplx_0(esk4_0,esk30_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_2946, plain, (k3_xcmplx_0(k5_prob_1(X1,X2,X3,X4),esk30_0)=k3_xcmplx_0(esk30_0,k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_2961, negated_conjecture, (k3_xcmplx_0(esk6_0,esk25_0)=k3_xcmplx_0(esk25_0,esk6_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_2962, negated_conjecture, (k3_xcmplx_0(esk25_0,esk5_0)=k3_xcmplx_0(esk5_0,esk25_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_2963, negated_conjecture, (k3_xcmplx_0(esk25_0,esk4_0)=k3_xcmplx_0(esk4_0,esk25_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_2970, plain, (k3_xcmplx_0(k5_prob_1(X1,X2,X3,X4),esk25_0)=k3_xcmplx_0(esk25_0,k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_2989, plain, (k3_xcmplx_0(k5_prob_1(X1,X2,X3,X4),np__1)=k3_xcmplx_0(np__1,k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_2993, negated_conjecture, (k3_xcmplx_0(esk6_0,esk30_0)=k3_xcmplx_0(esk30_0,esk6_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_2997, negated_conjecture, (k3_xcmplx_0(esk30_0,esk5_0)=k3_xcmplx_0(esk5_0,esk30_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_3001, negated_conjecture, (k3_xcmplx_0(esk30_0,esk4_0)=k3_xcmplx_0(esk4_0,esk30_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_2929, plain, (v6_ordinal1(k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3083, plain, (v6_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_1879, negated_conjecture, (v2_valued_0(esk3_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_885, plain, (k3_xcmplx_0(X1,k1_xboole_0)=k3_xcmplx_0(k1_xboole_0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2654, plain, (v3_valued_0(k5_numbers)|~v1_relat_1(k5_numbers)|~v4_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1899, negated_conjecture, (v4_valued_0(esk3_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2952, negated_conjecture, (v4_membered(esk2_0)|~v6_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_886, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_3020, negated_conjecture, (k3_xcmplx_0(esk6_0,esk25_0)=k3_xcmplx_0(esk25_0,esk6_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_3024, negated_conjecture, (k3_xcmplx_0(esk25_0,esk5_0)=k3_xcmplx_0(esk5_0,esk25_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_3028, negated_conjecture, (k3_xcmplx_0(esk25_0,esk4_0)=k3_xcmplx_0(esk4_0,esk25_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_3102, plain, (k3_xcmplx_0(k5_prob_1(X1,X2,X3,X4),k1_xboole_0)=k3_xcmplx_0(k1_xboole_0,k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_3128, plain, (k3_xcmplx_0(k5_prob_1(X1,X2,X3,X4),X5)=k3_xcmplx_0(X5,k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v1_xreal_0(X5)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_3174, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k3_xboole_0(X2,X3)))).#trainneg
cnf(i_0_889, plain, (r1_xxreal_0(X1,esk37_0)|r1_xxreal_0(esk37_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_3256, plain, (r1_xxreal_0(esk37_0,esk37_0))).#trainneg
cnf(i_0_1938, negated_conjecture, (v3_membered(esk6_0)|~v5_membered(esk1_0))).#trainneg
cnf(i_0_2114, negated_conjecture, (v3_membered(esk5_0)|~v5_membered(esk1_0))).#trainneg
cnf(i_0_890, plain, (r1_xxreal_0(X1,esk32_0)|r1_xxreal_0(esk32_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_891, plain, (r1_xxreal_0(X1,esk31_0)|r1_xxreal_0(esk31_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_3240, plain, (k3_xcmplx_0(X3,X4)=k3_xcmplx_0(X4,X3)|v1_xboole_0(X2)|~v1_xreal_0(X4)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_892, plain, (r1_xxreal_0(X1,esk30_0)|r1_xxreal_0(esk30_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_893, plain, (r1_xxreal_0(X1,esk26_0)|r1_xxreal_0(esk26_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_894, plain, (r1_xxreal_0(X1,esk25_0)|r1_xxreal_0(esk25_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_3252, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2)))).#trainneg
cnf(i_0_3279, plain, (v3_xxreal_0(esk37_0)|v2_xxreal_0(esk37_0))).#trainneg
cnf(i_0_895, plain, (r1_xxreal_0(X1,esk19_0)|r1_xxreal_0(esk19_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_3288, plain, (r1_xxreal_0(esk32_0,esk32_0))).#trainneg
cnf(i_0_2115, negated_conjecture, (v3_membered(esk4_0)|~v5_membered(esk1_0))).#trainneg
cnf(i_0_896, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_897, plain, (r1_xxreal_0(X1,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_900, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_910, negated_conjecture, (v1_xboole_0(esk3_0)|~v1_xboole_0(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_914, plain, (v1_xboole_0(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(k2_zfmisc_1(X2,k1_numbers))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_915, plain, (v1_xboole_0(k5_prob_1(X1,k1_zfmisc_1(X2),X3,X4))|~v1_xboole_0(X2)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k1_zfmisc_1(X2),X1))).#trainneg
cnf(i_0_916, plain, (v1_xboole_0(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_917, plain, (v5_membered(esk34_1(X1))|v1_xboole_0(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_918, plain, (v5_membered(esk29_1(X1))|~v5_membered(X1))).#trainneg
cnf(i_0_919, plain, (v5_membered(esk22_1(X1))|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_920, plain, (v5_membered(esk15_1(X1))|v1_xboole_0(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_921, plain, (v5_membered(esk13_1(X1))|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1546, negated_conjecture, (k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_922, plain, (v5_membered(esk8_1(k1_zfmisc_1(X1)))|~v5_membered(X1))).#trainneg
cnf(i_0_3311, plain, (r1_xxreal_0(esk31_0,esk31_0))).#trainneg
cnf(i_0_2605, negated_conjecture, (v3_membered(esk6_0)|~v4_membered(esk1_0))).#trainneg
cnf(i_0_2609, negated_conjecture, (v3_membered(esk5_0)|~v4_membered(esk1_0))).#trainneg
cnf(i_0_930, plain, (v5_membered(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v5_membered(k2_zfmisc_1(X2,k1_numbers))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_932, plain, (v5_membered(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v5_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_926, negated_conjecture, (v5_membered(esk3_0)|~v5_membered(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_931, plain, (v5_membered(k5_prob_1(X1,k1_zfmisc_1(X2),X3,X4))|~v5_membered(X2)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k1_zfmisc_1(X2),X1))).#trainneg
cnf(i_0_933, plain, (v6_ordinal1(esk34_1(X1))|v1_xboole_0(X1)|~v6_ordinal1(X1))).#trainneg
cnf(i_0_934, plain, (v6_ordinal1(esk29_1(X1))|~v6_ordinal1(X1))).#trainneg
cnf(i_0_935, plain, (v6_ordinal1(esk22_1(X1))|~v6_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_936, plain, (v6_ordinal1(esk15_1(X1))|v1_xboole_0(X1)|~v6_ordinal1(X1))).#trainneg
cnf(i_0_937, plain, (v6_ordinal1(esk13_1(X1))|~v6_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_938, plain, (v6_ordinal1(esk8_1(k1_zfmisc_1(X1)))|~v6_ordinal1(X1))).#trainneg
cnf(i_0_942, negated_conjecture, (v6_ordinal1(esk3_0)|~v6_ordinal1(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_946, plain, (v6_ordinal1(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v6_ordinal1(k2_zfmisc_1(X2,k1_numbers))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_947, plain, (v6_ordinal1(k5_prob_1(X1,k1_zfmisc_1(X2),X3,X4))|~v6_ordinal1(X2)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k1_zfmisc_1(X2),X1))).#trainneg
cnf(i_0_948, plain, (v6_ordinal1(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v6_ordinal1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_3366, plain, (r1_xxreal_0(esk30_0,esk30_0))).#trainneg
cnf(i_0_1575, negated_conjecture, (k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_949, plain, (v6_membered(esk34_1(X1))|v1_xboole_0(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_3390, plain, (r1_xxreal_0(esk26_0,esk26_0))).#trainneg
cnf(i_0_2610, negated_conjecture, (v3_membered(esk4_0)|~v4_membered(esk1_0))).#trainneg
cnf(i_0_950, plain, (v6_membered(esk29_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_951, plain, (v6_membered(esk22_1(X1))|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_952, plain, (v6_membered(esk15_1(X1))|v1_xboole_0(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_953, plain, (v6_membered(esk13_1(X1))|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_954, plain, (v6_membered(esk8_1(k1_zfmisc_1(X1)))|~v6_membered(X1))).#trainneg
cnf(i_0_958, negated_conjecture, (v6_membered(esk3_0)|~v6_membered(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_962, plain, (v6_membered(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v6_membered(k2_zfmisc_1(X2,k1_numbers))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_963, plain, (v6_membered(k5_prob_1(X1,k1_zfmisc_1(X2),X3,X4))|~v6_membered(X2)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k1_zfmisc_1(X2),X1))).#trainneg
cnf(i_0_964, plain, (v6_membered(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v6_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_965, plain, (v1_membered(esk34_1(X1))|v1_xboole_0(X1)|~v1_membered(X1))).#trainneg
cnf(i_0_966, plain, (v1_membered(esk29_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_967, plain, (v1_membered(esk22_1(X1))|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_968, plain, (v1_membered(esk15_1(X1))|v1_xboole_0(X1)|~v1_membered(X1))).#trainneg
cnf(i_0_3413, plain, (r1_xxreal_0(esk25_0,esk25_0))).#trainneg
cnf(i_0_969, plain, (v1_membered(esk13_1(X1))|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3445, plain, (r1_xxreal_0(esk19_0,esk19_0))).#trainneg
cnf(i_0_1147, negated_conjecture, (v1_valued_0(esk2_0)|~v1_valued_0(k1_zfmisc_1(esk1_0))|~v1_relat_1(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_3644, negated_conjecture, (v1_valued_0(esk2_0)|~v1_relat_1(k1_zfmisc_1(esk1_0))|~v3_valued_0(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_970, plain, (v1_membered(esk8_1(k1_zfmisc_1(X1)))|~v1_membered(X1))).#trainneg
cnf(i_0_974, negated_conjecture, (v1_membered(esk3_0)|~v1_membered(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_978, plain, (v1_membered(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_membered(k2_zfmisc_1(X2,k1_numbers))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_979, plain, (v1_membered(k5_prob_1(X1,k1_zfmisc_1(X2),X3,X4))|~v1_membered(X2)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k1_zfmisc_1(X2),X1))).#trainneg
cnf(i_0_980, plain, (v1_membered(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_995, plain, (esk29_1(X1)=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_997, plain, (X1=esk29_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1000, plain, (v2_membered(esk34_1(X1))|v1_xboole_0(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_1001, plain, (v2_membered(esk29_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_1002, plain, (v2_membered(esk22_1(X1))|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1003, plain, (v2_membered(esk15_1(X1))|v1_xboole_0(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_1004, plain, (v2_membered(esk13_1(X1))|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3501, plain, (r1_xxreal_0(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_1005, plain, (v2_membered(esk8_1(k1_zfmisc_1(X1)))|~v2_membered(X1))).#trainneg
cnf(i_0_3258, plain, (r1_xxreal_0(esk37_0,esk32_0)|r1_xxreal_0(esk32_0,esk37_0))).#trainneg
cnf(i_0_1164, negated_conjecture, (v2_valued_0(esk2_0)|~v2_valued_0(k1_zfmisc_1(esk1_0))|~v1_relat_1(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_3710, negated_conjecture, (v2_valued_0(esk2_0)|~v1_relat_1(k1_zfmisc_1(esk1_0))|~v3_valued_0(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1013, plain, (v2_membered(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v2_membered(k2_zfmisc_1(X2,k1_numbers))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1009, negated_conjecture, (v2_membered(esk3_0)|~v2_membered(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_1014, plain, (v2_membered(k5_prob_1(X1,k1_zfmisc_1(X2),X3,X4))|~v2_membered(X2)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k1_zfmisc_1(X2),X1))).#trainneg
cnf(i_0_1015, plain, (v2_membered(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1017, plain, (v3_membered(esk34_1(X1))|v1_xboole_0(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_1018, plain, (v3_membered(esk29_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_1019, plain, (v3_membered(esk22_1(X1))|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1020, plain, (v3_membered(esk15_1(X1))|v1_xboole_0(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_1021, plain, (v3_membered(esk13_1(X1))|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1022, plain, (v3_membered(esk8_1(k1_zfmisc_1(X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_1026, negated_conjecture, (v3_membered(esk3_0)|~v3_membered(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_1030, plain, (v3_membered(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v3_membered(k2_zfmisc_1(X2,k1_numbers))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_3159, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1032, plain, (v3_membered(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_3259, plain, (r1_xxreal_0(esk37_0,esk31_0)|r1_xxreal_0(esk31_0,esk37_0))).#trainneg
cnf(i_0_1181, negated_conjecture, (v3_valued_0(esk2_0)|~v3_valued_0(k1_zfmisc_1(esk1_0))|~v1_relat_1(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1034, plain, (v4_membered(esk34_1(X1))|v1_xboole_0(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_1035, plain, (v4_membered(esk29_1(X1))|~v4_membered(X1))).#trainneg
cnf(i_0_1031, plain, (v3_membered(k5_prob_1(X1,k1_zfmisc_1(X2),X3,X4))|~v3_membered(X2)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k1_zfmisc_1(X2),X1))).#trainneg
cnf(i_0_1036, plain, (v4_membered(esk22_1(X1))|~v4_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1037, plain, (v4_membered(esk15_1(X1))|v1_xboole_0(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_1038, plain, (v4_membered(esk13_1(X1))|~v4_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1039, plain, (v4_membered(esk8_1(k1_zfmisc_1(X1)))|~v4_membered(X1))).#trainneg
cnf(i_0_1043, negated_conjecture, (v4_membered(esk3_0)|~v4_membered(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_1047, plain, (v4_membered(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v4_membered(k2_zfmisc_1(X2,k1_numbers))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1048, plain, (v4_membered(k5_prob_1(X1,k1_zfmisc_1(X2),X3,X4))|~v4_membered(X2)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k1_zfmisc_1(X2),X1))).#trainneg
cnf(i_0_1049, plain, (v4_membered(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v4_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1062, plain, (v2_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X5,X6)))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v2_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X7,X8)))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X5,X6)))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X7,X8)))|~r1_tarski(k3_xboole_0(X7,X8),k3_xboole_0(X5,X6))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~m1_subset_1(X8,X2)|~m1_subset_1(X7,X2)|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1063, plain, (v2_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v2_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X6,X7)))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X6,X7)))|~r1_tarski(k3_xboole_0(X6,X7),X5)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X5,X2)|~m1_subset_1(X7,X2)|~m1_subset_1(X6,X2)|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1870, plain, (v1_valued_0(k1_xboole_0)|~v1_membered(X1))).#trainneg
cnf(i_0_3743, plain, (v1_xboole_0(k3_xboole_0(X2,X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1115, plain, (k3_xcmplx_0(X1,esk34_1(X2))=k3_xcmplx_0(esk34_1(X2),X1)|v1_xboole_0(X2)|~v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_3260, plain, (r1_xxreal_0(esk37_0,esk30_0)|r1_xxreal_0(esk30_0,esk37_0))).#trainneg
cnf(i_0_3282, negated_conjecture, (v3_membered(esk6_0)|~v6_membered(esk1_0))).#trainneg
cnf(i_0_1198, negated_conjecture, (v4_valued_0(esk2_0)|~v4_valued_0(k1_zfmisc_1(esk1_0))|~v1_relat_1(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1064, plain, (v2_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X5,X6)))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v2_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X7))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X5,X6)))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X7))|~r1_tarski(X7,k3_xboole_0(X5,X6))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~m1_subset_1(X7,X2)|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1065, plain, (v2_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v2_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X6))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X6))|~r1_tarski(X6,X5)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X5,X2)|~m1_subset_1(X6,X2)|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_3854, plain, (r1_xxreal_0(esk30_0,esk37_0)|~v2_xxreal_0(esk37_0))).#trainneg
cnf(i_0_1066, negated_conjecture, (v2_xxreal_0(k1_seq_1(X1,X2))|~v2_xxreal_0(k1_seq_1(X1,esk6_0))|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,esk6_0))|~r1_tarski(esk6_0,X2)|~m2_prob_1(X1,esk1_0,esk2_0)|~m1_subset_1(X2,esk2_0))).#trainneg
cnf(i_0_1067, negated_conjecture, (v2_xxreal_0(k1_seq_1(X1,X2))|~v2_xxreal_0(k1_seq_1(X1,esk5_0))|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,esk5_0))|~r1_tarski(esk5_0,X2)|~m2_prob_1(X1,esk1_0,esk2_0)|~m1_subset_1(X2,esk2_0))).#trainneg
cnf(i_0_1068, negated_conjecture, (v2_xxreal_0(k1_seq_1(X1,X2))|~v2_xxreal_0(k1_seq_1(X1,esk4_0))|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,esk4_0))|~r1_tarski(esk4_0,X2)|~m2_prob_1(X1,esk1_0,esk2_0)|~m1_subset_1(X2,esk2_0))).#trainneg
cnf(i_0_1069, plain, (v2_xxreal_0(k1_seq_1(X1,k3_xboole_0(X2,X3)))|v1_xboole_0(X4)|v1_xboole_0(X5)|~v2_xxreal_0(k1_seq_1(X1,k5_prob_1(X4,X5,X6,X7)))|~v1_xreal_0(k1_seq_1(X1,k3_xboole_0(X2,X3)))|~v1_xreal_0(k1_seq_1(X1,k5_prob_1(X4,X5,X6,X7)))|~r1_tarski(k5_prob_1(X4,X5,X6,X7),k3_xboole_0(X2,X3))|~m2_prob_1(X1,X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(k1_zfmisc_1(X4)))|~m1_subset_1(X3,X5)|~m1_subset_1(X2,X5)|~m1_subset_1(X7,X5)|~m1_subset_1(X6,X5)|~v4_prob_1(X5,X4)|~v1_prob_1(X5,X4))).#trainneg
cnf(i_0_1070, plain, (v2_xxreal_0(k1_seq_1(X1,X2))|v1_xboole_0(X3)|v1_xboole_0(X4)|~v2_xxreal_0(k1_seq_1(X1,k5_prob_1(X3,X4,X5,X6)))|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,k5_prob_1(X3,X4,X5,X6)))|~r1_tarski(k5_prob_1(X3,X4,X5,X6),X2)|~m2_prob_1(X1,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X2,X4)|~m1_subset_1(X6,X4)|~m1_subset_1(X5,X4)|~v4_prob_1(X4,X3)|~v1_prob_1(X4,X3))).#trainneg
cnf(i_0_1071, negated_conjecture, (v2_xxreal_0(k1_seq_1(X1,esk6_0))|~v2_xxreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,esk6_0))|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_tarski(X2,esk6_0)|~m1_prob_1(X2,esk1_0,esk2_0)|~m2_prob_1(X1,esk1_0,esk2_0))).#trainneg
cnf(i_0_3886, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk6_0))|~r1_tarski(k3_xboole_0(esk4_0,esk5_0),esk6_0)|~v1_xreal_0(k1_seq_1(esk3_0,esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0)))|~m1_prob_1(k3_xboole_0(esk4_0,esk5_0),esk1_0,esk2_0))).#trainneg
cnf(i_0_3887, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk6_0))|~r1_tarski(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0)|~v1_xreal_0(k1_seq_1(esk3_0,esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0))).#trainneg
cnf(i_0_3891, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_xreal_0(k1_seq_1(esk3_0,esk6_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk6_0)))).#trainneg
cnf(i_0_1880, plain, (v2_valued_0(k1_xboole_0)|~v2_membered(X1))).#trainneg
cnf(i_0_3284, negated_conjecture, (v3_membered(esk5_0)|~v6_membered(esk1_0))).#trainneg
cnf(i_0_1116, plain, (k3_xcmplx_0(esk34_1(X1),k1_xboole_0)=k1_xboole_0|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3262, plain, (r1_xxreal_0(esk37_0,esk26_0)|r1_xxreal_0(esk26_0,esk37_0))).#trainneg
cnf(i_0_3462, negated_conjecture, (v3_membered(esk4_0)|~v6_membered(esk1_0))).#trainneg
cnf(i_0_1218, plain, (v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_1225, plain, (v1_xxreal_0(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v2_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_1074, plain, (v2_xxreal_0(k1_seq_1(X1,k5_prob_1(X2,X3,X4,X5)))|v1_xboole_0(X3)|v1_xboole_0(X2)|~v2_xxreal_0(k1_seq_1(X1,X6))|~v1_xreal_0(k1_seq_1(X1,k5_prob_1(X2,X3,X4,X5)))|~v1_xreal_0(k1_seq_1(X1,X6))|~r1_tarski(X6,k5_prob_1(X2,X3,X4,X5))|~m1_prob_1(X6,X2,X3)|~m2_prob_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_1078, plain, (v3_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X5,X6)))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v3_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X7,X8)))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X7,X8)))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X5,X6)))|~r1_tarski(k3_xboole_0(X5,X6),k3_xboole_0(X7,X8))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X8,X2)|~m1_subset_1(X7,X2)|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_3894, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk6_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers))).#trainneg
cnf(i_0_3895, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk6_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k5_numbers))).#trainneg
cnf(i_0_3896, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk6_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk6_0))|~v1_xboole_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0))))).#trainneg
cnf(i_0_3897, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk6_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk6_0)))).#trainneg
cnf(i_0_3978, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk6_0))|~m1_prob_1(k3_xboole_0(esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k3_xboole_0(esk4_0,esk5_0),k1_zfmisc_1(esk6_0)))).#trainneg
cnf(i_0_3980, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk6_0))|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk6_0))|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0))).#trainneg
cnf(i_0_3981, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk6_0))|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk6_0)))).#trainneg
cnf(i_0_3983, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk6_0))|~m1_subset_1(k3_xboole_0(esk4_0,esk5_0),k1_zfmisc_1(esk6_0)))).#trainneg
cnf(i_0_3988, negated_conjecture, (~v1_xboole_0(k1_seq_1(esk3_0,esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk6_0))|~m1_subset_1(k3_xboole_0(esk4_0,esk5_0),k1_zfmisc_1(esk6_0)))).#trainneg
cnf(i_0_1890, plain, (v3_valued_0(k1_xboole_0)|~v3_membered(X1))).#trainneg
cnf(i_0_4003, plain, (v3_valued_0(k1_xboole_0))).#trainneg
cnf(i_0_1117, plain, (k3_xcmplx_0(np__1,esk34_1(X1))=esk34_1(X1)|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3263, plain, (r1_xxreal_0(esk37_0,esk25_0)|r1_xxreal_0(esk25_0,esk37_0))).#trainneg
cnf(i_0_1258, plain, (v7_ordinal1(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v4_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_4031, plain, (v1_xxreal_0(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1079, plain, (v3_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X5,X6)))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v3_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X7))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X7))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X5,X6)))|~r1_tarski(k3_xboole_0(X5,X6),X7)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X7,X2)|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1080, plain, (v3_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v3_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X6,X7)))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X6,X7)))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|~r1_tarski(X5,k3_xboole_0(X6,X7))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X7,X2)|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1072, negated_conjecture, (v2_xxreal_0(k1_seq_1(X1,esk5_0))|~v2_xxreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,esk5_0))|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_tarski(X2,esk5_0)|~m1_prob_1(X2,esk1_0,esk2_0)|~m2_prob_1(X1,esk1_0,esk2_0))).#trainneg
cnf(i_0_4056, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk5_0))|~r1_tarski(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk5_0)|~v1_xreal_0(k1_seq_1(esk3_0,esk5_0))|~v1_xreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0))).#trainneg
cnf(i_0_4059, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk5_0))|~v1_xreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_xreal_0(k1_seq_1(esk3_0,esk5_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk5_0)))).#trainneg
cnf(i_0_4062, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk5_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk5_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers))).#trainneg
cnf(i_0_4063, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk5_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk5_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k5_numbers))).#trainneg
cnf(i_0_4064, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk5_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk5_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk5_0))|~v1_xboole_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0))))).#trainneg
cnf(i_0_4065, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk5_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk5_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk5_0)))).#trainneg
cnf(i_0_4078, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk5_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk5_0))|~m1_prob_1(k3_xboole_0(esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k3_xboole_0(esk4_0,esk5_0),k1_zfmisc_1(esk5_0)))).#trainneg
cnf(i_0_4080, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk5_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk5_0))|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk5_0))|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0))).#trainneg
cnf(i_0_4081, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk5_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk5_0))|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk5_0)))).#trainneg
cnf(i_0_1900, plain, (v4_valued_0(k1_xboole_0)|~v6_membered(X1))).#trainneg
cnf(i_0_4090, plain, (v4_valued_0(k1_xboole_0))).#trainneg
cnf(i_0_1118, plain, (k3_xcmplx_0(X1,esk29_1(X2))=k3_xcmplx_0(esk29_1(X2),X1)|~v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_3265, plain, (r1_xxreal_0(esk37_0,esk19_0)|r1_xxreal_0(esk19_0,esk37_0))).#trainneg
cnf(i_0_4032, plain, (v1_xreal_0(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4033, plain, (v3_ordinal1(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1081, plain, (v3_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v3_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X6))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X6))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|~r1_tarski(X5,X6)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1085, plain, (v3_xxreal_0(k1_seq_1(X1,k5_prob_1(X2,X3,X4,X5)))|v1_xboole_0(X2)|v1_xboole_0(X3)|~v3_xxreal_0(k1_seq_1(X1,k3_xboole_0(X6,X7)))|~v1_xreal_0(k1_seq_1(X1,k3_xboole_0(X6,X7)))|~v1_xreal_0(k1_seq_1(X1,k5_prob_1(X2,X3,X4,X5)))|~r1_tarski(k5_prob_1(X2,X3,X4,X5),k3_xboole_0(X6,X7))|~m2_prob_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X7,X3)|~m1_subset_1(X6,X3)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_4082, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk5_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk5_0))|~m1_subset_1(k3_xboole_0(esk4_0,esk5_0),k1_zfmisc_1(esk5_0))|~m1_subset_1(k3_xboole_0(esk4_0,esk5_0),esk2_0))).#trainneg
cnf(i_0_4083, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk5_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk5_0))|~m1_subset_1(k3_xboole_0(esk4_0,esk5_0),k1_zfmisc_1(esk5_0)))).#trainneg
cnf(i_0_4156, negated_conjecture, (~v1_xboole_0(k1_seq_1(esk3_0,esk5_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk5_0))|~m1_subset_1(k3_xboole_0(esk4_0,esk5_0),k1_zfmisc_1(esk5_0)))).#trainneg
cnf(i_0_4161, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~r1_tarski(esk5_0,esk4_0)|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk5_0))|~m1_subset_1(k3_xboole_0(esk4_0,esk5_0),k1_zfmisc_1(esk5_0)))).#trainneg
cnf(i_0_4164, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk5_0))|~m1_subset_1(k3_xboole_0(esk4_0,esk5_0),k1_zfmisc_1(esk5_0))|~m1_subset_1(esk5_0,k1_zfmisc_1(esk4_0)))).#trainneg
cnf(i_0_4168, negated_conjecture, (~v1_xboole_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk5_0))|~m1_subset_1(k3_xboole_0(esk4_0,esk5_0),k1_zfmisc_1(esk5_0))|~m1_subset_1(esk5_0,k1_zfmisc_1(esk4_0)))).#trainneg
cnf(i_0_2646, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~m1_subset_1(k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_1119, plain, (k3_xcmplx_0(esk29_1(X1),k1_xboole_0)=k1_xboole_0|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3269, plain, (r1_xxreal_0(esk37_0,np__1)|r1_xxreal_0(np__1,esk37_0))).#trainneg
cnf(i_0_2917, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_1086, plain, (v3_xxreal_0(k1_seq_1(X1,k5_prob_1(X2,X3,X4,X5)))|v1_xboole_0(X2)|v1_xboole_0(X3)|~v3_xxreal_0(k1_seq_1(X1,X6))|~v1_xreal_0(k1_seq_1(X1,X6))|~v1_xreal_0(k1_seq_1(X1,k5_prob_1(X2,X3,X4,X5)))|~r1_tarski(k5_prob_1(X2,X3,X4,X5),X6)|~m2_prob_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X6,X3)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_1090, plain, (v3_xxreal_0(k1_seq_1(X1,X2))|v1_xboole_0(X4)|v1_xboole_0(X3)|~v3_xxreal_0(k1_seq_1(X1,k5_prob_1(X3,X4,X5,X6)))|~v1_xreal_0(k1_seq_1(X1,k5_prob_1(X3,X4,X5,X6)))|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_tarski(X2,k5_prob_1(X3,X4,X5,X6))|~m1_prob_1(X2,X3,X4)|~m2_prob_1(X1,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X6,X4)|~m1_subset_1(X5,X4)|~v4_prob_1(X4,X3)|~v1_prob_1(X4,X3))).#trainneg
cnf(i_0_4175, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk4_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers))).#trainneg
cnf(i_0_4176, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk4_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k5_numbers))).#trainneg
cnf(i_0_4177, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk4_0))|~v1_xboole_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0))))).#trainneg
cnf(i_0_4178, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk4_0))|~v4_valued_0(esk3_0))).#trainneg
cnf(i_0_4180, negated_conjecture, (~m1_subset_1(k3_xboole_0(esk4_0,esk5_0),k1_zfmisc_1(esk5_0))|~m1_subset_1(esk5_0,k1_zfmisc_1(esk4_0))|~v1_xboole_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(esk3_0,esk5_0),k1_numbers))).#trainneg
cnf(i_0_4181, negated_conjecture, (~m1_subset_1(k3_xboole_0(esk4_0,esk5_0),k1_zfmisc_1(esk5_0))|~m1_subset_1(esk5_0,k1_zfmisc_1(esk4_0))|~v1_xboole_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(esk3_0,esk5_0),k5_numbers))).#trainneg
cnf(i_0_4183, negated_conjecture, (~m1_subset_1(k3_xboole_0(esk4_0,esk5_0),k1_zfmisc_1(esk5_0))|~m1_subset_1(esk5_0,k1_zfmisc_1(esk4_0))|~v1_xboole_0(k1_seq_1(esk3_0,esk4_0))|~v4_valued_0(esk3_0))).#trainneg
cnf(i_0_4184, negated_conjecture, (~m1_subset_1(k3_xboole_0(esk4_0,esk5_0),k1_zfmisc_1(esk5_0))|~m1_subset_1(esk5_0,k1_zfmisc_1(esk4_0))|~v1_xboole_0(k1_seq_1(esk3_0,esk4_0)))).#trainneg
cnf(i_0_1564, negated_conjecture, (k3_xcmplx_0(k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_numbers))).#trainneg
cnf(i_0_1120, plain, (k3_xcmplx_0(np__1,esk29_1(X1))=esk29_1(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3271, plain, (r1_xxreal_0(esk37_0,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,esk37_0))).#trainneg
cnf(i_0_1094, plain, (v3_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X5,X6)))|v2_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X7,X8)))|v1_xboole_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X7,X8)))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X7,X8)))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X5,X6)))|~r1_tarski(k3_xboole_0(X5,X6),k3_xboole_0(X7,X8))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X8,X2)|~m1_subset_1(X7,X2)|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1095, plain, (v3_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X5,X6)))|v2_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X7))|v1_xboole_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X7))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X7))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X5,X6)))|~r1_tarski(k3_xboole_0(X5,X6),X7)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X7,X2)|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_4243, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk4_0))|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0))).#trainneg
cnf(i_0_4244, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),k1_zfmisc_1(esk4_0)))).#trainneg
cnf(i_0_4245, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k3_xboole_0(esk4_0,esk5_0),esk2_0))).#trainneg
cnf(i_0_4294, negated_conjecture, (~v1_xboole_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0)))).#trainneg
cnf(i_0_4295, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk6_0))|~r1_tarski(esk4_0,esk6_0)|~v1_xreal_0(k1_seq_1(esk3_0,esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0)))).#trainneg
cnf(i_0_4296, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk5_0))|~r1_tarski(esk4_0,esk5_0)|~v1_xreal_0(k1_seq_1(esk3_0,esk5_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0)))).#trainneg
cnf(i_0_4299, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(esk4_0,k1_zfmisc_1(esk6_0)))).#trainneg
cnf(i_0_4300, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk5_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk5_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(esk4_0,k1_zfmisc_1(esk5_0)))).#trainneg
cnf(i_0_4301, negated_conjecture, (~v1_xboole_0(k1_seq_1(esk3_0,esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(esk4_0,k1_zfmisc_1(esk6_0)))).#trainneg
cnf(i_0_2728, plain, (v1_membered(k1_xboole_0)|~v3_membered(X1))).#trainneg
cnf(i_0_4328, plain, (v1_membered(k1_xboole_0))).#trainneg
cnf(i_0_4346, plain, (v1_valued_0(k1_xboole_0))).#trainneg
cnf(i_0_1137, plain, (v1_valued_0(esk34_1(X1))|v1_xboole_0(X1)|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4315, negated_conjecture, (~m1_subset_1(esk4_0,k1_zfmisc_1(esk6_0))|~v1_xboole_0(k1_seq_1(esk3_0,esk6_0)))).#trainneg
cnf(i_0_3558, plain, (r1_xxreal_0(k1_seq_1(X1,X2),X3)|r1_xxreal_0(X3,k1_seq_1(X1,X2))|~v1_xboole_0(X3)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1096, plain, (v3_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|v2_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X6,X7)))|v1_xboole_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X6,X7)))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X6,X7)))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|~r1_tarski(X5,k3_xboole_0(X6,X7))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X7,X2)|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1097, plain, (v3_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|v2_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X6))|v1_xboole_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X6))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X6))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|~r1_tarski(X5,X6)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_4306, negated_conjecture, (~v1_xboole_0(k1_seq_1(esk3_0,esk5_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk5_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(esk4_0,k1_zfmisc_1(esk5_0)))).#trainneg
cnf(i_0_4366, negated_conjecture, (~m1_subset_1(esk4_0,k1_zfmisc_1(esk5_0))|~v1_xboole_0(k1_seq_1(esk3_0,esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_4367, negated_conjecture, (~m1_subset_1(esk4_0,k1_zfmisc_1(esk5_0))|~v1_xboole_0(k1_seq_1(esk3_0,esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k5_numbers))).#trainneg
cnf(i_0_4369, negated_conjecture, (~m1_subset_1(esk4_0,k1_zfmisc_1(esk5_0))|~v1_xboole_0(k1_seq_1(esk3_0,esk5_0))|~v4_valued_0(esk3_0))).#trainneg
cnf(i_0_4370, negated_conjecture, (~m1_subset_1(esk4_0,k1_zfmisc_1(esk5_0))|~v1_xboole_0(k1_seq_1(esk3_0,esk5_0)))).#trainneg
cnf(i_0_1082, negated_conjecture, (v3_xxreal_0(k1_seq_1(X1,esk6_0))|~v3_xxreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,esk6_0))|~r1_tarski(esk6_0,X2)|~m2_prob_1(X1,esk1_0,esk2_0)|~m1_subset_1(X2,esk2_0))).#trainneg
cnf(i_0_1083, negated_conjecture, (v3_xxreal_0(k1_seq_1(X1,esk5_0))|~v3_xxreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,esk5_0))|~r1_tarski(esk5_0,X2)|~m2_prob_1(X1,esk1_0,esk2_0)|~m1_subset_1(X2,esk2_0))).#trainneg
cnf(i_0_1084, negated_conjecture, (v3_xxreal_0(k1_seq_1(X1,esk4_0))|~v3_xxreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,esk4_0))|~r1_tarski(esk4_0,X2)|~m2_prob_1(X1,esk1_0,esk2_0)|~m1_subset_1(X2,esk2_0))).#trainneg
cnf(i_0_1087, negated_conjecture, (v3_xxreal_0(k1_seq_1(X1,X2))|~v3_xxreal_0(k1_seq_1(X1,esk6_0))|~v1_xreal_0(k1_seq_1(X1,esk6_0))|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_tarski(X2,esk6_0)|~m1_prob_1(X2,esk1_0,esk2_0)|~m2_prob_1(X1,esk1_0,esk2_0))).#trainneg
cnf(i_0_2790, plain, (v5_relat_1(k2_zfmisc_1(X1,k5_numbers),k3_numbers))).#trainneg
cnf(i_0_1909, negated_conjecture, (v5_relat_1(esk3_0,k1_numbers))).#trainneg
cnf(i_0_2237, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_1138, plain, (v1_valued_0(esk29_1(X1))|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3289, plain, (r1_xxreal_0(esk32_0,esk31_0)|r1_xxreal_0(esk31_0,esk32_0))).#trainneg
cnf(i_0_1254, plain, (v1_xcmplx_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1511, plain, (v1_xxreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4035, plain, (r1_xxreal_0(X1,k1_seq_1(X2,X3))|r1_xxreal_0(k1_seq_1(X2,X3),X1)|~v1_xboole_0(X1)|~v1_funct_1(X2)|~v4_valued_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1101, plain, (v3_xxreal_0(k1_seq_1(X1,k5_prob_1(X2,X3,X4,X5)))|v2_xxreal_0(k1_seq_1(X1,k3_xboole_0(X6,X7)))|v1_xboole_0(k1_seq_1(X1,k3_xboole_0(X6,X7)))|v1_xboole_0(X2)|v1_xboole_0(X3)|~v1_xreal_0(k1_seq_1(X1,k3_xboole_0(X6,X7)))|~v1_xreal_0(k1_seq_1(X1,k5_prob_1(X2,X3,X4,X5)))|~r1_tarski(k5_prob_1(X2,X3,X4,X5),k3_xboole_0(X6,X7))|~m2_prob_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X7,X3)|~m1_subset_1(X6,X3)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_1102, plain, (v3_xxreal_0(k1_seq_1(X1,k5_prob_1(X2,X3,X4,X5)))|v2_xxreal_0(k1_seq_1(X1,X6))|v1_xboole_0(k1_seq_1(X1,X6))|v1_xboole_0(X2)|v1_xboole_0(X3)|~v1_xreal_0(k1_seq_1(X1,X6))|~v1_xreal_0(k1_seq_1(X1,k5_prob_1(X2,X3,X4,X5)))|~r1_tarski(k5_prob_1(X2,X3,X4,X5),X6)|~m2_prob_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X6,X3)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_1088, negated_conjecture, (v3_xxreal_0(k1_seq_1(X1,X2))|~v3_xxreal_0(k1_seq_1(X1,esk5_0))|~v1_xreal_0(k1_seq_1(X1,esk5_0))|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_tarski(X2,esk5_0)|~m1_prob_1(X2,esk1_0,esk2_0)|~m2_prob_1(X1,esk1_0,esk2_0))).#trainneg
cnf(i_0_1089, negated_conjecture, (v3_xxreal_0(k1_seq_1(X1,X2))|~v3_xxreal_0(k1_seq_1(X1,esk4_0))|~v1_xreal_0(k1_seq_1(X1,esk4_0))|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_tarski(X2,esk4_0)|~m1_prob_1(X2,esk1_0,esk2_0)|~m2_prob_1(X1,esk1_0,esk2_0))).#trainneg
cnf(i_0_1098, negated_conjecture, (v3_xxreal_0(k1_seq_1(X1,esk6_0))|v2_xxreal_0(k1_seq_1(X1,X2))|v1_xboole_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,esk6_0))|~r1_tarski(esk6_0,X2)|~m2_prob_1(X1,esk1_0,esk2_0)|~m1_subset_1(X2,esk2_0))).#trainneg
cnf(i_0_1099, negated_conjecture, (v3_xxreal_0(k1_seq_1(X1,esk5_0))|v2_xxreal_0(k1_seq_1(X1,X2))|v1_xboole_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,esk5_0))|~r1_tarski(esk5_0,X2)|~m2_prob_1(X1,esk1_0,esk2_0)|~m1_subset_1(X2,esk2_0))).#trainneg
cnf(i_0_1100, negated_conjecture, (v3_xxreal_0(k1_seq_1(X1,esk4_0))|v2_xxreal_0(k1_seq_1(X1,X2))|v1_xboole_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,esk4_0))|~r1_tarski(esk4_0,X2)|~m2_prob_1(X1,esk1_0,esk2_0)|~m1_subset_1(X2,esk2_0))).#trainneg
cnf(i_0_1103, negated_conjecture, (v3_xxreal_0(k1_seq_1(X1,X2))|v2_xxreal_0(k1_seq_1(X1,esk6_0))|v1_xboole_0(k1_seq_1(X1,esk6_0))|~v1_xreal_0(k1_seq_1(X1,esk6_0))|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_tarski(X2,esk6_0)|~m1_prob_1(X2,esk1_0,esk2_0)|~m2_prob_1(X1,esk1_0,esk2_0))).#trainneg
cnf(i_0_1104, negated_conjecture, (v3_xxreal_0(k1_seq_1(X1,X2))|v2_xxreal_0(k1_seq_1(X1,esk5_0))|v1_xboole_0(k1_seq_1(X1,esk5_0))|~v1_xreal_0(k1_seq_1(X1,esk5_0))|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_tarski(X2,esk5_0)|~m1_prob_1(X2,esk1_0,esk2_0)|~m2_prob_1(X1,esk1_0,esk2_0))).#trainneg
cnf(i_0_1105, negated_conjecture, (v3_xxreal_0(k1_seq_1(X1,X2))|v2_xxreal_0(k1_seq_1(X1,esk4_0))|v1_xboole_0(k1_seq_1(X1,esk4_0))|~v1_xreal_0(k1_seq_1(X1,esk4_0))|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_tarski(X2,esk4_0)|~m1_prob_1(X2,esk1_0,esk2_0)|~m2_prob_1(X1,esk1_0,esk2_0))).#trainneg
cnf(i_0_1106, plain, (v3_xxreal_0(k1_seq_1(X1,X2))|v2_xxreal_0(k1_seq_1(X1,k5_prob_1(X3,X4,X5,X6)))|v1_xboole_0(k1_seq_1(X1,k5_prob_1(X3,X4,X5,X6)))|v1_xboole_0(X4)|v1_xboole_0(X3)|~v1_xreal_0(k1_seq_1(X1,k5_prob_1(X3,X4,X5,X6)))|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_tarski(X2,k5_prob_1(X3,X4,X5,X6))|~m1_prob_1(X2,X3,X4)|~m2_prob_1(X1,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X6,X4)|~m1_subset_1(X5,X4)|~v4_prob_1(X4,X3)|~v1_prob_1(X4,X3))).#trainneg
cnf(i_0_1122, plain, (v3_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X5,X6)))|v2_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X7,X8)))|v1_xboole_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X5,X6)))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X7,X8)))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X5,X6)))|~r1_tarski(k3_xboole_0(X5,X6),k3_xboole_0(X7,X8))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X8,X2)|~m1_subset_1(X7,X2)|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1123, plain, (v3_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X5,X6)))|v2_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X7))|v1_xboole_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X5,X6)))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X7))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X5,X6)))|~r1_tarski(k3_xboole_0(X5,X6),X7)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X7,X2)|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1124, plain, (v3_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|v2_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X6,X7)))|v1_xboole_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),k3_xboole_0(X6,X7)))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|~r1_tarski(X5,k3_xboole_0(X6,X7))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X7,X2)|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_2496, plain, (m1_subset_1(k3_xboole_0(X1,X2),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_2019, plain, (m1_subset_1(k1_xboole_0,k1_numbers))).#trainneg
cnf(i_0_4156, negated_conjecture, (~v1_xboole_0(k1_seq_1(esk3_0,esk5_0)))).#trainneg
cnf(i_0_1139, plain, (v1_valued_0(esk22_1(X1))|~v1_valued_0(k1_zfmisc_1(X1))|~v1_relat_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3290, plain, (r1_xxreal_0(esk32_0,esk30_0)|r1_xxreal_0(esk30_0,esk32_0))).#trainneg
cnf(i_0_4576, plain, (v1_rat_1(k1_xboole_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_3557, plain, (r1_xxreal_0(k1_seq_1(X1,X2),X3)|r1_xxreal_0(X3,k1_seq_1(X1,X2))|~v1_xxreal_0(X3)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1125, plain, (v3_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|v2_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X6))|v1_xboole_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X6))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|~r1_tarski(X5,X6)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1129, plain, (v3_xxreal_0(k1_seq_1(X1,k5_prob_1(X2,X3,X4,X5)))|v2_xxreal_0(k1_seq_1(X1,k3_xboole_0(X6,X7)))|v1_xboole_0(k1_seq_1(X1,k5_prob_1(X2,X3,X4,X5)))|v1_xboole_0(X2)|v1_xboole_0(X3)|~v1_xreal_0(k1_seq_1(X1,k3_xboole_0(X6,X7)))|~v1_xreal_0(k1_seq_1(X1,k5_prob_1(X2,X3,X4,X5)))|~r1_tarski(k5_prob_1(X2,X3,X4,X5),k3_xboole_0(X6,X7))|~m2_prob_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X7,X3)|~m1_subset_1(X6,X3)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_4083, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk5_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk5_0)))).#trainneg
cnf(i_0_4590, plain, (r1_xxreal_0(esk30_0,esk32_0)|~v2_xxreal_0(esk32_0))).#trainneg
cnf(i_0_4616, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk6_0))|~r1_tarski(esk5_0,esk6_0)|~v1_xreal_0(k1_seq_1(esk3_0,esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk5_0)))).#trainneg
cnf(i_0_4624, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk3_0,esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk5_0))|~m1_subset_1(esk5_0,k1_zfmisc_1(esk6_0)))).#trainneg
cnf(i_0_4625, negated_conjecture, (~v1_xboole_0(k1_seq_1(esk3_0,esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk5_0))|~m1_subset_1(esk5_0,k1_zfmisc_1(esk6_0)))).#trainneg
cnf(i_0_4630, negated_conjecture, (~m1_subset_1(esk5_0,k1_zfmisc_1(esk6_0))|~v1_xboole_0(k1_seq_1(esk3_0,esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,esk5_0),k1_numbers))).#trainneg
cnf(i_0_4631, negated_conjecture, (~m1_subset_1(esk5_0,k1_zfmisc_1(esk6_0))|~v1_xboole_0(k1_seq_1(esk3_0,esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,esk5_0),k5_numbers))).#trainneg
cnf(i_0_4577, plain, (v1_int_1(k1_xboole_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1140, plain, (v1_valued_0(esk15_1(X1))|v1_xboole_0(X1)|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3292, plain, (r1_xxreal_0(esk32_0,esk26_0)|r1_xxreal_0(esk26_0,esk32_0))).#trainneg
cnf(i_0_4634, negated_conjecture, (~m1_subset_1(esk5_0,k1_zfmisc_1(esk6_0))|~v1_xboole_0(k1_seq_1(esk3_0,esk6_0)))).#trainneg
cnf(i_0_1840, plain, (v1_xreal_0(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4034, plain, (r1_xxreal_0(X1,k1_seq_1(X2,X3))|r1_xxreal_0(k1_seq_1(X2,X3),X1)|~v1_xxreal_0(X1)|~v1_funct_1(X2)|~v4_valued_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1130, plain, (v3_xxreal_0(k1_seq_1(X1,k5_prob_1(X2,X3,X4,X5)))|v2_xxreal_0(k1_seq_1(X1,X6))|v1_xboole_0(k1_seq_1(X1,k5_prob_1(X2,X3,X4,X5)))|v1_xboole_0(X2)|v1_xboole_0(X3)|~v1_xreal_0(k1_seq_1(X1,X6))|~v1_xreal_0(k1_seq_1(X1,k5_prob_1(X2,X3,X4,X5)))|~r1_tarski(k5_prob_1(X2,X3,X4,X5),X6)|~m2_prob_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X6,X3)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_1134, plain, (v3_xxreal_0(k1_seq_1(X1,X2))|v2_xxreal_0(k1_seq_1(X1,k5_prob_1(X3,X4,X5,X6)))|v1_xboole_0(k1_seq_1(X1,X2))|v1_xboole_0(X4)|v1_xboole_0(X3)|~v1_xreal_0(k1_seq_1(X1,k5_prob_1(X3,X4,X5,X6)))|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_tarski(X2,k5_prob_1(X3,X4,X5,X6))|~m1_prob_1(X2,X3,X4)|~m2_prob_1(X1,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X6,X4)|~m1_subset_1(X5,X4)|~v4_prob_1(X4,X3)|~v1_prob_1(X4,X3))).#trainneg
cnf(i_0_1126, negated_conjecture, (v3_xxreal_0(k1_seq_1(X1,esk6_0))|v2_xxreal_0(k1_seq_1(X1,X2))|v1_xboole_0(k1_seq_1(X1,esk6_0))|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,esk6_0))|~r1_tarski(esk6_0,X2)|~m2_prob_1(X1,esk1_0,esk2_0)|~m1_subset_1(X2,esk2_0))).#trainneg
cnf(i_0_1127, negated_conjecture, (v3_xxreal_0(k1_seq_1(X1,esk5_0))|v2_xxreal_0(k1_seq_1(X1,X2))|v1_xboole_0(k1_seq_1(X1,esk5_0))|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,esk5_0))|~r1_tarski(esk5_0,X2)|~m2_prob_1(X1,esk1_0,esk2_0)|~m1_subset_1(X2,esk2_0))).#trainneg
cnf(i_0_1128, negated_conjecture, (v3_xxreal_0(k1_seq_1(X1,esk4_0))|v2_xxreal_0(k1_seq_1(X1,X2))|v1_xboole_0(k1_seq_1(X1,esk4_0))|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,esk4_0))|~r1_tarski(esk4_0,X2)|~m2_prob_1(X1,esk1_0,esk2_0)|~m1_subset_1(X2,esk2_0))).#trainneg
cnf(i_0_1131, negated_conjecture, (v3_xxreal_0(k1_seq_1(X1,X2))|v2_xxreal_0(k1_seq_1(X1,esk6_0))|v1_xboole_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,esk6_0))|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_tarski(X2,esk6_0)|~m1_prob_1(X2,esk1_0,esk2_0)|~m2_prob_1(X1,esk1_0,esk2_0))).#trainneg
cnf(i_0_1132, negated_conjecture, (v3_xxreal_0(k1_seq_1(X1,X2))|v2_xxreal_0(k1_seq_1(X1,esk5_0))|v1_xboole_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,esk5_0))|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_tarski(X2,esk5_0)|~m1_prob_1(X2,esk1_0,esk2_0)|~m2_prob_1(X1,esk1_0,esk2_0))).#trainneg
cnf(i_0_1133, negated_conjecture, (v3_xxreal_0(k1_seq_1(X1,X2))|v2_xxreal_0(k1_seq_1(X1,esk4_0))|v1_xboole_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,esk4_0))|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_tarski(X2,esk4_0)|~m1_prob_1(X2,esk1_0,esk2_0)|~m2_prob_1(X1,esk1_0,esk2_0))).#trainneg
cnf(i_0_1141, plain, (v1_valued_0(esk13_1(X1))|~v1_valued_0(k1_zfmisc_1(X1))|~v1_relat_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1142, plain, (v1_valued_0(esk8_1(k1_zfmisc_1(X1)))|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1143, negated_conjecture, (v1_valued_0(esk6_0)|~v1_valued_0(esk1_0)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_1144, negated_conjecture, (v1_valued_0(esk5_0)|~v1_valued_0(esk1_0)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_1145, negated_conjecture, (v1_valued_0(esk4_0)|~v1_valued_0(esk1_0)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_1146, negated_conjecture, (v1_valued_0(esk3_0)|~v1_valued_0(k2_zfmisc_1(esk2_0,k1_numbers))|~v1_relat_1(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_2817, plain, (v2_membered(k1_xboole_0)|~v3_membered(X1))).#trainneg
cnf(i_0_4730, plain, (v2_membered(k1_xboole_0))).#trainneg
cnf(i_0_4748, plain, (v2_valued_0(k1_xboole_0))).#trainneg
cnf(i_0_3293, plain, (r1_xxreal_0(esk32_0,esk25_0)|r1_xxreal_0(esk25_0,esk32_0))).#trainneg
cnf(i_0_1940, negated_conjecture, (v1_xcmplx_0(esk6_0)|~v3_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_2249, plain, (v3_ordinal1(X1)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_1150, plain, (v1_valued_0(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_valued_0(k2_zfmisc_1(X2,k1_numbers))|~v1_relat_1(k2_zfmisc_1(X2,k1_numbers))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1152, plain, (v1_valued_0(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v1_valued_0(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1151, plain, (v1_valued_0(k5_prob_1(X1,k1_zfmisc_1(X2),X3,X4))|~v1_valued_0(X2)|~v1_relat_1(X2)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k1_zfmisc_1(X2),X1))).#trainneg
cnf(i_0_1154, plain, (v2_valued_0(esk34_1(X1))|v1_xboole_0(X1)|~v2_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1155, plain, (v2_valued_0(esk29_1(X1))|~v2_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1156, plain, (v2_valued_0(esk22_1(X1))|~v2_valued_0(k1_zfmisc_1(X1))|~v1_relat_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1157, plain, (v2_valued_0(esk15_1(X1))|v1_xboole_0(X1)|~v2_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1158, plain, (v2_valued_0(esk13_1(X1))|~v2_valued_0(k1_zfmisc_1(X1))|~v1_relat_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1159, plain, (v2_valued_0(esk8_1(k1_zfmisc_1(X1)))|~v2_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1160, negated_conjecture, (v2_valued_0(esk6_0)|~v2_valued_0(esk1_0)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_1161, negated_conjecture, (v2_valued_0(esk5_0)|~v2_valued_0(esk1_0)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_1162, negated_conjecture, (v2_valued_0(esk4_0)|~v2_valued_0(esk1_0)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_1163, negated_conjecture, (v2_valued_0(esk3_0)|~v2_valued_0(k2_zfmisc_1(esk2_0,k1_numbers))|~v1_relat_1(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_4250, negated_conjecture, (k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))|~v1_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~v1_xcmplx_0(k1_seq_1(esk3_0,esk4_0)))).#trainneg
cnf(i_0_1167, plain, (v2_valued_0(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v2_valued_0(k2_zfmisc_1(X2,k1_numbers))|~v1_relat_1(k2_zfmisc_1(X2,k1_numbers))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_3295, plain, (r1_xxreal_0(esk32_0,esk19_0)|r1_xxreal_0(esk19_0,esk32_0))).#trainneg
cnf(i_0_1941, negated_conjecture, (v1_xcmplx_0(esk5_0)|~v3_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1942, negated_conjecture, (v1_xcmplx_0(esk4_0)|~v3_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_2766, negated_conjecture, (v3_ordinal1(esk2_0)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)))))).#trainneg
cnf(i_0_1169, plain, (v2_valued_0(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v2_valued_0(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1171, plain, (v3_valued_0(esk34_1(X1))|v1_xboole_0(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1168, plain, (v2_valued_0(k5_prob_1(X1,k1_zfmisc_1(X2),X3,X4))|~v2_valued_0(X2)|~v1_relat_1(X2)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k1_zfmisc_1(X2),X1))).#trainneg
cnf(i_0_1172, plain, (v3_valued_0(esk29_1(X1))|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1173, plain, (v3_valued_0(esk22_1(X1))|~v3_valued_0(k1_zfmisc_1(X1))|~v1_relat_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1174, plain, (v3_valued_0(esk15_1(X1))|v1_xboole_0(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1175, plain, (v3_valued_0(esk13_1(X1))|~v3_valued_0(k1_zfmisc_1(X1))|~v1_relat_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1176, plain, (v3_valued_0(esk8_1(k1_zfmisc_1(X1)))|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1177, negated_conjecture, (v3_valued_0(esk6_0)|~v3_valued_0(esk1_0)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_1178, negated_conjecture, (v3_valued_0(esk5_0)|~v3_valued_0(esk1_0)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_1179, negated_conjecture, (v3_valued_0(esk4_0)|~v3_valued_0(esk1_0)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_1184, plain, (v3_valued_0(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v3_valued_0(k2_zfmisc_1(X2,k1_numbers))|~v1_relat_1(k2_zfmisc_1(X2,k1_numbers))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_3459, plain, (v3_xxreal_0(esk32_0)|v2_xxreal_0(esk32_0)|v1_xboole_0(esk32_0))).#trainneg
cnf(i_0_2688, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~m1_subset_1(k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_1186, plain, (v3_valued_0(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v3_valued_0(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_3299, plain, (r1_xxreal_0(esk32_0,np__1)|r1_xxreal_0(np__1,esk32_0))).#trainneg
cnf(i_0_2391, negated_conjecture, (v1_xxreal_0(esk6_0)|~v3_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_2392, negated_conjecture, (v1_xxreal_0(esk5_0)|~v3_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_3540, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1188, plain, (v4_valued_0(esk34_1(X1))|v1_xboole_0(X1)|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1189, plain, (v4_valued_0(esk29_1(X1))|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1185, plain, (v3_valued_0(k5_prob_1(X1,k1_zfmisc_1(X2),X3,X4))|~v3_valued_0(X2)|~v1_relat_1(X2)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k1_zfmisc_1(X2),X1))).#trainneg
cnf(i_0_1190, plain, (v4_valued_0(esk22_1(X1))|~v4_valued_0(k1_zfmisc_1(X1))|~v1_relat_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1191, plain, (v4_valued_0(esk15_1(X1))|v1_xboole_0(X1)|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1192, plain, (v4_valued_0(esk13_1(X1))|~v4_valued_0(k1_zfmisc_1(X1))|~v1_relat_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1193, plain, (v4_valued_0(esk8_1(k1_zfmisc_1(X1)))|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1194, negated_conjecture, (v4_valued_0(esk6_0)|~v4_valued_0(esk1_0)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_1195, negated_conjecture, (v4_valued_0(esk5_0)|~v4_valued_0(esk1_0)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_1196, negated_conjecture, (v4_valued_0(esk4_0)|~v4_valued_0(esk1_0)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_1197, negated_conjecture, (v4_valued_0(esk3_0)|~v4_valued_0(k2_zfmisc_1(esk2_0,k1_numbers))|~v1_relat_1(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_1201, plain, (v4_valued_0(k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v4_valued_0(k2_zfmisc_1(X2,k1_numbers))|~v1_relat_1(k2_zfmisc_1(X2,k1_numbers))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1202, plain, (v4_valued_0(k5_prob_1(X1,k1_zfmisc_1(X2),X3,X4))|~v4_valued_0(X2)|~v1_relat_1(X2)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k1_zfmisc_1(X2),X1))).#trainneg
cnf(i_0_1858, plain, (v1_xboole_0(esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1203, plain, (v4_valued_0(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v4_valued_0(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_3301, plain, (r1_xxreal_0(esk32_0,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,esk32_0))).#trainneg
cnf(i_0_2393, negated_conjecture, (v1_xxreal_0(esk4_0)|~v3_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1207, plain, (v1_valued_0(k3_xboole_0(X2,X1))|~v1_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1212, plain, (v2_valued_0(k3_xboole_0(X2,X1))|~v2_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1215, plain, (k3_xcmplx_0(X1,k1_funct_1(X2,X3))=k3_xcmplx_0(k1_funct_1(X2,X3),X1)|~v1_xcmplx_0(X1)|~v1_funct_1(X2)|~v1_valued_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1216, plain, (k3_xcmplx_0(k1_funct_1(X1,X2),k1_xboole_0)=k1_xboole_0|~v1_funct_1(X1)|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1217, plain, (k3_xcmplx_0(np__1,k1_funct_1(X1,X2))=k1_funct_1(X1,X2)|~v1_funct_1(X1)|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1221, plain, (v3_valued_0(k3_xboole_0(X2,X1))|~v3_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1224, plain, (r1_xxreal_0(X1,k1_funct_1(X2,X3))|r1_xxreal_0(k1_funct_1(X2,X3),X1)|~v1_xxreal_0(X1)|~v1_funct_1(X2)|~v2_valued_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1227, plain, (v3_valued_0(k3_xboole_0(X1,X2))|~v1_relat_1(k3_xboole_0(X1,X2))|~v4_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4845, plain, (v1_xboole_0(X1)|v1_xboole_0(esk8_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(X1,X2)|~v1_prob_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3147, plain, (v1_xboole_0(k3_xboole_0(k2_zfmisc_1(X1,X2),X3))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1228, plain, (v5_relat_1(k3_xboole_0(X1,X2),k4_numbers)|~v1_relat_1(k3_xboole_0(X1,X2))|~v4_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1229, plain, (v5_relat_1(k3_xboole_0(X1,X2),k3_numbers)|~v1_relat_1(k3_xboole_0(X1,X2))|~v4_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1232, plain, (v4_valued_0(k3_xboole_0(X2,X1))|~v4_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1238, plain, (k3_xcmplx_0(X1,esk15_1(X2))=k3_xcmplx_0(esk15_1(X2),X1)|v1_xboole_0(X2)|~v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1239, plain, (k3_xcmplx_0(esk15_1(X1),k1_xboole_0)=k1_xboole_0|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1240, plain, (k3_xcmplx_0(np__1,esk15_1(X1))=esk15_1(X1)|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1241, negated_conjecture, (v1_xreal_0(esk6_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1243, negated_conjecture, (v1_xreal_0(esk6_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1244, negated_conjecture, (v1_xreal_0(esk6_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1245, negated_conjecture, (v1_xreal_0(esk5_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1247, negated_conjecture, (v1_xreal_0(esk5_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1248, negated_conjecture, (v1_xreal_0(esk5_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1251, plain, (k3_xcmplx_0(X1,esk8_1(X2))=k3_xcmplx_0(esk8_1(X2),X1)|~v1_xcmplx_0(X1)|~v1_membered(X2))).#trainneg
cnf(i_0_2152, plain, (v1_rat_1(k1_xboole_0)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2636, negated_conjecture, (v3_ordinal1(esk6_0)|~v6_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_2731, negated_conjecture, (k3_xcmplx_0(k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_numbers))).#trainneg
cnf(i_0_1252, plain, (k3_xcmplx_0(esk8_1(X1),k1_xboole_0)=k1_xboole_0|~v1_membered(X1))).#trainneg
cnf(i_0_1253, plain, (k3_xcmplx_0(np__1,esk8_1(X1))=esk8_1(X1)|~v1_membered(X1))).#trainneg
cnf(i_0_1255, plain, (v1_xxreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1256, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1257, plain, (v3_ordinal1(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1271, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_1275, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_1281, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_1285, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_2506, plain, (v1_int_1(k1_xboole_0)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5013, plain, (v1_xboole_0(k3_xboole_0(X3,k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2639, negated_conjecture, (v3_ordinal1(esk5_0)|~v6_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1295, plain, (esk8_1(k1_zfmisc_1(X1))=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3312, plain, (r1_xxreal_0(esk31_0,esk30_0)|r1_xxreal_0(esk30_0,esk31_0))).#trainneg
cnf(i_0_2574, negated_conjecture, (v1_xreal_0(esk6_0)|~v6_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1296, plain, (X1=esk8_1(k1_zfmisc_1(X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1301, plain, (k3_xcmplx_0(X1,k3_prob_2(X2,X3,X4,X5))=k3_xcmplx_0(k3_prob_2(X2,X3,X4,X5),X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|~v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X3,k1_numbers)))|~m2_prob_1(X4,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_1293, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers))).#trainneg
cnf(i_0_5239, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers))).#trainneg
cnf(i_0_5240, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k5_numbers))).#trainneg
cnf(i_0_1294, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers))).#trainneg
cnf(i_0_1297, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_1298, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_1302, plain, (k3_xcmplx_0(k3_prob_2(X1,X2,X3,X4),k1_xboole_0)=k1_xboole_0|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1303, plain, (k3_xcmplx_0(np__1,k3_prob_2(X1,X2,X3,X4))=k3_prob_2(X1,X2,X3,X4)|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1304, negated_conjecture, (v1_xreal_0(esk4_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1306, negated_conjecture, (v1_xreal_0(esk4_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_2642, negated_conjecture, (v3_ordinal1(esk4_0)|~v6_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_2867, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_1308, plain, (r1_xxreal_0(X1,esk37_0)|r1_xxreal_0(esk37_0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3314, plain, (r1_xxreal_0(esk31_0,esk26_0)|r1_xxreal_0(esk26_0,esk31_0))).#trainneg
cnf(i_0_2578, negated_conjecture, (v1_xreal_0(esk5_0)|~v6_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_2582, negated_conjecture, (v1_xreal_0(esk4_0)|~v6_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1309, plain, (r1_xxreal_0(X1,esk32_0)|r1_xxreal_0(esk32_0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1310, plain, (r1_xxreal_0(X1,esk31_0)|r1_xxreal_0(esk31_0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1307, negated_conjecture, (v1_xreal_0(esk4_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1311, plain, (r1_xxreal_0(X1,esk30_0)|r1_xxreal_0(esk30_0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1312, plain, (r1_xxreal_0(X1,esk26_0)|r1_xxreal_0(esk26_0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1313, plain, (r1_xxreal_0(X1,esk25_0)|r1_xxreal_0(esk25_0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1314, plain, (r1_xxreal_0(X1,esk19_0)|r1_xxreal_0(esk19_0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1315, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1316, plain, (r1_xxreal_0(X1,k1_funct_1(X2,X3))|r1_xxreal_0(k1_funct_1(X2,X3),X1)|~v1_xboole_0(X1)|~v1_funct_1(X2)|~v2_valued_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1317, plain, (r1_xxreal_0(X1,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1320, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1321, plain, (r1_xxreal_0(X1,esk29_1(X2))|r1_xxreal_0(esk29_1(X2),X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1322, plain, (r1_xxreal_0(X1,esk8_1(k1_zfmisc_1(X2)))|r1_xxreal_0(esk8_1(k1_zfmisc_1(X2)),X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2736, negated_conjecture, (v1_xreal_0(esk6_0)|~m1_subset_1(esk2_0,k5_numbers))).#trainneg
cnf(i_0_1324, plain, (k3_xcmplx_0(X1,k5_prob_1(X2,X3,X4,X5))=k3_xcmplx_0(k5_prob_1(X2,X3,X4,X5),X1)|v1_xboole_0(X3)|~v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_3315, plain, (r1_xxreal_0(esk31_0,esk25_0)|r1_xxreal_0(esk25_0,esk31_0))).#trainneg
cnf(i_0_2634, negated_conjecture, (v1_xxreal_0(esk6_0)|~v6_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_3569, plain, (v5_membered(X3)|v1_xboole_0(X2)|~v5_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_3598, plain, (v6_ordinal1(X3)|v1_xboole_0(X2)|~v6_ordinal1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1325, plain, (k3_xcmplx_0(k5_prob_1(X1,X2,X3,X4),k1_xboole_0)=k1_xboole_0|v1_xboole_0(X2)|~v1_membered(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1326, plain, (k3_xcmplx_0(np__1,k5_prob_1(X1,X2,X3,X4))=k5_prob_1(X1,X2,X3,X4)|v1_xboole_0(X2)|~v1_membered(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1327, negated_conjecture, (v1_xcmplx_0(X2)|~v1_membered(k1_zfmisc_1(X1))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X1)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X1))).#trainneg
cnf(i_0_1328, plain, (v1_xcmplx_0(k3_xboole_0(X3,X4))|v1_xboole_0(X2)|~v1_membered(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1332, plain, (v1_xcmplx_0(X3)|v1_xboole_0(X2)|~v1_membered(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1337, plain, (v1_xcmplx_0(k3_xboole_0(X3,X4))|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_2651, plain, (v1_xxreal_0(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2739, negated_conjecture, (v1_xreal_0(esk5_0)|~m1_subset_1(esk2_0,k5_numbers))).#trainneg
cnf(i_0_3317, plain, (r1_xxreal_0(esk31_0,esk19_0)|r1_xxreal_0(esk19_0,esk31_0))).#trainneg
cnf(i_0_2637, negated_conjecture, (v1_xxreal_0(esk5_0)|~v6_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_3630, plain, (v6_membered(X3)|v1_xboole_0(X2)|~v6_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_3666, plain, (v1_membered(X3)|v1_xboole_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1348, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers))).#trainneg
cnf(i_0_5737, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers))).#trainneg
cnf(i_0_5738, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k5_numbers))).#trainneg
cnf(i_0_1349, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~v1_xreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers))).#trainneg
cnf(i_0_1355, negated_conjecture, (k3_xcmplx_0(X1,X3)=k3_xcmplx_0(X3,X1)|~v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X2)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X2))).#trainneg
cnf(i_0_1846, plain, (v1_xboole_0(esk29_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1356, plain, (k3_xcmplx_0(X1,k3_xboole_0(X4,X5))=k3_xcmplx_0(k3_xboole_0(X4,X5),X1)|v1_xboole_0(X3)|~v1_xcmplx_0(X1)|~v1_membered(X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_3321, plain, (r1_xxreal_0(esk31_0,np__1)|r1_xxreal_0(np__1,esk31_0))).#trainneg
cnf(i_0_2640, negated_conjecture, (v1_xxreal_0(esk4_0)|~v6_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_3731, plain, (v2_membered(X3)|v1_xboole_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_3754, plain, (v3_membered(X3)|v1_xboole_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1369, negated_conjecture, (k3_xcmplx_0(X2,k1_xboole_0)=k1_xboole_0|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X1)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X1))).#trainneg
cnf(i_0_1374, plain, (k3_xcmplx_0(X3,k1_xboole_0)=k1_xboole_0|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1376, negated_conjecture, (k3_xcmplx_0(np__1,X2)=X2|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X1)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X1))).#trainneg
cnf(i_0_1382, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_1383, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_numbers)|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_1386, plain, (r1_xxreal_0(X1,esk34_1(X2))|r1_xxreal_0(esk34_1(X2),X1)|v1_xboole_0(X2)|~v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1387, plain, (r1_xxreal_0(X1,esk34_1(X2))|r1_xxreal_0(esk34_1(X2),X1)|v1_xboole_0(X2)|~v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1388, negated_conjecture, (k3_xcmplx_0(X1,esk6_0)=k3_xcmplx_0(esk6_0,X1)|~v1_xcmplx_0(X1)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_4388, plain, (v1_xcmplx_0(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3687, plain, (esk29_1(X1)=esk29_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3323, plain, (r1_xxreal_0(esk31_0,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,esk31_0))).#trainneg
cnf(i_0_2797, negated_conjecture, (v1_xreal_0(esk4_0)|~m1_subset_1(esk2_0,k5_numbers))).#trainneg
cnf(i_0_3800, plain, (v4_membered(X3)|v1_xboole_0(X2)|~v4_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_2767, negated_conjecture, (v3_ordinal1(esk2_0)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))))).#trainneg
cnf(i_0_5941, negated_conjecture, (k3_xcmplx_0(esk6_0,esk5_0)=k3_xcmplx_0(esk5_0,esk6_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_5942, negated_conjecture, (k3_xcmplx_0(esk6_0,esk4_0)=k3_xcmplx_0(esk4_0,esk6_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_6070, negated_conjecture, (k3_xcmplx_0(esk6_0,esk5_0)=k3_xcmplx_0(esk5_0,esk6_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_6074, negated_conjecture, (k3_xcmplx_0(esk6_0,esk4_0)=k3_xcmplx_0(esk4_0,esk6_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_1389, negated_conjecture, (k3_xcmplx_0(esk6_0,k1_xboole_0)=k1_xboole_0|~v3_membered(esk2_0))).#trainneg
cnf(i_0_1390, negated_conjecture, (k3_xcmplx_0(np__1,esk6_0)=esk6_0|~v3_membered(esk2_0))).#trainneg
cnf(i_0_1398, plain, (r1_xxreal_0(esk18_0,X1)|r1_xxreal_0(X1,esk18_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1400, plain, (r1_xxreal_0(k1_funct_1(X1,X2),X3)|r1_xxreal_0(X3,k1_funct_1(X1,X2))|~v1_xboole_0(X3)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1402, plain, (r1_xxreal_0(k3_xcmplx_0(X1,X2),X3)|r1_xxreal_0(X3,k3_xcmplx_0(X1,X2))|~v1_xboole_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2800, negated_conjecture, (v1_xxreal_0(esk6_0)|~m1_subset_1(esk2_0,k5_numbers))).#trainneg
cnf(i_0_1403, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2845, negated_conjecture, (v1_xxreal_0(esk5_0)|~m1_subset_1(esk2_0,k5_numbers))).#trainneg
cnf(i_0_2874, negated_conjecture, (v1_xxreal_0(esk4_0)|~m1_subset_1(esk2_0,k5_numbers))).#trainneg
cnf(i_0_1406, plain, (r1_xxreal_0(X1,esk29_1(X2))|r1_xxreal_0(esk29_1(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1407, plain, (r1_xxreal_0(X1,esk29_1(X2))|r1_xxreal_0(esk29_1(X2),X1)|~v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1408, plain, (r1_xxreal_0(X1,esk15_1(X2))|r1_xxreal_0(esk15_1(X2),X1)|v1_xboole_0(X2)|~v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1409, plain, (r1_xxreal_0(X1,esk15_1(X2))|r1_xxreal_0(esk15_1(X2),X1)|v1_xboole_0(X2)|~v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1412, plain, (k3_xcmplx_0(X1,esk22_1(X2))=k3_xcmplx_0(esk22_1(X2),X1)|~v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X2)),k1_zfmisc_1(k1_zfmisc_1(X3)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(X2)),X3)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(X2)),X3))).#trainneg
cnf(i_0_1414, plain, (k3_xcmplx_0(X1,esk13_1(X2))=k3_xcmplx_0(esk13_1(X2),X1)|~v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X2)),k1_zfmisc_1(k1_zfmisc_1(X3)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(X2)),X3)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(X2)),X3))).#trainneg
cnf(i_0_1416, negated_conjecture, (k3_xcmplx_0(X1,esk6_0)=k3_xcmplx_0(esk6_0,X1)|~v1_xcmplx_0(X1)|~v1_membered(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(esk2_0,X2)|~v1_prob_1(esk2_0,X2))).#trainneg
cnf(i_0_1417, negated_conjecture, (k3_xcmplx_0(X1,esk6_0)=k3_xcmplx_0(esk6_0,X1)|~v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(esk1_0))|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(k1_zfmisc_1(esk1_0),X2)|~v1_prob_1(k1_zfmisc_1(esk1_0),X2))).#trainneg
cnf(i_0_1418, negated_conjecture, (k3_xcmplx_0(X1,esk5_0)=k3_xcmplx_0(esk5_0,X1)|~v1_xcmplx_0(X1)|~v1_membered(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(esk2_0,X2)|~v1_prob_1(esk2_0,X2))).#trainneg
cnf(i_0_5965, plain, (m1_subset_1(esk29_1(X2),k1_zfmisc_1(X1))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_4754, negated_conjecture, (v1_xcmplx_0(esk6_0)|~v6_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_4780, negated_conjecture, (k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~v1_xcmplx_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_1436, plain, (k3_xcmplx_0(k1_xboole_0,k3_xboole_0(X3,X4))=k1_xboole_0|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_4787, negated_conjecture, (v1_xcmplx_0(esk5_0)|~v6_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_4768, plain, (v1_valued_0(X3)|v1_xboole_0(X2)|~v1_valued_0(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_4805, plain, (v2_valued_0(X3)|v1_xboole_0(X2)|~v2_valued_0(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1419, negated_conjecture, (k3_xcmplx_0(X1,esk5_0)=k3_xcmplx_0(esk5_0,X1)|~v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(esk1_0))|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(k1_zfmisc_1(esk1_0),X2)|~v1_prob_1(k1_zfmisc_1(esk1_0),X2))).#trainneg
cnf(i_0_1420, negated_conjecture, (k3_xcmplx_0(X1,esk4_0)=k3_xcmplx_0(esk4_0,X1)|~v1_xcmplx_0(X1)|~v1_membered(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(esk2_0,X2)|~v1_prob_1(esk2_0,X2))).#trainneg
cnf(i_0_1421, negated_conjecture, (k3_xcmplx_0(X1,esk4_0)=k3_xcmplx_0(esk4_0,X1)|~v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(esk1_0))|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(k1_zfmisc_1(esk1_0),X2)|~v1_prob_1(k1_zfmisc_1(esk1_0),X2))).#trainneg
cnf(i_0_1422, negated_conjecture, (k3_xcmplx_0(X1,esk3_0)=k3_xcmplx_0(esk3_0,X1)|~v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),X2)|~v1_prob_1(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),X2))).#trainneg
cnf(i_0_1423, negated_conjecture, (k3_xcmplx_0(X1,esk2_0)=k3_xcmplx_0(esk2_0,X1)|~v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X2)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X2))).#trainneg
cnf(i_0_1428, negated_conjecture, (k3_xcmplx_0(X1,k3_xboole_0(X2,X3))=k3_xcmplx_0(k3_xboole_0(X2,X3),X1)|~v1_xcmplx_0(X1)|~v1_membered(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(X4)))|~v4_prob_1(esk2_0,X4)|~v1_prob_1(esk2_0,X4)|~m1_subset_1(X3,esk2_0)|~m1_subset_1(X2,esk2_0))).#trainneg
cnf(i_0_1429, plain, (k3_xcmplx_0(X1,k5_numbers)=k3_xcmplx_0(k5_numbers,X1)|~v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(k1_numbers))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(k1_zfmisc_1(k1_numbers),X2)|~v1_prob_1(k1_zfmisc_1(k1_numbers),X2))).#trainneg
cnf(i_0_4791, negated_conjecture, (v1_xcmplx_0(esk4_0)|~v6_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1440, plain, (k3_xcmplx_0(k1_xboole_0,X3)=k1_xboole_0|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_4826, plain, (v3_valued_0(X3)|v1_xboole_0(X2)|~v3_valued_0(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_4941, plain, (v4_valued_0(X3)|v1_xboole_0(X2)|~v4_valued_0(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1446, plain, (k3_xcmplx_0(k3_xboole_0(X1,esk8_1(X2)),k1_xboole_0)=k1_xboole_0|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,X2)|~v4_prob_1(X2,X3)|~v1_prob_1(X2,X3))).#trainneg
cnf(i_0_1463, plain, (k3_xcmplx_0(k3_xboole_0(X1,k5_prob_1(X2,X3,X4,X5)),k1_xboole_0)=k1_xboole_0|v1_xboole_0(X3)|~v1_membered(X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X6)))|~m1_subset_1(X1,X3)|~v4_prob_1(X3,X6)|~v1_prob_1(X3,X6)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_1433, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X2))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X3)))|~v4_prob_1(k1_zfmisc_1(X2),X3)|~v1_prob_1(k1_zfmisc_1(X2),X3))).#trainneg
cnf(i_0_1435, negated_conjecture, (k3_xcmplx_0(k1_xboole_0,X2)=k1_xboole_0|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X1)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X1))).#trainneg
cnf(i_0_1441, plain, (k3_xcmplx_0(k3_xboole_0(X1,esk34_1(X2)),k1_xboole_0)=k1_xboole_0|v1_xboole_0(X2)|~v1_membered(k1_zfmisc_1(X2))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X3)|~v1_prob_1(k1_zfmisc_1(X2),X3))).#trainneg
cnf(i_0_1442, plain, (k3_xcmplx_0(k3_xboole_0(X1,esk29_1(X2)),k1_xboole_0)=k1_xboole_0|~v1_membered(k1_zfmisc_1(X2))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X3)|~v1_prob_1(k1_zfmisc_1(X2),X3))).#trainneg
cnf(i_0_1443, plain, (k3_xcmplx_0(k3_xboole_0(X1,esk22_1(X2)),k1_xboole_0)=k1_xboole_0|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X2)),k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(X2)),X3)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(X2)),X3))).#trainneg
cnf(i_0_1444, plain, (k3_xcmplx_0(k3_xboole_0(X1,esk15_1(X2)),k1_xboole_0)=k1_xboole_0|v1_xboole_0(X2)|~v1_membered(k1_zfmisc_1(X2))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X3)|~v1_prob_1(k1_zfmisc_1(X2),X3))).#trainneg
cnf(i_0_1445, plain, (k3_xcmplx_0(k3_xboole_0(X1,esk13_1(X2)),k1_xboole_0)=k1_xboole_0|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X2)),k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(X2)),X3)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(X2)),X3))).#trainneg
cnf(i_0_1447, negated_conjecture, (k3_xcmplx_0(k3_xboole_0(X1,esk6_0),k1_xboole_0)=k1_xboole_0|~v1_membered(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,esk2_0)|~v4_prob_1(esk2_0,X2)|~v1_prob_1(esk2_0,X2))).#trainneg
cnf(i_0_1448, negated_conjecture, (k3_xcmplx_0(k3_xboole_0(X1,esk6_0),k1_xboole_0)=k1_xboole_0|~v1_membered(k1_zfmisc_1(esk1_0))|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))|~v4_prob_1(k1_zfmisc_1(esk1_0),X2)|~v1_prob_1(k1_zfmisc_1(esk1_0),X2))).#trainneg
cnf(i_0_1449, negated_conjecture, (k3_xcmplx_0(k3_xboole_0(X1,esk5_0),k1_xboole_0)=k1_xboole_0|~v1_membered(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,esk2_0)|~v4_prob_1(esk2_0,X2)|~v1_prob_1(esk2_0,X2))).#trainneg
cnf(i_0_1450, negated_conjecture, (k3_xcmplx_0(k3_xboole_0(X1,esk5_0),k1_xboole_0)=k1_xboole_0|~v1_membered(k1_zfmisc_1(esk1_0))|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))|~v4_prob_1(k1_zfmisc_1(esk1_0),X2)|~v1_prob_1(k1_zfmisc_1(esk1_0),X2))).#trainneg
cnf(i_0_2617, plain, (v1_xxreal_0(esk8_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_1470, plain, (k3_xcmplx_0(np__1,k3_xboole_0(X1,esk8_1(X2)))=k3_xboole_0(X1,esk8_1(X2))|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,X2)|~v4_prob_1(X2,X3)|~v1_prob_1(X2,X3))).#trainneg
cnf(i_0_3368, plain, (r1_xxreal_0(esk30_0,esk26_0)|r1_xxreal_0(esk26_0,esk30_0))).#trainneg
cnf(i_0_1586, plain, (v1_xxreal_0(X3)|v1_xboole_0(X2)|~v2_membered(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1937, plain, (v1_xreal_0(X3)|v1_xboole_0(X2)|~v3_membered(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1487, plain, (k3_xcmplx_0(np__1,k3_xboole_0(X1,k5_prob_1(X2,X3,X4,X5)))=k3_xboole_0(X1,k5_prob_1(X2,X3,X4,X5))|v1_xboole_0(X3)|~v1_membered(X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X6)))|~m1_subset_1(X1,X3)|~v4_prob_1(X3,X6)|~v1_prob_1(X3,X6)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_1494, plain, (r1_xxreal_0(X1,esk8_1(X2))|r1_xxreal_0(esk8_1(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(X2))).#trainneg
cnf(i_0_6544, plain, (r1_xxreal_0(esk30_0,esk26_0)|~v1_xreal_0(esk26_0))).#trainneg
cnf(i_0_1451, negated_conjecture, (k3_xcmplx_0(k3_xboole_0(X1,esk4_0),k1_xboole_0)=k1_xboole_0|~v1_membered(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,esk2_0)|~v4_prob_1(esk2_0,X2)|~v1_prob_1(esk2_0,X2))).#trainneg
cnf(i_0_1452, negated_conjecture, (k3_xcmplx_0(k3_xboole_0(X1,esk4_0),k1_xboole_0)=k1_xboole_0|~v1_membered(k1_zfmisc_1(esk1_0))|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))|~v4_prob_1(k1_zfmisc_1(esk1_0),X2)|~v1_prob_1(k1_zfmisc_1(esk1_0),X2))).#trainneg
cnf(i_0_1453, negated_conjecture, (k3_xcmplx_0(k3_xboole_0(X1,esk3_0),k1_xboole_0)=k1_xboole_0|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)))|~v4_prob_1(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),X2)|~v1_prob_1(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),X2))).#trainneg
cnf(i_0_1454, negated_conjecture, (k3_xcmplx_0(k3_xboole_0(X1,esk2_0),k1_xboole_0)=k1_xboole_0|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X2)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X2))).#trainneg
cnf(i_0_1455, plain, (k3_xcmplx_0(k3_xboole_0(X1,np__1),k1_xboole_0)=k1_xboole_0|~v1_membered(k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_numbers)|~v4_prob_1(k1_numbers,X2)|~v1_prob_1(k1_numbers,X2))).#trainneg
cnf(i_0_1456, plain, (k3_xcmplx_0(k3_xboole_0(X1,np__1),k1_xboole_0)=k1_xboole_0|~v1_membered(k5_numbers)|~m1_subset_1(k5_numbers,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k5_numbers)|~v4_prob_1(k5_numbers,X2)|~v1_prob_1(k5_numbers,X2))).#trainneg
cnf(i_0_1459, negated_conjecture, (k3_xcmplx_0(k3_xboole_0(X1,k3_xboole_0(X2,X3)),k1_xboole_0)=k1_xboole_0|~v1_membered(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(X4)))|~m1_subset_1(X1,esk2_0)|~v4_prob_1(esk2_0,X4)|~v1_prob_1(esk2_0,X4)|~m1_subset_1(X3,esk2_0)|~m1_subset_1(X2,esk2_0))).#trainneg
cnf(i_0_1460, plain, (k3_xcmplx_0(k3_xboole_0(X1,k5_numbers),k1_xboole_0)=k1_xboole_0|~v1_membered(k1_zfmisc_1(k1_numbers))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v4_prob_1(k1_zfmisc_1(k1_numbers),X2)|~v1_prob_1(k1_zfmisc_1(k1_numbers),X2))).#trainneg
cnf(i_0_1461, plain, (k3_xcmplx_0(k3_xboole_0(X1,k3_prob_2(X2,X3,X4,X5)),k1_xboole_0)=k1_xboole_0|v1_xboole_0(X2)|v1_xboole_0(X3)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X3,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X3,k1_numbers)),k1_zfmisc_1(k1_zfmisc_1(X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,k1_numbers)))|~v4_prob_1(k1_zfmisc_1(k2_zfmisc_1(X3,k1_numbers)),X6)|~v1_prob_1(k1_zfmisc_1(k2_zfmisc_1(X3,k1_numbers)),X6)|~m2_prob_1(X4,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_1462, plain, (k3_xcmplx_0(k3_xboole_0(X1,k5_prob_1(X2,X3,X4,X5)),k1_xboole_0)=k1_xboole_0|v1_xboole_0(X3)|~v1_membered(k1_zfmisc_1(X2))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X6)))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X6)|~v1_prob_1(k1_zfmisc_1(X2),X6)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_1464, plain, (k3_xcmplx_0(k3_xboole_0(X1,X2),k1_xboole_0)=k1_xboole_0|~v1_membered(k1_zfmisc_1(X2))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X3)|~v1_prob_1(k1_zfmisc_1(X2),X3))).#trainneg
cnf(i_0_2618, plain, (v1_xreal_0(esk8_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_1495, plain, (r1_xxreal_0(X1,esk8_1(X2))|r1_xxreal_0(esk8_1(X2),X1)|~v1_xboole_0(X1)|~v2_membered(X2))).#trainneg
cnf(i_0_3369, plain, (r1_xxreal_0(esk30_0,esk25_0)|r1_xxreal_0(esk25_0,esk30_0))).#trainneg
cnf(i_0_2248, plain, (v1_rat_1(X3)|v1_xboole_0(X2)|~v4_membered(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_2553, plain, (v1_int_1(X3)|v1_xboole_0(X2)|~v5_membered(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1465, plain, (k3_xcmplx_0(np__1,k3_xboole_0(X1,esk34_1(X2)))=k3_xboole_0(X1,esk34_1(X2))|v1_xboole_0(X2)|~v1_membered(k1_zfmisc_1(X2))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X3)|~v1_prob_1(k1_zfmisc_1(X2),X3))).#trainneg
cnf(i_0_1466, plain, (k3_xcmplx_0(np__1,k3_xboole_0(X1,esk29_1(X2)))=k3_xboole_0(X1,esk29_1(X2))|~v1_membered(k1_zfmisc_1(X2))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X3)|~v1_prob_1(k1_zfmisc_1(X2),X3))).#trainneg
cnf(i_0_1467, plain, (k3_xcmplx_0(np__1,k3_xboole_0(X1,esk22_1(X2)))=k3_xboole_0(X1,esk22_1(X2))|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X2)),k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(X2)),X3)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(X2)),X3))).#trainneg
cnf(i_0_1468, plain, (k3_xcmplx_0(np__1,k3_xboole_0(X1,esk15_1(X2)))=k3_xboole_0(X1,esk15_1(X2))|v1_xboole_0(X2)|~v1_membered(k1_zfmisc_1(X2))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X3)|~v1_prob_1(k1_zfmisc_1(X2),X3))).#trainneg
cnf(i_0_1469, plain, (k3_xcmplx_0(np__1,k3_xboole_0(X1,esk13_1(X2)))=k3_xboole_0(X1,esk13_1(X2))|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X2)),k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(X2)),X3)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(X2)),X3))).#trainneg
cnf(i_0_1471, negated_conjecture, (k3_xcmplx_0(np__1,k3_xboole_0(X1,esk6_0))=k3_xboole_0(X1,esk6_0)|~v1_membered(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,esk2_0)|~v4_prob_1(esk2_0,X2)|~v1_prob_1(esk2_0,X2))).#trainneg
cnf(i_0_1472, negated_conjecture, (k3_xcmplx_0(np__1,k3_xboole_0(X1,esk6_0))=k3_xboole_0(X1,esk6_0)|~v1_membered(k1_zfmisc_1(esk1_0))|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))|~v4_prob_1(k1_zfmisc_1(esk1_0),X2)|~v1_prob_1(k1_zfmisc_1(esk1_0),X2))).#trainneg
cnf(i_0_1473, negated_conjecture, (k3_xcmplx_0(np__1,k3_xboole_0(X1,esk5_0))=k3_xboole_0(X1,esk5_0)|~v1_membered(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,esk2_0)|~v4_prob_1(esk2_0,X2)|~v1_prob_1(esk2_0,X2))).#trainneg
cnf(i_0_1474, negated_conjecture, (k3_xcmplx_0(np__1,k3_xboole_0(X1,esk5_0))=k3_xboole_0(X1,esk5_0)|~v1_membered(k1_zfmisc_1(esk1_0))|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))|~v4_prob_1(k1_zfmisc_1(esk1_0),X2)|~v1_prob_1(k1_zfmisc_1(esk1_0),X2))).#trainneg
cnf(i_0_1475, negated_conjecture, (k3_xcmplx_0(np__1,k3_xboole_0(X1,esk4_0))=k3_xboole_0(X1,esk4_0)|~v1_membered(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,esk2_0)|~v4_prob_1(esk2_0,X2)|~v1_prob_1(esk2_0,X2))).#trainneg
cnf(i_0_1476, negated_conjecture, (k3_xcmplx_0(np__1,k3_xboole_0(X1,esk4_0))=k3_xboole_0(X1,esk4_0)|~v1_membered(k1_zfmisc_1(esk1_0))|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))|~v4_prob_1(k1_zfmisc_1(esk1_0),X2)|~v1_prob_1(k1_zfmisc_1(esk1_0),X2))).#trainneg
cnf(i_0_1477, negated_conjecture, (k3_xcmplx_0(np__1,k3_xboole_0(X1,esk3_0))=k3_xboole_0(X1,esk3_0)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)))|~v4_prob_1(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),X2)|~v1_prob_1(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),X2))).#trainneg
cnf(i_0_5119, plain, (v1_rat_1(k1_xboole_0)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2619, plain, (v3_ordinal1(esk8_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_3371, plain, (r1_xxreal_0(esk30_0,esk19_0)|r1_xxreal_0(esk19_0,esk30_0))).#trainneg
cnf(i_0_1837, negated_conjecture, (v6_membered(esk2_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_2386, negated_conjecture, (v1_xxreal_0(esk6_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2681, plain, (v7_ordinal1(X3)|v1_xboole_0(X2)|~v6_membered(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1503, plain, (r1_xxreal_0(esk18_0,X1)|r1_xxreal_0(X1,esk18_0)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1478, negated_conjecture, (k3_xcmplx_0(np__1,k3_xboole_0(X1,esk2_0))=k3_xboole_0(X1,esk2_0)|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X2)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X2))).#trainneg
cnf(i_0_1479, plain, (k3_xcmplx_0(np__1,k3_xboole_0(X1,np__1))=k3_xboole_0(X1,np__1)|~v1_membered(k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_numbers)|~v4_prob_1(k1_numbers,X2)|~v1_prob_1(k1_numbers,X2))).#trainneg
cnf(i_0_1480, plain, (k3_xcmplx_0(np__1,k3_xboole_0(X1,np__1))=k3_xboole_0(X1,np__1)|~v1_membered(k5_numbers)|~m1_subset_1(k5_numbers,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k5_numbers)|~v4_prob_1(k5_numbers,X2)|~v1_prob_1(k5_numbers,X2))).#trainneg
cnf(i_0_1483, negated_conjecture, (k3_xcmplx_0(np__1,k3_xboole_0(X1,k3_xboole_0(X2,X3)))=k3_xboole_0(X1,k3_xboole_0(X2,X3))|~v1_membered(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(X4)))|~m1_subset_1(X1,esk2_0)|~v4_prob_1(esk2_0,X4)|~v1_prob_1(esk2_0,X4)|~m1_subset_1(X3,esk2_0)|~m1_subset_1(X2,esk2_0))).#trainneg
cnf(i_0_1484, plain, (k3_xcmplx_0(np__1,k3_xboole_0(X1,k5_numbers))=k3_xboole_0(X1,k5_numbers)|~v1_membered(k1_zfmisc_1(k1_numbers))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v4_prob_1(k1_zfmisc_1(k1_numbers),X2)|~v1_prob_1(k1_zfmisc_1(k1_numbers),X2))).#trainneg
cnf(i_0_1485, plain, (k3_xcmplx_0(np__1,k3_xboole_0(X1,k3_prob_2(X2,X3,X4,X5)))=k3_xboole_0(X1,k3_prob_2(X2,X3,X4,X5))|v1_xboole_0(X2)|v1_xboole_0(X3)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X3,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X3,k1_numbers)),k1_zfmisc_1(k1_zfmisc_1(X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,k1_numbers)))|~v4_prob_1(k1_zfmisc_1(k2_zfmisc_1(X3,k1_numbers)),X6)|~v1_prob_1(k1_zfmisc_1(k2_zfmisc_1(X3,k1_numbers)),X6)|~m2_prob_1(X4,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_1486, plain, (k3_xcmplx_0(np__1,k3_xboole_0(X1,k5_prob_1(X2,X3,X4,X5)))=k3_xboole_0(X1,k5_prob_1(X2,X3,X4,X5))|v1_xboole_0(X3)|~v1_membered(k1_zfmisc_1(X2))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X6)))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X6)|~v1_prob_1(k1_zfmisc_1(X2),X6)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_1488, plain, (k3_xcmplx_0(np__1,k3_xboole_0(X1,X2))=k3_xboole_0(X1,X2)|~v1_membered(k1_zfmisc_1(X2))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X3)|~v1_prob_1(k1_zfmisc_1(X2),X3))).#trainneg
cnf(i_0_1492, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_1493, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers)|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_1505, plain, (r1_xxreal_0(k1_funct_1(X1,X2),X3)|r1_xxreal_0(X3,k1_funct_1(X1,X2))|~v1_xxreal_0(X3)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3542, plain, (v3_membered(esk29_1(X1))|~v5_membered(X1))).#trainneg
cnf(i_0_2388, negated_conjecture, (v1_xxreal_0(esk6_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_3375, plain, (r1_xxreal_0(esk30_0,np__1)|r1_xxreal_0(np__1,esk30_0))).#trainneg
cnf(i_0_2389, negated_conjecture, (v1_xxreal_0(esk6_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2657, negated_conjecture, (v1_xxreal_0(esk5_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2780, plain, (v3_ordinal1(X3)|v1_xboole_0(X2)|~v3_ordinal1(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1507, plain, (r1_xxreal_0(k3_xcmplx_0(X1,X2),X3)|r1_xxreal_0(X3,k3_xcmplx_0(X1,X2))|~v1_xxreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1508, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1512, plain, (r1_xxreal_0(X1,k3_prob_2(X2,X3,X4,X5))|r1_xxreal_0(k3_prob_2(X2,X3,X4,X5),X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|~v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X3,k1_numbers)))|~m2_prob_1(X4,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_1513, plain, (r1_xxreal_0(X1,k3_prob_2(X2,X3,X4,X5))|r1_xxreal_0(k3_prob_2(X2,X3,X4,X5),X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|~v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X3,k1_numbers)))|~m2_prob_1(X4,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_1516, plain, (k3_xcmplx_0(np__1,esk22_1(X1))=esk22_1(X1)|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X1)),k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(X1)),X2)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(X1)),X2))).#trainneg
cnf(i_0_1518, plain, (k3_xcmplx_0(np__1,esk13_1(X1))=esk13_1(X1)|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X1)),k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(X1)),X2)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(X1)),X2))).#trainneg
cnf(i_0_1520, negated_conjecture, (k3_xcmplx_0(np__1,esk6_0)=esk6_0|~v1_membered(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(esk2_0,X1)|~v1_prob_1(esk2_0,X1))).#trainneg
cnf(i_0_1521, negated_conjecture, (k3_xcmplx_0(np__1,esk6_0)=esk6_0|~v1_membered(k1_zfmisc_1(esk1_0))|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(k1_zfmisc_1(esk1_0),X1)|~v1_prob_1(k1_zfmisc_1(esk1_0),X1))).#trainneg
cnf(i_0_1522, negated_conjecture, (k3_xcmplx_0(np__1,esk5_0)=esk5_0|~v1_membered(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(esk2_0,X1)|~v1_prob_1(esk2_0,X1))).#trainneg
cnf(i_0_2159, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_3765, plain, (v3_membered(esk29_1(X1))|~v4_membered(X1))).#trainneg
cnf(i_0_2659, negated_conjecture, (v1_xxreal_0(esk5_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_5253, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_1559, plain, (k8_real_1(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3377, plain, (r1_xxreal_0(esk30_0,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,esk30_0))).#trainneg
cnf(i_0_2660, negated_conjecture, (v1_xxreal_0(esk5_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2661, negated_conjecture, (v1_xxreal_0(esk4_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5620, plain, (v1_xcmplx_0(X1)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,X2)|~v4_prob_1(X2,X3)|~v1_prob_1(X2,X3)|~v3_membered(k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_1558, plain, (k3_xcmplx_0(k1_seq_1(X1,X2),k1_seq_1(k3_prob_2(X3,X4,X1,X2),X5))=k1_seq_1(X1,k5_prob_1(X3,X4,X5,X2))|r1_xxreal_0(k1_seq_1(X1,X2),k1_xboole_0)|v1_xboole_0(X4)|v1_xboole_0(X3)|~v1_xreal_0(k1_seq_1(X1,X2))|~m1_prob_1(X5,X3,X4)|~m1_prob_1(X2,X3,X4)|~m2_prob_1(X1,X3,X4)|~m1_subset_1(k1_seq_1(k3_prob_2(X3,X4,X1,X2),X5),k1_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3)))|~v4_prob_1(X4,X3)|~v1_prob_1(X4,X3)|~v1_xreal_0(k1_seq_1(k3_prob_2(X3,X4,X1,X2),X5))|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers))).#trainneg
cnf(i_0_6866, negated_conjecture, (k8_real_1(k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_6969, negated_conjecture, (k8_real_1(k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_6976, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~m1_subset_1(k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_6979, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~m1_subset_1(k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_2555, plain, (v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_3541, plain, (v3_membered(esk34_1(X1))|v1_xboole_0(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_2663, negated_conjecture, (v1_xxreal_0(esk4_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1579, plain, (r1_xxreal_0(X1,k5_prob_1(X2,X3,X4,X5))|r1_xxreal_0(k5_prob_1(X2,X3,X4,X5),X1)|v1_xboole_0(X3)|~v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_2664, negated_conjecture, (v1_xxreal_0(esk4_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2666, negated_conjecture, (v3_membered(esk2_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_6545, plain, (v1_xxreal_0(X1)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,X2)|~v4_prob_1(X2,X3)|~v1_prob_1(X2,X3)|~v3_membered(k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_6547, plain, (v1_xreal_0(X1)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,X2)|~v4_prob_1(X2,X3)|~v1_prob_1(X2,X3)|~v6_membered(k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_1566, plain, (v1_xreal_0(k1_seq_1(X3,k5_prob_1(X1,X2,X5,X4)))|r1_xxreal_0(k1_seq_1(X3,X4),k1_xboole_0)|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xreal_0(k1_seq_1(X3,X4))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|~m1_prob_1(X5,X1,X2)|~m1_prob_1(X4,X1,X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1567, plain, (v1_xreal_0(k1_seq_1(X1,k5_prob_1(X3,X4,X5,X2)))|r1_xxreal_0(k1_seq_1(X1,X2),k1_xboole_0)|v1_xboole_0(X3)|v1_xboole_0(X4)|~v1_xreal_0(k1_seq_1(k3_prob_2(X3,X4,X1,X2),X5))|~v1_xreal_0(k1_seq_1(X1,X2))|~m1_prob_1(X5,X3,X4)|~m1_prob_1(X2,X3,X4)|~m2_prob_1(X1,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v4_prob_1(X4,X3)|~v1_prob_1(X4,X3))).#trainneg
cnf(i_0_7000, plain, (k3_xcmplx_0(k1_seq_1(k3_prob_2(X3,X4,X1,X2),X5),k1_seq_1(X1,X2))=k1_seq_1(X1,k5_prob_1(X3,X4,X5,X2))|r1_xxreal_0(k1_seq_1(X1,X2),k1_xboole_0)|v1_xboole_0(X4)|v1_xboole_0(X3)|~v1_xreal_0(k1_seq_1(k3_prob_2(X3,X4,X1,X2),X5))|~m1_prob_1(X5,X3,X4)|~m1_prob_1(X2,X3,X4)|~m2_prob_1(X1,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v4_prob_1(X4,X3)|~v1_prob_1(X4,X3)|~m1_subset_1(k1_seq_1(k3_prob_2(X3,X4,X1,X2),X5),k1_numbers))).#trainneg
cnf(i_0_7022, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0)))).#trainneg
cnf(i_0_7028, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0)))).#trainneg
cnf(i_0_2612, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_3544, plain, (v3_membered(esk15_1(X1))|v1_xboole_0(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_2737, negated_conjecture, (v1_xreal_0(esk6_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1580, plain, (r1_xxreal_0(X1,k5_prob_1(X2,X3,X4,X5))|r1_xxreal_0(k5_prob_1(X2,X3,X4,X5),X1)|v1_xboole_0(X3)|~v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_2740, negated_conjecture, (v1_xreal_0(esk5_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2749, negated_conjecture, (v3_ordinal1(esk6_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_6621, plain, (v1_rat_1(X1)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,X2)|~v4_prob_1(X2,X3)|~v1_prob_1(X2,X3)|~v5_membered(k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_1582, plain, (v1_xxreal_0(k3_xboole_0(X3,X4))|v1_xboole_0(X2)|~v2_membered(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_7031, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0)))).#trainneg
cnf(i_0_7096, negated_conjecture, (k3_xcmplx_0(esk6_0,esk30_0)=k3_xcmplx_0(esk30_0,esk6_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_7097, negated_conjecture, (k3_xcmplx_0(esk30_0,esk5_0)=k3_xcmplx_0(esk5_0,esk30_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_7098, negated_conjecture, (k3_xcmplx_0(esk30_0,esk4_0)=k3_xcmplx_0(esk4_0,esk30_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_7099, negated_conjecture, (k3_xcmplx_0(esk6_0,esk25_0)=k3_xcmplx_0(esk25_0,esk6_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_7100, negated_conjecture, (k3_xcmplx_0(esk25_0,esk5_0)=k3_xcmplx_0(esk5_0,esk25_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_7101, negated_conjecture, (k3_xcmplx_0(esk25_0,esk4_0)=k3_xcmplx_0(esk4_0,esk25_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_7102, negated_conjecture, (k3_xcmplx_0(esk6_0,esk5_0)=k3_xcmplx_0(esk5_0,esk6_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_7103, negated_conjecture, (k3_xcmplx_0(esk6_0,esk4_0)=k3_xcmplx_0(esk4_0,esk6_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_4640, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_3764, plain, (v3_membered(esk34_1(X1))|v1_xboole_0(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_2750, negated_conjecture, (v3_ordinal1(esk6_0)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_6966, negated_conjecture, (k3_xcmplx_0(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_numbers))).#trainneg
cnf(i_0_2753, negated_conjecture, (v3_ordinal1(esk5_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_2754, negated_conjecture, (v3_ordinal1(esk5_0)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_6623, plain, (v1_int_1(X1)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,X2)|~v4_prob_1(X2,X3)|~v1_prob_1(X2,X3)|~v6_membered(k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_1589, negated_conjecture, (v5_membered(esk2_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_7131, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_5165, plain, (v1_int_1(k1_xboole_0)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_3776, plain, (v3_membered(esk15_1(X1))|v1_xboole_0(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_2757, negated_conjecture, (v3_ordinal1(esk4_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1591, plain, (r1_xxreal_0(X1,k5_prob_1(X2,X3,X4,X5))|r1_xxreal_0(k5_prob_1(X2,X3,X4,X5),X1)|v1_xboole_0(X3)|~v1_xxreal_0(X1)|~v2_membered(X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_2758, negated_conjecture, (v3_ordinal1(esk4_0)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_2798, negated_conjecture, (v1_xreal_0(esk4_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_6935, plain, (v1_xcmplx_0(X1)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,X2)|~v4_prob_1(X2,X3)|~v1_prob_1(X2,X3)|~v6_membered(k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_1592, plain, (r1_xxreal_0(X1,k5_prob_1(X2,X3,X4,X5))|r1_xxreal_0(k5_prob_1(X2,X3,X4,X5),X1)|v1_xboole_0(X3)|~v1_xboole_0(X1)|~v2_membered(X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_1594, plain, (v1_xxreal_0(k3_xboole_0(X3,X4))|v1_xboole_0(X2)|~v2_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_7138, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_7139, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_7145, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_7146, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_7185, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_5656, plain, (v1_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_1823, plain, (v1_relat_1(esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_2801, negated_conjecture, (v1_xxreal_0(esk6_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_6970, negated_conjecture, (k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_3391, plain, (r1_xxreal_0(esk26_0,esk25_0)|r1_xxreal_0(esk25_0,esk26_0))).#trainneg
cnf(i_0_2846, negated_conjecture, (v1_xxreal_0(esk5_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2875, negated_conjecture, (v1_xxreal_0(esk4_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_7105, plain, (v1_xxreal_0(X1)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,X2)|~v4_prob_1(X2,X3)|~v1_prob_1(X2,X3)|~v6_membered(k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_7186, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_7241, negated_conjecture, (k3_xcmplx_0(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_numbers))).#trainneg
cnf(i_0_7245, negated_conjecture, (k3_xcmplx_0(k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_7247, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_numbers)|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_6638, plain, (v1_rat_1(k1_xboole_0)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_1906, plain, (v5_relat_1(esk29_1(k2_zfmisc_1(X1,X2)),X2))).#trainneg
cnf(i_0_3283, negated_conjecture, (v3_membered(esk6_0)|~m1_subset_1(esk1_0,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_6972, negated_conjecture, (k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_1599, plain, (r1_xxreal_0(esk34_1(X1),X2)|r1_xxreal_0(X2,esk34_1(X1))|v1_xboole_0(X1)|~v1_xxreal_0(X2)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3393, plain, (r1_xxreal_0(esk26_0,esk19_0)|r1_xxreal_0(esk19_0,esk26_0))).#trainneg
cnf(i_0_3285, negated_conjecture, (v3_membered(esk5_0)|~m1_subset_1(esk1_0,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_3463, negated_conjecture, (v3_membered(esk4_0)|~m1_subset_1(esk1_0,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_7182, plain, (v1_rat_1(X1)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,X2)|~v4_prob_1(X2,X3)|~v1_prob_1(X2,X3)|~v6_membered(k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_1600, plain, (r1_xxreal_0(esk34_1(X1),X2)|r1_xxreal_0(X2,esk34_1(X1))|v1_xboole_0(X1)|~v1_xboole_0(X2)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_7265, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers))).#trainneg
cnf(i_0_1830, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk29_1(X2))|~m1_subset_1(X1,esk29_1(X2)))).#trainneg
cnf(i_0_1916, plain, (v4_relat_1(esk29_1(k2_zfmisc_1(X1,X2)),X1))).#trainneg
cnf(i_0_3554, negated_conjecture, (v3_membered(esk6_0)|~m1_subset_1(esk1_0,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_3397, plain, (r1_xxreal_0(esk26_0,np__1)|r1_xxreal_0(np__1,esk26_0))).#trainneg
cnf(i_0_3556, negated_conjecture, (v3_membered(esk5_0)|~m1_subset_1(esk1_0,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_3570, negated_conjecture, (v5_membered(esk3_0)|~v6_membered(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_3530, plain, (X1=k5_prob_1(X2,X3,X4,X5)|v1_xboole_0(X3)|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1605, plain, (v1_rat_1(k1_seq_1(X1,X2))|~v4_membered(k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1606, plain, (v1_int_1(k1_seq_1(X1,X2))|~v5_membered(k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7356, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xreal_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0))).#trainneg
cnf(i_0_7431, negated_conjecture, (k3_xcmplx_0(k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~m1_subset_1(k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_7433, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers)|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_7436, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0)))).#trainneg
cnf(i_0_3145, plain, (v1_relat_1(k3_xboole_0(k2_zfmisc_1(X1,X2),X3)))).#trainneg
cnf(i_0_3581, negated_conjecture, (v6_ordinal1(esk3_0)|~v1_xboole_0(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_1607, plain, (v7_ordinal1(k1_seq_1(X1,X2))|~v6_membered(k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3399, plain, (r1_xxreal_0(esk26_0,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,esk26_0))).#trainneg
cnf(i_0_3610, negated_conjecture, (v3_membered(esk4_0)|~m1_subset_1(esk1_0,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_3611, negated_conjecture, (v6_membered(esk3_0)|~v1_xboole_0(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_6768, plain, (v3_ordinal1(X1)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X1,X2)|~v4_prob_1(X2,X3)|~v1_prob_1(X2,X3)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(X3))))).#trainneg
cnf(i_0_1608, plain, (v3_ordinal1(k1_seq_1(X1,X2))|~v3_ordinal1(k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1616, plain, (v1_rat_1(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_7438, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k3_xcmplx_0(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))|~v1_xcmplx_0(k1_seq_1(esk3_0,esk4_0))|~v1_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))).#trainneg
cnf(i_0_7442, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_7475, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k5_numbers))).#trainneg
cnf(i_0_7476, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xboole_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0)))).#trainneg
cnf(i_0_7478, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_funct_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v4_valued_0(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0))))).#trainneg
cnf(i_0_4536, plain, (v1_relat_1(k3_xboole_0(X1,k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_6868, negated_conjecture, (k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_1617, plain, (v1_int_1(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_3646, negated_conjecture, (v1_membered(esk3_0)|~v3_membered(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_1618, plain, (v1_xcmplx_0(k8_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1619, plain, (v1_xxreal_0(k8_real_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7518, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0)))).#trainneg
cnf(i_0_7522, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_7527, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_7528, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_7561, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v4_valued_0(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0))).#trainneg
cnf(i_0_7445, plain, (v5_relat_1(esk29_1(k2_zfmisc_1(X1,k5_numbers)),k4_numbers))).#trainneg
cnf(i_0_4909, negated_conjecture, (v4_valued_0(esk3_0)|~v1_xboole_0(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_6870, negated_conjecture, (k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0),k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)),k1_numbers)|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_1620, plain, (v1_xreal_0(k8_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3712, negated_conjecture, (v2_membered(esk3_0)|~v3_membered(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_3734, negated_conjecture, (v3_membered(esk3_0)|~v6_membered(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_1621, plain, (v1_rat_1(k8_real_1(X1,X2))|~v4_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1622, plain, (v1_int_1(k8_real_1(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7618, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_7619, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_7446, plain, (v5_relat_1(esk29_1(k2_zfmisc_1(X1,k5_numbers)),k3_numbers))).#trainneg
cnf(i_0_5206, negated_conjecture, (v1_xreal_0(esk6_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1623, plain, (v7_ordinal1(k8_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3779, negated_conjecture, (v4_membered(esk3_0)|~v5_membered(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_4716, negated_conjecture, (v1_valued_0(esk3_0)|~v3_valued_0(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_1624, plain, (v3_ordinal1(k8_real_1(X1,X2))|~v3_ordinal1(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1630, plain, (m1_subset_1(k1_seq_1(X3,k5_prob_1(X1,X2,X5,X4)),k1_numbers)|r1_xxreal_0(k1_seq_1(X3,X4),k1_xboole_0)|v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xreal_0(k1_seq_1(X3,X4))|~m1_subset_1(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5),k1_numbers)|~m1_prob_1(X5,X1,X2)|~m1_prob_1(X4,X1,X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_7634, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xboole_0(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0))))).#trainneg
cnf(i_0_7635, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v5_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k5_numbers))).#trainneg
cnf(i_0_7689, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v5_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k4_numbers))).#trainneg
cnf(i_0_7690, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk6_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v5_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k3_numbers))).#trainneg
cnf(i_0_7738, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xboole_0(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0))))).#trainneg
cnf(i_0_7744, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v5_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k5_numbers)|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_7721, negated_conjecture, (v1_valued_0(esk3_0))).#trainneg
cnf(i_0_1633, plain, (m1_subset_1(k8_real_1(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4778, negated_conjecture, (v2_valued_0(esk3_0)|~v3_valued_0(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_7813, negated_conjecture, (v2_valued_0(esk3_0))).#trainneg
cnf(i_0_2710, plain, (v1_xxreal_0(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v6_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1631, plain, (m1_subset_1(k1_seq_1(X1,k5_prob_1(X3,X4,X5,X2)),k1_numbers)|r1_xxreal_0(k1_seq_1(X1,X2),k1_xboole_0)|v1_xboole_0(X4)|v1_xboole_0(X3)|~v1_xreal_0(k1_seq_1(k3_prob_2(X3,X4,X1,X2),X5))|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~m1_prob_1(X5,X3,X4)|~m1_prob_1(X2,X3,X4)|~m2_prob_1(X1,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3)))|~v4_prob_1(X4,X3)|~v1_prob_1(X4,X3))).#trainneg
cnf(i_0_7750, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v5_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k4_numbers)|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_7756, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v5_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k3_numbers)|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_1523, negated_conjecture, (k3_xcmplx_0(np__1,esk5_0)=esk5_0|~v1_membered(k1_zfmisc_1(esk1_0))|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(k1_zfmisc_1(esk1_0),X1)|~v1_prob_1(k1_zfmisc_1(esk1_0),X1))).#trainneg
cnf(i_0_1524, negated_conjecture, (k3_xcmplx_0(np__1,esk4_0)=esk4_0|~v1_membered(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(esk2_0,X1)|~v1_prob_1(esk2_0,X1))).#trainneg
cnf(i_0_1525, negated_conjecture, (k3_xcmplx_0(np__1,esk4_0)=esk4_0|~v1_membered(k1_zfmisc_1(esk1_0))|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(k1_zfmisc_1(esk1_0),X1)|~v1_prob_1(k1_zfmisc_1(esk1_0),X1))).#trainneg
cnf(i_0_1526, negated_conjecture, (k3_xcmplx_0(np__1,esk3_0)=esk3_0|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),X1)|~v1_prob_1(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers)),X1))).#trainneg
cnf(i_0_1527, negated_conjecture, (k3_xcmplx_0(np__1,esk2_0)=esk2_0|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X1)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X1))).#trainneg
cnf(i_0_1838, negated_conjecture, (v6_membered(esk2_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_5308, negated_conjecture, (v1_xreal_0(esk5_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_7935, plain, (m1_subset_1(esk30_0,k1_numbers))).#trainneg
cnf(i_0_7981, plain, (v1_rat_1(esk30_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_7982, plain, (v1_int_1(esk30_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_7839, plain, (v1_xxreal_0(X3)|v1_xboole_0(X2)|~v6_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_2711, plain, (v1_xreal_0(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v6_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1649, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,k1_numbers))|v1_xboole_0(X3)|v1_xboole_0(X2)|~r2_hidden(X1,k3_prob_2(X3,X2,X4,X5))|~m2_prob_1(X4,X3,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X5,X2)|~v4_prob_1(X2,X3)|~v1_prob_1(X2,X3))).#trainneg
cnf(i_0_1532, negated_conjecture, (k3_xcmplx_0(np__1,k3_xboole_0(X1,X2))=k3_xboole_0(X1,X2)|~v1_membered(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(X3)))|~v4_prob_1(esk2_0,X3)|~v1_prob_1(esk2_0,X3)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_1533, plain, (k3_xcmplx_0(np__1,k5_numbers)=k5_numbers|~v1_membered(k1_zfmisc_1(k1_numbers))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(k1_zfmisc_1(k1_numbers),X1)|~v1_prob_1(k1_zfmisc_1(k1_numbers),X1))).#trainneg
cnf(i_0_1537, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_membered(k1_zfmisc_1(X1))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(k1_zfmisc_1(X1),X2)|~v1_prob_1(k1_zfmisc_1(X1),X2))).#trainneg
cnf(i_0_1548, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_numbers))).#trainneg
cnf(i_0_1549, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_numbers))).#trainneg
cnf(i_0_1550, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_numbers))).#trainneg
cnf(i_0_1551, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_numbers))).#trainneg
cnf(i_0_1552, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers))).#trainneg
cnf(i_0_1553, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers))).#trainneg
cnf(i_0_1554, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers))).#trainneg
cnf(i_0_7937, plain, (m1_subset_1(esk25_0,k1_numbers))).#trainneg
cnf(i_0_7983, plain, (v7_ordinal1(esk30_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_6867, negated_conjecture, (k8_real_1(k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_1651, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X3)|~r2_hidden(X1,k5_prob_1(X2,X3,X4,X5))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_8100, plain, (m1_subset_1(esk37_0,k1_numbers))).#trainneg
cnf(i_0_7984, plain, (v3_ordinal1(esk30_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_8150, plain, (v1_rat_1(esk25_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_8061, plain, (v1_xreal_0(X3)|v1_xboole_0(X2)|~v6_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_8199, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X3)|~r2_hidden(X1,X4)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_1663, plain, (v2_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|v2_xxreal_0(k1_seq_1(X3,X4))|r1_xxreal_0(k1_seq_1(X3,X4),k1_xboole_0)|v1_xboole_0(X2)|v1_xboole_0(X1)|~v3_xxreal_0(k1_seq_1(X3,k5_prob_1(X1,X2,X5,X4)))|~v1_xreal_0(k1_seq_1(X3,X4))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|~m1_prob_1(X5,X1,X2)|~m1_prob_1(X4,X1,X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_8140, plain, (m1_subset_1(k5_prob_1(X1,X2,X3,X4),k1_numbers)|v1_xboole_0(X2)|~v3_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_8141, plain, (m1_subset_1(k5_prob_1(X1,X2,X3,X4),k1_numbers)|v1_xboole_0(X2)|~v6_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_8329, plain, (m1_subset_1(k3_xboole_0(X3,X4),k1_numbers)|v1_xboole_0(X2)|~v3_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_8374, plain, (m1_subset_1(k3_xboole_0(X3,X4),k1_numbers)|v1_xboole_0(X2)|~v6_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_8144, plain, (m1_subset_1(X1,k1_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_8102, plain, (m1_subset_1(esk32_0,k1_numbers))).#trainneg
cnf(i_0_8151, plain, (v1_int_1(esk25_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_8152, plain, (v7_ordinal1(esk25_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_8153, plain, (v3_ordinal1(esk25_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_8276, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X3)|v1_xboole_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2)|~m1_subset_1(X1,X4))).#trainneg
cnf(i_0_2712, plain, (v3_ordinal1(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v6_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1555, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k3_xcmplx_0(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers))).#trainneg
cnf(i_0_1576, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0))|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_1577, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_numbers)|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_1581, negated_conjecture, (v1_xxreal_0(X2)|~v2_membered(k1_zfmisc_1(X1))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~v4_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X1)|~v1_prob_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),X1))).#trainneg
cnf(i_0_1601, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk29_1(X1)))).#trainneg
cnf(i_0_1609, plain, (k3_xcmplx_0(np__1,k1_seq_1(X1,X2))=k1_seq_1(X1,X2)|~v1_membered(k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k1_zfmisc_1(X3)))|~v4_prob_1(k1_numbers,X3)|~v1_prob_1(k1_numbers,X3)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1611, plain, (k3_xcmplx_0(X1,k1_seq_1(X2,X3))=k3_xcmplx_0(k1_seq_1(X2,X3),X1)|~v1_xcmplx_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k1_zfmisc_1(X4)))|~v4_prob_1(k1_numbers,X4)|~v1_prob_1(k1_numbers,X4)|~v1_funct_1(X2)|~v3_valued_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1612, plain, (k3_xcmplx_0(k3_xboole_0(X1,k1_seq_1(X2,X3)),k1_xboole_0)=k1_xboole_0|~v1_membered(k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k1_zfmisc_1(X4)))|~m1_subset_1(X1,k1_numbers)|~v4_prob_1(k1_numbers,X4)|~v1_prob_1(k1_numbers,X4)|~v1_funct_1(X2)|~v3_valued_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1613, plain, (k3_xcmplx_0(np__1,k3_xboole_0(X1,k1_seq_1(X2,X3)))=k3_xboole_0(X1,k1_seq_1(X2,X3))|~v1_membered(k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k1_zfmisc_1(X4)))|~m1_subset_1(X1,k1_numbers)|~v4_prob_1(k1_numbers,X4)|~v1_prob_1(k1_numbers,X4)|~v1_funct_1(X2)|~v3_valued_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1625, plain, (k3_xcmplx_0(np__1,k8_real_1(X1,X2))=k8_real_1(X1,X2)|~v1_membered(k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k1_zfmisc_1(X3)))|~v4_prob_1(k1_numbers,X3)|~v1_prob_1(k1_numbers,X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1627, plain, (k3_xcmplx_0(X1,k8_real_1(X2,X3))=k3_xcmplx_0(k8_real_1(X2,X3),X1)|~v1_xcmplx_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k1_zfmisc_1(X4)))|~v4_prob_1(k1_numbers,X4)|~v1_prob_1(k1_numbers,X4)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_8143, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_8107, plain, (m1_subset_1(esk18_0,k1_numbers))).#trainneg
cnf(i_0_8183, plain, (v3_ordinal1(esk30_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_8202, plain, (v1_xcmplx_0(esk37_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_8520, plain, (v1_xcmplx_0(esk37_0))).#trainneg
cnf(i_0_8451, plain, (v3_ordinal1(X3)|v1_xboole_0(X2)|~v6_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1669, plain, (v2_xxreal_0(X1)|v2_xxreal_0(X2)|v2_xxreal_0(k3_xcmplx_0(X1,X2))|v1_xboole_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xxreal_0(k3_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_8529, negated_conjecture, (k3_xcmplx_0(esk37_0,esk6_0)=k3_xcmplx_0(esk6_0,esk37_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_8647, negated_conjecture, (k3_xcmplx_0(esk37_0,esk6_0)=k3_xcmplx_0(esk6_0,esk37_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_8650, negated_conjecture, (k3_xcmplx_0(esk37_0,esk6_0)=k3_xcmplx_0(esk6_0,esk37_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1628, plain, (k3_xcmplx_0(k3_xboole_0(X1,k8_real_1(X2,X3)),k1_xboole_0)=k1_xboole_0|~v1_membered(k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k1_zfmisc_1(X4)))|~m1_subset_1(X1,k1_numbers)|~v4_prob_1(k1_numbers,X4)|~v1_prob_1(k1_numbers,X4)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1629, plain, (k3_xcmplx_0(np__1,k3_xboole_0(X1,k8_real_1(X2,X3)))=k3_xboole_0(X1,k8_real_1(X2,X3))|~v1_membered(k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k1_zfmisc_1(X4)))|~m1_subset_1(X1,k1_numbers)|~v4_prob_1(k1_numbers,X4)|~v1_prob_1(k1_numbers,X4)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1642, negated_conjecture, (m1_subset_1(X1,esk1_0)|~r2_hidden(X1,esk6_0))).#trainneg
cnf(i_0_1643, negated_conjecture, (m1_subset_1(X1,esk1_0)|~r2_hidden(X1,esk5_0))).#trainneg
cnf(i_0_1644, negated_conjecture, (m1_subset_1(X1,esk1_0)|~r2_hidden(X1,esk4_0))).#trainneg
cnf(i_0_1645, negated_conjecture, (m1_subset_1(X1,k2_zfmisc_1(esk2_0,k1_numbers))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_1650, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k5_prob_1(X3,k1_zfmisc_1(X2),X4,X5))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X5,k1_zfmisc_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X3)|~v1_prob_1(k1_zfmisc_1(X2),X3))).#trainneg
cnf(i_0_1670, plain, (v2_xxreal_0(X1)|v2_xxreal_0(X2)|v2_xxreal_0(k3_xcmplx_0(X1,X2))|v1_xboole_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(k3_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_8145, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_8480, plain, (m1_subset_1(esk8_1(k5_numbers),k1_numbers))).#trainneg
cnf(i_0_8205, plain, (v1_rat_1(esk37_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_3415, plain, (r1_xxreal_0(esk25_0,esk19_0)|r1_xxreal_0(esk19_0,esk25_0))).#trainneg
cnf(i_0_8206, plain, (v1_int_1(esk37_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_8402, plain, (v1_xcmplx_0(esk32_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1681, plain, (v3_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|v2_xxreal_0(k1_seq_1(X3,X4))|r1_xxreal_0(k1_seq_1(X3,X4),k1_xboole_0)|v1_xboole_0(X2)|v1_xboole_0(X1)|~v2_xxreal_0(k1_seq_1(X3,k5_prob_1(X1,X2,X5,X4)))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|~v1_xreal_0(k1_seq_1(X3,X4))|~m1_prob_1(X5,X1,X2)|~m1_prob_1(X4,X1,X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1682, plain, (v3_xxreal_0(k1_seq_1(X1,X2))|v2_xxreal_0(k1_seq_1(k3_prob_2(X3,X4,X1,X2),X5))|r1_xxreal_0(k1_seq_1(X1,X2),k1_xboole_0)|v1_xboole_0(X3)|v1_xboole_0(X4)|~v2_xxreal_0(k1_seq_1(X1,k5_prob_1(X3,X4,X5,X2)))|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(k3_prob_2(X3,X4,X1,X2),X5))|~m1_prob_1(X5,X3,X4)|~m1_prob_1(X2,X3,X4)|~m2_prob_1(X1,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v4_prob_1(X4,X3)|~v1_prob_1(X4,X3))).#trainneg
cnf(i_0_1713, plain, (v3_xxreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|v3_xxreal_0(k1_seq_1(X3,X4))|r1_xxreal_0(k1_seq_1(X3,X4),k1_xboole_0)|v1_xboole_0(X2)|v1_xboole_0(X1)|~v3_xxreal_0(k1_seq_1(X3,k5_prob_1(X1,X2,X5,X4)))|~v1_xreal_0(k1_seq_1(X3,X4))|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|~m1_prob_1(X5,X1,X2)|~m1_prob_1(X4,X1,X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1719, plain, (v3_xxreal_0(X1)|v3_xxreal_0(X2)|v2_xxreal_0(k3_xcmplx_0(X1,X2))|v1_xboole_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xxreal_0(k3_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_8795, plain, (v1_xcmplx_0(esk32_0))).#trainneg
cnf(i_0_8757, plain, (v1_xreal_0(esk8_1(k5_numbers)))).#trainneg
cnf(i_0_8405, plain, (v1_rat_1(esk32_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_1720, plain, (v3_xxreal_0(X1)|v3_xxreal_0(X2)|v2_xxreal_0(k3_xcmplx_0(X1,X2))|v1_xboole_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(k3_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_3419, plain, (r1_xxreal_0(esk25_0,np__1)|r1_xxreal_0(np__1,esk25_0))).#trainneg
cnf(i_0_8406, plain, (v1_int_1(esk32_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_8439, plain, (v3_ordinal1(esk25_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_3560, plain, (v3_membered(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v5_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1725, plain, (r1_xxreal_0(esk29_1(X1),X2)|r1_xxreal_0(X2,esk29_1(X1))|~v1_xxreal_0(X2)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1726, plain, (r1_xxreal_0(esk29_1(X1),X2)|r1_xxreal_0(X2,esk29_1(X1))|~v1_xboole_0(X2)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_8916, negated_conjecture, (k3_xcmplx_0(esk32_0,esk6_0)=k3_xcmplx_0(esk6_0,esk32_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_9018, negated_conjecture, (k3_xcmplx_0(esk32_0,esk6_0)=k3_xcmplx_0(esk6_0,esk32_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_9021, negated_conjecture, (k3_xcmplx_0(esk32_0,esk6_0)=k3_xcmplx_0(esk6_0,esk32_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1722, negated_conjecture, (k3_xcmplx_0(X1,esk5_0)=k3_xcmplx_0(esk5_0,X1)|~v1_xcmplx_0(X1)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_9040, negated_conjecture, (k3_xcmplx_0(esk37_0,esk5_0)=k3_xcmplx_0(esk5_0,esk37_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_9042, negated_conjecture, (k3_xcmplx_0(esk32_0,esk5_0)=k3_xcmplx_0(esk5_0,esk32_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_9052, negated_conjecture, (k3_xcmplx_0(esk5_0,esk4_0)=k3_xcmplx_0(esk4_0,esk5_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_9067, negated_conjecture, (k3_xcmplx_0(esk37_0,esk5_0)=k3_xcmplx_0(esk5_0,esk37_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_9070, negated_conjecture, (k3_xcmplx_0(esk37_0,esk5_0)=k3_xcmplx_0(esk5_0,esk37_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_8485, plain, (v1_xcmplx_0(esk18_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_9099, plain, (v1_xcmplx_0(esk18_0))).#trainneg
cnf(i_0_8488, plain, (v1_rat_1(esk18_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_6964, negated_conjecture, (k3_xcmplx_0(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k8_real_1(k1_seq_1(esk3_0,esk4_0),k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0)),k1_numbers))).#trainneg
cnf(i_0_1733, plain, (r1_xxreal_0(esk15_1(X1),X2)|r1_xxreal_0(X2,esk15_1(X1))|v1_xboole_0(X1)|~v1_xxreal_0(X2)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3421, plain, (r1_xxreal_0(esk25_0,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,esk25_0))).#trainneg
cnf(i_0_8486, plain, (v1_xxreal_0(esk18_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_9123, plain, (v1_xxreal_0(esk18_0))).#trainneg
cnf(i_0_9008, plain, (v3_membered(X3)|v1_xboole_0(X2)|~v5_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1734, plain, (r1_xxreal_0(esk15_1(X1),X2)|r1_xxreal_0(X2,esk15_1(X1))|v1_xboole_0(X1)|~v1_xboole_0(X2)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1735, plain, (r1_xxreal_0(esk8_1(X1),X2)|r1_xxreal_0(X2,esk8_1(X1))|~v1_xxreal_0(X2)|~v3_membered(X1))).#trainneg
cnf(i_0_9072, negated_conjecture, (k3_xcmplx_0(esk32_0,esk5_0)=k3_xcmplx_0(esk5_0,esk32_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_9075, negated_conjecture, (k3_xcmplx_0(esk32_0,esk5_0)=k3_xcmplx_0(esk5_0,esk32_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_9077, negated_conjecture, (k3_xcmplx_0(esk5_0,esk4_0)=k3_xcmplx_0(esk4_0,esk5_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_9080, negated_conjecture, (k3_xcmplx_0(esk5_0,esk4_0)=k3_xcmplx_0(esk4_0,esk5_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_9108, negated_conjecture, (k3_xcmplx_0(esk18_0,esk6_0)=k3_xcmplx_0(esk6_0,esk18_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_9109, negated_conjecture, (k3_xcmplx_0(esk18_0,esk5_0)=k3_xcmplx_0(esk5_0,esk18_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_9122, plain, (r1_xxreal_0(k1_xboole_0,esk25_0))).#trainneg
cnf(i_0_9267, negated_conjecture, (k3_xcmplx_0(esk18_0,esk6_0)=k3_xcmplx_0(esk6_0,esk18_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_9270, negated_conjecture, (k3_xcmplx_0(esk18_0,esk6_0)=k3_xcmplx_0(esk6_0,esk18_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_9272, negated_conjecture, (k3_xcmplx_0(esk18_0,esk5_0)=k3_xcmplx_0(esk5_0,esk18_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_9275, negated_conjecture, (k3_xcmplx_0(esk18_0,esk5_0)=k3_xcmplx_0(esk5_0,esk18_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_9276, plain, (~v3_xxreal_0(esk25_0))).#trainneg
cnf(i_0_9136, plain, (r1_xxreal_0(esk18_0,esk18_0))).#trainneg
cnf(i_0_8388, plain, (m1_subset_1(esk29_1(X1),k1_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_8489, plain, (v1_int_1(esk18_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1736, plain, (r1_xxreal_0(esk8_1(X1),X2)|r1_xxreal_0(X2,esk8_1(X1))|~v1_xboole_0(X2)|~v3_membered(X1))).#trainneg
cnf(i_0_8490, plain, (v7_ordinal1(esk18_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_8491, plain, (v3_ordinal1(esk18_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_1737, plain, (v5_relat_1(esk34_1(X1),k4_numbers)|v1_xboole_0(X1)|~v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1738, plain, (v5_relat_1(esk29_1(X1),k4_numbers)|~v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1723, negated_conjecture, (k3_xcmplx_0(esk5_0,k1_xboole_0)=k1_xboole_0|~v3_membered(esk2_0))).#trainneg
cnf(i_0_1724, negated_conjecture, (k3_xcmplx_0(np__1,esk5_0)=esk5_0|~v3_membered(esk2_0))).#trainneg
cnf(i_0_1729, negated_conjecture, (k3_xcmplx_0(X1,esk4_0)=k3_xcmplx_0(esk4_0,X1)|~v1_xcmplx_0(X1)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_9423, negated_conjecture, (k3_xcmplx_0(esk37_0,esk4_0)=k3_xcmplx_0(esk4_0,esk37_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_9425, negated_conjecture, (k3_xcmplx_0(esk32_0,esk4_0)=k3_xcmplx_0(esk4_0,esk32_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_9430, negated_conjecture, (k3_xcmplx_0(esk18_0,esk4_0)=k3_xcmplx_0(esk4_0,esk18_0)|~v3_membered(esk2_0))).#trainneg
cnf(i_0_9451, negated_conjecture, (k3_xcmplx_0(esk37_0,esk4_0)=k3_xcmplx_0(esk4_0,esk37_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_9454, negated_conjecture, (k3_xcmplx_0(esk37_0,esk4_0)=k3_xcmplx_0(esk4_0,esk37_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_9409, plain, (v3_ordinal1(esk18_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_9325, plain, (v1_xreal_0(esk29_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_9319, plain, (v3_xxreal_0(esk18_0)|v2_xxreal_0(esk18_0)|v1_xboole_0(esk18_0))).#trainneg
cnf(i_0_6971, negated_conjecture, (k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_1739, plain, (v5_relat_1(esk22_1(X1),k4_numbers)|~v5_relat_1(k1_zfmisc_1(X1),k4_numbers)|~v1_relat_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5311, negated_conjecture, (v1_xreal_0(esk4_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_5514, negated_conjecture, (v1_xxreal_0(esk6_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1740, plain, (v5_relat_1(esk15_1(X1),k4_numbers)|v1_xboole_0(X1)|~v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1741, plain, (v5_relat_1(esk13_1(X1),k4_numbers)|~v5_relat_1(k1_zfmisc_1(X1),k4_numbers)|~v1_relat_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_9456, negated_conjecture, (k3_xcmplx_0(esk32_0,esk4_0)=k3_xcmplx_0(esk4_0,esk32_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_9459, negated_conjecture, (k3_xcmplx_0(esk32_0,esk4_0)=k3_xcmplx_0(esk4_0,esk32_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_9461, negated_conjecture, (k3_xcmplx_0(esk18_0,esk4_0)=k3_xcmplx_0(esk4_0,esk18_0)|~v6_membered(esk2_0))).#trainneg
cnf(i_0_9464, negated_conjecture, (k3_xcmplx_0(esk18_0,esk4_0)=k3_xcmplx_0(esk4_0,esk18_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1730, negated_conjecture, (k3_xcmplx_0(esk4_0,k1_xboole_0)=k1_xboole_0|~v3_membered(esk2_0))).#trainneg
cnf(i_0_1731, negated_conjecture, (k3_xcmplx_0(np__1,esk4_0)=esk4_0|~v3_membered(esk2_0))).#trainneg
cnf(i_0_1742, plain, (v5_relat_1(esk8_1(k1_zfmisc_1(X1)),k4_numbers)|~v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1743, negated_conjecture, (v5_relat_1(esk6_0,k4_numbers)|~v5_relat_1(esk1_0,k4_numbers)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_1744, negated_conjecture, (v5_relat_1(esk5_0,k4_numbers)|~v5_relat_1(esk1_0,k4_numbers)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_1745, negated_conjecture, (v5_relat_1(esk4_0,k4_numbers)|~v5_relat_1(esk1_0,k4_numbers)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_1746, negated_conjecture, (v5_relat_1(esk3_0,k4_numbers)|~v5_relat_1(k2_zfmisc_1(esk2_0,k1_numbers),k4_numbers)|~v1_relat_1(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_1747, negated_conjecture, (v5_relat_1(esk2_0,k4_numbers)|~v5_relat_1(k1_zfmisc_1(esk1_0),k4_numbers)|~v1_relat_1(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_8134, plain, (m1_subset_1(k5_numbers,k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_6973, negated_conjecture, (k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)),k1_numbers)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_1750, plain, (v5_relat_1(k3_prob_2(X1,X2,X3,X4),k4_numbers)|v1_xboole_0(X1)|v1_xboole_0(X2)|~v5_relat_1(k2_zfmisc_1(X2,k1_numbers),k4_numbers)|~v1_relat_1(k2_zfmisc_1(X2,k1_numbers))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_5665, negated_conjecture, (v1_xxreal_0(esk5_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_5807, negated_conjecture, (v1_xxreal_0(esk4_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1752, plain, (v5_relat_1(k5_prob_1(X1,X2,X3,X4),k4_numbers)|v1_xboole_0(X2)|~v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1754, plain, (v5_relat_1(esk34_1(X1),k3_numbers)|v1_xboole_0(X1)|~v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1751, plain, (v5_relat_1(k5_prob_1(X1,k1_zfmisc_1(X2),X3,X4),k4_numbers)|~v5_relat_1(X2,k4_numbers)|~v1_relat_1(X2)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k1_zfmisc_1(X2),X1))).#trainneg
cnf(i_0_1755, plain, (v5_relat_1(esk29_1(X1),k3_numbers)|~v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1756, plain, (v5_relat_1(esk22_1(X1),k3_numbers)|~v5_relat_1(k1_zfmisc_1(X1),k3_numbers)|~v1_relat_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1757, plain, (v5_relat_1(esk15_1(X1),k3_numbers)|v1_xboole_0(X1)|~v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1758, plain, (v5_relat_1(esk13_1(X1),k3_numbers)|~v5_relat_1(k1_zfmisc_1(X1),k3_numbers)|~v1_relat_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1759, plain, (v5_relat_1(esk8_1(k1_zfmisc_1(X1)),k3_numbers)|~v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1760, negated_conjecture, (v5_relat_1(esk6_0,k3_numbers)|~v5_relat_1(esk1_0,k3_numbers)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_1761, negated_conjecture, (v5_relat_1(esk5_0,k3_numbers)|~v5_relat_1(esk1_0,k3_numbers)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_1762, negated_conjecture, (v5_relat_1(esk4_0,k3_numbers)|~v5_relat_1(esk1_0,k3_numbers)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_1763, negated_conjecture, (v5_relat_1(esk3_0,k3_numbers)|~v5_relat_1(k2_zfmisc_1(esk2_0,k1_numbers),k3_numbers)|~v1_relat_1(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_1764, negated_conjecture, (v5_relat_1(esk2_0,k3_numbers)|~v5_relat_1(k1_zfmisc_1(esk1_0),k3_numbers)|~v1_relat_1(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1767, plain, (v5_relat_1(k3_prob_2(X1,X2,X3,X4),k3_numbers)|v1_xboole_0(X1)|v1_xboole_0(X2)|~v5_relat_1(k2_zfmisc_1(X2,k1_numbers),k3_numbers)|~v1_relat_1(k2_zfmisc_1(X2,k1_numbers))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1839, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8389, plain, (m1_subset_1(esk8_1(k1_zfmisc_1(X1)),k1_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1769, plain, (v5_relat_1(k5_prob_1(X1,X2,X3,X4),k3_numbers)|v1_xboole_0(X2)|~v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1781, plain, (v5_relat_1(k3_xboole_0(X1,X2),k4_numbers)|~v1_relat_1(k3_xboole_0(X1,X2))|~v5_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1782, plain, (v5_relat_1(k3_xboole_0(X1,X2),k3_numbers)|~v1_relat_1(k3_xboole_0(X1,X2))|~v5_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1768, plain, (v5_relat_1(k5_prob_1(X1,k1_zfmisc_1(X2),X3,X4),k3_numbers)|~v5_relat_1(X2,k3_numbers)|~v1_relat_1(X2)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v4_prob_1(k1_zfmisc_1(X2),X1)|~v1_prob_1(k1_zfmisc_1(X2),X1))).#trainneg
cnf(i_0_1785, plain, (v5_relat_1(k3_xboole_0(X2,X1),X3)|~v5_relat_1(X1,X3)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1790, plain, (v4_relat_1(k3_xboole_0(X2,X1),X3)|~v4_relat_1(X1,X3)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1793, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,k1_xboole_0))=k1_xboole_0|~v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1794, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,k5_prob_1(X3,X4,X5,X6)))=k3_xcmplx_0(k5_prob_1(X3,X4,X5,X6),k3_xcmplx_0(X1,X2))|v1_xboole_0(X4)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_membered(X4)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X6,X4)|~m1_subset_1(X5,X4)|~v4_prob_1(X4,X3)|~v1_prob_1(X4,X3)|~v1_xcmplx_0(k5_prob_1(X3,X4,X5,X6))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9831, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(esk30_0,k5_prob_1(X1,X2,X3,X4)))=k3_xcmplx_0(k5_prob_1(X1,X2,X3,X4),esk30_0)|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_9833, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(esk25_0,k5_prob_1(X1,X2,X3,X4)))=k3_xcmplx_0(k5_prob_1(X1,X2,X3,X4),esk25_0)|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_9836, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(np__1,k5_prob_1(X1,X2,X3,X4)))=k3_xcmplx_0(k5_prob_1(X1,X2,X3,X4),np__1)|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_9838, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(k1_xboole_0,k5_prob_1(X1,X2,X3,X4)))=k3_xcmplx_0(k5_prob_1(X1,X2,X3,X4),k1_xboole_0)|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_9842, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(X1,k5_prob_1(X2,X3,X4,X5)))=k3_xcmplx_0(k5_prob_1(X2,X3,X4,X5),X1)|v1_xboole_0(X3)|~v1_xcmplx_0(X1)|~v1_membered(X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1841, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_9628, plain, (v1_xreal_0(esk8_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1795, plain, (v1_xreal_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)))|~v1_xreal_0(X3)|~v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1796, plain, (v2_xxreal_0(k3_xcmplx_0(X1,X2))|v2_xxreal_0(X3)|~v3_xxreal_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)))|~v1_xreal_0(X3)|~v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1797, plain, (v3_xxreal_0(k3_xcmplx_0(X1,X2))|v2_xxreal_0(X3)|~v2_xxreal_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)))|~v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9845, plain, (k3_xcmplx_0(k1_xboole_0,k3_xcmplx_0(k1_xboole_0,k5_prob_1(X1,X2,X3,X4)))=k3_xcmplx_0(k5_prob_1(X1,X2,X3,X4),k1_xboole_0)|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1842, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_8394, plain, (m1_subset_1(k3_xboole_0(X1,X2),k1_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1798, plain, (v3_xxreal_0(k3_xcmplx_0(X1,X2))|v3_xxreal_0(X3)|~v3_xxreal_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)))|~v1_xreal_0(X3)|~v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3791, plain, (v3_membered(k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v4_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1799, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)),X4)=k3_xcmplx_0(k3_xcmplx_0(X1,X2),k3_xcmplx_0(X3,X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1800, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k3_xcmplx_0(X2,X3))|~v2_xxreal_0(k3_xcmplx_0(X2,k3_xcmplx_0(X3,X1)))|~v1_xreal_0(X1)|~v1_xreal_0(k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_10028, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(k5_prob_1(X2,X3,X4,X5),X1))=k3_xcmplx_0(k5_prob_1(X2,X3,X4,X5),X1)|v1_xboole_0(X3)|~v1_xreal_0(X1)|~v1_membered(X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10914, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(esk30_0,k5_prob_1(X1,X2,X3,X4)))=k3_xcmplx_0(esk30_0,k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_10915, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(esk25_0,k5_prob_1(X1,X2,X3,X4)))=k3_xcmplx_0(esk25_0,k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_10422, plain, (v1_xreal_0(k3_xboole_0(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1801, plain, (k3_xcmplx_0(esk30_0,k3_xcmplx_0(k1_xboole_0,X1))=k3_xcmplx_0(k1_xboole_0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3449, plain, (r1_xxreal_0(esk19_0,np__1)|r1_xxreal_0(np__1,esk19_0))).#trainneg
cnf(i_0_10592, plain, (v3_membered(X3)|v1_xboole_0(X2)|~v4_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1802, plain, (k3_xcmplx_0(esk25_0,k3_xcmplx_0(k1_xboole_0,X1))=k3_xcmplx_0(k1_xboole_0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1803, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(esk30_0,X1))=k3_xcmplx_0(esk30_0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10916, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(np__1,k5_prob_1(X1,X2,X3,X4)))=k3_xcmplx_0(np__1,k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_10919, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(k1_xboole_0,k5_prob_1(X1,X2,X3,X4)))=k3_xcmplx_0(k1_xboole_0,k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_10922, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(X5,k5_prob_1(X1,X2,X3,X4)))=k3_xcmplx_0(X5,k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v1_xreal_0(X5)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_xcmplx_0(X5))).#trainneg
cnf(i_0_2153, plain, (v1_rat_1(k1_xboole_0)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_10997, plain, (v1_xreal_0(k3_xboole_0(X2,X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7496, negated_conjecture, (v5_membered(esk3_0)|~v1_xboole_0(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_1804, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(esk25_0,X1))=k3_xcmplx_0(esk25_0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3451, plain, (r1_xxreal_0(esk19_0,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,esk19_0))).#trainneg
cnf(i_0_7577, negated_conjecture, (v1_membered(esk3_0)|~v6_membered(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_7665, negated_conjecture, (v2_membered(esk3_0)|~v6_membered(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_1805, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(np__1,X1))=k3_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1806, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(k1_xboole_0,X1))=k3_xcmplx_0(k1_xboole_0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2154, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7668, negated_conjecture, (v3_membered(esk3_0)|~v1_xboole_0(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_1807, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(k5_prob_1(X1,X2,X3,X4),X5))=k3_xcmplx_0(k5_prob_1(X1,X2,X3,X4),X5)|v1_xboole_0(X2)|~v1_xcmplx_0(X5)|~v1_xcmplx_0(k5_prob_1(X1,X2,X3,X4))|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_7717, negated_conjecture, (v4_membered(esk3_0)|~v6_membered(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_1808, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1809, plain, (k3_xcmplx_0(k1_xboole_0,k3_xcmplx_0(k1_xboole_0,X1))=k3_xcmplx_0(k1_xboole_0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2156, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_8397, plain, (m1_subset_1(k2_zfmisc_1(X1,X2),k1_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_11313, negated_conjecture, (v1_membered(esk3_0)|~v1_xboole_0(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_6869, negated_conjecture, (k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k1_numbers)|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_1810, plain, (k3_xcmplx_0(k1_xboole_0,k3_xcmplx_0(k5_prob_1(X1,X2,X3,X4),X5))=k3_xcmplx_0(k1_xboole_0,X5)|v1_xboole_0(X2)|~v1_xcmplx_0(X5)|~v1_xcmplx_0(k5_prob_1(X1,X2,X3,X4))|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_11316, negated_conjecture, (v2_membered(esk3_0)|~v1_xboole_0(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_11462, negated_conjecture, (v4_membered(esk3_0)|~v1_xboole_0(k2_zfmisc_1(esk2_0,k1_numbers)))).#trainneg
cnf(i_0_1811, plain, (k3_xcmplx_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5),k3_xcmplx_0(k1_seq_1(X3,X4),X6))=k3_xcmplx_0(k1_seq_1(X3,k5_prob_1(X1,X2,X5,X4)),X6)|r1_xxreal_0(k1_seq_1(X3,X4),k1_xboole_0)|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xcmplx_0(X6)|~v1_xcmplx_0(k1_seq_1(X3,X4))|~v1_xcmplx_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|~v1_xreal_0(k1_seq_1(X3,X4))|~m1_prob_1(X5,X1,X2)|~m1_prob_1(X4,X1,X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1812, plain, (k3_xcmplx_0(k1_seq_1(X1,X2),k3_xcmplx_0(k1_seq_1(k3_prob_2(X3,X4,X1,X2),X5),X6))=k3_xcmplx_0(k1_seq_1(X1,k5_prob_1(X3,X4,X5,X2)),X6)|r1_xxreal_0(k1_seq_1(X1,X2),k1_xboole_0)|v1_xboole_0(X3)|v1_xboole_0(X4)|~v1_xcmplx_0(X6)|~v1_xcmplx_0(k1_seq_1(k3_prob_2(X3,X4,X1,X2),X5))|~v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(k3_prob_2(X3,X4,X1,X2),X5))|~m1_prob_1(X5,X3,X4)|~m1_prob_1(X2,X3,X4)|~m2_prob_1(X1,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v4_prob_1(X4,X3)|~v1_prob_1(X4,X3))).#trainneg
cnf(i_0_11460, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(X5,k5_prob_1(X1,X2,X3,X4)))=k3_xcmplx_0(X5,k5_prob_1(X1,X2,X3,X4))|v1_xboole_0(X2)|~v1_xcmplx_0(X5)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_11581, plain, (k3_xcmplx_0(np__1,k1_seq_1(X3,k5_prob_1(X1,X2,X5,X4)))=k1_seq_1(X3,k5_prob_1(X1,X2,X5,X4))|r1_xxreal_0(k1_seq_1(X3,X4),k1_xboole_0)|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|~v1_xcmplx_0(k1_seq_1(X3,X4))|~v1_xreal_0(k1_seq_1(X3,X4))|~m1_prob_1(X5,X1,X2)|~m1_prob_1(X4,X1,X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_11582, plain, (k3_xcmplx_0(np__1,k1_seq_1(X3,k5_prob_1(X1,X2,X5,X4)))=k1_seq_1(X3,k5_prob_1(X1,X2,X5,X4))|r1_xxreal_0(k1_seq_1(X3,X4),k1_xboole_0)|v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xreal_0(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5))|~v1_xcmplx_0(k1_seq_1(X3,X4))|~m1_prob_1(X5,X1,X2)|~m1_prob_1(X4,X1,X2)|~m2_prob_1(X3,X1,X2)|~m1_subset_1(k1_seq_1(k3_prob_2(X1,X2,X3,X4),X5),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(k1_seq_1(X3,X4),k1_numbers)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_11584, plain, (k3_xcmplx_0(np__1,k1_seq_1(X1,k5_prob_1(X3,X4,X5,X2)))=k1_seq_1(X1,k5_prob_1(X3,X4,X5,X2))|r1_xxreal_0(k1_seq_1(X1,X2),k1_xboole_0)|v1_xboole_0(X3)|v1_xboole_0(X4)|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xcmplx_0(k1_seq_1(k3_prob_2(X3,X4,X1,X2),X5))|~m1_prob_1(X5,X3,X4)|~m1_prob_1(X2,X3,X4)|~m2_prob_1(X1,X3,X4)|~m1_subset_1(k1_seq_1(k3_prob_2(X3,X4,X1,X2),X5),k1_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v4_prob_1(X4,X3)|~v1_prob_1(X4,X3))).#trainneg
cnf(i_0_2157, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_11644, plain, (v1_xreal_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6871, negated_conjecture, (k8_real_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0),k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)))!=k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,k3_xboole_0(esk4_0,esk5_0),esk6_0))|~v1_xreal_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,k3_xboole_0(esk4_0,esk5_0)),esk6_0))|~m1_subset_1(k3_xcmplx_0(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_seq_1(esk3_0,esk4_0)),k1_numbers)|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),esk5_0),k1_numbers))).#trainneg
cnf(i_0_1813, plain, (k3_xcmplx_0(k5_prob_1(X1,X2,X3,X4),k3_xcmplx_0(k1_xboole_0,X5))=k3_xcmplx_0(k1_xboole_0,X5)|v1_xboole_0(X2)|~v1_xcmplx_0(X5)|~v1_xcmplx_0(k5_prob_1(X1,X2,X3,X4))|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_2508, negated_conjecture, (v3_membered(esk2_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2510, negated_conjecture, (v3_membered(esk2_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1814, plain, (k3_xcmplx_0(k3_xcmplx_0(X5,k5_prob_1(X1,X2,X3,X4)),X6)=k3_xcmplx_0(k5_prob_1(X1,X2,X3,X4),k3_xcmplx_0(X5,X6))|v1_xboole_0(X2)|~v1_xcmplx_0(X6)|~v1_xcmplx_0(X5)|~v1_xcmplx_0(k5_prob_1(X1,X2,X3,X4))|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1815, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(k1_xboole_0,X2))=k3_xcmplx_0(k1_xboole_0,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_11911, plain, (k3_xcmplx_0(np__1,k1_seq_1(X1,k5_prob_1(X2,X3,X5,X4)))=k1_seq_1(X1,k5_prob_1(X2,X3,X5,X4))|r1_xxreal_0(k1_seq_1(X1,X5),k1_xboole_0)|v1_xboole_0(X2)|v1_xboole_0(X3)|~v1_xreal_0(k1_seq_1(k3_prob_2(X2,X3,X1,X5),X4))|~v1_xreal_0(k1_seq_1(X1,X5))|~m1_prob_1(X4,X2,X3)|~m1_prob_1(X5,X2,X3)|~m2_prob_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_2507, plain, (v1_int_1(k1_xboole_0)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_2940, plain, (k3_xcmplx_0(esk30_0,np__1)=esk30_0)).#trainneg
cnf(i_0_1820, plain, (v1_relat_1(esk34_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2511, negated_conjecture, (v3_membered(esk2_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2513, negated_conjecture, (v4_membered(esk2_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1816, plain, (k3_xcmplx_0(k3_xcmplx_0(k5_prob_1(X2,X3,X4,X5),X1),X6)=k3_xcmplx_0(X1,k3_xcmplx_0(k5_prob_1(X2,X3,X4,X5),X6))|v1_xboole_0(X3)|~v1_xcmplx_0(X6)|~v1_xcmplx_0(k5_prob_1(X2,X3,X4,X5))|~v1_xcmplx_0(X1)|~v1_membered(X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2))).#trainneg
cnf(i_0_1822, plain, (v1_relat_1(esk15_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_11942, plain, (k3_xcmplx_0(np__1,k1_seq_1(X1,k5_prob_1(X2,X3,X4,X5)))=k1_seq_1(X1,k5_prob_1(X2,X3,X4,X5))|r1_xxreal_0(k1_seq_1(X1,X5),k1_xboole_0)|v1_xboole_0(X3)|v1_xboole_0(X2)|~v1_xreal_0(k1_seq_1(X1,X5))|~m1_prob_1(X4,X2,X3)|~m1_prob_1(X5,X2,X3)|~m2_prob_1(X1,X2,X3)|~m1_subset_1(k1_seq_1(k3_prob_2(X2,X3,X1,X5),X4),k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2)|~v1_xreal_0(k1_seq_1(k3_prob_2(X2,X3,X1,X5),X4)))).#trainneg
cnf(i_0_12125, plain, (k3_xcmplx_0(k5_prob_1(X1,X2,X3,X4),k3_xcmplx_0(np__1,X5))=k3_xcmplx_0(k5_prob_1(X1,X2,X3,X4),X5)|v1_xboole_0(X2)|~v1_xcmplx_0(X5)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_2556, plain, (v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_10459, plain, (m1_subset_1(k3_xboole_0(X2,X1),k1_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2942, plain, (k3_xcmplx_0(k1_xboole_0,esk30_0)=k1_xboole_0)).#trainneg
cnf(i_0_1828, plain, (v1_relat_1(k5_prob_1(k2_zfmisc_1(X1,X2),X3,X4,X5))|v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,k2_zfmisc_1(X1,X2))|~v1_prob_1(X3,k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1832, negated_conjecture, (v1_xreal_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_12232, plain, (k3_xcmplx_0(esk30_0,k3_xcmplx_0(np__1,k5_prob_1(X1,X2,X3,X4)))=k3_xcmplx_0(k5_prob_1(X1,X2,X3,X4),esk30_0)|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_12377, plain, (k3_xcmplx_0(np__1,k1_seq_1(X1,k5_prob_1(X2,X3,X5,X4)))=k1_seq_1(X1,k5_prob_1(X2,X3,X5,X4))|r1_xxreal_0(k1_seq_1(X1,X5),k1_xboole_0)|v1_xboole_0(X2)|v1_xboole_0(X3)|~v1_xreal_0(k1_seq_1(X1,X5))|~m1_prob_1(X4,X2,X3)|~m1_prob_1(X5,X2,X3)|~m2_prob_1(X1,X2,X3)|~m1_subset_1(k1_seq_1(k3_prob_2(X2,X3,X1,X5),X4),k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(X3,X2)|~v1_prob_1(X3,X2)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_2964, plain, (k3_xcmplx_0(esk25_0,np__1)=esk25_0)).#trainneg
cnf(i_0_1834, negated_conjecture, (v1_xreal_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_2536, negated_conjecture, (v1_xcmplx_0(esk3_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers))))).#trainneg
cnf(i_0_2570, negated_conjecture, (v1_xxreal_0(esk3_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_numbers))))).#trainneg
cnf(i_0_1835, negated_conjecture, (v1_xreal_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1843, plain, (r1_xxreal_0(k3_prob_2(X1,X2,X3,X4),X5)|r1_xxreal_0(X5,k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xxreal_0(X5)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_12590, plain, (k3_xcmplx_0(esk25_0,k3_xcmplx_0(np__1,k5_prob_1(X1,X2,X3,X4)))=k3_xcmplx_0(k5_prob_1(X1,X2,X3,X4),esk25_0)|v1_xboole_0(X2)|~v1_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_2613, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2966, plain, (k3_xcmplx_0(k1_xboole_0,esk25_0)=k1_xboole_0)).#trainneg
cnf(i_0_1844, plain, (r1_xxreal_0(k3_prob_2(X1,X2,X3,X4),X5)|r1_xxreal_0(X5,k3_prob_2(X1,X2,X3,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xboole_0(X5)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~m2_prob_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_2573, negated_conjecture, (v1_xreal_0(esk6_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2575, negated_conjecture, (v1_xreal_0(esk6_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_12650, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v5_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),k4_numbers))).#trainneg
cnf(i_0_12651, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v5_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),k3_numbers))).#trainneg
cnf(i_0_12652, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v5_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_12653, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v1_xboole_0(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0)))).#trainneg
cnf(i_0_12654, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_prob_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers)|~v5_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0),k5_numbers))).#trainneg
cnf(i_0_12704, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_prob_1(k3_xboole_0(esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k3_xboole_0(esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_12710, negated_conjecture, (v3_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_prob_1(k3_xboole_0(esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k3_xboole_0(esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_12712, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk3_0,esk4_0),k1_xboole_0)|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k3_xboole_0(esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_2667, negated_conjecture, (v3_membered(esk2_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1848, plain, (v1_xboole_0(esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3479, plain, (r1_xxreal_0(np__1,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,np__1))).#trainneg
cnf(i_0_2576, negated_conjecture, (v1_xreal_0(esk6_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2577, negated_conjecture, (v1_xreal_0(esk5_0)|~m1_subset_1(k1_zfmisc_1(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1853, plain, (v1_xboole_0(k5_prob_1(k2_zfmisc_1(X1,X2),X3,X4,X5))|v1_xboole_0(X3)|~v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~v4_prob_1(X3,k2_zfmisc_1(X1,X2))|~v1_prob_1(X3,k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_12715, negated_conjecture, (v3_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_prob_1(k3_xboole_0(esk4_0,esk5_0),esk1_0,esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k3_xboole_0(esk4_0,esk5_0),esk2_0))).#trainneg
cnf(i_0_12718, negated_conjecture, (v3_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_subset_1(k1_seq_1(esk3_0,k3_xboole_0(esk4_0,esk5_0)),k1_numbers)|~m1_subset_1(k3_xboole_0(esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_12807, negated_conjecture, (v3_xxreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~v1_relat_1(k3_prob_2(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_subset_1(k5_prob_1(esk1_0,esk2_0,esk4_0,esk5_0),esk2_0)|~m1_subset_1(k1_seq_1(esk3_0,esk4_0),k1_numbers))).#trainneg
cnf(i_0_12814, negated_conjecture, (~v1_xreal_0(k1_seq_1(esk3_0,esk4_0))|~m1_subset_1(k3_xboole_0(esk4_0,esk5_0),esk2_0))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 233
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 317
# Removed in clause preprocessing      : 31
# Initial clauses in saturation        : 286
# Processed clauses                    : 2746
# ...of these trivial                  : 41
# ...subsumed                          : 849
# ...remaining for further processing  : 1856
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 89
# Backward-rewritten                   : 40
# Generated clauses                    : 12511
# ...of the previous two non-trivial   : 11325
# Contextual simplify-reflections      : 550
# Paramodulations                      : 12511
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
# Propositional unsat check successes  : 0
# Current number of processed clauses  : 1727
#    Positive orientable unit clauses  : 186
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 36
#    Non-unit-clauses                  : 1504
# Current number of unprocessed clauses: 8354
# ...number of literals in the above   : 61318
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1985
# Proof object given clauses           : 113
# Proof search given clauses           : 1856
# Clause-clause subsumption calls (NU) : 1653601
# Rec. Clause-clause subsumption calls : 173915
# Non-unit clause-clause subsumptions  : 1403
# Unit Clause-clause subsumption calls : 40665
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 127
# BW rewrite match successes           : 26
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 352129

# -------------------------------------------------
# User time                : 3.117 s
# System time              : 0.100 s
# Total time               : 3.217 s
# ...preprocessing         : 0.295 s
# ...main loop             : 2.922 s
# Maximum resident set size: 36844 pages
