# ENIGMA: LightGBM model 'Enigma/mizar40-all-T10/mega-VHSLCAXPh2e15/lgb-t150-d30-l900-e0.15/model.lgb' loaded. (hash_base: 32768; conj_feats: 95; version: 991; iters: 150)
# Preprocessing time       : 0.307 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', t6_boole)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', redefinition_k6_numbers)).
fof(t32_lopban_4, axiom, ![X1]:((((((((((((((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v5_rlvect_1(X1))&v6_rlvect_1(X1))&v7_rlvect_1(X1))&v8_rlvect_1(X1))&v3_normsp_0(X1))&v4_normsp_0(X1))&v2_normsp_1(X1))&v2_funcsdom(X1))&v3_group_1(X1))&v1_vectsp_1(X1))&v3_vectsp_1(X1))&v5_lopban_2(X1))&l1_lopban_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(v1_xreal_0(X4)=>~((~(r1_xxreal_0(X4,k6_numbers))&![X5]:(m2_subset_1(X5,k1_numbers,k5_numbers)=>?[X6]:((m2_subset_1(X6,k1_numbers,k5_numbers)&r1_xxreal_0(X5,X6))&r1_xxreal_0(X4,k18_complex1(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,k8_lopban_4(X1,X2,X3,X6))),X6))))))))))), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', t32_lopban_4)).
fof(spc0_boole, axiom, v1_xboole_0(np__0), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', spc0_boole)).
fof(d7_seq_2, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>(v2_comseq_2(X1)=>![X2]:(v1_xreal_0(X2)=>(X2=k1_seq_2(X1)<=>![X3]:(v1_xreal_0(X3)=>~((~(r1_xxreal_0(X3,k6_numbers))&![X4]:(m2_subset_1(X4,k1_numbers,k5_numbers)=>?[X5]:((m2_subset_1(X5,k1_numbers,k5_numbers)&r1_xxreal_0(X4,X5))&r1_xxreal_0(X3,k18_complex1(k9_real_1(k1_seq_1(X1,X5),X2)))))))))))), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', d7_seq_2)).
fof(redefinition_k1_seq_1, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v3_valued_0(X1))=>k1_seq_1(X1,X2)=k1_funct_1(X1,X2)), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', redefinition_k1_seq_1)).
fof(rc4_xreal_0, axiom, ?[X1]:(((v1_xboole_0(X1)&v1_xcmplx_0(X1))&v1_xxreal_0(X1))&v1_xreal_0(X1)), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', rc4_xreal_0)).
fof(s1_seq_1__e1_46__lopban_4, axiom, ![X1, X2, X3]:((((((((((((((((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v5_rlvect_1(X1))&v6_rlvect_1(X1))&v7_rlvect_1(X1))&v8_rlvect_1(X1))&v3_normsp_0(X1))&v4_normsp_0(X1))&v2_normsp_1(X1))&v2_funcsdom(X1))&v3_group_1(X1))&v1_vectsp_1(X1))&v3_vectsp_1(X1))&v5_lopban_2(X1))&l1_lopban_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>?[X4]:(((v1_funct_1(X4)&v1_funct_2(X4,k5_numbers,k1_numbers))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))&![X5]:(m2_subset_1(X5,k1_numbers,k5_numbers)=>k1_seq_1(X4,X5)=k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,k8_lopban_4(X1,X2,X3,X5))),X5)))), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', s1_seq_1__e1_46__lopban_4)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', cc1_relset_1)).
fof(redefinition_k9_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k9_real_1(X1,X2)=k6_xcmplx_0(X1,X2)), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', redefinition_k9_real_1)).
fof(t4_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k6_xcmplx_0(X1,k6_numbers)=X1), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', t4_arithm)).
fof(d6_seq_2, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>(v2_comseq_2(X1)<=>?[X2]:(v1_xreal_0(X2)&![X3]:(v1_xreal_0(X3)=>~((~(r1_xxreal_0(X3,k6_numbers))&![X4]:(m2_subset_1(X4,k1_numbers,k5_numbers)=>?[X5]:((m2_subset_1(X5,k1_numbers,k5_numbers)&r1_xxreal_0(X4,X5))&r1_xxreal_0(X3,k18_complex1(k9_real_1(k1_seq_1(X1,X5),X2))))))))))), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', d6_seq_2)).
fof(t10_lopban_4, axiom, ![X1]:((((((((((((((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v5_rlvect_1(X1))&v6_rlvect_1(X1))&v7_rlvect_1(X1))&v8_rlvect_1(X1))&v3_normsp_0(X1))&v4_normsp_0(X1))&v2_normsp_1(X1))&v2_funcsdom(X1))&v3_group_1(X1))&v1_vectsp_1(X1))&v3_vectsp_1(X1))&v5_lopban_2(X1))&l1_lopban_2(X1))=>![X2]:(m2_subset_1(X2,k1_numbers,k5_numbers)=>![X3]:(((v1_funct_1(X3)&v1_funct_2(X3,k5_numbers,u1_struct_0(X1)))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))=>r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,X3),X2)),k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X3)),X2))))), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', t10_lopban_4)).
fof(t33_lopban_4, conjecture, ![X1]:((((((((((((((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v5_rlvect_1(X1))&v6_rlvect_1(X1))&v7_rlvect_1(X1))&v8_rlvect_1(X1))&v3_normsp_0(X1))&v4_normsp_0(X1))&v2_normsp_1(X1))&v2_funcsdom(X1))&v3_group_1(X1))&v1_vectsp_1(X1))&v3_vectsp_1(X1))&v5_lopban_2(X1))&l1_lopban_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(((v1_funct_1(X4)&v1_funct_2(X4,k5_numbers,u1_struct_0(X1)))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))=>(![X5]:(m2_subset_1(X5,k1_numbers,k5_numbers)=>k8_nat_1(u1_struct_0(X1),X4,X5)=k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X5)),X5))=>(v3_normsp_1(X4,X1)&k6_normsp_1(X1,X4)=k4_struct_0(X1))))))), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', t33_lopban_4)).
fof(cc18_valued_0, axiom, ![X1, X2]:(v3_membered(X2)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v3_valued_0(X3))), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', cc18_valued_0)).
fof(cc3_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xcmplx_0(X1)), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', cc3_xreal_0)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', fc3_membered)).
fof(redefinition_k8_nat_1, axiom, ![X1, X2, X3]:((((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,X1))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))&v7_ordinal1(X3))=>k8_nat_1(X1,X2,X3)=k1_funct_1(X2,X3)), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', redefinition_k8_nat_1)).
fof(cc1_xreal_0, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>v1_xreal_0(X1)), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', cc1_xreal_0)).
fof(dt_k8_nat_1, axiom, ![X1, X2, X3]:((((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,X1))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))&v7_ordinal1(X3))=>m1_subset_1(k8_nat_1(X1,X2,X3),X1)), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', dt_k8_nat_1)).
fof(fc6_ordinal1, axiom, (~(v1_xboole_0(k4_ordinal1))&v3_ordinal1(k4_ordinal1)), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', fc6_ordinal1)).
fof(redefinition_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', redefinition_m2_subset_1)).
fof(fc1_numbers, axiom, ~(v1_xboole_0(k1_numbers)), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', fc1_numbers)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', redefinition_k5_numbers)).
fof(t12_lopban_4, axiom, ![X1]:((((((((((((((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v5_rlvect_1(X1))&v6_rlvect_1(X1))&v7_rlvect_1(X1))&v8_rlvect_1(X1))&v3_normsp_0(X1))&v4_normsp_0(X1))&v2_normsp_1(X1))&v2_funcsdom(X1))&v3_group_1(X1))&v1_vectsp_1(X1))&v3_vectsp_1(X1))&v5_lopban_2(X1))&l1_lopban_2(X1))=>![X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,u1_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))=>![X3]:(((v1_funct_1(X3)&v1_funct_2(X3,k5_numbers,k1_numbers))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>(((![X4]:(m2_subset_1(X4,k1_numbers,k5_numbers)=>r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),X2,X4)),k8_nat_1(k1_numbers,X3,X4)))&v2_comseq_2(X3))&k2_seq_2(X3)=k6_numbers)=>(v3_normsp_1(X2,X1)&k6_normsp_1(X1,X2)=k4_struct_0(X1)))))), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', t12_lopban_4)).
fof(cc17_membered, axiom, ![X1]:(v6_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v7_ordinal1(X2))), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', cc17_membered)).
fof(dt_k5_numbers, axiom, m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', dt_k5_numbers)).
fof(fc6_membered, axiom, v6_membered(k4_ordinal1), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', fc6_membered)).
fof(dt_k8_lopban_4, axiom, ![X1, X2, X3, X4]:(((((((((((((((((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v5_rlvect_1(X1))&v6_rlvect_1(X1))&v7_rlvect_1(X1))&v8_rlvect_1(X1))&v3_normsp_0(X1))&v4_normsp_0(X1))&v2_normsp_1(X1))&v2_funcsdom(X1))&v3_group_1(X1))&v1_vectsp_1(X1))&v3_vectsp_1(X1))&v5_lopban_2(X1))&l1_lopban_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))&m1_subset_1(X4,k5_numbers))=>((v1_funct_1(k8_lopban_4(X1,X2,X3,X4))&v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1)))&m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', dt_k8_lopban_4)).
fof(cc8_ordinal1, axiom, ![X1]:(m1_subset_1(X1,k4_ordinal1)=>v7_ordinal1(X1)), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', cc8_ordinal1)).
fof(redefinition_k2_seq_2, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>k2_seq_2(X1)=k1_seq_2(X1)), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', redefinition_k2_seq_2)).
fof(cc30_valued_0, axiom, ![X1]:((v1_relat_1(X1)&v5_relat_1(X1,k1_numbers))=>(v1_relat_1(X1)&v3_valued_0(X1))), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', cc30_valued_0)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('/home/yan/atp/benchmarks/mizar40/all/t33_lopban_4', cc2_relset_1)).
fof(c_0_33, plain, ![X405]:(~v1_xboole_0(X405)|X405=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
cnf(c_0_34, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_35, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
fof(c_0_36, plain, ![X384, X385, X386, X387, X389]:((m2_subset_1(esk63_4(X384,X385,X386,X387),k1_numbers,k5_numbers)|r1_xxreal_0(X387,k6_numbers)|~v1_xreal_0(X387)|~m1_subset_1(X386,u1_struct_0(X384))|~m1_subset_1(X385,u1_struct_0(X384))|(v2_struct_0(X384)|~v13_algstr_0(X384)|~v2_rlvect_1(X384)|~v3_rlvect_1(X384)|~v4_rlvect_1(X384)|~v5_rlvect_1(X384)|~v6_rlvect_1(X384)|~v7_rlvect_1(X384)|~v8_rlvect_1(X384)|~v3_normsp_0(X384)|~v4_normsp_0(X384)|~v2_normsp_1(X384)|~v2_funcsdom(X384)|~v3_group_1(X384)|~v1_vectsp_1(X384)|~v3_vectsp_1(X384)|~v5_lopban_2(X384)|~l1_lopban_2(X384)))&(~m2_subset_1(X389,k1_numbers,k5_numbers)|~r1_xxreal_0(esk63_4(X384,X385,X386,X387),X389)|~r1_xxreal_0(X387,k18_complex1(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X384,k8_lopban_4(X384,X385,X386,X389))),X389)))|r1_xxreal_0(X387,k6_numbers)|~v1_xreal_0(X387)|~m1_subset_1(X386,u1_struct_0(X384))|~m1_subset_1(X385,u1_struct_0(X384))|(v2_struct_0(X384)|~v13_algstr_0(X384)|~v2_rlvect_1(X384)|~v3_rlvect_1(X384)|~v4_rlvect_1(X384)|~v5_rlvect_1(X384)|~v6_rlvect_1(X384)|~v7_rlvect_1(X384)|~v8_rlvect_1(X384)|~v3_normsp_0(X384)|~v4_normsp_0(X384)|~v2_normsp_1(X384)|~v2_funcsdom(X384)|~v3_group_1(X384)|~v1_vectsp_1(X384)|~v3_vectsp_1(X384)|~v5_lopban_2(X384)|~l1_lopban_2(X384)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t32_lopban_4])])])])])])])])).
cnf(c_0_37, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_34])).
cnf(c_0_38, plain, (k1_xboole_0=k6_numbers), inference(evalgc,[status(thm)],[c_0_35])).
cnf(c_0_39, plain, (v1_xboole_0(np__0)), inference(split_conjunct,[status(thm)],[spc0_boole])).
fof(c_0_40, plain, ![X174, X175, X176, X178, X180]:(((m2_subset_1(esk9_3(X174,X175,X176),k1_numbers,k5_numbers)|r1_xxreal_0(X176,k6_numbers)|~v1_xreal_0(X176)|X175!=k1_seq_2(X174)|~v1_xreal_0(X175)|~v2_comseq_2(X174)|(~v1_funct_1(X174)|~v1_funct_2(X174,k5_numbers,k1_numbers)|~m1_subset_1(X174,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(~m2_subset_1(X178,k1_numbers,k5_numbers)|~r1_xxreal_0(esk9_3(X174,X175,X176),X178)|~r1_xxreal_0(X176,k18_complex1(k9_real_1(k1_seq_1(X174,X178),X175)))|r1_xxreal_0(X176,k6_numbers)|~v1_xreal_0(X176)|X175!=k1_seq_2(X174)|~v1_xreal_0(X175)|~v2_comseq_2(X174)|(~v1_funct_1(X174)|~v1_funct_2(X174,k5_numbers,k1_numbers)|~m1_subset_1(X174,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))&((v1_xreal_0(esk10_2(X174,X175))|X175=k1_seq_2(X174)|~v1_xreal_0(X175)|~v2_comseq_2(X174)|(~v1_funct_1(X174)|~v1_funct_2(X174,k5_numbers,k1_numbers)|~m1_subset_1(X174,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&((~r1_xxreal_0(esk10_2(X174,X175),k6_numbers)|X175=k1_seq_2(X174)|~v1_xreal_0(X175)|~v2_comseq_2(X174)|(~v1_funct_1(X174)|~v1_funct_2(X174,k5_numbers,k1_numbers)|~m1_subset_1(X174,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(((m2_subset_1(esk11_3(X174,X175,X180),k1_numbers,k5_numbers)|~m2_subset_1(X180,k1_numbers,k5_numbers)|X175=k1_seq_2(X174)|~v1_xreal_0(X175)|~v2_comseq_2(X174)|(~v1_funct_1(X174)|~v1_funct_2(X174,k5_numbers,k1_numbers)|~m1_subset_1(X174,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(r1_xxreal_0(X180,esk11_3(X174,X175,X180))|~m2_subset_1(X180,k1_numbers,k5_numbers)|X175=k1_seq_2(X174)|~v1_xreal_0(X175)|~v2_comseq_2(X174)|(~v1_funct_1(X174)|~v1_funct_2(X174,k5_numbers,k1_numbers)|~m1_subset_1(X174,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))&(r1_xxreal_0(esk10_2(X174,X175),k18_complex1(k9_real_1(k1_seq_1(X174,esk11_3(X174,X175,X180)),X175)))|~m2_subset_1(X180,k1_numbers,k5_numbers)|X175=k1_seq_2(X174)|~v1_xreal_0(X175)|~v2_comseq_2(X174)|(~v1_funct_1(X174)|~v1_funct_2(X174,k5_numbers,k1_numbers)|~m1_subset_1(X174,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_seq_2])])])])])])])])).
fof(c_0_41, plain, ![X344, X345]:(~v1_relat_1(X344)|~v1_funct_1(X344)|~v3_valued_0(X344)|k1_seq_1(X344,X345)=k1_funct_1(X344,X345)), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_seq_1])])])])).
fof(c_0_42, plain, (((v1_xboole_0(esk57_0)&v1_xcmplx_0(esk57_0))&v1_xxreal_0(esk57_0))&v1_xreal_0(esk57_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc4_xreal_0])])).
cnf(c_0_43, plain, (r1_xxreal_0(X5,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(esk63_4(X2,X3,X4,X5),X1)|~r1_xxreal_0(X5,k18_complex1(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X1))),X1)))|~v1_xreal_0(X5)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v5_rlvect_1(X2)|~v6_rlvect_1(X2)|~v7_rlvect_1(X2)|~v8_rlvect_1(X2)|~v3_normsp_0(X2)|~v4_normsp_0(X2)|~v2_normsp_1(X2)|~v2_funcsdom(X2)|~v3_group_1(X2)|~v1_vectsp_1(X2)|~v3_vectsp_1(X2)|~v5_lopban_2(X2)|~l1_lopban_2(X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_44, plain, (X1=k6_numbers|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_45, plain, (v1_xboole_0(np__0)), inference(evalgc,[status(thm)],[c_0_39])).
fof(c_0_46, plain, ![X362, X363, X364, X366]:((((v1_funct_1(esk61_3(X362,X363,X364))|(v2_struct_0(X362)|~v13_algstr_0(X362)|~v2_rlvect_1(X362)|~v3_rlvect_1(X362)|~v4_rlvect_1(X362)|~v5_rlvect_1(X362)|~v6_rlvect_1(X362)|~v7_rlvect_1(X362)|~v8_rlvect_1(X362)|~v3_normsp_0(X362)|~v4_normsp_0(X362)|~v2_normsp_1(X362)|~v2_funcsdom(X362)|~v3_group_1(X362)|~v1_vectsp_1(X362)|~v3_vectsp_1(X362)|~v5_lopban_2(X362)|~l1_lopban_2(X362)|~m1_subset_1(X363,u1_struct_0(X362))|~m1_subset_1(X364,u1_struct_0(X362))))&(v1_funct_2(esk61_3(X362,X363,X364),k5_numbers,k1_numbers)|(v2_struct_0(X362)|~v13_algstr_0(X362)|~v2_rlvect_1(X362)|~v3_rlvect_1(X362)|~v4_rlvect_1(X362)|~v5_rlvect_1(X362)|~v6_rlvect_1(X362)|~v7_rlvect_1(X362)|~v8_rlvect_1(X362)|~v3_normsp_0(X362)|~v4_normsp_0(X362)|~v2_normsp_1(X362)|~v2_funcsdom(X362)|~v3_group_1(X362)|~v1_vectsp_1(X362)|~v3_vectsp_1(X362)|~v5_lopban_2(X362)|~l1_lopban_2(X362)|~m1_subset_1(X363,u1_struct_0(X362))|~m1_subset_1(X364,u1_struct_0(X362)))))&(m1_subset_1(esk61_3(X362,X363,X364),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|(v2_struct_0(X362)|~v13_algstr_0(X362)|~v2_rlvect_1(X362)|~v3_rlvect_1(X362)|~v4_rlvect_1(X362)|~v5_rlvect_1(X362)|~v6_rlvect_1(X362)|~v7_rlvect_1(X362)|~v8_rlvect_1(X362)|~v3_normsp_0(X362)|~v4_normsp_0(X362)|~v2_normsp_1(X362)|~v2_funcsdom(X362)|~v3_group_1(X362)|~v1_vectsp_1(X362)|~v3_vectsp_1(X362)|~v5_lopban_2(X362)|~l1_lopban_2(X362)|~m1_subset_1(X363,u1_struct_0(X362))|~m1_subset_1(X364,u1_struct_0(X362)))))&(~m2_subset_1(X366,k1_numbers,k5_numbers)|k1_seq_1(esk61_3(X362,X363,X364),X366)=k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X362,k8_lopban_4(X362,X363,X364,X366))),X366)|(v2_struct_0(X362)|~v13_algstr_0(X362)|~v2_rlvect_1(X362)|~v3_rlvect_1(X362)|~v4_rlvect_1(X362)|~v5_rlvect_1(X362)|~v6_rlvect_1(X362)|~v7_rlvect_1(X362)|~v8_rlvect_1(X362)|~v3_normsp_0(X362)|~v4_normsp_0(X362)|~v2_normsp_1(X362)|~v2_funcsdom(X362)|~v3_group_1(X362)|~v1_vectsp_1(X362)|~v3_vectsp_1(X362)|~v5_lopban_2(X362)|~l1_lopban_2(X362)|~m1_subset_1(X363,u1_struct_0(X362))|~m1_subset_1(X364,u1_struct_0(X362))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[s1_seq_1__e1_46__lopban_4])])])])])])])])).
cnf(c_0_47, plain, (r1_xxreal_0(esk10_2(X1,X2),k18_complex1(k9_real_1(k1_seq_1(X1,esk11_3(X1,X2,X3)),X2)))|X2=k1_seq_2(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_48, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_49, plain, ![X67, X68, X69]:(~m1_subset_1(X69,k1_zfmisc_1(k2_zfmisc_1(X67,X68)))|v1_relat_1(X69)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_50, plain, (v1_xboole_0(esk57_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_51, plain, (v2_struct_0(X2)|r1_xxreal_0(X5,k6_numbers)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v5_rlvect_1(X2)|~v6_rlvect_1(X2)|~v7_rlvect_1(X2)|~v8_rlvect_1(X2)|~v3_normsp_0(X2)|~v4_normsp_0(X2)|~v2_normsp_1(X2)|~v2_funcsdom(X2)|~v3_group_1(X2)|~v1_vectsp_1(X2)|~v3_vectsp_1(X2)|~v5_lopban_2(X2)|~l1_lopban_2(X2)|~v1_xreal_0(X5)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(esk63_4(X2,X3,X4,X5),X1)|~r1_xxreal_0(X5,k18_complex1(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X1))),X1)))), inference(evalgc,[status(thm)],[c_0_43])).
cnf(c_0_52, plain, (k6_numbers=np__0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45])])).
cnf(c_0_53, plain, (k1_seq_1(esk61_3(X2,X3,X4),X1)=k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X1))),X1)|v2_struct_0(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v5_rlvect_1(X2)|~v6_rlvect_1(X2)|~v7_rlvect_1(X2)|~v8_rlvect_1(X2)|~v3_normsp_0(X2)|~v4_normsp_0(X2)|~v2_normsp_1(X2)|~v2_funcsdom(X2)|~v3_group_1(X2)|~v1_vectsp_1(X2)|~v3_vectsp_1(X2)|~v5_lopban_2(X2)|~l1_lopban_2(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_54, plain, (X2=k1_seq_2(X1)|r1_xxreal_0(esk10_2(X1,X2),k18_complex1(k9_real_1(k1_seq_1(X1,esk11_3(X1,X2,X3)),X2)))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_47])).
cnf(c_0_55, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_48])).
cnf(c_0_56, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_49])).
fof(c_0_57, plain, ![X353, X354]:(~m1_subset_1(X353,k1_numbers)|~v1_xreal_0(X354)|k9_real_1(X353,X354)=k6_xcmplx_0(X353,X354)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k9_real_1])])).
fof(c_0_58, plain, ![X394]:(~v1_xcmplx_0(X394)|k6_xcmplx_0(X394,k6_numbers)=X394), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_arithm])])).
cnf(c_0_59, plain, (v1_xreal_0(esk57_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_60, plain, (v1_xboole_0(esk57_0)), inference(evalgc,[status(thm)],[c_0_50])).
cnf(c_0_61, plain, (r1_xxreal_0(X1,np__0)|v2_struct_0(X2)|~r1_xxreal_0(X1,k18_complex1(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),X5)))|~r1_xxreal_0(esk63_4(X2,X3,X4,X1),X5)|~v1_xreal_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(rw,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_62, plain, (k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X1))),X1)=k1_seq_1(esk61_3(X2,X3,X4),X1)|v2_struct_0(X2)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v5_rlvect_1(X2)|~v6_rlvect_1(X2)|~v7_rlvect_1(X2)|~v8_rlvect_1(X2)|~v3_normsp_0(X2)|~v4_normsp_0(X2)|~v2_normsp_1(X2)|~v2_funcsdom(X2)|~v3_group_1(X2)|~v1_vectsp_1(X2)|~v3_vectsp_1(X2)|~v5_lopban_2(X2)|~l1_lopban_2(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_53])).
cnf(c_0_63, plain, (m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|v2_struct_0(X1)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_64, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk10_2(X2,X1),k18_complex1(k9_real_1(k1_funct_1(X2,esk11_3(X2,X1,X3)),X1)))|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v3_valued_0(X2)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55])])).
cnf(c_0_65, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_56])).
cnf(c_0_66, plain, (k9_real_1(X1,X2)=k6_xcmplx_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_67, plain, (k6_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_68, plain, (v1_xreal_0(esk57_0)), inference(evalgc,[status(thm)],[c_0_59])).
cnf(c_0_69, plain, (esk57_0=k6_numbers), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_60])])).
fof(c_0_70, plain, ![X165, X167, X169, X170, X172]:(((v1_xreal_0(esk5_1(X165))|~v2_comseq_2(X165)|(~v1_funct_1(X165)|~v1_funct_2(X165,k5_numbers,k1_numbers)|~m1_subset_1(X165,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&((m2_subset_1(esk6_2(X165,X167),k1_numbers,k5_numbers)|r1_xxreal_0(X167,k6_numbers)|~v1_xreal_0(X167)|~v2_comseq_2(X165)|(~v1_funct_1(X165)|~v1_funct_2(X165,k5_numbers,k1_numbers)|~m1_subset_1(X165,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(~m2_subset_1(X169,k1_numbers,k5_numbers)|~r1_xxreal_0(esk6_2(X165,X167),X169)|~r1_xxreal_0(X167,k18_complex1(k9_real_1(k1_seq_1(X165,X169),esk5_1(X165))))|r1_xxreal_0(X167,k6_numbers)|~v1_xreal_0(X167)|~v2_comseq_2(X165)|(~v1_funct_1(X165)|~v1_funct_2(X165,k5_numbers,k1_numbers)|~m1_subset_1(X165,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))))&((v1_xreal_0(esk7_2(X165,X170))|~v1_xreal_0(X170)|v2_comseq_2(X165)|(~v1_funct_1(X165)|~v1_funct_2(X165,k5_numbers,k1_numbers)|~m1_subset_1(X165,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&((~r1_xxreal_0(esk7_2(X165,X170),k6_numbers)|~v1_xreal_0(X170)|v2_comseq_2(X165)|(~v1_funct_1(X165)|~v1_funct_2(X165,k5_numbers,k1_numbers)|~m1_subset_1(X165,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(((m2_subset_1(esk8_3(X165,X170,X172),k1_numbers,k5_numbers)|~m2_subset_1(X172,k1_numbers,k5_numbers)|~v1_xreal_0(X170)|v2_comseq_2(X165)|(~v1_funct_1(X165)|~v1_funct_2(X165,k5_numbers,k1_numbers)|~m1_subset_1(X165,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(r1_xxreal_0(X172,esk8_3(X165,X170,X172))|~m2_subset_1(X172,k1_numbers,k5_numbers)|~v1_xreal_0(X170)|v2_comseq_2(X165)|(~v1_funct_1(X165)|~v1_funct_2(X165,k5_numbers,k1_numbers)|~m1_subset_1(X165,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))&(r1_xxreal_0(esk7_2(X165,X170),k18_complex1(k9_real_1(k1_seq_1(X165,esk8_3(X165,X170,X172)),X170)))|~m2_subset_1(X172,k1_numbers,k5_numbers)|~v1_xreal_0(X170)|v2_comseq_2(X165)|(~v1_funct_1(X165)|~v1_funct_2(X165,k5_numbers,k1_numbers)|~m1_subset_1(X165,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_seq_2])])])])])])])])).
cnf(c_0_71, plain, (r1_xxreal_0(X1,np__0)|v2_struct_0(X2)|~r1_xxreal_0(X1,k18_complex1(k1_seq_1(esk61_3(X2,X3,X4),X5)))|~r1_xxreal_0(esk63_4(X2,X3,X4,X1),X5)|~v1_xreal_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62])])).
cnf(c_0_72, plain, (v1_funct_1(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_73, plain, (v2_struct_0(X1)|m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_63])).
cnf(c_0_74, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk10_2(X2,X1),k18_complex1(k9_real_1(k1_funct_1(X2,esk11_3(X2,X1,X3)),X1)))|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(csr,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_75, plain, (k9_real_1(X1,X2)=k6_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_66])).
cnf(c_0_76, plain, (k6_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_67])).
cnf(c_0_77, plain, (v1_xreal_0(k6_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_69])])).
cnf(c_0_78, plain, (r1_xxreal_0(esk7_2(X1,X2),k18_complex1(k9_real_1(k1_seq_1(X1,esk8_3(X1,X2,X3)),X2)))|v2_comseq_2(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_79, plain, (r1_xxreal_0(X1,np__0)|v2_struct_0(X2)|~r1_xxreal_0(X1,k18_complex1(k1_funct_1(esk61_3(X2,X3,X4),X5)))|~r1_xxreal_0(esk63_4(X2,X3,X4,X1),X5)|~v1_xreal_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)|~v3_valued_0(esk61_3(X2,X3,X4))|~v1_relat_1(esk61_3(X2,X3,X4))|~v1_funct_1(esk61_3(X2,X3,X4))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_55])])).
cnf(c_0_80, plain, (v2_struct_0(X1)|v1_funct_1(esk61_3(X1,X2,X3))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_72])).
cnf(c_0_81, plain, (v1_relat_1(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_73])])).
cnf(c_0_82, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk10_2(X2,X1),k18_complex1(k6_xcmplx_0(k1_funct_1(X2,esk11_3(X2,X1,X3)),X1)))|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_funct_1(X2,esk11_3(X2,X1,X3)),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75])])).
cnf(c_0_83, plain, (k6_xcmplx_0(X1,np__0)=X1|~v1_xcmplx_0(X1)), inference(rw,[status(thm)],[c_0_76, c_0_52])).
cnf(c_0_84, plain, (v1_xreal_0(np__0)), inference(rw,[status(thm)],[c_0_77, c_0_52])).
cnf(c_0_85, plain, (v2_comseq_2(X1)|r1_xxreal_0(esk7_2(X1,X2),k18_complex1(k9_real_1(k1_seq_1(X1,esk8_3(X1,X2,X3)),X2)))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_78])).
fof(c_0_86, plain, ![X369, X370, X371]:(v2_struct_0(X369)|~v13_algstr_0(X369)|~v2_rlvect_1(X369)|~v3_rlvect_1(X369)|~v4_rlvect_1(X369)|~v5_rlvect_1(X369)|~v6_rlvect_1(X369)|~v7_rlvect_1(X369)|~v8_rlvect_1(X369)|~v3_normsp_0(X369)|~v4_normsp_0(X369)|~v2_normsp_1(X369)|~v2_funcsdom(X369)|~v3_group_1(X369)|~v1_vectsp_1(X369)|~v3_vectsp_1(X369)|~v5_lopban_2(X369)|~l1_lopban_2(X369)|(~m2_subset_1(X370,k1_numbers,k5_numbers)|(~v1_funct_1(X371)|~v1_funct_2(X371,k5_numbers,u1_struct_0(X369))|~m1_subset_1(X371,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X369))))|r1_xxreal_0(k1_normsp_0(X369,k8_nat_1(u1_struct_0(X369),k1_bhsp_4(X369,X371),X370)),k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X369,X371)),X370))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t10_lopban_4])])])])])])).
fof(c_0_87, negated_conjecture, ~(![X1]:((((((((((((((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v5_rlvect_1(X1))&v6_rlvect_1(X1))&v7_rlvect_1(X1))&v8_rlvect_1(X1))&v3_normsp_0(X1))&v4_normsp_0(X1))&v2_normsp_1(X1))&v2_funcsdom(X1))&v3_group_1(X1))&v1_vectsp_1(X1))&v3_vectsp_1(X1))&v5_lopban_2(X1))&l1_lopban_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(((v1_funct_1(X4)&v1_funct_2(X4,k5_numbers,u1_struct_0(X1)))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))=>(![X5]:(m2_subset_1(X5,k1_numbers,k5_numbers)=>k8_nat_1(u1_struct_0(X1),X4,X5)=k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X5)),X5))=>(v3_normsp_1(X4,X1)&k6_normsp_1(X1,X4)=k4_struct_0(X1)))))))), inference(assume_negation,[status(cth)],[t33_lopban_4])).
cnf(c_0_88, plain, (r1_xxreal_0(X1,np__0)|v2_struct_0(X2)|~r1_xxreal_0(X1,k18_complex1(k1_funct_1(esk61_3(X2,X3,X4),X5)))|~r1_xxreal_0(esk63_4(X2,X3,X4,X1),X5)|~v1_xreal_0(X1)|~v3_valued_0(esk61_3(X2,X3,X4))|~m2_subset_1(X5,k1_numbers,k5_numbers)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_79, c_0_80]), c_0_81])).
cnf(c_0_89, plain, (k1_seq_2(X1)=np__0|r1_xxreal_0(esk10_2(X1,np__0),k18_complex1(k1_funct_1(X1,esk11_3(X1,np__0,X2))))|~v2_comseq_2(X1)|~v3_valued_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk11_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xcmplx_0(k1_funct_1(X1,esk11_3(X1,np__0,X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_84])])])).
cnf(c_0_90, plain, (v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|v2_struct_0(X1)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_91, plain, (r1_xxreal_0(esk7_2(X1,X2),k18_complex1(k9_real_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v3_valued_0(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_55])])).
cnf(c_0_92, plain, (v2_struct_0(X1)|r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,X3),X2)),k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X3)),X2))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X3)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))), inference(split_conjunct,[status(thm)],[c_0_86])).
fof(c_0_93, negated_conjecture, ![X11]:((((((((((((((((((~v2_struct_0(esk1_0)&v13_algstr_0(esk1_0))&v2_rlvect_1(esk1_0))&v3_rlvect_1(esk1_0))&v4_rlvect_1(esk1_0))&v5_rlvect_1(esk1_0))&v6_rlvect_1(esk1_0))&v7_rlvect_1(esk1_0))&v8_rlvect_1(esk1_0))&v3_normsp_0(esk1_0))&v4_normsp_0(esk1_0))&v2_normsp_1(esk1_0))&v2_funcsdom(esk1_0))&v3_group_1(esk1_0))&v1_vectsp_1(esk1_0))&v3_vectsp_1(esk1_0))&v5_lopban_2(esk1_0))&l1_lopban_2(esk1_0))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(((v1_funct_1(esk4_0)&v1_funct_2(esk4_0,k5_numbers,u1_struct_0(esk1_0)))&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))&((~m2_subset_1(X11,k1_numbers,k5_numbers)|k8_nat_1(u1_struct_0(esk1_0),esk4_0,X11)=k8_nat_1(u1_struct_0(esk1_0),k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X11)),X11))&(~v3_normsp_1(esk4_0,esk1_0)|k6_normsp_1(esk1_0,esk4_0)!=k4_struct_0(esk1_0))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_87])])])])])])])).
cnf(c_0_94, plain, (r1_xxreal_0(esk10_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|k1_seq_2(esk61_3(X1,X2,X3))=np__0|~r1_xxreal_0(esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)),esk11_3(esk61_3(X1,X2,X3),np__0,X4))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~v3_valued_0(esk61_3(X1,X2,X3))|~m2_subset_1(esk11_3(esk61_3(X1,X2,X3),np__0,X4),k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_xcmplx_0(k1_funct_1(esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,X4)))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(k1_funct_1(esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,X4)),k1_numbers)|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89])])).
cnf(c_0_95, plain, (v2_struct_0(X1)|v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_90])).
cnf(c_0_96, plain, (r1_xxreal_0(X1,esk11_3(X2,X3,X1))|X3=k1_seq_2(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xreal_0(X3)|~v2_comseq_2(X2)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_97, plain, (m2_subset_1(esk63_4(X1,X2,X3,X4),k1_numbers,k5_numbers)|r1_xxreal_0(X4,k6_numbers)|v2_struct_0(X1)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_98, plain, (r1_xxreal_0(esk7_2(X1,X2),k18_complex1(k9_real_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(csr,[status(thm)],[c_0_91, c_0_65])).
fof(c_0_99, plain, ![X52, X53, X54]:(~v3_membered(X53)|(~m1_subset_1(X54,k1_zfmisc_1(k2_zfmisc_1(X52,X53)))|v3_valued_0(X54))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc18_valued_0])])])])])).
cnf(c_0_100, plain, (v2_struct_0(X1)|r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,X3),X2)),k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X3)),X2))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~v1_funct_1(X3)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_92])).
cnf(c_0_101, negated_conjecture, (k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)=k8_nat_1(u1_struct_0(esk1_0),k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_102, negated_conjecture, (l1_lopban_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_103, negated_conjecture, (v5_lopban_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_104, negated_conjecture, (v3_vectsp_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_105, negated_conjecture, (v1_vectsp_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_106, negated_conjecture, (v3_group_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_107, negated_conjecture, (v2_funcsdom(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_108, negated_conjecture, (v2_normsp_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_109, negated_conjecture, (v4_normsp_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_110, negated_conjecture, (v3_normsp_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_111, negated_conjecture, (v8_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_112, negated_conjecture, (v7_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_113, negated_conjecture, (v6_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_114, negated_conjecture, (v5_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_115, negated_conjecture, (v4_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_116, negated_conjecture, (v3_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_117, negated_conjecture, (v2_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_118, negated_conjecture, (v13_algstr_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_119, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_120, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|r1_xxreal_0(esk10_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|~r1_xxreal_0(esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)),esk11_3(esk61_3(X1,X2,X3),np__0,X4))|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~v3_valued_0(esk61_3(X1,X2,X3))|~v1_xcmplx_0(k1_funct_1(esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,X4)))|~m2_subset_1(esk11_3(esk61_3(X1,X2,X3),np__0,X4),k1_numbers,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(k1_funct_1(esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,X4)),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_94, c_0_73]), c_0_80]), c_0_95])).
cnf(c_0_121, plain, (X3=k1_seq_2(X2)|r1_xxreal_0(X1,esk11_3(X2,X3,X1))|~v1_funct_1(X2)|~v1_xreal_0(X3)|~v2_comseq_2(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_96])).
cnf(c_0_122, plain, (v2_struct_0(X1)|r1_xxreal_0(X4,k6_numbers)|m2_subset_1(esk63_4(X1,X2,X3,X4),k1_numbers,k5_numbers)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_97])).
fof(c_0_123, plain, ![X123]:(~v1_xreal_0(X123)|v1_xcmplx_0(X123)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_xreal_0])])).
cnf(c_0_124, plain, (r1_xxreal_0(esk7_2(X1,X2),k18_complex1(k6_xcmplx_0(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_75])])).
cnf(c_0_125, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_126, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
cnf(c_0_127, plain, (r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)),k1_seq_1(esk61_3(X1,X2,X3),X4))|v2_struct_0(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_62])])).
fof(c_0_128, plain, ![X350, X351, X352]:(~v1_funct_1(X351)|~v1_funct_2(X351,k5_numbers,X350)|~m1_subset_1(X351,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X350)))|~v7_ordinal1(X352)|k8_nat_1(X350,X351,X352)=k1_funct_1(X351,X352)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_nat_1])])).
cnf(c_0_129, negated_conjecture, (k8_nat_1(u1_struct_0(esk1_0),k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),X1)=k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_101])).
cnf(c_0_130, negated_conjecture, (l1_lopban_2(esk1_0)), inference(evalgc,[status(thm)],[c_0_102])).
cnf(c_0_131, negated_conjecture, (v5_lopban_2(esk1_0)), inference(evalgc,[status(thm)],[c_0_103])).
cnf(c_0_132, negated_conjecture, (v3_vectsp_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_104])).
cnf(c_0_133, negated_conjecture, (v1_vectsp_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_134, negated_conjecture, (v3_group_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_106])).
cnf(c_0_135, negated_conjecture, (v2_funcsdom(esk1_0)), inference(evalgc,[status(thm)],[c_0_107])).
cnf(c_0_136, negated_conjecture, (v2_normsp_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_108])).
cnf(c_0_137, negated_conjecture, (v4_normsp_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_109])).
cnf(c_0_138, negated_conjecture, (v3_normsp_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_110])).
cnf(c_0_139, negated_conjecture, (v8_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_111])).
cnf(c_0_140, negated_conjecture, (v7_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_112])).
cnf(c_0_141, negated_conjecture, (v6_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_113])).
cnf(c_0_142, negated_conjecture, (v5_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_114])).
cnf(c_0_143, negated_conjecture, (v4_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_115])).
cnf(c_0_144, negated_conjecture, (v3_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_116])).
cnf(c_0_145, negated_conjecture, (v2_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_117])).
cnf(c_0_146, negated_conjecture, (v13_algstr_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_118])).
cnf(c_0_147, negated_conjecture, (~v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_119])).
cnf(c_0_148, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_149, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_150, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|r1_xxreal_0(esk10_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~v3_valued_0(esk61_3(X1,X2,X3))|~v1_xcmplx_0(k1_funct_1(esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)))))|~m2_subset_1(esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m2_subset_1(esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)),k1_numbers,k5_numbers)|~m1_subset_1(k1_funct_1(esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)))),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121]), c_0_84])])])).
cnf(c_0_151, plain, (r1_xxreal_0(X1,np__0)|m2_subset_1(esk63_4(X2,X3,X4,X1),k1_numbers,k5_numbers)|v2_struct_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(rw,[status(thm)],[c_0_122, c_0_52])).
cnf(c_0_152, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_123])).
fof(c_0_153, plain, ![X73]:(~m1_subset_1(X73,k1_numbers)|v1_xreal_0(X73)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_xreal_0])])).
fof(c_0_154, plain, ![X206, X207, X208]:(~v1_funct_1(X207)|~v1_funct_2(X207,k5_numbers,X206)|~m1_subset_1(X207,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X206)))|~v7_ordinal1(X208)|m1_subset_1(k8_nat_1(X206,X207,X208),X206)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_nat_1])])).
fof(c_0_155, plain, (~v1_xboole_0(k4_ordinal1)&v3_ordinal1(k4_ordinal1)), inference(fof_simplification,[status(thm)],[fc6_ordinal1])).
cnf(c_0_156, plain, (r1_xxreal_0(esk7_2(X1,np__0),k18_complex1(k1_funct_1(X1,esk8_3(X1,np__0,X2))))|v2_comseq_2(X1)|~v3_valued_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk8_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xcmplx_0(k1_funct_1(X1,esk8_3(X1,np__0,X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_83]), c_0_84])])])).
cnf(c_0_157, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_125])).
cnf(c_0_158, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_126])).
cnf(c_0_159, plain, (r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)),k1_funct_1(esk61_3(X1,X2,X3),X4))|v2_struct_0(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v3_valued_0(esk61_3(X1,X2,X3))|~v1_relat_1(esk61_3(X1,X2,X3))|~v1_funct_1(esk61_3(X1,X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_55])])).
cnf(c_0_160, plain, (k8_nat_1(X2,X1,X3)=k1_funct_1(X1,X3)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v7_ordinal1(X3)), inference(split_conjunct,[status(thm)],[c_0_128])).
cnf(c_0_161, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_129]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_162, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_148])).
cnf(c_0_163, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_149])).
cnf(c_0_164, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|r1_xxreal_0(esk10_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~v3_valued_0(esk61_3(X1,X2,X3))|~v1_xcmplx_0(k1_funct_1(esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)))))|~m2_subset_1(esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(k1_funct_1(esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)))),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_150, c_0_73]), c_0_80]), c_0_95]), c_0_151])).
cnf(c_0_165, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_152])).
cnf(c_0_166, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_167, plain, (m1_subset_1(k8_nat_1(X2,X1,X3),X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v7_ordinal1(X3)), inference(split_conjunct,[status(thm)],[c_0_154])).
fof(c_0_168, plain, ![X355, X356, X357, X357]:((~m2_subset_1(X357,X355,X356)|m1_subset_1(X357,X356)|(v1_xboole_0(X355)|v1_xboole_0(X356)|~m1_subset_1(X356,k1_zfmisc_1(X355))))&(~m1_subset_1(X357,X356)|m2_subset_1(X357,X355,X356)|(v1_xboole_0(X355)|v1_xboole_0(X356)|~m1_subset_1(X356,k1_zfmisc_1(X355))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m2_subset_1])])])])])])])).
fof(c_0_169, plain, ~v1_xboole_0(k1_numbers), inference(fof_simplification,[status(thm)],[fc1_numbers])).
cnf(c_0_170, plain, (~v1_xboole_0(k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_155])).
cnf(c_0_171, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
cnf(c_0_172, plain, (r1_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|v2_comseq_2(esk61_3(X1,X2,X3))|~r1_xxreal_0(esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)),esk8_3(esk61_3(X1,X2,X3),np__0,X4))|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~v3_valued_0(esk61_3(X1,X2,X3))|~m2_subset_1(esk8_3(esk61_3(X1,X2,X3),np__0,X4),k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_xcmplx_0(k1_funct_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,X4)))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(k1_funct_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,X4)),k1_numbers)|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_156])])).
cnf(c_0_173, plain, (v3_valued_0(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_73]), c_0_158])])])).
cnf(c_0_174, plain, (r1_xxreal_0(X1,esk8_3(X2,X3,X1))|v2_comseq_2(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xreal_0(X3)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_70])).
fof(c_0_175, plain, ![X372, X373, X374]:(((v3_normsp_1(X373,X372)|(m2_subset_1(esk62_3(X372,X373,X374),k1_numbers,k5_numbers)|~v2_comseq_2(X374)|k2_seq_2(X374)!=k6_numbers)|(~v1_funct_1(X374)|~v1_funct_2(X374,k5_numbers,k1_numbers)|~m1_subset_1(X374,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|(~v1_funct_1(X373)|~v1_funct_2(X373,k5_numbers,u1_struct_0(X372))|~m1_subset_1(X373,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X372)))))|(v2_struct_0(X372)|~v13_algstr_0(X372)|~v2_rlvect_1(X372)|~v3_rlvect_1(X372)|~v4_rlvect_1(X372)|~v5_rlvect_1(X372)|~v6_rlvect_1(X372)|~v7_rlvect_1(X372)|~v8_rlvect_1(X372)|~v3_normsp_0(X372)|~v4_normsp_0(X372)|~v2_normsp_1(X372)|~v2_funcsdom(X372)|~v3_group_1(X372)|~v1_vectsp_1(X372)|~v3_vectsp_1(X372)|~v5_lopban_2(X372)|~l1_lopban_2(X372)))&(k6_normsp_1(X372,X373)=k4_struct_0(X372)|(m2_subset_1(esk62_3(X372,X373,X374),k1_numbers,k5_numbers)|~v2_comseq_2(X374)|k2_seq_2(X374)!=k6_numbers)|(~v1_funct_1(X374)|~v1_funct_2(X374,k5_numbers,k1_numbers)|~m1_subset_1(X374,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|(~v1_funct_1(X373)|~v1_funct_2(X373,k5_numbers,u1_struct_0(X372))|~m1_subset_1(X373,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X372)))))|(v2_struct_0(X372)|~v13_algstr_0(X372)|~v2_rlvect_1(X372)|~v3_rlvect_1(X372)|~v4_rlvect_1(X372)|~v5_rlvect_1(X372)|~v6_rlvect_1(X372)|~v7_rlvect_1(X372)|~v8_rlvect_1(X372)|~v3_normsp_0(X372)|~v4_normsp_0(X372)|~v2_normsp_1(X372)|~v2_funcsdom(X372)|~v3_group_1(X372)|~v1_vectsp_1(X372)|~v3_vectsp_1(X372)|~v5_lopban_2(X372)|~l1_lopban_2(X372))))&((v3_normsp_1(X373,X372)|(~r1_xxreal_0(k1_normsp_0(X372,k8_nat_1(u1_struct_0(X372),X373,esk62_3(X372,X373,X374))),k8_nat_1(k1_numbers,X374,esk62_3(X372,X373,X374)))|~v2_comseq_2(X374)|k2_seq_2(X374)!=k6_numbers)|(~v1_funct_1(X374)|~v1_funct_2(X374,k5_numbers,k1_numbers)|~m1_subset_1(X374,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|(~v1_funct_1(X373)|~v1_funct_2(X373,k5_numbers,u1_struct_0(X372))|~m1_subset_1(X373,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X372)))))|(v2_struct_0(X372)|~v13_algstr_0(X372)|~v2_rlvect_1(X372)|~v3_rlvect_1(X372)|~v4_rlvect_1(X372)|~v5_rlvect_1(X372)|~v6_rlvect_1(X372)|~v7_rlvect_1(X372)|~v8_rlvect_1(X372)|~v3_normsp_0(X372)|~v4_normsp_0(X372)|~v2_normsp_1(X372)|~v2_funcsdom(X372)|~v3_group_1(X372)|~v1_vectsp_1(X372)|~v3_vectsp_1(X372)|~v5_lopban_2(X372)|~l1_lopban_2(X372)))&(k6_normsp_1(X372,X373)=k4_struct_0(X372)|(~r1_xxreal_0(k1_normsp_0(X372,k8_nat_1(u1_struct_0(X372),X373,esk62_3(X372,X373,X374))),k8_nat_1(k1_numbers,X374,esk62_3(X372,X373,X374)))|~v2_comseq_2(X374)|k2_seq_2(X374)!=k6_numbers)|(~v1_funct_1(X374)|~v1_funct_2(X374,k5_numbers,k1_numbers)|~m1_subset_1(X374,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|(~v1_funct_1(X373)|~v1_funct_2(X373,k5_numbers,u1_struct_0(X372))|~m1_subset_1(X373,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X372)))))|(v2_struct_0(X372)|~v13_algstr_0(X372)|~v2_rlvect_1(X372)|~v3_rlvect_1(X372)|~v4_rlvect_1(X372)|~v5_rlvect_1(X372)|~v6_rlvect_1(X372)|~v7_rlvect_1(X372)|~v8_rlvect_1(X372)|~v3_normsp_0(X372)|~v4_normsp_0(X372)|~v2_normsp_1(X372)|~v2_funcsdom(X372)|~v3_group_1(X372)|~v1_vectsp_1(X372)|~v3_vectsp_1(X372)|~v5_lopban_2(X372)|~l1_lopban_2(X372))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t12_lopban_4])])])])])])])])).
cnf(c_0_176, plain, (r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)),k1_funct_1(esk61_3(X1,X2,X3),X4))|v2_struct_0(X1)|~v3_valued_0(esk61_3(X1,X2,X3))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_159, c_0_80]), c_0_81])).
cnf(c_0_177, plain, (k8_nat_1(X2,X1,X3)=k1_funct_1(X1,X3)|~v1_funct_1(X1)|~v7_ordinal1(X3)|~v1_funct_2(X1,k5_numbers,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))), inference(evalgc,[status(thm)],[c_0_160])).
cnf(c_0_178, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_62]), c_0_162]), c_0_163]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_179, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|r1_xxreal_0(esk10_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~v3_valued_0(esk61_3(X1,X2,X3))|~m2_subset_1(esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(k1_funct_1(esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)))),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_xreal_0(k1_funct_1(esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_165])])).
cnf(c_0_180, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_166])).
cnf(c_0_181, plain, (m1_subset_1(k8_nat_1(X2,X1,X3),X2)|~v1_funct_1(X1)|~v7_ordinal1(X3)|~v1_funct_2(X1,k5_numbers,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))), inference(evalgc,[status(thm)],[c_0_167])).
fof(c_0_182, plain, ![X45, X46]:(~v6_membered(X45)|(~m1_subset_1(X46,X45)|v7_ordinal1(X46))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc17_membered])])])])])).
cnf(c_0_183, plain, (m1_subset_1(X1,X3)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_168])).
cnf(c_0_184, plain, (m2_subset_1(esk11_3(X1,X2,X3),k1_numbers,k5_numbers)|X2=k1_seq_2(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_185, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[dt_k5_numbers])).
cnf(c_0_186, plain, (~v1_xboole_0(k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_169])).
cnf(c_0_187, plain, (~v1_xboole_0(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_170])).
cnf(c_0_188, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_171])).
cnf(c_0_189, plain, (v6_membered(k4_ordinal1)), inference(split_conjunct,[status(thm)],[fc6_membered])).
cnf(c_0_190, plain, (r1_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~r1_xxreal_0(esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)),esk8_3(esk61_3(X1,X2,X3),np__0,X4))|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~v1_xcmplx_0(k1_funct_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,X4)))|~m2_subset_1(esk8_3(esk61_3(X1,X2,X3),np__0,X4),k1_numbers,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(k1_funct_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,X4)),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_172, c_0_73]), c_0_80]), c_0_95]), c_0_173])).
cnf(c_0_191, plain, (v2_comseq_2(X2)|r1_xxreal_0(X1,esk8_3(X2,X3,X1))|~v1_funct_1(X2)|~v1_xreal_0(X3)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_174])).
cnf(c_0_192, plain, (v3_normsp_1(X1,X2)|v2_struct_0(X2)|~r1_xxreal_0(k1_normsp_0(X2,k8_nat_1(u1_struct_0(X2),X1,esk62_3(X2,X1,X3))),k8_nat_1(k1_numbers,X3,esk62_3(X2,X1,X3)))|~v2_comseq_2(X3)|k2_seq_2(X3)!=k6_numbers|~v1_funct_1(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v5_rlvect_1(X2)|~v6_rlvect_1(X2)|~v7_rlvect_1(X2)|~v8_rlvect_1(X2)|~v3_normsp_0(X2)|~v4_normsp_0(X2)|~v2_normsp_1(X2)|~v2_funcsdom(X2)|~v3_group_1(X2)|~v1_vectsp_1(X2)|~v3_vectsp_1(X2)|~v5_lopban_2(X2)|~l1_lopban_2(X2)), inference(split_conjunct,[status(thm)],[c_0_175])).
cnf(c_0_193, plain, (r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)),k8_nat_1(X5,esk61_3(X1,X2,X3),X4))|v2_struct_0(X1)|~v3_valued_0(esk61_3(X1,X2,X3))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,X5)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_177])])).
cnf(c_0_194, plain, (v3_normsp_1(X1,X2)|m2_subset_1(esk62_3(X2,X1,X3),k1_numbers,k5_numbers)|v2_struct_0(X2)|~v2_comseq_2(X3)|k2_seq_2(X3)!=k6_numbers|~v1_funct_1(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v5_rlvect_1(X2)|~v6_rlvect_1(X2)|~v7_rlvect_1(X2)|~v8_rlvect_1(X2)|~v3_normsp_0(X2)|~v4_normsp_0(X2)|~v2_normsp_1(X2)|~v2_funcsdom(X2)|~v3_group_1(X2)|~v1_vectsp_1(X2)|~v3_vectsp_1(X2)|~v5_lopban_2(X2)|~l1_lopban_2(X2)), inference(split_conjunct,[status(thm)],[c_0_175])).
cnf(c_0_195, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|~r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),X2,esk62_3(X1,X2,X3))),k8_nat_1(k1_numbers,X3,esk62_3(X1,X2,X3)))|~v2_comseq_2(X3)|k2_seq_2(X3)!=k6_numbers|~v1_funct_1(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)), inference(split_conjunct,[status(thm)],[c_0_175])).
cnf(c_0_196, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_relat_1(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_55])])).
cnf(c_0_197, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|r1_xxreal_0(esk10_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~m2_subset_1(esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(k1_funct_1(esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)))),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_179, c_0_180]), c_0_173])).
cnf(c_0_198, plain, (m1_subset_1(k1_funct_1(X2,X3),X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_177])])).
cnf(c_0_199, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_182])).
cnf(c_0_200, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_183])).
cnf(c_0_201, plain, (X2=k1_seq_2(X1)|m2_subset_1(esk11_3(X1,X2,X3),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_184])).
cnf(c_0_202, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_185])).
cnf(c_0_203, plain, (~v1_xboole_0(k1_numbers)), inference(evalgc,[status(thm)],[c_0_186])).
cnf(c_0_204, plain, (~v1_xboole_0(k5_numbers)), inference(rw,[status(thm)],[c_0_187, c_0_188])).
cnf(c_0_205, plain, (v6_membered(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_189])).
cnf(c_0_206, plain, (r1_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~v1_xcmplx_0(k1_funct_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)))))|~m2_subset_1(esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m2_subset_1(esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)),k1_numbers,k5_numbers)|~m1_subset_1(k1_funct_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)))),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_190, c_0_191]), c_0_84])])])).
cnf(c_0_207, plain, (v2_struct_0(X2)|v3_normsp_1(X1,X2)|k2_seq_2(X3)!=k6_numbers|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v5_rlvect_1(X2)|~v6_rlvect_1(X2)|~v7_rlvect_1(X2)|~v8_rlvect_1(X2)|~v3_normsp_0(X2)|~v4_normsp_0(X2)|~v2_normsp_1(X2)|~v2_funcsdom(X2)|~v3_group_1(X2)|~v1_vectsp_1(X2)|~v3_vectsp_1(X2)|~v5_lopban_2(X2)|~l1_lopban_2(X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~v2_comseq_2(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~r1_xxreal_0(k1_normsp_0(X2,k8_nat_1(u1_struct_0(X2),X1,esk62_3(X2,X1,X3))),k8_nat_1(k1_numbers,X3,esk62_3(X2,X1,X3)))), inference(evalgc,[status(thm)],[c_0_192])).
cnf(c_0_208, plain, (r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)),k8_nat_1(X5,esk61_3(X1,X2,X3),X4))|v2_struct_0(X1)|~v7_ordinal1(X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,X5)|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_193, c_0_80]), c_0_173])).
cnf(c_0_209, negated_conjecture, (v1_funct_2(esk4_0,k5_numbers,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_210, negated_conjecture, (v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_211, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_93])).
fof(c_0_212, plain, ![X202, X203, X204, X205]:(((v1_funct_1(k8_lopban_4(X202,X203,X204,X205))|(v2_struct_0(X202)|~v13_algstr_0(X202)|~v2_rlvect_1(X202)|~v3_rlvect_1(X202)|~v4_rlvect_1(X202)|~v5_rlvect_1(X202)|~v6_rlvect_1(X202)|~v7_rlvect_1(X202)|~v8_rlvect_1(X202)|~v3_normsp_0(X202)|~v4_normsp_0(X202)|~v2_normsp_1(X202)|~v2_funcsdom(X202)|~v3_group_1(X202)|~v1_vectsp_1(X202)|~v3_vectsp_1(X202)|~v5_lopban_2(X202)|~l1_lopban_2(X202)|~m1_subset_1(X203,u1_struct_0(X202))|~m1_subset_1(X204,u1_struct_0(X202))|~m1_subset_1(X205,k5_numbers)))&(v1_funct_2(k8_lopban_4(X202,X203,X204,X205),k5_numbers,u1_struct_0(X202))|(v2_struct_0(X202)|~v13_algstr_0(X202)|~v2_rlvect_1(X202)|~v3_rlvect_1(X202)|~v4_rlvect_1(X202)|~v5_rlvect_1(X202)|~v6_rlvect_1(X202)|~v7_rlvect_1(X202)|~v8_rlvect_1(X202)|~v3_normsp_0(X202)|~v4_normsp_0(X202)|~v2_normsp_1(X202)|~v2_funcsdom(X202)|~v3_group_1(X202)|~v1_vectsp_1(X202)|~v3_vectsp_1(X202)|~v5_lopban_2(X202)|~l1_lopban_2(X202)|~m1_subset_1(X203,u1_struct_0(X202))|~m1_subset_1(X204,u1_struct_0(X202))|~m1_subset_1(X205,k5_numbers))))&(m1_subset_1(k8_lopban_4(X202,X203,X204,X205),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X202))))|(v2_struct_0(X202)|~v13_algstr_0(X202)|~v2_rlvect_1(X202)|~v3_rlvect_1(X202)|~v4_rlvect_1(X202)|~v5_rlvect_1(X202)|~v6_rlvect_1(X202)|~v7_rlvect_1(X202)|~v8_rlvect_1(X202)|~v3_normsp_0(X202)|~v4_normsp_0(X202)|~v2_normsp_1(X202)|~v2_funcsdom(X202)|~v3_group_1(X202)|~v1_vectsp_1(X202)|~v3_vectsp_1(X202)|~v5_lopban_2(X202)|~l1_lopban_2(X202)|~m1_subset_1(X203,u1_struct_0(X202))|~m1_subset_1(X204,u1_struct_0(X202))|~m1_subset_1(X205,k5_numbers)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k8_lopban_4])])])])).
cnf(c_0_213, plain, (v2_struct_0(X2)|v3_normsp_1(X1,X2)|m2_subset_1(esk62_3(X2,X1,X3),k1_numbers,k5_numbers)|k2_seq_2(X3)!=k6_numbers|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v5_rlvect_1(X2)|~v6_rlvect_1(X2)|~v7_rlvect_1(X2)|~v8_rlvect_1(X2)|~v3_normsp_0(X2)|~v4_normsp_0(X2)|~v2_normsp_1(X2)|~v2_funcsdom(X2)|~v3_group_1(X2)|~v1_vectsp_1(X2)|~v3_vectsp_1(X2)|~v5_lopban_2(X2)|~l1_lopban_2(X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~v2_comseq_2(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_194])).
cnf(c_0_214, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=k6_numbers|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v2_comseq_2(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),X2,esk62_3(X1,X2,X3))),k8_nat_1(k1_numbers,X3,esk62_3(X1,X2,X3)))), inference(evalgc,[status(thm)],[c_0_195])).
cnf(c_0_215, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k8_nat_1(X2,esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_relat_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v7_ordinal1(X1)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X2)|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_177])])).
cnf(c_0_216, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|r1_xxreal_0(esk10_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~m2_subset_1(esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0))))|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_198])])).
cnf(c_0_217, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_199])).
cnf(c_0_218, plain, (m1_subset_1(esk11_3(X1,X2,X3),k5_numbers)|X2=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_201]), c_0_202])]), c_0_203]), c_0_204])])).
cnf(c_0_219, plain, (v6_membered(k5_numbers)), inference(rw,[status(thm)],[c_0_205, c_0_188])).
cnf(c_0_220, plain, (r1_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~v1_xcmplx_0(k1_funct_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)))))|~m2_subset_1(esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(k1_funct_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)))),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_206, c_0_73]), c_0_80]), c_0_95]), c_0_151])).
cnf(c_0_221, plain, (v3_normsp_1(X1,X2)|v2_struct_0(X2)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X2,k8_nat_1(u1_struct_0(X2),X1,esk62_3(X2,X1,X3))),k8_nat_1(k1_numbers,X3,esk62_3(X2,X1,X3)))|~v2_comseq_2(X3)|~v1_funct_2(X1,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(rw,[status(thm)],[c_0_207, c_0_52])).
cnf(c_0_222, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k8_nat_1(X2,esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X2)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_129]), c_0_162]), c_0_163]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_223, negated_conjecture, (v1_funct_2(esk4_0,k5_numbers,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_209])).
cnf(c_0_224, negated_conjecture, (v1_funct_1(esk4_0)), inference(evalgc,[status(thm)],[c_0_210])).
cnf(c_0_225, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))), inference(evalgc,[status(thm)],[c_0_211])).
cnf(c_0_226, plain, (m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|v2_struct_0(X1)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_212])).
cnf(c_0_227, plain, (v3_normsp_1(X1,X2)|m2_subset_1(esk62_3(X2,X1,X3),k1_numbers,k5_numbers)|v2_struct_0(X2)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(X1,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(rw,[status(thm)],[c_0_213, c_0_52])).
cnf(c_0_228, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),X2,esk62_3(X1,X2,X3))),k8_nat_1(k1_numbers,X3,esk62_3(X1,X2,X3)))|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(rw,[status(thm)],[c_0_214, c_0_52])).
cnf(c_0_229, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k8_nat_1(X2,esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X2)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))), inference(csr,[status(thm)],[c_0_215, c_0_65])).
cnf(c_0_230, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|r1_xxreal_0(esk10_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v7_ordinal1(esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0))))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~m2_subset_1(esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_216, c_0_73]), c_0_80]), c_0_95])).
cnf(c_0_231, plain, (v7_ordinal1(esk11_3(X1,X2,X3))|X2=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_218]), c_0_219])])])).
cnf(c_0_232, plain, (r1_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~m2_subset_1(esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(k1_funct_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)))),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_xreal_0(k1_funct_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_165])])).
cnf(c_0_233, plain, (m2_subset_1(esk8_3(X1,X2,X3),k1_numbers,k5_numbers)|v2_comseq_2(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_234, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_222]), c_0_223]), c_0_224]), c_0_225]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_235, plain, (v2_struct_0(X1)|m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_226])).
cnf(c_0_236, plain, (m1_subset_1(esk62_3(X1,X2,X3),k5_numbers)|v3_normsp_1(X2,X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_227]), c_0_202])]), c_0_203]), c_0_204])])).
cnf(c_0_237, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_229]), c_0_223]), c_0_224]), c_0_225]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_238, plain, (v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|v2_struct_0(X1)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_212])).
fof(c_0_239, plain, ![X155]:(~m1_subset_1(X155,k4_ordinal1)|v7_ordinal1(X155)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_ordinal1])])).
cnf(c_0_240, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|r1_xxreal_0(esk10_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~m2_subset_1(esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m2_subset_1(esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_231]), c_0_84])])])).
cnf(c_0_241, plain, (r1_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~m2_subset_1(esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(k1_funct_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)))),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_232, c_0_180])).
cnf(c_0_242, plain, (v2_comseq_2(X1)|m2_subset_1(esk8_3(X1,X2,X3),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_233])).
cnf(c_0_243, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_235]), c_0_162]), c_0_163]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_244, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|m1_subset_1(esk62_3(esk1_0,esk4_0,X1),k5_numbers)|k2_seq_2(X1)!=np__0|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_236, c_0_225]), c_0_223]), c_0_224]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_245, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_235]), c_0_162]), c_0_163]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_246, plain, (v2_struct_0(X1)|v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_238])).
cnf(c_0_247, plain, (v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_212])).
cnf(c_0_248, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_239])).
cnf(c_0_249, plain, (X2=k1_seq_2(X1)|~r1_xxreal_0(esk10_2(X1,X2),k6_numbers)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_250, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|r1_xxreal_0(esk10_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~m2_subset_1(esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_240, c_0_73]), c_0_80]), c_0_95]), c_0_151])).
cnf(c_0_251, plain, (r1_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~m2_subset_1(esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0))))|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_198])])).
cnf(c_0_252, plain, (m1_subset_1(esk8_3(X1,X2,X3),k5_numbers)|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_242]), c_0_202])]), c_0_203]), c_0_204])])).
cnf(c_0_253, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(csr,[status(thm)],[c_0_243, c_0_244])).
cnf(c_0_254, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_246]), c_0_162]), c_0_163]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_255, plain, (v2_struct_0(X1)|v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_247])).
cnf(c_0_256, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_248])).
cnf(c_0_257, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|m2_subset_1(esk62_3(X1,X2,X3),k1_numbers,k5_numbers)|v2_struct_0(X1)|~v2_comseq_2(X3)|k2_seq_2(X3)!=k6_numbers|~v1_funct_1(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)), inference(split_conjunct,[status(thm)],[c_0_175])).
cnf(c_0_258, plain, (X2=k1_seq_2(X1)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~r1_xxreal_0(esk10_2(X1,X2),k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_249])).
cnf(c_0_259, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|r1_xxreal_0(esk10_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m2_subset_1(esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_201]), c_0_84])])])).
cnf(c_0_260, plain, (r1_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v7_ordinal1(esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0))))|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~m2_subset_1(esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_251, c_0_73]), c_0_80]), c_0_95])).
cnf(c_0_261, plain, (v7_ordinal1(esk8_3(X1,X2,X3))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_252]), c_0_219])])])).
cnf(c_0_262, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_246]), c_0_162]), c_0_163]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_263, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_255]), c_0_162]), c_0_163]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_264, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)), inference(rw,[status(thm)],[c_0_256, c_0_188])).
cnf(c_0_265, plain, (m2_subset_1(X1,X3,X2)|v1_xboole_0(X3)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_168])).
cnf(c_0_266, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|m2_subset_1(esk62_3(X1,X2,X3),k1_numbers,k5_numbers)|k2_seq_2(X3)!=k6_numbers|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v2_comseq_2(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_257])).
cnf(c_0_267, plain, (X1=k1_seq_2(X2)|~r1_xxreal_0(esk10_2(X2,X1),np__0)|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(rw,[status(thm)],[c_0_258, c_0_52])).
cnf(c_0_268, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|r1_xxreal_0(esk10_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_259, c_0_73]), c_0_80]), c_0_95]), c_0_151])).
cnf(c_0_269, plain, (r1_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~m2_subset_1(esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m2_subset_1(esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_261]), c_0_84])])])).
cnf(c_0_270, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(csr,[status(thm)],[c_0_262, c_0_244])).
cnf(c_0_271, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_264])])).
cnf(c_0_272, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_265])).
cnf(c_0_273, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|m2_subset_1(esk62_3(X1,X2,X3),k1_numbers,k5_numbers)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(rw,[status(thm)],[c_0_266, c_0_52])).
fof(c_0_274, plain, ![X346]:(~v1_funct_1(X346)|~v1_funct_2(X346,k5_numbers,k1_numbers)|~m1_subset_1(X346,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|k2_seq_2(X346)=k1_seq_2(X346)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_seq_2])])).
cnf(c_0_275, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_268]), c_0_84])])])).
cnf(c_0_276, plain, (v1_xreal_0(esk10_2(X1,X2))|X2=k1_seq_2(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_277, plain, (v2_comseq_2(X1)|~r1_xxreal_0(esk7_2(X1,X2),k6_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_278, plain, (r1_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~m2_subset_1(esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_269, c_0_73]), c_0_80]), c_0_95]), c_0_151])).
cnf(c_0_279, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_270, c_0_255]), c_0_162]), c_0_163]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_280, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_272]), c_0_202])]), c_0_204]), c_0_203])])).
cnf(c_0_281, plain, (m1_subset_1(esk62_3(X1,X2,X3),k5_numbers)|k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_273]), c_0_202])]), c_0_203]), c_0_204])])).
cnf(c_0_282, plain, (k2_seq_2(X1)=k1_seq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_274])).
cnf(c_0_283, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_275, c_0_73]), c_0_80]), c_0_95])).
cnf(c_0_284, plain, (X2=k1_seq_2(X1)|v1_xreal_0(esk10_2(X1,X2))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_276])).
cnf(c_0_285, plain, (v2_comseq_2(X1)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~r1_xxreal_0(esk7_2(X1,X2),k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_277])).
cnf(c_0_286, plain, (r1_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m2_subset_1(esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_242]), c_0_84])])])).
cnf(c_0_287, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(csr,[status(thm)],[c_0_279, c_0_244])).
cnf(c_0_288, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_281]), c_0_223]), c_0_224]), c_0_225]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_289, plain, (k2_seq_2(X1)=k1_seq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_282])).
cnf(c_0_290, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_284]), c_0_84])])])).
fof(c_0_291, plain, ![X110]:((v1_relat_1(X110)|(~v1_relat_1(X110)|~v5_relat_1(X110,k1_numbers)))&(v3_valued_0(X110)|(~v1_relat_1(X110)|~v5_relat_1(X110,k1_numbers)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc30_valued_0])])])).
cnf(c_0_292, plain, (v2_comseq_2(X1)|~r1_xxreal_0(esk7_2(X1,X2),np__0)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(rw,[status(thm)],[c_0_285, c_0_52])).
cnf(c_0_293, plain, (r1_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_286, c_0_73]), c_0_80]), c_0_95]), c_0_151])).
cnf(c_0_294, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_287, c_0_264])])).
cnf(c_0_295, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|k1_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_289])])).
cnf(c_0_296, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_290, c_0_73]), c_0_80]), c_0_95])).
cnf(c_0_297, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_291])).
fof(c_0_298, plain, ![X102, X103, X104]:((v4_relat_1(X104,X102)|~m1_subset_1(X104,k1_zfmisc_1(k2_zfmisc_1(X102,X103))))&(v5_relat_1(X104,X103)|~m1_subset_1(X104,k1_zfmisc_1(k2_zfmisc_1(X102,X103))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
cnf(c_0_299, plain, (v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_293]), c_0_84])])])).
cnf(c_0_300, plain, (v1_xreal_0(esk7_2(X1,X2))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_301, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(csr,[status(thm)],[c_0_294, c_0_244])).
cnf(c_0_302, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_295, c_0_296]), c_0_162]), c_0_163]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_303, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_297])).
cnf(c_0_304, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_298])).
cnf(c_0_305, plain, (v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_299, c_0_73]), c_0_80]), c_0_95])).
cnf(c_0_306, plain, (v2_comseq_2(X1)|v1_xreal_0(esk7_2(X1,X2))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_300])).
cnf(c_0_307, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_301, c_0_272]), c_0_202])]), c_0_204]), c_0_203])])).
cnf(c_0_308, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v5_relat_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_relat_1(esk61_3(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_303])])).
cnf(c_0_309, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(evalgc,[status(thm)],[c_0_304])).
cnf(c_0_310, plain, (v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_306]), c_0_84])])])).
cnf(c_0_311, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(csr,[status(thm)],[c_0_307, c_0_244])).
cnf(c_0_312, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_308, c_0_65]), c_0_309])).
cnf(c_0_313, plain, (v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_310, c_0_73]), c_0_80]), c_0_95])).
cnf(c_0_314, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k1_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_289])])).
cnf(c_0_315, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_312, c_0_313]), c_0_162]), c_0_163]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_316, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_314, c_0_296]), c_0_162]), c_0_163]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_317, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_73]), c_0_162]), c_0_163]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_318, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_316, c_0_313]), c_0_162]), c_0_163]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_319, negated_conjecture, (~v3_normsp_1(esk4_0,esk1_0)|k6_normsp_1(esk1_0,esk4_0)!=k4_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_320, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_317, c_0_95]), c_0_162]), c_0_163]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_321, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_73]), c_0_162]), c_0_163]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_322, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)!=k4_struct_0(esk1_0)|~v3_normsp_1(esk4_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_319])).
cnf(c_0_323, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_320, c_0_80]), c_0_162]), c_0_163]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_324, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_321, c_0_95]), c_0_162]), c_0_163]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_325, negated_conjecture, (~v3_normsp_1(esk4_0,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_322, c_0_323])])])).
cnf(c_0_326, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_324, c_0_80]), c_0_162]), c_0_163]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134]), c_0_135]), c_0_136]), c_0_137]), c_0_138]), c_0_139]), c_0_140]), c_0_141]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_146])]), c_0_147])])).
cnf(c_0_327, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_325, c_0_326])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 328
# Proof object clause steps            : 267
# Proof object formula steps           : 61
# Proof object conjectures             : 94
# Proof object clause conjectures      : 91
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 76
# Proof object initial formulas used   : 33
# Proof object generating inferences   : 70
# Proof object simplifying inferences  : 625
# Training examples: 148 positive, 2712 negative
# Training: Positive examples begin
cnf(c_0_188, plain, (k4_ordinal1=k5_numbers)).#trainpos
cnf(c_0_38, plain, (k1_xboole_0=k6_numbers)).#trainpos
cnf(c_0_146, negated_conjecture, (v13_algstr_0(esk1_0))).#trainpos
cnf(c_0_145, negated_conjecture, (v2_rlvect_1(esk1_0))).#trainpos
cnf(c_0_144, negated_conjecture, (v3_rlvect_1(esk1_0))).#trainpos
cnf(c_0_143, negated_conjecture, (v4_rlvect_1(esk1_0))).#trainpos
cnf(c_0_142, negated_conjecture, (v5_rlvect_1(esk1_0))).#trainpos
cnf(c_0_141, negated_conjecture, (v6_rlvect_1(esk1_0))).#trainpos
cnf(c_0_140, negated_conjecture, (v7_rlvect_1(esk1_0))).#trainpos
cnf(c_0_139, negated_conjecture, (v8_rlvect_1(esk1_0))).#trainpos
cnf(c_0_138, negated_conjecture, (v3_normsp_0(esk1_0))).#trainpos
cnf(c_0_137, negated_conjecture, (v4_normsp_0(esk1_0))).#trainpos
cnf(c_0_147, negated_conjecture, (~v2_struct_0(esk1_0))).#trainpos
cnf(c_0_136, negated_conjecture, (v2_normsp_1(esk1_0))).#trainpos
cnf(c_0_135, negated_conjecture, (v2_funcsdom(esk1_0))).#trainpos
cnf(c_0_322, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)!=k4_struct_0(esk1_0)|~v3_normsp_1(esk4_0,esk1_0))).#trainpos
cnf(c_0_134, negated_conjecture, (v3_group_1(esk1_0))).#trainpos
cnf(c_0_133, negated_conjecture, (v1_vectsp_1(esk1_0))).#trainpos
cnf(c_0_132, negated_conjecture, (v3_vectsp_1(esk1_0))).#trainpos
cnf(c_0_131, negated_conjecture, (v5_lopban_2(esk1_0))).#trainpos
cnf(c_0_130, negated_conjecture, (l1_lopban_2(esk1_0))).#trainpos
cnf(c_0_224, negated_conjecture, (v1_funct_1(esk4_0))).#trainpos
cnf(c_0_163, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_162, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_203, plain, (~v1_xboole_0(k1_numbers))).#trainpos
cnf(c_0_205, plain, (v6_membered(k4_ordinal1))).#trainpos
cnf(c_0_158, plain, (v3_membered(k1_numbers))).#trainpos
cnf(c_0_223, negated_conjecture, (v1_funct_2(esk4_0,k5_numbers,u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_225, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainpos
cnf(c_0_187, plain, (~v1_xboole_0(k4_ordinal1))).#trainpos
cnf(c_0_68, plain, (v1_xreal_0(esk57_0))).#trainpos
cnf(c_0_60, plain, (v1_xboole_0(esk57_0))).#trainpos
cnf(c_0_129, negated_conjecture, (k8_nat_1(u1_struct_0(esk1_0),k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),X1)=k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainpos
cnf(c_0_45, plain, (v1_xboole_0(np__0))).#trainpos
cnf(c_0_37, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainpos
cnf(c_0_69, plain, (esk57_0=k6_numbers)).#trainpos
cnf(c_0_52, plain, (k6_numbers=np__0)).#trainpos
cnf(c_0_77, plain, (v1_xreal_0(k6_numbers))).#trainpos
cnf(c_0_165, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainpos
cnf(c_0_202, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainpos
cnf(c_0_217, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).#trainpos
cnf(c_0_55, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainpos
cnf(c_0_75, plain, (k9_real_1(X1,X2)=k6_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_65, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainpos
cnf(c_0_272, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainpos
cnf(c_0_200, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainpos
cnf(c_0_289, plain, (k2_seq_2(X1)=k1_seq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_284, plain, (X2=k1_seq_2(X1)|v1_xreal_0(esk10_2(X1,X2))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_285, plain, (v2_comseq_2(X1)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~r1_xxreal_0(esk7_2(X1,X2),k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_258, plain, (X2=k1_seq_2(X1)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~r1_xxreal_0(esk10_2(X1,X2),k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_177, plain, (k8_nat_1(X2,X1,X3)=k1_funct_1(X1,X3)|~v1_funct_1(X1)|~v7_ordinal1(X3)|~v1_funct_2(X1,k5_numbers,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainpos
cnf(c_0_80, plain, (v2_struct_0(X1)|v1_funct_1(esk61_3(X1,X2,X3))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainpos
cnf(c_0_95, plain, (v2_struct_0(X1)|v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainpos
cnf(c_0_73, plain, (v2_struct_0(X1)|m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainpos
cnf(c_0_191, plain, (v2_comseq_2(X2)|r1_xxreal_0(X1,esk8_3(X2,X3,X1))|~v1_funct_1(X2)|~v1_xreal_0(X3)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_121, plain, (X3=k1_seq_2(X2)|r1_xxreal_0(X1,esk11_3(X2,X3,X1))|~v1_funct_1(X2)|~v1_xreal_0(X3)|~v2_comseq_2(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_242, plain, (v2_comseq_2(X1)|m2_subset_1(esk8_3(X1,X2,X3),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_201, plain, (X2=k1_seq_2(X1)|m2_subset_1(esk11_3(X1,X2,X3),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_255, plain, (v2_struct_0(X1)|v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainpos
cnf(c_0_213, plain, (v2_struct_0(X2)|v3_normsp_1(X1,X2)|m2_subset_1(esk62_3(X2,X1,X3),k1_numbers,k5_numbers)|k2_seq_2(X3)!=k6_numbers|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v5_rlvect_1(X2)|~v6_rlvect_1(X2)|~v7_rlvect_1(X2)|~v8_rlvect_1(X2)|~v3_normsp_0(X2)|~v4_normsp_0(X2)|~v2_normsp_1(X2)|~v2_funcsdom(X2)|~v3_group_1(X2)|~v1_vectsp_1(X2)|~v3_vectsp_1(X2)|~v5_lopban_2(X2)|~l1_lopban_2(X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~v2_comseq_2(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))))).#trainpos
cnf(c_0_266, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|m2_subset_1(esk62_3(X1,X2,X3),k1_numbers,k5_numbers)|k2_seq_2(X3)!=k6_numbers|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v2_comseq_2(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainpos
cnf(c_0_122, plain, (v2_struct_0(X1)|r1_xxreal_0(X4,k6_numbers)|m2_subset_1(esk63_4(X1,X2,X3,X4),k1_numbers,k5_numbers)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainpos
cnf(c_0_246, plain, (v2_struct_0(X1)|v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainpos
cnf(c_0_235, plain, (v2_struct_0(X1)|m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainpos
cnf(c_0_100, plain, (v2_struct_0(X1)|r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,X3),X2)),k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X3)),X2))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~v1_funct_1(X3)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainpos
cnf(c_0_85, plain, (v2_comseq_2(X1)|r1_xxreal_0(esk7_2(X1,X2),k18_complex1(k9_real_1(k1_seq_1(X1,esk8_3(X1,X2,X3)),X2)))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_54, plain, (X2=k1_seq_2(X1)|r1_xxreal_0(esk10_2(X1,X2),k18_complex1(k9_real_1(k1_seq_1(X1,esk11_3(X1,X2,X3)),X2)))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_207, plain, (v2_struct_0(X2)|v3_normsp_1(X1,X2)|k2_seq_2(X3)!=k6_numbers|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v5_rlvect_1(X2)|~v6_rlvect_1(X2)|~v7_rlvect_1(X2)|~v8_rlvect_1(X2)|~v3_normsp_0(X2)|~v4_normsp_0(X2)|~v2_normsp_1(X2)|~v2_funcsdom(X2)|~v3_group_1(X2)|~v1_vectsp_1(X2)|~v3_vectsp_1(X2)|~v5_lopban_2(X2)|~l1_lopban_2(X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~v2_comseq_2(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~r1_xxreal_0(k1_normsp_0(X2,k8_nat_1(u1_struct_0(X2),X1,esk62_3(X2,X1,X3))),k8_nat_1(k1_numbers,X3,esk62_3(X2,X1,X3))))).#trainpos
cnf(c_0_214, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=k6_numbers|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v2_comseq_2(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),X2,esk62_3(X1,X2,X3))),k8_nat_1(k1_numbers,X3,esk62_3(X1,X2,X3))))).#trainpos
cnf(c_0_62, plain, (k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X1))),X1)=k1_seq_1(esk61_3(X2,X3,X4),X1)|v2_struct_0(X2)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v5_rlvect_1(X2)|~v6_rlvect_1(X2)|~v7_rlvect_1(X2)|~v8_rlvect_1(X2)|~v3_normsp_0(X2)|~v4_normsp_0(X2)|~v2_normsp_1(X2)|~v2_funcsdom(X2)|~v3_group_1(X2)|~v1_vectsp_1(X2)|~v3_vectsp_1(X2)|~v5_lopban_2(X2)|~l1_lopban_2(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainpos
cnf(c_0_51, plain, (v2_struct_0(X2)|r1_xxreal_0(X5,k6_numbers)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v5_rlvect_1(X2)|~v6_rlvect_1(X2)|~v7_rlvect_1(X2)|~v8_rlvect_1(X2)|~v3_normsp_0(X2)|~v4_normsp_0(X2)|~v2_normsp_1(X2)|~v2_funcsdom(X2)|~v3_group_1(X2)|~v1_vectsp_1(X2)|~v3_vectsp_1(X2)|~v5_lopban_2(X2)|~l1_lopban_2(X2)|~v1_xreal_0(X5)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(esk63_4(X2,X3,X4,X5),X1)|~r1_xxreal_0(X5,k18_complex1(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X1))),X1))))).#trainpos
cnf(c_0_81, plain, (v1_relat_1(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainpos
cnf(c_0_252, plain, (m1_subset_1(esk8_3(X1,X2,X3),k5_numbers)|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_218, plain, (m1_subset_1(esk11_3(X1,X2,X3),k5_numbers)|X2=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_236, plain, (m1_subset_1(esk62_3(X1,X2,X3),k5_numbers)|v3_normsp_1(X2,X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainpos
cnf(c_0_281, plain, (m1_subset_1(esk62_3(X1,X2,X3),k5_numbers)|k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainpos
cnf(c_0_161, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainpos
cnf(c_0_91, plain, (r1_xxreal_0(esk7_2(X1,X2),k18_complex1(k9_real_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainpos
cnf(c_0_64, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk10_2(X2,X1),k18_complex1(k9_real_1(k1_funct_1(X2,esk11_3(X2,X1,X3)),X1)))|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v3_valued_0(X2)|~v1_relat_1(X2))).#trainpos
cnf(c_0_127, plain, (r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)),k1_seq_1(esk61_3(X1,X2,X3),X4))|v2_struct_0(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainpos
cnf(c_0_71, plain, (r1_xxreal_0(X1,np__0)|v2_struct_0(X2)|~r1_xxreal_0(X1,k18_complex1(k1_seq_1(esk61_3(X2,X3,X4),X5)))|~r1_xxreal_0(esk63_4(X2,X3,X4,X1),X5)|~v1_xreal_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainpos
cnf(c_0_261, plain, (v7_ordinal1(esk8_3(X1,X2,X3))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_231, plain, (v7_ordinal1(esk11_3(X1,X2,X3))|X2=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_178, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainpos
cnf(c_0_124, plain, (r1_xxreal_0(esk7_2(X1,X2),k18_complex1(k6_xcmplx_0(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),k1_numbers))).#trainpos
cnf(c_0_82, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk10_2(X2,X1),k18_complex1(k6_xcmplx_0(k1_funct_1(X2,esk11_3(X2,X1,X3)),X1)))|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_funct_1(X2,esk11_3(X2,X1,X3)),k1_numbers))).#trainpos
cnf(c_0_79, plain, (r1_xxreal_0(X1,np__0)|v2_struct_0(X2)|~r1_xxreal_0(X1,k18_complex1(k1_funct_1(esk61_3(X2,X3,X4),X5)))|~r1_xxreal_0(esk63_4(X2,X3,X4,X1),X5)|~v1_xreal_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)|~v3_valued_0(esk61_3(X2,X3,X4))|~v1_relat_1(esk61_3(X2,X3,X4))|~v1_funct_1(esk61_3(X2,X3,X4)))).#trainpos
cnf(c_0_196, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_relat_1(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0)))).#trainpos
cnf(c_0_215, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k8_nat_1(X2,esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_relat_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v7_ordinal1(X1)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X2)|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainpos
cnf(c_0_237, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainpos
cnf(c_0_245, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers))).#trainpos
cnf(c_0_254, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers))).#trainpos
cnf(c_0_263, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers))).#trainpos
cnf(c_0_76, plain, (k6_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).#trainpos
cnf(c_0_89, plain, (k1_seq_2(X1)=np__0|r1_xxreal_0(esk10_2(X1,np__0),k18_complex1(k1_funct_1(X1,esk11_3(X1,np__0,X2))))|~v2_comseq_2(X1)|~v3_valued_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk11_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xcmplx_0(k1_funct_1(X1,esk11_3(X1,np__0,X2))))).#trainpos
cnf(c_0_156, plain, (r1_xxreal_0(esk7_2(X1,np__0),k18_complex1(k1_funct_1(X1,esk8_3(X1,np__0,X2))))|v2_comseq_2(X1)|~v3_valued_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk8_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xcmplx_0(k1_funct_1(X1,esk8_3(X1,np__0,X2))))).#trainpos
cnf(c_0_94, plain, (r1_xxreal_0(esk10_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|k1_seq_2(esk61_3(X1,X2,X3))=np__0|~r1_xxreal_0(esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)),esk11_3(esk61_3(X1,X2,X3),np__0,X4))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~v3_valued_0(esk61_3(X1,X2,X3))|~m2_subset_1(esk11_3(esk61_3(X1,X2,X3),np__0,X4),k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_xcmplx_0(k1_funct_1(esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,X4)))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(k1_funct_1(esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,X4)),k1_numbers)|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_150, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|r1_xxreal_0(esk10_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~v3_valued_0(esk61_3(X1,X2,X3))|~v1_xcmplx_0(k1_funct_1(esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)))))|~m2_subset_1(esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m2_subset_1(esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)),k1_numbers,k5_numbers)|~m1_subset_1(k1_funct_1(esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)))),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_180, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_256, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainpos
cnf(c_0_271, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers))).#trainpos
cnf(c_0_280, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers))).#trainpos
cnf(c_0_288, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_295, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|k1_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_303, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainpos
cnf(c_0_159, plain, (r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)),k1_funct_1(esk61_3(X1,X2,X3),X4))|v2_struct_0(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v3_valued_0(esk61_3(X1,X2,X3))|~v1_relat_1(esk61_3(X1,X2,X3))|~v1_funct_1(esk61_3(X1,X2,X3)))).#trainpos
cnf(c_0_157, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainpos
cnf(c_0_173, plain, (v3_valued_0(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainpos
cnf(c_0_309, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainpos
cnf(c_0_306, plain, (v2_comseq_2(X1)|v1_xreal_0(esk7_2(X1,X2))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_181, plain, (m1_subset_1(k8_nat_1(X2,X1,X3),X2)|~v1_funct_1(X1)|~v7_ordinal1(X3)|~v1_funct_2(X1,k5_numbers,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainpos
cnf(c_0_198, plain, (m1_subset_1(k1_funct_1(X2,X3),X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainpos
cnf(c_0_244, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|m1_subset_1(esk62_3(esk1_0,esk4_0,X1),k5_numbers)|k2_seq_2(X1)!=np__0|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_172, plain, (r1_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|v2_comseq_2(esk61_3(X1,X2,X3))|~r1_xxreal_0(esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)),esk8_3(esk61_3(X1,X2,X3),np__0,X4))|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~v3_valued_0(esk61_3(X1,X2,X3))|~m2_subset_1(esk8_3(esk61_3(X1,X2,X3),np__0,X4),k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_xcmplx_0(k1_funct_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,X4)))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(k1_funct_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,X4)),k1_numbers)|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_206, plain, (r1_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~v1_xcmplx_0(k1_funct_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)))))|~m2_subset_1(esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m2_subset_1(esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)),k1_numbers,k5_numbers)|~m1_subset_1(k1_funct_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)))),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_179, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|r1_xxreal_0(esk10_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~v3_valued_0(esk61_3(X1,X2,X3))|~m2_subset_1(esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(k1_funct_1(esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)))),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_xreal_0(k1_funct_1(esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0))))))).#trainpos
cnf(c_0_216, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|r1_xxreal_0(esk10_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~m2_subset_1(esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0))))|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_240, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|r1_xxreal_0(esk10_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~m2_subset_1(esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m2_subset_1(esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_259, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|r1_xxreal_0(esk10_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m2_subset_1(esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_275, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainpos
cnf(c_0_290, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_302, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_308, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v5_relat_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_relat_1(esk61_3(esk1_0,esk2_0,esk3_0)))).#trainpos
cnf(c_0_193, plain, (r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)),k8_nat_1(X5,esk61_3(X1,X2,X3),X4))|v2_struct_0(X1)|~v3_valued_0(esk61_3(X1,X2,X3))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,X5)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5))))).#trainpos
cnf(c_0_222, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k8_nat_1(X2,esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X2)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainpos
cnf(c_0_234, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainpos
cnf(c_0_243, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers))).#trainpos
cnf(c_0_262, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers))).#trainpos
cnf(c_0_279, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers))).#trainpos
cnf(c_0_294, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers))).#trainpos
cnf(c_0_307, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers))).#trainpos
cnf(c_0_314, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k1_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_316, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_232, plain, (r1_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~m2_subset_1(esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(k1_funct_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)))),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_xreal_0(k1_funct_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0))))))).#trainpos
cnf(c_0_251, plain, (r1_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~m2_subset_1(esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0))))|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_269, plain, (r1_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~m2_subset_1(esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m2_subset_1(esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_286, plain, (r1_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m2_subset_1(esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_299, plain, (v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainpos
cnf(c_0_310, plain, (v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_315, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_318, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_317, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0)))).#trainpos
cnf(c_0_321, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0)))).#trainpos
cnf(c_0_320, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0)))).#trainpos
cnf(c_0_324, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0)))).#trainpos
cnf(c_0_323, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0))).#trainpos
cnf(c_0_326, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0))).#trainpos
cnf(c_0_325, negated_conjecture, (~v3_normsp_1(esk4_0,esk1_0))).#trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_504, plain, (k4_xcmplx_0(np__0)=np__0)).#trainneg
cnf(i_0_446, plain, (v1_funct_2(esk52_0,k5_numbers,k5_numbers))).#trainneg
cnf(i_0_409, plain, (~v2_struct_0(esk40_0))).#trainneg
cnf(i_0_432, plain, (~v2_struct_0(esk49_0))).#trainneg
cnf(i_0_316, plain, (l1_lopban_2(esk14_0))).#trainneg
cnf(i_0_433, plain, (l1_lopban_2(esk49_0))).#trainneg
cnf(i_0_403, plain, (v1_funct_1(esk36_0))).#trainneg
cnf(i_0_448, plain, (v1_funct_1(esk52_0))).#trainneg
cnf(i_0_320, plain, (l1_struct_0(esk18_0))).#trainneg
cnf(i_0_408, plain, (v7_struct_0(esk40_0))).#trainneg
cnf(i_0_404, plain, (v1_relat_1(esk36_0))).#trainneg
cnf(i_0_451, plain, (v1_relat_1(esk52_0))).#trainneg
cnf(i_0_445, plain, (v1_valued_0(esk52_0))).#trainneg
cnf(i_0_394, plain, (v6_membered(esk33_0))).#trainneg
cnf(i_0_357, plain, (v5_membered(k4_numbers))).#trainneg
cnf(i_0_518, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_352, plain, (~v1_xboole_0(k4_numbers))).#trainneg
cnf(i_0_515, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_353, plain, (v7_membered(k1_numbers))).#trainneg
cnf(i_0_416, plain, (v6_membered(esk43_0))).#trainneg
cnf(i_0_435, plain, (v6_membered(esk50_0))).#trainneg
cnf(i_0_444, plain, (v2_valued_0(esk52_0))).#trainneg
cnf(i_0_429, plain, (v1_xcmplx_0(esk47_0))).#trainneg
cnf(i_0_457, plain, (v1_xcmplx_0(esk53_0))).#trainneg
cnf(i_0_465, plain, (v1_xcmplx_0(esk57_0))).#trainneg
cnf(i_0_443, plain, (v3_valued_0(esk52_0))).#trainneg
cnf(i_0_407, plain, (v1_xxreal_0(esk39_0))).#trainneg
cnf(i_0_428, plain, (v1_xxreal_0(esk47_0))).#trainneg
cnf(i_0_431, plain, (v1_xxreal_0(esk48_0))).#trainneg
cnf(i_0_456, plain, (v1_xxreal_0(esk53_0))).#trainneg
cnf(i_0_459, plain, (v1_xxreal_0(esk54_0))).#trainneg
cnf(i_0_464, plain, (v1_xxreal_0(esk57_0))).#trainneg
cnf(i_0_467, plain, (v1_xxreal_0(esk58_0))).#trainneg
cnf(i_0_406, plain, (v1_xreal_0(esk38_0))).#trainneg
cnf(i_0_377, plain, (~v1_finset_1(k1_numbers))).#trainneg
cnf(i_0_426, plain, (v1_xreal_0(esk47_0))).#trainneg
cnf(i_0_279, plain, (m2_subset_1(k6_numbers,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_348, plain, (~v1_xboole_0(k3_numbers))).#trainneg
cnf(i_0_289, plain, (l1_normsp_1(X1)|~l1_lopban_2(X1))).#trainneg
cnf(i_0_290, plain, (l1_funcsdom(X1)|~l1_lopban_2(X1))).#trainneg
cnf(i_0_454, plain, (v1_xreal_0(esk53_0))).#trainneg
cnf(i_0_402, plain, (v4_valued_0(esk36_0))).#trainneg
cnf(i_0_442, plain, (v4_valued_0(esk52_0))).#trainneg
cnf(i_0_460, plain, (v7_ordinal1(esk55_0))).#trainneg
cnf(i_0_469, plain, (v7_ordinal1(esk59_0))).#trainneg
cnf(i_0_405, plain, (v1_xboole_0(esk37_0))).#trainneg
cnf(i_0_468, plain, (v1_xboole_0(esk58_0))).#trainneg
cnf(i_0_396, plain, (v3_ordinal1(esk34_0))).#trainneg
cnf(i_0_418, plain, (v3_ordinal1(esk44_0))).#trainneg
cnf(i_0_421, plain, (v3_ordinal1(esk45_0))).#trainneg
cnf(i_0_423, plain, (v1_ordinal1(esk45_0))).#trainneg
cnf(i_0_422, plain, (v2_ordinal1(esk45_0))).#trainneg
cnf(i_0_395, plain, (~v1_xboole_0(esk33_0))).#trainneg
cnf(i_0_351, plain, (v4_membered(k3_numbers))).#trainneg
cnf(i_0_419, plain, (~v1_xboole_0(esk44_0))).#trainneg
cnf(i_0_342, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_514, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_417, plain, (~v1_xboole_0(esk43_0))).#trainneg
cnf(i_0_424, plain, (~v1_xboole_0(esk45_0))).#trainneg
cnf(i_0_75, plain, (v5_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_225, plain, (v6_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_366, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_441, plain, (v5_valued_0(esk52_0))).#trainneg
cnf(i_0_434, plain, (v7_membered(esk50_0))).#trainneg
cnf(i_0_455, plain, (v3_xxreal_0(esk53_0))).#trainneg
cnf(i_0_458, plain, (v3_xxreal_0(esk54_0))).#trainneg
cnf(i_0_427, plain, (v2_xxreal_0(esk47_0))).#trainneg
cnf(i_0_430, plain, (v2_xxreal_0(esk48_0))).#trainneg
cnf(i_0_321, plain, (l2_algstr_0(esk19_0))).#trainneg
cnf(i_0_317, plain, (l1_normsp_0(esk15_0))).#trainneg
cnf(i_0_323, plain, (l2_struct_0(esk21_0))).#trainneg
cnf(i_0_410, plain, (l2_struct_0(esk40_0))).#trainneg
cnf(i_0_318, plain, (l1_normsp_1(esk16_0))).#trainneg
cnf(i_0_314, plain, (l1_algstr_0(esk12_0))).#trainneg
cnf(i_0_425, plain, (~v1_xboole_0(esk46_0))).#trainneg
cnf(i_0_436, plain, (~v1_xboole_0(esk50_0))).#trainneg
cnf(i_0_355, plain, (v7_membered(k4_numbers))).#trainneg
cnf(i_0_452, plain, (~v1_xboole_0(esk52_0))).#trainneg
cnf(i_0_354, plain, (v7_membered(k3_numbers))).#trainneg
cnf(i_0_492, plain, (~r1_xxreal_0(np__1,np__0))).#trainneg
cnf(i_0_470, plain, (~v1_xboole_0(esk59_0))).#trainneg
cnf(i_0_365, plain, (~v1_finset_1(k4_numbers))).#trainneg
cnf(i_0_48, plain, (v6_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_192, plain, (v1_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_356, plain, (v7_membered(k4_ordinal1))).#trainneg
cnf(i_0_519, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_315, plain, (l1_funcsdom(esk13_0))).#trainneg
cnf(i_0_329, plain, (l6_algstr_0(esk27_0))).#trainneg
cnf(i_0_319, plain, (l1_rlvect_1(esk17_0))).#trainneg
cnf(i_0_322, plain, (l2_normsp_0(esk20_0))).#trainneg
cnf(i_0_324, plain, (l3_algstr_0(esk22_0))).#trainneg
cnf(i_0_325, plain, (l3_struct_0(esk23_0))).#trainneg
cnf(i_0_326, plain, (l4_algstr_0(esk24_0))).#trainneg
cnf(i_0_327, plain, (l4_struct_0(esk25_0))).#trainneg
cnf(i_0_328, plain, (l5_algstr_0(esk26_0))).#trainneg
cnf(i_0_369, plain, (~v1_finset_1(k3_numbers))).#trainneg
cnf(i_0_291, plain, (l1_struct_0(X1)|~l1_normsp_0(X1))).#trainneg
cnf(i_0_551, plain, (l1_struct_0(esk15_0))).#trainneg
cnf(i_0_547, plain, (esk58_0=k6_numbers)).#trainneg
cnf(i_0_467, plain, (v1_xxreal_0(k6_numbers))).#trainneg
cnf(i_0_549, plain, (esk37_0=k6_numbers)).#trainneg
cnf(i_0_494, plain, (~r1_xxreal_0(np__1,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_491, plain, (~r1_xxreal_0(np__0,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_541, plain, (X1=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_465, plain, (v1_xcmplx_0(k6_numbers))).#trainneg
cnf(i_0_464, plain, (v1_xxreal_0(k6_numbers))).#trainneg
cnf(i_0_300, plain, (l1_struct_0(X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_286, plain, (l1_struct_0(X1)|~l1_algstr_0(X1))).#trainneg
cnf(i_0_301, plain, (l1_struct_0(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_302, plain, (l1_struct_0(X1)|~l3_struct_0(X1))).#trainneg
cnf(i_0_179, plain, (v2_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_189, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_143, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_167, plain, (v3_membered(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_142, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_130, plain, (v4_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_207, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_552, plain, (l1_struct_0(esk40_0))).#trainneg
cnf(i_0_549, plain, (esk37_0=np__0)).#trainneg
cnf(i_0_548, plain, (esk57_0=np__0)).#trainneg
cnf(i_0_547, plain, (esk58_0=np__0)).#trainneg
cnf(i_0_482, plain, (k1_xboole_0=np__0)).#trainneg
cnf(i_0_553, plain, (l1_struct_0(esk21_0))).#trainneg
cnf(i_0_554, plain, (l1_struct_0(esk12_0))).#trainneg
cnf(i_0_563, plain, (v1_xxreal_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_562, plain, (v3_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_555, plain, (l1_struct_0(esk22_0))).#trainneg
cnf(i_0_556, plain, (l1_struct_0(esk23_0))).#trainneg
cnf(i_0_77, plain, (v3_ordinal1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_170, plain, (v3_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_79, plain, (v1_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_78, plain, (v2_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_111, plain, (v7_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_180, plain, (v5_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_310, plain, (l2_algstr_0(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_294, plain, (l2_algstr_0(X1)|~l1_rlvect_1(X1))).#trainneg
cnf(i_0_299, plain, (l1_normsp_0(X1)|~l2_normsp_0(X1))).#trainneg
cnf(i_0_297, plain, (l2_struct_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_298, plain, (l2_struct_0(X1)|~l2_normsp_0(X1))).#trainneg
cnf(i_0_557, plain, (v1_xxreal_0(esk59_0))).#trainneg
cnf(i_0_558, plain, (v1_xxreal_0(esk55_0))).#trainneg
cnf(i_0_559, plain, (v3_membered(k3_numbers))).#trainneg
cnf(i_0_131, plain, (~v3_xxreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_560, plain, (v1_xreal_0(esk59_0))).#trainneg
cnf(i_0_561, plain, (v1_xreal_0(esk55_0))).#trainneg
cnf(i_0_572, plain, (~m1_subset_1(esk54_0,k5_numbers))).#trainneg
cnf(i_0_573, plain, (~m1_subset_1(esk53_0,k5_numbers))).#trainneg
cnf(i_0_566, plain, (v3_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_564, plain, (v1_xreal_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_565, plain, (v3_membered(k4_numbers))).#trainneg
cnf(i_0_567, plain, (l2_algstr_0(esk27_0))).#trainneg
cnf(i_0_306, plain, (l2_struct_0(X1)|~l4_struct_0(X1))).#trainneg
cnf(i_0_296, plain, (l1_algstr_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_288, plain, (l6_algstr_0(X1)|~l1_funcsdom(X1))).#trainneg
cnf(i_0_293, plain, (l1_rlvect_1(X1)|~l1_normsp_1(X1))).#trainneg
cnf(i_0_287, plain, (l1_rlvect_1(X1)|~l1_funcsdom(X1))).#trainneg
cnf(i_0_292, plain, (l2_normsp_0(X1)|~l1_normsp_1(X1))).#trainneg
cnf(i_0_303, plain, (l3_algstr_0(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_304, plain, (l3_struct_0(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_305, plain, (l3_struct_0(X1)|~l4_struct_0(X1))).#trainneg
cnf(i_0_308, plain, (l4_algstr_0(X1)|~l5_algstr_0(X1))).#trainneg
cnf(i_0_307, plain, (l4_struct_0(X1)|~l5_algstr_0(X1))).#trainneg
cnf(i_0_309, plain, (l5_algstr_0(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_568, plain, (l2_algstr_0(esk17_0))).#trainneg
cnf(i_0_569, plain, (l1_normsp_0(esk20_0))).#trainneg
cnf(i_0_570, plain, (l2_struct_0(esk19_0))).#trainneg
cnf(i_0_168, plain, (~v7_ordinal1(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_571, plain, (l2_struct_0(esk20_0))).#trainneg
cnf(i_0_575, plain, (l2_struct_0(esk27_0))).#trainneg
cnf(i_0_593, plain, (~v3_xxreal_0(esk59_0))).#trainneg
cnf(i_0_594, plain, (~v3_xxreal_0(esk55_0))).#trainneg
cnf(i_0_588, plain, (l4_algstr_0(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_589, plain, (l4_struct_0(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_576, plain, (l2_struct_0(esk25_0))).#trainneg
cnf(i_0_577, plain, (l1_struct_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_498, plain, (k6_xcmplx_0(np__0,np__1)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_476, plain, (k16_complex1(X1)=k18_complex1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_381, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_336, plain, (v3_xxreal_0(X1)|v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_334, plain, (v2_xxreal_0(X1)|v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_543, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_275, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_35, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_37, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_39, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_41, plain, (v1_rat_1(X2)|~v4_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_584, plain, (l3_struct_0(esk24_0))).#trainneg
cnf(i_0_585, plain, (l3_struct_0(esk25_0))).#trainneg
cnf(i_0_586, plain, (l4_algstr_0(esk26_0))).#trainneg
cnf(i_0_595, plain, (~v3_xxreal_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_578, plain, (l2_algstr_0(X1)|~l1_funcsdom(X1))).#trainneg
cnf(i_0_587, plain, (l4_struct_0(esk26_0))).#trainneg
cnf(i_0_633, plain, (~v1_xboole_0(esk54_0))).#trainneg
cnf(i_0_634, plain, (~v1_xboole_0(esk53_0))).#trainneg
cnf(i_0_636, plain, (l2_algstr_0(X1)|~l1_lopban_2(X1))).#trainneg
cnf(i_0_639, plain, (l2_struct_0(X1)|~l1_lopban_2(X1))).#trainneg
cnf(i_0_579, plain, (l2_algstr_0(X1)|~l1_normsp_1(X1))).#trainneg
cnf(i_0_43, plain, (v1_int_1(X2)|~v5_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_193, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_528, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_530, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_534, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_538, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_662, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_590, plain, (l2_struct_0(esk17_0))).#trainneg
cnf(i_0_591, plain, (l1_struct_0(esk20_0))).#trainneg
cnf(i_0_581, plain, (l1_normsp_0(X1)|~l1_normsp_1(X1))).#trainneg
cnf(i_0_592, plain, (l1_struct_0(esk19_0))).#trainneg
cnf(i_0_597, plain, (l1_struct_0(esk27_0))).#trainneg
cnf(i_0_663, negated_conjecture, (v1_relat_1(esk4_0))).#trainneg
cnf(i_0_582, plain, (l2_struct_0(X1)|~l1_normsp_1(X1))).#trainneg
cnf(i_0_583, plain, (l1_struct_0(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_312, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_479, plain, (k2_series_1(X1)=k3_series_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_269, plain, (v1_funct_1(k3_series_1(X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_254, plain, (v1_xreal_0(k1_seq_2(X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_268, plain, (v1_funct_2(k3_series_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_267, plain, (m1_subset_1(k3_series_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_480, plain, (k2_normsp_0(X1,X2)=k4_normsp_0(X1,X2)|v2_struct_0(X1)|~v1_funct_1(X2)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_601, plain, (l1_struct_0(esk25_0))).#trainneg
cnf(i_0_598, plain, (l3_struct_0(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_604, plain, (l1_struct_0(esk17_0))).#trainneg
cnf(i_0_630, plain, (l1_struct_0(esk24_0))).#trainneg
cnf(i_0_670, plain, (l1_normsp_0(X1)|~l1_lopban_2(X1))).#trainneg
cnf(i_0_599, plain, (l2_struct_0(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_249, plain, (v2_struct_0(X1)|v1_funct_1(k1_bhsp_4(X1,X2))|~v1_funct_1(X2)|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_272, plain, (v2_struct_0(X1)|v1_funct_1(k4_normsp_0(X1,X2))|~v1_funct_1(X2)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_236, plain, (r1_xxreal_0(X2,k6_numbers)|m2_subset_1(esk6_2(X1,X2),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_271, plain, (v2_struct_0(X1)|v1_funct_2(k4_normsp_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(X2)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_248, plain, (v2_struct_0(X1)|v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_270, plain, (v2_struct_0(X1)|m1_subset_1(k4_normsp_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_247, plain, (v2_struct_0(X1)|m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_1(X2)|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_278, plain, (v2_struct_0(X1)|m1_subset_1(k6_normsp_1(X1,X2),u1_struct_0(X1))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v1_funct_1(X2)|~l1_normsp_1(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_244, plain, (r1_xxreal_0(X3,k6_numbers)|m2_subset_1(esk9_3(X1,X2,X3),k1_numbers,k5_numbers)|X2!=k1_seq_2(X1)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_632, plain, (l3_struct_0(esk26_0))).#trainneg
cnf(i_0_635, plain, (l2_algstr_0(esk13_0))).#trainneg
cnf(i_0_637, plain, (l2_struct_0(esk26_0))).#trainneg
cnf(i_0_177, plain, (~v1_xxreal_0(X1)|~v1_xboole_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_610, negated_conjecture, (v1_xcmplx_0(esk4_0)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_736, plain, (~v1_xboole_0(esk48_0))).#trainneg
cnf(i_0_737, plain, (~v1_xboole_0(esk47_0))).#trainneg
cnf(i_0_640, plain, (l1_struct_0(X1)|~l1_lopban_2(X1))).#trainneg
cnf(i_0_742, negated_conjecture, (l1_struct_0(esk1_0))).#trainneg
cnf(i_0_611, negated_conjecture, (v1_xcmplx_0(esk3_0)|~v1_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_612, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v1_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_235, plain, (r1_xxreal_0(X3,k6_numbers)|~v1_funct_1(X2)|~v1_xreal_0(X3)|~v2_comseq_2(X2)|~r1_xxreal_0(esk6_2(X2,X3),X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r1_xxreal_0(X3,k18_complex1(k9_real_1(k1_seq_1(X2,X1),esk5_1(X2)))))).#trainneg
cnf(i_0_243, plain, (r1_xxreal_0(X4,k6_numbers)|X3!=k1_seq_2(X2)|~v1_funct_1(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v2_comseq_2(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(esk9_3(X2,X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r1_xxreal_0(X4,k18_complex1(k9_real_1(k1_seq_1(X2,X1),X3))))).#trainneg
cnf(i_0_642, plain, (l2_algstr_0(esk16_0))).#trainneg
cnf(i_0_669, plain, (l1_normsp_0(esk16_0))).#trainneg
cnf(i_0_613, plain, (v1_xcmplx_0(k4_struct_0(X1))|~v1_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_671, plain, (l2_struct_0(esk16_0))).#trainneg
cnf(i_0_673, plain, (l1_struct_0(esk26_0))).#trainneg
cnf(i_0_614, plain, (v1_xcmplx_0(k5_numbers)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_615, negated_conjecture, (v1_xxreal_0(esk4_0)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_675, plain, (l1_struct_0(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_687, plain, (v1_relat_1(k3_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_689, negated_conjecture, (k1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),X1)=k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_693, plain, (l2_struct_0(X1)|~l1_funcsdom(X1))).#trainneg
cnf(i_0_733, plain, (l2_struct_0(esk13_0))).#trainneg
cnf(i_0_734, plain, (l1_struct_0(esk13_0))).#trainneg
cnf(i_0_616, negated_conjecture, (v1_xxreal_0(esk3_0)|~v2_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_740, plain, (l1_struct_0(esk49_0))).#trainneg
cnf(i_0_741, plain, (l1_struct_0(esk14_0))).#trainneg
cnf(i_0_626, negated_conjecture, (v1_rat_1(esk3_0)|~v4_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_617, negated_conjecture, (v1_xxreal_0(esk2_0)|~v2_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_618, plain, (v1_xxreal_0(k4_struct_0(X1))|~v2_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_704, plain, (v1_relat_1(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_722, plain, (v1_relat_1(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_746, plain, (m1_subset_1(esk8_3(X1,X2,X3),k1_numbers)|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_748, plain, (m1_subset_1(esk11_3(X1,X2,X3),k1_numbers)|X2=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_749, plain, (r1_xxreal_0(X1,np__0)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_funct_1(X2,X3),esk5_1(X2))))|~r1_xxreal_0(esk6_2(X2,X1),X3)|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v3_valued_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_752, plain, (m1_subset_1(esk62_3(X1,X2,X3),k1_numbers)|v3_normsp_1(X2,X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_754, plain, (m1_subset_1(esk62_3(X1,X2,X3),k1_numbers)|k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_769, plain, (l1_struct_0(esk16_0))).#trainneg
cnf(i_0_175, plain, (~v1_xxreal_0(X1)|~v3_xxreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_619, plain, (v1_xxreal_0(k5_numbers)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_851, plain, (~v2_xxreal_0(esk54_0))).#trainneg
cnf(i_0_852, plain, (~v2_xxreal_0(esk53_0))).#trainneg
cnf(i_0_627, negated_conjecture, (v1_rat_1(esk2_0)|~v4_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_629, plain, (v1_rat_1(k5_numbers)|~v4_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_620, negated_conjecture, (v1_xreal_0(esk4_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_621, negated_conjecture, (v1_xreal_0(esk3_0)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_757, plain, (r1_xxreal_0(X1,np__0)|X2!=k1_seq_2(X3)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_funct_1(X3,X4),X2)))|~r1_xxreal_0(esk9_3(X3,X2,X1),X4)|~v2_comseq_2(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v3_valued_0(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_767, plain, (v1_relat_1(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_779, plain, (r1_xxreal_0(esk7_2(X1,X2),k18_complex1(k6_xcmplx_0(k1_seq_1(X1,esk8_3(X1,X2,X3)),X2)))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X1,esk8_3(X1,X2,X3)),k1_numbers))).#trainneg
cnf(i_0_783, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk10_2(X2,X1),k18_complex1(k6_xcmplx_0(k1_seq_1(X2,esk11_3(X2,X1,X3)),X1)))|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X2,esk11_3(X2,X1,X3)),k1_numbers))).#trainneg
cnf(i_0_791, plain, (l1_struct_0(X1)|~l1_funcsdom(X1))).#trainneg
cnf(i_0_792, negated_conjecture, (k8_nat_1(X1,k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),X2)=k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2)|~v7_ordinal1(X2)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k5_numbers,X1)|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k5_numbers,u1_struct_0(esk1_0))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_645, negated_conjecture, (v1_int_1(esk3_0)|~v5_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_853, plain, (~v3_xxreal_0(esk48_0))).#trainneg
cnf(i_0_622, negated_conjecture, (v1_xreal_0(esk2_0)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_646, negated_conjecture, (v1_int_1(esk2_0)|~v5_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_854, plain, (~v3_xxreal_0(esk47_0))).#trainneg
cnf(i_0_648, plain, (v1_int_1(k5_numbers)|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_623, plain, (v1_xreal_0(k4_struct_0(X1))|~v3_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_624, plain, (v1_xreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_804, plain, (v3_ordinal1(esk8_3(X1,X2,X3))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_818, plain, (v3_ordinal1(esk11_3(X1,X2,X3))|X2=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_821, plain, (v1_xreal_0(esk11_3(X1,X2,X3))|X2=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_830, plain, (v3_normsp_1(k1_bhsp_4(X1,X2),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,X2),X3),k5_numbers)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(k1_bhsp_4(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_833, plain, (v3_normsp_1(k1_bhsp_4(X1,X2),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,X2),X3),k1_numbers)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(k1_bhsp_4(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_650, negated_conjecture, (v7_ordinal1(esk3_0)|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_905, negated_conjecture, (v7_ordinal1(esk3_0)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_517, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_908, negated_conjecture, (~v3_xxreal_0(esk3_0)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_625, negated_conjecture, (v1_rat_1(esk4_0)|~v4_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_914, plain, (v7_ordinal1(np__1))).#trainneg
cnf(i_0_919, plain, (~v3_xxreal_0(np__1))).#trainneg
cnf(i_0_906, negated_conjecture, (v1_xxreal_0(esk3_0)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_628, plain, (v1_rat_1(k4_struct_0(X1))|~v4_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_644, negated_conjecture, (v1_int_1(esk4_0)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_647, plain, (v1_int_1(k4_struct_0(X1))|~v5_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_889, plain, (v1_xxreal_0(esk8_3(X1,X2,X3))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_891, plain, (v2_comseq_2(X1)|~v3_xxreal_0(esk8_3(X1,X2,X3))|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_892, plain, (v1_xxreal_0(esk11_3(X1,X2,X3))|X2=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_894, plain, (X2=k1_seq_2(X1)|~v3_xxreal_0(esk11_3(X1,X2,X3))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_902, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_915, plain, (v3_ordinal1(np__1))).#trainneg
cnf(i_0_917, plain, (v1_xxreal_0(np__1))).#trainneg
cnf(i_0_918, plain, (v1_xreal_0(np__1))).#trainneg
cnf(i_0_649, negated_conjecture, (v7_ordinal1(esk4_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_909, plain, (v1_xcmplx_0(np__1)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_855, plain, (~v2_xxreal_0(esk39_0)|~v3_xxreal_0(esk39_0))).#trainneg
cnf(i_0_856, plain, (~v2_xxreal_0(np__0)|~v3_xxreal_0(np__0))).#trainneg
cnf(i_0_652, plain, (v7_ordinal1(k4_struct_0(X1))|~v6_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_657, plain, (v3_ordinal1(k4_struct_0(X1))|~v3_ordinal1(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_651, negated_conjecture, (v7_ordinal1(esk2_0)|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_653, plain, (v7_ordinal1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_926, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_relat_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_928, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)),k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v1_funct_2(esk4_0,k5_numbers,X2)|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_938, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_939, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X2)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_940, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),k8_nat_1(X2,esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X2)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_942, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_relat_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v1_funct_2(esk4_0,k5_numbers,X2)|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_912, plain, (v1_rat_1(np__1)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_913, plain, (v1_int_1(np__1)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_654, negated_conjecture, (v3_ordinal1(esk4_0)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_931, plain, (v1_xcmplx_0(np__1)|~v3_membered(k5_numbers))).#trainneg
cnf(i_0_935, negated_conjecture, (v7_ordinal1(esk2_0)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_936, plain, (v7_ordinal1(k5_numbers)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_655, negated_conjecture, (v3_ordinal1(esk3_0)|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_656, negated_conjecture, (v3_ordinal1(esk2_0)|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_946, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers))).#trainneg
cnf(i_0_947, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_948, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X2)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_949, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X3,esk4_0,X1)),k8_nat_1(X2,esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X2)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk4_0,k5_numbers,X3)|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_966, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers))).#trainneg
cnf(i_0_968, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_955, plain, (v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_956, negated_conjecture, (v1_xxreal_0(esk2_0)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_516, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_958, negated_conjecture, (~v3_xxreal_0(esk2_0)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_661, negated_conjecture, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~r2_hidden(X1,esk4_0))).#trainneg
cnf(i_0_982, plain, (v1_rat_1(np__1)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_961, plain, (~v3_xxreal_0(k5_numbers)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_959, plain, (v1_xxreal_0(k5_numbers)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_963, negated_conjecture, (v3_ordinal1(esk3_0)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_678, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_679, plain, (v1_xcmplx_0(k3_series_1(X1))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_972, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v7_ordinal1(X2)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X3)|~v1_funct_2(esk4_0,k5_numbers,X1)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_975, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v7_ordinal1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers))).#trainneg
cnf(i_0_977, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v1_funct_2(esk4_0,k5_numbers,X2)|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_996, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers)|~v1_xboole_0(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))))).#trainneg
cnf(i_0_997, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers)|~v1_xboole_0(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))))).#trainneg
cnf(i_0_1000, negated_conjecture, (k6_normsp_1(esk1_0,esk4_0)=k4_struct_0(esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_xboole_0(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers))).#trainneg
cnf(i_0_1003, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_xboole_0(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers))).#trainneg
cnf(i_0_506, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_983, plain, (v1_int_1(np__1)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_965, negated_conjecture, (v3_ordinal1(esk2_0)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_526, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_680, plain, (v1_xxreal_0(k3_series_1(X1))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1013, plain, (v7_ordinal1(np__0))).#trainneg
cnf(i_0_1018, plain, (~v3_xxreal_0(np__0))).#trainneg
cnf(i_0_932, plain, (v1_xxreal_0(k4_struct_0(X1))|~l2_struct_0(X1)|~v6_membered(u1_struct_0(X1)))).#trainneg
cnf(i_0_933, plain, (v1_xreal_0(k4_struct_0(X1))|~l2_struct_0(X1)|~v6_membered(u1_struct_0(X1)))).#trainneg
cnf(i_0_681, plain, (v1_xreal_0(k3_series_1(X1))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_682, plain, (v1_rat_1(k3_series_1(X1))|~v4_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_438, plain, (v1_funct_1(esk51_1(X1)))).#trainneg
cnf(i_0_473, plain, (v1_funct_1(esk60_1(X1)))).#trainneg
cnf(i_0_440, plain, (v1_relat_1(esk51_1(X1)))).#trainneg
cnf(i_0_475, plain, (v1_relat_1(esk60_1(X1)))).#trainneg
cnf(i_0_471, plain, (v4_valued_0(esk60_1(X1)))).#trainneg
cnf(i_0_437, plain, (v5_ordinal1(esk51_1(X1)))).#trainneg
cnf(i_0_501, plain, (k6_xcmplx_0(np__1,np__1)=np__0)).#trainneg
cnf(i_0_500, plain, (k6_xcmplx_0(np__1,np__0)=np__1)).#trainneg
cnf(i_0_512, plain, (m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_449, plain, (v5_relat_1(esk52_0,k5_numbers))).#trainneg
cnf(i_0_447, plain, (v1_partfun1(esk52_0,k5_numbers))).#trainneg
cnf(i_0_1014, plain, (v3_ordinal1(np__0))).#trainneg
cnf(i_0_1011, plain, (v1_rat_1(np__0)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_1012, plain, (v1_int_1(np__0)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_934, plain, (~v3_xxreal_0(k4_struct_0(X1))|~l2_struct_0(X1)|~v6_membered(u1_struct_0(X1)))).#trainneg
cnf(i_0_683, plain, (v1_int_1(k3_series_1(X1))|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1024, plain, (v1_rat_1(np__0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_1025, plain, (v1_int_1(np__0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_930, negated_conjecture, (v7_ordinal1(esk4_0)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_954, negated_conjecture, (v3_ordinal1(esk4_0)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_684, plain, (v7_ordinal1(k3_series_1(X1))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_685, plain, (v3_ordinal1(k3_series_1(X1))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_450, plain, (v4_relat_1(esk52_0,k5_numbers))).#trainneg
cnf(i_0_493, plain, (r1_xxreal_0(np__1,np__1))).#trainneg
cnf(i_0_490, plain, (r1_xxreal_0(np__0,np__1))).#trainneg
cnf(i_0_489, plain, (r1_xxreal_0(np__0,np__0))).#trainneg
cnf(i_0_487, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_499, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_420, plain, (m1_subset_1(esk44_0,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_496, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__1))).#trainneg
cnf(i_0_495, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__0))).#trainneg
cnf(i_0_545, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_276, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1034, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1042, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,esk44_0))).#trainneg
cnf(i_0_686, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k1_numbers))|~r2_hidden(X1,k3_series_1(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1038, plain, (v1_rat_1(esk44_0)|~v4_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1030, negated_conjecture, (~v3_xxreal_0(esk4_0)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_743, negated_conjecture, (v1_xcmplx_0(esk3_0)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1051, plain, (v1_rat_1(X1)|~v4_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1052, plain, (v1_int_1(X1)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_688, plain, (k8_nat_1(X1,X2,X3)=k8_nat_1(X4,X2,X3)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4))))).#trainneg
cnf(i_0_696, plain, (v1_xcmplx_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1068, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),X2,esk62_3(X1,X2,X3))),k8_nat_1(X4,X3,esk62_3(X1,X2,X3)))|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk62_3(X1,X2,X3))|~v1_funct_2(X3,k5_numbers,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4))))).#trainneg
cnf(i_0_1069, plain, (v3_normsp_1(X1,X2)|v2_struct_0(X2)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X2,k8_nat_1(u1_struct_0(X2),X1,esk62_3(X2,X1,X3))),k8_nat_1(X4,X3,esk62_3(X2,X1,X3)))|~v2_comseq_2(X3)|~v1_funct_2(X1,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)|~v7_ordinal1(esk62_3(X2,X1,X3))|~v1_funct_2(X3,k5_numbers,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4))))).#trainneg
cnf(i_0_1070, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)),k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v7_ordinal1(X1)|~v1_funct_2(esk4_0,k5_numbers,X2)|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_1078, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(X4,X2,esk62_3(X1,X2,X3))),k8_nat_1(k1_numbers,X3,esk62_3(X1,X2,X3)))|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk62_3(X1,X2,X3))|~v1_funct_2(X2,k5_numbers,X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4))))).#trainneg
cnf(i_0_1079, plain, (v3_normsp_1(X1,X2)|v2_struct_0(X2)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X2,k8_nat_1(X4,X1,esk62_3(X2,X1,X3))),k8_nat_1(k1_numbers,X3,esk62_3(X2,X1,X3)))|~v2_comseq_2(X3)|~v1_funct_2(X1,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)|~v7_ordinal1(esk62_3(X2,X1,X3))|~v1_funct_2(X1,k5_numbers,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4))))).#trainneg
cnf(i_0_1043, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk44_0))).#trainneg
cnf(i_0_1057, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_697, plain, (v1_xxreal_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1039, plain, (v1_int_1(esk44_0)|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_744, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_772, plain, (v1_xcmplx_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_796, negated_conjecture, (v1_xxreal_0(esk3_0)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_797, negated_conjecture, (v1_xxreal_0(esk2_0)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_698, plain, (v1_xreal_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_699, plain, (v1_rat_1(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v4_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_860, plain, (v1_xxreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_862, negated_conjecture, (v1_xreal_0(esk3_0)|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_700, plain, (v1_int_1(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_887, negated_conjecture, (v1_xreal_0(esk2_0)|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_888, plain, (v1_xreal_0(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_907, negated_conjecture, (v1_xreal_0(esk3_0)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_957, negated_conjecture, (v1_xreal_0(esk2_0)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_960, plain, (v1_xreal_0(k5_numbers)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_701, plain, (v7_ordinal1(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_702, plain, (v3_ordinal1(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1118, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),X2,esk62_3(X1,X2,X3))),k1_funct_1(X3,esk62_3(X1,X2,X3)))|~v2_comseq_2(X3)|~v7_ordinal1(esk62_3(X1,X2,X3))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,X4)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1122, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(X5,X2,esk62_3(X1,X2,X3))),k8_nat_1(X4,X3,esk62_3(X1,X2,X3)))|~v2_comseq_2(X3)|~v7_ordinal1(esk62_3(X1,X2,X3))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,X4)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,X5)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5))))).#trainneg
cnf(i_0_1127, plain, (v3_normsp_1(X1,X2)|v2_struct_0(X2)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X2,k8_nat_1(u1_struct_0(X2),X1,esk62_3(X2,X1,X3))),k1_funct_1(X3,esk62_3(X2,X1,X3)))|~v2_comseq_2(X3)|~v7_ordinal1(esk62_3(X2,X1,X3))|~v1_funct_2(X1,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,X4)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_962, negated_conjecture, (v3_ordinal1(esk3_0)|~v7_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_703, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k1_numbers))|v2_struct_0(X2)|~r2_hidden(X1,k4_normsp_0(X2,X3))|~l1_normsp_0(X2)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))))).#trainneg
cnf(i_0_964, negated_conjecture, (v3_ordinal1(esk2_0)|~v7_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1035, plain, (v1_xcmplx_0(esk44_0)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1036, plain, (v1_xxreal_0(esk44_0)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1048, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_705, plain, (v1_xcmplx_0(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_706, plain, (v1_xxreal_0(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_350, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_245, plain, (v1_xreal_0(k16_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_349, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_502, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),np__0)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_46, plain, (v7_struct_0(X1)|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_203, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_196, plain, (v2_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_220, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_1213, plain, (v3_valued_0(esk36_0))).#trainneg
cnf(i_0_227, plain, (v4_valued_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1037, plain, (v1_xreal_0(esk44_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1211, plain, (v3_valued_0(esk60_1(X1)))).#trainneg
cnf(i_0_1040, plain, (v7_ordinal1(esk44_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_707, plain, (v1_xreal_0(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1216, plain, (v7_ordinal1(esk44_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1219, plain, (~v3_xxreal_0(esk44_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1063, negated_conjecture, (v1_xcmplx_0(esk3_0)|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1220, negated_conjecture, (v1_xcmplx_0(esk3_0)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1214, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_708, plain, (v1_rat_1(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v4_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_709, plain, (v1_int_1(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_132, plain, (v3_ordinal1(X1)|~v1_ordinal1(X1)|~v2_ordinal1(X1))).#trainneg
cnf(i_0_185, plain, (v8_struct_0(X1)|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_202, plain, (v8_struct_0(X1)|~l1_struct_0(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_330, plain, (m1_subset_1(esk28_1(X1),X1))).#trainneg
cnf(i_0_439, plain, (v5_relat_1(esk51_1(X1),X1))).#trainneg
cnf(i_0_472, plain, (v1_partfun1(esk60_1(X1),X1))).#trainneg
cnf(i_0_474, plain, (v4_relat_1(esk60_1(X1),X1))).#trainneg
cnf(i_0_503, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))=np__0)).#trainneg
cnf(i_0_497, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1230, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk28_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_1207, plain, (v1_xreal_0(k18_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1217, plain, (v1_xxreal_0(esk44_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_710, plain, (v7_ordinal1(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1160, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1161, plain, (v1_xcmplx_0(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1247, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1248, plain, (v1_xcmplx_0(k5_numbers)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_711, plain, (v3_ordinal1(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_712, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k1_numbers))|v2_struct_0(X2)|~r2_hidden(X1,esk61_3(X2,X3,X4))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1240, plain, (k1_seq_2(X1)=np__0|r1_xxreal_0(esk10_2(X1,np__0),k18_complex1(k1_seq_1(X1,esk11_3(X1,np__0,X2))))|~v2_comseq_2(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_1(X1,esk11_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xcmplx_0(k1_seq_1(X1,esk11_3(X1,np__0,X2))))).#trainneg
cnf(i_0_1241, plain, (r1_xxreal_0(esk7_2(X1,np__0),k18_complex1(k1_seq_1(X1,esk8_3(X1,np__0,X2))))|v2_comseq_2(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_1(X1,esk8_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xcmplx_0(k1_seq_1(X1,esk8_3(X1,np__0,X2))))).#trainneg
cnf(i_0_1254, plain, (r1_xxreal_0(esk10_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|k1_seq_2(esk61_3(X1,X2,X3))=np__0|~r1_xxreal_0(esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)),esk11_3(esk61_3(X1,X2,X3),np__0,X4))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~m2_subset_1(esk11_3(esk61_3(X1,X2,X3),np__0,X4),k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_xcmplx_0(k1_seq_1(esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,X4)))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(k1_seq_1(esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,X4)),k1_numbers)|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1256, plain, (r1_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|v2_comseq_2(esk61_3(X1,X2,X3))|~r1_xxreal_0(esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)),esk8_3(esk61_3(X1,X2,X3),np__0,X4))|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~m2_subset_1(esk8_3(esk61_3(X1,X2,X3),np__0,X4),k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_xcmplx_0(k1_seq_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,X4)))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(k1_seq_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,X4)),k1_numbers)|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1259, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|r1_xxreal_0(esk10_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~v1_xcmplx_0(k1_seq_1(esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)))))|~m2_subset_1(esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m2_subset_1(esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)),k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)))),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1260, plain, (r1_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~v1_xcmplx_0(k1_seq_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)))))|~m2_subset_1(esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m2_subset_1(esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)),k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)))),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1049, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1231, plain, (v1_xcmplx_0(esk28_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_1223, plain, (m1_subset_1(esk28_1(esk44_0),k1_numbers))).#trainneg
cnf(i_0_714, plain, (v1_xcmplx_0(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1162, negated_conjecture, (v1_xxreal_0(esk3_0)|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1163, negated_conjecture, (v1_xxreal_0(esk2_0)|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1165, plain, (v1_xxreal_0(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_715, plain, (v1_xxreal_0(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_716, plain, (v1_xreal_0(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1270, plain, (v1_xreal_0(esk28_1(esk44_0)))).#trainneg
cnf(i_0_1274, plain, (v3_ordinal1(esk28_1(esk44_0))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_190, plain, (v1_xboole_0(X1)|v3_xxreal_0(X1)|v2_xxreal_0(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1281, plain, (v2_xxreal_0(esk59_0))).#trainneg
cnf(i_0_383, plain, (k18_complex1(k18_complex1(X1))=k18_complex1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_382, plain, (k16_complex1(k16_complex1(X1))=k16_complex1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1309, plain, (k16_complex1(k18_complex1(X1))=k18_complex1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_411, plain, (v2_struct_0(X1)|v1_zfmisc_1(esk41_1(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_110, plain, (v6_membered(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1050, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1232, plain, (v1_xxreal_0(esk28_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_1282, plain, (v2_xxreal_0(esk55_0)|v1_xboole_0(esk55_0))).#trainneg
cnf(i_0_717, plain, (v1_rat_1(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v4_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1288, plain, (v2_xxreal_0(esk39_0)|v3_xxreal_0(esk39_0)|v1_xboole_0(esk39_0))).#trainneg
cnf(i_0_1301, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X1)|v1_xboole_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_718, plain, (v1_int_1(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_719, plain, (v7_ordinal1(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1343, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers))).#trainneg
cnf(i_0_1352, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k5_numbers))).#trainneg
cnf(i_0_341, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_370, plain, (v1_zfmisc_1(u1_struct_0(X1))|~l1_struct_0(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_378, plain, (v1_finset_1(u1_struct_0(X1))|~l1_struct_0(X1)|~v8_struct_0(X1))).#trainneg
cnf(i_0_1346, plain, (v1_xxreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1314, plain, (v7_ordinal1(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_720, plain, (v3_ordinal1(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1315, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_1381, plain, (~v3_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_1347, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1317, plain, (v1_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_721, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))|v2_struct_0(X2)|~r2_hidden(X1,k1_bhsp_4(X2,X3))|~l2_algstr_0(X2)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))))).#trainneg
cnf(i_0_723, plain, (v1_xcmplx_0(k6_normsp_1(X1,X2))|v2_struct_0(X1)|~v1_membered(u1_struct_0(X1))|~l1_normsp_1(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_246, plain, (m1_subset_1(k18_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_219, plain, (v13_struct_0(X1,k1_xboole_0)|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_87, plain, (v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_209, plain, (v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_332, plain, (v9_struct_0(k4_struct_0(X1),X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_384, plain, (v9_struct_0(esk30_1(X1),X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_346, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_412, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(esk41_1(X1)))).#trainneg
cnf(i_0_461, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(esk56_1(X1)))).#trainneg
cnf(i_0_368, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v1_zfmisc_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_414, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v1_zfmisc_1(esk42_1(X1)))).#trainneg
cnf(i_0_335, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1403, plain, (v2_struct_0(esk27_0)|~v1_xboole_0(u1_struct_0(esk27_0)))).#trainneg
cnf(i_0_1394, plain, (v5_relat_1(esk52_0,k4_numbers))).#trainneg
cnf(i_0_1421, plain, (~v2_xxreal_0(np__0))).#trainneg
cnf(i_0_724, plain, (v1_xxreal_0(k6_normsp_1(X1,X2))|v2_struct_0(X1)|~v2_membered(u1_struct_0(X1))|~l1_normsp_1(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1395, plain, (v5_relat_1(esk36_0,k4_numbers))).#trainneg
cnf(i_0_1401, plain, (~v1_xboole_0(u1_struct_0(esk49_0)))).#trainneg
cnf(i_0_1402, plain, (~v1_xboole_0(u1_struct_0(esk40_0)))).#trainneg
cnf(i_0_1318, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_1320, negated_conjecture, (v7_ordinal1(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_725, plain, (v1_xreal_0(k6_normsp_1(X1,X2))|v2_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l1_normsp_1(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_726, plain, (v1_rat_1(k6_normsp_1(X1,X2))|v2_struct_0(X1)|~v4_membered(u1_struct_0(X1))|~l1_normsp_1(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_333, plain, (v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v3_xxreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_380, plain, (v8_struct_0(X1)|~l1_struct_0(X1)|~v1_finset_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_488, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1436, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1472, plain, (r1_xxreal_0(esk11_3(X1,X2,X3),esk11_3(X1,X2,X3))|X2=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1473, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),esk8_3(X1,X2,X3))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1486, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1487, plain, (r1_xxreal_0(X1,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_389, plain, (v1_funct_1(esk32_2(X1,X2)))).#trainneg
cnf(i_0_392, plain, (v1_relat_1(esk32_2(X1,X2)))).#trainneg
cnf(i_0_399, plain, (v1_relat_1(esk35_2(X1,X2)))).#trainneg
cnf(i_0_1404, plain, (v2_struct_0(esk26_0)|~v1_xboole_0(u1_struct_0(esk26_0)))).#trainneg
cnf(i_0_1398, plain, (v5_relat_1(esk52_0,k3_numbers))).#trainneg
cnf(i_0_1399, plain, (v5_relat_1(esk36_0,k3_numbers))).#trainneg
cnf(i_0_727, plain, (v1_int_1(k6_normsp_1(X1,X2))|v2_struct_0(X1)|~v5_membered(u1_struct_0(X1))|~l1_normsp_1(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1463, plain, (r1_xxreal_0(esk59_0,esk59_0))).#trainneg
cnf(i_0_1419, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1430, negated_conjecture, (~v3_xxreal_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1321, negated_conjecture, (v7_ordinal1(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1324, negated_conjecture, (v1_xcmplx_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_728, plain, (v7_ordinal1(k6_normsp_1(X1,X2))|v2_struct_0(X1)|~v6_membered(u1_struct_0(X1))|~l1_normsp_1(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_729, plain, (v3_ordinal1(k6_normsp_1(X1,X2))|v2_struct_0(X1)|~v3_ordinal1(u1_struct_0(X1))|~l1_normsp_1(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_400, plain, (v1_xboole_0(esk35_2(X1,X2)))).#trainneg
cnf(i_0_29, plain, (v2_struct_0(X1)|v13_struct_0(X1,np__1)|~l1_struct_0(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_27, plain, (v5_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_31, plain, (v6_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_217, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1535, plain, (v1_xreal_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_224, plain, (v4_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_226, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v13_struct_0(X1,k1_xboole_0))).#trainneg
cnf(i_0_32, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v13_struct_0(X1,np__1))).#trainneg
cnf(i_0_172, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_187, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k3_numbers))).#trainneg
cnf(i_0_1405, plain, (v2_struct_0(esk25_0)|~v1_xboole_0(u1_struct_0(esk25_0)))).#trainneg
cnf(i_0_1464, plain, (r1_xxreal_0(esk55_0,esk55_0))).#trainneg
cnf(i_0_1465, plain, (r1_xxreal_0(esk54_0,esk54_0))).#trainneg
cnf(i_0_739, negated_conjecture, (v1_xcmplx_0(esk4_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1466, plain, (r1_xxreal_0(esk53_0,esk53_0))).#trainneg
cnf(i_0_1505, negated_conjecture, (~v3_xxreal_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1325, negated_conjecture, (v1_xcmplx_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1326, negated_conjecture, (v1_xxreal_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_750, plain, (r1_xxreal_0(X1,np__0)|~r1_xxreal_0(X1,k18_complex1(k6_xcmplx_0(k1_seq_1(X2,X3),esk5_1(X2))))|~r1_xxreal_0(esk6_2(X2,X1),X3)|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(esk5_1(X2))|~m1_subset_1(k1_seq_1(X2,X3),k1_numbers))).#trainneg
cnf(i_0_758, plain, (r1_xxreal_0(X1,np__0)|X2!=k1_seq_2(X3)|~r1_xxreal_0(X1,k18_complex1(k6_xcmplx_0(k1_seq_1(X3,X4),X2)))|~r1_xxreal_0(esk9_3(X3,X2,X1),X4)|~v2_comseq_2(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X4),k1_numbers))).#trainneg
cnf(i_0_1551, negated_conjecture, (r1_xxreal_0(esk3_0,esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_146, plain, (v4_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_82, plain, (v1_seq_2(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_comseq_2(X1))).#trainneg
cnf(i_0_81, plain, (v2_seq_2(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_comseq_2(X1))).#trainneg
cnf(i_0_96, plain, (v5_valued_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_valued_0(X1)|~v1_zfmisc_1(X1))).#trainneg
cnf(i_0_95, plain, (v6_valued_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_valued_0(X1)|~v1_zfmisc_1(X1))).#trainneg
cnf(i_0_101, plain, (v7_valued_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_valued_0(X1)|~v5_valued_0(X1))).#trainneg
cnf(i_0_106, plain, (v8_valued_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_valued_0(X1)|~v6_valued_0(X1))).#trainneg
cnf(i_0_385, plain, (m1_subset_1(esk30_1(X1),u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_33, plain, (~v2_struct_0(X1)|~l1_struct_0(X1)|~v13_struct_0(X1,np__1))).#trainneg
cnf(i_0_372, plain, (v1_valued_0(k2_zfmisc_1(X2,X1))|~v1_membered(X1))).#trainneg
cnf(i_0_373, plain, (v2_valued_0(k2_zfmisc_1(X2,X1))|~v2_membered(X1))).#trainneg
cnf(i_0_1406, plain, (v2_struct_0(esk24_0)|~v1_xboole_0(u1_struct_0(esk24_0)))).#trainneg
cnf(i_0_1467, plain, (r1_xxreal_0(esk48_0,esk48_0))).#trainneg
cnf(i_0_1468, plain, (r1_xxreal_0(esk47_0,esk47_0))).#trainneg
cnf(i_0_759, plain, (v1_xcmplx_0(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1470, plain, (r1_xxreal_0(esk39_0,esk39_0))).#trainneg
cnf(i_0_1200, plain, (v1_xcmplx_0(esk44_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1327, negated_conjecture, (v1_xxreal_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1344, negated_conjecture, (v3_ordinal1(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_760, plain, (v1_xxreal_0(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_761, plain, (v1_xreal_0(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1579, negated_conjecture, (r1_xxreal_0(esk2_0,esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_374, plain, (v3_valued_0(k2_zfmisc_1(X2,X1))|~v3_membered(X1))).#trainneg
cnf(i_0_379, plain, (v4_valued_0(k2_zfmisc_1(X2,X1))|~v6_membered(X1))).#trainneg
cnf(i_0_544, plain, (v3_xxreal_0(X2)|v2_xxreal_0(X1)|r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1601, plain, (r1_xxreal_0(X1,k1_seq_2(X2))|v2_xxreal_0(X1)|v3_xxreal_0(k1_seq_2(X2))|~v1_xreal_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_387, plain, (v7_struct_0(X1)|m1_subset_1(esk31_1(X1),u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_390, plain, (v5_relat_1(esk32_2(X1,X2),X2))).#trainneg
cnf(i_0_397, plain, (v5_relat_1(esk35_2(X1,X2),X2))).#trainneg
cnf(i_0_391, plain, (v4_relat_1(esk32_2(X1,X2),X1))).#trainneg
cnf(i_0_398, plain, (v4_relat_1(esk35_2(X1,X2),X1))).#trainneg
cnf(i_0_140, plain, (v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_1407, plain, (v2_struct_0(esk23_0)|~v1_xboole_0(u1_struct_0(esk23_0)))).#trainneg
cnf(i_0_1513, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1316, plain, (v7_ordinal1(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_762, plain, (v1_rat_1(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v4_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1485, plain, (r1_xxreal_0(k5_numbers,k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_1642, plain, (~v3_xxreal_0(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_1345, negated_conjecture, (v3_ordinal1(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_763, plain, (v1_int_1(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_764, plain, (v7_ordinal1(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1632, plain, (v1_xreal_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_135, plain, (v1_comseq_2(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_seq_2(X1)|~v2_seq_2(X1))).#trainneg
cnf(i_0_229, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_100, plain, (v5_membered(X2)|~v5_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_28, plain, (v6_ordinal1(X2)|~v6_ordinal1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_105, plain, (v6_membered(X2)|~v6_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_50, plain, (v1_membered(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_90, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_92, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1729, plain, (v3_membered(esk44_0))).#trainneg
cnf(i_0_94, plain, (v4_membered(X2)|~v4_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_386, plain, (v7_struct_0(X1)|~l2_struct_0(X1)|~v9_struct_0(esk31_1(X1),X1))).#trainneg
cnf(i_0_1408, plain, (v2_struct_0(esk22_0)|~v1_xboole_0(u1_struct_0(esk22_0)))).#trainneg
cnf(i_0_1736, plain, (v3_membered(k5_numbers))).#trainneg
cnf(i_0_1679, plain, (v5_membered(esk44_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_765, plain, (v3_ordinal1(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1686, plain, (v5_membered(k5_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1689, plain, (v6_ordinal1(esk44_0)|~v6_ordinal1(k1_numbers))).#trainneg
cnf(i_0_1696, plain, (v6_ordinal1(k5_numbers)|~v6_ordinal1(k1_numbers))).#trainneg
cnf(i_0_1749, plain, (v5_membered(k5_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_766, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))|v2_struct_0(X2)|~r2_hidden(X1,k8_lopban_4(X2,X3,X4,X5))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,k5_numbers)|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_773, negated_conjecture, (v1_xxreal_0(esk4_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_371, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_26, plain, (~r2_hidden(X2,X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_376, plain, (v5_relat_1(k2_zfmisc_1(X2,X1),k4_numbers)|~v5_membered(X1))).#trainneg
cnf(i_0_375, plain, (v5_relat_1(k2_zfmisc_1(X2,X1),k3_numbers)|~v4_membered(X1))).#trainneg
cnf(i_0_535, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_529, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(X2)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_527, plain, (v2_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_1871, plain, (v2_xxreal_0(esk11_3(X1,X2,X3))|X2=k1_seq_2(X1)|~v2_xxreal_0(X3)|~v1_xreal_0(esk11_3(X1,X2,X3))|~v1_xreal_0(X3)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1409, plain, (v2_struct_0(esk21_0)|~v1_xboole_0(u1_struct_0(esk21_0)))).#trainneg
cnf(i_0_1759, plain, (v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1699, plain, (v6_membered(esk44_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_780, plain, (r1_xxreal_0(esk10_2(X1,esk5_1(X1)),np__0)|esk5_1(X1)=k1_seq_2(X1)|~r1_xxreal_0(esk6_2(X1,esk10_2(X1,esk5_1(X1))),esk11_3(X1,esk5_1(X1),X2))|~v2_comseq_2(X1)|~v1_xreal_0(esk10_2(X1,esk5_1(X1)))|~m2_subset_1(esk11_3(X1,esk5_1(X1),X2),k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(esk5_1(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1716, plain, (v1_membered(k5_numbers)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1850, plain, (~v2_xxreal_0(k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1878, plain, (v1_membered(k5_numbers))).#trainneg
cnf(i_0_1726, plain, (v2_membered(k5_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1881, plain, (v2_membered(k5_numbers))).#trainneg
cnf(i_0_774, plain, (r1_xxreal_0(k1_normsp_0(X1,k1_funct_1(k1_bhsp_4(X1,X2),X3)),k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|v2_struct_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X3)|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,X2))|~m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_784, plain, (v3_normsp_1(X1,X2)|v2_struct_0(X2)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X2,k1_funct_1(X1,esk62_3(X2,X1,X3))),k8_nat_1(k1_numbers,X3,esk62_3(X2,X1,X3)))|~v2_comseq_2(X3)|~v1_funct_2(X1,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)|~v7_ordinal1(esk62_3(X2,X1,X3)))).#trainneg
cnf(i_0_1875, plain, (X2=k1_seq_2(X1)|~v1_xboole_0(esk11_3(X1,X2,X3))|~v2_comseq_2(X1)|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1877, plain, (esk5_1(X1)=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(esk5_1(X1))|~m2_subset_1(esk6_2(X1,esk10_2(X1,esk5_1(X1))),k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_413, plain, (v2_struct_0(X1)|m1_subset_1(esk41_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_462, plain, (v2_struct_0(X1)|m1_subset_1(esk56_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_415, plain, (v7_struct_0(X1)|m1_subset_1(esk42_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_539, plain, (v1_xboole_0(X2)|v3_xxreal_0(X1)|v2_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_542, plain, (v1_xboole_0(X1)|v3_xxreal_0(X1)|v2_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_30, plain, (v1_valued_0(X2)|~v1_relat_1(X1)|~v1_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_34, plain, (v2_valued_0(X2)|~v1_relat_1(X1)|~v2_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1410, plain, (v2_struct_0(esk20_0)|~v1_xboole_0(u1_struct_0(esk20_0)))).#trainneg
cnf(i_0_1544, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_827, plain, (r1_xxreal_0(X1,np__0)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k8_nat_1(X4,X2,X3),esk5_1(X2))))|~r1_xxreal_0(esk6_2(X2,X1),X3)|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4))))).#trainneg
cnf(i_0_1746, plain, (v4_membered(k5_numbers)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_1709, plain, (v1_membered(esk44_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1719, plain, (v2_membered(esk44_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_2100, plain, (v4_membered(k5_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_785, plain, (v3_normsp_1(k1_bhsp_4(X1,X2),X1)|v2_struct_0(X1)|k2_seq_2(k3_series_1(k4_normsp_0(X1,X2)))!=np__0|~v2_comseq_2(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k3_series_1(k4_normsp_0(X1,X2)),k5_numbers,k1_numbers)|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k1_bhsp_4(X1,X2))|~m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k3_series_1(k4_normsp_0(X1,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m2_subset_1(esk62_3(X1,k1_bhsp_4(X1,X2),k3_series_1(k4_normsp_0(X1,X2))),k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_786, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X1,k1_funct_1(X2,esk62_3(X1,X2,X3))),k8_nat_1(k1_numbers,X3,esk62_3(X1,X2,X3)))|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk62_3(X1,X2,X3)))).#trainneg
cnf(i_0_2083, plain, (v1_xreal_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_36, plain, (v3_valued_0(X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_42, plain, (v4_valued_0(X2)|~v1_relat_1(X1)|~v4_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_358, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_359, plain, (v1_xxreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_valued_0(X1))).#trainneg
cnf(i_0_360, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_363, plain, (v7_ordinal1(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_520, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_344, plain, (v3_xxreal_0(X2)|v3_xxreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_345, plain, (v3_xxreal_0(X2)|v2_xxreal_0(k6_xcmplx_0(X2,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_1411, plain, (v2_struct_0(esk19_0)|~v1_xboole_0(u1_struct_0(esk19_0)))).#trainneg
cnf(i_0_1393, plain, (v5_relat_1(esk60_1(X1),k4_numbers))).#trainneg
cnf(i_0_828, plain, (r1_xxreal_0(X1,np__0)|~r1_xxreal_0(X1,k18_complex1(k6_xcmplx_0(k1_funct_1(X2,X3),esk5_1(X2))))|~r1_xxreal_0(esk6_2(X2,X1),X3)|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(esk5_1(X2))|~m1_subset_1(k1_funct_1(X2,X3),k1_numbers))).#trainneg
cnf(i_0_1739, plain, (v4_membered(esk44_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_1201, plain, (v1_xxreal_0(esk44_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_787, plain, (k6_normsp_1(X1,k1_bhsp_4(X1,X2))=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(k3_series_1(k4_normsp_0(X1,X2)))!=np__0|~v2_comseq_2(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k3_series_1(k4_normsp_0(X1,X2)),k5_numbers,k1_numbers)|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k1_bhsp_4(X1,X2))|~m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k3_series_1(k4_normsp_0(X1,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m2_subset_1(esk62_3(X1,k1_bhsp_4(X1,X2),k3_series_1(k4_normsp_0(X1,X2))),k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_826, negated_conjecture, (r1_xxreal_0(X1,np__0)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2),esk5_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))))))|~r1_xxreal_0(esk6_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),X1),X2)|~v2_comseq_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)))|~v1_xreal_0(X1)|~v3_valued_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v7_ordinal1(X2)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k5_numbers,u1_struct_0(esk1_0))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_343, plain, (v2_xxreal_0(X2)|v3_xxreal_0(k6_xcmplx_0(X2,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_339, plain, (v2_xxreal_0(X2)|v2_xxreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_453, plain, (m1_subset_1(esk52_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_253, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_540, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_361, plain, (v1_rat_1(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k3_numbers))).#trainneg
cnf(i_0_362, plain, (v1_int_1(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_285, plain, (m1_subset_1(k9_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_337, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(k6_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_338, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(k6_xcmplx_0(X2,X1)))).#trainneg
cnf(i_0_1412, plain, (v2_struct_0(esk18_0)|~v1_xboole_0(u1_struct_0(esk18_0)))).#trainneg
cnf(i_0_1397, plain, (v5_relat_1(esk60_1(X1),k3_numbers))).#trainneg
cnf(i_0_2204, plain, (v3_xxreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_2296, plain, (~m1_subset_1(k4_xcmplx_0(np__1),k5_numbers))).#trainneg
cnf(i_0_861, negated_conjecture, (v1_xreal_0(esk4_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2057, plain, (v1_valued_0(k5_numbers)|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2302, plain, (v1_valued_0(k5_numbers)|~v1_relat_1(k1_numbers)|~v3_valued_0(k1_numbers))).#trainneg
cnf(i_0_863, negated_conjecture, (r1_xxreal_0(X1,np__0)|X2!=k1_seq_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)))|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X3),X2)))|~r1_xxreal_0(esk9_3(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)),X2,X1),X3)|~v2_comseq_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v3_valued_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)),k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v7_ordinal1(X3)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)),k5_numbers,u1_struct_0(esk1_0))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_864, plain, (r1_xxreal_0(X1,np__0)|X2!=k1_seq_2(X3)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k8_nat_1(X5,X3,X4),X2)))|~r1_xxreal_0(esk9_3(X3,X2,X1),X4)|~v2_comseq_2(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v3_valued_0(X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,X5)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5))))).#trainneg
cnf(i_0_40, plain, (v5_relat_1(X2,k4_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2312, plain, (v5_relat_1(esk61_3(X1,X2,X3),k4_numbers)|v2_struct_0(X1)|~v5_relat_1(k2_zfmisc_1(k5_numbers,k1_numbers),k4_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_38, plain, (v5_relat_1(X2,k3_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k3_numbers)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2326, plain, (v5_relat_1(esk61_3(X1,X2,X3),k3_numbers)|v2_struct_0(X1)|~v5_relat_1(k2_zfmisc_1(k5_numbers,k1_numbers),k3_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_262, plain, (v1_funct_1(k2_series_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_264, plain, (v1_relat_1(k2_series_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_260, plain, (v1_valued_0(k2_series_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_2342, plain, (v1_valued_0(k3_series_1(X1))|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_388, plain, (v1_funct_2(esk32_2(X1,X2),X1,X2))).#trainneg
cnf(i_0_261, plain, (v1_partfun1(k2_series_1(X1),k5_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_1413, plain, (v2_struct_0(esk17_0)|~v1_xboole_0(u1_struct_0(esk17_0)))).#trainneg
cnf(i_0_1510, plain, (v1_xreal_0(esk35_2(X1,X2)))).#trainneg
cnf(i_0_1894, plain, (esk5_1(X1)=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(esk5_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk6_2(X1,esk10_2(X1,esk5_1(X1))),k5_numbers))).#trainneg
cnf(i_0_865, plain, (r1_xxreal_0(X1,np__0)|X2!=k1_seq_2(X3)|~r1_xxreal_0(X1,k18_complex1(k6_xcmplx_0(k1_funct_1(X3,X4),X2)))|~r1_xxreal_0(esk9_3(X3,X2,X1),X4)|~v2_comseq_2(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v3_valued_0(X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_funct_1(X3,X4),k1_numbers))).#trainneg
cnf(i_0_2324, plain, (v5_relat_1(k5_numbers,k4_numbers)|~v5_relat_1(k1_numbers,k4_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2297, plain, (~v1_xboole_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1215, plain, (v1_xreal_0(esk44_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1880, plain, (v1_valued_0(k2_zfmisc_1(X1,k5_numbers)))).#trainneg
cnf(i_0_870, plain, (r1_xxreal_0(esk7_2(X1,X2),k18_complex1(k9_real_1(k8_nat_1(X4,X1,esk8_3(X1,X2,X3)),X2)))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v7_ordinal1(esk8_3(X1,X2,X3))|~v1_funct_2(X1,k5_numbers,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4))))).#trainneg
cnf(i_0_2343, plain, (v1_partfun1(k3_series_1(X1),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_263, plain, (v4_relat_1(k2_series_1(X1),k5_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_2365, plain, (v4_relat_1(k3_series_1(X1),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_251, plain, (v2_struct_0(X1)|m1_subset_1(k1_normsp_0(X1,X2),k1_numbers)|~l1_normsp_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2368, plain, (v1_xreal_0(k1_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_257, plain, (v2_struct_0(X1)|v1_funct_1(k2_normsp_0(X1,X2))|~v1_funct_1(X2)|~v1_relat_1(X2)|~l1_normsp_0(X1)|~v5_relat_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_258, plain, (v2_struct_0(X1)|v1_relat_1(k2_normsp_0(X1,X2))|~v1_funct_1(X2)|~v1_relat_1(X2)|~l1_normsp_0(X1)|~v5_relat_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_393, plain, (m1_subset_1(esk32_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_401, plain, (m1_subset_1(esk35_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_331, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m2_subset_1(esk29_2(X1,X2),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1414, plain, (v2_struct_0(esk16_0)|~v1_xboole_0(u1_struct_0(esk16_0)))).#trainneg
cnf(i_0_1640, plain, (v1_xxreal_0(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_875, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk10_2(X2,X1),k18_complex1(k9_real_1(k8_nat_1(X4,X2,esk11_3(X2,X1,X3)),X1)))|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v7_ordinal1(esk11_3(X2,X1,X3))|~v1_funct_2(X2,k5_numbers,X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4))))).#trainneg
cnf(i_0_2338, plain, (v5_relat_1(k5_numbers,k3_numbers)|~v5_relat_1(k1_numbers,k3_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1218, plain, (v1_xreal_0(esk44_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1377, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X1)|v1_xboole_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1488, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_879, plain, (r1_xxreal_0(k1_normsp_0(X1,k1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)),k1_seq_1(esk61_3(X1,X2,X3),X4))|v2_struct_0(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v1_funct_2(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)))|~m1_subset_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2439, plain, (r1_xxreal_0(esk44_0,esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_2486, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v7_ordinal1(X1)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2487, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2516, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2517, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v7_ordinal1(X1)|~v1_funct_2(esk4_0,k5_numbers,X2)|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_1415, plain, (v2_struct_0(esk15_0)|~v1_xboole_0(u1_struct_0(esk15_0)))).#trainneg
cnf(i_0_1641, plain, (v1_xreal_0(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_903, plain, (r1_xxreal_0(esk7_2(X1,X2),k18_complex1(k6_xcmplx_0(k8_nat_1(X4,X1,esk8_3(X1,X2,X3)),X2)))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k8_nat_1(X4,X1,esk8_3(X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v7_ordinal1(esk8_3(X1,X2,X3))|~v1_funct_2(X1,k5_numbers,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4))))).#trainneg
cnf(i_0_2450, plain, (v5_relat_1(k5_numbers,k3_numbers)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k4_numbers))).#trainneg
cnf(i_0_904, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk10_2(X2,X1),k18_complex1(k6_xcmplx_0(k8_nat_1(X4,X2,esk11_3(X2,X1,X3)),X1)))|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(k8_nat_1(X4,X2,esk11_3(X2,X1,X3)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v7_ordinal1(esk11_3(X2,X1,X3))|~v1_funct_2(X2,k5_numbers,X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4))))).#trainneg
cnf(i_0_916, negated_conjecture, (v1_rat_1(esk4_0)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_44, plain, (v1_valued_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_47, plain, (v2_valued_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_2582, plain, (v3_valued_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2583, plain, (v3_valued_0(k3_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_93, plain, (v4_valued_0(X2)|~v6_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_181, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_171, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_1416, plain, (v2_struct_0(esk14_0)|~v1_xboole_0(u1_struct_0(esk14_0)))).#trainneg
cnf(i_0_1360, plain, (v1_xboole_0(u1_struct_0(esk27_0))|~v2_struct_0(esk27_0))).#trainneg
cnf(i_0_921, negated_conjecture, (v1_int_1(esk4_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2315, plain, (v5_relat_1(esk44_0,k4_numbers)|~v5_relat_1(k1_numbers,k4_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1417, plain, (v2_struct_0(esk13_0)|~v1_xboole_0(u1_struct_0(esk13_0)))).#trainneg
cnf(i_0_1418, plain, (v2_struct_0(esk12_0)|~v1_xboole_0(u1_struct_0(esk12_0)))).#trainneg
cnf(i_0_1563, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_922, plain, (r1_xxreal_0(X1,np__0)|v2_struct_0(X2)|~r1_xxreal_0(X1,k18_complex1(k8_nat_1(X6,esk61_3(X2,X3,X4),X5)))|~r1_xxreal_0(esk63_4(X2,X3,X4,X1),X5)|~v1_xreal_0(X1)|~v3_valued_0(esk61_3(X2,X3,X4))|~m2_subset_1(X5,k1_numbers,k5_numbers)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X5)|~v1_funct_2(esk61_3(X2,X3,X4),k5_numbers,X6)|~v1_funct_1(esk61_3(X2,X3,X4))|~m1_subset_1(esk61_3(X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X6))))).#trainneg
cnf(i_0_953, negated_conjecture, (v3_ordinal1(esk4_0)|~v7_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2635, plain, (v5_relat_1(esk61_3(X1,X2,X3),k1_numbers)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2640, negated_conjecture, (v5_relat_1(esk4_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2643, plain, (v5_relat_1(k1_bhsp_4(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2644, plain, (v5_relat_1(k8_lopban_4(X1,X2,X3,X4),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_134, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2646, plain, (v4_relat_1(esk61_3(X1,X2,X3),k5_numbers)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_59, plain, (v2_struct_0(X1)|v2_vectsp_1(X1)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1))).#trainneg
cnf(i_0_57, plain, (v2_struct_0(X1)|v6_vectsp_1(X1)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1))).#trainneg
cnf(i_0_56, plain, (v2_struct_0(X1)|v3_lopban_1(X1)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1))).#trainneg
cnf(i_0_55, plain, (v2_struct_0(X1)|v2_lopban_2(X1)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1))).#trainneg
cnf(i_0_54, plain, (v2_struct_0(X1)|v3_lopban_2(X1)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1))).#trainneg
cnf(i_0_2622, plain, (X1=u1_struct_0(esk27_0)|~v1_xboole_0(X1)|~v2_struct_0(esk27_0))).#trainneg
cnf(i_0_2651, negated_conjecture, (v4_relat_1(esk4_0,k5_numbers))).#trainneg
cnf(i_0_1750, plain, (v5_membered(k5_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_986, negated_conjecture, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|v1_xboole_0(esk4_0)|~m1_subset_1(X1,esk4_0))).#trainneg
cnf(i_0_2329, plain, (v5_relat_1(esk44_0,k3_numbers)|~v5_relat_1(k1_numbers,k3_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1576, plain, (v1_xcmplx_0(esk44_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2191, plain, (v1_xxreal_0(esk44_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1484, plain, (r1_xxreal_0(k5_numbers,k5_numbers)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1514, plain, (v7_ordinal1(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_988, plain, (v1_xboole_0(esk61_3(X1,X2,X3))|v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|m1_subset_1(X4,k2_zfmisc_1(k5_numbers,k1_numbers))|v2_struct_0(X1)|~m1_subset_1(X4,esk61_3(X1,X2,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_53, plain, (v2_struct_0(X1)|v4_lopban_2(X1)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1))).#trainneg
cnf(i_0_150, plain, (v2_struct_0(X1)|v4_vectsp_1(X1)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~v5_lopban_2(X1)|~l1_lopban_2(X1))).#trainneg
cnf(i_0_91, plain, (v5_relat_1(X2,k4_numbers)|~v5_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_51, plain, (v5_relat_1(X2,k3_numbers)|~v4_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_113, plain, (v2_struct_0(X1)|v5_lopban_2(X1)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_normsp_1(X1)|~v2_funcsdom(X1)|~v3_group_1(X1)|~v1_vectsp_1(X1)|~v3_vectsp_1(X1)|~l1_lopban_2(X1)|~v2_vectsp_1(X1)|~v6_vectsp_1(X1)|~v3_lopban_1(X1)|~v2_lopban_2(X1)|~v3_lopban_2(X1)|~v4_lopban_2(X1))).#trainneg
cnf(i_0_52, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_182, plain, (v1_comseq_2(X1)|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1053, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1361, plain, (v1_xboole_0(u1_struct_0(esk26_0))|~v2_struct_0(esk26_0))).#trainneg
cnf(i_0_2666, plain, (v1_xcmplx_0(esk44_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2620, plain, (u1_struct_0(esk27_0)=np__0|~v2_struct_0(esk27_0))).#trainneg
cnf(i_0_990, plain, (v1_xboole_0(k4_normsp_0(X1,X2))|v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|m1_subset_1(X3,k2_zfmisc_1(k5_numbers,k1_numbers))|v2_struct_0(X1)|~m1_subset_1(X3,k4_normsp_0(X1,X2))|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2070, plain, (v2_valued_0(k5_numbers)|~v2_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2678, plain, (~v3_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2714, plain, (v7_ordinal1(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2726, plain, (v1_xxreal_0(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_991, plain, (v1_xboole_0(k3_series_1(X1))|v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|m1_subset_1(X2,k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X2,k3_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_992, plain, (v1_xboole_0(k1_bhsp_4(X1,X2))|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|m1_subset_1(X3,k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X3,k1_bhsp_4(X1,X2))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_237, plain, (v1_xreal_0(esk5_1(X1))|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_259, plain, (m1_subset_1(k2_seq_2(X1),k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2738, plain, (v1_xreal_0(k2_seq_2(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2744, plain, (r1_xxreal_0(X1,k2_seq_2(X2))|v2_xxreal_0(X1)|v3_xxreal_0(k2_seq_2(X2))|~v1_xreal_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_178, plain, (v1_partfun1(X1,X2)|~v1_funct_2(X1,X2,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2))))).#trainneg
cnf(i_0_148, plain, (v1_xboole_0(X1)|v1_partfun1(X2,X3)|~v1_funct_2(X2,X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_2758, plain, (v1_partfun1(esk61_3(X1,X2,X3),k5_numbers)|v2_struct_0(X1)|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_112, plain, (v1_partfun1(X2,X1)|~v1_xboole_0(X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_215, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|~v1_funct_1(X1)|~v1_xboole_0(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2780, plain, (v2_struct_0(X1)|~v1_xboole_0(esk61_3(X1,X2,X3))|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2719, plain, (X1=u1_struct_0(esk26_0)|~v1_xboole_0(X1)|~v2_struct_0(esk26_0))).#trainneg
cnf(i_0_1362, plain, (v1_xboole_0(u1_struct_0(esk25_0))|~v2_struct_0(esk25_0))).#trainneg
cnf(i_0_2791, plain, (u1_struct_0(esk25_0)=np__0|~v2_struct_0(esk25_0))).#trainneg
cnf(i_0_993, plain, (v1_xboole_0(k8_lopban_4(X1,X2,X3,X4))|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|m1_subset_1(X5,k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X5,k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2785, negated_conjecture, (~v1_xboole_0(esk4_0))).#trainneg
cnf(i_0_2559, negated_conjecture, (v1_valued_0(esk4_0)|~v1_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2722, plain, (v2_valued_0(k5_numbers)|~v1_relat_1(k1_numbers)|~v3_valued_0(k1_numbers))).#trainneg
cnf(i_0_2715, plain, (v7_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_1020, plain, (v1_xreal_0(k3_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1028, negated_conjecture, (v1_xxreal_0(esk4_0)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2857, negated_conjecture, (m1_subset_1(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(X1)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2858, negated_conjecture, (m1_subset_1(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2),X1)|~v7_ordinal1(X2)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k5_numbers,X1)|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k5_numbers,u1_struct_0(esk1_0))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2862, negated_conjecture, (m1_subset_1(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~l2_algstr_0(esk1_0)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2906, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_funct_1(X3,X4))|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X2)))))).#trainneg
cnf(i_0_667, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_694, plain, (m1_subset_1(esk6_2(X1,X2),k5_numbers)|r1_xxreal_0(X2,np__0)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_695, plain, (m1_subset_1(esk6_2(X1,X2),k1_numbers)|r1_xxreal_0(X2,np__0)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_730, plain, (m1_subset_1(esk9_3(X1,X2,X3),k5_numbers)|r1_xxreal_0(X3,np__0)|X2!=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2919, plain, (m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~r2_hidden(X1,k1_funct_1(esk32_2(k5_numbers,k1_zfmisc_1(X2)),X3)))).#trainneg
cnf(i_0_1363, plain, (v1_xboole_0(u1_struct_0(esk24_0))|~v2_struct_0(esk24_0))).#trainneg
cnf(i_0_2763, negated_conjecture, (v1_partfun1(esk4_0,k5_numbers))).#trainneg
cnf(i_0_2717, plain, (u1_struct_0(esk26_0)=np__0|~v2_struct_0(esk26_0))).#trainneg
cnf(i_0_1029, negated_conjecture, (v1_xreal_0(esk4_0)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2123, plain, (v3_valued_0(k5_numbers)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2570, negated_conjecture, (v2_valued_0(esk4_0)|~v2_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2581, negated_conjecture, (v3_valued_0(esk4_0)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2800, plain, (v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_1031, plain, (v1_xxreal_0(k3_series_1(X1))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_731, plain, (m1_subset_1(esk9_3(X1,X2,X3),k1_numbers)|r1_xxreal_0(X3,np__0)|X2!=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_755, plain, (m1_subset_1(esk63_4(X1,X2,X3,X4),k5_numbers)|r1_xxreal_0(X4,np__0)|v2_struct_0(X1)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_756, plain, (m1_subset_1(esk63_4(X1,X2,X3,X4),k1_numbers)|r1_xxreal_0(X4,np__0)|v2_struct_0(X1)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_798, plain, (v1_xcmplx_0(esk8_3(X1,X2,X3))|v2_comseq_2(X1)|~v1_membered(k5_numbers)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_800, plain, (v1_xreal_0(esk8_3(X1,X2,X3))|v2_comseq_2(X1)|~v3_membered(k5_numbers)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_801, plain, (v1_rat_1(esk8_3(X1,X2,X3))|v2_comseq_2(X1)|~v4_membered(k5_numbers)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_802, plain, (v1_int_1(esk8_3(X1,X2,X3))|v2_comseq_2(X1)|~v5_membered(k5_numbers)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_808, plain, (v1_rat_1(esk8_3(X1,X2,X3))|v2_comseq_2(X1)|~v4_membered(k1_numbers)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2793, plain, (X1=u1_struct_0(esk25_0)|~v1_xboole_0(X1)|~v2_struct_0(esk25_0))).#trainneg
cnf(i_0_1364, plain, (v1_xboole_0(u1_struct_0(esk23_0))|~v2_struct_0(esk23_0))).#trainneg
cnf(i_0_2999, plain, (u1_struct_0(esk23_0)=np__0|~v2_struct_0(esk23_0))).#trainneg
cnf(i_0_1054, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2959, plain, (v3_valued_0(k5_numbers)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k1_numbers))).#trainneg
cnf(i_0_2592, negated_conjecture, (v4_valued_0(esk4_0)|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3004, plain, (v3_ordinal1(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1062, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k1_numbers))|v1_xboole_0(k3_series_1(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k3_series_1(X2)))).#trainneg
cnf(i_0_1064, plain, (r1_xxreal_0(X1,np__0)|v2_struct_0(X2)|~r1_xxreal_0(X1,k18_complex1(k8_nat_1(X6,k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),X5)))|~r1_xxreal_0(esk63_4(X2,X3,X4,X1),X5)|~v1_xreal_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X5)|~v1_funct_2(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),k5_numbers,X6)|~v1_funct_2(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),k5_numbers,k1_numbers)|~v1_funct_1(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))))|~m1_subset_1(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X6)))|~m1_subset_1(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_809, plain, (v1_int_1(esk8_3(X1,X2,X3))|v2_comseq_2(X1)|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_812, plain, (v1_xcmplx_0(esk11_3(X1,X2,X3))|X2=k1_seq_2(X1)|~v1_membered(k5_numbers)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_815, plain, (v1_rat_1(esk11_3(X1,X2,X3))|X2=k1_seq_2(X1)|~v4_membered(k5_numbers)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_816, plain, (v1_int_1(esk11_3(X1,X2,X3))|X2=k1_seq_2(X1)|~v5_membered(k5_numbers)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_822, plain, (v1_rat_1(esk11_3(X1,X2,X3))|X2=k1_seq_2(X1)|~v4_membered(k1_numbers)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_823, plain, (v1_int_1(esk11_3(X1,X2,X3))|X2=k1_seq_2(X1)|~v5_membered(k1_numbers)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2957, plain, (X1=u1_struct_0(esk24_0)|~v1_xboole_0(X1)|~v2_struct_0(esk24_0))).#trainneg
cnf(i_0_1365, plain, (v1_xboole_0(u1_struct_0(esk22_0))|~v2_struct_0(esk22_0))).#trainneg
cnf(i_0_2048, plain, (v1_valued_0(esk44_0)|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_3022, plain, (u1_struct_0(esk22_0)=np__0|~v2_struct_0(esk22_0))).#trainneg
cnf(i_0_2960, plain, (v3_valued_0(k5_numbers)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k4_numbers))).#trainneg
cnf(i_0_2796, negated_conjecture, (v1_valued_0(esk4_0)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2964, negated_conjecture, (v2_valued_0(esk4_0)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2801, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_3006, negated_conjecture, (v4_valued_0(esk4_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1066, plain, (k8_nat_1(X5,k3_series_1(k4_normsp_0(X1,k8_lopban_4(X1,X2,X3,X4))),X4)=k1_seq_1(esk61_3(X1,X2,X3),X4)|v2_struct_0(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v1_funct_2(k3_series_1(k4_normsp_0(X1,k8_lopban_4(X1,X2,X3,X4))),k5_numbers,X5)|~v1_funct_2(k3_series_1(k4_normsp_0(X1,k8_lopban_4(X1,X2,X3,X4))),k5_numbers,k1_numbers)|~v1_funct_1(k3_series_1(k4_normsp_0(X1,k8_lopban_4(X1,X2,X3,X4))))|~m1_subset_1(k3_series_1(k4_normsp_0(X1,k8_lopban_4(X1,X2,X3,X4))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(k3_series_1(k4_normsp_0(X1,k8_lopban_4(X1,X2,X3,X4))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1067, plain, (r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,X2),X3)),k8_nat_1(X4,k3_series_1(k4_normsp_0(X1,X2)),X3))|v2_struct_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X3)|~v1_funct_2(k3_series_1(k4_normsp_0(X1,X2)),k5_numbers,X4)|~v1_funct_2(k3_series_1(k4_normsp_0(X1,X2)),k5_numbers,k1_numbers)|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2)))|~m1_subset_1(k3_series_1(k4_normsp_0(X1,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(k3_series_1(k4_normsp_0(X1,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_831, plain, (v3_normsp_1(k8_lopban_4(X1,X2,X3,X4),X1)|m1_subset_1(esk62_3(X1,k8_lopban_4(X1,X2,X3,X4),X5),k5_numbers)|v2_struct_0(X1)|k2_seq_2(X5)!=np__0|~v2_comseq_2(X5)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X5,k5_numbers,k1_numbers)|~v1_funct_1(X5)|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_832, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|m1_subset_1(esk62_3(esk1_0,esk4_0,X1),k1_numbers)|k2_seq_2(X1)!=np__0|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_834, plain, (v3_normsp_1(k8_lopban_4(X1,X2,X3,X4),X1)|m1_subset_1(esk62_3(X1,k8_lopban_4(X1,X2,X3,X4),X5),k1_numbers)|v2_struct_0(X1)|k2_seq_2(X5)!=np__0|~v2_comseq_2(X5)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X5,k5_numbers,k1_numbers)|~v1_funct_1(X5)|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_835, plain, (v1_xcmplx_0(esk62_3(X1,X2,X3))|k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|~v1_membered(k5_numbers)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_836, plain, (v1_xxreal_0(esk62_3(X1,X2,X3))|k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|~v2_membered(k5_numbers)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3109, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|~v2_xxreal_0(esk62_3(X1,X2,X3))|~v3_xxreal_0(esk62_3(X1,X2,X3))|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3110, plain, (r1_xxreal_0(esk62_3(X1,X2,X3),esk62_3(X1,X2,X3))|k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3001, plain, (X1=u1_struct_0(esk23_0)|~v1_xboole_0(X1)|~v2_struct_0(esk23_0))).#trainneg
cnf(i_0_1366, plain, (v1_xboole_0(u1_struct_0(esk21_0))|~v2_struct_0(esk21_0))).#trainneg
cnf(i_0_2061, plain, (v2_valued_0(esk44_0)|~v2_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_3117, plain, (u1_struct_0(esk21_0)=np__0|~v2_struct_0(esk21_0))).#trainneg
cnf(i_0_2961, plain, (v3_valued_0(k5_numbers)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k3_numbers))).#trainneg
cnf(i_0_2966, negated_conjecture, (v3_valued_0(esk4_0)|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3027, negated_conjecture, (v1_valued_0(esk4_0)|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1076, plain, (r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(X5,k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)),k1_seq_1(esk61_3(X1,X2,X3),X4))|v2_struct_0(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v1_funct_2(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k5_numbers,X5)|~v1_funct_2(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)))|~m1_subset_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1077, plain, (r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(X4,k1_bhsp_4(X1,X2),X3)),k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|v2_struct_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X3)|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,X4)|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,X2))|~m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_837, plain, (v1_xreal_0(esk62_3(X1,X2,X3))|k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|~v3_membered(k5_numbers)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_838, plain, (v1_rat_1(esk62_3(X1,X2,X3))|k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|~v4_membered(k5_numbers)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_839, plain, (v1_int_1(esk62_3(X1,X2,X3))|k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|~v5_membered(k5_numbers)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_840, plain, (v7_ordinal1(esk62_3(X1,X2,X3))|k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3160, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|~v3_xxreal_0(esk62_3(X1,X2,X3))|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_841, plain, (v3_ordinal1(esk62_3(X1,X2,X3))|k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_845, plain, (v1_rat_1(esk62_3(X1,X2,X3))|k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|~v4_membered(k1_numbers)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3024, plain, (X1=u1_struct_0(esk22_0)|~v1_xboole_0(X1)|~v2_struct_0(esk22_0))).#trainneg
cnf(i_0_1367, plain, (v1_xboole_0(u1_struct_0(esk20_0))|~v2_struct_0(esk20_0))).#trainneg
cnf(i_0_2101, plain, (v4_membered(k5_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_3161, plain, (u1_struct_0(esk20_0)=np__0|~v2_struct_0(esk20_0))).#trainneg
cnf(i_0_2963, plain, (v3_valued_0(k5_numbers)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3031, negated_conjecture, (v2_valued_0(esk4_0)|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3035, negated_conjecture, (v3_valued_0(esk4_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_846, plain, (v1_int_1(esk62_3(X1,X2,X3))|k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|~v5_membered(k1_numbers)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_857, plain, (~v2_xxreal_0(k4_struct_0(X1))|~v3_xxreal_0(k4_struct_0(X1))|~l2_struct_0(X1)|~v2_membered(u1_struct_0(X1)))).#trainneg
cnf(i_0_858, plain, (~v2_xxreal_0(X1)|~v3_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_867, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3119, plain, (X1=u1_struct_0(esk21_0)|~v1_xboole_0(X1)|~v2_struct_0(esk21_0))).#trainneg
cnf(i_0_1368, plain, (v1_xboole_0(u1_struct_0(esk19_0))|~v2_struct_0(esk19_0))).#trainneg
cnf(i_0_3192, plain, (u1_struct_0(esk19_0)=np__0|~v2_struct_0(esk19_0))).#trainneg
cnf(i_0_2136, plain, (v4_valued_0(k5_numbers)|~v4_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1322, negated_conjecture, (v1_xreal_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_3125, negated_conjecture, (v1_valued_0(esk4_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_895, negated_conjecture, (v3_normsp_1(k1_bhsp_4(esk1_0,esk4_0),esk1_0)|m1_subset_1(esk62_3(esk1_0,k1_bhsp_4(esk1_0,esk4_0),X1),k5_numbers)|k2_seq_2(X1)!=np__0|~v2_comseq_2(X1)|~v1_funct_2(k1_bhsp_4(esk1_0,esk4_0),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(esk1_0,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_896, plain, (v3_normsp_1(k1_bhsp_4(X1,k1_bhsp_4(X1,X2)),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,k1_bhsp_4(X1,X2)),X3),k5_numbers)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(k1_bhsp_4(X1,k1_bhsp_4(X1,X2)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(X1,k1_bhsp_4(X1,X2)))|~v1_funct_1(X3)|~v1_funct_1(k1_bhsp_4(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_897, plain, (v3_normsp_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X5),k5_numbers)|v2_struct_0(X1)|k2_seq_2(X5)!=np__0|~v2_comseq_2(X5)|~v1_funct_2(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X5,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)))|~v1_funct_1(X5)|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_898, negated_conjecture, (v3_normsp_1(k1_bhsp_4(esk1_0,esk4_0),esk1_0)|m1_subset_1(esk62_3(esk1_0,k1_bhsp_4(esk1_0,esk4_0),X1),k1_numbers)|k2_seq_2(X1)!=np__0|~v2_comseq_2(X1)|~v1_funct_2(k1_bhsp_4(esk1_0,esk4_0),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(esk1_0,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_899, plain, (v3_normsp_1(k1_bhsp_4(X1,k1_bhsp_4(X1,X2)),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,k1_bhsp_4(X1,X2)),X3),k1_numbers)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(k1_bhsp_4(X1,k1_bhsp_4(X1,X2)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(X1,k1_bhsp_4(X1,X2)))|~v1_funct_1(X3)|~v1_funct_1(k1_bhsp_4(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_900, plain, (v3_normsp_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X5),k1_numbers)|v2_struct_0(X1)|k2_seq_2(X5)!=np__0|~v2_comseq_2(X5)|~v1_funct_2(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X5,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)))|~v1_funct_1(X5)|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_3163, plain, (X1=u1_struct_0(esk20_0)|~v1_xboole_0(X1)|~v2_struct_0(esk20_0))).#trainneg
cnf(i_0_1369, plain, (v1_xboole_0(u1_struct_0(esk18_0))|~v2_struct_0(esk18_0))).#trainneg
cnf(i_0_2114, plain, (v3_valued_0(esk44_0)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_3204, plain, (u1_struct_0(esk18_0)=np__0|~v2_struct_0(esk18_0))).#trainneg
cnf(i_0_3197, plain, (v4_valued_0(k5_numbers)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1323, negated_conjecture, (v1_xreal_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_3168, negated_conjecture, (v2_valued_0(esk4_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_2724, negated_conjecture, (v3_ordinal1(esk3_0)|~v1_xboole_0(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1015, plain, (~v2_xxreal_0(k3_series_1(X1))|~v3_xxreal_0(k3_series_1(X1))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1033, plain, (~v3_xxreal_0(k3_series_1(X1))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1055, plain, (v3_normsp_1(k1_bhsp_4(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2),k5_numbers)|v2_struct_0(X1)|k2_seq_2(X2)!=np__0|~v2_comseq_2(X2)|~v1_funct_2(k1_bhsp_4(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_1(X2)|~v1_funct_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1056, plain, (v3_normsp_1(k1_bhsp_4(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2),k1_numbers)|v2_struct_0(X1)|k2_seq_2(X2)!=np__0|~v2_comseq_2(X2)|~v1_funct_2(k1_bhsp_4(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_1(X2)|~v1_funct_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1058, plain, (v3_normsp_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),X1)|m1_subset_1(esk62_3(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),X2),k5_numbers)|v2_struct_0(X1)|k2_seq_2(X2)!=np__0|~v2_comseq_2(X2)|~v1_funct_2(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1059, plain, (v3_normsp_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),X1)|m1_subset_1(esk62_3(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),X2),k1_numbers)|v2_struct_0(X1)|k2_seq_2(X2)!=np__0|~v2_comseq_2(X2)|~v1_funct_2(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1065, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k8_nat_1(X2,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v7_ordinal1(X1)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),k5_numbers,X2)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),k5_numbers,k1_numbers)|~v1_funct_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_3194, plain, (X1=u1_struct_0(esk19_0)|~v1_xboole_0(X1)|~v2_struct_0(esk19_0))).#trainneg
cnf(i_0_1370, plain, (v1_xboole_0(u1_struct_0(esk17_0))|~v2_struct_0(esk17_0))).#trainneg
cnf(i_0_2127, plain, (v4_valued_0(esk44_0)|~v4_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_3227, plain, (u1_struct_0(esk17_0)=np__0|~v2_struct_0(esk17_0))).#trainneg
cnf(i_0_3209, plain, (v5_relat_1(k5_numbers,k4_numbers)|~v1_relat_1(k5_numbers)|~v5_relat_1(k1_numbers,k5_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2227, plain, (v5_membered(esk52_0)|~v5_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_2725, negated_conjecture, (v3_ordinal1(esk2_0)|~v1_xboole_0(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3009, plain, (v3_ordinal1(X1)|~v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_1119, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),X2,esk62_3(X1,X2,X3))),k8_nat_1(X5,X3,esk62_3(X1,X2,X3)))|~v2_comseq_2(X3)|~v7_ordinal1(esk62_3(X1,X2,X3))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,X4)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_funct_2(X3,k5_numbers,X5)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5))))).#trainneg
cnf(i_0_3240, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),X2,esk62_3(X1,X2,X3))),k1_funct_1(X3,esk62_3(X1,X2,X3)))|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,X4)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk62_3(X1,X2,X3)))).#trainneg
cnf(i_0_3244, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(X5,X2,esk62_3(X1,X2,X3))),k8_nat_1(X4,X3,esk62_3(X1,X2,X3)))|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,X4)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk62_3(X1,X2,X3))|~v1_funct_2(X2,k5_numbers,X5)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5))))).#trainneg
cnf(i_0_1130, plain, (v3_normsp_1(X1,X2)|v2_struct_0(X2)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X2,k1_funct_1(X1,esk62_3(X2,X1,X3))),k8_nat_1(X4,X3,esk62_3(X2,X1,X3)))|~v2_comseq_2(X3)|~v7_ordinal1(esk62_3(X2,X1,X3))|~v1_funct_2(X1,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,X4)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1131, plain, (v3_normsp_1(X1,X2)|v2_struct_0(X2)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X2,k8_nat_1(X5,X1,esk62_3(X2,X1,X3))),k8_nat_1(X4,X3,esk62_3(X2,X1,X3)))|~v2_comseq_2(X3)|~v7_ordinal1(esk62_3(X2,X1,X3))|~v1_funct_2(X1,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,X4)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)|~v1_funct_2(X1,k5_numbers,X5)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5))))).#trainneg
cnf(i_0_3206, plain, (X1=u1_struct_0(esk18_0)|~v1_xboole_0(X1)|~v2_struct_0(esk18_0))).#trainneg
cnf(i_0_1371, plain, (v1_xboole_0(u1_struct_0(esk16_0))|~v2_struct_0(esk16_0))).#trainneg
cnf(i_0_3285, plain, (u1_struct_0(esk16_0)=np__0|~v2_struct_0(esk16_0))).#trainneg
cnf(i_0_1164, plain, (v1_xreal_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_3210, plain, (v5_relat_1(k5_numbers,k3_numbers)|~v1_relat_1(k5_numbers)|~v5_relat_1(k1_numbers,k5_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2228, plain, (v6_ordinal1(esk52_0)|~v6_ordinal1(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_2229, plain, (v6_membered(esk52_0)|~v6_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1144, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X1,k1_funct_1(X2,esk62_3(X1,X2,X3))),k8_nat_1(k1_numbers,X3,esk62_3(X1,X2,X3)))|~v2_comseq_2(X3)|~v7_ordinal1(esk62_3(X1,X2,X3))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,X4)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1145, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(X5,X2,esk62_3(X1,X2,X3))),k8_nat_1(k1_numbers,X3,esk62_3(X1,X2,X3)))|~v2_comseq_2(X3)|~v7_ordinal1(esk62_3(X1,X2,X3))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,X4)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,X5)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5))))).#trainneg
cnf(i_0_1136, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)),k8_nat_1(X3,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),X2))|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X1)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k5_numbers,k1_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k5_numbers,X3)|~v1_funct_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_3229, plain, (X1=u1_struct_0(esk17_0)|~v1_xboole_0(X1)|~v2_struct_0(esk17_0))).#trainneg
cnf(i_0_1372, plain, (v1_xboole_0(u1_struct_0(esk15_0))|~v2_struct_0(esk15_0))).#trainneg
cnf(i_0_2667, plain, (v1_xcmplx_0(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_3321, plain, (u1_struct_0(esk15_0)=np__0|~v2_struct_0(esk15_0))).#trainneg
cnf(i_0_1169, plain, (v1_xxreal_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_3232, plain, (v5_relat_1(esk44_0,k4_numbers)|~v1_relat_1(esk44_0)|~v4_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2230, plain, (v1_membered(esk52_0)|~v1_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_2231, plain, (v2_membered(esk52_0)|~v2_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_2730, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X1)|v1_xboole_0(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2732, plain, (r1_xxreal_0(X1,X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1159, plain, (v2_struct_0(X1)|~v2_xxreal_0(k4_normsp_0(X1,X2))|~v3_xxreal_0(k4_normsp_0(X1,X2))|~l1_normsp_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1171, plain, (v2_struct_0(X1)|~v3_xxreal_0(k4_normsp_0(X1,X2))|~l1_normsp_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1190, plain, (v3_normsp_1(X1,X2)|v2_struct_0(X2)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X2,k1_funct_1(X1,esk62_3(X2,X1,X3))),k1_funct_1(X3,esk62_3(X2,X1,X3)))|~v2_comseq_2(X3)|~v7_ordinal1(esk62_3(X2,X1,X3))|~v1_funct_2(X1,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,X4)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1191, plain, (v3_normsp_1(X1,X2)|v2_struct_0(X2)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X2,k8_nat_1(X5,X1,esk62_3(X2,X1,X3))),k1_funct_1(X3,esk62_3(X2,X1,X3)))|~v2_comseq_2(X3)|~v7_ordinal1(esk62_3(X2,X1,X3))|~v1_funct_2(X1,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,X4)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)|~v1_funct_2(X1,k5_numbers,X5)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5))))).#trainneg
cnf(i_0_1198, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k1_numbers))|v2_struct_0(X2)|v1_xboole_0(k4_normsp_0(X2,X3))|~l1_normsp_0(X2)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k4_normsp_0(X2,X3)))).#trainneg
cnf(i_0_3287, plain, (X1=u1_struct_0(esk16_0)|~v1_xboole_0(X1)|~v2_struct_0(esk16_0))).#trainneg
cnf(i_0_1373, plain, (v1_xboole_0(u1_struct_0(esk14_0))|~v2_struct_0(esk14_0))).#trainneg
cnf(i_0_3380, plain, (u1_struct_0(esk14_0)=np__0|~v2_struct_0(esk14_0))).#trainneg
cnf(i_0_1202, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3330, plain, (v5_relat_1(esk44_0,k4_numbers)|~v1_relat_1(esk44_0)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2232, plain, (v3_membered(esk52_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_2233, plain, (v4_membered(esk52_0)|~v4_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1517, plain, (v1_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1518, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1210, plain, (v1_xreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1226, plain, (v1_relat_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_1203, plain, (v2_struct_0(X1)|~v2_xxreal_0(esk61_3(X1,X2,X3))|~v3_xxreal_0(esk61_3(X1,X2,X3))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1224, plain, (v3_normsp_1(k1_bhsp_4(X1,esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))),X2),k5_numbers)|v2_struct_0(X1)|k2_seq_2(X2)!=np__0|~v2_comseq_2(X2)|~v1_funct_2(k1_bhsp_4(X1,esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))),k5_numbers,u1_struct_0(X1))|~v1_funct_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(X1,esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))))|~v1_funct_1(X2)|~v1_funct_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1225, plain, (v3_normsp_1(k1_bhsp_4(X1,esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))),X2),k1_numbers)|v2_struct_0(X1)|k2_seq_2(X2)!=np__0|~v2_comseq_2(X2)|~v1_funct_2(k1_bhsp_4(X1,esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))),k5_numbers,u1_struct_0(X1))|~v1_funct_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(X1,esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))))|~v1_funct_1(X2)|~v1_funct_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1227, plain, (v3_normsp_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),X1)|m1_subset_1(esk62_3(X1,esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),X2),k5_numbers)|v2_struct_0(X1)|k2_seq_2(X2)!=np__0|~v2_comseq_2(X2)|~v1_funct_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1228, plain, (v3_normsp_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),X1)|m1_subset_1(esk62_3(X1,esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),X2),k1_numbers)|v2_struct_0(X1)|k2_seq_2(X2)!=np__0|~v2_comseq_2(X2)|~v1_funct_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1229, plain, (v1_xboole_0(esk28_1(k1_zfmisc_1(X1)))|v1_xboole_0(X1)|m1_subset_1(X2,X1)|~m1_subset_1(X2,esk28_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_1233, plain, (v1_xreal_0(esk28_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_1234, plain, (v1_rat_1(esk28_1(X1))|~v4_membered(X1))).#trainneg
cnf(i_0_1235, plain, (v1_int_1(esk28_1(X1))|~v5_membered(X1))).#trainneg
cnf(i_0_1236, plain, (v7_ordinal1(esk28_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_1237, plain, (v3_ordinal1(esk28_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1242, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk28_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,esk28_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_3323, plain, (X1=u1_struct_0(esk15_0)|~v1_xboole_0(X1)|~v2_struct_0(esk15_0))).#trainneg
cnf(i_0_1374, plain, (v1_xboole_0(u1_struct_0(esk13_0))|~v2_struct_0(esk13_0))).#trainneg
cnf(i_0_3290, plain, (v6_ordinal1(esk52_0)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_3421, plain, (u1_struct_0(esk13_0)=np__0|~v2_struct_0(esk13_0))).#trainneg
cnf(i_0_1244, plain, (v1_xxreal_0(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3234, plain, (v5_membered(esk52_0)|~v6_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1520, negated_conjecture, (v7_ordinal1(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1521, negated_conjecture, (v7_ordinal1(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1245, plain, (v1_xreal_0(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1249, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k1_numbers))|v2_struct_0(X2)|v1_xboole_0(esk61_3(X2,X3,X4))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)|~m1_subset_1(X1,esk61_3(X2,X3,X4)))).#trainneg
cnf(i_0_1246, plain, (v2_struct_0(X1)|~v3_xxreal_0(esk61_3(X1,X2,X3))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1251, plain, (k1_seq_2(X1)=np__0|r1_xxreal_0(esk10_2(X1,np__0),k18_complex1(k8_nat_1(X3,X1,esk11_3(X1,np__0,X2))))|~v2_comseq_2(X1)|~v3_valued_0(X1)|~v1_xcmplx_0(k8_nat_1(X3,X1,esk11_3(X1,np__0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k8_nat_1(X3,X1,esk11_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v7_ordinal1(esk11_3(X1,np__0,X2))|~v1_funct_2(X1,k5_numbers,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_1253, plain, (r1_xxreal_0(esk7_2(X1,np__0),k18_complex1(k8_nat_1(X3,X1,esk8_3(X1,np__0,X2))))|v2_comseq_2(X1)|~v3_valued_0(X1)|~v1_xcmplx_0(k8_nat_1(X3,X1,esk8_3(X1,np__0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k8_nat_1(X3,X1,esk8_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v7_ordinal1(esk8_3(X1,np__0,X2))|~v1_funct_2(X1,k5_numbers,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_1261, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|r1_xxreal_0(esk10_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~v3_valued_0(esk61_3(X1,X2,X3))|~v1_xcmplx_0(k8_nat_1(X4,esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)))))|~m2_subset_1(esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(k8_nat_1(X4,esk61_3(X1,X2,X3),esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0)))),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk11_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk10_2(esk61_3(X1,X2,X3),np__0))))|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,X4)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4))))).#trainneg
cnf(i_0_3477, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|v2_struct_0(X1)|~v2_xxreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~v1_xreal_0(esk10_2(esk61_3(X1,X2,X3),np__0))|~v2_comseq_2(esk61_3(X1,X2,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3382, plain, (X1=u1_struct_0(esk14_0)|~v1_xboole_0(X1)|~v2_struct_0(esk14_0))).#trainneg
cnf(i_0_1375, plain, (v1_xboole_0(u1_struct_0(esk12_0))|~v2_struct_0(esk12_0))).#trainneg
cnf(i_0_3291, plain, (v6_membered(esk52_0)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1267, plain, (v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3233, plain, (v5_relat_1(esk44_0,k3_numbers)|~v1_relat_1(esk44_0)|~v4_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_3434, negated_conjecture, (~v3_xxreal_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3437, negated_conjecture, (~v3_xxreal_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1524, negated_conjecture, (v1_xcmplx_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1525, negated_conjecture, (v1_xcmplx_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1266, plain, (r1_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~m2_subset_1(esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)))),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_xreal_0(k1_seq_1(esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0))))))).#trainneg
cnf(i_0_3484, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(esk10_2(esk61_3(X1,X2,X3),np__0),k1_numbers))).#trainneg
cnf(i_0_3485, plain, (k1_seq_2(esk61_3(X1,X2,X3))=np__0|v2_struct_0(X1)|~v2_comseq_2(esk61_3(X1,X2,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(esk10_2(esk61_3(X1,X2,X3),np__0),k5_numbers))).#trainneg
cnf(i_0_1268, plain, (v1_xcmplx_0(esk28_1(esk44_0))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1269, plain, (v1_xxreal_0(esk28_1(esk44_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_3512, plain, (~v2_xxreal_0(esk28_1(esk44_0))|~v3_xxreal_0(esk28_1(esk44_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1271, plain, (v1_rat_1(esk28_1(esk44_0))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_1272, plain, (v1_int_1(esk28_1(esk44_0))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1273, plain, (v7_ordinal1(esk28_1(esk44_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_3423, plain, (X1=u1_struct_0(esk13_0)|~v1_xboole_0(X1)|~v2_struct_0(esk13_0))).#trainneg
cnf(i_0_2621, plain, (v1_xreal_0(u1_struct_0(esk27_0))|~v2_struct_0(esk27_0))).#trainneg
cnf(i_0_3429, plain, (v5_membered(esk52_0)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1279, plain, (v1_xreal_0(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_3497, plain, (v5_relat_1(esk44_0,k3_numbers)|~v1_relat_1(esk44_0)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3519, plain, (~v3_xxreal_0(esk28_1(esk44_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_3331, plain, (v1_membered(esk52_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1526, negated_conjecture, (v1_xxreal_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1527, negated_conjecture, (v1_xxreal_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1280, plain, (v2_xxreal_0(esk61_3(X1,X2,X3))|v3_xxreal_0(esk61_3(X1,X2,X3))|v1_xboole_0(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1294, plain, (v2_xxreal_0(k4_normsp_0(X1,X2))|v3_xxreal_0(k4_normsp_0(X1,X2))|v1_xboole_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_3524, plain, (v1_xreal_0(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers))).#trainneg
cnf(i_0_1278, plain, (v2_struct_0(X1)|~v2_xxreal_0(k1_bhsp_4(X1,X2))|~v3_xxreal_0(k1_bhsp_4(X1,X2))|~l2_algstr_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1287, plain, (v2_xxreal_0(esk44_0)|v3_xxreal_0(esk44_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1289, plain, (v2_xxreal_0(esk11_3(X1,X2,X3))|v3_xxreal_0(esk11_3(X1,X2,X3))|v1_xboole_0(esk11_3(X1,X2,X3))|X2=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1290, plain, (v2_xxreal_0(esk8_3(X1,X2,X3))|v3_xxreal_0(esk8_3(X1,X2,X3))|v1_xboole_0(esk8_3(X1,X2,X3))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1295, plain, (v2_xxreal_0(k3_series_1(X1))|v3_xxreal_0(k3_series_1(X1))|v1_xboole_0(k3_series_1(X1))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1297, plain, (v2_xxreal_0(k4_struct_0(X1))|v3_xxreal_0(k4_struct_0(X1))|v1_xboole_0(k4_struct_0(X1))|~l2_struct_0(X1)|~v2_membered(u1_struct_0(X1)))).#trainneg
cnf(i_0_1298, plain, (v2_xxreal_0(k4_struct_0(X1))|v3_xxreal_0(k4_struct_0(X1))|v1_xboole_0(k4_struct_0(X1))|~l2_struct_0(X1)|~v6_membered(u1_struct_0(X1)))).#trainneg
cnf(i_0_1299, plain, (v2_xxreal_0(k1_bhsp_4(X1,X2))|v3_xxreal_0(k1_bhsp_4(X1,X2))|v1_xboole_0(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1300, plain, (v2_xxreal_0(k5_numbers)|v3_xxreal_0(k5_numbers)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3490, plain, (X1=u1_struct_0(esk12_0)|~v1_xboole_0(X1)|~v2_struct_0(esk12_0))).#trainneg
cnf(i_0_2718, plain, (v1_xreal_0(u1_struct_0(esk26_0))|~v2_struct_0(esk26_0))).#trainneg
cnf(i_0_3521, plain, (v3_membered(esk52_0)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1529, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1530, plain, (v1_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1306, plain, (k18_complex1(k16_complex1(X1))=k16_complex1(X1)|~v1_xcmplx_0(k16_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1319, negated_conjecture, (v7_ordinal1(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k5_numbers))).#trainneg
cnf(i_0_1328, plain, (v1_xreal_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1329, plain, (v2_xxreal_0(esk28_1(X1))|v3_xxreal_0(esk28_1(X1))|v1_xboole_0(esk28_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_1330, plain, (~v2_xxreal_0(esk28_1(X1))|~v3_xxreal_0(esk28_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_1339, plain, (v1_xxreal_0(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2727, plain, (v1_xreal_0(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2792, plain, (v1_xreal_0(u1_struct_0(esk25_0))|~v2_struct_0(esk25_0))).#trainneg
cnf(i_0_3527, plain, (v1_membered(esk52_0)|~v6_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_3578, negated_conjecture, (~v3_xxreal_0(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k5_numbers))).#trainneg
cnf(i_0_1384, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))|v2_struct_0(X2)|v1_xboole_0(k1_bhsp_4(X2,X3))|~l2_algstr_0(X2)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_bhsp_4(X2,X3)))).#trainneg
cnf(i_0_1539, negated_conjecture, (v1_xcmplx_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1385, plain, (v1_xcmplx_0(k18_complex1(X1))|~v1_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1386, plain, (v1_xxreal_0(k18_complex1(X1))|~v2_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1341, plain, (v2_struct_0(X1)|~v3_xxreal_0(k1_bhsp_4(X1,X2))|~l2_algstr_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1382, plain, (v2_xxreal_0(k5_numbers)|v3_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_1388, plain, (v1_rat_1(k18_complex1(X1))|~v4_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1389, plain, (v1_int_1(k18_complex1(X1))|~v5_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1390, plain, (v7_ordinal1(k18_complex1(X1))|~v6_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1391, plain, (v3_ordinal1(k18_complex1(X1))|~v3_ordinal1(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1396, plain, (v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1400, plain, (v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3003, plain, (v3_ordinal1(X1)|~v7_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2956, plain, (v1_xreal_0(u1_struct_0(esk24_0))|~v2_struct_0(esk24_0))).#trainneg
cnf(i_0_2955, plain, (u1_struct_0(esk24_0)=np__0|~v2_struct_0(esk24_0))).#trainneg
cnf(i_0_3623, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_1423, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|~v2_xxreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1424, plain, (v3_xxreal_0(X1)|v3_xxreal_0(k4_xcmplx_0(X1))|v1_xboole_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1426, plain, (v2_xxreal_0(k6_normsp_1(X1,X2))|v3_xxreal_0(k6_normsp_1(X1,X2))|v1_xboole_0(k6_normsp_1(X1,X2))|v2_struct_0(X1)|~l1_normsp_1(X1)|~v2_membered(u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1427, plain, (v2_struct_0(X1)|~v2_xxreal_0(k6_normsp_1(X1,X2))|~v3_xxreal_0(k6_normsp_1(X1,X2))|~l1_normsp_1(X1)|~v2_membered(u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1433, plain, (v2_xxreal_0(k4_xcmplx_0(X1))|~v3_xxreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1462, plain, (r1_xxreal_0(esk61_3(X1,X2,X3),esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1469, plain, (r1_xxreal_0(esk44_0,esk44_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1471, plain, (r1_xxreal_0(esk28_1(X1),esk28_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_3385, plain, (v1_xcmplx_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3000, plain, (v1_xreal_0(u1_struct_0(esk23_0))|~v2_struct_0(esk23_0))).#trainneg
cnf(i_0_3488, plain, (u1_struct_0(esk12_0)=np__0|~v2_struct_0(esk12_0))).#trainneg
cnf(i_0_1477, plain, (r1_xxreal_0(k4_normsp_0(X1,X2),k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_3668, plain, (v1_xcmplx_0(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3669, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_1478, plain, (r1_xxreal_0(k3_series_1(X1),k3_series_1(X1))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1480, plain, (r1_xxreal_0(k4_struct_0(X1),k4_struct_0(X1))|~l2_struct_0(X1)|~v2_membered(u1_struct_0(X1)))).#trainneg
cnf(i_0_1481, plain, (r1_xxreal_0(k4_struct_0(X1),k4_struct_0(X1))|~l2_struct_0(X1)|~v6_membered(u1_struct_0(X1)))).#trainneg
cnf(i_0_1482, plain, (r1_xxreal_0(k6_normsp_1(X1,X2),k6_normsp_1(X1,X2))|v2_struct_0(X1)|~l1_normsp_1(X1)|~v2_membered(u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1483, plain, (r1_xxreal_0(k1_bhsp_4(X1,X2),k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1489, plain, (r1_xxreal_0(k18_complex1(k1_funct_1(esk61_3(X1,X2,X3),X4)),np__0)|v2_struct_0(X1)|~r1_xxreal_0(esk63_4(X1,X2,X3,k18_complex1(k1_funct_1(esk61_3(X1,X2,X3),X4))),X4)|~v1_xreal_0(k18_complex1(k1_funct_1(esk61_3(X1,X2,X3),X4)))|~v3_valued_0(esk61_3(X1,X2,X3))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1490, plain, (r1_xxreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,X2),esk5_1(X1))),np__0)|~r1_xxreal_0(esk6_2(X1,k18_complex1(k9_real_1(k1_seq_1(X1,X2),esk5_1(X1)))),X2)|~v2_comseq_2(X1)|~v1_xreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,X2),esk5_1(X1))))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1491, plain, (r1_xxreal_0(k18_complex1(k9_real_1(k1_funct_1(X1,X2),X3)),np__0)|X3!=k1_seq_2(X1)|~r1_xxreal_0(esk9_3(X1,X3,k18_complex1(k9_real_1(k1_funct_1(X1,X2),X3))),X2)|~v2_comseq_2(X1)|~v1_xreal_0(k18_complex1(k9_real_1(k1_funct_1(X1,X2),X3)))|~v1_xreal_0(X3)|~v3_valued_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1492, plain, (r1_xxreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,X2),X3)),np__0)|X3!=k1_seq_2(X1)|~r1_xxreal_0(esk9_3(X1,X3,k18_complex1(k9_real_1(k1_seq_1(X1,X2),X3))),X2)|~v2_comseq_2(X1)|~v1_xreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,X2),X3)))|~v1_xreal_0(X3)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1493, plain, (r1_xxreal_0(k18_complex1(k9_real_1(k1_funct_1(X1,X2),esk5_1(X1))),np__0)|~r1_xxreal_0(esk6_2(X1,k18_complex1(k9_real_1(k1_funct_1(X1,X2),esk5_1(X1)))),X2)|~v2_comseq_2(X1)|~v1_xreal_0(k18_complex1(k9_real_1(k1_funct_1(X1,X2),esk5_1(X1))))|~v3_valued_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1494, plain, (r1_xxreal_0(k18_complex1(k1_seq_1(esk61_3(X1,X2,X3),X4)),np__0)|v2_struct_0(X1)|~r1_xxreal_0(esk63_4(X1,X2,X3,k18_complex1(k1_seq_1(esk61_3(X1,X2,X3),X4))),X4)|~v1_xreal_0(k18_complex1(k1_seq_1(esk61_3(X1,X2,X3),X4)))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1495, plain, (r1_xxreal_0(k18_complex1(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,k8_lopban_4(X1,X2,X3,X4))),X4)),np__0)|v2_struct_0(X1)|~r1_xxreal_0(esk63_4(X1,X2,X3,k18_complex1(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,k8_lopban_4(X1,X2,X3,X4))),X4))),X4)|~v1_xreal_0(k18_complex1(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,k8_lopban_4(X1,X2,X3,X4))),X4)))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3493, plain, (v1_xxreal_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3023, plain, (v1_xreal_0(u1_struct_0(esk22_0))|~v2_struct_0(esk22_0))).#trainneg
cnf(i_0_3567, plain, (~v2_xxreal_0(k5_numbers)|~v3_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1506, plain, (v1_xxreal_0(k6_normsp_1(X1,X2))|v2_struct_0(X1)|~l1_normsp_1(X1)|~v6_membered(u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1507, plain, (v1_xreal_0(k6_normsp_1(X1,X2))|v2_struct_0(X1)|~l1_normsp_1(X1)|~v6_membered(u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1508, plain, (v2_struct_0(X1)|~v3_xxreal_0(k6_normsp_1(X1,X2))|~l1_normsp_1(X1)|~v6_membered(u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1509, plain, (esk35_2(X1,X2)=np__0)).#trainneg
cnf(i_0_1512, negated_conjecture, (v1_int_1(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1515, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1516, plain, (v7_ordinal1(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1519, negated_conjecture, (v7_ordinal1(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1522, negated_conjecture, (v1_xreal_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1523, negated_conjecture, (v1_xreal_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_399, plain, (v1_relat_1(np__0))).#trainneg
cnf(i_0_398, plain, (v4_relat_1(np__0,X1))).#trainneg
cnf(i_0_401, plain, (m1_subset_1(np__0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1528, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3231, plain, (v3_valued_0(esk44_0)|~v1_relat_1(esk44_0)|~v4_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_3730, plain, (~v3_xxreal_0(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1540, negated_conjecture, (v1_xcmplx_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1531, plain, (v1_xreal_0(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1532, plain, (v1_xreal_0(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1533, negated_conjecture, (v1_xreal_0(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1534, plain, (v1_xreal_0(k3_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1536, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1537, negated_conjecture, (v1_xreal_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1538, negated_conjecture, (v1_xreal_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1541, negated_conjecture, (v1_xxreal_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1542, negated_conjecture, (v1_xxreal_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1543, negated_conjecture, (v1_rat_1(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1547, negated_conjecture, (v1_xcmplx_0(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1548, negated_conjecture, (v1_xcmplx_0(esk4_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1549, negated_conjecture, (v2_xxreal_0(esk3_0)|v3_xxreal_0(esk3_0)|v1_xboole_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1552, plain, (r1_xxreal_0(X1,np__0)|~r1_xxreal_0(X1,k18_complex1(k6_xcmplx_0(k1_funct_1(X2,X3),esk5_1(X2))))|~r1_xxreal_0(esk6_2(X2,X1),X3)|~v2_comseq_2(X2)|~v1_xreal_0(esk5_1(X2))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_funct_1(X2,X3),k1_numbers)|~v3_valued_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_3768, plain, (v1_valued_0(np__0)|~v1_membered(X1))).#trainneg
cnf(i_0_3811, plain, (v1_valued_0(np__0))).#trainneg
cnf(i_0_397, plain, (v5_relat_1(np__0,X2))).#trainneg
cnf(i_0_1553, plain, (r1_xxreal_0(esk10_2(X1,esk5_1(X1)),np__0)|esk5_1(X1)=k1_seq_2(X1)|~r1_xxreal_0(esk6_2(X1,esk10_2(X1,esk5_1(X1))),esk11_3(X1,esk5_1(X1),X2))|~v2_comseq_2(X1)|~v1_xreal_0(esk5_1(X1))|~v1_xreal_0(esk10_2(X1,esk5_1(X1)))|~m2_subset_1(esk11_3(X1,esk5_1(X1),X2),k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X1,esk11_3(X1,esk5_1(X1),X2)),k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3785, plain, (v3_valued_0(esk44_0)|~v1_relat_1(esk44_0)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3790, negated_conjecture, (~v2_xxreal_0(esk3_0)|~v3_xxreal_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3794, negated_conjecture, (~v2_xxreal_0(esk2_0)|~v3_xxreal_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2716, plain, (v7_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1555, plain, (r1_xxreal_0(k18_complex1(k6_xcmplx_0(k1_seq_1(X1,X2),esk5_1(X1))),np__0)|~r1_xxreal_0(esk6_2(X1,k18_complex1(k6_xcmplx_0(k1_seq_1(X1,X2),esk5_1(X1)))),X2)|~v2_comseq_2(X1)|~v1_xreal_0(esk5_1(X1))|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k1_seq_1(X1,X2),esk5_1(X1))))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers))).#trainneg
cnf(i_0_3813, plain, (esk5_1(X1)=k1_seq_2(X1)|r1_xxreal_0(esk10_2(X1,esk5_1(X1)),np__0)|~v2_comseq_2(X1)|~m2_subset_1(esk6_2(X1,esk10_2(X1,esk5_1(X1))),k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(esk5_1(X1)))).#trainneg
cnf(i_0_1558, plain, (r1_xxreal_0(X1,np__0)|k1_seq_2(X2)!=np__0|~r1_xxreal_0(X1,k18_complex1(k1_seq_1(X2,X3)))|~r1_xxreal_0(esk9_3(X2,np__0,X1),X3)|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X2,X3),k1_numbers)|~v1_xcmplx_0(k1_seq_1(X2,X3)))).#trainneg
cnf(i_0_1561, plain, (r1_xxreal_0(k18_complex1(k6_xcmplx_0(k1_seq_1(X1,X2),X3)),np__0)|X3!=k1_seq_2(X1)|~r1_xxreal_0(esk9_3(X1,X3,k18_complex1(k6_xcmplx_0(k1_seq_1(X1,X2),X3))),X2)|~v2_comseq_2(X1)|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k1_seq_1(X1,X2),X3)))|~v1_xreal_0(X3)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers))).#trainneg
cnf(i_0_1564, plain, (v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_1565, plain, (v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_1566, plain, (v1_xcmplx_0(esk30_1(X1))|~v1_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1567, plain, (v1_xxreal_0(esk30_1(X1))|~v2_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1568, plain, (v1_xreal_0(esk30_1(X1))|~v3_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1569, plain, (v1_rat_1(esk30_1(X1))|~v4_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1570, plain, (v1_int_1(esk30_1(X1))|~v5_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_3769, plain, (v2_valued_0(np__0)|~v2_membered(X1))).#trainneg
cnf(i_0_3849, plain, (v2_valued_0(np__0))).#trainneg
cnf(i_0_3733, negated_conjecture, (~v3_xxreal_0(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1571, plain, (v7_ordinal1(esk30_1(X1))|~v6_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_3840, plain, (v5_relat_1(k5_numbers,k4_numbers)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2798, negated_conjecture, (v3_ordinal1(esk3_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_2799, negated_conjecture, (v3_ordinal1(esk2_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_1572, plain, (v3_ordinal1(esk30_1(X1))|~v3_ordinal1(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1574, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_1575, plain, (v1_xcmplx_0(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1577, negated_conjecture, (v2_xxreal_0(esk2_0)|v3_xxreal_0(esk2_0)|v1_xboole_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1580, plain, (v2_xxreal_0(k8_lopban_4(X1,X2,X3,X4))|v3_xxreal_0(k8_lopban_4(X1,X2,X3,X4))|v1_xboole_0(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1581, plain, (v2_struct_0(X1)|~v2_xxreal_0(k8_lopban_4(X1,X2,X3,X4))|~v3_xxreal_0(k8_lopban_4(X1,X2,X3,X4))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1582, plain, (r1_xxreal_0(k8_lopban_4(X1,X2,X3,X4),k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1583, plain, (v1_xreal_0(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1584, plain, (v1_xreal_0(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1585, plain, (v3_valued_0(k2_zfmisc_1(X1,X2))|~v6_membered(X2))).#trainneg
cnf(i_0_1586, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),k4_numbers)|~v6_membered(X2))).#trainneg
cnf(i_0_1587, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),k3_numbers)|~v6_membered(X2))).#trainneg
cnf(i_0_1588, plain, (r1_xxreal_0(X1,esk61_3(X2,X3,X4))|v2_xxreal_0(X1)|v3_xxreal_0(esk61_3(X2,X3,X4))|v2_struct_0(X2)|~v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1589, plain, (r1_xxreal_0(X1,esk59_0)|v2_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3770, plain, (v3_valued_0(np__0)|~v3_membered(X1))).#trainneg
cnf(i_0_3910, plain, (v3_valued_0(np__0))).#trainneg
cnf(i_0_3891, plain, (r1_xxreal_0(esk53_0,esk59_0))).#trainneg
cnf(i_0_1593, plain, (r1_xxreal_0(X1,esk38_0)|v2_xxreal_0(X1)|v3_xxreal_0(esk38_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3841, plain, (v5_relat_1(k5_numbers,k3_numbers)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3332, plain, (v2_membered(esk52_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_3389, plain, (v3_membered(esk52_0)|~v6_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_2965, negated_conjecture, (v3_valued_0(esk4_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3007, negated_conjecture, (v4_valued_0(esk4_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1596, plain, (r1_xxreal_0(X1,esk10_2(X2,X3))|v2_xxreal_0(X1)|v3_xxreal_0(esk10_2(X2,X3))|X3=k1_seq_2(X2)|~v1_xreal_0(X1)|~v2_comseq_2(X2)|~v1_xreal_0(X3)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1598, plain, (r1_xxreal_0(X1,k4_xcmplx_0(X2))|v2_xxreal_0(X1)|v3_xxreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_3890, plain, (r1_xxreal_0(esk54_0,esk59_0)|~v1_xreal_0(esk54_0))).#trainneg
cnf(i_0_3892, plain, (r1_xxreal_0(esk39_0,esk59_0)|~v3_xxreal_0(esk39_0)|~v1_xreal_0(esk39_0))).#trainneg
cnf(i_0_3893, plain, (r1_xxreal_0(esk28_1(esk44_0),esk59_0)|~v3_xxreal_0(esk28_1(esk44_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_3925, plain, (r1_xxreal_0(esk53_0,esk38_0)|v3_xxreal_0(esk38_0))).#trainneg
cnf(i_0_3927, plain, (r1_xxreal_0(esk28_1(esk44_0),esk38_0)|v3_xxreal_0(esk38_0)|~v3_xxreal_0(esk28_1(esk44_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_3990, plain, (r1_xxreal_0(k2_seq_2(X1),k4_xcmplx_0(X2))|v2_xxreal_0(k2_seq_2(X1))|v3_xxreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_3993, plain, (r1_xxreal_0(k1_seq_2(X1),k4_xcmplx_0(X2))|v2_xxreal_0(k1_seq_2(X1))|v3_xxreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4017, plain, (r1_xxreal_0(esk53_0,esk38_0)|~m1_subset_1(esk38_0,k5_numbers))).#trainneg
cnf(i_0_4018, plain, (r1_xxreal_0(esk53_0,esk38_0)|~v1_xboole_0(esk38_0))).#trainneg
cnf(i_0_1590, plain, (r1_xxreal_0(X1,esk55_0)|v2_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4044, plain, (r1_xxreal_0(esk54_0,esk55_0)|~v1_xreal_0(esk54_0))).#trainneg
cnf(i_0_3771, plain, (v4_valued_0(np__0)|~v6_membered(X1))).#trainneg
cnf(i_0_4068, plain, (v4_valued_0(np__0))).#trainneg
cnf(i_0_3932, plain, (r1_xxreal_0(np__0,esk38_0)|v3_xxreal_0(esk38_0))).#trainneg
cnf(i_0_2787, plain, (~v1_xboole_0(k3_series_1(X1))|~v1_funct_2(k3_series_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k3_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1599, plain, (r1_xxreal_0(X1,k4_normsp_0(X2,X3))|v2_xxreal_0(X1)|v3_xxreal_0(k4_normsp_0(X2,X3))|v2_struct_0(X2)|~v1_xreal_0(X1)|~l1_normsp_0(X2)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))))).#trainneg
cnf(i_0_3392, plain, (v4_membered(esk52_0)|~v5_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_3898, plain, (r1_xxreal_0(np__0,esk59_0))).#trainneg
cnf(i_0_3026, negated_conjecture, (v1_valued_0(esk4_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3030, negated_conjecture, (v2_valued_0(esk4_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1600, plain, (r1_xxreal_0(X1,k16_complex1(X2))|v2_xxreal_0(X1)|v3_xxreal_0(k16_complex1(X2))|~v1_xreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1602, plain, (r1_xxreal_0(X1,k18_complex1(X2))|v2_xxreal_0(X1)|v3_xxreal_0(k18_complex1(X2))|~v1_xreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_4045, plain, (r1_xxreal_0(esk53_0,esk55_0))).#trainneg
cnf(i_0_4046, plain, (r1_xxreal_0(esk39_0,esk55_0)|~v3_xxreal_0(esk39_0)|~v1_xreal_0(esk39_0))).#trainneg
cnf(i_0_4047, plain, (r1_xxreal_0(esk28_1(esk44_0),esk55_0)|~v3_xxreal_0(esk28_1(esk44_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_4115, plain, (r1_xxreal_0(k2_seq_2(X1),k16_complex1(X2))|v2_xxreal_0(k2_seq_2(X1))|v3_xxreal_0(k16_complex1(X2))|~v1_xcmplx_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4118, plain, (r1_xxreal_0(k1_seq_2(X1),k16_complex1(X2))|v2_xxreal_0(k1_seq_2(X1))|v3_xxreal_0(k16_complex1(X2))|~v1_xcmplx_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4166, plain, (r1_xxreal_0(k2_seq_2(X1),k18_complex1(X2))|v2_xxreal_0(k2_seq_2(X1))|v3_xxreal_0(k18_complex1(X2))|~v1_xcmplx_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4169, plain, (r1_xxreal_0(k1_seq_2(X1),k18_complex1(X2))|v2_xxreal_0(k1_seq_2(X1))|v3_xxreal_0(k18_complex1(X2))|~v1_xcmplx_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2907, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k1_funct_1(X3,X4))|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_1604, plain, (r1_xxreal_0(X1,k4_struct_0(X2))|v2_xxreal_0(X1)|v3_xxreal_0(k4_struct_0(X2))|~v1_xreal_0(X1)|~l2_struct_0(X2)|~v3_membered(u1_struct_0(X2)))).#trainneg
cnf(i_0_3944, plain, (v2_membered(esk52_0)|~v6_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_3123, negated_conjecture, (v3_valued_0(esk4_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3126, negated_conjecture, (v1_valued_0(esk4_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1605, plain, (r1_xxreal_0(X1,k4_struct_0(X2))|v2_xxreal_0(X1)|v3_xxreal_0(k4_struct_0(X2))|~v1_xreal_0(X1)|~l2_struct_0(X2)|~v6_membered(u1_struct_0(X2)))).#trainneg
cnf(i_0_1606, plain, (r1_xxreal_0(X1,k6_normsp_1(X2,X3))|v2_xxreal_0(X1)|v3_xxreal_0(k6_normsp_1(X2,X3))|v2_struct_0(X2)|~v1_xreal_0(X1)|~l1_normsp_1(X2)|~v3_membered(u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1592, plain, (r1_xxreal_0(X1,esk47_0)|v2_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4291, plain, (r1_xxreal_0(esk54_0,esk47_0)|~v1_xreal_0(esk54_0))).#trainneg
cnf(i_0_4292, plain, (r1_xxreal_0(esk53_0,esk47_0))).#trainneg
cnf(i_0_4293, plain, (r1_xxreal_0(esk39_0,esk47_0)|~v3_xxreal_0(esk39_0)|~v1_xreal_0(esk39_0))).#trainneg
cnf(i_0_4294, plain, (r1_xxreal_0(esk28_1(esk44_0),esk47_0)|~v3_xxreal_0(esk28_1(esk44_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1594, plain, (r1_xxreal_0(X1,esk28_1(esk44_0))|v2_xxreal_0(X1)|v3_xxreal_0(esk28_1(esk44_0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4323, plain, (r1_xxreal_0(esk53_0,esk28_1(esk44_0))|v3_xxreal_0(esk28_1(esk44_0)))).#trainneg
cnf(i_0_4334, plain, (r1_xxreal_0(esk5_1(X1),esk28_1(esk44_0))|v2_xxreal_0(esk5_1(X1))|v3_xxreal_0(esk28_1(esk44_0))|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4370, plain, (r1_xxreal_0(esk53_0,esk28_1(esk44_0))|~v1_xboole_0(esk28_1(esk44_0)))).#trainneg
cnf(i_0_4378, plain, (v2_xxreal_0(esk28_1(esk44_0))|v3_xxreal_0(esk5_1(X1))|v1_xboole_0(esk28_1(esk44_0))|v2_xxreal_0(esk5_1(X1))|v3_xxreal_0(esk28_1(esk44_0))|~v1_xreal_0(esk5_1(X1))|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1911, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,esk41_1(X2))|~l1_struct_0(X2))).#trainneg
cnf(i_0_3118, plain, (v1_xreal_0(u1_struct_0(esk21_0))|~v2_struct_0(esk21_0))).#trainneg
cnf(i_0_4052, plain, (r1_xxreal_0(np__0,esk55_0))).#trainneg
cnf(i_0_4259, plain, (~v1_xboole_0(X1)|~v7_ordinal1(X2)|~r2_hidden(X3,k1_funct_1(np__0,X2))|~v1_funct_2(np__0,k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(np__0))).#trainneg
cnf(i_0_1607, plain, (r1_xxreal_0(X1,k5_numbers)|v2_xxreal_0(X1)|v3_xxreal_0(k5_numbers)|~v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_3036, negated_conjecture, (v5_relat_1(esk4_0,k4_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_3169, negated_conjecture, (v2_valued_0(esk4_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3384, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1608, plain, (r1_xxreal_0(X1,X2)|v2_xxreal_0(X1)|v3_xxreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1609, plain, (r1_xxreal_0(X1,X2)|v2_xxreal_0(X1)|v3_xxreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4407, plain, (r1_xxreal_0(esk28_1(esk44_0),k5_numbers)|~v3_xxreal_0(esk28_1(esk44_0))|~v2_membered(k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_4443, plain, (r1_xxreal_0(X1,esk62_3(X2,X3,X4))|v2_xxreal_0(X1)|k6_normsp_1(X2,X3)=k4_struct_0(X2)|v2_struct_0(X2)|~v1_xreal_0(X1)|k2_seq_2(X4)!=np__0|~v2_comseq_2(X4)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X4,k5_numbers,k1_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_4445, plain, (r1_xxreal_0(X1,esk11_3(X2,X3,X4))|v2_xxreal_0(X1)|X3=k1_seq_2(X2)|~v1_xreal_0(X1)|~v2_comseq_2(X2)|~v1_xreal_0(X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4446, plain, (r1_xxreal_0(X1,esk9_3(X2,X3,X4))|v2_xxreal_0(X1)|r1_xxreal_0(X4,np__0)|~v1_xreal_0(X1)|X3!=k1_seq_2(X2)|~v2_comseq_2(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4448, plain, (r1_xxreal_0(X1,esk6_2(X2,X3))|v2_xxreal_0(X1)|r1_xxreal_0(X3,np__0)|~v1_xreal_0(X1)|~v2_comseq_2(X2)|~v1_xreal_0(X3)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1597, plain, (r1_xxreal_0(X1,np__0)|v2_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4516, plain, (r1_xxreal_0(esk28_1(esk44_0),np__0)|~v3_xxreal_0(esk28_1(esk44_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1603, plain, (r1_xxreal_0(X1,np__1)|v2_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4537, plain, (r1_xxreal_0(esk28_1(esk44_0),np__1)|~v3_xxreal_0(esk28_1(esk44_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1927, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,esk56_1(X2))|~l1_struct_0(X2))).#trainneg
cnf(i_0_3162, plain, (v1_xreal_0(u1_struct_0(esk20_0))|~v2_struct_0(esk20_0))).#trainneg
cnf(i_0_4262, plain, (~v1_xboole_0(X1)|~v7_ordinal1(X2)|~r2_hidden(X3,k1_funct_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1)),X2))|~v1_funct_2(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1)),k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_1610, plain, (r1_xxreal_0(X1,k1_seq_2(esk61_3(X2,X3,X4)))|v2_xxreal_0(X1)|v3_xxreal_0(k1_seq_2(esk61_3(X2,X3,X4)))|v2_struct_0(X2)|~v1_xreal_0(X1)|~v1_funct_2(esk61_3(X2,X3,X4),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X2,X3,X4))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_3037, negated_conjecture, (v5_relat_1(esk4_0,k3_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1611, plain, (r1_xxreal_0(X1,k1_seq_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))|v2_xxreal_0(X1)|v3_xxreal_0(k1_seq_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))|~v1_xreal_0(X1)|~v1_funct_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))),k5_numbers,k1_numbers)|~v1_funct_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))).#trainneg
cnf(i_0_1612, plain, (r1_xxreal_0(X1,k1_seq_2(k4_normsp_0(X2,X3)))|v2_xxreal_0(X1)|v3_xxreal_0(k1_seq_2(k4_normsp_0(X2,X3)))|v2_struct_0(X2)|~v1_xreal_0(X1)|~v1_funct_2(k4_normsp_0(X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(k4_normsp_0(X2,X3))|~l1_normsp_0(X2)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))))).#trainneg
cnf(i_0_1613, plain, (r1_xxreal_0(X1,k1_seq_2(k3_series_1(X2)))|v2_xxreal_0(X1)|v3_xxreal_0(k1_seq_2(k3_series_1(X2)))|~v1_xreal_0(X1)|~v1_funct_2(k3_series_1(X2),k5_numbers,k1_numbers)|~v1_funct_1(k3_series_1(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1614, plain, (r1_xxreal_0(X1,k1_seq_2(k2_zfmisc_1(k5_numbers,k1_numbers)))|v2_xxreal_0(X1)|v3_xxreal_0(k1_seq_2(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(X1)|~v1_funct_2(k2_zfmisc_1(k5_numbers,k1_numbers),k5_numbers,k1_numbers)|~v1_funct_1(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_1615, plain, (v1_xcmplx_0(esk31_1(X1))|v7_struct_0(X1)|~v1_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1616, plain, (v1_xxreal_0(esk31_1(X1))|v7_struct_0(X1)|~v2_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1617, plain, (v1_xreal_0(esk31_1(X1))|v7_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1618, plain, (v1_rat_1(esk31_1(X1))|v7_struct_0(X1)|~v4_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1619, plain, (v1_int_1(esk31_1(X1))|v7_struct_0(X1)|~v5_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1620, plain, (v7_ordinal1(esk31_1(X1))|v7_struct_0(X1)|~v6_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1621, plain, (v3_ordinal1(esk31_1(X1))|v7_struct_0(X1)|~v3_ordinal1(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1622, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1623, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1624, plain, (v1_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_3741, plain, (v5_membered(np__0)|~v5_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_3193, plain, (v1_xreal_0(u1_struct_0(esk19_0))|~v2_struct_0(esk19_0))).#trainneg
cnf(i_0_4299, plain, (r1_xxreal_0(np__0,esk47_0))).#trainneg
cnf(i_0_4258, plain, (~v1_xboole_0(X1)|~v7_ordinal1(X2)|~r2_hidden(X3,k1_funct_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1)))),X2))|~v1_funct_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1)))),k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1))))))).#trainneg
cnf(i_0_1625, plain, (v1_xreal_0(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_4601, plain, (~v2_xxreal_0(k5_numbers)|~v3_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_3492, plain, (v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3624, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X1)),k5_numbers))).#trainneg
cnf(i_0_1626, plain, (v1_xcmplx_0(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1627, plain, (v1_xreal_0(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1628, plain, (v1_xreal_0(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1629, negated_conjecture, (v1_xreal_0(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1630, negated_conjecture, (v1_xcmplx_0(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1631, plain, (v1_xreal_0(k3_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1633, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1634, negated_conjecture, (v1_xreal_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1635, negated_conjecture, (v1_xreal_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1636, negated_conjecture, (v1_xcmplx_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1637, negated_conjecture, (v1_xcmplx_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1638, negated_conjecture, (v1_xxreal_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1639, negated_conjecture, (v1_xxreal_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1643, plain, (v1_xxreal_0(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3742, plain, (v6_ordinal1(np__0)|~v6_ordinal1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_3205, plain, (v1_xreal_0(u1_struct_0(esk18_0))|~v2_struct_0(esk18_0))).#trainneg
cnf(i_0_4514, plain, (r1_xxreal_0(esk53_0,np__0))).#trainneg
cnf(i_0_4260, plain, (~v1_xboole_0(X1)|~v7_ordinal1(X2)|~r2_hidden(X3,k1_funct_1(k1_funct_1(X4,X5),X2))|~v1_funct_2(k1_funct_1(X4,X5),k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(k1_funct_1(X4,X5))|~v7_ordinal1(X5)|~v1_funct_2(X4,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1))))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1)))))))).#trainneg
cnf(i_0_4651, negated_conjecture, (~v2_xxreal_0(esk3_0)|~v3_xxreal_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1646, plain, (r1_xxreal_0(X1,k4_normsp_0(X2,X3))|v2_xxreal_0(X1)|v3_xxreal_0(k4_normsp_0(X2,X3))|v2_struct_0(X2)|~v1_xreal_0(X1)|~l1_normsp_0(X2)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k4_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))))).#trainneg
cnf(i_0_1648, plain, (r1_xxreal_0(X1,esk61_3(X2,X3,X4))|r1_xxreal_0(esk61_3(X2,X3,X4),X1)|v2_struct_0(X2)|~v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1645, plain, (v2_struct_0(X1)|~v3_xxreal_0(k8_lopban_4(X1,X2,X3,X4))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1649, plain, (r1_xxreal_0(X1,esk59_0)|r1_xxreal_0(esk59_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_4701, plain, (r1_xxreal_0(esk59_0,esk11_3(X1,X2,X3))|r1_xxreal_0(esk11_3(X1,X2,X3),esk59_0)|X2=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1650, plain, (r1_xxreal_0(X1,esk55_0)|r1_xxreal_0(esk55_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_4760, plain, (r1_xxreal_0(esk55_0,esk11_3(X1,X2,X3))|r1_xxreal_0(esk11_3(X1,X2,X3),esk55_0)|X2=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1651, plain, (r1_xxreal_0(X1,esk54_0)|r1_xxreal_0(esk54_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_4819, plain, (r1_xxreal_0(esk54_0,esk11_3(X1,X2,X3))|r1_xxreal_0(esk11_3(X1,X2,X3),esk54_0)|X2=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1652, plain, (r1_xxreal_0(X1,esk53_0)|r1_xxreal_0(esk53_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_4878, plain, (r1_xxreal_0(esk53_0,esk11_3(X1,X2,X3))|r1_xxreal_0(esk11_3(X1,X2,X3),esk53_0)|X2=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4918, plain, (X2=k1_seq_2(X1)|r1_xxreal_0(esk53_0,esk11_3(X1,X2,X3))|~v2_xxreal_0(esk11_3(X1,X2,X3))|~v1_xreal_0(esk11_3(X1,X2,X3))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4923, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk53_0,esk11_3(X2,X1,X3))|r1_xxreal_0(esk11_3(X2,X1,X3),esk59_0)|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(esk11_3(X2,X1,X3)))).#trainneg
cnf(i_0_4924, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk53_0,esk11_3(X2,X1,X3))|r1_xxreal_0(esk11_3(X2,X1,X3),esk55_0)|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(esk11_3(X2,X1,X3)))).#trainneg
cnf(i_0_3743, plain, (v6_membered(np__0)|~v6_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_3228, plain, (v1_xreal_0(u1_struct_0(esk17_0))|~v2_struct_0(esk17_0))).#trainneg
cnf(i_0_4535, plain, (r1_xxreal_0(esk53_0,np__1))).#trainneg
cnf(i_0_4671, negated_conjecture, (~v1_xboole_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~r2_hidden(X4,k1_funct_1(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X3),X2))|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)),k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1))))|~v1_funct_2(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X3),k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X3))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1))))))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)),k5_numbers,u1_struct_0(esk1_0))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1653, plain, (r1_xxreal_0(X1,esk48_0)|r1_xxreal_0(esk48_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_4689, plain, (r1_xxreal_0(esk59_0,esk55_0)|r1_xxreal_0(esk55_0,esk59_0))).#trainneg
cnf(i_0_4655, negated_conjecture, (~v2_xxreal_0(esk2_0)|~v3_xxreal_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_3670, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1654, plain, (r1_xxreal_0(X1,esk47_0)|r1_xxreal_0(esk47_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1655, plain, (r1_xxreal_0(X1,esk44_0)|r1_xxreal_0(esk44_0,X1)|~v1_xxreal_0(X1)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4925, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk53_0,esk11_3(X2,X1,X3))|r1_xxreal_0(esk11_3(X2,X1,X3),esk47_0)|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(esk11_3(X2,X1,X3)))).#trainneg
cnf(i_0_4932, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk53_0,esk11_3(X2,X1,X3))|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_xxreal_0(X3)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_4969, plain, (r1_xxreal_0(esk48_0,esk11_3(X1,X2,X3))|r1_xxreal_0(esk11_3(X1,X2,X3),esk48_0)|X2=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_5040, plain, (r1_xxreal_0(esk47_0,esk11_3(X1,X2,X3))|r1_xxreal_0(esk11_3(X1,X2,X3),esk47_0)|X2=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1656, plain, (r1_xxreal_0(X1,esk39_0)|r1_xxreal_0(esk39_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_5166, plain, (r1_xxreal_0(esk39_0,esk11_3(X1,X2,X3))|r1_xxreal_0(esk11_3(X1,X2,X3),esk39_0)|X2=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1657, plain, (r1_xxreal_0(X1,esk28_1(X2))|r1_xxreal_0(esk28_1(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(X2))).#trainneg
cnf(i_0_1658, plain, (r1_xxreal_0(X1,esk11_3(X2,X3,X4))|r1_xxreal_0(esk11_3(X2,X3,X4),X1)|X3=k1_seq_2(X2)|~v1_xxreal_0(X1)|~v2_comseq_2(X2)|~v1_xreal_0(X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1659, plain, (r1_xxreal_0(X1,esk8_3(X2,X3,X4))|r1_xxreal_0(esk8_3(X2,X3,X4),X1)|v2_comseq_2(X2)|~v1_xxreal_0(X1)|~v1_xreal_0(X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1661, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|r1_xxreal_0(esk3_0,X1)|~v1_xxreal_0(X1)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_3744, plain, (v1_membered(np__0)|~v1_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_3286, plain, (v1_xreal_0(u1_struct_0(esk16_0))|~v2_struct_0(esk16_0))).#trainneg
cnf(i_0_3924, plain, (r1_xxreal_0(esk54_0,esk38_0)|v3_xxreal_0(esk38_0)|~v1_xreal_0(esk54_0))).#trainneg
cnf(i_0_4672, plain, (~v1_xboole_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~r2_hidden(X4,k1_funct_1(k8_nat_1(X6,X5,X3),X2))|~v1_funct_2(X5,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1))))|~v1_funct_2(k8_nat_1(X6,X5,X3),k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(k8_nat_1(X6,X5,X3))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1))))))|~v1_funct_2(X5,k5_numbers,X6)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X6))))).#trainneg
cnf(i_0_4690, plain, (r1_xxreal_0(esk59_0,esk54_0)|r1_xxreal_0(esk54_0,esk59_0))).#trainneg
cnf(i_0_3716, plain, (v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1663, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|r1_xxreal_0(esk2_0,X1)|~v1_xxreal_0(X1)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1664, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_5343, plain, (r1_xxreal_0(esk54_0,esk38_0)|~m1_subset_1(esk38_0,k5_numbers)|~v1_xreal_0(esk54_0))).#trainneg
cnf(i_0_5344, plain, (r1_xxreal_0(esk54_0,esk38_0)|~v1_xboole_0(esk38_0)|~v1_xreal_0(esk54_0))).#trainneg
cnf(i_0_1665, plain, (r1_xxreal_0(X1,k4_normsp_0(X2,X3))|r1_xxreal_0(k4_normsp_0(X2,X3),X1)|v2_struct_0(X2)|~v1_xxreal_0(X1)|~l1_normsp_0(X2)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))))).#trainneg
cnf(i_0_1666, plain, (r1_xxreal_0(X1,k3_series_1(X2))|r1_xxreal_0(k3_series_1(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1667, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1668, plain, (r1_xxreal_0(X1,k4_struct_0(X2))|r1_xxreal_0(k4_struct_0(X2),X1)|~v1_xxreal_0(X1)|~l2_struct_0(X2)|~v2_membered(u1_struct_0(X2)))).#trainneg
cnf(i_0_3745, plain, (v2_membered(np__0)|~v2_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_3322, plain, (v1_xreal_0(u1_struct_0(esk15_0))|~v2_struct_0(esk15_0))).#trainneg
cnf(i_0_1669, plain, (r1_xxreal_0(X1,k4_struct_0(X2))|r1_xxreal_0(k4_struct_0(X2),X1)|~v1_xxreal_0(X1)|~l2_struct_0(X2)|~v6_membered(u1_struct_0(X2)))).#trainneg
cnf(i_0_4692, plain, (r1_xxreal_0(esk59_0,esk48_0)|r1_xxreal_0(esk48_0,esk59_0))).#trainneg
cnf(i_0_4081, plain, (v4_membered(esk52_0)|~v6_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_5562, plain, (v4_membered(esk52_0)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_2970, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X1)|v1_xboole_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_1670, plain, (r1_xxreal_0(X1,k6_normsp_1(X2,X3))|r1_xxreal_0(k6_normsp_1(X2,X3),X1)|v2_struct_0(X2)|~v1_xxreal_0(X1)|~l1_normsp_1(X2)|~v2_membered(u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1671, plain, (r1_xxreal_0(X1,k1_bhsp_4(X2,X3))|r1_xxreal_0(k1_bhsp_4(X2,X3),X1)|v2_struct_0(X2)|~v1_xxreal_0(X1)|~l2_algstr_0(X2)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))))).#trainneg
cnf(i_0_1672, plain, (r1_xxreal_0(X1,k8_lopban_4(X2,X3,X4,X5))|r1_xxreal_0(k8_lopban_4(X2,X3,X4,X5),X1)|v2_struct_0(X2)|~v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,k5_numbers)|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1673, plain, (r1_xxreal_0(X1,k5_numbers)|r1_xxreal_0(k5_numbers,X1)|~v1_xxreal_0(X1)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1674, plain, (r1_xxreal_0(X1,k5_numbers)|r1_xxreal_0(k5_numbers,X1)|~v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_1675, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1676, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1677, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_5842, plain, (r1_xxreal_0(X1,esk11_3(X2,X3,X4))|r1_xxreal_0(esk11_3(X2,X3,X4),X1)|X3=k1_seq_2(X2)|~m1_subset_1(X1,k5_numbers)|~v2_comseq_2(X2)|~v1_xreal_0(X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1678, plain, (v5_membered(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v5_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1680, plain, (v5_membered(esk28_1(k1_zfmisc_1(X1)))|~v5_membered(X1))).#trainneg
cnf(i_0_1681, negated_conjecture, (v5_membered(esk4_0)|~v5_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1682, plain, (v5_membered(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v5_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_3746, plain, (v3_membered(np__0)|~v3_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_3381, plain, (v1_xreal_0(u1_struct_0(esk14_0))|~v2_struct_0(esk14_0))).#trainneg
cnf(i_0_1683, plain, (v5_membered(k3_series_1(X1))|~v5_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4693, plain, (r1_xxreal_0(esk59_0,esk47_0)|r1_xxreal_0(esk47_0,esk59_0))).#trainneg
cnf(i_0_1691, negated_conjecture, (v6_ordinal1(esk4_0)|~v6_ordinal1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1701, negated_conjecture, (v6_membered(esk4_0)|~v6_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1684, plain, (v5_membered(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v5_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1685, plain, (v5_membered(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v5_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1688, plain, (v6_ordinal1(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v6_ordinal1(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1690, plain, (v6_ordinal1(esk28_1(k1_zfmisc_1(X1)))|~v6_ordinal1(X1))).#trainneg
cnf(i_0_1692, plain, (v6_ordinal1(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v6_ordinal1(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1693, plain, (v6_ordinal1(k3_series_1(X1))|~v6_ordinal1(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1694, plain, (v6_ordinal1(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v6_ordinal1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1695, plain, (v6_ordinal1(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v6_ordinal1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1698, plain, (v6_membered(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v6_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1700, plain, (v6_membered(esk28_1(k1_zfmisc_1(X1)))|~v6_membered(X1))).#trainneg
cnf(i_0_1702, plain, (v6_membered(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v6_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1703, plain, (v6_membered(k3_series_1(X1))|~v6_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1704, plain, (v6_membered(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v6_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_3747, plain, (v4_membered(np__0)|~v4_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_3422, plain, (v1_xreal_0(u1_struct_0(esk13_0))|~v2_struct_0(esk13_0))).#trainneg
cnf(i_0_1705, plain, (v6_membered(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v6_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_4696, plain, (r1_xxreal_0(esk59_0,esk39_0)|r1_xxreal_0(esk39_0,esk59_0))).#trainneg
cnf(i_0_1711, negated_conjecture, (v1_membered(esk4_0)|~v1_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1721, negated_conjecture, (v2_membered(esk4_0)|~v2_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1708, plain, (v1_membered(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1710, plain, (v1_membered(esk28_1(k1_zfmisc_1(X1)))|~v1_membered(X1))).#trainneg
cnf(i_0_1712, plain, (v1_membered(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v1_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1713, plain, (v1_membered(k3_series_1(X1))|~v1_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1714, plain, (v1_membered(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v1_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1715, plain, (v1_membered(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v1_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1718, plain, (v2_membered(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v2_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1720, plain, (v2_membered(esk28_1(k1_zfmisc_1(X1)))|~v2_membered(X1))).#trainneg
cnf(i_0_1722, plain, (v2_membered(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v2_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1723, plain, (v2_membered(k3_series_1(X1))|~v2_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1724, plain, (v2_membered(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v2_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1725, plain, (v2_membered(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v2_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1728, plain, (v3_membered(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1730, plain, (v3_membered(esk28_1(k1_zfmisc_1(X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_3489, plain, (v1_xreal_0(u1_struct_0(esk12_0))|~v2_struct_0(esk12_0))).#trainneg
cnf(i_0_1731, negated_conjecture, (v3_membered(esk4_0)|~v3_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_4720, plain, (r1_xxreal_0(esk59_0,np__1)|r1_xxreal_0(np__1,esk59_0))).#trainneg
cnf(i_0_1741, negated_conjecture, (v4_membered(esk4_0)|~v4_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2052, negated_conjecture, (v1_valued_0(esk4_0)|~v1_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2608, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2619, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1732, plain, (v3_membered(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1733, plain, (v3_membered(k3_series_1(X1))|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1734, plain, (v3_membered(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v3_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1735, plain, (v3_membered(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v3_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1738, plain, (v4_membered(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v4_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1740, plain, (v4_membered(esk28_1(k1_zfmisc_1(X1)))|~v4_membered(X1))).#trainneg
cnf(i_0_1742, plain, (v4_membered(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v4_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1743, plain, (v4_membered(k3_series_1(X1))|~v4_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1744, plain, (v4_membered(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v4_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1745, plain, (v4_membered(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v4_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1753, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))|v2_struct_0(X2)|v1_xboole_0(k8_lopban_4(X2,X3,X4,X5))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,k5_numbers)|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)|~m1_subset_1(X1,k8_lopban_4(X2,X3,X4,X5)))).#trainneg
cnf(i_0_1754, negated_conjecture, (v1_xxreal_0(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1755, negated_conjecture, (v1_xxreal_0(esk4_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1756, negated_conjecture, (v1_xxreal_0(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1767, plain, (v5_relat_1(k2_zfmisc_1(X1,k4_numbers),k4_numbers))).#trainneg
cnf(i_0_1757, plain, (r1_xxreal_0(X1,k6_xcmplx_0(X2,X3))|v2_xxreal_0(X1)|v3_xxreal_0(k6_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2065, negated_conjecture, (v2_valued_0(esk4_0)|~v2_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3237, plain, (v3_ordinal1(X1)|~v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_3238, negated_conjecture, (v3_ordinal1(esk3_0)|~v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1766, plain, (v5_relat_1(k2_zfmisc_1(X1,esk44_0),k4_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1768, plain, (v5_relat_1(k2_zfmisc_1(X1,k5_numbers),k4_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1765, plain, (v1_xboole_0(X2)|~r2_hidden(X2,X1)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_1770, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),k4_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1771, plain, (v5_relat_1(k2_zfmisc_1(X1,k3_numbers),k3_numbers))).#trainneg
cnf(i_0_1772, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),k3_numbers)|~v5_membered(X2))).#trainneg
cnf(i_0_1773, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),k3_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1783, plain, (v3_xxreal_0(esk10_2(X1,np__0))|k1_seq_2(X1)=np__0|~v3_xxreal_0(k18_complex1(k1_funct_1(X1,esk11_3(X1,np__0,X2))))|~v1_xreal_0(k18_complex1(k1_funct_1(X1,esk11_3(X1,np__0,X2))))|~v1_xreal_0(esk10_2(X1,np__0))|~v2_comseq_2(X1)|~v3_valued_0(X1)|~v1_xcmplx_0(k1_funct_1(X1,esk11_3(X1,np__0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk11_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1784, plain, (v3_xxreal_0(esk10_2(X1,np__0))|k1_seq_2(X1)=np__0|~v3_xxreal_0(k18_complex1(k1_seq_1(X1,esk11_3(X1,np__0,X2))))|~v1_xreal_0(k18_complex1(k1_seq_1(X1,esk11_3(X1,np__0,X2))))|~v1_xreal_0(esk10_2(X1,np__0))|~v2_comseq_2(X1)|~v1_xcmplx_0(k1_seq_1(X1,esk11_3(X1,np__0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_1(X1,esk11_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1785, plain, (v3_xxreal_0(esk10_2(X1,X2))|X2=k1_seq_2(X1)|~v3_xxreal_0(k18_complex1(k6_xcmplx_0(k1_seq_1(X1,esk11_3(X1,X2,X3)),X2)))|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k1_seq_1(X1,esk11_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk10_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_1(X1,esk11_3(X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1786, plain, (v3_xxreal_0(esk10_2(X1,X2))|X2=k1_seq_2(X1)|~v3_xxreal_0(k18_complex1(k6_xcmplx_0(k1_funct_1(X1,esk11_3(X1,X2,X3)),X2)))|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k1_funct_1(X1,esk11_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk10_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk11_3(X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1787, plain, (v3_xxreal_0(esk10_2(X1,X2))|X2=k1_seq_2(X1)|~v3_xxreal_0(k18_complex1(k9_real_1(k1_funct_1(X1,esk11_3(X1,X2,X3)),X2)))|~v1_xreal_0(k18_complex1(k9_real_1(k1_funct_1(X1,esk11_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk10_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1788, plain, (v3_xxreal_0(esk10_2(X1,X2))|X2=k1_seq_2(X1)|~v3_xxreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk11_3(X1,X2,X3)),X2)))|~v1_xreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk11_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk10_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2645, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_4261, plain, (~v1_xboole_0(X1)|~v7_ordinal1(X2)|~r2_hidden(X3,k1_funct_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1))),X4,X5),X2))|~v1_funct_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1))),X4,X5),k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1))),X4,X5))|~v7_ordinal1(X5)|~v1_funct_2(X4,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1))))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1)))))))).#trainneg
cnf(i_0_1790, plain, (v3_xxreal_0(esk7_2(X1,np__0))|v2_comseq_2(X1)|~v3_xxreal_0(k18_complex1(k1_funct_1(X1,esk8_3(X1,np__0,X2))))|~v1_xreal_0(k18_complex1(k1_funct_1(X1,esk8_3(X1,np__0,X2))))|~v1_xreal_0(esk7_2(X1,np__0))|~v3_valued_0(X1)|~v1_xcmplx_0(k1_funct_1(X1,esk8_3(X1,np__0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk8_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4749, plain, (r1_xxreal_0(esk55_0,esk54_0)|r1_xxreal_0(esk54_0,esk55_0))).#trainneg
cnf(i_0_2072, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2073, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_3239, negated_conjecture, (v3_ordinal1(esk2_0)|~v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_5914, negated_conjecture, (v6_ordinal1(esk4_0)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1791, plain, (v3_xxreal_0(esk7_2(X1,np__0))|v2_comseq_2(X1)|~v3_xxreal_0(k18_complex1(k1_seq_1(X1,esk8_3(X1,np__0,X2))))|~v1_xreal_0(k18_complex1(k1_seq_1(X1,esk8_3(X1,np__0,X2))))|~v1_xreal_0(esk7_2(X1,np__0))|~v1_xcmplx_0(k1_seq_1(X1,esk8_3(X1,np__0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_1(X1,esk8_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1792, plain, (v3_xxreal_0(esk7_2(X1,X2))|v2_comseq_2(X1)|~v3_xxreal_0(k18_complex1(k6_xcmplx_0(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk7_2(X1,X2))|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1793, plain, (v3_xxreal_0(esk7_2(X1,X2))|v2_comseq_2(X1)|~v3_xxreal_0(k18_complex1(k6_xcmplx_0(k1_seq_1(X1,esk8_3(X1,X2,X3)),X2)))|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k1_seq_1(X1,esk8_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk7_2(X1,X2))|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_1(X1,esk8_3(X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1794, plain, (v3_xxreal_0(esk7_2(X1,X2))|v2_comseq_2(X1)|~v3_xxreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk8_3(X1,X2,X3)),X2)))|~v1_xreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk8_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk7_2(X1,X2))|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1795, plain, (v3_xxreal_0(esk7_2(X1,X2))|v2_comseq_2(X1)|~v3_xxreal_0(k18_complex1(k9_real_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|~v1_xreal_0(k18_complex1(k9_real_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk7_2(X1,X2))|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1809, negated_conjecture, (v3_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v3_xxreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1810, negated_conjecture, (v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v3_xxreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_6114, negated_conjecture, (~m1_subset_1(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k5_numbers)|~v3_xxreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v1_xreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_6115, negated_conjecture, (~v1_xboole_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v3_xxreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v1_xreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2656, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_1833, plain, (v2_xxreal_0(k18_complex1(k1_funct_1(X1,esk11_3(X1,np__0,X2))))|k1_seq_2(X1)=np__0|~v2_xxreal_0(esk10_2(X1,np__0))|~v1_xreal_0(k18_complex1(k1_funct_1(X1,esk11_3(X1,np__0,X2))))|~v1_xreal_0(esk10_2(X1,np__0))|~v2_comseq_2(X1)|~v3_valued_0(X1)|~v1_xcmplx_0(k1_funct_1(X1,esk11_3(X1,np__0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk11_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4751, plain, (r1_xxreal_0(esk55_0,esk48_0)|r1_xxreal_0(esk48_0,esk55_0))).#trainneg
cnf(i_0_2074, plain, (v1_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_6134, plain, (~v2_xxreal_0(k5_numbers)|~v3_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_5915, negated_conjecture, (v6_membered(esk4_0)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_5971, plain, (v1_xreal_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1812, plain, (v3_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|v2_struct_0(X1)|~v3_xxreal_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|~v1_xreal_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|~v1_xreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1818, plain, (v3_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,X2),X3)))|v2_struct_0(X1)|~v3_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|~v1_xreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,X2),X3)))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_6119, negated_conjecture, (~v3_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k5_numbers)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_relat_1(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_6120, negated_conjecture, (~v3_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xboole_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_relat_1(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_1811, negated_conjecture, (v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)))|~v3_xxreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v1_xreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)))|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X1)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_1813, negated_conjecture, (v3_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v3_xxreal_0(k8_nat_1(X2,esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k8_nat_1(X2,esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v7_ordinal1(X1)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X2)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_1814, negated_conjecture, (v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v3_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1815, negated_conjecture, (v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)))|~v3_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),X2))|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),X2))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)))|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X1)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_1816, negated_conjecture, (v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v3_xxreal_0(k8_nat_1(X2,esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k8_nat_1(X2,esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v7_ordinal1(X1)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X2)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_1817, negated_conjecture, (v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)))|~v3_xxreal_0(k8_nat_1(X3,esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v1_xreal_0(k8_nat_1(X3,esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)))|~v7_ordinal1(X2)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X3)|~v1_funct_2(esk4_0,k5_numbers,X1)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_1834, plain, (v2_xxreal_0(k18_complex1(k1_seq_1(X1,esk11_3(X1,np__0,X2))))|k1_seq_2(X1)=np__0|~v2_xxreal_0(esk10_2(X1,np__0))|~v1_xreal_0(k18_complex1(k1_seq_1(X1,esk11_3(X1,np__0,X2))))|~v1_xreal_0(esk10_2(X1,np__0))|~v2_comseq_2(X1)|~v1_xcmplx_0(k1_seq_1(X1,esk11_3(X1,np__0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_1(X1,esk11_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1835, plain, (v2_xxreal_0(k18_complex1(k6_xcmplx_0(k1_seq_1(X1,esk11_3(X1,X2,X3)),X2)))|X2=k1_seq_2(X1)|~v2_xxreal_0(esk10_2(X1,X2))|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k1_seq_1(X1,esk11_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk10_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_1(X1,esk11_3(X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2755, plain, (v1_partfun1(esk32_2(X1,X1),X1))).#trainneg
cnf(i_0_1836, plain, (v2_xxreal_0(k18_complex1(k6_xcmplx_0(k1_funct_1(X1,esk11_3(X1,X2,X3)),X2)))|X2=k1_seq_2(X1)|~v2_xxreal_0(esk10_2(X1,X2))|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k1_funct_1(X1,esk11_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk10_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk11_3(X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4752, plain, (r1_xxreal_0(esk55_0,esk47_0)|r1_xxreal_0(esk47_0,esk55_0))).#trainneg
cnf(i_0_5974, plain, (v1_xreal_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2972, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_1837, plain, (v2_xxreal_0(k18_complex1(k9_real_1(k1_funct_1(X1,esk11_3(X1,X2,X3)),X2)))|X2=k1_seq_2(X1)|~v2_xxreal_0(esk10_2(X1,X2))|~v1_xreal_0(k18_complex1(k9_real_1(k1_funct_1(X1,esk11_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk10_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1838, plain, (v2_xxreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk11_3(X1,X2,X3)),X2)))|X2=k1_seq_2(X1)|~v2_xxreal_0(esk10_2(X1,X2))|~v1_xreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk11_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk10_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1840, plain, (v2_xxreal_0(k18_complex1(k1_funct_1(X1,esk8_3(X1,np__0,X2))))|v2_comseq_2(X1)|~v2_xxreal_0(esk7_2(X1,np__0))|~v1_xreal_0(k18_complex1(k1_funct_1(X1,esk8_3(X1,np__0,X2))))|~v1_xreal_0(esk7_2(X1,np__0))|~v3_valued_0(X1)|~v1_xcmplx_0(k1_funct_1(X1,esk8_3(X1,np__0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk8_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1841, plain, (v2_xxreal_0(k18_complex1(k1_seq_1(X1,esk8_3(X1,np__0,X2))))|v2_comseq_2(X1)|~v2_xxreal_0(esk7_2(X1,np__0))|~v1_xreal_0(k18_complex1(k1_seq_1(X1,esk8_3(X1,np__0,X2))))|~v1_xreal_0(esk7_2(X1,np__0))|~v1_xcmplx_0(k1_seq_1(X1,esk8_3(X1,np__0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_1(X1,esk8_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1842, plain, (v2_xxreal_0(k18_complex1(k6_xcmplx_0(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|v2_comseq_2(X1)|~v2_xxreal_0(esk7_2(X1,X2))|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk7_2(X1,X2))|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1843, plain, (v2_xxreal_0(k18_complex1(k6_xcmplx_0(k1_seq_1(X1,esk8_3(X1,X2,X3)),X2)))|v2_comseq_2(X1)|~v2_xxreal_0(esk7_2(X1,X2))|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k1_seq_1(X1,esk8_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk7_2(X1,X2))|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_1(X1,esk8_3(X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1844, plain, (v2_xxreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk8_3(X1,X2,X3)),X2)))|v2_comseq_2(X1)|~v2_xxreal_0(esk7_2(X1,X2))|~v1_xreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk8_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk7_2(X1,X2))|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1845, plain, (v2_xxreal_0(k18_complex1(k9_real_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|v2_comseq_2(X1)|~v2_xxreal_0(esk7_2(X1,X2))|~v1_xreal_0(k18_complex1(k9_real_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk7_2(X1,X2))|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1859, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v2_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v1_xreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_6277, negated_conjecture, (~v1_xboole_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v2_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v7_ordinal1(X1)|~v1_xreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_6278, negated_conjecture, (~v3_xxreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v2_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v7_ordinal1(X1)|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_6280, negated_conjecture, (~v2_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~v1_xboole_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v1_funct_2(esk4_0,k5_numbers,X2)|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_6283, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),esk59_0)|~v1_xboole_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_4603, plain, (v5_membered(np__0)|~v6_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_3868, plain, (v5_relat_1(k2_zfmisc_1(X1,esk50_0),k4_numbers))).#trainneg
cnf(i_0_1872, plain, (v2_xxreal_0(esk8_3(X1,X2,X3))|v2_comseq_2(X1)|~v2_xxreal_0(X3)|~v1_xreal_0(esk8_3(X1,X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4755, plain, (r1_xxreal_0(esk55_0,esk39_0)|r1_xxreal_0(esk39_0,esk55_0))).#trainneg
cnf(i_0_2075, plain, (v1_xreal_0(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1862, plain, (v2_xxreal_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|v2_struct_0(X1)|~v2_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|~v1_xreal_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|~v1_xreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1868, plain, (v2_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|v2_struct_0(X1)|~v2_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,X2),X3)))|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|~v1_xreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,X2),X3)))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_6284, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),esk55_0)|~v1_xboole_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_6285, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),esk47_0)|~v1_xboole_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_6292, negated_conjecture, (~v2_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v3_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_relat_1(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_6298, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)),esk59_0)|~v1_xboole_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v7_ordinal1(X2)|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X1)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_6299, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)),esk55_0)|~v1_xboole_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v7_ordinal1(X2)|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X1)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_6300, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)),esk47_0)|~v1_xboole_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v7_ordinal1(X2)|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X1)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_6315, plain, (v2_comseq_2(X1)|~v1_xboole_0(esk8_3(X1,X2,X3))|~v2_xxreal_0(X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4661, plain, (v6_ordinal1(np__0)|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_3870, plain, (v5_relat_1(k2_zfmisc_1(X1,esk43_0),k4_numbers))).#trainneg
cnf(i_0_4779, plain, (r1_xxreal_0(esk55_0,np__1)|r1_xxreal_0(np__1,esk55_0))).#trainneg
cnf(i_0_2076, plain, (v1_xcmplx_0(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2085, negated_conjecture, (v1_xreal_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_3822, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X1)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1882, plain, (v3_xxreal_0(k1_normsp_0(X1,k1_funct_1(k1_bhsp_4(X1,X2),X3)))|v2_struct_0(X1)|~v3_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|~v1_xreal_0(k1_normsp_0(X1,k1_funct_1(k1_bhsp_4(X1,X2),X3)))|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1883, plain, (v2_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|v2_struct_0(X1)|~v2_xxreal_0(k1_normsp_0(X1,k1_funct_1(k1_bhsp_4(X1,X2),X3)))|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|~v1_xreal_0(k1_normsp_0(X1,k1_funct_1(k1_bhsp_4(X1,X2),X3)))|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_6363, negated_conjecture, (~v2_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~v3_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_relat_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v1_funct_2(esk4_0,k5_numbers,X2)|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_6366, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),esk59_0)|~v3_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_relat_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_6367, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),esk55_0)|~v3_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_relat_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_6368, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),esk47_0)|~v3_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~v1_relat_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1860, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v2_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v1_xreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1861, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v2_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~v1_xreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X2)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_1863, negated_conjecture, (v2_xxreal_0(k8_nat_1(X1,esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v2_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X2)))|~v1_xreal_0(k8_nat_1(X1,esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X2)))|~v7_ordinal1(X2)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X1)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_1864, negated_conjecture, (v2_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v2_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1865, negated_conjecture, (v2_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v2_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X2)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_6396, plain, (v6_ordinal1(np__0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_6487, plain, (v6_ordinal1(np__0))).#trainneg
cnf(i_0_1895, plain, (esk5_1(X1)=k1_seq_2(X1)|r1_xxreal_0(esk10_2(X1,esk5_1(X1)),np__0)|~v2_comseq_2(X1)|~v1_xreal_0(esk5_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(esk10_2(X1,esk5_1(X1))))).#trainneg
cnf(i_0_2086, negated_conjecture, (v1_xreal_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2087, negated_conjecture, (v1_xcmplx_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_3823, negated_conjecture, (v3_ordinal1(esk3_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3824, negated_conjecture, (v3_ordinal1(esk2_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1884, plain, (r1_xxreal_0(k1_normsp_0(X1,k1_funct_1(k1_bhsp_4(X1,X2),X3)),k8_nat_1(X4,k3_series_1(k4_normsp_0(X1,X2)),X3))|v2_struct_0(X1)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_funct_2(k3_series_1(k4_normsp_0(X1,X2)),k5_numbers,k1_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(X1,X2)),k5_numbers,X4)|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2)))|~m1_subset_1(k3_series_1(k4_normsp_0(X1,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_series_1(k4_normsp_0(X1,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4))))).#trainneg
cnf(i_0_1866, negated_conjecture, (v2_xxreal_0(k8_nat_1(X1,esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v2_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2)))|~v1_xreal_0(k8_nat_1(X1,esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2)))|~v7_ordinal1(X2)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X1)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_1867, negated_conjecture, (v2_xxreal_0(k8_nat_1(X1,esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v2_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X3,esk4_0,X2)))|~v1_xreal_0(k8_nat_1(X1,esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X3,esk4_0,X2)))|~v7_ordinal1(X2)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X1)|~v1_funct_2(esk4_0,k5_numbers,X3)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_1896, plain, (v1_xcmplx_0(esk41_1(X1))|v2_struct_0(X1)|~v1_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1897, plain, (v1_xxreal_0(esk41_1(X1))|v2_struct_0(X1)|~v2_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1898, plain, (v1_xreal_0(esk41_1(X1))|v2_struct_0(X1)|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_4941, plain, (v6_membered(np__0)|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_3871, plain, (v5_relat_1(k2_zfmisc_1(X1,esk33_0),k4_numbers))).#trainneg
cnf(i_0_3517, plain, (v1_xxreal_0(esk28_1(esk44_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1899, plain, (v1_rat_1(esk41_1(X1))|v2_struct_0(X1)|~v4_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_2088, negated_conjecture, (v1_xcmplx_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_4631, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1))),k5_numbers))).#trainneg
cnf(i_0_4632, negated_conjecture, (v3_ordinal1(esk3_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))),k5_numbers))).#trainneg
cnf(i_0_1900, plain, (v1_int_1(esk41_1(X1))|v2_struct_0(X1)|~v5_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1901, plain, (v7_ordinal1(esk41_1(X1))|v2_struct_0(X1)|~v6_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1902, plain, (v3_ordinal1(esk41_1(X1))|v2_struct_0(X1)|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1903, plain, (v5_membered(esk41_1(X1))|v2_struct_0(X1)|~v5_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1904, plain, (v6_ordinal1(esk41_1(X1))|v2_struct_0(X1)|~v6_ordinal1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1905, plain, (v6_membered(esk41_1(X1))|v2_struct_0(X1)|~v6_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1906, plain, (v1_membered(esk41_1(X1))|v2_struct_0(X1)|~v1_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1907, plain, (v2_membered(esk41_1(X1))|v2_struct_0(X1)|~v2_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1908, plain, (v3_membered(esk41_1(X1))|v2_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1909, plain, (v4_membered(esk41_1(X1))|v2_struct_0(X1)|~v4_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1910, plain, (v1_xboole_0(esk41_1(X1))|v1_xboole_0(u1_struct_0(X1))|m1_subset_1(X2,u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,esk41_1(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1912, plain, (v1_xcmplx_0(esk56_1(X1))|v2_struct_0(X1)|~v1_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1913, plain, (v1_xxreal_0(esk56_1(X1))|v2_struct_0(X1)|~v2_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1914, plain, (v1_xreal_0(esk56_1(X1))|v2_struct_0(X1)|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_6581, plain, (v6_membered(np__0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_6651, plain, (v6_membered(np__0))).#trainneg
cnf(i_0_4080, plain, (r1_xxreal_0(np__0,esk38_0)|~v1_xboole_0(esk38_0))).#trainneg
cnf(i_0_1915, plain, (v1_rat_1(esk56_1(X1))|v2_struct_0(X1)|~v4_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_4809, plain, (r1_xxreal_0(esk54_0,esk53_0)|r1_xxreal_0(esk53_0,esk54_0))).#trainneg
cnf(i_0_4633, negated_conjecture, (v3_ordinal1(esk2_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))),k5_numbers))).#trainneg
cnf(i_0_5916, negated_conjecture, (v6_membered(esk4_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_1916, plain, (v1_int_1(esk56_1(X1))|v2_struct_0(X1)|~v5_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1917, plain, (v7_ordinal1(esk56_1(X1))|v2_struct_0(X1)|~v6_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1918, plain, (v3_ordinal1(esk56_1(X1))|v2_struct_0(X1)|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1919, plain, (v5_membered(esk56_1(X1))|v2_struct_0(X1)|~v5_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1920, plain, (v6_ordinal1(esk56_1(X1))|v2_struct_0(X1)|~v6_ordinal1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1921, plain, (v6_membered(esk56_1(X1))|v2_struct_0(X1)|~v6_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1922, plain, (v1_membered(esk56_1(X1))|v2_struct_0(X1)|~v1_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1923, plain, (v2_membered(esk56_1(X1))|v2_struct_0(X1)|~v2_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1924, plain, (v3_membered(esk56_1(X1))|v2_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1925, plain, (v4_membered(esk56_1(X1))|v2_struct_0(X1)|~v4_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1926, plain, (v1_xboole_0(esk56_1(X1))|v1_xboole_0(u1_struct_0(X1))|m1_subset_1(X2,u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,esk56_1(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1928, plain, (v1_xcmplx_0(esk42_1(X1))|v7_struct_0(X1)|~v1_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1929, plain, (v1_xxreal_0(esk42_1(X1))|v7_struct_0(X1)|~v2_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_5337, plain, (v1_membered(np__0)|~v3_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_3872, plain, (v5_relat_1(k2_zfmisc_1(X1,k5_numbers),k4_numbers))).#trainneg
cnf(i_0_4513, plain, (r1_xxreal_0(esk54_0,np__0)|~v1_xreal_0(esk54_0))).#trainneg
cnf(i_0_1930, plain, (v1_xreal_0(esk42_1(X1))|v7_struct_0(X1)|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_4810, plain, (r1_xxreal_0(esk54_0,esk48_0)|r1_xxreal_0(esk48_0,esk54_0))).#trainneg
cnf(i_0_2089, negated_conjecture, (v1_xxreal_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1931, plain, (v1_rat_1(esk42_1(X1))|v7_struct_0(X1)|~v4_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1932, plain, (v1_int_1(esk42_1(X1))|v7_struct_0(X1)|~v5_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1933, plain, (v7_ordinal1(esk42_1(X1))|v7_struct_0(X1)|~v6_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1934, plain, (v3_ordinal1(esk42_1(X1))|v7_struct_0(X1)|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1935, plain, (v5_membered(esk42_1(X1))|v7_struct_0(X1)|~v5_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1936, plain, (v6_ordinal1(esk42_1(X1))|v7_struct_0(X1)|~v6_ordinal1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1937, plain, (v6_membered(esk42_1(X1))|v7_struct_0(X1)|~v6_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1938, plain, (v1_membered(esk42_1(X1))|v7_struct_0(X1)|~v1_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1939, plain, (v2_membered(esk42_1(X1))|v7_struct_0(X1)|~v2_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1940, plain, (v3_membered(esk42_1(X1))|v7_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1941, plain, (v4_membered(esk42_1(X1))|v7_struct_0(X1)|~v4_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1942, plain, (v1_xboole_0(esk42_1(X1))|v1_xboole_0(u1_struct_0(X1))|m1_subset_1(X2,u1_struct_0(X1))|v7_struct_0(X1)|~m1_subset_1(X2,esk42_1(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1943, plain, (m1_subset_1(X1,u1_struct_0(X2))|v7_struct_0(X2)|~r2_hidden(X1,esk42_1(X2))|~l1_struct_0(X2))).#trainneg
cnf(i_0_5549, plain, (v2_membered(np__0)|~v3_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_3877, plain, (v5_relat_1(k2_zfmisc_1(X1,esk50_0),k3_numbers))).#trainneg
cnf(i_0_4534, plain, (r1_xxreal_0(esk54_0,np__1)|~v1_xreal_0(esk54_0))).#trainneg
cnf(i_0_1952, plain, (v2_xxreal_0(k18_complex1(k1_funct_1(X1,esk11_3(X1,np__0,X2))))|v3_xxreal_0(esk10_2(X1,np__0))|v1_xboole_0(k18_complex1(k1_funct_1(X1,esk11_3(X1,np__0,X2))))|k1_seq_2(X1)=np__0|~v1_xreal_0(k18_complex1(k1_funct_1(X1,esk11_3(X1,np__0,X2))))|~v1_xreal_0(esk10_2(X1,np__0))|~v2_comseq_2(X1)|~v3_valued_0(X1)|~v1_xcmplx_0(k1_funct_1(X1,esk11_3(X1,np__0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk11_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4811, plain, (r1_xxreal_0(esk54_0,esk47_0)|r1_xxreal_0(esk47_0,esk54_0))).#trainneg
cnf(i_0_6742, negated_conjecture, (~v2_xxreal_0(esk3_0)|~v3_xxreal_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2090, negated_conjecture, (v1_xxreal_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_3531, negated_conjecture, (v2_xxreal_0(esk3_0)|v3_xxreal_0(esk3_0)|v1_xboole_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3535, negated_conjecture, (v2_xxreal_0(esk2_0)|v3_xxreal_0(esk2_0)|v1_xboole_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1953, plain, (v2_xxreal_0(k18_complex1(k1_seq_1(X1,esk11_3(X1,np__0,X2))))|v3_xxreal_0(esk10_2(X1,np__0))|v1_xboole_0(k18_complex1(k1_seq_1(X1,esk11_3(X1,np__0,X2))))|k1_seq_2(X1)=np__0|~v1_xreal_0(k18_complex1(k1_seq_1(X1,esk11_3(X1,np__0,X2))))|~v1_xreal_0(esk10_2(X1,np__0))|~v2_comseq_2(X1)|~v1_xcmplx_0(k1_seq_1(X1,esk11_3(X1,np__0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_1(X1,esk11_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1954, plain, (v2_xxreal_0(k18_complex1(k6_xcmplx_0(k1_seq_1(X1,esk11_3(X1,X2,X3)),X2)))|v3_xxreal_0(esk10_2(X1,X2))|v1_xboole_0(k18_complex1(k6_xcmplx_0(k1_seq_1(X1,esk11_3(X1,X2,X3)),X2)))|X2=k1_seq_2(X1)|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k1_seq_1(X1,esk11_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk10_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_1(X1,esk11_3(X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1955, plain, (v2_xxreal_0(k18_complex1(k6_xcmplx_0(k1_funct_1(X1,esk11_3(X1,X2,X3)),X2)))|v3_xxreal_0(esk10_2(X1,X2))|v1_xboole_0(k18_complex1(k6_xcmplx_0(k1_funct_1(X1,esk11_3(X1,X2,X3)),X2)))|X2=k1_seq_2(X1)|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k1_funct_1(X1,esk11_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk10_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk11_3(X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1956, plain, (v2_xxreal_0(k18_complex1(k9_real_1(k1_funct_1(X1,esk11_3(X1,X2,X3)),X2)))|v3_xxreal_0(esk10_2(X1,X2))|v1_xboole_0(k18_complex1(k9_real_1(k1_funct_1(X1,esk11_3(X1,X2,X3)),X2)))|X2=k1_seq_2(X1)|~v1_xreal_0(k18_complex1(k9_real_1(k1_funct_1(X1,esk11_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk10_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1957, plain, (v2_xxreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk11_3(X1,X2,X3)),X2)))|v3_xxreal_0(esk10_2(X1,X2))|v1_xboole_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk11_3(X1,X2,X3)),X2)))|X2=k1_seq_2(X1)|~v1_xreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk11_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk10_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1959, plain, (v2_xxreal_0(k18_complex1(k1_funct_1(X1,esk8_3(X1,np__0,X2))))|v3_xxreal_0(esk7_2(X1,np__0))|v1_xboole_0(k18_complex1(k1_funct_1(X1,esk8_3(X1,np__0,X2))))|v2_comseq_2(X1)|~v1_xreal_0(k18_complex1(k1_funct_1(X1,esk8_3(X1,np__0,X2))))|~v1_xreal_0(esk7_2(X1,np__0))|~v3_valued_0(X1)|~v1_xcmplx_0(k1_funct_1(X1,esk8_3(X1,np__0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk8_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1960, plain, (v2_xxreal_0(k18_complex1(k1_seq_1(X1,esk8_3(X1,np__0,X2))))|v3_xxreal_0(esk7_2(X1,np__0))|v1_xboole_0(k18_complex1(k1_seq_1(X1,esk8_3(X1,np__0,X2))))|v2_comseq_2(X1)|~v1_xreal_0(k18_complex1(k1_seq_1(X1,esk8_3(X1,np__0,X2))))|~v1_xreal_0(esk7_2(X1,np__0))|~v1_xcmplx_0(k1_seq_1(X1,esk8_3(X1,np__0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_1(X1,esk8_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1961, plain, (v2_xxreal_0(k18_complex1(k6_xcmplx_0(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|v3_xxreal_0(esk7_2(X1,X2))|v1_xboole_0(k18_complex1(k6_xcmplx_0(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|v2_comseq_2(X1)|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk7_2(X1,X2))|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1962, plain, (v2_xxreal_0(k18_complex1(k6_xcmplx_0(k1_seq_1(X1,esk8_3(X1,X2,X3)),X2)))|v3_xxreal_0(esk7_2(X1,X2))|v1_xboole_0(k18_complex1(k6_xcmplx_0(k1_seq_1(X1,esk8_3(X1,X2,X3)),X2)))|v2_comseq_2(X1)|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k1_seq_1(X1,esk8_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk7_2(X1,X2))|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_1(X1,esk8_3(X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1963, plain, (v2_xxreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk8_3(X1,X2,X3)),X2)))|v3_xxreal_0(esk7_2(X1,X2))|v1_xboole_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk8_3(X1,X2,X3)),X2)))|v2_comseq_2(X1)|~v1_xreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk8_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk7_2(X1,X2))|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1964, plain, (v2_xxreal_0(k18_complex1(k9_real_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|v3_xxreal_0(esk7_2(X1,X2))|v1_xboole_0(k18_complex1(k9_real_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|v2_comseq_2(X1)|~v1_xreal_0(k18_complex1(k9_real_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk7_2(X1,X2))|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_5900, plain, (v3_membered(np__0)|~v6_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_3879, plain, (v5_relat_1(k2_zfmisc_1(X1,esk43_0),k3_numbers))).#trainneg
cnf(i_0_4814, plain, (r1_xxreal_0(esk54_0,esk39_0)|r1_xxreal_0(esk39_0,esk54_0))).#trainneg
cnf(i_0_6806, negated_conjecture, (~v2_xxreal_0(esk2_0)|~v3_xxreal_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2118, negated_conjecture, (v3_valued_0(esk4_0)|~v3_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_6895, negated_conjecture, (v3_valued_0(esk4_0)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3789, negated_conjecture, (v2_xxreal_0(esk3_0)|v3_xxreal_0(esk3_0)|v1_xboole_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1981, plain, (v2_xxreal_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|v3_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|v1_xboole_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|v2_struct_0(X1)|~v1_xreal_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|~v1_xreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1987, plain, (v2_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|v3_xxreal_0(k1_normsp_0(X1,k1_funct_1(k1_bhsp_4(X1,X2),X3)))|v1_xboole_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|v2_struct_0(X1)|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|~v1_xreal_0(k1_normsp_0(X1,k1_funct_1(k1_bhsp_4(X1,X2),X3)))|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1978, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|v3_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|v1_xboole_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1979, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|v1_xboole_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1980, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|v1_xboole_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X2)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_1982, negated_conjecture, (v2_xxreal_0(k8_nat_1(X1,esk61_3(esk1_0,esk2_0,esk3_0),X2))|v3_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X2)))|v1_xboole_0(k8_nat_1(X1,esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v1_xreal_0(k8_nat_1(X1,esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X2)))|~v7_ordinal1(X2)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X1)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_1983, negated_conjecture, (v2_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|v1_xboole_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1984, negated_conjecture, (v2_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|v1_xboole_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X2)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_1985, negated_conjecture, (v2_xxreal_0(k8_nat_1(X1,esk61_3(esk1_0,esk2_0,esk3_0),X2))|v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2)))|v1_xboole_0(k8_nat_1(X1,esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v1_xreal_0(k8_nat_1(X1,esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2)))|~v7_ordinal1(X2)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X1)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_1986, negated_conjecture, (v2_xxreal_0(k8_nat_1(X1,esk61_3(esk1_0,esk2_0,esk3_0),X2))|v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X3,esk4_0,X2)))|v1_xboole_0(k8_nat_1(X1,esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v1_xreal_0(k8_nat_1(X1,esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X3,esk4_0,X2)))|~v7_ordinal1(X2)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X1)|~v1_funct_2(esk4_0,k5_numbers,X3)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_1988, plain, (v2_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|v3_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,X2),X3)))|v1_xboole_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|v2_struct_0(X1)|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|~v1_xreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,X2),X3)))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3880, plain, (v5_relat_1(k2_zfmisc_1(X1,esk33_0),k3_numbers))).#trainneg
cnf(i_0_2003, plain, (v2_xxreal_0(k18_complex1(k1_funct_1(X1,esk11_3(X1,np__0,X2))))|v3_xxreal_0(esk10_2(X1,np__0))|v1_xboole_0(esk10_2(X1,np__0))|k1_seq_2(X1)=np__0|~v1_xreal_0(k18_complex1(k1_funct_1(X1,esk11_3(X1,np__0,X2))))|~v1_xreal_0(esk10_2(X1,np__0))|~v2_comseq_2(X1)|~v3_valued_0(X1)|~v1_xcmplx_0(k1_funct_1(X1,esk11_3(X1,np__0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk11_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4831, plain, (r1_xxreal_0(esk54_0,np__0)|r1_xxreal_0(np__0,esk54_0))).#trainneg
cnf(i_0_2131, negated_conjecture, (v4_valued_0(esk4_0)|~v4_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_6967, negated_conjecture, (v4_valued_0(esk4_0)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2004, plain, (v2_xxreal_0(k18_complex1(k1_seq_1(X1,esk11_3(X1,np__0,X2))))|v3_xxreal_0(esk10_2(X1,np__0))|v1_xboole_0(esk10_2(X1,np__0))|k1_seq_2(X1)=np__0|~v1_xreal_0(k18_complex1(k1_seq_1(X1,esk11_3(X1,np__0,X2))))|~v1_xreal_0(esk10_2(X1,np__0))|~v2_comseq_2(X1)|~v1_xcmplx_0(k1_seq_1(X1,esk11_3(X1,np__0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_1(X1,esk11_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2005, plain, (v2_xxreal_0(k18_complex1(k6_xcmplx_0(k1_seq_1(X1,esk11_3(X1,X2,X3)),X2)))|v3_xxreal_0(esk10_2(X1,X2))|v1_xboole_0(esk10_2(X1,X2))|X2=k1_seq_2(X1)|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k1_seq_1(X1,esk11_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk10_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_1(X1,esk11_3(X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2006, plain, (v2_xxreal_0(k18_complex1(k6_xcmplx_0(k1_funct_1(X1,esk11_3(X1,X2,X3)),X2)))|v3_xxreal_0(esk10_2(X1,X2))|v1_xboole_0(esk10_2(X1,X2))|X2=k1_seq_2(X1)|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k1_funct_1(X1,esk11_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk10_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk11_3(X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2007, plain, (v2_xxreal_0(k18_complex1(k9_real_1(k1_funct_1(X1,esk11_3(X1,X2,X3)),X2)))|v3_xxreal_0(esk10_2(X1,X2))|v1_xboole_0(esk10_2(X1,X2))|X2=k1_seq_2(X1)|~v1_xreal_0(k18_complex1(k9_real_1(k1_funct_1(X1,esk11_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk10_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2008, plain, (v2_xxreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk11_3(X1,X2,X3)),X2)))|v3_xxreal_0(esk10_2(X1,X2))|v1_xboole_0(esk10_2(X1,X2))|X2=k1_seq_2(X1)|~v1_xreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk11_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk10_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2010, plain, (v2_xxreal_0(k18_complex1(k1_funct_1(X1,esk8_3(X1,np__0,X2))))|v3_xxreal_0(esk7_2(X1,np__0))|v1_xboole_0(esk7_2(X1,np__0))|v2_comseq_2(X1)|~v1_xreal_0(k18_complex1(k1_funct_1(X1,esk8_3(X1,np__0,X2))))|~v1_xreal_0(esk7_2(X1,np__0))|~v3_valued_0(X1)|~v1_xcmplx_0(k1_funct_1(X1,esk8_3(X1,np__0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk8_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2011, plain, (v2_xxreal_0(k18_complex1(k1_seq_1(X1,esk8_3(X1,np__0,X2))))|v3_xxreal_0(esk7_2(X1,np__0))|v1_xboole_0(esk7_2(X1,np__0))|v2_comseq_2(X1)|~v1_xreal_0(k18_complex1(k1_seq_1(X1,esk8_3(X1,np__0,X2))))|~v1_xreal_0(esk7_2(X1,np__0))|~v1_xcmplx_0(k1_seq_1(X1,esk8_3(X1,np__0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_1(X1,esk8_3(X1,np__0,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2012, plain, (v2_xxreal_0(k18_complex1(k6_xcmplx_0(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|v3_xxreal_0(esk7_2(X1,X2))|v1_xboole_0(esk7_2(X1,X2))|v2_comseq_2(X1)|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk7_2(X1,X2))|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2013, plain, (v2_xxreal_0(k18_complex1(k6_xcmplx_0(k1_seq_1(X1,esk8_3(X1,X2,X3)),X2)))|v3_xxreal_0(esk7_2(X1,X2))|v1_xboole_0(esk7_2(X1,X2))|v2_comseq_2(X1)|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k1_seq_1(X1,esk8_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk7_2(X1,X2))|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_1(X1,esk8_3(X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2014, plain, (v2_xxreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk8_3(X1,X2,X3)),X2)))|v3_xxreal_0(esk7_2(X1,X2))|v1_xboole_0(esk7_2(X1,X2))|v2_comseq_2(X1)|~v1_xreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk8_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk7_2(X1,X2))|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2015, plain, (v2_xxreal_0(k18_complex1(k9_real_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|v3_xxreal_0(esk7_2(X1,X2))|v1_xboole_0(esk7_2(X1,X2))|v2_comseq_2(X1)|~v1_xreal_0(k18_complex1(k9_real_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk7_2(X1,X2))|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_3881, plain, (v5_relat_1(k2_zfmisc_1(X1,k5_numbers),k3_numbers))).#trainneg
cnf(i_0_2042, plain, (v2_xxreal_0(esk11_3(X1,X2,X3))|v3_xxreal_0(X3)|v1_xboole_0(X3)|X2=k1_seq_2(X1)|~v1_xreal_0(esk11_3(X1,X2,X3))|~v1_xreal_0(X3)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4838, plain, (r1_xxreal_0(esk54_0,np__1)|r1_xxreal_0(np__1,esk54_0))).#trainneg
cnf(i_0_2190, plain, (v1_xxreal_0(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7027, plain, (~v2_xxreal_0(esk44_0)|~v3_xxreal_0(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3793, negated_conjecture, (v2_xxreal_0(esk2_0)|v3_xxreal_0(esk2_0)|v1_xboole_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4627, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X1)|v1_xboole_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2032, plain, (v2_xxreal_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|v3_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|v1_xboole_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|v2_struct_0(X1)|~v1_xreal_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|~v1_xreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2038, plain, (v2_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|v3_xxreal_0(k1_normsp_0(X1,k1_funct_1(k1_bhsp_4(X1,X2),X3)))|v1_xboole_0(k1_normsp_0(X1,k1_funct_1(k1_bhsp_4(X1,X2),X3)))|v2_struct_0(X1)|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|~v1_xreal_0(k1_normsp_0(X1,k1_funct_1(k1_bhsp_4(X1,X2),X3)))|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_7014, plain, (X2=k1_seq_2(X1)|v3_xxreal_0(X3)|v1_xboole_0(X3)|~v1_xboole_0(esk11_3(X1,X2,X3))|~v2_comseq_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_7016, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk53_0,esk11_3(X2,X1,X3))|v3_xxreal_0(X3)|v1_xboole_0(X3)|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_2029, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|v3_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|v1_xboole_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v1_xreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2030, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|v1_xboole_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v1_xreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2031, negated_conjecture, (v2_xxreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|v1_xboole_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~v1_xreal_0(k1_seq_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X2)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_2033, negated_conjecture, (v2_xxreal_0(k8_nat_1(X1,esk61_3(esk1_0,esk2_0,esk3_0),X2))|v3_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X2)))|v1_xboole_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X2)))|~v1_xreal_0(k8_nat_1(X1,esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X2)))|~v7_ordinal1(X2)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X1)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_2034, negated_conjecture, (v2_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|v1_xboole_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2035, negated_conjecture, (v2_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|v1_xboole_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X2)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_2036, negated_conjecture, (v2_xxreal_0(k8_nat_1(X1,esk61_3(esk1_0,esk2_0,esk3_0),X2))|v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2)))|v1_xboole_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2)))|~v1_xreal_0(k8_nat_1(X1,esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2)))|~v7_ordinal1(X2)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X1)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_2037, negated_conjecture, (v2_xxreal_0(k8_nat_1(X1,esk61_3(esk1_0,esk2_0,esk3_0),X2))|v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X3,esk4_0,X2)))|v1_xboole_0(k1_normsp_0(esk1_0,k8_nat_1(X3,esk4_0,X2)))|~v1_xreal_0(k8_nat_1(X1,esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X3,esk4_0,X2)))|~v7_ordinal1(X2)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X1)|~v1_funct_2(esk4_0,k5_numbers,X3)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_5932, plain, (v4_membered(np__0)|~v5_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_6077, plain, (v5_relat_1(k2_zfmisc_1(X1,k4_numbers),k3_numbers))).#trainneg
cnf(i_0_2917, negated_conjecture, (m1_subset_1(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2043, plain, (v2_xxreal_0(esk8_3(X1,X2,X3))|v3_xxreal_0(X3)|v1_xboole_0(X3)|v2_comseq_2(X1)|~v1_xreal_0(esk8_3(X1,X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4869, plain, (r1_xxreal_0(esk53_0,esk48_0)|r1_xxreal_0(esk48_0,esk53_0))).#trainneg
cnf(i_0_2192, plain, (v1_xxreal_0(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_5358, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X1)|v1_xboole_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2624, negated_conjecture, (v1_int_1(esk4_0)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2039, plain, (v2_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|v3_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,X2),X3)))|v1_xboole_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,X2),X3)))|v2_struct_0(X1)|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|~v1_xreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,X2),X3)))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2046, plain, (v1_valued_0(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_valued_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_7155, negated_conjecture, (m1_subset_1(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_7158, plain, (r1_xxreal_0(esk53_0,esk48_0)|~v1_xreal_0(esk48_0))).#trainneg
cnf(i_0_7231, negated_conjecture, (m1_subset_1(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2047, plain, (v1_valued_0(esk56_1(X1))|v2_struct_0(X1)|~v1_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_2049, plain, (v1_valued_0(esk42_1(X1))|v7_struct_0(X1)|~v1_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_2050, plain, (v1_valued_0(esk41_1(X1))|v2_struct_0(X1)|~v1_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_2051, plain, (v1_valued_0(esk28_1(k1_zfmisc_1(X1)))|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2053, plain, (v1_valued_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v1_valued_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_6312, plain, (v5_membered(np__0)|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_6670, plain, (v5_relat_1(k2_zfmisc_1(X1,np__0),k4_numbers))).#trainneg
cnf(i_0_4873, plain, (r1_xxreal_0(esk53_0,esk39_0)|r1_xxreal_0(esk39_0,esk53_0))).#trainneg
cnf(i_0_2194, plain, (k6_normsp_1(X1,k1_bhsp_4(X1,X2))=k4_struct_0(X1)|v2_struct_0(X1)|k1_seq_2(k3_series_1(k4_normsp_0(X1,X2)))!=np__0|~v2_comseq_2(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_2(k3_series_1(k4_normsp_0(X1,X2)),k5_numbers,k1_numbers)|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k1_bhsp_4(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(k3_series_1(k4_normsp_0(X1,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2054, plain, (v1_valued_0(k3_series_1(X1))|~v1_valued_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_7168, plain, (~v2_xxreal_0(esk44_0)|~v3_xxreal_0(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_6062, plain, (v3_ordinal1(X1)|~v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1))))))).#trainneg
cnf(i_0_2055, plain, (v1_valued_0(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v1_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2056, plain, (v1_valued_0(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v1_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_7240, plain, (r1_xxreal_0(esk53_0,esk39_0)|~v2_xxreal_0(esk39_0)|~v1_xreal_0(esk39_0))).#trainneg
cnf(i_0_7263, plain, (r1_xxreal_0(esk53_0,esk39_0)|r1_xxreal_0(esk39_0,esk59_0)|~v1_xreal_0(esk39_0))).#trainneg
cnf(i_0_7264, plain, (r1_xxreal_0(esk53_0,esk39_0)|r1_xxreal_0(esk39_0,esk55_0)|~v1_xreal_0(esk39_0))).#trainneg
cnf(i_0_7265, plain, (r1_xxreal_0(esk53_0,esk39_0)|r1_xxreal_0(esk39_0,esk47_0)|~v1_xreal_0(esk39_0))).#trainneg
cnf(i_0_2059, plain, (v2_valued_0(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v2_valued_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2060, plain, (v2_valued_0(esk56_1(X1))|v2_struct_0(X1)|~v2_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_2062, plain, (v2_valued_0(esk42_1(X1))|v7_struct_0(X1)|~v2_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_2063, plain, (v2_valued_0(esk41_1(X1))|v2_struct_0(X1)|~v2_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_2064, plain, (v2_valued_0(esk28_1(k1_zfmisc_1(X1)))|~v2_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2066, plain, (v2_valued_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v2_valued_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2067, plain, (v2_valued_0(k3_series_1(X1))|~v2_valued_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2068, plain, (v2_valued_0(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v2_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_7237, plain, (v5_membered(np__0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_7287, plain, (v5_membered(np__0))).#trainneg
cnf(i_0_7306, plain, (v3_membered(np__0))).#trainneg
cnf(i_0_2069, plain, (v2_valued_0(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v2_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2193, plain, (v1_xxreal_0(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_7315, plain, (~v2_xxreal_0(esk44_0)|~v3_xxreal_0(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_6063, negated_conjecture, (v3_ordinal1(esk3_0)|~v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))))).#trainneg
cnf(i_0_6064, negated_conjecture, (v3_ordinal1(esk2_0)|~v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))))).#trainneg
cnf(i_0_2077, plain, (v1_xreal_0(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2078, plain, (v1_xreal_0(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2079, negated_conjecture, (v1_xreal_0(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2080, negated_conjecture, (v1_xcmplx_0(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2081, negated_conjecture, (v1_xxreal_0(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2082, plain, (v1_xreal_0(k3_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2084, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2091, plain, (r1_xxreal_0(X1,np__0)|v2_struct_0(X2)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_seq_1(esk61_3(X2,X3,X4),X5),esk5_1(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5)))))))|~r1_xxreal_0(esk6_2(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),X1),X5)|~v2_comseq_2(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))))|~v7_ordinal1(X5)|~v1_xreal_0(X1)|~v3_valued_0(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))))|~m2_subset_1(X5,k1_numbers,k5_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),k5_numbers,k1_numbers)|~v1_funct_1(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))))|~m1_subset_1(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_2097, plain, (r1_xxreal_0(X1,np__0)|~r1_xxreal_0(X1,k18_complex1(k6_xcmplx_0(k8_nat_1(X2,X3,X4),esk5_1(X3))))|~r1_xxreal_0(esk6_2(X3,X1),X4)|~v2_comseq_2(X3)|~v7_ordinal1(X4)|~v1_xreal_0(X1)|~v3_valued_0(X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_xreal_0(esk5_1(X3))|~m1_subset_1(k8_nat_1(X2,X3,X4),k1_numbers))).#trainneg
cnf(i_0_6719, plain, (v1_membered(np__0)|~v6_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_6671, plain, (v5_relat_1(k2_zfmisc_1(X1,np__0),k3_numbers))).#trainneg
cnf(i_0_7239, plain, (v3_xxreal_0(esk39_0)|r1_xxreal_0(esk53_0,esk39_0)|~v1_xreal_0(esk39_0))).#trainneg
cnf(i_0_2098, plain, (r1_xxreal_0(k18_complex1(k9_real_1(k8_nat_1(X1,X2,X3),esk5_1(X2))),np__0)|~r1_xxreal_0(esk6_2(X2,k18_complex1(k9_real_1(k8_nat_1(X1,X2,X3),esk5_1(X2)))),X3)|~v2_comseq_2(X2)|~v7_ordinal1(X3)|~v1_xreal_0(k18_complex1(k9_real_1(k8_nat_1(X1,X2,X3),esk5_1(X2))))|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_2220, plain, (v1_xcmplx_0(esk52_0)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_3397, plain, (r1_xxreal_0(k5_numbers,k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3533, negated_conjecture, (r1_xxreal_0(esk3_0,esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2102, plain, (v1_valued_0(k2_zfmisc_1(X1,esk44_0))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_7391, plain, (r1_xxreal_0(esk53_0,esk39_0)|~m1_subset_1(esk39_0,k5_numbers)|~v1_xreal_0(esk39_0))).#trainneg
cnf(i_0_7392, plain, (r1_xxreal_0(esk53_0,esk39_0)|~v1_xboole_0(esk39_0)|~v1_xreal_0(esk39_0))).#trainneg
cnf(i_0_2105, plain, (v3_normsp_1(k1_bhsp_4(X1,X2),X1)|v2_struct_0(X1)|k1_seq_2(k3_series_1(k4_normsp_0(X1,X2)))!=np__0|~v2_comseq_2(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_2(k3_series_1(k4_normsp_0(X1,X2)),k5_numbers,k1_numbers)|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k1_bhsp_4(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(k3_series_1(k4_normsp_0(X1,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2111, plain, (r1_xxreal_0(X1,k4_normsp_0(X2,X3))|v2_xxreal_0(X1)|v3_xxreal_0(k4_normsp_0(X2,X3))|v2_struct_0(X2)|~v1_xreal_0(X1)|~l1_normsp_0(X2)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k3_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))))).#trainneg
cnf(i_0_2113, plain, (v3_valued_0(esk56_1(X1))|v2_struct_0(X1)|~v3_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_2115, plain, (v3_valued_0(esk42_1(X1))|v7_struct_0(X1)|~v3_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_2116, plain, (v3_valued_0(esk41_1(X1))|v2_struct_0(X1)|~v3_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_6774, plain, (v2_membered(np__0)|~v6_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_7418, plain, (v1_valued_0(k2_zfmisc_1(X1,esk44_0)))).#trainneg
cnf(i_0_3783, plain, (m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~r2_hidden(X1,k1_funct_1(np__0,X3))|~v1_funct_2(np__0,k5_numbers,k1_zfmisc_1(X2))|~v1_funct_1(np__0))).#trainneg
cnf(i_0_2117, plain, (v3_valued_0(esk28_1(k1_zfmisc_1(X1)))|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2221, plain, (v1_xxreal_0(esk52_0)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_2222, plain, (v1_xreal_0(esk52_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_3537, negated_conjecture, (r1_xxreal_0(esk2_0,esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3568, plain, (r1_xxreal_0(k5_numbers,k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2121, plain, (v3_valued_0(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v3_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2122, plain, (v3_valued_0(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v3_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2125, plain, (v4_valued_0(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v4_valued_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2126, plain, (v4_valued_0(esk56_1(X1))|v2_struct_0(X1)|~v4_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_2128, plain, (v4_valued_0(esk42_1(X1))|v7_struct_0(X1)|~v4_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_2129, plain, (v4_valued_0(esk41_1(X1))|v2_struct_0(X1)|~v4_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_2130, plain, (v4_valued_0(esk28_1(k1_zfmisc_1(X1)))|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2132, plain, (v4_valued_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v4_valued_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2133, plain, (v4_valued_0(k3_series_1(X1))|~v4_valued_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2134, plain, (v4_valued_0(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v4_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2135, plain, (v4_valued_0(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v4_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2139, negated_conjecture, (v1_xcmplx_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v1_valued_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_relat_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_7153, plain, (v4_membered(np__0)|~v6_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2786, plain, (v2_struct_0(X1)|~v1_xboole_0(k4_normsp_0(X1,X2))|~v1_funct_2(k4_normsp_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k4_normsp_0(X1,X2))|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2140, plain, (v1_xcmplx_0(k8_nat_1(X3,X1,X2))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_4961, plain, (r1_xxreal_0(esk48_0,esk47_0)|r1_xxreal_0(esk47_0,esk48_0))).#trainneg
cnf(i_0_2223, plain, (v1_rat_1(esk52_0)|~v4_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_2224, plain, (v1_int_1(esk52_0)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_3791, negated_conjecture, (r1_xxreal_0(esk3_0,esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3795, negated_conjecture, (r1_xxreal_0(esk2_0,esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2141, plain, (r1_xxreal_0(X1,k1_funct_1(X2,X3))|r1_xxreal_0(k1_funct_1(X2,X3),X1)|~v1_xxreal_0(X1)|~v2_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2142, plain, (v2_xxreal_0(k1_funct_1(X1,X2))|v3_xxreal_0(k1_funct_1(X1,X2))|v1_xboole_0(k1_funct_1(X1,X2))|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2143, plain, (~v2_xxreal_0(k1_funct_1(X1,X2))|~v3_xxreal_0(k1_funct_1(X1,X2))|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7568, negated_conjecture, (~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v2_valued_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_relat_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_7584, negated_conjecture, (~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v3_valued_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_relat_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))))).#trainneg
cnf(i_0_7589, negated_conjecture, (~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v1_xboole_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_relat_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))))).#trainneg
cnf(i_0_7592, negated_conjecture, (~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v1_xboole_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~l2_algstr_0(esk1_0)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2144, plain, (r1_xxreal_0(k1_funct_1(X1,X2),k1_funct_1(X1,X2))|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7598, negated_conjecture, (r1_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1),k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v2_valued_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_relat_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_7599, plain, (r1_xxreal_0(k8_nat_1(X3,X1,X2),k8_nat_1(X3,X1,X2))|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_7602, negated_conjecture, (r1_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1),k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v3_valued_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_relat_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))))).#trainneg
cnf(i_0_7388, plain, (v1_membered(np__0)|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_3415, plain, (v1_xxreal_0(esk28_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_4964, plain, (r1_xxreal_0(esk48_0,esk39_0)|r1_xxreal_0(esk39_0,esk48_0))).#trainneg
cnf(i_0_7585, negated_conjecture, (~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~l2_algstr_0(esk1_0)|~v2_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2145, negated_conjecture, (v1_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v2_valued_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_relat_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_4981, plain, (r1_xxreal_0(esk48_0,np__0)|r1_xxreal_0(np__0,esk48_0))).#trainneg
cnf(i_0_2225, plain, (v7_ordinal1(esk52_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_2226, plain, (v3_ordinal1(esk52_0)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_4629, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5360, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2146, plain, (v1_xxreal_0(k8_nat_1(X3,X1,X2))|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_2147, plain, (r1_xxreal_0(X1,k1_funct_1(X2,X3))|v2_xxreal_0(X1)|v3_xxreal_0(k1_funct_1(X2,X3))|~v1_xreal_0(X1)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_7612, plain, (r1_xxreal_0(k1_funct_1(X2,X3),k1_funct_1(X2,X3))|~v7_ordinal1(X3)|~v2_valued_0(X2)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_7625, negated_conjecture, (r1_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1),k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v1_xboole_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_relat_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))))).#trainneg
cnf(i_0_7839, negated_conjecture, (r1_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1),k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v1_xboole_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~l2_algstr_0(esk1_0)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_7840, negated_conjecture, (r1_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1),k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~l2_algstr_0(esk1_0)|~v1_xboole_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_7845, negated_conjecture, (r1_xxreal_0(k1_funct_1(esk4_0,X1),k1_funct_1(esk4_0,X1))|~l2_algstr_0(esk1_0)|~v1_xboole_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2148, negated_conjecture, (v1_xreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_valued_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_relat_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2149, plain, (v1_xreal_0(k8_nat_1(X3,X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_7628, plain, (v1_membered(np__0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_7873, plain, (v1_membered(np__0))).#trainneg
cnf(i_0_7661, plain, (r1_xxreal_0(np__0,esk48_0)|~v1_xreal_0(esk48_0))).#trainneg
cnf(i_0_7657, negated_conjecture, (~l2_algstr_0(esk1_0)|~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v2_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2150, plain, (v1_xxreal_0(k1_funct_1(X1,X2))|~v4_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4988, plain, (r1_xxreal_0(esk48_0,np__1)|r1_xxreal_0(np__1,esk48_0))).#trainneg
cnf(i_0_2354, plain, (v1_xreal_0(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_5972, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_5975, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2151, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v4_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2153, negated_conjecture, (v7_ordinal1(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v4_valued_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_relat_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2152, plain, (~v3_xxreal_0(k1_funct_1(X1,X2))|~v4_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7977, negated_conjecture, (~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v4_valued_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_relat_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_7979, negated_conjecture, (~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v1_xboole_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_relat_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))))).#trainneg
cnf(i_0_7980, negated_conjecture, (~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v5_relat_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers)|~v1_relat_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))))).#trainneg
cnf(i_0_7982, negated_conjecture, (~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v1_xboole_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~l2_algstr_0(esk1_0)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_7983, negated_conjecture, (~l2_algstr_0(esk1_0)|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v1_xboole_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2154, plain, (v7_ordinal1(k8_nat_1(X3,X1,X2))|~v4_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_2155, plain, (v1_xreal_0(k6_xcmplx_0(X2,X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2157, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),np__0)=k6_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2158, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),np__1)=k6_xcmplx_0(k4_xcmplx_0(np__1),X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2159, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),X2)=k6_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2160, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(X1))=k6_xcmplx_0(X1,np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7928, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X4)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_7892, plain, (v1_valued_0(k2_zfmisc_1(X1,np__0)))).#trainneg
cnf(i_0_7897, negated_conjecture, (~l2_algstr_0(esk1_0)|~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v2_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2161, plain, (k6_xcmplx_0(np__1,k4_xcmplx_0(X1))=k6_xcmplx_0(X1,k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_2668, plain, (v1_xcmplx_0(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_5306, negated_conjecture, (r1_xxreal_0(esk3_0,esk2_0)|r1_xxreal_0(esk2_0,esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_2162, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2169, plain, (v3_xxreal_0(k6_xcmplx_0(X2,X1))|v3_xxreal_0(k4_xcmplx_0(X2))|~v3_xxreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2163, plain, (v3_xxreal_0(X2)|~m1_subset_1(k6_xcmplx_0(X1,X2),k5_numbers)|~v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2164, plain, (v3_xxreal_0(X2)|~v1_xboole_0(k6_xcmplx_0(X1,X2))|~v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2173, plain, (v3_xxreal_0(X1)|~v1_xboole_0(k6_xcmplx_0(X1,X2))|~v3_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2178, plain, (v2_xxreal_0(k6_xcmplx_0(X2,X1))|v3_xxreal_0(k4_xcmplx_0(X1))|~v3_xxreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2189, plain, (v5_relat_1(k2_zfmisc_1(X1,esk44_0),k3_numbers)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_2195, negated_conjecture, (r1_xxreal_0(X1,np__0)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_funct_1(esk4_0,X2),esk5_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))))))|~r1_xxreal_0(esk6_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),X1),X2)|~v2_comseq_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)))|~v7_ordinal1(X2)|~v1_xreal_0(X1)|~v3_valued_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2196, negated_conjecture, (r1_xxreal_0(X1,np__0)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k8_nat_1(X3,esk4_0,X2),esk5_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))))))|~r1_xxreal_0(esk6_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),X1),X2)|~v2_comseq_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)))|~v7_ordinal1(X2)|~v1_xreal_0(X1)|~v3_valued_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(esk4_0,k5_numbers,X3)|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_2198, negated_conjecture, (r1_xxreal_0(X1,np__0)|~r1_xxreal_0(X1,k18_complex1(k6_xcmplx_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2),esk5_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))))))|~r1_xxreal_0(esk6_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),X1),X2)|~v2_comseq_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)))|~v7_ordinal1(X2)|~v1_xreal_0(X1)|~v3_valued_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(esk5_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2),k1_numbers))).#trainneg
cnf(i_0_8079, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,np__0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_8368, plain, (k2_zfmisc_1(X1,np__0)=k2_zfmisc_1(X2,np__0))).#trainneg
cnf(i_0_2207, plain, (v2_xxreal_0(k4_xcmplx_0(X1))|v3_xxreal_0(k6_xcmplx_0(X2,X1))|~v2_xxreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2199, negated_conjecture, (r1_xxreal_0(k18_complex1(k9_real_1(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1),esk5_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))))),np__0)|~r1_xxreal_0(esk6_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k18_complex1(k9_real_1(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1),esk5_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))))),X1)|~v2_comseq_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v7_ordinal1(X1)|~v1_xreal_0(k18_complex1(k9_real_1(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1),esk5_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))))))|~v3_valued_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2201, plain, (v2_xxreal_0(X1)|~m1_subset_1(k6_xcmplx_0(X1,X2),k5_numbers)|~v2_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2202, plain, (v2_xxreal_0(X1)|~v1_xboole_0(k6_xcmplx_0(X1,X2))|~v2_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2211, plain, (v2_xxreal_0(X2)|~v1_xboole_0(k6_xcmplx_0(X1,X2))|~v2_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2216, plain, (v2_xxreal_0(k6_xcmplx_0(X2,X1))|v2_xxreal_0(k4_xcmplx_0(X2))|~v2_xxreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2238, plain, (v1_xboole_0(k2_zfmisc_1(k5_numbers,k5_numbers))|m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k5_numbers))|~m1_subset_1(X1,esk52_0))).#trainneg
cnf(i_0_2239, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k5_numbers))|~r2_hidden(X1,esk52_0))).#trainneg
cnf(i_0_2241, plain, (v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2242, plain, (v1_xxreal_0(k1_seq_1(X1,X2))|~v2_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2243, plain, (v1_xreal_0(k1_seq_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2244, plain, (v1_rat_1(k1_seq_1(X1,X2))|~v4_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8415, plain, (v1_funct_2(X1,X2,np__0)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,np__0))))).#trainneg
cnf(i_0_2245, plain, (v1_int_1(k1_seq_1(X1,X2))|~v5_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5763, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_5764, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2246, plain, (v7_ordinal1(k1_seq_1(X1,X2))|~v6_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2247, plain, (v3_ordinal1(k1_seq_1(X1,X2))|~v3_ordinal1(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8658, plain, (r1_xxreal_0(X1,k4_normsp_0(X2,X3))|r1_xxreal_0(k4_normsp_0(X2,X3),X1)|v2_struct_0(X2)|~v1_xreal_0(X1)|~l1_normsp_0(X2)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))))).#trainneg
cnf(i_0_8676, plain, (r1_xxreal_0(X1,k1_bhsp_4(X2,X3))|r1_xxreal_0(k1_bhsp_4(X2,X3),X1)|v2_struct_0(X2)|~v1_xreal_0(X1)|~l2_algstr_0(X2)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))))).#trainneg
cnf(i_0_8677, plain, (r1_xxreal_0(X1,k1_bhsp_4(X2,X3))|r1_xxreal_0(k1_bhsp_4(X2,X3),X1)|v2_struct_0(X2)|~v1_xreal_0(X1)|~l2_algstr_0(X2)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))),k1_zfmisc_1(k4_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))))).#trainneg
cnf(i_0_8678, plain, (r1_xxreal_0(X1,k1_bhsp_4(X2,X3))|r1_xxreal_0(k1_bhsp_4(X2,X3),X1)|v2_struct_0(X2)|~v1_xreal_0(X1)|~l2_algstr_0(X2)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))),k1_zfmisc_1(k3_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))))).#trainneg
cnf(i_0_8679, plain, (r1_xxreal_0(X1,k1_bhsp_4(X2,X3))|r1_xxreal_0(k1_bhsp_4(X2,X3),X1)|v2_struct_0(X2)|~v1_xreal_0(X1)|~l2_algstr_0(X2)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))),k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))))).#trainneg
cnf(i_0_2248, plain, (m1_subset_1(k1_funct_1(X1,X2),k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2249, plain, (v2_struct_0(X2)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~r2_hidden(X1,esk61_3(X2,X3,X4))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_2251, plain, (~v1_xboole_0(k2_zfmisc_1(k5_numbers,k5_numbers))|~r2_hidden(X1,esk52_0))).#trainneg
cnf(i_0_2253, plain, (v7_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~r2_hidden(X2,esk42_1(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_2263, negated_conjecture, (v1_rat_1(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v5_relat_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k3_numbers)|~v1_relat_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_5820, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2264, plain, (v1_rat_1(k8_nat_1(X3,X1,X2))|~v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_2265, negated_conjecture, (v1_int_1(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v5_relat_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k4_numbers)|~v1_relat_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2255, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,esk28_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_2256, negated_conjecture, (~v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~r2_hidden(X1,esk4_0))).#trainneg
cnf(i_0_2257, plain, (v2_struct_0(X2)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~r2_hidden(X1,k4_normsp_0(X2,X3))|~l1_normsp_0(X2)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))))).#trainneg
cnf(i_0_2258, plain, (~v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~r2_hidden(X1,k3_series_1(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2259, plain, (v2_struct_0(X1)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~r2_hidden(X2,k1_bhsp_4(X1,X3))|~l2_algstr_0(X1)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2260, plain, (v2_struct_0(X1)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~r2_hidden(X2,k8_lopban_4(X1,X3,X4,X5))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X5,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2266, plain, (v1_int_1(k8_nat_1(X3,X1,X2))|~v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_2267, plain, (v1_xcmplx_0(k9_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2268, plain, (v1_xxreal_0(k9_real_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2269, plain, (v1_xreal_0(k9_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2270, plain, (v1_rat_1(k9_real_1(X1,X2))|~v4_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5035, plain, (r1_xxreal_0(esk47_0,esk39_0)|r1_xxreal_0(esk39_0,esk47_0))).#trainneg
cnf(i_0_2271, plain, (v1_int_1(k9_real_1(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5059, plain, (r1_xxreal_0(esk47_0,np__1)|r1_xxreal_0(np__1,esk47_0))).#trainneg
cnf(i_0_2625, negated_conjecture, (v1_int_1(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k5_numbers))).#trainneg
cnf(i_0_2634, negated_conjecture, (v3_ordinal1(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k5_numbers))).#trainneg
cnf(i_0_2272, plain, (v7_ordinal1(k9_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2273, plain, (v3_ordinal1(k9_real_1(X1,X2))|~v3_ordinal1(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2274, plain, (m1_subset_1(k6_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2279, plain, (v2_xxreal_0(k4_xcmplx_0(X1))|v3_xxreal_0(k4_xcmplx_0(X2))|~v3_xxreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2289, plain, (v2_xxreal_0(k4_xcmplx_0(X1))|v3_xxreal_0(k4_xcmplx_0(X2))|~v2_xxreal_0(k6_xcmplx_0(X2,X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2293, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X2)|v3_xxreal_0(k6_xcmplx_0(X1,X2))|v1_xboole_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(k6_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_2300, negated_conjecture, (v1_xreal_0(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k5_numbers))).#trainneg
cnf(i_0_2301, negated_conjecture, (v1_xreal_0(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2303, plain, (r1_xxreal_0(X1,np__0)|v2_struct_0(X3)|X2!=k1_seq_2(k3_series_1(k4_normsp_0(X3,k8_lopban_4(X3,X4,X5,X6))))|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_seq_1(esk61_3(X3,X4,X5),X6),X2)))|~r1_xxreal_0(esk9_3(k3_series_1(k4_normsp_0(X3,k8_lopban_4(X3,X4,X5,X6))),X2,X1),X6)|~v2_comseq_2(k3_series_1(k4_normsp_0(X3,k8_lopban_4(X3,X4,X5,X6))))|~v7_ordinal1(X6)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v3_valued_0(k3_series_1(k4_normsp_0(X3,k8_lopban_4(X3,X4,X5,X6))))|~m2_subset_1(X6,k1_numbers,k5_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(X3,k8_lopban_4(X3,X4,X5,X6))),k5_numbers,k1_numbers)|~v1_funct_1(k3_series_1(k4_normsp_0(X3,k8_lopban_4(X3,X4,X5,X6))))|~m1_subset_1(k3_series_1(k4_normsp_0(X3,k8_lopban_4(X3,X4,X5,X6))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X5,u1_struct_0(X3))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_lopban_2(X3)|~v5_lopban_2(X3)|~v3_vectsp_1(X3)|~v1_vectsp_1(X3)|~v3_group_1(X3)|~v2_funcsdom(X3)|~v2_normsp_1(X3)|~v4_normsp_0(X3)|~v3_normsp_0(X3)|~v8_rlvect_1(X3)|~v7_rlvect_1(X3)|~v6_rlvect_1(X3)|~v5_rlvect_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v13_algstr_0(X3))).#trainneg
cnf(i_0_8846, plain, (m1_subset_1(k4_xcmplx_0(np__1),k1_numbers))).#trainneg
cnf(i_0_8098, negated_conjecture, (~l2_algstr_0(esk1_0)|~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v2_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_3576, negated_conjecture, (v1_xxreal_0(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k5_numbers))).#trainneg
cnf(i_0_3797, negated_conjecture, (v1_xcmplx_0(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k5_numbers))).#trainneg
cnf(i_0_2309, plain, (r1_xxreal_0(X1,np__0)|X2!=k1_seq_2(X3)|~r1_xxreal_0(X1,k18_complex1(k6_xcmplx_0(k8_nat_1(X4,X3,X5),X2)))|~r1_xxreal_0(esk9_3(X3,X2,X1),X5)|~v2_comseq_2(X3)|~v7_ordinal1(X5)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v3_valued_0(X3)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,X4)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(k8_nat_1(X4,X3,X5),k1_numbers))).#trainneg
cnf(i_0_2310, plain, (r1_xxreal_0(k18_complex1(k9_real_1(k8_nat_1(X1,X2,X3),X4)),np__0)|X4!=k1_seq_2(X2)|~r1_xxreal_0(esk9_3(X2,X4,k18_complex1(k9_real_1(k8_nat_1(X1,X2,X3),X4))),X3)|~v2_comseq_2(X2)|~v7_ordinal1(X3)|~v1_xreal_0(k18_complex1(k9_real_1(k8_nat_1(X1,X2,X3),X4)))|~v1_xreal_0(X4)|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_2313, plain, (v5_relat_1(esk56_1(X1),k4_numbers)|v2_struct_0(X1)|~v5_relat_1(u1_struct_0(X1),k4_numbers)|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_2316, plain, (v5_relat_1(esk42_1(X1),k4_numbers)|v7_struct_0(X1)|~v5_relat_1(u1_struct_0(X1),k4_numbers)|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_2317, plain, (v5_relat_1(esk41_1(X1),k4_numbers)|v2_struct_0(X1)|~v5_relat_1(u1_struct_0(X1),k4_numbers)|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_2318, plain, (v5_relat_1(esk28_1(k1_zfmisc_1(X1)),k4_numbers)|~v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2319, negated_conjecture, (v5_relat_1(esk4_0,k4_numbers)|~v5_relat_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k4_numbers))).#trainneg
cnf(i_0_2320, plain, (v5_relat_1(k4_normsp_0(X1,X2),k4_numbers)|v2_struct_0(X1)|~v5_relat_1(k2_zfmisc_1(k5_numbers,k1_numbers),k4_numbers)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2321, plain, (v5_relat_1(k3_series_1(X1),k4_numbers)|~v5_relat_1(k2_zfmisc_1(k5_numbers,k1_numbers),k4_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2322, plain, (v5_relat_1(k1_bhsp_4(X1,X2),k4_numbers)|v2_struct_0(X1)|~v5_relat_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),k4_numbers)|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2323, plain, (v5_relat_1(k8_lopban_4(X1,X2,X3,X4),k4_numbers)|v2_struct_0(X1)|~v5_relat_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),k4_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2327, plain, (v5_relat_1(esk56_1(X1),k3_numbers)|v2_struct_0(X1)|~v5_relat_1(u1_struct_0(X1),k3_numbers)|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_8946, plain, (v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_2330, plain, (v5_relat_1(esk42_1(X1),k3_numbers)|v7_struct_0(X1)|~v5_relat_1(u1_struct_0(X1),k3_numbers)|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_2671, plain, (v1_xxreal_0(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_5917, negated_conjecture, (v6_membered(esk4_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_5958, negated_conjecture, (v3_membered(esk4_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2331, plain, (v5_relat_1(esk41_1(X1),k3_numbers)|v2_struct_0(X1)|~v5_relat_1(u1_struct_0(X1),k3_numbers)|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_2332, plain, (v5_relat_1(esk28_1(k1_zfmisc_1(X1)),k3_numbers)|~v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2333, negated_conjecture, (v5_relat_1(esk4_0,k3_numbers)|~v5_relat_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k3_numbers))).#trainneg
cnf(i_0_2334, plain, (v5_relat_1(k4_normsp_0(X1,X2),k3_numbers)|v2_struct_0(X1)|~v5_relat_1(k2_zfmisc_1(k5_numbers,k1_numbers),k3_numbers)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2335, plain, (v5_relat_1(k3_series_1(X1),k3_numbers)|~v5_relat_1(k2_zfmisc_1(k5_numbers,k1_numbers),k3_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2336, plain, (v5_relat_1(k1_bhsp_4(X1,X2),k3_numbers)|v2_struct_0(X1)|~v5_relat_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),k3_numbers)|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2337, plain, (v5_relat_1(k8_lopban_4(X1,X2,X3,X4),k3_numbers)|v2_struct_0(X1)|~v5_relat_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),k3_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2345, negated_conjecture, (r1_xxreal_0(X1,np__0)|X2!=k1_seq_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)))|~r1_xxreal_0(X1,k18_complex1(k6_xcmplx_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X3),X2)))|~r1_xxreal_0(esk9_3(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)),X2,X1),X3)|~v2_comseq_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v3_valued_0(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)),k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X3),k1_numbers)|~v7_ordinal1(X3)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)),k5_numbers,u1_struct_0(esk1_0))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X3)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2347, plain, (r1_xxreal_0(X1,np__0)|k1_seq_2(X2)!=np__0|~r1_xxreal_0(X1,k18_complex1(k1_funct_1(X2,X3)))|~r1_xxreal_0(esk9_3(X2,np__0,X1),X3)|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_funct_1(X2,X3),k1_numbers)|~v1_xcmplx_0(k1_funct_1(X2,X3)))).#trainneg
cnf(i_0_2350, plain, (r1_xxreal_0(k18_complex1(k6_xcmplx_0(k1_funct_1(X1,X2),X3)),np__0)|X3!=k1_seq_2(X1)|~r1_xxreal_0(esk9_3(X1,X3,k18_complex1(k6_xcmplx_0(k1_funct_1(X1,X2),X3))),X2)|~v2_comseq_2(X1)|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k1_funct_1(X1,X2),X3)))|~v1_xreal_0(X3)|~v3_valued_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_funct_1(X1,X2),k1_numbers))).#trainneg
cnf(i_0_2355, plain, (v3_xxreal_0(esk7_2(X1,X2))|v2_comseq_2(X1)|~v3_xxreal_0(k18_complex1(k9_real_1(k8_nat_1(X3,X1,esk8_3(X1,X2,X4)),X2)))|~v1_xreal_0(k18_complex1(k9_real_1(k8_nat_1(X3,X1,esk8_3(X1,X2,X4)),X2)))|~v1_xreal_0(esk7_2(X1,X2))|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_2356, plain, (v2_xxreal_0(k18_complex1(k9_real_1(k8_nat_1(X1,X2,esk8_3(X2,X3,X4)),X3)))|v2_comseq_2(X2)|~v2_xxreal_0(esk7_2(X2,X3))|~v1_xreal_0(k18_complex1(k9_real_1(k8_nat_1(X1,X2,esk8_3(X2,X3,X4)),X3)))|~v1_xreal_0(esk7_2(X2,X3))|~v1_xreal_0(X3)|~v3_valued_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_8418, plain, (v1_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,np__0))))).#trainneg
cnf(i_0_8389, plain, (v3_valued_0(k2_zfmisc_1(X2,np__0)))).#trainneg
cnf(i_0_8947, plain, (v1_xxreal_0(k4_xcmplx_0(np__1))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_8954, negated_conjecture, (~l2_algstr_0(esk1_0)|~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v2_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2366, plain, (v1_xcmplx_0(k1_normsp_0(X1,X2))|v2_struct_0(X1)|~v1_membered(k1_numbers)|~l1_normsp_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_6403, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1))),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2367, plain, (v1_xxreal_0(k1_normsp_0(X1,X2))|v2_struct_0(X1)|~v2_membered(k1_numbers)|~l1_normsp_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2369, plain, (v1_rat_1(k1_normsp_0(X1,X2))|v2_struct_0(X1)|~v4_membered(k1_numbers)|~l1_normsp_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2357, plain, (v2_xxreal_0(k18_complex1(k9_real_1(k8_nat_1(X1,X2,esk8_3(X2,X3,X4)),X3)))|v3_xxreal_0(esk7_2(X2,X3))|v1_xboole_0(k18_complex1(k9_real_1(k8_nat_1(X1,X2,esk8_3(X2,X3,X4)),X3)))|v2_comseq_2(X2)|~v1_xreal_0(k18_complex1(k9_real_1(k8_nat_1(X1,X2,esk8_3(X2,X3,X4)),X3)))|~v1_xreal_0(esk7_2(X2,X3))|~v1_xreal_0(X3)|~v3_valued_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_2358, plain, (v2_xxreal_0(k18_complex1(k9_real_1(k8_nat_1(X1,X2,esk8_3(X2,X3,X4)),X3)))|v3_xxreal_0(esk7_2(X2,X3))|v1_xboole_0(esk7_2(X2,X3))|v2_comseq_2(X2)|~v1_xreal_0(k18_complex1(k9_real_1(k8_nat_1(X1,X2,esk8_3(X2,X3,X4)),X3)))|~v1_xreal_0(esk7_2(X2,X3))|~v1_xreal_0(X3)|~v3_valued_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_2370, plain, (v1_int_1(k1_normsp_0(X1,X2))|v2_struct_0(X1)|~v5_membered(k1_numbers)|~l1_normsp_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2371, plain, (v7_ordinal1(k1_normsp_0(X1,X2))|v2_struct_0(X1)|~v6_membered(k1_numbers)|~l1_normsp_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2372, plain, (v3_ordinal1(k1_normsp_0(X1,X2))|v2_struct_0(X1)|~v3_ordinal1(k1_numbers)|~l1_normsp_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2373, plain, (r1_xxreal_0(X1,k1_normsp_0(X2,X3))|v2_xxreal_0(X1)|v3_xxreal_0(k1_normsp_0(X2,X3))|v2_struct_0(X2)|~v1_xreal_0(X1)|~l1_normsp_0(X2)|~m1_subset_1(X3,u1_struct_0(X2)))).#trainneg
cnf(i_0_2376, plain, (v1_xcmplx_0(esk32_2(X1,X2))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2377, plain, (v1_xxreal_0(esk32_2(X1,X2))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_8420, plain, (v3_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,np__0))))).#trainneg
cnf(i_0_8390, plain, (v4_valued_0(k2_zfmisc_1(X2,np__0)))).#trainneg
cnf(i_0_8949, plain, (v1_rat_1(k4_xcmplx_0(np__1))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_7981, negated_conjecture, (~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~l2_algstr_0(esk1_0)|~v4_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2378, plain, (v1_xreal_0(esk32_2(X1,X2))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2967, negated_conjecture, (v3_valued_0(esk4_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_6404, negated_conjecture, (v3_ordinal1(esk3_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_6405, negated_conjecture, (v3_ordinal1(esk2_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2379, plain, (v1_rat_1(esk32_2(X1,X2))|~v4_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2380, plain, (v1_int_1(esk32_2(X1,X2))|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2381, plain, (v7_ordinal1(esk32_2(X1,X2))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2382, plain, (v3_ordinal1(esk32_2(X1,X2))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2383, plain, (v5_membered(esk32_2(X1,X2))|~v5_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2384, plain, (v6_ordinal1(esk32_2(X1,X2))|~v6_ordinal1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2385, plain, (v6_membered(esk32_2(X1,X2))|~v6_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2386, plain, (v1_membered(esk32_2(X1,X2))|~v1_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2387, plain, (v2_membered(esk32_2(X1,X2))|~v2_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2388, plain, (v3_membered(esk32_2(X1,X2))|~v3_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2389, plain, (v4_membered(esk32_2(X1,X2))|~v4_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2390, plain, (v1_valued_0(esk32_2(X1,X2))|~v1_valued_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2391, plain, (v2_valued_0(esk32_2(X1,X2))|~v2_valued_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2392, plain, (v3_valued_0(esk32_2(X1,X2))|~v3_valued_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_8421, plain, (v4_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,np__0))))).#trainneg
cnf(i_0_8950, plain, (v1_int_1(k4_xcmplx_0(np__1))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_9208, negated_conjecture, (~l2_algstr_0(esk1_0)|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v4_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2393, plain, (v4_valued_0(esk32_2(X1,X2))|~v4_valued_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2968, negated_conjecture, (v3_valued_0(esk4_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_3028, negated_conjecture, (v1_valued_0(esk4_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_6606, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1)))),k5_numbers))).#trainneg
cnf(i_0_6607, negated_conjecture, (v3_ordinal1(esk3_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))),k5_numbers))).#trainneg
cnf(i_0_2394, plain, (v5_relat_1(esk32_2(X1,X2),k4_numbers)|~v5_relat_1(k2_zfmisc_1(X1,X2),k4_numbers))).#trainneg
cnf(i_0_2395, plain, (v5_relat_1(esk32_2(X1,X2),k3_numbers)|~v5_relat_1(k2_zfmisc_1(X1,X2),k3_numbers))).#trainneg
cnf(i_0_2396, plain, (v1_xboole_0(esk32_2(X1,X2))|v1_xboole_0(k2_zfmisc_1(X1,X2))|m1_subset_1(X3,k2_zfmisc_1(X1,X2))|~m1_subset_1(X3,esk32_2(X1,X2)))).#trainneg
cnf(i_0_2397, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X1,esk32_2(X2,X3)))).#trainneg
cnf(i_0_2398, plain, (~v1_xboole_0(k2_zfmisc_1(X1,X2))|~r2_hidden(X3,esk32_2(X1,X2)))).#trainneg
cnf(i_0_2399, plain, (v3_normsp_1(k1_bhsp_4(X1,esk32_2(k5_numbers,u1_struct_0(X1))),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,esk32_2(k5_numbers,u1_struct_0(X1))),X2),k5_numbers)|v2_struct_0(X1)|k2_seq_2(X2)!=np__0|~v2_comseq_2(X2)|~v1_funct_2(k1_bhsp_4(X1,esk32_2(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(X1,esk32_2(k5_numbers,u1_struct_0(X1))))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2400, plain, (v3_normsp_1(k1_bhsp_4(X1,esk32_2(k5_numbers,u1_struct_0(X1))),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,esk32_2(k5_numbers,u1_struct_0(X1))),X2),k1_numbers)|v2_struct_0(X1)|k2_seq_2(X2)!=np__0|~v2_comseq_2(X2)|~v1_funct_2(k1_bhsp_4(X1,esk32_2(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(X1,esk32_2(k5_numbers,u1_struct_0(X1))))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2402, plain, (r1_xxreal_0(X1,k1_seq_2(esk32_2(k5_numbers,k1_numbers)))|v2_xxreal_0(X1)|v3_xxreal_0(k1_seq_2(esk32_2(k5_numbers,k1_numbers)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2403, plain, (v3_normsp_1(esk32_2(k5_numbers,u1_struct_0(X1)),X1)|m1_subset_1(esk62_3(X1,esk32_2(k5_numbers,u1_struct_0(X1)),X2),k5_numbers)|v2_struct_0(X1)|k2_seq_2(X2)!=np__0|~v2_comseq_2(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2404, plain, (v3_normsp_1(esk32_2(k5_numbers,u1_struct_0(X1)),X1)|m1_subset_1(esk62_3(X1,esk32_2(k5_numbers,u1_struct_0(X1)),X2),k1_numbers)|v2_struct_0(X1)|k2_seq_2(X2)!=np__0|~v2_comseq_2(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2434, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(esk29_2(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2435, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(esk29_2(X1,X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2436, plain, (r1_xxreal_0(X1,esk44_0)|r1_xxreal_0(esk44_0,X1)|~v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_8422, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,np__0))))).#trainneg
cnf(i_0_8401, plain, (v1_xboole_0(k2_zfmisc_1(X2,np__0)))).#trainneg
cnf(i_0_8951, plain, (v7_ordinal1(k4_xcmplx_0(np__1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_9234, negated_conjecture, (~l2_algstr_0(esk1_0)|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v4_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_9611, plain, (~v6_membered(k1_numbers))).#trainneg
cnf(i_0_6608, negated_conjecture, (v3_ordinal1(esk2_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))),k5_numbers))).#trainneg
cnf(i_0_7258, plain, (v3_ordinal1(X1)|~v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1)))))))).#trainneg
cnf(i_0_2437, plain, (v2_xxreal_0(esk44_0)|v3_xxreal_0(esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_2440, plain, (v3_xxreal_0(esk10_2(X1,X2))|X2=k1_seq_2(X1)|~v3_xxreal_0(k18_complex1(k9_real_1(k8_nat_1(X3,X1,esk11_3(X1,X2,X4)),X2)))|~v1_xreal_0(k18_complex1(k9_real_1(k8_nat_1(X3,X1,esk11_3(X1,X2,X4)),X2)))|~v1_xreal_0(esk10_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_2441, plain, (v2_xxreal_0(k18_complex1(k9_real_1(k8_nat_1(X1,X2,esk11_3(X2,X3,X4)),X3)))|X3=k1_seq_2(X2)|~v2_xxreal_0(esk10_2(X2,X3))|~v1_xreal_0(k18_complex1(k9_real_1(k8_nat_1(X1,X2,esk11_3(X2,X3,X4)),X3)))|~v1_xreal_0(esk10_2(X2,X3))|~v2_comseq_2(X2)|~v1_xreal_0(X3)|~v3_valued_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_2442, plain, (v2_xxreal_0(k18_complex1(k9_real_1(k8_nat_1(X1,X2,esk11_3(X2,X3,X4)),X3)))|v3_xxreal_0(esk10_2(X2,X3))|v1_xboole_0(k18_complex1(k9_real_1(k8_nat_1(X1,X2,esk11_3(X2,X3,X4)),X3)))|X3=k1_seq_2(X2)|~v1_xreal_0(k18_complex1(k9_real_1(k8_nat_1(X1,X2,esk11_3(X2,X3,X4)),X3)))|~v1_xreal_0(esk10_2(X2,X3))|~v2_comseq_2(X2)|~v1_xreal_0(X3)|~v3_valued_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_2443, plain, (v2_xxreal_0(k18_complex1(k9_real_1(k8_nat_1(X1,X2,esk11_3(X2,X3,X4)),X3)))|v3_xxreal_0(esk10_2(X2,X3))|v1_xboole_0(esk10_2(X2,X3))|X3=k1_seq_2(X2)|~v1_xreal_0(k18_complex1(k9_real_1(k8_nat_1(X1,X2,esk11_3(X2,X3,X4)),X3)))|~v1_xreal_0(esk10_2(X2,X3))|~v2_comseq_2(X2)|~v1_xreal_0(X3)|~v3_valued_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_2456, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X2)|v1_xboole_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(k6_xcmplx_0(X1,X2),k5_numbers))).#trainneg
cnf(i_0_2458, plain, (v3_xxreal_0(X1)|v1_xboole_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers))).#trainneg
cnf(i_0_8408, plain, (v1_xreal_0(k2_zfmisc_1(X2,np__0)))).#trainneg
cnf(i_0_8952, plain, (v3_ordinal1(k4_xcmplx_0(np__1))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_9614, plain, (~m1_subset_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5178, plain, (r1_xxreal_0(esk39_0,np__0)|r1_xxreal_0(np__0,esk39_0))).#trainneg
cnf(i_0_9615, plain, (~m1_subset_1(k1_numbers,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_7929, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X4)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2478, plain, (v3_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|v2_struct_0(X1)|~v3_xxreal_0(k1_funct_1(esk61_3(X1,X2,X3),X4))|~v1_xreal_0(k1_funct_1(esk61_3(X1,X2,X3),X4))|~v1_xreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|~v3_valued_0(esk61_3(X1,X2,X3))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2479, plain, (v2_xxreal_0(k1_funct_1(esk61_3(X1,X2,X3),X4))|v2_struct_0(X1)|~v2_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|~v1_xreal_0(k1_funct_1(esk61_3(X1,X2,X3),X4))|~v1_xreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|~v3_valued_0(esk61_3(X1,X2,X3))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2480, plain, (v2_xxreal_0(k1_funct_1(esk61_3(X1,X2,X3),X4))|v3_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|v1_xboole_0(k1_funct_1(esk61_3(X1,X2,X3),X4))|v2_struct_0(X1)|~v1_xreal_0(k1_funct_1(esk61_3(X1,X2,X3),X4))|~v1_xreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|~v3_valued_0(esk61_3(X1,X2,X3))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2481, plain, (v2_xxreal_0(k1_funct_1(esk61_3(X1,X2,X3),X4))|v3_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|v1_xboole_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|v2_struct_0(X1)|~v1_xreal_0(k1_funct_1(esk61_3(X1,X2,X3),X4))|~v1_xreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|~v3_valued_0(esk61_3(X1,X2,X3))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_9707, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,np__0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_9780, plain, (k2_zfmisc_1(np__0,X1)=k2_zfmisc_1(X2,np__0))).#trainneg
cnf(i_0_2535, plain, (r1_xxreal_0(X1,esk44_0)|v2_xxreal_0(X1)|v3_xxreal_0(esk44_0)|~v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_5185, plain, (r1_xxreal_0(esk39_0,np__1)|r1_xxreal_0(np__1,esk39_0))).#trainneg
cnf(i_0_7948, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2483, plain, (r1_xxreal_0(k1_normsp_0(X1,k1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)),k1_funct_1(esk61_3(X1,X2,X3),X4))|v2_struct_0(X1)|~v3_valued_0(esk61_3(X1,X2,X3))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v1_funct_2(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)))|~m1_subset_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2484, plain, (r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(X5,k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)),k1_funct_1(esk61_3(X1,X2,X3),X4))|v2_struct_0(X1)|~v3_valued_0(esk61_3(X1,X2,X3))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v1_funct_2(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k5_numbers,X5)|~v1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)))|~m1_subset_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5))))).#trainneg
cnf(i_0_9927, plain, (r1_xxreal_0(esk53_0,esk44_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_9933, plain, (r1_xxreal_0(esk28_1(esk44_0),esk44_0)|~v3_xxreal_0(esk28_1(esk44_0))|~v2_membered(k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_10003, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_10026, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X2)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_2552, negated_conjecture, (v1_rat_1(esk4_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_10052, negated_conjecture, (v1_rat_1(esk4_0)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_10053, negated_conjecture, (v1_rat_1(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k5_numbers))).#trainneg
cnf(i_0_2488, plain, (v3_xxreal_0(k1_normsp_0(X1,k1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|v2_struct_0(X1)|~v3_xxreal_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|~v1_xreal_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|~v1_xreal_0(k1_normsp_0(X1,k1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|~v7_ordinal1(X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_10029, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),k8_nat_1(X2,esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X2)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_10030, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X3,esk4_0,X1)),k8_nat_1(X2,esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X2)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk4_0,k5_numbers,X3)|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_10045, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_10046, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v1_funct_2(esk4_0,k5_numbers,X2)|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_10049, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,np__0)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(X2,np__0))))).#trainneg
cnf(i_0_10051, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2553, negated_conjecture, (v1_rat_1(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_2489, plain, (v2_xxreal_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|v2_struct_0(X1)|~v2_xxreal_0(k1_normsp_0(X1,k1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|~v1_xreal_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|~v1_xreal_0(k1_normsp_0(X1,k1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|~v7_ordinal1(X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2490, plain, (v2_xxreal_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|v3_xxreal_0(k1_normsp_0(X1,k1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|v1_xboole_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|v2_struct_0(X1)|~v1_xreal_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|~v1_xreal_0(k1_normsp_0(X1,k1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|~v7_ordinal1(X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_10066, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X2)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_10082, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,X3)|~v1_funct_2(esk4_0,k5_numbers,X1)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_2554, plain, (v1_valued_0(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3029, negated_conjecture, (v1_valued_0(esk4_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_3032, negated_conjecture, (v2_valued_0(esk4_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_2491, plain, (v2_xxreal_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|v3_xxreal_0(k1_normsp_0(X1,k1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|v1_xboole_0(k1_normsp_0(X1,k1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|v2_struct_0(X1)|~v1_xreal_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|~v1_xreal_0(k1_normsp_0(X1,k1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|~v7_ordinal1(X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_10110, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,np__0)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(np__0,X3))))).#trainneg
cnf(i_0_10119, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_10120, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v1_funct_2(esk4_0,k5_numbers,X2)|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_10145, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,np__0)|~v1_funct_2(esk4_0,k5_numbers,X1)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(X3,np__0)))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_2557, plain, (v1_valued_0(esk32_2(X1,X2))|~v1_membered(X2))).#trainneg
cnf(i_0_3033, negated_conjecture, (v2_valued_0(esk4_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_3292, plain, (v6_membered(esk52_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k5_numbers))).#trainneg
cnf(i_0_2558, plain, (v1_valued_0(esk28_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_membered(X2))).#trainneg
cnf(i_0_10178, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)),k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,np__0)|~v1_funct_2(esk4_0,k5_numbers,X1)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(np__0,X4)))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_9831, plain, (v1_funct_2(X1,X2,np__0)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(np__0,X3))))).#trainneg
cnf(i_0_2560, plain, (v1_valued_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v1_membered(k1_numbers)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_3430, plain, (v5_membered(esk52_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k5_numbers))).#trainneg
cnf(i_0_2561, plain, (v1_valued_0(k3_series_1(X1))|~v1_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2562, plain, (v1_valued_0(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v1_membered(u1_struct_0(X1))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_10179, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|k2_seq_2(esk61_3(esk1_0,esk2_0,esk3_0))!=np__0|~v2_comseq_2(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0)),k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(esk1_0,esk2_0,esk3_0),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk61_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xboole_0(esk62_3(esk1_0,esk4_0,esk61_3(esk1_0,esk2_0,esk3_0))))).#trainneg
cnf(i_0_2509, negated_conjecture, (v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v3_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_9857, plain, (k2_zfmisc_1(np__0,X1)=k2_zfmisc_1(np__0,X3))).#trainneg
cnf(i_0_2563, plain, (v1_valued_0(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v1_membered(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2565, plain, (v2_valued_0(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2568, plain, (v2_valued_0(esk32_2(X1,X2))|~v2_membered(X2))).#trainneg
cnf(i_0_10219, negated_conjecture, (~m1_subset_1(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k5_numbers)|~v3_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_10220, negated_conjecture, (~v1_xboole_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v3_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2510, negated_conjecture, (v2_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v2_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2511, negated_conjecture, (v2_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|v1_xboole_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_3948, plain, (v3_membered(esk52_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k5_numbers))).#trainneg
cnf(i_0_2569, plain, (v2_valued_0(esk28_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v2_membered(X2))).#trainneg
cnf(i_0_6875, plain, (v3_xxreal_0(esk39_0)|r1_xxreal_0(esk54_0,esk39_0)|~v1_xreal_0(esk54_0)|~v1_xreal_0(esk39_0))).#trainneg
cnf(i_0_5563, plain, (v4_membered(esk52_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k5_numbers))).#trainneg
cnf(i_0_2571, plain, (v2_valued_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v2_membered(k1_numbers)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2572, plain, (v2_valued_0(k3_series_1(X1))|~v2_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_10404, plain, (r1_xxreal_0(esk54_0,esk39_0)|~m1_subset_1(esk39_0,k5_numbers)|~v1_xreal_0(esk54_0)|~v1_xreal_0(esk39_0))).#trainneg
cnf(i_0_2512, negated_conjecture, (v2_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|v1_xboole_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2519, negated_conjecture, (v3_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v3_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v7_ordinal1(X1)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2520, negated_conjecture, (v2_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v2_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v7_ordinal1(X1)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_10421, negated_conjecture, (~v1_xboole_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v2_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v7_ordinal1(X1)|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_10422, negated_conjecture, (~v3_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v2_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v7_ordinal1(X1)|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_10425, negated_conjecture, (~v2_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~v1_xboole_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v1_funct_2(esk4_0,k5_numbers,X2)|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_10428, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),esk59_0)|~v1_xboole_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_10429, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),esk55_0)|~v1_xboole_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_10430, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),esk47_0)|~v1_xboole_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_10440, negated_conjecture, (~v2_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~v3_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v1_funct_2(esk4_0,k5_numbers,X2)|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_9975, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(np__0,X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7488, plain, (v1_int_1(esk52_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_2573, plain, (v2_valued_0(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v2_membered(u1_struct_0(X1))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_5895, negated_conjecture, (v5_membered(esk4_0)|~v6_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_5945, negated_conjecture, (v1_membered(esk4_0)|~v3_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_10513, negated_conjecture, (v5_membered(esk4_0)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_10520, negated_conjecture, (v3_membered(esk4_0)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2574, plain, (v2_valued_0(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v2_membered(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2579, plain, (v3_valued_0(esk32_2(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_10443, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),esk59_0)|~v3_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_10444, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),esk55_0)|~v3_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_10445, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),esk47_0)|~v3_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_10460, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)),esk59_0)|~v1_xboole_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v7_ordinal1(X2)|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X1)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_10461, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)),esk55_0)|~v1_xboole_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v7_ordinal1(X2)|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X1)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_10462, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)),esk47_0)|~v1_xboole_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v7_ordinal1(X2)|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)))|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X1)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_7664, plain, (v7_ordinal1(esk52_0)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_2580, plain, (v3_valued_0(esk28_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v3_membered(X2))).#trainneg
cnf(i_0_10549, plain, (~v3_xxreal_0(esk52_0)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_5946, negated_conjecture, (v2_membered(esk4_0)|~v3_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_10514, negated_conjecture, (v5_membered(esk4_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_10560, negated_conjecture, (v3_membered(esk4_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_2584, plain, (v3_valued_0(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2585, plain, (v3_valued_0(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2521, negated_conjecture, (v2_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|v3_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|v1_xboole_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v7_ordinal1(X1)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2522, negated_conjecture, (v2_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|v3_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|v1_xboole_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v7_ordinal1(X1)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2525, negated_conjecture, (v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)))|~v3_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X2))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)))|~v7_ordinal1(X2)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X1)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_2526, negated_conjecture, (v2_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v2_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~v7_ordinal1(X1)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X2)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_2527, negated_conjecture, (v2_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|v1_xboole_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~v7_ordinal1(X1)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X2)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_2528, negated_conjecture, (v2_xxreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|v1_xboole_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~v1_xreal_0(k1_funct_1(esk61_3(esk1_0,esk2_0,esk3_0),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X2,esk4_0,X1)))|~v7_ordinal1(X1)|~v3_valued_0(esk61_3(esk1_0,esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X2)|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_2536, plain, (v3_xxreal_0(esk7_2(X1,X2))|v2_comseq_2(X1)|~v3_xxreal_0(k18_complex1(k6_xcmplx_0(k8_nat_1(X3,X1,esk8_3(X1,X2,X4)),X2)))|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k8_nat_1(X3,X1,esk8_3(X1,X2,X4)),X2)))|~v1_xreal_0(esk7_2(X1,X2))|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1)|~m1_subset_1(k8_nat_1(X3,X1,esk8_3(X1,X2,X4)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_2537, plain, (v2_xxreal_0(k18_complex1(k6_xcmplx_0(k8_nat_1(X1,X2,esk8_3(X2,X3,X4)),X3)))|v2_comseq_2(X2)|~v2_xxreal_0(esk7_2(X2,X3))|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k8_nat_1(X1,X2,esk8_3(X2,X3,X4)),X3)))|~v1_xreal_0(esk7_2(X2,X3))|~v1_xreal_0(X3)|~v3_valued_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(k8_nat_1(X1,X2,esk8_3(X2,X3,X4)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_2538, plain, (v2_xxreal_0(k18_complex1(k6_xcmplx_0(k8_nat_1(X1,X2,esk8_3(X2,X3,X4)),X3)))|v3_xxreal_0(esk7_2(X2,X3))|v1_xboole_0(k18_complex1(k6_xcmplx_0(k8_nat_1(X1,X2,esk8_3(X2,X3,X4)),X3)))|v2_comseq_2(X2)|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k8_nat_1(X1,X2,esk8_3(X2,X3,X4)),X3)))|~v1_xreal_0(esk7_2(X2,X3))|~v1_xreal_0(X3)|~v3_valued_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(k8_nat_1(X1,X2,esk8_3(X2,X3,X4)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_2539, plain, (v2_xxreal_0(k18_complex1(k6_xcmplx_0(k8_nat_1(X1,X2,esk8_3(X2,X3,X4)),X3)))|v3_xxreal_0(esk7_2(X2,X3))|v1_xboole_0(esk7_2(X2,X3))|v2_comseq_2(X2)|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k8_nat_1(X1,X2,esk8_3(X2,X3,X4)),X3)))|~v1_xreal_0(esk7_2(X2,X3))|~v1_xreal_0(X3)|~v3_valued_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(k8_nat_1(X1,X2,esk8_3(X2,X3,X4)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_7668, plain, (v3_ordinal1(esk52_0)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_9612, negated_conjecture, (~l2_algstr_0(esk1_0)|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v4_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_5959, negated_conjecture, (v3_membered(esk4_0)|~v6_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_5966, negated_conjecture, (v4_membered(esk4_0)|~v5_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2590, plain, (v4_valued_0(esk32_2(X1,X2))|~v6_membered(X2))).#trainneg
cnf(i_0_2591, plain, (v4_valued_0(esk28_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v6_membered(X2))).#trainneg
cnf(i_0_2543, plain, (r1_xxreal_0(esk7_2(X1,np__0),k18_complex1(k8_nat_1(X2,X1,esk8_3(X1,np__0,X3))))|v2_comseq_2(X1)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,X2)|~v1_funct_1(X1)|~m1_subset_1(k8_nat_1(X2,X1,esk8_3(X1,np__0,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_xcmplx_0(k8_nat_1(X2,X1,esk8_3(X1,np__0,X3))))).#trainneg
cnf(i_0_2544, plain, (v3_xxreal_0(esk10_2(X1,X2))|X2=k1_seq_2(X1)|~v3_xxreal_0(k18_complex1(k6_xcmplx_0(k8_nat_1(X3,X1,esk11_3(X1,X2,X4)),X2)))|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k8_nat_1(X3,X1,esk11_3(X1,X2,X4)),X2)))|~v1_xreal_0(esk10_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v3_valued_0(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1)|~m1_subset_1(k8_nat_1(X3,X1,esk11_3(X1,X2,X4)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_2545, plain, (v2_xxreal_0(k18_complex1(k6_xcmplx_0(k8_nat_1(X1,X2,esk11_3(X2,X3,X4)),X3)))|X3=k1_seq_2(X2)|~v2_xxreal_0(esk10_2(X2,X3))|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k8_nat_1(X1,X2,esk11_3(X2,X3,X4)),X3)))|~v1_xreal_0(esk10_2(X2,X3))|~v2_comseq_2(X2)|~v1_xreal_0(X3)|~v3_valued_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(k8_nat_1(X1,X2,esk11_3(X2,X3,X4)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_2546, plain, (v2_xxreal_0(k18_complex1(k6_xcmplx_0(k8_nat_1(X1,X2,esk11_3(X2,X3,X4)),X3)))|v3_xxreal_0(esk10_2(X2,X3))|v1_xboole_0(k18_complex1(k6_xcmplx_0(k8_nat_1(X1,X2,esk11_3(X2,X3,X4)),X3)))|X3=k1_seq_2(X2)|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k8_nat_1(X1,X2,esk11_3(X2,X3,X4)),X3)))|~v1_xreal_0(esk10_2(X2,X3))|~v2_comseq_2(X2)|~v1_xreal_0(X3)|~v3_valued_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(k8_nat_1(X1,X2,esk11_3(X2,X3,X4)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_2547, plain, (v2_xxreal_0(k18_complex1(k6_xcmplx_0(k8_nat_1(X1,X2,esk11_3(X2,X3,X4)),X3)))|v3_xxreal_0(esk10_2(X2,X3))|v1_xboole_0(esk10_2(X2,X3))|X3=k1_seq_2(X2)|~v1_xreal_0(k18_complex1(k6_xcmplx_0(k8_nat_1(X1,X2,esk11_3(X2,X3,X4)),X3)))|~v1_xreal_0(esk10_2(X2,X3))|~v2_comseq_2(X2)|~v1_xreal_0(X3)|~v3_valued_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(k8_nat_1(X1,X2,esk11_3(X2,X3,X4)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_2551, plain, (k1_seq_2(X1)=np__0|r1_xxreal_0(esk10_2(X1,np__0),k18_complex1(k8_nat_1(X2,X1,esk11_3(X1,np__0,X3))))|~v2_comseq_2(X1)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,X2)|~v1_funct_1(X1)|~m1_subset_1(k8_nat_1(X2,X1,esk11_3(X1,np__0,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_xcmplx_0(k8_nat_1(X2,X1,esk11_3(X1,np__0,X3))))).#trainneg
cnf(i_0_10638, negated_conjecture, (~l2_algstr_0(esk1_0)|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v4_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5968, negated_conjecture, (v1_valued_0(esk4_0)|~v3_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_10710, negated_conjecture, (v1_valued_0(esk4_0)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2595, plain, (v4_valued_0(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v6_membered(u1_struct_0(X1))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2596, plain, (v4_valued_0(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v6_membered(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2601, plain, (v1_xboole_0(esk32_2(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2602, plain, (v1_xboole_0(esk28_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2606, plain, (v1_xboole_0(k1_bhsp_4(X1,X2))|v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2607, plain, (v1_xboole_0(k8_lopban_4(X1,X2,X3,X4))|v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2612, plain, (v1_xboole_0(esk32_2(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2613, plain, (v1_xboole_0(esk28_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4695, plain, (r1_xxreal_0(esk59_0,esk44_0)|r1_xxreal_0(esk44_0,esk59_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_6059, negated_conjecture, (v2_valued_0(esk4_0)|~v3_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_10821, negated_conjecture, (v2_valued_0(esk4_0)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2626, negated_conjecture, (v1_int_1(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2627, plain, (r1_xxreal_0(X1,np__0)|v2_struct_0(X2)|~r1_xxreal_0(X1,k18_complex1(k1_funct_1(esk61_3(X2,X4,X5),X6)))|~r1_xxreal_0(esk63_4(X2,X4,X5,X1),X6)|~v7_ordinal1(X6)|~v1_xreal_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(X2,X4,X5),k5_numbers,X3)|~m1_subset_1(esk61_3(X2,X4,X5),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)|~v1_funct_1(esk61_3(X2,X4,X5)))).#trainneg
cnf(i_0_10825, plain, (r1_xxreal_0(X1,np__0)|v2_struct_0(X2)|~r1_xxreal_0(X1,k18_complex1(k1_funct_1(esk61_3(X2,X3,X4),X5)))|~r1_xxreal_0(esk63_4(X2,X3,X4,X1),X5)|~v7_ordinal1(X5)|~v1_xreal_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(X2,X3,X4),k5_numbers,np__0)|~m1_subset_1(esk61_3(X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(X6,np__0)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_10827, plain, (r1_xxreal_0(X1,np__0)|v2_struct_0(X2)|~r1_xxreal_0(X1,k18_complex1(k1_funct_1(esk61_3(X2,X3,X4),X5)))|~r1_xxreal_0(esk63_4(X2,X3,X4,X1),X5)|~v7_ordinal1(X5)|~v1_xreal_0(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(X2,X3,X4),k5_numbers,k1_numbers)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_2630, plain, (r1_xxreal_0(k18_complex1(k8_nat_1(X1,esk61_3(X2,X3,X4),X5)),np__0)|v2_struct_0(X2)|~r1_xxreal_0(esk63_4(X2,X3,X4,k18_complex1(k8_nat_1(X1,esk61_3(X2,X3,X4),X5))),X5)|~v7_ordinal1(X5)|~v1_xreal_0(k18_complex1(k8_nat_1(X1,esk61_3(X2,X3,X4),X5)))|~m2_subset_1(X5,k1_numbers,k5_numbers)|~v1_funct_2(esk61_3(X2,X3,X4),k5_numbers,X1)|~m1_subset_1(esk61_3(X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_2639, plain, (v5_relat_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_2641, plain, (v5_relat_1(k4_normsp_0(X1,X2),k1_numbers)|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_10547, plain, (v1_xxreal_0(esk52_0)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_10718, negated_conjecture, (~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~l2_algstr_0(esk1_0)|~v6_membered(u1_struct_0(esk1_0))|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2642, plain, (v5_relat_1(k3_series_1(X1),k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4704, negated_conjecture, (r1_xxreal_0(esk59_0,esk3_0)|r1_xxreal_0(esk3_0,esk59_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_7409, plain, (v1_xcmplx_0(esk52_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_10515, negated_conjecture, (v5_membered(esk4_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_10516, negated_conjecture, (v1_membered(esk4_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2650, plain, (v4_relat_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_2652, plain, (v4_relat_1(k4_normsp_0(X1,X2),k5_numbers)|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2653, plain, (v4_relat_1(k3_series_1(X1),k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2654, plain, (v4_relat_1(k1_bhsp_4(X1,X2),k5_numbers)|v2_struct_0(X1)|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2655, plain, (v4_relat_1(k8_lopban_4(X1,X2,X3,X4),k5_numbers)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2659, negated_conjecture, (v1_xcmplx_0(X1)|v1_xboole_0(esk4_0)|~v1_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~m1_subset_1(X1,esk4_0))).#trainneg
cnf(i_0_2660, negated_conjecture, (v1_xxreal_0(X1)|v1_xboole_0(esk4_0)|~v2_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~m1_subset_1(X1,esk4_0))).#trainneg
cnf(i_0_2661, negated_conjecture, (v1_xreal_0(X1)|v1_xboole_0(esk4_0)|~v3_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~m1_subset_1(X1,esk4_0))).#trainneg
cnf(i_0_2662, negated_conjecture, (v1_rat_1(X1)|v1_xboole_0(esk4_0)|~v4_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~m1_subset_1(X1,esk4_0))).#trainneg
cnf(i_0_2663, negated_conjecture, (v1_int_1(X1)|v1_xboole_0(esk4_0)|~v5_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~m1_subset_1(X1,esk4_0))).#trainneg
cnf(i_0_2664, negated_conjecture, (v7_ordinal1(X1)|v1_xboole_0(esk4_0)|~v6_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~m1_subset_1(X1,esk4_0))).#trainneg
cnf(i_0_2665, negated_conjecture, (v3_ordinal1(X1)|v1_xboole_0(esk4_0)|~v3_ordinal1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~m1_subset_1(X1,esk4_0))).#trainneg
cnf(i_0_10884, negated_conjecture, (~l2_algstr_0(esk1_0)|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v6_membered(u1_struct_0(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2680, plain, (v5_relat_1(esk61_3(X1,X2,X3),k4_numbers)|v2_struct_0(X1)|~v5_membered(k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_4709, negated_conjecture, (r1_xxreal_0(esk59_0,esk2_0)|r1_xxreal_0(esk2_0,esk59_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_7430, plain, (v1_xxreal_0(esk52_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_7432, plain, (v1_xreal_0(esk52_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_10556, negated_conjecture, (v2_membered(esk4_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_10641, negated_conjecture, (v3_membered(esk4_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2683, plain, (v5_relat_1(esk32_2(X1,X2),k4_numbers)|~v5_membered(X2))).#trainneg
cnf(i_0_2684, plain, (v5_relat_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),k4_numbers)|~v5_membered(X2))).#trainneg
cnf(i_0_2685, negated_conjecture, (v5_relat_1(esk4_0,k4_numbers)|~v5_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2686, plain, (v5_relat_1(k4_normsp_0(X1,X2),k4_numbers)|v2_struct_0(X1)|~v5_membered(k1_numbers)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2687, plain, (v5_relat_1(k3_series_1(X1),k4_numbers)|~v5_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2688, plain, (v5_relat_1(k1_bhsp_4(X1,X2),k4_numbers)|v2_struct_0(X1)|~v5_membered(u1_struct_0(X1))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2689, plain, (v5_relat_1(k8_lopban_4(X1,X2,X3,X4),k4_numbers)|v2_struct_0(X1)|~v5_membered(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2691, plain, (v5_relat_1(esk61_3(X1,X2,X3),k3_numbers)|v2_struct_0(X1)|~v4_membered(k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2694, plain, (v5_relat_1(esk32_2(X1,X2),k3_numbers)|~v4_membered(X2))).#trainneg
cnf(i_0_2695, plain, (v5_relat_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),k3_numbers)|~v4_membered(X2))).#trainneg
cnf(i_0_2696, negated_conjecture, (v5_relat_1(esk4_0,k3_numbers)|~v4_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2697, plain, (v5_relat_1(k4_normsp_0(X1,X2),k3_numbers)|v2_struct_0(X1)|~v4_membered(k1_numbers)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_10929, negated_conjecture, (~l2_algstr_0(esk1_0)|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v6_membered(u1_struct_0(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2698, plain, (v5_relat_1(k3_series_1(X1),k3_numbers)|~v4_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4730, plain, (r1_xxreal_0(esk59_0,k5_numbers)|r1_xxreal_0(k5_numbers,esk59_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_7486, plain, (v1_rat_1(esk52_0)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_7667, plain, (v3_ordinal1(esk52_0)|~v7_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_2699, plain, (v5_relat_1(k1_bhsp_4(X1,X2),k3_numbers)|v2_struct_0(X1)|~v4_membered(u1_struct_0(X1))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2700, plain, (v5_relat_1(k8_lopban_4(X1,X2,X3,X4),k3_numbers)|v2_struct_0(X1)|~v4_membered(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2707, plain, (v1_funct_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1,X2)|~v1_partfun1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_2713, plain, (v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v1_partfun1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_2723, negated_conjecture, (v3_ordinal1(esk4_0)|~v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))))).#trainneg
cnf(i_0_2728, plain, (~v3_xxreal_0(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_10987, plain, (r1_xxreal_0(esk53_0,esk38_0)|~v1_xboole_0(k1_zfmisc_1(esk38_0)))).#trainneg
cnf(i_0_2729, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xboole_0(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_2735, plain, (r1_xxreal_0(X1,esk5_1(X2))|v2_xxreal_0(X1)|v3_xxreal_0(esk5_1(X2))|~v1_xreal_0(X1)|~v2_comseq_2(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2736, plain, (v1_xcmplx_0(k2_seq_2(X1))|~v1_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2737, plain, (v1_xxreal_0(k2_seq_2(X1))|~v2_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2739, plain, (v1_rat_1(k2_seq_2(X1))|~v4_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_10963, plain, (v1_rat_1(esk52_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_10958, negated_conjecture, (~l2_algstr_0(esk1_0)|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v6_membered(u1_struct_0(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2740, plain, (v1_int_1(k2_seq_2(X1))|~v5_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4754, plain, (r1_xxreal_0(esk55_0,esk44_0)|r1_xxreal_0(esk44_0,esk55_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_10983, plain, (~v1_xboole_0(k1_zfmisc_1(esk54_0)))).#trainneg
cnf(i_0_10984, plain, (~v1_xboole_0(k1_zfmisc_1(esk53_0)))).#trainneg
cnf(i_0_2797, negated_conjecture, (v3_ordinal1(esk4_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))),k5_numbers))).#trainneg
cnf(i_0_2742, plain, (v3_ordinal1(k2_seq_2(X1))|~v3_ordinal1(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2743, plain, (m1_subset_1(k1_seq_2(X1),k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2746, plain, (r1_xxreal_0(X1,k2_seq_2(esk61_3(X2,X3,X4)))|v2_xxreal_0(X1)|v3_xxreal_0(k2_seq_2(esk61_3(X2,X3,X4)))|v2_struct_0(X2)|~v1_xreal_0(X1)|~v1_funct_2(esk61_3(X2,X3,X4),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X2,X3,X4))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_2748, plain, (r1_xxreal_0(X1,k2_seq_2(esk32_2(k5_numbers,k1_numbers)))|v2_xxreal_0(X1)|v3_xxreal_0(k2_seq_2(esk32_2(k5_numbers,k1_numbers)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2749, plain, (r1_xxreal_0(X1,k2_seq_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))|v2_xxreal_0(X1)|v3_xxreal_0(k2_seq_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))|~v1_xreal_0(X1)|~v1_funct_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))),k5_numbers,k1_numbers)|~v1_funct_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))).#trainneg
cnf(i_0_2750, plain, (r1_xxreal_0(X1,k2_seq_2(k4_normsp_0(X2,X3)))|v2_xxreal_0(X1)|v3_xxreal_0(k2_seq_2(k4_normsp_0(X2,X3)))|v2_struct_0(X2)|~v1_xreal_0(X1)|~v1_funct_2(k4_normsp_0(X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(k4_normsp_0(X2,X3))|~l1_normsp_0(X2)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))))).#trainneg
cnf(i_0_2751, plain, (r1_xxreal_0(X1,k2_seq_2(k3_series_1(X2)))|v2_xxreal_0(X1)|v3_xxreal_0(k2_seq_2(k3_series_1(X2)))|~v1_xreal_0(X1)|~v1_funct_2(k3_series_1(X2),k5_numbers,k1_numbers)|~v1_funct_1(k3_series_1(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2752, plain, (r1_xxreal_0(X1,k2_seq_2(k2_zfmisc_1(k5_numbers,k1_numbers)))|v2_xxreal_0(X1)|v3_xxreal_0(k2_seq_2(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(X1)|~v1_funct_2(k2_zfmisc_1(k5_numbers,k1_numbers),k5_numbers,k1_numbers)|~v1_funct_1(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_2756, plain, (v1_partfun1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))),X1)|~v1_funct_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))),X1,X1))).#trainneg
cnf(i_0_2757, plain, (v1_partfun1(k2_zfmisc_1(X1,X1),X1)|~v1_funct_2(k2_zfmisc_1(X1,X1),X1,X1))).#trainneg
cnf(i_0_2761, plain, (v1_partfun1(esk32_2(X1,X2),X1)|v1_xboole_0(X2))).#trainneg
cnf(i_0_11138, negated_conjecture, (~l2_algstr_0(esk1_0)|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v6_membered(u1_struct_0(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2762, plain, (v1_partfun1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1)|v1_xboole_0(X2)|~v1_funct_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1,X2))).#trainneg
cnf(i_0_4763, negated_conjecture, (r1_xxreal_0(esk55_0,esk3_0)|r1_xxreal_0(esk3_0,esk55_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_10996, plain, (~v1_xboole_0(k1_zfmisc_1(k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_10517, negated_conjecture, (v1_membered(esk4_0)|~v6_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_11289, negated_conjecture, (v1_membered(esk4_0)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_11290, negated_conjecture, (v1_membered(esk4_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_2766, plain, (v1_partfun1(k1_bhsp_4(X1,X2),k5_numbers)|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2767, plain, (v1_partfun1(k8_lopban_4(X1,X2,X3,X4),k5_numbers)|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2764, plain, (v1_partfun1(k4_normsp_0(X1,X2),k5_numbers)|v2_struct_0(X1)|~v1_funct_2(k4_normsp_0(X1,X2),k5_numbers,k1_numbers)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2765, plain, (v1_partfun1(k3_series_1(X1),k5_numbers)|~v1_funct_2(k3_series_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2768, plain, (v1_partfun1(k2_zfmisc_1(X1,X2),X1)|v1_xboole_0(X2)|~v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2))).#trainneg
cnf(i_0_2772, plain, (v1_partfun1(esk32_2(X1,X2),X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2773, plain, (v1_partfun1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1)|~v1_xboole_0(X1)|~v1_funct_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1,X2))).#trainneg
cnf(i_0_2779, plain, (v1_partfun1(k2_zfmisc_1(X1,X2),X1)|~v1_xboole_0(X1)|~v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2))).#trainneg
cnf(i_0_2783, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(esk32_2(X1,X2)))).#trainneg
cnf(i_0_2784, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(esk28_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_funct_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1,X2)|~v1_funct_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_10511, negated_conjecture, (~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~l2_algstr_0(esk1_0)|~v2_membered(u1_struct_0(esk1_0))|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2790, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v1_funct_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_4768, negated_conjecture, (r1_xxreal_0(esk55_0,esk2_0)|r1_xxreal_0(esk2_0,esk55_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_10548, plain, (v1_xreal_0(esk52_0)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_10557, negated_conjecture, (v2_membered(esk4_0)|~v6_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_11331, negated_conjecture, (v2_membered(esk4_0)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_11332, negated_conjecture, (v2_membered(esk4_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_2803, plain, (r1_xxreal_0(X1,k3_series_1(X2))|v2_xxreal_0(X1)|v3_xxreal_0(k3_series_1(X2))|~v1_xreal_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2804, plain, (r1_xxreal_0(X1,esk7_2(X2,X3))|v2_xxreal_0(X1)|v3_xxreal_0(esk7_2(X2,X3))|v2_comseq_2(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2788, plain, (v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v1_xboole_0(k1_bhsp_4(X1,X2))|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,X2))|~l2_algstr_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2789, plain, (v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v1_xboole_0(k8_lopban_4(X1,X2,X3,X4))|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2802, plain, (~v3_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_11371, plain, (r1_xxreal_0(esk53_0,esk39_0)|~m1_subset_1(k1_zfmisc_1(esk39_0),k5_numbers)|~v1_xreal_0(esk39_0))).#trainneg
cnf(i_0_11372, plain, (r1_xxreal_0(esk53_0,esk38_0)|~m1_subset_1(k1_zfmisc_1(esk38_0),k5_numbers))).#trainneg
cnf(i_0_11373, plain, (r1_xxreal_0(esk54_0,esk38_0)|~m1_subset_1(k1_zfmisc_1(esk38_0),k5_numbers)|~v1_xreal_0(esk54_0))).#trainneg
cnf(i_0_2806, plain, (m1_subset_1(k8_nat_1(esk44_0,X1,X2),k1_numbers)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,esk44_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,esk44_0))))).#trainneg
cnf(i_0_2811, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4))|~v1_funct_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),X1,X2)|~v1_funct_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4))|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))).#trainneg
cnf(i_0_2812, plain, (v3_normsp_1(k1_bhsp_4(X1,k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2,X3)),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2,X3)),X4),k5_numbers)|v2_struct_0(X1)|k2_seq_2(X4)!=np__0|~v2_comseq_2(X4)|~v1_funct_2(k1_bhsp_4(X1,k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2,X3)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2,X3),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X4,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(X1,k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2,X3)))|~v1_funct_1(X4)|~v1_funct_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))))).#trainneg
cnf(i_0_2813, plain, (v3_normsp_1(k1_bhsp_4(X1,k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2,X3)),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2,X3)),X4),k1_numbers)|v2_struct_0(X1)|k2_seq_2(X4)!=np__0|~v2_comseq_2(X4)|~v1_funct_2(k1_bhsp_4(X1,k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2,X3)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2,X3),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X4,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(X1,k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2,X3)))|~v1_funct_1(X4)|~v1_funct_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))))).#trainneg
cnf(i_0_2814, plain, (v1_relat_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4))|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))).#trainneg
cnf(i_0_2815, plain, (v1_xboole_0(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4))|~v1_xboole_0(X1)|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))).#trainneg
cnf(i_0_11368, plain, (~m1_subset_1(k1_zfmisc_1(esk54_0),k5_numbers))).#trainneg
cnf(i_0_2816, plain, (v4_relat_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),X1)|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))).#trainneg
cnf(i_0_4789, plain, (r1_xxreal_0(esk55_0,k5_numbers)|r1_xxreal_0(k5_numbers,esk55_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_11369, plain, (~m1_subset_1(k1_zfmisc_1(esk53_0),k5_numbers))).#trainneg
cnf(i_0_11381, plain, (~m1_subset_1(k1_zfmisc_1(k4_xcmplx_0(np__1)),k5_numbers))).#trainneg
cnf(i_0_11291, negated_conjecture, (v1_membered(esk4_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_11333, negated_conjecture, (v2_membered(esk4_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2817, plain, (v1_funct_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),X1,X2)|~v1_partfun1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),X1)|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))).#trainneg
cnf(i_0_2818, plain, (v1_partfun1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),X1)|v1_xboole_0(X2)|~v1_funct_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),X1,X2)|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))).#trainneg
cnf(i_0_2819, plain, (v1_partfun1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),X1)|~v1_xboole_0(X1)|~v1_funct_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),X1,X2)|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))).#trainneg
cnf(i_0_2820, plain, (v1_valued_0(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4))|~v1_membered(X2)|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))).#trainneg
cnf(i_0_2821, plain, (v2_valued_0(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4))|~v2_membered(X2)|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))).#trainneg
cnf(i_0_2822, plain, (v3_valued_0(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4))|~v3_membered(X2)|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))).#trainneg
cnf(i_0_2823, plain, (v4_valued_0(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4))|~v6_membered(X2)|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))).#trainneg
cnf(i_0_2824, plain, (v1_xboole_0(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4))|~v1_xboole_0(X2)|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))).#trainneg
cnf(i_0_2825, plain, (v5_relat_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),X2)|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))).#trainneg
cnf(i_0_2826, plain, (v5_relat_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),k4_numbers)|~v5_membered(X2)|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))).#trainneg
cnf(i_0_2827, plain, (v5_relat_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),k3_numbers)|~v4_membered(X2)|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))).#trainneg
cnf(i_0_2828, plain, (v1_partfun1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1)),X2,X3),X1)|~v1_funct_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1)),X2,X3),X1,X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))).#trainneg
cnf(i_0_2829, plain, (r1_xxreal_0(X1,k1_seq_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X2,X3)))|v2_xxreal_0(X1)|v3_xxreal_0(k1_seq_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X2,X3)))|~v1_xreal_0(X1)|~v1_funct_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X2,X3))|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))).#trainneg
cnf(i_0_2830, plain, (r1_xxreal_0(X1,k2_seq_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X2,X3)))|v2_xxreal_0(X1)|v3_xxreal_0(k2_seq_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X2,X3)))|~v1_xreal_0(X1)|~v1_funct_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X2,X3))|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))).#trainneg
cnf(i_0_11320, negated_conjecture, (~l2_algstr_0(esk1_0)|~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v2_membered(u1_struct_0(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2833, plain, (v5_membered(k8_nat_1(k1_zfmisc_1(X1),X2,X3))|~v5_membered(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_4813, plain, (r1_xxreal_0(esk54_0,esk44_0)|r1_xxreal_0(esk44_0,esk54_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_10642, negated_conjecture, (v4_membered(esk4_0)|~v6_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_10890, plain, (v1_xcmplx_0(esk52_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_11672, negated_conjecture, (v4_membered(esk4_0)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2834, plain, (v6_ordinal1(k8_nat_1(k1_zfmisc_1(X1),X2,X3))|~v6_ordinal1(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_2835, plain, (v6_membered(k8_nat_1(k1_zfmisc_1(X1),X2,X3))|~v6_membered(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_2831, plain, (v3_normsp_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2,X3),X1)|m1_subset_1(esk62_3(X1,k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2,X3),X4),k5_numbers)|v2_struct_0(X1)|k2_seq_2(X4)!=np__0|~v2_comseq_2(X4)|~v1_funct_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2,X3),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X4,k5_numbers,k1_numbers)|~v1_funct_1(X4)|~v1_funct_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))))).#trainneg
cnf(i_0_2832, plain, (v3_normsp_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2,X3),X1)|m1_subset_1(esk62_3(X1,k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2,X3),X4),k1_numbers)|v2_struct_0(X1)|k2_seq_2(X4)!=np__0|~v2_comseq_2(X4)|~v1_funct_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2,X3),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X4,k5_numbers,k1_numbers)|~v1_funct_1(X4)|~v1_funct_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))))).#trainneg
cnf(i_0_2836, plain, (v1_membered(k8_nat_1(k1_zfmisc_1(X1),X2,X3))|~v1_membered(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_2837, plain, (v2_membered(k8_nat_1(k1_zfmisc_1(X1),X2,X3))|~v2_membered(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_2838, plain, (v3_membered(k8_nat_1(k1_zfmisc_1(X1),X2,X3))|~v3_membered(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_2839, plain, (v4_membered(k8_nat_1(k1_zfmisc_1(X1),X2,X3))|~v4_membered(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_2840, plain, (v1_valued_0(k8_nat_1(k1_zfmisc_1(X1),X2,X3))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_2841, plain, (v2_valued_0(k8_nat_1(k1_zfmisc_1(X1),X2,X3))|~v2_valued_0(X1)|~v1_relat_1(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_2842, plain, (v3_valued_0(k8_nat_1(k1_zfmisc_1(X1),X2,X3))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_2843, plain, (v4_valued_0(k8_nat_1(k1_zfmisc_1(X1),X2,X3))|~v4_valued_0(X1)|~v1_relat_1(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_2844, plain, (v5_relat_1(k8_nat_1(k1_zfmisc_1(X1),X2,X3),k4_numbers)|~v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_2845, plain, (v5_relat_1(k8_nat_1(k1_zfmisc_1(X1),X2,X3),k3_numbers)|~v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_11675, plain, (v1_xcmplx_0(esk52_0)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_11659, negated_conjecture, (~l2_algstr_0(esk1_0)|~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v2_membered(u1_struct_0(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2846, plain, (v1_xboole_0(k8_nat_1(k1_zfmisc_1(X1),X2,X3))|v1_xboole_0(X1)|m1_subset_1(X4,X1)|~m1_subset_1(X4,k8_nat_1(k1_zfmisc_1(X1),X2,X3))|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_4822, negated_conjecture, (r1_xxreal_0(esk54_0,esk3_0)|r1_xxreal_0(esk3_0,esk54_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_10935, plain, (v1_xxreal_0(esk52_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_11673, negated_conjecture, (v4_membered(esk4_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_2847, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k8_nat_1(k1_zfmisc_1(X2),X3,X4))|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X2)))))).#trainneg
cnf(i_0_2849, plain, (v1_xcmplx_0(k8_nat_1(X1,X2,X3))|~v1_membered(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_2848, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k8_nat_1(k1_zfmisc_1(X1),X3,X4))|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_2850, plain, (v1_xxreal_0(k8_nat_1(X1,X2,X3))|~v2_membered(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_11805, plain, (r1_xxreal_0(k8_nat_1(X1,X2,X3),k8_nat_1(X1,X2,X3))|~v7_ordinal1(X3)|~v2_membered(X1)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_11836, negated_conjecture, (r1_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2),k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2))|~v7_ordinal1(X2)|~v2_membered(X1)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k5_numbers,X1)|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k5_numbers,u1_struct_0(esk1_0))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_11837, plain, (r1_xxreal_0(k1_funct_1(X2,X3),k1_funct_1(X2,X3))|~v7_ordinal1(X3)|~v2_membered(X1)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_2851, plain, (v1_xreal_0(k8_nat_1(X1,X2,X3))|~v3_membered(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_2852, plain, (v1_rat_1(k8_nat_1(X1,X2,X3))|~v4_membered(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_2853, plain, (v1_int_1(k8_nat_1(X1,X2,X3))|~v5_membered(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_2854, plain, (v7_ordinal1(k8_nat_1(X1,X2,X3))|~v6_membered(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_11758, negated_conjecture, (~l2_algstr_0(esk1_0)|~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v2_membered(u1_struct_0(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2855, plain, (v3_ordinal1(k8_nat_1(X1,X2,X3))|~v3_ordinal1(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_4827, negated_conjecture, (r1_xxreal_0(esk54_0,esk2_0)|r1_xxreal_0(esk2_0,esk54_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_3235, plain, (v5_membered(esk52_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_3293, plain, (v6_membered(esk52_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_11674, negated_conjecture, (v4_membered(esk4_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2856, plain, (m1_subset_1(k1_seq_1(esk61_3(X1,X2,X3),X4),k1_numbers)|v2_struct_0(X1)|~v7_ordinal1(X4)|~v1_funct_2(k3_series_1(k4_normsp_0(X1,k8_lopban_4(X1,X2,X3,X4))),k5_numbers,k1_numbers)|~v1_funct_1(k3_series_1(k4_normsp_0(X1,k8_lopban_4(X1,X2,X3,X4))))|~m1_subset_1(k3_series_1(k4_normsp_0(X1,k8_lopban_4(X1,X2,X3,X4))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2860, plain, (m1_subset_1(k8_nat_1(X4,X2,X3),X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4))))).#trainneg
cnf(i_0_2863, negated_conjecture, (m1_subset_1(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1),X2)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,X2)|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~l2_algstr_0(esk1_0)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2865, plain, (m1_subset_1(k1_funct_1(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,esk44_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,esk44_0))))).#trainneg
cnf(i_0_12032, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,esk44_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,esk44_0))))).#trainneg
cnf(i_0_2870, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(k1_funct_1(X3,X4))|~v1_funct_2(k1_funct_1(X3,X4),X1,X2)|~v1_funct_1(k1_funct_1(X3,X4))|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))).#trainneg
cnf(i_0_2871, plain, (v3_normsp_1(k1_bhsp_4(X1,k1_funct_1(X2,X3)),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,k1_funct_1(X2,X3)),X4),k5_numbers)|v2_struct_0(X1)|k2_seq_2(X4)!=np__0|~v2_comseq_2(X4)|~v1_funct_2(k1_bhsp_4(X1,k1_funct_1(X2,X3)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k1_funct_1(X2,X3),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X4,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(X1,k1_funct_1(X2,X3)))|~v1_funct_1(X4)|~v1_funct_1(k1_funct_1(X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))))).#trainneg
cnf(i_0_2872, plain, (v3_normsp_1(k1_bhsp_4(X1,k1_funct_1(X2,X3)),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,k1_funct_1(X2,X3)),X4),k1_numbers)|v2_struct_0(X1)|k2_seq_2(X4)!=np__0|~v2_comseq_2(X4)|~v1_funct_2(k1_bhsp_4(X1,k1_funct_1(X2,X3)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k1_funct_1(X2,X3),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X4,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(X1,k1_funct_1(X2,X3)))|~v1_funct_1(X4)|~v1_funct_1(k1_funct_1(X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))))).#trainneg
cnf(i_0_2873, plain, (v1_relat_1(k1_funct_1(X1,X2))|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))))).#trainneg
cnf(i_0_2874, plain, (v1_xboole_0(k1_funct_1(X1,X2))|~v1_xboole_0(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))))).#trainneg
cnf(i_0_2875, plain, (v4_relat_1(k1_funct_1(X1,X2),X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))))).#trainneg
cnf(i_0_11908, negated_conjecture, (~l2_algstr_0(esk1_0)|~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v2_membered(u1_struct_0(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2876, plain, (v1_funct_2(k1_funct_1(X1,X2),X3,X4)|~v1_partfun1(k1_funct_1(X1,X2),X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))))).#trainneg
cnf(i_0_4848, plain, (r1_xxreal_0(esk54_0,k5_numbers)|r1_xxreal_0(k5_numbers,esk54_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_3388, plain, (v3_membered(esk52_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3390, plain, (v3_membered(esk52_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_9122, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1)))),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_9123, negated_conjecture, (v3_ordinal1(esk3_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2877, plain, (v1_partfun1(k1_funct_1(X1,X2),X3)|v1_xboole_0(X4)|~v1_funct_2(k1_funct_1(X1,X2),X3,X4)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))))).#trainneg
cnf(i_0_2878, plain, (v1_partfun1(k1_funct_1(X1,X2),X3)|~v1_xboole_0(X3)|~v1_funct_2(k1_funct_1(X1,X2),X3,X4)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))))).#trainneg
cnf(i_0_2879, plain, (v1_valued_0(k1_funct_1(X1,X2))|~v1_membered(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X4,X3)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X4,X3))))))).#trainneg
cnf(i_0_2880, plain, (v2_valued_0(k1_funct_1(X1,X2))|~v2_membered(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X4,X3)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X4,X3))))))).#trainneg
cnf(i_0_2881, plain, (v3_valued_0(k1_funct_1(X1,X2))|~v3_membered(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X4,X3)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X4,X3))))))).#trainneg
cnf(i_0_2882, plain, (v4_valued_0(k1_funct_1(X1,X2))|~v6_membered(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X4,X3)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X4,X3))))))).#trainneg
cnf(i_0_2883, plain, (v1_xboole_0(k1_funct_1(X1,X2))|~v1_xboole_0(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X4,X3)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X4,X3))))))).#trainneg
cnf(i_0_2884, plain, (v5_relat_1(k1_funct_1(X1,X2),X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X4,X3)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X4,X3))))))).#trainneg
cnf(i_0_2885, plain, (v5_relat_1(k1_funct_1(X1,X2),k4_numbers)|~v5_membered(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X4,X3)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X4,X3))))))).#trainneg
cnf(i_0_2886, plain, (v5_relat_1(k1_funct_1(X1,X2),k3_numbers)|~v4_membered(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X4,X3)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X4,X3))))))).#trainneg
cnf(i_0_2887, plain, (v1_partfun1(k1_funct_1(X1,X2),X3)|~v1_funct_2(k1_funct_1(X1,X2),X3,X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X3,X3)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X3,X3))))))).#trainneg
cnf(i_0_2888, plain, (r1_xxreal_0(X1,k1_seq_2(k1_funct_1(X2,X3)))|v2_xxreal_0(X1)|v3_xxreal_0(k1_seq_2(k1_funct_1(X2,X3)))|~v1_xreal_0(X1)|~v1_funct_2(k1_funct_1(X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(k1_funct_1(X2,X3))|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))).#trainneg
cnf(i_0_2889, plain, (r1_xxreal_0(X1,k2_seq_2(k1_funct_1(X2,X3)))|v2_xxreal_0(X1)|v3_xxreal_0(k2_seq_2(k1_funct_1(X2,X3)))|~v1_xreal_0(X1)|~v1_funct_2(k1_funct_1(X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(k1_funct_1(X2,X3))|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))).#trainneg
cnf(i_0_2890, plain, (v3_normsp_1(k1_funct_1(X1,X2),X3)|m1_subset_1(esk62_3(X3,k1_funct_1(X1,X2),X4),k5_numbers)|v2_struct_0(X3)|k2_seq_2(X4)!=np__0|~v2_comseq_2(X4)|~v1_funct_2(k1_funct_1(X1,X2),k5_numbers,u1_struct_0(X3))|~v1_funct_2(X4,k5_numbers,k1_numbers)|~v1_funct_1(X4)|~v1_funct_1(k1_funct_1(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X3)|~v5_lopban_2(X3)|~v3_vectsp_1(X3)|~v1_vectsp_1(X3)|~v3_group_1(X3)|~v2_funcsdom(X3)|~v2_normsp_1(X3)|~v4_normsp_0(X3)|~v3_normsp_0(X3)|~v8_rlvect_1(X3)|~v7_rlvect_1(X3)|~v6_rlvect_1(X3)|~v5_rlvect_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X3))))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X3)))))))).#trainneg
cnf(i_0_9706, plain, (v3_xxreal_0(esk39_0)|v1_xboole_0(esk39_0)|r1_xxreal_0(np__0,esk39_0)|~v1_xreal_0(esk39_0))).#trainneg
cnf(i_0_11792, negated_conjecture, (~v1_xboole_0(X1)|~v7_ordinal1(X2)|~r2_hidden(X3,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2))|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1))))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k5_numbers,u1_struct_0(esk1_0))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2892, plain, (v5_membered(k1_funct_1(X1,X2))|~v5_membered(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X3)))))).#trainneg
cnf(i_0_3391, plain, (v3_membered(esk52_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_3393, plain, (v4_membered(esk52_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_9124, negated_conjecture, (v3_ordinal1(esk2_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_9240, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1))))),k5_numbers))).#trainneg
cnf(i_0_2893, plain, (v6_ordinal1(k1_funct_1(X1,X2))|~v6_ordinal1(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X3)))))).#trainneg
cnf(i_0_2894, plain, (v6_membered(k1_funct_1(X1,X2))|~v6_membered(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X3)))))).#trainneg
cnf(i_0_2891, plain, (v3_normsp_1(k1_funct_1(X1,X2),X3)|m1_subset_1(esk62_3(X3,k1_funct_1(X1,X2),X4),k1_numbers)|v2_struct_0(X3)|k2_seq_2(X4)!=np__0|~v2_comseq_2(X4)|~v1_funct_2(k1_funct_1(X1,X2),k5_numbers,u1_struct_0(X3))|~v1_funct_2(X4,k5_numbers,k1_numbers)|~v1_funct_1(X4)|~v1_funct_1(k1_funct_1(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X3)|~v5_lopban_2(X3)|~v3_vectsp_1(X3)|~v1_vectsp_1(X3)|~v3_group_1(X3)|~v2_funcsdom(X3)|~v2_normsp_1(X3)|~v4_normsp_0(X3)|~v3_normsp_0(X3)|~v8_rlvect_1(X3)|~v7_rlvect_1(X3)|~v6_rlvect_1(X3)|~v5_rlvect_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v13_algstr_0(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X3))))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X3)))))))).#trainneg
cnf(i_0_2895, plain, (v1_membered(k1_funct_1(X1,X2))|~v1_membered(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X3)))))).#trainneg
cnf(i_0_2896, plain, (v2_membered(k1_funct_1(X1,X2))|~v2_membered(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X3)))))).#trainneg
cnf(i_0_2897, plain, (v3_membered(k1_funct_1(X1,X2))|~v3_membered(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X3)))))).#trainneg
cnf(i_0_2898, plain, (v4_membered(k1_funct_1(X1,X2))|~v4_membered(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X3)))))).#trainneg
cnf(i_0_2899, plain, (v1_valued_0(k1_funct_1(X1,X2))|~v1_valued_0(X3)|~v1_relat_1(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X3)))))).#trainneg
cnf(i_0_2900, plain, (v2_valued_0(k1_funct_1(X1,X2))|~v2_valued_0(X3)|~v1_relat_1(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X3)))))).#trainneg
cnf(i_0_2901, plain, (v3_valued_0(k1_funct_1(X1,X2))|~v3_valued_0(X3)|~v1_relat_1(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X3)))))).#trainneg
cnf(i_0_2902, plain, (v4_valued_0(k1_funct_1(X1,X2))|~v4_valued_0(X3)|~v1_relat_1(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X3)))))).#trainneg
cnf(i_0_2903, plain, (v5_relat_1(k1_funct_1(X1,X2),k4_numbers)|~v5_relat_1(X3,k4_numbers)|~v1_relat_1(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X3)))))).#trainneg
cnf(i_0_2904, plain, (v5_relat_1(k1_funct_1(X1,X2),k3_numbers)|~v5_relat_1(X3,k3_numbers)|~v1_relat_1(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X3)))))).#trainneg
cnf(i_0_2905, plain, (v1_xboole_0(k1_funct_1(X1,X2))|v1_xboole_0(X3)|m1_subset_1(X4,X3)|~m1_subset_1(X4,k1_funct_1(X1,X2))|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X3)))))).#trainneg
cnf(i_0_3900, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),esk59_0))).#trainneg
cnf(i_0_10564, negated_conjecture, (~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~l2_algstr_0(esk1_0)|~v3_membered(u1_struct_0(esk1_0))|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2908, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_membered(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_4881, negated_conjecture, (r1_xxreal_0(esk53_0,esk3_0)|r1_xxreal_0(esk3_0,esk53_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_3431, plain, (v5_membered(esk52_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3526, plain, (v1_membered(esk52_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_9241, negated_conjecture, (v3_ordinal1(esk3_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))),k5_numbers))).#trainneg
cnf(i_0_9242, negated_conjecture, (v3_ordinal1(esk2_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))),k5_numbers))).#trainneg
cnf(i_0_2909, plain, (v1_xxreal_0(k1_funct_1(X1,X2))|~v2_membered(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_2910, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v3_membered(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_2911, plain, (v1_rat_1(k1_funct_1(X1,X2))|~v4_membered(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_2912, plain, (v1_int_1(k1_funct_1(X1,X2))|~v5_membered(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_2913, plain, (v7_ordinal1(k1_funct_1(X1,X2))|~v6_membered(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_2914, plain, (v3_ordinal1(k1_funct_1(X1,X2))|~v3_ordinal1(X3)|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_2920, plain, (m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~r2_hidden(X1,k1_funct_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X2)))),X3))|~v1_funct_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X2)))),k5_numbers,k1_zfmisc_1(X2))|~v1_funct_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X2))))))).#trainneg
cnf(i_0_2921, plain, (m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~r2_hidden(X1,k1_funct_1(k1_funct_1(X4,X5),X3))|~v1_funct_2(k1_funct_1(X4,X5),k5_numbers,k1_zfmisc_1(X2))|~v1_funct_1(k1_funct_1(X4,X5))|~v7_ordinal1(X5)|~v1_funct_2(X4,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X2))))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X2)))))))).#trainneg
cnf(i_0_2922, plain, (m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~r2_hidden(X1,k1_funct_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X2))),X4,X5),X3))|~v1_funct_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X2))),X4,X5),k5_numbers,k1_zfmisc_1(X2))|~v1_funct_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X2))),X4,X5))|~v7_ordinal1(X5)|~v1_funct_2(X4,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X2))))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X2)))))))).#trainneg
cnf(i_0_2923, plain, (m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~r2_hidden(X1,k1_funct_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X2)),X3))|~v1_funct_2(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X2)),k5_numbers,k1_zfmisc_1(X2))|~v1_funct_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_2924, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4054, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),esk55_0))).#trainneg
cnf(i_0_12438, negated_conjecture, (~l2_algstr_0(esk1_0)|~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v3_membered(u1_struct_0(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2953, plain, (m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~r2_hidden(X1,k8_nat_1(X4,esk32_2(k5_numbers,k1_zfmisc_1(X2)),X3))|~v1_funct_2(esk32_2(k5_numbers,k1_zfmisc_1(X2)),k5_numbers,X4)|~m1_subset_1(esk32_2(k5_numbers,k1_zfmisc_1(X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4))))).#trainneg
cnf(i_0_4886, negated_conjecture, (r1_xxreal_0(esk53_0,esk2_0)|r1_xxreal_0(esk2_0,esk53_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_3528, plain, (v1_membered(esk52_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_10054, negated_conjecture, (v1_rat_1(esk4_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_11292, negated_conjecture, (v1_valued_0(k2_zfmisc_1(X1,esk4_0))|~v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2954, plain, (m1_subset_1(X1,X2)|v1_xboole_0(k1_funct_1(esk32_2(k5_numbers,k1_zfmisc_1(X2)),X3))|~v7_ordinal1(X3)|~m1_subset_1(X1,k1_funct_1(esk32_2(k5_numbers,k1_zfmisc_1(X2)),X3)))).#trainneg
cnf(i_0_2969, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X2),k5_numbers))).#trainneg
cnf(i_0_2927, plain, (v1_rat_1(X1)|~v4_membered(k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2928, plain, (v1_int_1(X1)|~v5_membered(k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2930, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2931, plain, (v1_xcmplx_0(esk6_2(X1,X2))|r1_xxreal_0(X2,np__0)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2932, plain, (v1_xxreal_0(esk6_2(X1,X2))|r1_xxreal_0(X2,np__0)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2933, plain, (v1_xreal_0(esk6_2(X1,X2))|r1_xxreal_0(X2,np__0)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2934, plain, (v1_rat_1(esk6_2(X1,X2))|r1_xxreal_0(X2,np__0)|~v4_membered(k5_numbers)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2935, plain, (v1_int_1(esk6_2(X1,X2))|r1_xxreal_0(X2,np__0)|~v5_membered(k5_numbers)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2936, plain, (v7_ordinal1(esk6_2(X1,X2))|r1_xxreal_0(X2,np__0)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_12737, plain, (r1_xxreal_0(X2,np__0)|~v3_xxreal_0(esk6_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2937, plain, (v3_ordinal1(esk6_2(X1,X2))|r1_xxreal_0(X2,np__0)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_9088, plain, (v1_valued_0(esk32_2(X1,np__0)))).#trainneg
cnf(i_0_4301, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),esk47_0))).#trainneg
cnf(i_0_12585, negated_conjecture, (~l2_algstr_0(esk1_0)|~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v3_membered(u1_struct_0(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2973, plain, (r1_xxreal_0(X1,k3_series_1(X2))|r1_xxreal_0(k3_series_1(X2),X1)|~v1_xxreal_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4907, plain, (r1_xxreal_0(esk53_0,k5_numbers)|r1_xxreal_0(k5_numbers,esk53_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_3529, plain, (v1_membered(esk52_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_3796, negated_conjecture, (v1_xcmplx_0(esk4_0)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_8080, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5))|~v1_xboole_0(X2)|~v1_xboole_0(X5))).#trainneg
cnf(i_0_8081, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5))|~v1_xboole_0(X2)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_2974, plain, (v2_xxreal_0(k3_series_1(X1))|v3_xxreal_0(k3_series_1(X1))|v1_xboole_0(k3_series_1(X1))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2976, plain, (r1_xxreal_0(k3_series_1(X1),k3_series_1(X1))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2942, plain, (v1_rat_1(esk6_2(X1,X2))|r1_xxreal_0(X2,np__0)|~v4_membered(k1_numbers)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2943, plain, (v1_int_1(esk6_2(X1,X2))|r1_xxreal_0(X2,np__0)|~v5_membered(k1_numbers)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2946, plain, (v1_xcmplx_0(esk9_3(X1,X2,X3))|r1_xxreal_0(X3,np__0)|X2!=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2947, plain, (v1_xxreal_0(esk9_3(X1,X2,X3))|r1_xxreal_0(X3,np__0)|X2!=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2948, plain, (v1_xreal_0(esk9_3(X1,X2,X3))|r1_xxreal_0(X3,np__0)|X2!=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2949, plain, (v1_rat_1(esk9_3(X1,X2,X3))|r1_xxreal_0(X3,np__0)|~v4_membered(k5_numbers)|X2!=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_9197, plain, (v3_valued_0(esk32_2(X1,np__0)))).#trainneg
cnf(i_0_4335, plain, (r1_xxreal_0(np__0,esk28_1(esk44_0))|v3_xxreal_0(esk28_1(esk44_0)))).#trainneg
cnf(i_0_12742, negated_conjecture, (~l2_algstr_0(esk1_0)|~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v3_membered(u1_struct_0(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_4963, plain, (r1_xxreal_0(esk48_0,esk44_0)|r1_xxreal_0(esk44_0,esk48_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_3943, plain, (v2_membered(esk52_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3945, plain, (v2_membered(esk52_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_9711, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5))|~v1_xboole_0(X1)|~v1_xboole_0(X5))).#trainneg
cnf(i_0_9712, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,k2_zfmisc_1(X4,X5))|~v1_xboole_0(X1)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_3012, plain, (v1_xboole_0(k3_series_1(X1))|m1_subset_1(esk28_1(k3_series_1(X1)),k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_3013, plain, (v1_xboole_0(k3_series_1(X1))|m1_subset_1(k1_funct_1(X2,X3),k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k3_series_1(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k3_series_1(X1)))))).#trainneg
cnf(i_0_12855, plain, (r1_xxreal_0(np__0,esk28_1(esk44_0))|~v1_xboole_0(esk28_1(esk44_0)))).#trainneg
cnf(i_0_2950, plain, (v1_int_1(esk9_3(X1,X2,X3))|r1_xxreal_0(X3,np__0)|~v5_membered(k5_numbers)|X2!=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2951, plain, (v7_ordinal1(esk9_3(X1,X2,X3))|r1_xxreal_0(X3,np__0)|X2!=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_12954, plain, (r1_xxreal_0(X3,np__0)|~v3_xxreal_0(esk9_3(X1,X2,X3))|X2!=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2952, plain, (v3_ordinal1(esk9_3(X1,X2,X3))|r1_xxreal_0(X3,np__0)|X2!=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_9228, plain, (v4_valued_0(esk32_2(X1,np__0)))).#trainneg
cnf(i_0_4079, plain, (r1_xxreal_0(np__0,esk38_0)|~m1_subset_1(esk38_0,k5_numbers))).#trainneg
cnf(i_0_12867, negated_conjecture, (~l2_algstr_0(esk1_0)|~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v3_membered(u1_struct_0(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3014, plain, (v1_xboole_0(k3_series_1(X1))|m1_subset_1(k8_nat_1(k3_series_1(X1),X2,X3),k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,k3_series_1(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k3_series_1(X1)))))).#trainneg
cnf(i_0_4972, negated_conjecture, (r1_xxreal_0(esk48_0,esk3_0)|r1_xxreal_0(esk3_0,esk48_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_9979, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(k2_zfmisc_1(X3,X4),X5)|~v1_xboole_0(X1)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_9980, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(k2_zfmisc_1(X3,X4),X5)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_3016, plain, (r1_xxreal_0(X1,np__0)|v2_struct_0(X2)|~r1_xxreal_0(X1,k18_complex1(k1_funct_1(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X4,X5,X6))),X6)))|~r1_xxreal_0(esk63_4(X2,X4,X5,X1),X6)|~v7_ordinal1(X6)|~v1_xreal_0(X1)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X4,X5,X6))),k5_numbers,k1_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X4,X5,X6))),k5_numbers,X3)|~v1_funct_1(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X4,X5,X6))))|~m1_subset_1(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X4,X5,X6))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X4,X5,X6))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_12971, negated_conjecture, (m1_subset_1(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2),k2_zfmisc_1(k5_numbers,k1_numbers))|~v7_ordinal1(X2)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k5_numbers,k3_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k3_series_1(X1))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k5_numbers,u1_struct_0(esk1_0))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2980, plain, (v1_rat_1(esk9_3(X1,X2,X3))|r1_xxreal_0(X3,np__0)|~v4_membered(k1_numbers)|X2!=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2981, plain, (v1_int_1(esk9_3(X1,X2,X3))|r1_xxreal_0(X3,np__0)|~v5_membered(k1_numbers)|X2!=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2984, plain, (v1_xcmplx_0(esk63_4(X1,X2,X3,X4))|r1_xxreal_0(X4,np__0)|v2_struct_0(X1)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2985, plain, (v1_xxreal_0(esk63_4(X1,X2,X3,X4))|r1_xxreal_0(X4,np__0)|v2_struct_0(X1)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_13050, plain, (r1_xxreal_0(X4,np__0)|v2_struct_0(X1)|~v2_xxreal_0(esk63_4(X1,X2,X3,X4))|~v3_xxreal_0(esk63_4(X1,X2,X3,X4))|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_13051, plain, (r1_xxreal_0(esk63_4(X1,X2,X3,X4),esk63_4(X1,X2,X3,X4))|r1_xxreal_0(X4,np__0)|v2_struct_0(X1)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2986, plain, (v1_xreal_0(esk63_4(X1,X2,X3,X4))|r1_xxreal_0(X4,np__0)|v2_struct_0(X1)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2987, plain, (v1_rat_1(esk63_4(X1,X2,X3,X4))|r1_xxreal_0(X4,np__0)|v2_struct_0(X1)|~v4_membered(k5_numbers)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2988, plain, (v1_int_1(esk63_4(X1,X2,X3,X4))|r1_xxreal_0(X4,np__0)|v2_struct_0(X1)|~v5_membered(k5_numbers)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_9834, plain, (v1_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(np__0,X3))))).#trainneg
cnf(i_0_9590, plain, (v1_xboole_0(esk32_2(X1,np__0)))).#trainneg
cnf(i_0_7591, negated_conjecture, (~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~m1_subset_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~l2_algstr_0(esk1_0)|~v3_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_3034, plain, (r1_xxreal_0(X1,X2)|v2_xxreal_0(X1)|v3_xxreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X2),k5_numbers))).#trainneg
cnf(i_0_4977, negated_conjecture, (r1_xxreal_0(esk48_0,esk2_0)|r1_xxreal_0(esk2_0,esk48_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_3946, plain, (v2_membered(esk52_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_3949, plain, (v3_membered(esk52_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_5765, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_5821, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2989, plain, (v7_ordinal1(esk63_4(X1,X2,X3,X4))|r1_xxreal_0(X4,np__0)|v2_struct_0(X1)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_13166, plain, (r1_xxreal_0(X4,np__0)|v2_struct_0(X1)|~v3_xxreal_0(esk63_4(X1,X2,X3,X4))|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2990, plain, (v3_ordinal1(esk63_4(X1,X2,X3,X4))|r1_xxreal_0(X4,np__0)|v2_struct_0(X1)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2994, plain, (v1_rat_1(esk63_4(X1,X2,X3,X4))|r1_xxreal_0(X4,np__0)|v2_struct_0(X1)|~v4_membered(k1_numbers)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2995, plain, (v1_int_1(esk63_4(X1,X2,X3,X4))|r1_xxreal_0(X4,np__0)|v2_struct_0(X1)|~v5_membered(k1_numbers)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2998, plain, (r1_xxreal_0(X1,esk8_3(X2,X3,X4))|v2_xxreal_0(X1)|v3_xxreal_0(esk8_3(X2,X3,X4))|v2_comseq_2(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_3019, plain, (r1_xxreal_0(k18_complex1(k8_nat_1(X1,k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),X5)),np__0)|v2_struct_0(X2)|~r1_xxreal_0(esk63_4(X2,X3,X4,k18_complex1(k8_nat_1(X1,k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),X5))),X5)|~v7_ordinal1(X5)|~v1_xreal_0(k18_complex1(k8_nat_1(X1,k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),X5)))|~m2_subset_1(X5,k1_numbers,k5_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),k5_numbers,k1_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),k5_numbers,X1)|~v1_funct_1(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))))|~m1_subset_1(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_13131, negated_conjecture, (~l2_algstr_0(esk1_0)|~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v3_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_4998, plain, (r1_xxreal_0(esk48_0,k5_numbers)|r1_xxreal_0(k5_numbers,esk48_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_4082, plain, (v4_membered(esk52_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_3043, plain, (k1_funct_1(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),X5)=k1_seq_1(esk61_3(X2,X3,X4),X5)|v2_struct_0(X2)|~v7_ordinal1(X5)|~v1_funct_2(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),k5_numbers,X1)|~v1_funct_1(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))))|~m1_subset_1(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m2_subset_1(X5,k1_numbers,k5_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),k5_numbers,k1_numbers)|~m1_subset_1(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_3044, plain, (m1_subset_1(k1_seq_1(esk61_3(X2,X3,X4),X5),X1)|v2_struct_0(X2)|~v7_ordinal1(X5)|~v1_funct_2(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),k5_numbers,X1)|~v1_funct_1(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))))|~m1_subset_1(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m2_subset_1(X5,k1_numbers,k5_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),k5_numbers,k1_numbers)|~m1_subset_1(k3_series_1(k4_normsp_0(X2,k8_lopban_4(X2,X3,X4,X5))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_3050, plain, (v3_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,X2),X3)))|v2_struct_0(X1)|~v3_xxreal_0(k8_nat_1(X4,k3_series_1(k4_normsp_0(X1,X2)),X3))|~v1_xreal_0(k8_nat_1(X4,k3_series_1(k4_normsp_0(X1,X2)),X3))|~v1_xreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,X2),X3)))|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(X1,X2)),k5_numbers,k1_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(X1,X2)),k5_numbers,X4)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(X2)|~m1_subset_1(k3_series_1(k4_normsp_0(X1,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_series_1(k4_normsp_0(X1,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3051, plain, (v2_xxreal_0(k8_nat_1(X1,k3_series_1(k4_normsp_0(X2,X3)),X4))|v2_struct_0(X2)|~v2_xxreal_0(k1_normsp_0(X2,k8_nat_1(u1_struct_0(X2),k1_bhsp_4(X2,X3),X4)))|~v1_xreal_0(k8_nat_1(X1,k3_series_1(k4_normsp_0(X2,X3)),X4))|~v1_xreal_0(k1_normsp_0(X2,k8_nat_1(u1_struct_0(X2),k1_bhsp_4(X2,X3),X4)))|~v7_ordinal1(X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(X2,X3)),k5_numbers,k1_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(X2,X3)),k5_numbers,X1)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(k3_series_1(k4_normsp_0(X2,X3)))|~v1_funct_1(X3)|~m1_subset_1(k3_series_1(k4_normsp_0(X2,X3)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_series_1(k4_normsp_0(X2,X3)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_10989, plain, (r1_xxreal_0(np__0,esk38_0)|~v1_xboole_0(k1_zfmisc_1(esk38_0)))).#trainneg
cnf(i_0_13202, negated_conjecture, (~l2_algstr_0(esk1_0)|~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v3_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k1_bhsp_4(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1)))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_3236, negated_conjecture, (v3_ordinal1(esk4_0)|~v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))))).#trainneg
cnf(i_0_5034, plain, (r1_xxreal_0(esk47_0,esk44_0)|r1_xxreal_0(esk44_0,esk47_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_8703, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3052, plain, (v2_xxreal_0(k8_nat_1(X1,k3_series_1(k4_normsp_0(X2,X3)),X4))|v3_xxreal_0(k1_normsp_0(X2,k8_nat_1(u1_struct_0(X2),k1_bhsp_4(X2,X3),X4)))|v1_xboole_0(k8_nat_1(X1,k3_series_1(k4_normsp_0(X2,X3)),X4))|v2_struct_0(X2)|~v1_xreal_0(k8_nat_1(X1,k3_series_1(k4_normsp_0(X2,X3)),X4))|~v1_xreal_0(k1_normsp_0(X2,k8_nat_1(u1_struct_0(X2),k1_bhsp_4(X2,X3),X4)))|~v7_ordinal1(X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(X2,X3)),k5_numbers,k1_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(X2,X3)),k5_numbers,X1)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(k3_series_1(k4_normsp_0(X2,X3)))|~v1_funct_1(X3)|~m1_subset_1(k3_series_1(k4_normsp_0(X2,X3)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_series_1(k4_normsp_0(X2,X3)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_3053, plain, (v2_xxreal_0(k8_nat_1(X1,k3_series_1(k4_normsp_0(X2,X3)),X4))|v3_xxreal_0(k1_normsp_0(X2,k8_nat_1(u1_struct_0(X2),k1_bhsp_4(X2,X3),X4)))|v1_xboole_0(k1_normsp_0(X2,k8_nat_1(u1_struct_0(X2),k1_bhsp_4(X2,X3),X4)))|v2_struct_0(X2)|~v1_xreal_0(k8_nat_1(X1,k3_series_1(k4_normsp_0(X2,X3)),X4))|~v1_xreal_0(k1_normsp_0(X2,k8_nat_1(u1_struct_0(X2),k1_bhsp_4(X2,X3),X4)))|~v7_ordinal1(X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(X2,X3)),k5_numbers,k1_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(X2,X3)),k5_numbers,X1)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(k3_series_1(k4_normsp_0(X2,X3)))|~v1_funct_1(X3)|~m1_subset_1(k3_series_1(k4_normsp_0(X2,X3)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_series_1(k4_normsp_0(X2,X3)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_3062, plain, (r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(u1_struct_0(X1),k1_bhsp_4(X1,X2),X3)),k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),X3))|v2_struct_0(X1)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(X1,X2)),k5_numbers,k1_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(X1,X2)),k5_numbers,X4)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(X2)|~m1_subset_1(k3_series_1(k4_normsp_0(X1,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_series_1(k4_normsp_0(X1,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3067, plain, (r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(X5,k1_bhsp_4(X1,X2),X3)),k8_nat_1(X4,k3_series_1(k4_normsp_0(X1,X2)),X3))|v2_struct_0(X1)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(X1,X2)),k5_numbers,k1_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(X1,X2)),k5_numbers,X4)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(X2)|~m1_subset_1(k3_series_1(k4_normsp_0(X1,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_series_1(k4_normsp_0(X1,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,X5)|~v1_funct_1(k1_bhsp_4(X1,X2))|~m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5))))).#trainneg
cnf(i_0_9836, plain, (v3_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(np__0,X3))))).#trainneg
cnf(i_0_9804, plain, (v3_valued_0(k2_zfmisc_1(np__0,X2)))).#trainneg
cnf(i_0_13239, plain, (v2_xxreal_0(esk38_0)|v1_xboole_0(esk38_0)|~v1_xboole_0(k1_zfmisc_1(esk38_0)))).#trainneg
cnf(i_0_13241, negated_conjecture, (~l2_algstr_0(esk1_0)|~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v3_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_3289, plain, (r1_xxreal_0(X1,k4_normsp_0(X2,X3))|v2_xxreal_0(X1)|v3_xxreal_0(k4_normsp_0(X2,X3))|v2_struct_0(X2)|~v1_xreal_0(X1)|~l1_normsp_0(X2)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))))).#trainneg
cnf(i_0_5043, negated_conjecture, (r1_xxreal_0(esk47_0,esk3_0)|r1_xxreal_0(esk3_0,esk47_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_5564, plain, (v4_membered(esk52_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3127, plain, (v3_xxreal_0(k1_normsp_0(X1,k8_nat_1(X2,k1_bhsp_4(X1,k8_lopban_4(X1,X3,X4,X5)),X5)))|v2_struct_0(X1)|~v3_xxreal_0(k1_seq_1(esk61_3(X1,X3,X4),X5))|~v1_xreal_0(k1_seq_1(esk61_3(X1,X3,X4),X5))|~v1_xreal_0(k1_normsp_0(X1,k8_nat_1(X2,k1_bhsp_4(X1,k8_lopban_4(X1,X3,X4,X5)),X5)))|~v7_ordinal1(X5)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(X1,k8_lopban_4(X1,X3,X4,X5)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k1_bhsp_4(X1,k8_lopban_4(X1,X3,X4,X5)),k5_numbers,X2)|~v1_funct_2(k8_lopban_4(X1,X3,X4,X5),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X3,X4,X5)))|~v1_funct_1(k8_lopban_4(X1,X3,X4,X5))|~m1_subset_1(k1_bhsp_4(X1,k8_lopban_4(X1,X3,X4,X5)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_bhsp_4(X1,k8_lopban_4(X1,X3,X4,X5)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(k8_lopban_4(X1,X3,X4,X5),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3128, plain, (v2_xxreal_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|v2_struct_0(X1)|~v2_xxreal_0(k1_normsp_0(X1,k8_nat_1(X5,k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|~v1_xreal_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|~v1_xreal_0(k1_normsp_0(X1,k8_nat_1(X5,k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|~v7_ordinal1(X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k5_numbers,X5)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3071, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(X1,X2,X3)),k5_numbers)|v2_struct_0(X1)|k2_seq_2(esk61_3(X1,X2,X3))!=np__0|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3073, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|m1_subset_1(esk62_3(esk1_0,esk4_0,esk32_2(k5_numbers,k1_numbers)),k5_numbers)|k2_seq_2(esk32_2(k5_numbers,k1_numbers))!=np__0|~v2_comseq_2(esk32_2(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_3074, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|m1_subset_1(esk62_3(esk1_0,esk4_0,esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))),k5_numbers)|k2_seq_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))!=np__0|~v2_comseq_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|~v1_funct_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))),k5_numbers,k1_numbers)|~v1_funct_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))).#trainneg
cnf(i_0_3075, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|m1_subset_1(esk62_3(esk1_0,esk4_0,k1_funct_1(X1,X2)),k5_numbers)|k2_seq_2(k1_funct_1(X1,X2))!=np__0|~v2_comseq_2(k1_funct_1(X1,X2))|~v1_funct_2(k1_funct_1(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_funct_1(X1,X2))|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))).#trainneg
cnf(i_0_3076, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|m1_subset_1(esk62_3(esk1_0,esk4_0,k4_normsp_0(X1,X2)),k5_numbers)|v2_struct_0(X1)|k2_seq_2(k4_normsp_0(X1,X2))!=np__0|~v2_comseq_2(k4_normsp_0(X1,X2))|~v1_funct_2(k4_normsp_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k4_normsp_0(X1,X2))|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_3077, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|m1_subset_1(esk62_3(esk1_0,esk4_0,k3_series_1(X1)),k5_numbers)|k2_seq_2(k3_series_1(X1))!=np__0|~v2_comseq_2(k3_series_1(X1))|~v1_funct_2(k3_series_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k3_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_3078, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|m1_subset_1(esk62_3(esk1_0,esk4_0,k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X1,X2)),k5_numbers)|k2_seq_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X1,X2))!=np__0|~v2_comseq_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X1,X2))|~v1_funct_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X1,X2))|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))).#trainneg
cnf(i_0_3079, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|m1_subset_1(esk62_3(esk1_0,esk4_0,k2_zfmisc_1(k5_numbers,k1_numbers)),k5_numbers)|k2_seq_2(k2_zfmisc_1(k5_numbers,k1_numbers))!=np__0|~v2_comseq_2(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(k2_zfmisc_1(k5_numbers,k1_numbers),k5_numbers,k1_numbers)|~v1_funct_1(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_3080, plain, (v3_normsp_1(k8_lopban_4(X1,X2,X3,X4),X1)|m1_subset_1(esk62_3(X1,k8_lopban_4(X1,X2,X3,X4),esk61_3(X5,X6,X7)),k5_numbers)|v2_struct_0(X1)|v2_struct_0(X5)|k2_seq_2(esk61_3(X5,X6,X7))!=np__0|~v2_comseq_2(esk61_3(X5,X6,X7))|~v1_funct_2(esk61_3(X5,X6,X7),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X5,X6,X7))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X7,u1_struct_0(X5))|~m1_subset_1(X6,u1_struct_0(X5))|~l1_lopban_2(X5)|~v5_lopban_2(X5)|~v3_vectsp_1(X5)|~v1_vectsp_1(X5)|~v3_group_1(X5)|~v2_funcsdom(X5)|~v2_normsp_1(X5)|~v4_normsp_0(X5)|~v3_normsp_0(X5)|~v8_rlvect_1(X5)|~v7_rlvect_1(X5)|~v6_rlvect_1(X5)|~v5_rlvect_1(X5)|~v4_rlvect_1(X5)|~v3_rlvect_1(X5)|~v2_rlvect_1(X5)|~v13_algstr_0(X5))).#trainneg
cnf(i_0_3082, plain, (v3_normsp_1(k8_lopban_4(X1,X2,X3,X4),X1)|m1_subset_1(esk62_3(X1,k8_lopban_4(X1,X2,X3,X4),esk32_2(k5_numbers,k1_numbers)),k5_numbers)|v2_struct_0(X1)|k2_seq_2(esk32_2(k5_numbers,k1_numbers))!=np__0|~v2_comseq_2(esk32_2(k5_numbers,k1_numbers))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_9805, plain, (v4_valued_0(k2_zfmisc_1(np__0,X2)))).#trainneg
cnf(i_0_13323, negated_conjecture, (~l2_algstr_0(esk1_0)|~v2_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v3_xxreal_0(k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1))|~v7_ordinal1(X1)|~v3_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3325, plain, (r1_xxreal_0(X1,k4_normsp_0(X2,X3))|r1_xxreal_0(k4_normsp_0(X2,X3),X1)|v2_struct_0(X2)|~v1_xxreal_0(X1)|~l1_normsp_0(X2)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))))).#trainneg
cnf(i_0_5048, negated_conjecture, (r1_xxreal_0(esk47_0,esk2_0)|r1_xxreal_0(esk2_0,esk47_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_6892, negated_conjecture, (v3_valued_0(esk4_0)|~v5_relat_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k1_numbers))).#trainneg
cnf(i_0_8684, plain, (r1_xxreal_0(X1,k2_zfmisc_1(X2,X3))|r1_xxreal_0(k2_zfmisc_1(X2,X3),X1)|~v1_xreal_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_8685, plain, (r1_xxreal_0(X1,k2_zfmisc_1(X2,X3))|r1_xxreal_0(k2_zfmisc_1(X2,X3),X1)|~v1_xreal_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3129, plain, (v2_xxreal_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|v3_xxreal_0(k1_normsp_0(X1,k8_nat_1(X5,k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|v1_xboole_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|v2_struct_0(X1)|~v1_xreal_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|~v1_xreal_0(k1_normsp_0(X1,k8_nat_1(X5,k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|~v7_ordinal1(X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k5_numbers,X5)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3130, plain, (v2_xxreal_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|v3_xxreal_0(k1_normsp_0(X1,k8_nat_1(X5,k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|v1_xboole_0(k1_normsp_0(X1,k8_nat_1(X5,k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|v2_struct_0(X1)|~v1_xreal_0(k1_seq_1(esk61_3(X1,X2,X3),X4))|~v1_xreal_0(k1_normsp_0(X1,k8_nat_1(X5,k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X4)))|~v7_ordinal1(X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k5_numbers,X5)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3083, plain, (v3_normsp_1(k8_lopban_4(X1,X2,X3,X4),X1)|m1_subset_1(esk62_3(X1,k8_lopban_4(X1,X2,X3,X4),esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))),k5_numbers)|v2_struct_0(X1)|k2_seq_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))!=np__0|~v2_comseq_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|~v1_funct_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))),k5_numbers,k1_numbers)|~v1_funct_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3084, plain, (v3_normsp_1(k8_lopban_4(X1,X2,X3,X4),X1)|m1_subset_1(esk62_3(X1,k8_lopban_4(X1,X2,X3,X4),k1_funct_1(X5,X6)),k5_numbers)|v2_struct_0(X1)|k2_seq_2(k1_funct_1(X5,X6))!=np__0|~v2_comseq_2(k1_funct_1(X5,X6))|~v1_funct_2(k1_funct_1(X5,X6),k5_numbers,k1_numbers)|~v1_funct_1(k1_funct_1(X5,X6))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v1_funct_2(X5,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))).#trainneg
cnf(i_0_3085, plain, (v3_normsp_1(k8_lopban_4(X1,X2,X3,X4),X1)|m1_subset_1(esk62_3(X1,k8_lopban_4(X1,X2,X3,X4),k4_normsp_0(X5,X6)),k5_numbers)|v2_struct_0(X1)|v2_struct_0(X5)|k2_seq_2(k4_normsp_0(X5,X6))!=np__0|~v2_comseq_2(k4_normsp_0(X5,X6))|~v1_funct_2(k4_normsp_0(X5,X6),k5_numbers,k1_numbers)|~v1_funct_1(k4_normsp_0(X5,X6))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~l1_normsp_0(X5)|~v1_funct_2(X6,k5_numbers,u1_struct_0(X5))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X5)))))).#trainneg
cnf(i_0_3086, plain, (v3_normsp_1(k8_lopban_4(X1,X2,X3,X4),X1)|m1_subset_1(esk62_3(X1,k8_lopban_4(X1,X2,X3,X4),k3_series_1(X5)),k5_numbers)|v2_struct_0(X1)|k2_seq_2(k3_series_1(X5))!=np__0|~v2_comseq_2(k3_series_1(X5))|~v1_funct_2(k3_series_1(X5),k5_numbers,k1_numbers)|~v1_funct_1(k3_series_1(X5))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_funct_2(X5,k5_numbers,k1_numbers)|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_3087, plain, (v3_normsp_1(k8_lopban_4(X1,X2,X3,X4),X1)|m1_subset_1(esk62_3(X1,k8_lopban_4(X1,X2,X3,X4),k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X5,X6)),k5_numbers)|v2_struct_0(X1)|k2_seq_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X5,X6))!=np__0|~v2_comseq_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X5,X6))|~v1_funct_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X5,X6),k5_numbers,k1_numbers)|~v1_funct_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X5,X6))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v1_funct_2(X5,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))).#trainneg
cnf(i_0_3088, plain, (v3_normsp_1(k8_lopban_4(X1,X2,X3,X4),X1)|m1_subset_1(esk62_3(X1,k8_lopban_4(X1,X2,X3,X4),k2_zfmisc_1(k5_numbers,k1_numbers)),k5_numbers)|v2_struct_0(X1)|k2_seq_2(k2_zfmisc_1(k5_numbers,k1_numbers))!=np__0|~v2_comseq_2(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(k2_zfmisc_1(k5_numbers,k1_numbers),k5_numbers,k1_numbers)|~v1_funct_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3089, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|m1_subset_1(esk62_3(esk1_0,esk4_0,esk61_3(X1,X2,X3)),k1_numbers)|v2_struct_0(X1)|k2_seq_2(esk61_3(X1,X2,X3))!=np__0|~v2_comseq_2(esk61_3(X1,X2,X3))|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3091, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|m1_subset_1(esk62_3(esk1_0,esk4_0,esk32_2(k5_numbers,k1_numbers)),k1_numbers)|k2_seq_2(esk32_2(k5_numbers,k1_numbers))!=np__0|~v2_comseq_2(esk32_2(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_3092, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|m1_subset_1(esk62_3(esk1_0,esk4_0,esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))),k1_numbers)|k2_seq_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))!=np__0|~v2_comseq_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|~v1_funct_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))),k5_numbers,k1_numbers)|~v1_funct_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))).#trainneg
cnf(i_0_3093, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|m1_subset_1(esk62_3(esk1_0,esk4_0,k1_funct_1(X1,X2)),k1_numbers)|k2_seq_2(k1_funct_1(X1,X2))!=np__0|~v2_comseq_2(k1_funct_1(X1,X2))|~v1_funct_2(k1_funct_1(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_funct_1(X1,X2))|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))).#trainneg
cnf(i_0_3094, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|m1_subset_1(esk62_3(esk1_0,esk4_0,k4_normsp_0(X1,X2)),k1_numbers)|v2_struct_0(X1)|k2_seq_2(k4_normsp_0(X1,X2))!=np__0|~v2_comseq_2(k4_normsp_0(X1,X2))|~v1_funct_2(k4_normsp_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k4_normsp_0(X1,X2))|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_3095, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|m1_subset_1(esk62_3(esk1_0,esk4_0,k3_series_1(X1)),k1_numbers)|k2_seq_2(k3_series_1(X1))!=np__0|~v2_comseq_2(k3_series_1(X1))|~v1_funct_2(k3_series_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k3_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_9810, plain, (v1_valued_0(k2_zfmisc_1(np__0,X2)))).#trainneg
cnf(i_0_7665, plain, (v7_ordinal1(esk52_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)),k5_numbers))).#trainneg
cnf(i_0_8262, plain, (~v3_xxreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(k6_xcmplx_0(k4_xcmplx_0(np__1),X1),k5_numbers)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3326, plain, (v2_xxreal_0(k4_normsp_0(X1,X2))|v3_xxreal_0(k4_normsp_0(X1,X2))|v1_xboole_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_5069, plain, (r1_xxreal_0(esk47_0,k5_numbers)|r1_xxreal_0(k5_numbers,esk47_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_3096, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|m1_subset_1(esk62_3(esk1_0,esk4_0,k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X1,X2)),k1_numbers)|k2_seq_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X1,X2))!=np__0|~v2_comseq_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X1,X2))|~v1_funct_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X1,X2))|~v7_ordinal1(X2)|~v1_funct_2(X1,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))).#trainneg
cnf(i_0_3097, negated_conjecture, (v3_normsp_1(esk4_0,esk1_0)|m1_subset_1(esk62_3(esk1_0,esk4_0,k2_zfmisc_1(k5_numbers,k1_numbers)),k1_numbers)|k2_seq_2(k2_zfmisc_1(k5_numbers,k1_numbers))!=np__0|~v2_comseq_2(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(k2_zfmisc_1(k5_numbers,k1_numbers),k5_numbers,k1_numbers)|~v1_funct_1(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_3098, plain, (v3_normsp_1(k8_lopban_4(X1,X2,X3,X4),X1)|m1_subset_1(esk62_3(X1,k8_lopban_4(X1,X2,X3,X4),esk61_3(X5,X6,X7)),k1_numbers)|v2_struct_0(X1)|v2_struct_0(X5)|k2_seq_2(esk61_3(X5,X6,X7))!=np__0|~v2_comseq_2(esk61_3(X5,X6,X7))|~v1_funct_2(esk61_3(X5,X6,X7),k5_numbers,k1_numbers)|~v1_funct_1(esk61_3(X5,X6,X7))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X7,u1_struct_0(X5))|~m1_subset_1(X6,u1_struct_0(X5))|~l1_lopban_2(X5)|~v5_lopban_2(X5)|~v3_vectsp_1(X5)|~v1_vectsp_1(X5)|~v3_group_1(X5)|~v2_funcsdom(X5)|~v2_normsp_1(X5)|~v4_normsp_0(X5)|~v3_normsp_0(X5)|~v8_rlvect_1(X5)|~v7_rlvect_1(X5)|~v6_rlvect_1(X5)|~v5_rlvect_1(X5)|~v4_rlvect_1(X5)|~v3_rlvect_1(X5)|~v2_rlvect_1(X5)|~v13_algstr_0(X5))).#trainneg
cnf(i_0_3100, plain, (v3_normsp_1(k8_lopban_4(X1,X2,X3,X4),X1)|m1_subset_1(esk62_3(X1,k8_lopban_4(X1,X2,X3,X4),esk32_2(k5_numbers,k1_numbers)),k1_numbers)|v2_struct_0(X1)|k2_seq_2(esk32_2(k5_numbers,k1_numbers))!=np__0|~v2_comseq_2(esk32_2(k5_numbers,k1_numbers))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3101, plain, (v3_normsp_1(k8_lopban_4(X1,X2,X3,X4),X1)|m1_subset_1(esk62_3(X1,k8_lopban_4(X1,X2,X3,X4),esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))),k1_numbers)|v2_struct_0(X1)|k2_seq_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))!=np__0|~v2_comseq_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|~v1_funct_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))),k5_numbers,k1_numbers)|~v1_funct_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3102, plain, (v3_normsp_1(k8_lopban_4(X1,X2,X3,X4),X1)|m1_subset_1(esk62_3(X1,k8_lopban_4(X1,X2,X3,X4),k1_funct_1(X5,X6)),k1_numbers)|v2_struct_0(X1)|k2_seq_2(k1_funct_1(X5,X6))!=np__0|~v2_comseq_2(k1_funct_1(X5,X6))|~v1_funct_2(k1_funct_1(X5,X6),k5_numbers,k1_numbers)|~v1_funct_1(k1_funct_1(X5,X6))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v1_funct_2(X5,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))).#trainneg
cnf(i_0_3103, plain, (v3_normsp_1(k8_lopban_4(X1,X2,X3,X4),X1)|m1_subset_1(esk62_3(X1,k8_lopban_4(X1,X2,X3,X4),k4_normsp_0(X5,X6)),k1_numbers)|v2_struct_0(X1)|v2_struct_0(X5)|k2_seq_2(k4_normsp_0(X5,X6))!=np__0|~v2_comseq_2(k4_normsp_0(X5,X6))|~v1_funct_2(k4_normsp_0(X5,X6),k5_numbers,k1_numbers)|~v1_funct_1(k4_normsp_0(X5,X6))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~l1_normsp_0(X5)|~v1_funct_2(X6,k5_numbers,u1_struct_0(X5))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X5)))))).#trainneg
cnf(i_0_3104, plain, (v3_normsp_1(k8_lopban_4(X1,X2,X3,X4),X1)|m1_subset_1(esk62_3(X1,k8_lopban_4(X1,X2,X3,X4),k3_series_1(X5)),k1_numbers)|v2_struct_0(X1)|k2_seq_2(k3_series_1(X5))!=np__0|~v2_comseq_2(k3_series_1(X5))|~v1_funct_2(k3_series_1(X5),k5_numbers,k1_numbers)|~v1_funct_1(k3_series_1(X5))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_funct_2(X5,k5_numbers,k1_numbers)|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_3105, plain, (v3_normsp_1(k8_lopban_4(X1,X2,X3,X4),X1)|m1_subset_1(esk62_3(X1,k8_lopban_4(X1,X2,X3,X4),k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X5,X6)),k1_numbers)|v2_struct_0(X1)|k2_seq_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X5,X6))!=np__0|~v2_comseq_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X5,X6))|~v1_funct_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X5,X6),k5_numbers,k1_numbers)|~v1_funct_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X5,X6))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X6)|~v1_funct_2(X5,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))).#trainneg
cnf(i_0_3106, plain, (v3_normsp_1(k8_lopban_4(X1,X2,X3,X4),X1)|m1_subset_1(esk62_3(X1,k8_lopban_4(X1,X2,X3,X4),k2_zfmisc_1(k5_numbers,k1_numbers)),k1_numbers)|v2_struct_0(X1)|k2_seq_2(k2_zfmisc_1(k5_numbers,k1_numbers))!=np__0|~v2_comseq_2(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(k2_zfmisc_1(k5_numbers,k1_numbers),k5_numbers,k1_numbers)|~v1_funct_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3107, plain, (r1_xxreal_0(X1,esk62_3(X2,X3,X4))|r1_xxreal_0(esk62_3(X2,X3,X4),X1)|k6_normsp_1(X2,X3)=k4_struct_0(X2)|v2_struct_0(X2)|~v1_xxreal_0(X1)|k2_seq_2(X4)!=np__0|~v2_comseq_2(X4)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X4,k5_numbers,k1_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_3108, plain, (v2_xxreal_0(esk62_3(X1,X2,X3))|v3_xxreal_0(esk62_3(X1,X2,X3))|v1_xboole_0(esk62_3(X1,X2,X3))|k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_9837, plain, (v4_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(np__0,X3))))).#trainneg
cnf(i_0_7669, plain, (v3_ordinal1(esk52_0)|~v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)))))).#trainneg
cnf(i_0_13724, plain, (~v1_xcmplx_0(k4_xcmplx_0(np__1))|~m1_subset_1(k6_xcmplx_0(k4_xcmplx_0(np__1),np__1),k5_numbers))).#trainneg
cnf(i_0_3328, plain, (r1_xxreal_0(k4_normsp_0(X1,X2),k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_5160, plain, (r1_xxreal_0(esk39_0,esk44_0)|r1_xxreal_0(esk44_0,esk39_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_13715, plain, (~v3_xxreal_0(esk52_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)),k5_numbers))).#trainneg
cnf(i_0_8747, plain, (r1_xxreal_0(X1,k2_zfmisc_1(X2,X3))|r1_xxreal_0(k2_zfmisc_1(X2,X3),X1)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_8748, plain, (r1_xxreal_0(X1,k2_zfmisc_1(X2,X3))|r1_xxreal_0(k2_zfmisc_1(X2,X3),X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3137, plain, (v3_xxreal_0(k1_normsp_0(X1,k8_nat_1(X2,k1_bhsp_4(X1,X3),X4)))|v2_struct_0(X1)|~v3_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X3)),X4))|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X3)),X4))|~v1_xreal_0(k1_normsp_0(X1,k8_nat_1(X2,k1_bhsp_4(X1,X3),X4)))|~v7_ordinal1(X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(X1,X3),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k1_bhsp_4(X1,X3),k5_numbers,X2)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,X3))|~v1_funct_1(X3)|~m1_subset_1(k1_bhsp_4(X1,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_bhsp_4(X1,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3138, plain, (v2_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|v2_struct_0(X1)|~v2_xxreal_0(k1_normsp_0(X1,k8_nat_1(X4,k1_bhsp_4(X1,X2),X3)))|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|~v1_xreal_0(k1_normsp_0(X1,k8_nat_1(X4,k1_bhsp_4(X1,X2),X3)))|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,X4)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3139, plain, (v2_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|v3_xxreal_0(k1_normsp_0(X1,k8_nat_1(X4,k1_bhsp_4(X1,X2),X3)))|v1_xboole_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|v2_struct_0(X1)|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|~v1_xreal_0(k1_normsp_0(X1,k8_nat_1(X4,k1_bhsp_4(X1,X2),X3)))|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,X4)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3140, plain, (v2_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|v3_xxreal_0(k1_normsp_0(X1,k8_nat_1(X4,k1_bhsp_4(X1,X2),X3)))|v1_xboole_0(k1_normsp_0(X1,k8_nat_1(X4,k1_bhsp_4(X1,X2),X3)))|v2_struct_0(X1)|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(X1,X2)),X3))|~v1_xreal_0(k1_normsp_0(X1,k8_nat_1(X4,k1_bhsp_4(X1,X2),X3)))|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,X4)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(k1_bhsp_4(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_9817, plain, (v1_xboole_0(k2_zfmisc_1(np__0,X2)))).#trainneg
cnf(i_0_10966, plain, (v3_ordinal1(esk52_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)),k5_numbers))).#trainneg
cnf(i_0_5169, negated_conjecture, (r1_xxreal_0(esk39_0,esk3_0)|r1_xxreal_0(esk3_0,esk39_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_3821, negated_conjecture, (v3_ordinal1(esk4_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_4630, negated_conjecture, (v3_ordinal1(esk4_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))),k5_numbers))).#trainneg
cnf(i_0_3178, plain, (v3_xxreal_0(X1)|~v3_xxreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(k6_xcmplx_0(X1,X2))|~v3_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_9824, plain, (v1_xreal_0(k2_zfmisc_1(np__0,X2)))).#trainneg
cnf(i_0_5174, negated_conjecture, (r1_xxreal_0(esk39_0,esk2_0)|r1_xxreal_0(esk2_0,esk39_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_11293, negated_conjecture, (v1_valued_0(k2_zfmisc_1(X1,esk4_0))|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_12126, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1))))),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3179, plain, (v2_xxreal_0(X2)|~v3_xxreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(k6_xcmplx_0(X1,X2))|~v2_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3183, negated_conjecture, (v3_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v3_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_3184, negated_conjecture, (v2_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v2_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_3185, negated_conjecture, (v2_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|v3_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|v1_xboole_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_3186, negated_conjecture, (v2_xxreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|v3_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|v1_xboole_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v1_xreal_0(k8_nat_1(k1_numbers,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_3187, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk4_0,X1)),k8_nat_1(X2,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),k5_numbers,k1_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),k5_numbers,X2)|~v1_funct_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_3198, negated_conjecture, (v3_normsp_1(k1_bhsp_4(esk1_0,esk4_0),esk1_0)|m1_subset_1(esk62_3(esk1_0,k1_bhsp_4(esk1_0,esk4_0),X1),k5_numbers)|k2_seq_2(X1)!=np__0|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(esk1_0,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_3199, plain, (v3_normsp_1(k1_bhsp_4(X1,k1_bhsp_4(X1,X2)),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,k1_bhsp_4(X1,X2)),X3),k5_numbers)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(X1,k1_bhsp_4(X1,X2)))|~v1_funct_1(k1_bhsp_4(X1,X2))|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_9838, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(np__0,X3))))).#trainneg
cnf(i_0_13725, plain, (~v3_xxreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1)|~m1_subset_1(k6_xcmplx_0(k4_xcmplx_0(X1),np__1),k5_numbers))).#trainneg
cnf(i_0_5195, plain, (r1_xxreal_0(esk39_0,k5_numbers)|r1_xxreal_0(k5_numbers,esk39_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_12304, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1)))))),k5_numbers))).#trainneg
cnf(i_0_5326, negated_conjecture, (r1_xxreal_0(esk3_0,k5_numbers)|r1_xxreal_0(k5_numbers,esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3200, plain, (v3_normsp_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X5),k5_numbers)|v2_struct_0(X1)|k2_seq_2(X5)!=np__0|~v2_comseq_2(X5)|~v1_funct_2(X5,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~l2_algstr_0(X1)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_3201, negated_conjecture, (v3_normsp_1(k1_bhsp_4(esk1_0,esk4_0),esk1_0)|m1_subset_1(esk62_3(esk1_0,k1_bhsp_4(esk1_0,esk4_0),X1),k1_numbers)|k2_seq_2(X1)!=np__0|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(esk1_0,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_3202, plain, (v3_normsp_1(k1_bhsp_4(X1,k1_bhsp_4(X1,X2)),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,k1_bhsp_4(X1,X2)),X3),k1_numbers)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~v2_comseq_2(X3)|~v1_funct_2(k1_bhsp_4(X1,X2),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(X1,k1_bhsp_4(X1,X2)))|~v1_funct_1(k1_bhsp_4(X1,X2))|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_bhsp_4(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_3203, plain, (v3_normsp_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)),X5),k1_numbers)|v2_struct_0(X1)|k2_seq_2(X5)!=np__0|~v2_comseq_2(X5)|~v1_funct_2(X5,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(X1,k8_lopban_4(X1,X2,X3,X4)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,k5_numbers)|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~l2_algstr_0(X1)|~v1_funct_2(k8_lopban_4(X1,X2,X3,X4),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k8_lopban_4(X1,X2,X3,X4))|~m1_subset_1(k8_lopban_4(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_3213, plain, (v3_normsp_1(k1_bhsp_4(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2),k5_numbers)|v2_struct_0(X1)|k2_seq_2(X2)!=np__0|~v2_comseq_2(X2)|~v1_funct_2(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_3214, plain, (v3_normsp_1(k1_bhsp_4(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2),k1_numbers)|v2_struct_0(X1)|k2_seq_2(X2)!=np__0|~v2_comseq_2(X2)|~v1_funct_2(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_3215, negated_conjecture, (v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v3_xxreal_0(k8_nat_1(X2,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v1_xreal_0(k8_nat_1(X2,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),k5_numbers,k1_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),k5_numbers,X2)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_3216, negated_conjecture, (v2_xxreal_0(k8_nat_1(X1,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),X2))|~v2_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2)))|~v1_xreal_0(k8_nat_1(X1,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),X2))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2)))|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k5_numbers,k1_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k5_numbers,X1)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_3217, negated_conjecture, (v2_xxreal_0(k8_nat_1(X1,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),X2))|v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2)))|v1_xboole_0(k8_nat_1(X1,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),X2))|~v1_xreal_0(k8_nat_1(X1,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),X2))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2)))|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k5_numbers,k1_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k5_numbers,X1)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_3218, negated_conjecture, (v2_xxreal_0(k8_nat_1(X1,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),X2))|v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2)))|v1_xboole_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2)))|~v1_xreal_0(k8_nat_1(X1,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),X2))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X2)))|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k5_numbers,k1_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k5_numbers,X1)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_3220, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(u1_struct_0(esk1_0),esk4_0,X1)),k1_funct_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),k5_numbers,k1_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),k5_numbers,X2)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X1))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_6893, negated_conjecture, (v3_valued_0(esk4_0)|~v5_relat_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k4_numbers))).#trainneg
cnf(i_0_5414, negated_conjecture, (r1_xxreal_0(esk2_0,k5_numbers)|r1_xxreal_0(k5_numbers,esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3243, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X1,k1_funct_1(X2,esk62_3(X1,X2,X3))),k8_nat_1(X4,X3,esk62_3(X1,X2,X3)))|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,X4)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk62_3(X1,X2,X3)))).#trainneg
cnf(i_0_6894, negated_conjecture, (v3_valued_0(esk4_0)|~v5_relat_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k3_numbers))).#trainneg
cnf(i_0_3252, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X1,k1_funct_1(X2,esk62_3(X1,X2,X3))),k1_funct_1(X3,esk62_3(X1,X2,X3)))|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,X4)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk62_3(X1,X2,X3)))).#trainneg
cnf(i_0_3253, plain, (k6_normsp_1(X1,X2)=k4_struct_0(X1)|v2_struct_0(X1)|k2_seq_2(X3)!=np__0|~r1_xxreal_0(k1_normsp_0(X1,k8_nat_1(X5,X2,esk62_3(X1,X2,X3))),k1_funct_1(X3,esk62_3(X1,X2,X3)))|~v2_comseq_2(X3)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,X4)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk62_3(X1,X2,X3))|~v1_funct_2(X2,k5_numbers,X5)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5))))).#trainneg
cnf(i_0_6896, negated_conjecture, (v3_valued_0(esk4_0)|~v5_relat_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_6968, negated_conjecture, (v4_valued_0(esk4_0)|~v5_relat_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_6165, plain, (r1_xxreal_0(X1,k2_zfmisc_1(X2,X3))|v2_xxreal_0(X1)|v3_xxreal_0(k2_zfmisc_1(X2,X3))|~v1_xreal_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_6220, plain, (r1_xxreal_0(X1,k2_zfmisc_1(X2,X3))|v2_xxreal_0(X1)|v3_xxreal_0(k2_zfmisc_1(X2,X3))|~v1_xreal_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_11676, plain, (v1_xcmplx_0(esk52_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)),k5_numbers))).#trainneg
cnf(i_0_10707, negated_conjecture, (v1_valued_0(esk4_0)|~v5_relat_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k1_numbers))).#trainneg
cnf(i_0_10708, negated_conjecture, (v1_valued_0(esk4_0)|~v5_relat_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k4_numbers))).#trainneg
cnf(i_0_5877, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3309, negated_conjecture, (v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)))|~v3_xxreal_0(k8_nat_1(X3,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),X2))|~v1_xreal_0(k8_nat_1(X3,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),X2))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)))|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k5_numbers,k1_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k5_numbers,X3)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X1)|~v1_funct_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_3310, negated_conjecture, (v2_xxreal_0(k8_nat_1(X1,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),X2))|~v2_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X3,esk4_0,X2)))|~v1_xreal_0(k8_nat_1(X1,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),X2))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X3,esk4_0,X2)))|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k5_numbers,k1_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k5_numbers,X1)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X3)|~v1_funct_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_3311, negated_conjecture, (v2_xxreal_0(k8_nat_1(X1,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),X2))|v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X3,esk4_0,X2)))|v1_xboole_0(k8_nat_1(X1,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),X2))|~v1_xreal_0(k8_nat_1(X1,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),X2))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X3,esk4_0,X2)))|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k5_numbers,k1_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k5_numbers,X1)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X3)|~v1_funct_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_3312, negated_conjecture, (v2_xxreal_0(k8_nat_1(X1,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),X2))|v3_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X3,esk4_0,X2)))|v1_xboole_0(k1_normsp_0(esk1_0,k8_nat_1(X3,esk4_0,X2)))|~v1_xreal_0(k8_nat_1(X1,k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),X2))|~v1_xreal_0(k1_normsp_0(esk1_0,k8_nat_1(X3,esk4_0,X2)))|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k5_numbers,k1_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k5_numbers,X1)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X3)|~v1_funct_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).#trainneg
cnf(i_0_3314, negated_conjecture, (r1_xxreal_0(k1_normsp_0(esk1_0,k8_nat_1(X1,esk4_0,X2)),k1_funct_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),X2))|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k5_numbers,k1_numbers)|~v1_funct_2(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k5_numbers,X3)|~v1_funct_2(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k5_numbers,u1_struct_0(esk1_0))|~v1_funct_2(esk4_0,k5_numbers,X1)|~v1_funct_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))))|~v1_funct_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_series_1(k4_normsp_0(esk1_0,k8_lopban_4(esk1_0,esk2_0,esk3_0,X2))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(k8_lopban_4(esk1_0,esk2_0,esk3_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_10709, negated_conjecture, (v1_valued_0(esk4_0)|~v5_relat_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k3_numbers))).#trainneg
cnf(i_0_3377, plain, (v1_xboole_0(k4_normsp_0(X1,X2))|m1_subset_1(esk28_1(k4_normsp_0(X1,X2)),k2_zfmisc_1(k5_numbers,k1_numbers))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_10711, negated_conjecture, (v1_valued_0(esk4_0)|~v5_relat_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_10818, negated_conjecture, (v2_valued_0(esk4_0)|~v5_relat_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k1_numbers))).#trainneg
cnf(i_0_3378, plain, (v1_xboole_0(k4_normsp_0(X1,X2))|m1_subset_1(k1_funct_1(X3,X4),k2_zfmisc_1(k5_numbers,k1_numbers))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k4_normsp_0(X1,X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_normsp_0(X1,X2)))))).#trainneg
cnf(i_0_3379, plain, (v1_xboole_0(k4_normsp_0(X1,X2))|m1_subset_1(k8_nat_1(k4_normsp_0(X1,X2),X3,X4),k2_zfmisc_1(k5_numbers,k1_numbers))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k4_normsp_0(X1,X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_normsp_0(X1,X2)))))).#trainneg
cnf(i_0_3386, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_3387, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_3394, plain, (r1_xxreal_0(X1,k5_numbers)|r1_xxreal_0(k5_numbers,X1)|~v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3395, plain, (v2_xxreal_0(k5_numbers)|v3_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3401, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_3402, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers))).#trainneg
cnf(i_0_11774, plain, (v1_xxreal_0(esk52_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)),k5_numbers))).#trainneg
cnf(i_0_3403, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(k4_xcmplx_0(X1)),k5_numbers))).#trainneg
cnf(i_0_3414, plain, (r1_xxreal_0(X1,esk28_1(X2))|v2_xxreal_0(X1)|v3_xxreal_0(esk28_1(X2))|~v1_xreal_0(X1)|~v3_membered(X2))).#trainneg
cnf(i_0_3416, plain, (v1_xreal_0(esk28_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_3409, plain, (v3_normsp_1(k1_bhsp_4(X1,esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))),X2),k5_numbers)|v2_struct_0(X1)|k2_seq_2(X2)!=np__0|~v2_comseq_2(X2)|~v1_funct_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(X1,esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))))|~v1_funct_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_3410, plain, (v3_normsp_1(k1_bhsp_4(X1,esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))),X1)|m1_subset_1(esk62_3(X1,k1_bhsp_4(X1,esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))),X2),k1_numbers)|v2_struct_0(X1)|k2_seq_2(X2)!=np__0|~v2_comseq_2(X2)|~v1_funct_2(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(k1_bhsp_4(X1,esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))))|~v1_funct_1(esk28_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_3417, plain, (~v3_xxreal_0(esk28_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_3418, plain, (v1_xboole_0(esk28_1(k1_zfmisc_1(X1)))|m1_subset_1(esk28_1(esk28_1(k1_zfmisc_1(X1))),X1))).#trainneg
cnf(i_0_3419, plain, (v1_xboole_0(esk28_1(k1_zfmisc_1(X1)))|m1_subset_1(k1_funct_1(X2,X3),X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,esk28_1(k1_zfmisc_1(X1)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,esk28_1(k1_zfmisc_1(X1))))))).#trainneg
cnf(i_0_3420, plain, (v1_xboole_0(esk28_1(k1_zfmisc_1(X1)))|m1_subset_1(k8_nat_1(esk28_1(k1_zfmisc_1(X1)),X2,X3),X1)|~v7_ordinal1(X3)|~v1_funct_2(X2,k5_numbers,esk28_1(k1_zfmisc_1(X1)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,esk28_1(k1_zfmisc_1(X1))))))).#trainneg
cnf(i_0_3425, plain, (r1_xxreal_0(X1,esk61_3(X2,X3,X4))|r1_xxreal_0(esk61_3(X2,X3,X4),X1)|v2_struct_0(X2)|~v1_xxreal_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_3426, plain, (v2_xxreal_0(esk61_3(X1,X2,X3))|v3_xxreal_0(esk61_3(X1,X2,X3))|v1_xboole_0(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3428, plain, (r1_xxreal_0(esk61_3(X1,X2,X3),esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3438, plain, (r1_xxreal_0(X1,esk61_3(X2,X3,X4))|v2_xxreal_0(X1)|v3_xxreal_0(esk61_3(X2,X3,X4))|v2_struct_0(X2)|~v1_xreal_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_lopban_2(X2)|~v5_lopban_2(X2)|~v3_vectsp_1(X2)|~v1_vectsp_1(X2)|~v3_group_1(X2)|~v2_funcsdom(X2)|~v2_normsp_1(X2)|~v4_normsp_0(X2)|~v3_normsp_0(X2)|~v8_rlvect_1(X2)|~v7_rlvect_1(X2)|~v6_rlvect_1(X2)|~v5_rlvect_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_5310, negated_conjecture, (r1_xxreal_0(esk3_0,np__0)|r1_xxreal_0(np__0,esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_3451, plain, (r1_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0),np__0)|v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~v1_xcmplx_0(k8_nat_1(X4,esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)))))|~m2_subset_1(esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0))),k1_numbers,k5_numbers)|~m1_subset_1(k8_nat_1(X4,esk61_3(X1,X2,X3),esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0)))),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk8_3(esk61_3(X1,X2,X3),np__0,esk63_4(X1,X2,X3,esk7_2(esk61_3(X1,X2,X3),np__0))))|~v1_funct_2(esk61_3(X1,X2,X3),k5_numbers,X4)|~v1_funct_1(esk61_3(X1,X2,X3))|~m1_subset_1(esk61_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4))))).#trainneg
cnf(i_0_15133, plain, (v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~v2_xxreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~v1_xreal_0(esk7_2(esk61_3(X1,X2,X3),np__0))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_15151, plain, (v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(esk7_2(esk61_3(X1,X2,X3),np__0),k1_numbers))).#trainneg
cnf(i_0_15152, plain, (v2_comseq_2(esk61_3(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lopban_2(X1)|~v5_lopban_2(X1)|~v3_vectsp_1(X1)|~v1_vectsp_1(X1)|~v3_group_1(X1)|~v2_funcsdom(X1)|~v2_normsp_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v8_rlvect_1(X1)|~v7_rlvect_1(X1)|~v6_rlvect_1(X1)|~v5_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(esk7_2(esk61_3(X1,X2,X3),np__0),k5_numbers))).#trainneg
cnf(i_0_3494, plain, (v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_5317, negated_conjecture, (r1_xxreal_0(esk3_0,np__1)|r1_xxreal_0(np__1,esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_10819, negated_conjecture, (v2_valued_0(esk4_0)|~v5_relat_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k4_numbers))).#trainneg
cnf(i_0_3495, plain, (v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_3520, plain, (r1_xxreal_0(X1,u1_struct_0(esk27_0))|v2_xxreal_0(X1)|v3_xxreal_0(u1_struct_0(esk27_0))|~v1_xreal_0(X1)|~v2_struct_0(esk27_0))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 306
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 546
# Removed in clause preprocessing      : 102
# Initial clauses in saturation        : 444
# Processed clauses                    : 4012
# ...of these trivial                  : 72
# ...subsumed                          : 1080
# ...remaining for further processing  : 2859
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 136
# Backward-rewritten                   : 67
# Generated clauses                    : 14692
# ...of the previous two non-trivial   : 13478
# Contextual simplify-reflections      : 1626
# Paramodulations                      : 14692
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
# Propositional unsat check successes  : 0
# Current number of processed clauses  : 2656
#    Positive orientable unit clauses  : 281
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 53
#    Non-unit-clauses                  : 2319
# Current number of unprocessed clauses: 9319
# ...number of literals in the above   : 94515
# Current number of archived formulas  : 0
# Current number of archived clauses   : 3063
# Proof object given clauses           : 148
# Proof search given clauses           : 2860
# Clause-clause subsumption calls (NU) : 15601211
# Rec. Clause-clause subsumption calls : 225874
# Non-unit clause-clause subsumptions  : 2666
# Unit Clause-clause subsumption calls : 45503
# Rewrite failures with RHS unbound    : 40
# BW rewrite match attempts            : 73
# BW rewrite match successes           : 36
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 589270

# -------------------------------------------------
# User time                : 6.668 s
# System time              : 0.088 s
# Total time               : 6.756 s
# ...preprocessing         : 0.307 s
# ...main loop             : 6.449 s
# Maximum resident set size: 44044 pages
