2020-06-04 14:22:10.293333: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_8_bb1/premsel_enigma_01_2020_T30_loop02_epoch_339
2020-06-04 14:22:10.357801: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-06-04 14:22:10.521610: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-06-04 14:22:11.003972: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-06-04 14:22:13.065473: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 2772168 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_8_bb1/premsel_enigma_01_2020_T30_loop02_epoch_339' loaded (query=256; context=512; weigths=real; len_mult=0.000000).
# Preprocessing time       : 3.671 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t187_xcmplx_1, axiom, ![X1]:(v1_xcmplx_0(X1)=>![X2]:(v1_xcmplx_0(X2)=>k4_xcmplx_0(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k4_xcmplx_0(X1),X2))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', t187_xcmplx_1)).
fof(involutiveness_k4_xcmplx_0, axiom, ![X1]:(v1_xcmplx_0(X1)=>k4_xcmplx_0(k4_xcmplx_0(X1))=X1), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', involutiveness_k4_xcmplx_0)).
fof(redefinition_k1_real_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>k1_real_1(X1)=k4_xcmplx_0(X1)), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', redefinition_k1_real_1)).
fof(involutiveness_k1_real_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>k1_real_1(k1_real_1(X1))=X1), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', involutiveness_k1_real_1)).
fof(t9_fdiff_1, axiom, ![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>![X2]:((v1_funct_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(r2_fdiff_1(X2,X1)<=>(r1_tarski(X1,k1_relset_1(k1_numbers,X2))&![X3]:(m1_subset_1(X3,k1_numbers)=>(r2_hidden(X3,X1)=>r1_fdiff_1(X2,X3))))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', t9_fdiff_1)).
fof(t13_fdiff_2, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>![X2]:((v1_funct_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>((r1_fdiff_1(X2,X1)&r1_fdiff_1(X3,k1_seq_1(X2,X1)))=>(r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),X1)&k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),X1)=k8_real_1(k1_fdiff_1(X3,k1_seq_1(X2,X1)),k1_fdiff_1(X2,X1))))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', t13_fdiff_2)).
fof(dt_k4_xcmplx_0, axiom, ![X1]:(v1_xcmplx_0(X1)=>v1_xcmplx_0(k4_xcmplx_0(X1))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', dt_k4_xcmplx_0)).
fof(dt_k1_real_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>m1_subset_1(k1_real_1(X1),k1_numbers)), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', dt_k1_real_1)).
fof(dt_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:((((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&v1_funct_1(X6))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))=>(v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))&m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', dt_k1_partfun1)).
fof(t18_taylor_1, axiom, (((((((r2_relset_1(k1_numbers,k1_numbers,k3_taylor_1,k2_partfun2(k1_numbers,k1_numbers,k24_sin_cos))&v2_funct_1(k3_taylor_1))&k1_relset_1(k1_numbers,k3_taylor_1)=k3_limfunc1(k6_numbers))&k2_relset_1(k1_numbers,k3_taylor_1)=k1_numbers)&r2_fdiff_1(k3_taylor_1,k3_limfunc1(k6_numbers)))&![X1]:(m1_subset_1(X1,k1_numbers)=>(~(r1_xxreal_0(X1,k6_numbers))=>r1_fdiff_1(k3_taylor_1,X1))))&![X1]:(m2_subset_1(X1,k1_numbers,k3_limfunc1(k6_numbers))=>k1_fdiff_1(k3_taylor_1,X1)=k10_real_1(np__1,X1)))&![X1]:(m2_subset_1(X1,k1_numbers,k3_limfunc1(k6_numbers))=>~(r1_xxreal_0(k1_fdiff_1(k3_taylor_1,X1),k6_numbers)))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', t18_taylor_1)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v5_relat_1(X2,X1))=>k2_relset_1(X1,X2)=k10_xtuple_0(X2)), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', redefinition_k2_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', cc1_relset_1)).
fof(t63_sin_cos, axiom, ![X1]:(v1_xreal_0(X1)=>(r1_fdiff_1(k19_sin_cos,X1)&k1_fdiff_1(k19_sin_cos,X1)=k1_real_1(k1_seq_1(k16_sin_cos,X1)))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', t63_sin_cos)).
fof(cc14_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xreal_0(X2))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', cc14_membered)).
fof(dt_k1_seq_1, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v3_valued_0(X1))=>m1_subset_1(k1_seq_1(X1,X2),k1_numbers)), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', dt_k1_seq_1)).
fof(dt_k3_taylor_1, axiom, (v1_funct_1(k3_taylor_1)&m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', dt_k3_taylor_1)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', cc2_relset_1)).
fof(dt_k19_sin_cos, axiom, ((v1_funct_1(k19_sin_cos)&v1_funct_2(k19_sin_cos,k1_numbers,k1_numbers))&m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', dt_k19_sin_cos)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', fc3_membered)).
fof(t33_fdiff_7, conjecture, ![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>((r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))&![X2]:(m1_subset_1(X2,k1_numbers)=>~((r2_hidden(X2,X1)&r1_xxreal_0(X2,k6_numbers)))))=>(r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),X1)&![X2]:(m1_subset_1(X2,k1_numbers)=>(r2_hidden(X2,X1)=>k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),X1),X2)=k1_real_1(k10_real_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,X2)),X2))))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', t33_fdiff_7)).
fof(cc18_valued_0, axiom, ![X1, X2]:(v3_membered(X2)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v3_valued_0(X3))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', cc18_valued_0)).
fof(dt_k10_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>m1_subset_1(k10_real_1(X1,X2),k1_numbers)), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', dt_k10_real_1)).
fof(redefinition_k10_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', redefinition_k10_real_1)).
fof(fc1_valued_0, axiom, ![X1]:((v1_relat_1(X1)&v1_valued_0(X1))=>v1_membered(k10_xtuple_0(X1))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', fc1_valued_0)).
fof(cc12_membered, axiom, ![X1]:(v1_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xcmplx_0(X2))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', cc12_membered)).
fof(cc6_valued_0, axiom, ![X1]:((v1_relat_1(X1)&v3_valued_0(X1))=>(v1_relat_1(X1)&v1_valued_0(X1))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', cc6_valued_0)).
fof(redefinition_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:((((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&v1_funct_1(X6))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))=>k1_partfun1(X1,X2,X3,X4,X5,X6)=k3_relat_1(X5,X6)), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', redefinition_k1_partfun1)).
fof(commutativity_k8_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k8_real_1(X1,X2)=k8_real_1(X2,X1)), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', commutativity_k8_real_1)).
fof(redefinition_k8_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', redefinition_k8_real_1)).
fof(t99_xcmplx_1, axiom, ![X1]:(v1_xcmplx_0(X1)=>![X2]:(v1_xcmplx_0(X2)=>k3_xcmplx_0(X1,k7_xcmplx_0(np__1,X2))=k7_xcmplx_0(X1,X2))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', t99_xcmplx_1)).
fof(l34_fdiff_7, axiom, k3_limfunc1(k6_numbers)=a_0_0_fdiff_7, file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', l34_fdiff_7)).
fof(d7_fdiff_1, axiom, ![X1]:((v1_funct_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>![X2]:(r2_fdiff_1(X1,X2)=>![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(X3=k2_fdiff_1(X1,X2)<=>(k1_relset_1(k1_numbers,X3)=X2&![X4]:(m1_subset_1(X4,k1_numbers)=>(r2_hidden(X4,X2)=>k1_seq_1(X3,X4)=k1_fdiff_1(X1,X4)))))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', d7_fdiff_1)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', spc1_numerals)).
fof(dt_k2_fdiff_1, axiom, ![X1, X2]:((v1_funct_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(v1_funct_1(k2_fdiff_1(X1,X2))&m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', dt_k2_fdiff_1)).
fof(dt_k8_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>m1_subset_1(k8_real_1(X1,X2),k1_numbers)), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', dt_k8_real_1)).
fof(t3_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k3_xcmplx_0(np__1,X1)=X1), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', t3_arithm)).
fof(t1_numerals, axiom, m1_subset_1(k1_xboole_0,k4_ordinal1), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', t1_numerals)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', redefinition_k6_numbers)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', redefinition_k5_numbers)).
fof(cc21_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v3_membered(X2))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', cc21_membered)).
fof(fraenkel_a_0_0_fdiff_7, axiom, ![X1]:(r2_hidden(X1,a_0_0_fdiff_7)<=>?[X2]:((m1_subset_1(X2,k1_numbers)&X1=X2)&~(r1_xxreal_0(X2,k6_numbers)))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', fraenkel_a_0_0_fdiff_7)).
fof(dt_k16_sin_cos, axiom, ((v1_funct_1(k16_sin_cos)&v1_funct_2(k16_sin_cos,k1_numbers,k1_numbers))&m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', dt_k16_sin_cos)).
fof(dt_k3_limfunc1, axiom, ![X1]:(v1_xreal_0(X1)=>m1_subset_1(k3_limfunc1(X1),k1_zfmisc_1(k1_numbers))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', dt_k3_limfunc1)).
fof(dt_k5_numbers, axiom, m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', dt_k5_numbers)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', t1_subset)).
fof(redefinition_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', redefinition_m2_subset_1)).
fof(fc1_numbers, axiom, ~(v1_xboole_0(k1_numbers)), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', fc1_numbers)).
fof(t7_boole, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2))), file('/local1/mptp/loop2/hardpred/preds__0.05/t33_fdiff_7', t7_boole)).
fof(c_0_48, plain, ![X59, X60]:(~v1_xcmplx_0(X59)|(~v1_xcmplx_0(X60)|k4_xcmplx_0(k7_xcmplx_0(X59,X60))=k7_xcmplx_0(k4_xcmplx_0(X59),X60))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t187_xcmplx_1])])])).
fof(c_0_49, plain, ![X112]:(~v1_xcmplx_0(X112)|k4_xcmplx_0(k4_xcmplx_0(X112))=X112), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k4_xcmplx_0])])).
fof(c_0_50, plain, ![X45]:(~m1_subset_1(X45,k1_numbers)|k1_real_1(X45)=k4_xcmplx_0(X45)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_real_1])])).
fof(c_0_51, plain, ![X40]:(~m1_subset_1(X40,k1_numbers)|k1_real_1(k1_real_1(X40))=X40), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k1_real_1])])).
fof(c_0_52, plain, ![X10, X11, X12]:(((r1_tarski(X10,k1_relset_1(k1_numbers,X11))|~r2_fdiff_1(X11,X10)|(~v1_funct_1(X11)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X10)|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers))))&(~m1_subset_1(X12,k1_numbers)|(~r2_hidden(X12,X10)|r1_fdiff_1(X11,X12))|~r2_fdiff_1(X11,X10)|(~v1_funct_1(X11)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X10)|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers)))))&((m1_subset_1(esk3_2(X10,X11),k1_numbers)|~r1_tarski(X10,k1_relset_1(k1_numbers,X11))|r2_fdiff_1(X11,X10)|(~v1_funct_1(X11)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X10)|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers))))&((r2_hidden(esk3_2(X10,X11),X10)|~r1_tarski(X10,k1_relset_1(k1_numbers,X11))|r2_fdiff_1(X11,X10)|(~v1_funct_1(X11)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X10)|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers))))&(~r1_fdiff_1(X11,esk3_2(X10,X11))|~r1_tarski(X10,k1_relset_1(k1_numbers,X11))|r2_fdiff_1(X11,X10)|(~v1_funct_1(X11)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X10)|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t9_fdiff_1])])])])])).
fof(c_0_53, plain, ![X22, X23, X24]:((r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X23,X24),X22)|(~r1_fdiff_1(X23,X22)|~r1_fdiff_1(X24,k1_seq_1(X23,X22)))|(~v1_funct_1(X24)|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v1_funct_1(X23)|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~m1_subset_1(X22,k1_numbers))&(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X23,X24),X22)=k8_real_1(k1_fdiff_1(X24,k1_seq_1(X23,X22)),k1_fdiff_1(X23,X22))|(~r1_fdiff_1(X23,X22)|~r1_fdiff_1(X24,k1_seq_1(X23,X22)))|(~v1_funct_1(X24)|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v1_funct_1(X23)|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~m1_subset_1(X22,k1_numbers))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_fdiff_2])])])])).
cnf(c_0_54, plain, (k4_xcmplx_0(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_55, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
fof(c_0_56, plain, ![X95]:(~v1_xcmplx_0(X95)|v1_xcmplx_0(k4_xcmplx_0(X95))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_xcmplx_0])])).
cnf(c_0_57, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_58, plain, (k1_real_1(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_51])).
fof(c_0_59, plain, ![X72]:(~m1_subset_1(X72,k1_numbers)|m1_subset_1(k1_real_1(X72),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_real_1])])).
cnf(c_0_60, plain, (r2_fdiff_1(X1,X2)|~r1_fdiff_1(X1,esk3_2(X2,X1))|~r1_tarski(X2,k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_61, plain, (r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~r1_fdiff_1(X1,X3)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_53])).
fof(c_0_62, plain, ![X34, X35, X36, X37, X38, X39]:((v1_funct_1(k1_partfun1(X34,X35,X36,X37,X38,X39))|(~v1_funct_1(X38)|~m1_subset_1(X38,k1_zfmisc_1(k2_zfmisc_1(X34,X35)))|~v1_funct_1(X39)|~m1_subset_1(X39,k1_zfmisc_1(k2_zfmisc_1(X36,X37)))))&(m1_subset_1(k1_partfun1(X34,X35,X36,X37,X38,X39),k1_zfmisc_1(k2_zfmisc_1(X34,X37)))|(~v1_funct_1(X38)|~m1_subset_1(X38,k1_zfmisc_1(k2_zfmisc_1(X34,X35)))|~v1_funct_1(X39)|~m1_subset_1(X39,k1_zfmisc_1(k2_zfmisc_1(X36,X37)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_partfun1])])])).
cnf(c_0_63, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_54])).
cnf(c_0_64, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_55])).
cnf(c_0_65, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_66, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_57])).
cnf(c_0_67, plain, (k1_real_1(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_58])).
cnf(c_0_68, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_59])).
fof(c_0_69, plain, ![X19, X20, X21]:(((((((r2_relset_1(k1_numbers,k1_numbers,k3_taylor_1,k2_partfun2(k1_numbers,k1_numbers,k24_sin_cos))&v2_funct_1(k3_taylor_1))&k1_relset_1(k1_numbers,k3_taylor_1)=k3_limfunc1(k6_numbers))&k2_relset_1(k1_numbers,k3_taylor_1)=k1_numbers)&r2_fdiff_1(k3_taylor_1,k3_limfunc1(k6_numbers)))&(~m1_subset_1(X19,k1_numbers)|(r1_xxreal_0(X19,k6_numbers)|r1_fdiff_1(k3_taylor_1,X19))))&(~m2_subset_1(X20,k1_numbers,k3_limfunc1(k6_numbers))|k1_fdiff_1(k3_taylor_1,X20)=k10_real_1(np__1,X20)))&(~m2_subset_1(X21,k1_numbers,k3_limfunc1(k6_numbers))|~r1_xxreal_0(k1_fdiff_1(k3_taylor_1,X21),k6_numbers))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t18_taylor_1])])])])).
fof(c_0_70, plain, ![X77, X78]:(~v1_relat_1(X78)|~v5_relat_1(X78,X77)|k2_relset_1(X77,X78)=k10_xtuple_0(X78)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
fof(c_0_71, plain, ![X53, X54, X55]:(~m1_subset_1(X55,k1_zfmisc_1(k2_zfmisc_1(X53,X54)))|v1_relat_1(X55)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_72, plain, (r2_fdiff_1(X1,X2)|~v3_rcomp_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r1_fdiff_1(X1,esk3_2(X2,X1))|~r1_tarski(X2,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_60])).
cnf(c_0_73, plain, (r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_fdiff_1(X1,X3)|~m1_subset_1(X3,k1_numbers)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_61])).
cnf(c_0_74, plain, (m1_subset_1(esk3_2(X1,X2),k1_numbers)|r2_fdiff_1(X2,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_75, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_76, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_62])).
fof(c_0_77, plain, ![X27]:((r1_fdiff_1(k19_sin_cos,X27)|~v1_xreal_0(X27))&(k1_fdiff_1(k19_sin_cos,X27)=k1_real_1(k1_seq_1(k16_sin_cos,X27))|~v1_xreal_0(X27))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_sin_cos])])])).
fof(c_0_78, plain, ![X41, X42]:(~v3_membered(X41)|(~m1_subset_1(X42,X41)|v1_xreal_0(X42))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc14_membered])])])).
fof(c_0_79, plain, ![X87, X88]:(~v1_relat_1(X87)|~v1_funct_1(X87)|~v3_valued_0(X87)|m1_subset_1(k1_seq_1(X87,X88),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_seq_1])])).
cnf(c_0_80, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(X1),X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64])])).
cnf(c_0_81, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_65])).
cnf(c_0_82, plain, (k4_xcmplx_0(k1_real_1(X1))=X1|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67])])).
cnf(c_0_83, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_68])).
cnf(c_0_84, plain, (k2_relset_1(k1_numbers,k3_taylor_1)=k1_numbers), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_85, plain, (k2_relset_1(X2,X1)=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_86, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_87, plain, (m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[dt_k3_taylor_1])).
fof(c_0_88, plain, ![X67, X68, X69]:((v4_relat_1(X69,X67)|~m1_subset_1(X69,k1_zfmisc_1(k2_zfmisc_1(X67,X68))))&(v5_relat_1(X69,X68)|~m1_subset_1(X69,k1_zfmisc_1(k2_zfmisc_1(X67,X68))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
cnf(c_0_89, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~r1_fdiff_1(X2,k1_seq_1(X1,esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))))|~r1_fdiff_1(X1,esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73])])).
cnf(c_0_90, plain, (r2_fdiff_1(X2,X1)|m1_subset_1(esk3_2(X1,X2),k1_numbers)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_74])).
cnf(c_0_91, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_75])).
cnf(c_0_92, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_76])).
cnf(c_0_93, plain, (r1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_94, plain, (v1_funct_1(k19_sin_cos)), inference(split_conjunct,[status(thm)],[dt_k19_sin_cos])).
cnf(c_0_95, plain, (m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[dt_k19_sin_cos])).
cnf(c_0_96, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_97, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_98, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
cnf(c_0_99, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(X1),X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_100, plain, (k4_xcmplx_0(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_101, plain, (k2_relset_1(k1_numbers,k3_taylor_1)=k1_numbers), inference(evalgc,[status(thm)],[c_0_84])).
cnf(c_0_102, plain, (k2_relset_1(X2,X1)=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_85])).
cnf(c_0_103, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_86])).
cnf(c_0_104, plain, (m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_87])).
cnf(c_0_105, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_88])).
fof(c_0_106, negated_conjecture, ~(![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>((r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))&![X2]:(m1_subset_1(X2,k1_numbers)=>~((r2_hidden(X2,X1)&r1_xxreal_0(X2,k6_numbers)))))=>(r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),X1)&![X2]:(m1_subset_1(X2,k1_numbers)=>(r2_hidden(X2,X1)=>k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),X1),X2)=k1_real_1(k10_real_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,X2)),X2)))))))), inference(assume_negation,[status(cth)],[t33_fdiff_7])).
cnf(c_0_107, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~r1_fdiff_1(X2,k1_seq_1(X1,esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))))|~r1_fdiff_1(X1,esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_89, c_0_90]), c_0_91]), c_0_92])).
cnf(c_0_108, plain, (r1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_93])).
cnf(c_0_109, plain, (v1_funct_1(k19_sin_cos)), inference(evalgc,[status(thm)],[c_0_94])).
cnf(c_0_110, plain, (m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_95])).
cnf(c_0_111, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_96])).
cnf(c_0_112, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_97])).
cnf(c_0_113, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_98])).
fof(c_0_114, plain, ![X79, X80, X81]:(~v3_membered(X80)|(~m1_subset_1(X81,k1_zfmisc_1(k2_zfmisc_1(X79,X80)))|v3_valued_0(X81))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc18_valued_0])])])).
cnf(c_0_115, plain, (k7_xcmplx_0(k1_real_1(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100])])).
fof(c_0_116, plain, ![X70, X71]:(~m1_subset_1(X70,k1_numbers)|~v1_xreal_0(X71)|m1_subset_1(k10_real_1(X70,X71),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k10_real_1])])).
fof(c_0_117, plain, ![X43, X44]:(~m1_subset_1(X43,k1_numbers)|~v1_xreal_0(X44)|k10_real_1(X43,X44)=k7_xcmplx_0(X43,X44)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k10_real_1])])).
fof(c_0_118, plain, ![X107]:(~v1_relat_1(X107)|~v1_valued_0(X107)|v1_membered(k10_xtuple_0(X107))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_valued_0])])).
cnf(c_0_119, plain, (k10_xtuple_0(k3_taylor_1)=k1_numbers|~v5_relat_1(k3_taylor_1,k1_numbers)|~v1_relat_1(k3_taylor_1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102])])).
cnf(c_0_120, plain, (v1_relat_1(k3_taylor_1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104])])).
cnf(c_0_121, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(evalgc,[status(thm)],[c_0_105])).
fof(c_0_122, negated_conjecture, ![X8]:((v3_rcomp_1(esk1_0)&m1_subset_1(esk1_0,k1_zfmisc_1(k1_numbers)))&((r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))&(~m1_subset_1(X8,k1_numbers)|(~r2_hidden(X8,esk1_0)|~r1_xxreal_0(X8,k6_numbers))))&((m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0))&((r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0))&(k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)!=k1_real_1(k10_real_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_106])])])])])).
cnf(c_0_123, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),X2)|~r1_fdiff_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_xreal_0(k1_seq_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_109]), c_0_110])])])).
cnf(c_0_124, plain, (v1_xreal_0(k1_seq_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_113])])])).
cnf(c_0_125, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_114])).
cnf(c_0_126, plain, (k4_xcmplx_0(k7_xcmplx_0(k1_real_1(X1),X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_67])])).
cnf(c_0_127, plain, (m1_subset_1(k10_real_1(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_116])).
cnf(c_0_128, plain, (k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_117])).
fof(c_0_129, plain, ![X101, X102]:(~v1_membered(X101)|(~m1_subset_1(X102,X101)|v1_xcmplx_0(X102))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc12_membered])])])).
cnf(c_0_130, plain, (v1_membered(k10_xtuple_0(X1))|~v1_relat_1(X1)|~v1_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_131, plain, (k10_xtuple_0(k3_taylor_1)=k1_numbers|~v5_relat_1(k3_taylor_1,k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_119, c_0_120])])).
cnf(c_0_132, plain, (v5_relat_1(k3_taylor_1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_104])])).
fof(c_0_133, plain, ![X122]:((v1_relat_1(X122)|(~v1_relat_1(X122)|~v3_valued_0(X122)))&(v1_valued_0(X122)|(~v1_relat_1(X122)|~v3_valued_0(X122)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc6_valued_0])])])).
cnf(c_0_134, negated_conjecture, (~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk1_0)|~r1_xxreal_0(X1,k6_numbers)), inference(split_conjunct,[status(thm)],[c_0_122])).
cnf(c_0_135, plain, (r2_hidden(esk3_2(X1,X2),X1)|r2_fdiff_1(X2,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_136, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_122])).
cnf(c_0_137, negated_conjecture, (v3_rcomp_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_122])).
cnf(c_0_138, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),X2)|~r1_fdiff_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v3_valued_0(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_124])])).
cnf(c_0_139, plain, (r1_xxreal_0(X1,k6_numbers)|r1_fdiff_1(k3_taylor_1,X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_140, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_125])).
cnf(c_0_141, plain, (v1_funct_1(k3_taylor_1)), inference(split_conjunct,[status(thm)],[dt_k3_taylor_1])).
cnf(c_0_142, plain, (k1_real_1(k4_xcmplx_0(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_66])])).
cnf(c_0_143, plain, (k4_xcmplx_0(k7_xcmplx_0(k1_real_1(X1),X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_126, c_0_83])).
cnf(c_0_144, plain, (m1_subset_1(k10_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_127])).
cnf(c_0_145, plain, (k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_128])).
cnf(c_0_146, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_147, plain, (v1_membered(k10_xtuple_0(X1))|~v1_relat_1(X1)|~v1_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_130])).
cnf(c_0_148, plain, (k10_xtuple_0(k3_taylor_1)=k1_numbers), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_131, c_0_132])])])).
cnf(c_0_149, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_133])).
fof(c_0_150, plain, ![X28, X29, X30, X31, X32, X33]:(~v1_funct_1(X32)|~m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(X28,X29)))|~v1_funct_1(X33)|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X30,X31)))|k1_partfun1(X28,X29,X30,X31,X32,X33)=k3_relat_1(X32,X33)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_partfun1])])).
cnf(c_0_151, negated_conjecture, (~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk1_0)|~r1_xxreal_0(X1,k6_numbers)), inference(evalgc,[status(thm)],[c_0_134])).
cnf(c_0_152, plain, (r2_fdiff_1(X2,X1)|r2_hidden(esk3_2(X1,X2),X1)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_135])).
cnf(c_0_153, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_136])).
cnf(c_0_154, negated_conjecture, (v3_rcomp_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_137])).
cnf(c_0_155, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),X2)|~v3_valued_0(X1)|~r1_fdiff_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)), inference(csr,[status(thm)],[c_0_138, c_0_103])).
cnf(c_0_156, plain, (r1_xxreal_0(X1,k6_numbers)|r1_fdiff_1(k3_taylor_1,X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_139])).
cnf(c_0_157, plain, (v3_valued_0(k3_taylor_1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_104]), c_0_113])])])).
cnf(c_0_158, plain, (v1_funct_1(k3_taylor_1)), inference(evalgc,[status(thm)],[c_0_141])).
cnf(c_0_159, negated_conjecture, (r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))), inference(split_conjunct,[status(thm)],[c_0_122])).
fof(c_0_160, plain, ![X108, X109]:(~m1_subset_1(X108,k1_numbers)|~v1_xreal_0(X109)|k8_real_1(X108,X109)=k8_real_1(X109,X108)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k8_real_1])])).
fof(c_0_161, plain, ![X65, X66]:(~m1_subset_1(X65,k1_numbers)|~v1_xreal_0(X66)|k8_real_1(X65,X66)=k3_xcmplx_0(X65,X66)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_real_1])])).
cnf(c_0_162, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)!=k1_real_1(k10_real_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_122])).
cnf(c_0_163, plain, (k1_real_1(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k1_real_1(X1),X2)|~m1_subset_1(k7_xcmplx_0(k1_real_1(X1),X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_143])])).
cnf(c_0_164, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_145])])).
cnf(c_0_165, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_146])).
cnf(c_0_166, plain, (v1_membered(k1_numbers)|~v1_valued_0(k3_taylor_1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_148]), c_0_120])])])).
cnf(c_0_167, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_149])).
cnf(c_0_168, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k3_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(split_conjunct,[status(thm)],[c_0_150])).
cnf(c_0_169, negated_conjecture, (r2_fdiff_1(X1,esk1_0)|~r1_xxreal_0(esk3_2(esk1_0,X1),k6_numbers)|~m1_subset_1(esk3_2(esk1_0,X1),k1_numbers)|~v1_funct_1(X1)|~r1_tarski(esk1_0,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152]), c_0_153]), c_0_154])])])).
cnf(c_0_170, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),X1)|r1_xxreal_0(esk3_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k6_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(esk3_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_156]), c_0_157]), c_0_158]), c_0_104])])])).
cnf(c_0_171, negated_conjecture, (r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))), inference(evalgc,[status(thm)],[c_0_159])).
cnf(c_0_172, plain, (k8_real_1(X1,X2)=k8_real_1(X2,X1)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_160])).
cnf(c_0_173, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_161])).
cnf(c_0_174, negated_conjecture, (k1_real_1(k10_real_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)), inference(evalgc,[status(thm)],[c_0_162])).
cnf(c_0_175, plain, (k1_real_1(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k1_real_1(X1),X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_164])])).
cnf(c_0_176, plain, (v1_xcmplx_0(k1_real_1(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_83])])).
cnf(c_0_177, plain, (v1_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_167]), c_0_157]), c_0_120])])])).
cnf(c_0_178, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_122])).
cnf(c_0_179, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k3_relat_1(X1,X4)|~v1_funct_1(X4)|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_168])).
cnf(c_0_180, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_170]), c_0_171]), c_0_153]), c_0_154])])])).
cnf(c_0_181, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|m1_subset_1(esk3_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_171]), c_0_153]), c_0_154])])])).
cnf(c_0_182, plain, (k8_real_1(X1,X2)=k8_real_1(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_172])).
cnf(c_0_183, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_173])).
cnf(c_0_184, plain, (v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_83]), c_0_113])])])).
cnf(c_0_185, negated_conjecture, (k1_real_1(k7_xcmplx_0(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~v1_xreal_0(esk2_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_145])])).
cnf(c_0_186, plain, (k1_real_1(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k1_real_1(X1),X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_175, c_0_83])).
cnf(c_0_187, plain, (v1_xcmplx_0(X1)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_100])])).
cnf(c_0_188, plain, (v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_176, c_0_177])])])).
cnf(c_0_189, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)), inference(evalgc,[status(thm)],[c_0_178])).
cnf(c_0_190, plain, (k1_partfun1(X1,X2,X3,X4,X5,X6)=k1_partfun1(X7,X8,X9,X10,X5,X6)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_179])])).
cnf(c_0_191, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(csr,[status(thm)],[c_0_180, c_0_181])).
fof(c_0_192, plain, ![X83, X84]:(~v1_xcmplx_0(X83)|(~v1_xcmplx_0(X84)|k3_xcmplx_0(X83,k7_xcmplx_0(np__1,X84))=k7_xcmplx_0(X83,X84))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t99_xcmplx_1])])])).
cnf(c_0_193, plain, (k8_real_1(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_183])])).
cnf(c_0_194, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_67])])).
cnf(c_0_195, negated_conjecture, (k7_xcmplx_0(k1_real_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0))),esk2_0)!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)))|~v1_xcmplx_0(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_186])])).
cnf(c_0_196, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_188])])).
cnf(c_0_197, plain, (k1_fdiff_1(k19_sin_cos,X1)=k1_real_1(k1_seq_1(k16_sin_cos,X1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_198, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_190]), c_0_109]), c_0_158]), c_0_110]), c_0_104])])])).
cnf(c_0_199, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_91]), c_0_109]), c_0_158]), c_0_110]), c_0_104])])])).
cnf(c_0_200, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(np__1,X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_192])).
cnf(c_0_201, plain, (k3_xcmplx_0(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_193])])).
cnf(c_0_202, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_194, c_0_83])).
cnf(c_0_203, plain, (k1_fdiff_1(k3_taylor_1,X1)=k10_real_1(np__1,X1)|~m2_subset_1(X1,k1_numbers,k3_limfunc1(k6_numbers))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_204, plain, (k3_limfunc1(k6_numbers)=a_0_0_fdiff_7), inference(split_conjunct,[status(thm)],[l34_fdiff_7])).
cnf(c_0_205, negated_conjecture, (k7_xcmplx_0(k1_real_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0))),esk2_0)!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)|~v1_xcmplx_0(esk2_0)|~v1_xreal_0(esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)), inference(csr,[status(thm)],[c_0_195, c_0_196])).
cnf(c_0_206, plain, (k1_real_1(k1_seq_1(k16_sin_cos,X1))=k1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_197])).
cnf(c_0_207, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_199]), c_0_110]), c_0_104])])])).
fof(c_0_208, plain, ![X14, X15, X16, X17]:(((k1_relset_1(k1_numbers,X16)=X15|X16!=k2_fdiff_1(X14,X15)|(~v1_funct_1(X16)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X14,X15)|(~v1_funct_1(X14)|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(~m1_subset_1(X17,k1_numbers)|(~r2_hidden(X17,X15)|k1_seq_1(X16,X17)=k1_fdiff_1(X14,X17))|X16!=k2_fdiff_1(X14,X15)|(~v1_funct_1(X16)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X14,X15)|(~v1_funct_1(X14)|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))))&((m1_subset_1(esk4_3(X14,X15,X16),k1_numbers)|k1_relset_1(k1_numbers,X16)!=X15|X16=k2_fdiff_1(X14,X15)|(~v1_funct_1(X16)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X14,X15)|(~v1_funct_1(X14)|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&((r2_hidden(esk4_3(X14,X15,X16),X15)|k1_relset_1(k1_numbers,X16)!=X15|X16=k2_fdiff_1(X14,X15)|(~v1_funct_1(X16)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X14,X15)|(~v1_funct_1(X14)|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(k1_seq_1(X16,esk4_3(X14,X15,X16))!=k1_fdiff_1(X14,esk4_3(X14,X15,X16))|k1_relset_1(k1_numbers,X16)!=X15|X16=k2_fdiff_1(X14,X15)|(~v1_funct_1(X16)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X14,X15)|(~v1_funct_1(X14)|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_fdiff_1])])])])])).
cnf(c_0_209, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_122])).
cnf(c_0_210, plain, (k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)=k8_real_1(k1_fdiff_1(X2,k1_seq_1(X1,X3)),k1_fdiff_1(X1,X3))|~r1_fdiff_1(X1,X3)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_211, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(np__1,X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_200])).
cnf(c_0_212, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_201, c_0_202]), c_0_202])).
cnf(c_0_213, plain, (k10_real_1(np__1,X1)=k1_fdiff_1(k3_taylor_1,X1)|~m2_subset_1(X1,k1_numbers,k3_limfunc1(k6_numbers))), inference(evalgc,[status(thm)],[c_0_203])).
cnf(c_0_214, plain, (k3_limfunc1(k6_numbers)=a_0_0_fdiff_7), inference(evalgc,[status(thm)],[c_0_204])).
cnf(c_0_215, plain, (m1_subset_1(np__1,k1_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_216, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)!=k7_xcmplx_0(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0)|~v1_xcmplx_0(esk2_0)|~v1_xreal_0(esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_206])])).
cnf(c_0_217, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_92]), c_0_109]), c_0_158]), c_0_110]), c_0_104])])])).
cnf(c_0_218, plain, (k1_seq_1(X3,X1)=k1_fdiff_1(X4,X1)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,X2)|X3!=k2_fdiff_1(X4,X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_fdiff_1(X4,X2)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_208])).
fof(c_0_219, plain, ![X25, X26]:((v1_funct_1(k2_fdiff_1(X25,X26))|(~v1_funct_1(X25)|~m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(m1_subset_1(k2_fdiff_1(X25,X26),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|(~v1_funct_1(X25)|~m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_fdiff_1])])])).
cnf(c_0_220, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)), inference(evalgc,[status(thm)],[c_0_209])).
cnf(c_0_221, plain, (k8_real_1(k1_fdiff_1(X2,k1_seq_1(X1,X3)),k1_fdiff_1(X1,X3))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_fdiff_1(X1,X3)|~m1_subset_1(X3,k1_numbers)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_210])).
cnf(c_0_222, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,X2),X1)=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,X2),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_212])])).
cnf(c_0_223, plain, (k10_real_1(np__1,X1)=k1_fdiff_1(k3_taylor_1,X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)), inference(rw,[status(thm)],[c_0_213, c_0_214])).
cnf(c_0_224, plain, (m1_subset_1(np__1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_215])).
fof(c_0_225, plain, ![X61, X62]:(~m1_subset_1(X61,k1_numbers)|~v1_xreal_0(X62)|m1_subset_1(k8_real_1(X61,X62),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_real_1])])).
cnf(c_0_226, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)!=k7_xcmplx_0(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0)|~v1_xcmplx_0(esk2_0)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~v1_xreal_0(esk2_0)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_190]), c_0_109]), c_0_158]), c_0_110]), c_0_104])])])).
cnf(c_0_227, negated_conjecture, (v1_xreal_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_217]), c_0_113])])])).
cnf(c_0_228, plain, (k1_fdiff_1(X4,X1)=k1_seq_1(X3,X1)|X3!=k2_fdiff_1(X4,X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~r2_hidden(X1,X2)|~r2_fdiff_1(X4,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_218])).
cnf(c_0_229, plain, (m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_219])).
cnf(c_0_230, plain, (v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_219])).
cnf(c_0_231, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_190]), c_0_109]), c_0_158]), c_0_110]), c_0_104])])])).
cnf(c_0_232, plain, (k3_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(X2,X3)),k1_fdiff_1(X2,X3))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3)|~v1_xreal_0(k1_fdiff_1(X2,X3))|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(X2,X3)),k1_numbers)|~r1_fdiff_1(X1,k1_seq_1(X2,X3))|~r1_fdiff_1(X2,X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_221])])).
cnf(c_0_233, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,X1),X2)=k7_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_222, c_0_196])).
cnf(c_0_234, plain, (k7_xcmplx_0(np__1,X1)=k1_fdiff_1(k3_taylor_1,X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_145]), c_0_224])])])).
cnf(c_0_235, plain, (m1_subset_1(k8_real_1(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_225])).
fof(c_0_236, plain, ![X110]:(~v1_xcmplx_0(X110)|k3_xcmplx_0(np__1,X110)=X110), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_arithm])])).
cnf(c_0_237, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)!=k7_xcmplx_0(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0)|~v1_xcmplx_0(esk2_0)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_226, c_0_227])])])).
cnf(c_0_238, negated_conjecture, (v1_xcmplx_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_217]), c_0_177])])])).
cnf(c_0_239, plain, (k1_seq_1(k2_fdiff_1(X1,X3),X2)=k1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(X1,X3))|~r2_fdiff_1(X1,X3)|~r2_hidden(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_228])])).
cnf(c_0_240, plain, (m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_229])).
cnf(c_0_241, plain, (v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_230])).
cnf(c_0_242, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_199]), c_0_110]), c_0_104])])])).
cnf(c_0_243, plain, (m1_subset_1(k1_partfun1(X7,X8,X9,X10,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_190])])).
cnf(c_0_244, plain, (k3_xcmplx_0(k1_fdiff_1(X2,X3),k1_fdiff_1(X1,k1_seq_1(X2,X3)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3)|~v1_xreal_0(k1_fdiff_1(X2,X3))|~r1_fdiff_1(X1,k1_seq_1(X2,X3))|~r1_fdiff_1(X2,X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(k1_fdiff_1(X2,X3),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_212])])).
cnf(c_0_245, plain, (k3_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2)=k7_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_234])])).
cnf(c_0_246, plain, (m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_234]), c_0_224])])])).
cnf(c_0_247, plain, (m1_subset_1(k8_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_235])).
cnf(c_0_248, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_236])).
cnf(c_0_249, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)!=k7_xcmplx_0(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_237, c_0_238])])).
cnf(c_0_250, plain, (k1_seq_1(k2_fdiff_1(X1,X2),X3)=k1_fdiff_1(X1,X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_239, c_0_240]), c_0_241])).
cnf(c_0_251, negated_conjecture, (r2_hidden(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_92]), c_0_109]), c_0_158]), c_0_110]), c_0_104])])])).
cnf(c_0_252, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X6,k1_numbers)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_110]), c_0_109])])])).
cnf(c_0_253, plain, (k3_xcmplx_0(k1_fdiff_1(X1,X2),k1_fdiff_1(X3,k1_seq_1(X1,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X3),X2)|~r1_fdiff_1(X3,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(k1_fdiff_1(X3,k1_seq_1(X1,X2)),k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_244, c_0_202])).
cnf(c_0_254, plain, (k3_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2)=k7_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_245, c_0_246])).
cnf(c_0_255, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_183])])).
cnf(c_0_256, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_248])).
cnf(c_0_257, plain, (v1_funct_1(k3_relat_1(X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_179])])).
cnf(c_0_258, negated_conjecture, (k1_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk2_0)!=k7_xcmplx_0(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_249, c_0_250]), c_0_251]), c_0_217])])])).
cnf(c_0_259, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_104]), c_0_158])])])).
cnf(c_0_260, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_190]), c_0_109]), c_0_158]), c_0_110]), c_0_104])])])).
cnf(c_0_261, plain, (k7_xcmplx_0(k1_fdiff_1(X2,k1_seq_1(k3_taylor_1,X1)),X1)=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X2),X1)|~r1_fdiff_1(X2,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~v1_funct_1(X2)|~m1_subset_1(k1_fdiff_1(X2,k1_seq_1(k3_taylor_1,X1)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_254]), c_0_158]), c_0_104])])])).
cnf(c_0_262, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_256]), c_0_224])])])).
cnf(c_0_263, plain, (v1_funct_1(k3_relat_1(X1,k19_sin_cos))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_110]), c_0_109])])])).
cnf(c_0_264, negated_conjecture, (k1_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk2_0)!=k7_xcmplx_0(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_258, c_0_259]), c_0_260])).
cnf(c_0_265, plain, (k7_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2)),X2)=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m2_subset_1(X2,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2))|~r1_fdiff_1(k3_taylor_1,X2)|~v1_funct_1(X1)|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_261, c_0_262]), c_0_246])).
cnf(c_0_266, plain, (v1_funct_1(k3_relat_1(k3_taylor_1,k19_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_104]), c_0_158])])])).
cnf(c_0_267, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1)), inference(split_conjunct,[status(thm)],[t1_numerals])).
cnf(c_0_268, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
cnf(c_0_269, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
fof(c_0_270, plain, ![X93, X94]:(~v3_membered(X93)|(~m1_subset_1(X94,k1_zfmisc_1(X93))|v3_membered(X94))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc21_membered])])])).
fof(c_0_271, plain, ![X48, X50, X51]:((((m1_subset_1(esk5_1(X48),k1_numbers)|~r2_hidden(X48,a_0_0_fdiff_7))&(X48=esk5_1(X48)|~r2_hidden(X48,a_0_0_fdiff_7)))&(~r1_xxreal_0(esk5_1(X48),k6_numbers)|~r2_hidden(X48,a_0_0_fdiff_7)))&(~m1_subset_1(X51,k1_numbers)|X50!=X51|r1_xxreal_0(X51,k6_numbers)|r2_hidden(X50,a_0_0_fdiff_7))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_0_0_fdiff_7])])])])])])])).
cnf(c_0_272, negated_conjecture, (k1_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk2_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk2_0)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m2_subset_1(esk2_0,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0))|~r1_fdiff_1(k3_taylor_1,esk2_0)|~m1_subset_1(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_265]), c_0_238]), c_0_227]), c_0_109]), c_0_110])])])).
cnf(c_0_273, plain, (m1_subset_1(k1_fdiff_1(k19_sin_cos,X1),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_206])])).
cnf(c_0_274, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_179]), c_0_109]), c_0_158])])])).
cnf(c_0_275, plain, (m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[dt_k16_sin_cos])).
fof(c_0_276, plain, ![X52]:(~v1_xreal_0(X52)|m1_subset_1(k3_limfunc1(X52),k1_zfmisc_1(k1_numbers))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_limfunc1])])).
cnf(c_0_277, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_267])).
cnf(c_0_278, plain, (k1_xboole_0=k6_numbers), inference(evalgc,[status(thm)],[c_0_268])).
cnf(c_0_279, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_269])).
cnf(c_0_280, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_270])).
cnf(c_0_281, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[dt_k5_numbers])).
fof(c_0_282, plain, ![X63, X64]:(~r2_hidden(X63,X64)|m1_subset_1(X63,X64)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_283, plain, (r1_xxreal_0(X1,k6_numbers)|r2_hidden(X2,a_0_0_fdiff_7)|~m1_subset_1(X1,k1_numbers)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_271])).
cnf(c_0_284, negated_conjecture, (k1_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk2_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk2_0)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~m2_subset_1(esk2_0,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(k3_taylor_1,esk2_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_272, c_0_273]), c_0_274]), c_0_108])).
cnf(c_0_285, plain, (m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_275])).
cnf(c_0_286, plain, (v1_funct_1(k16_sin_cos)), inference(split_conjunct,[status(thm)],[dt_k16_sin_cos])).
cnf(c_0_287, plain, (m1_subset_1(k3_limfunc1(X1),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_276])).
cnf(c_0_288, plain, (m1_subset_1(k6_numbers,k5_numbers)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_277, c_0_278]), c_0_279])).
cnf(c_0_289, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_280])).
cnf(c_0_290, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_281])).
cnf(c_0_291, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_282])).
cnf(c_0_292, plain, (r2_hidden(X2,a_0_0_fdiff_7)|r1_xxreal_0(X1,k6_numbers)|X1!=X2|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_283])).
cnf(c_0_293, negated_conjecture, (~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~m2_subset_1(esk2_0,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(k3_taylor_1,esk2_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_284]), c_0_110]), c_0_104])])])).
cnf(c_0_294, plain, (v3_valued_0(k16_sin_cos)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_285]), c_0_113])])])).
cnf(c_0_295, plain, (v1_relat_1(k16_sin_cos)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_285])])).
cnf(c_0_296, plain, (v1_funct_1(k16_sin_cos)), inference(evalgc,[status(thm)],[c_0_286])).
fof(c_0_297, plain, ![X125, X126, X127]:((~m2_subset_1(X127,X125,X126)|m1_subset_1(X127,X126)|(v1_xboole_0(X125)|v1_xboole_0(X126)|~m1_subset_1(X126,k1_zfmisc_1(X125))))&(~m1_subset_1(X127,X126)|m2_subset_1(X127,X125,X126)|(v1_xboole_0(X125)|v1_xboole_0(X126)|~m1_subset_1(X126,k1_zfmisc_1(X125))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m2_subset_1])])])])])).
cnf(c_0_298, plain, (m1_subset_1(k3_limfunc1(X1),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_287])).
cnf(c_0_299, plain, (v1_xreal_0(k6_numbers)|~v3_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_288])])).
cnf(c_0_300, plain, (v3_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_289, c_0_290]), c_0_113])])])).
cnf(c_0_301, negated_conjecture, (~r1_xxreal_0(esk2_0,k6_numbers)|~m1_subset_1(esk2_0,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_251])])).
cnf(c_0_302, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_291])).
cnf(c_0_303, plain, (r1_xxreal_0(X1,k6_numbers)|r2_hidden(X1,a_0_0_fdiff_7)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_292])])).
fof(c_0_304, plain, ~v1_xboole_0(k1_numbers), inference(fof_simplification,[status(thm)],[fc1_numbers])).
fof(c_0_305, plain, ![X105, X106]:(~r2_hidden(X105,X106)|~v1_xboole_0(X106)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
cnf(c_0_306, negated_conjecture, (~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~m2_subset_1(esk2_0,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(k3_taylor_1,esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_293, c_0_112]), c_0_294]), c_0_295]), c_0_296])])])).
cnf(c_0_307, plain, (m2_subset_1(X1,X3,X2)|v1_xboole_0(X3)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_297])).
cnf(c_0_308, plain, (m1_subset_1(a_0_0_fdiff_7,k1_zfmisc_1(k1_numbers))|~v1_xreal_0(k6_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_214])])).
cnf(c_0_309, plain, (v1_xreal_0(k6_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_299, c_0_300])])).
cnf(c_0_310, negated_conjecture, (~r1_xxreal_0(esk2_0,k6_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_301, c_0_217])])).
cnf(c_0_311, plain, (m1_subset_1(X1,a_0_0_fdiff_7)|r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_303])])).
cnf(c_0_312, plain, (~v1_xboole_0(k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_304])).
cnf(c_0_313, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_305])).
cnf(c_0_314, negated_conjecture, (~m2_subset_1(esk2_0,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(k3_taylor_1,esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_306, c_0_124]), c_0_157]), c_0_120]), c_0_158])])])).
cnf(c_0_315, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_307])).
cnf(c_0_316, plain, (m1_subset_1(a_0_0_fdiff_7,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_308, c_0_309])])])).
cnf(c_0_317, negated_conjecture, (m1_subset_1(esk2_0,a_0_0_fdiff_7)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_311]), c_0_217])])])).
cnf(c_0_318, plain, (~v1_xboole_0(k1_numbers)), inference(evalgc,[status(thm)],[c_0_312])).
cnf(c_0_319, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_313])).
cnf(c_0_320, negated_conjecture, (v1_xboole_0(a_0_0_fdiff_7)|~r1_fdiff_1(k3_taylor_1,esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_314, c_0_315]), c_0_316]), c_0_317])]), c_0_318])])).
cnf(c_0_321, plain, (r1_xxreal_0(X1,k6_numbers)|~v1_xboole_0(a_0_0_fdiff_7)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_319, c_0_303])])).
cnf(c_0_322, negated_conjecture, (v1_xboole_0(a_0_0_fdiff_7)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_320, c_0_156]), c_0_217])]), c_0_310])])).
cnf(c_0_323, plain, (r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_321, c_0_322])])])).
cnf(c_0_324, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_323]), c_0_217])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 325
# Proof object clause steps            : 238
# Proof object formula steps           : 87
# Proof object conjectures             : 51
# Proof object clause conjectures      : 48
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 65
# Proof object initial formulas used   : 48
# Proof object generating inferences   : 79
# Proof object simplifying inferences  : 174
# Training examples: 149 positive, 1563 negative
# Training: Positive examples begin
cnf(c_0_278, plain, (k1_xboole_0=k6_numbers)).# trainpos
cnf(c_0_279, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_214, plain, (k3_limfunc1(k6_numbers)=a_0_0_fdiff_7)).# trainpos
cnf(c_0_154, negated_conjecture, (v3_rcomp_1(esk1_0))).# trainpos
cnf(c_0_158, plain, (v1_funct_1(k3_taylor_1))).# trainpos
cnf(c_0_109, plain, (v1_funct_1(k19_sin_cos))).# trainpos
cnf(c_0_296, plain, (v1_funct_1(k16_sin_cos))).# trainpos
cnf(c_0_113, plain, (v3_membered(k1_numbers))).# trainpos
cnf(c_0_151, negated_conjecture, (~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk1_0)|~r1_xxreal_0(X1,k6_numbers))).# trainpos
cnf(c_0_153, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_174, negated_conjecture, (k1_real_1(k10_real_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0))).# trainpos
cnf(c_0_318, plain, (~v1_xboole_0(k1_numbers))).# trainpos
cnf(c_0_101, plain, (k2_relset_1(k1_numbers,k3_taylor_1)=k1_numbers)).# trainpos
cnf(c_0_224, plain, (m1_subset_1(np__1,k1_numbers))).# trainpos
cnf(c_0_277, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).# trainpos
cnf(c_0_302, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_290, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_81, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_189, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0))).# trainpos
cnf(c_0_104, plain, (m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_220, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0))).# trainpos
cnf(c_0_110, plain, (m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_285, plain, (m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_167, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).# trainpos
cnf(c_0_108, plain, (r1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_64, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_256, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_319, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_66, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_67, plain, (k1_real_1(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_298, plain, (m1_subset_1(k3_limfunc1(X1),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_83, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_171, negated_conjecture, (r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))))).# trainpos
cnf(c_0_156, plain, (r1_xxreal_0(X1,k6_numbers)|r1_fdiff_1(k3_taylor_1,X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_147, plain, (v1_membered(k10_xtuple_0(X1))|~v1_relat_1(X1)|~v1_valued_0(X1))).# trainpos
cnf(c_0_111, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_165, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_206, plain, (k1_real_1(k1_seq_1(k16_sin_cos,X1))=k1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_103, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_289, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_292, plain, (r2_hidden(X2,a_0_0_fdiff_7)|r1_xxreal_0(X1,k6_numbers)|X1!=X2|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_213, plain, (k10_real_1(np__1,X1)=k1_fdiff_1(k3_taylor_1,X1)|~m2_subset_1(X1,k1_numbers,k3_limfunc1(k6_numbers)))).# trainpos
cnf(c_0_90, plain, (r2_fdiff_1(X2,X1)|m1_subset_1(esk3_2(X1,X2),k1_numbers)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_145, plain, (k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_140, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).# trainpos
cnf(c_0_144, plain, (m1_subset_1(k10_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_112, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).# trainpos
cnf(c_0_121, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).# trainpos
cnf(c_0_102, plain, (k2_relset_1(X2,X1)=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).# trainpos
cnf(c_0_182, plain, (k8_real_1(X1,X2)=k8_real_1(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_183, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_247, plain, (m1_subset_1(k8_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_63, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_211, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(np__1,X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_241, plain, (v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_72, plain, (r2_fdiff_1(X1,X2)|~v3_rcomp_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r1_fdiff_1(X1,esk3_2(X2,X1))|~r1_tarski(X2,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_240, plain, (m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_315, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_92, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_179, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k3_relat_1(X1,X4)|~v1_funct_1(X4)|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_91, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_152, plain, (r2_fdiff_1(X2,X1)|r2_hidden(esk3_2(X1,X2),X1)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_228, plain, (k1_fdiff_1(X4,X1)=k1_seq_1(X3,X1)|X3!=k2_fdiff_1(X4,X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~r2_hidden(X1,X2)|~r2_fdiff_1(X4,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_73, plain, (r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_fdiff_1(X1,X3)|~m1_subset_1(X3,k1_numbers)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_221, plain, (k8_real_1(k1_fdiff_1(X2,k1_seq_1(X1,X3)),k1_fdiff_1(X1,X3))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_fdiff_1(X1,X3)|~m1_subset_1(X3,k1_numbers)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_303, plain, (r1_xxreal_0(X1,k6_numbers)|r2_hidden(X1,a_0_0_fdiff_7)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_295, plain, (v1_relat_1(k16_sin_cos))).# trainpos
cnf(c_0_185, negated_conjecture, (k1_real_1(k7_xcmplx_0(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~v1_xreal_0(esk2_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers))).# trainpos
cnf(c_0_164, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_120, plain, (v1_relat_1(k3_taylor_1))).# trainpos
cnf(c_0_157, plain, (v3_valued_0(k3_taylor_1))).# trainpos
cnf(c_0_255, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_181, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|m1_subset_1(esk3_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_119, plain, (k10_xtuple_0(k3_taylor_1)=k1_numbers|~v5_relat_1(k3_taylor_1,k1_numbers)|~v1_relat_1(k3_taylor_1))).# trainpos
cnf(c_0_132, plain, (v5_relat_1(k3_taylor_1,k1_numbers))).# trainpos
cnf(c_0_300, plain, (v3_membered(k5_numbers))).# trainpos
cnf(c_0_193, plain, (k8_real_1(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_308, plain, (m1_subset_1(a_0_0_fdiff_7,k1_zfmisc_1(k1_numbers))|~v1_xreal_0(k6_numbers))).# trainpos
cnf(c_0_294, plain, (v3_valued_0(k16_sin_cos))).# trainpos
cnf(c_0_82, plain, (k4_xcmplx_0(k1_real_1(X1))=X1|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_142, plain, (k1_real_1(k4_xcmplx_0(X1))=X1|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_124, plain, (v1_xreal_0(k1_seq_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).# trainpos
cnf(c_0_234, plain, (k7_xcmplx_0(np__1,X1)=k1_fdiff_1(k3_taylor_1,X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_184, plain, (v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_299, plain, (v1_xreal_0(k6_numbers)|~v3_membered(k5_numbers))).# trainpos
cnf(c_0_80, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(X1),X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_190, plain, (k1_partfun1(X1,X2,X3,X4,X5,X6)=k1_partfun1(X7,X8,X9,X10,X5,X6)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).# trainpos
cnf(c_0_89, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~r1_fdiff_1(X2,k1_seq_1(X1,esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))))|~r1_fdiff_1(X1,esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)),k1_numbers))).# trainpos
cnf(c_0_316, plain, (m1_subset_1(a_0_0_fdiff_7,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_169, negated_conjecture, (r2_fdiff_1(X1,esk1_0)|~r1_xxreal_0(esk3_2(esk1_0,X1),k6_numbers)|~m1_subset_1(esk3_2(esk1_0,X1),k1_numbers)|~v1_funct_1(X1)|~r1_tarski(esk1_0,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_194, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_262, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_176, plain, (v1_xcmplx_0(k1_real_1(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_239, plain, (k1_seq_1(k2_fdiff_1(X1,X3),X2)=k1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(X1,X3))|~r2_fdiff_1(X1,X3)|~r2_hidden(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_148, plain, (k10_xtuple_0(k3_taylor_1)=k1_numbers)).# trainpos
cnf(c_0_246, plain, (m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).# trainpos
cnf(c_0_311, plain, (m1_subset_1(X1,a_0_0_fdiff_7)|r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_115, plain, (k7_xcmplx_0(k1_real_1(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_232, plain, (k3_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(X2,X3)),k1_fdiff_1(X2,X3))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3)|~v1_xreal_0(k1_fdiff_1(X2,X3))|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(X2,X3)),k1_numbers)|~r1_fdiff_1(X1,k1_seq_1(X2,X3))|~r1_fdiff_1(X2,X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).# trainpos
cnf(c_0_187, plain, (v1_xcmplx_0(X1)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_321, plain, (r1_xxreal_0(X1,k6_numbers)|~v1_xboole_0(a_0_0_fdiff_7)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_243, plain, (m1_subset_1(k1_partfun1(X7,X8,X9,X10,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).# trainpos
cnf(c_0_198, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_123, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),X2)|~r1_fdiff_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_xreal_0(k1_seq_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))))).# trainpos
cnf(c_0_231, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_166, plain, (v1_membered(k1_numbers)|~v1_valued_0(k3_taylor_1))).# trainpos
cnf(c_0_126, plain, (k4_xcmplx_0(k7_xcmplx_0(k1_real_1(X1),X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_201, plain, (k3_xcmplx_0(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_177, plain, (v1_membered(k1_numbers))).# trainpos
cnf(c_0_138, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),X2)|~r1_fdiff_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v3_valued_0(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_244, plain, (k3_xcmplx_0(k1_fdiff_1(X2,X3),k1_fdiff_1(X1,k1_seq_1(X2,X3)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3)|~v1_xreal_0(k1_fdiff_1(X2,X3))|~r1_fdiff_1(X1,k1_seq_1(X2,X3))|~r1_fdiff_1(X2,X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(k1_fdiff_1(X2,X3),k1_numbers))).# trainpos
cnf(c_0_163, plain, (k1_real_1(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k1_real_1(X1),X2)|~m1_subset_1(k7_xcmplx_0(k1_real_1(X1),X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_252, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X6,k1_numbers)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_273, plain, (m1_subset_1(k1_fdiff_1(k19_sin_cos,X1),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_170, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),X1)|r1_xxreal_0(esk3_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k6_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(esk3_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k1_numbers))).# trainpos
cnf(c_0_188, plain, (v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_175, plain, (k1_real_1(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k1_real_1(X1),X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers))).# trainpos
cnf(c_0_259, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_196, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_180, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k1_numbers))).# trainpos
cnf(c_0_195, negated_conjecture, (k7_xcmplx_0(k1_real_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0))),esk2_0)!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)))|~v1_xcmplx_0(esk2_0))).# trainpos
cnf(c_0_199, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))).# trainpos
cnf(c_0_216, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)!=k7_xcmplx_0(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0)|~v1_xcmplx_0(esk2_0)|~v1_xreal_0(esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0)))).# trainpos
cnf(c_0_242, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))).# trainpos
cnf(c_0_207, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))).# trainpos
cnf(c_0_251, negated_conjecture, (r2_hidden(esk2_0,esk1_0))).# trainpos
cnf(c_0_217, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers))).# trainpos
cnf(c_0_226, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)!=k7_xcmplx_0(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0)|~v1_xcmplx_0(esk2_0)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~v1_xreal_0(esk2_0)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_260, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_227, negated_conjecture, (v1_xreal_0(esk2_0))).# trainpos
cnf(c_0_238, negated_conjecture, (v1_xcmplx_0(esk2_0))).# trainpos
cnf(c_0_301, negated_conjecture, (~r1_xxreal_0(esk2_0,k6_numbers)|~m1_subset_1(esk2_0,k1_numbers))).# trainpos
cnf(c_0_317, negated_conjecture, (m1_subset_1(esk2_0,a_0_0_fdiff_7))).# trainpos
cnf(c_0_237, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)!=k7_xcmplx_0(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0)|~v1_xcmplx_0(esk2_0)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_257, plain, (v1_funct_1(k3_relat_1(X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_263, plain, (v1_funct_1(k3_relat_1(X1,k19_sin_cos))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_258, negated_conjecture, (k1_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk2_0)!=k7_xcmplx_0(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_266, plain, (v1_funct_1(k3_relat_1(k3_taylor_1,k19_sin_cos)))).# trainpos
cnf(c_0_274, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_222, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,X2),X1)=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,X2),k1_numbers))).# trainpos
cnf(c_0_245, plain, (k3_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2)=k7_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).# trainpos
cnf(c_0_261, plain, (k7_xcmplx_0(k1_fdiff_1(X2,k1_seq_1(k3_taylor_1,X1)),X1)=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X2),X1)|~r1_fdiff_1(X2,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~v1_funct_1(X2)|~m1_subset_1(k1_fdiff_1(X2,k1_seq_1(k3_taylor_1,X1)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).# trainpos
cnf(c_0_272, negated_conjecture, (k1_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),esk2_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk2_0)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m2_subset_1(esk2_0,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0))|~r1_fdiff_1(k3_taylor_1,esk2_0)|~m1_subset_1(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers))).# trainpos
cnf(c_0_293, negated_conjecture, (~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~m2_subset_1(esk2_0,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(k3_taylor_1,esk2_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers))).# trainpos
cnf(c_0_306, negated_conjecture, (~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~m2_subset_1(esk2_0,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(k3_taylor_1,esk2_0))).# trainpos
cnf(c_0_314, negated_conjecture, (~m2_subset_1(esk2_0,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(k3_taylor_1,esk2_0))).# trainpos
cnf(c_0_320, negated_conjecture, (v1_xboole_0(a_0_0_fdiff_7)|~r1_fdiff_1(k3_taylor_1,esk2_0))).# trainpos
cnf(c_0_322, negated_conjecture, (v1_xboole_0(a_0_0_fdiff_7))).# trainpos
cnf(c_0_323, plain, (r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_151, plain, (k4_xcmplx_0(np__0)=np__0)).#trainneg
cnf(i_0_82, plain, (v1_funct_1(k24_sin_cos))).#trainneg
cnf(i_0_24, plain, (v2_funct_1(k3_taylor_1))).#trainneg
cnf(i_0_109, plain, (v1_xreal_0(esk6_0))).#trainneg
cnf(i_0_44, plain, (v1_xboole_0(np__0))).#trainneg
cnf(i_0_110, plain, (v1_xcmplx_0(esk6_0))).#trainneg
cnf(i_0_77, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_132, plain, (~r1_xxreal_0(np__1,np__0))).#trainneg
cnf(i_0_131, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_66, plain, (v1_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_108, plain, (v1_int_1(esk6_0))).#trainneg
cnf(i_0_78, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_107, plain, (v2_xxreal_0(np__2))).#trainneg
cnf(i_0_97, plain, (v6_membered(k4_ordinal1))).#trainneg
cnf(i_0_111, plain, (v1_xxreal_0(esk6_0))).#trainneg
cnf(i_0_130, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_76, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_91, plain, (m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_92, plain, (m1_subset_1(np__0,k5_numbers))).#trainneg
cnf(i_0_104, plain, (m1_subset_1(np__2,k1_numbers))).#trainneg
cnf(i_0_105, plain, (m1_subset_1(np__2,k5_numbers))).#trainneg
cnf(i_0_43, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_112, plain, (m1_subset_1(esk6_0,k1_numbers))).#trainneg
cnf(i_0_158, plain, (~r1_xxreal_0(np__2,np__0))).#trainneg
cnf(i_0_65, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_93, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_121, plain, (~r1_xxreal_0(np__0,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_149, plain, (r1_xxreal_0(np__0,np__0))).#trainneg
cnf(i_0_114, plain, (v2_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_115, plain, (v5_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_70, plain, (v1_xreal_0(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_86, plain, (v1_int_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_23, plain, (k3_limfunc1(k6_numbers)=k1_relset_1(k1_numbers,k3_taylor_1))).#trainneg
cnf(i_0_21, plain, (r2_fdiff_1(k3_taylor_1,k3_limfunc1(k6_numbers)))).#trainneg
cnf(i_0_155, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__0))).#trainneg
cnf(i_0_87, plain, (v1_relat_1(k3_relat_1(X1,X2)))).#trainneg
cnf(i_0_106, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_102, plain, (v1_funct_2(k19_sin_cos,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_52, plain, (v1_funct_2(k16_sin_cos,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_81, plain, (v1_funct_2(k24_sin_cos,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_154, plain, (r1_xxreal_0(np__0,k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_117, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_42, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_46, plain, (~r2_hidden(X1,a_0_0_fdiff_7)|~r1_xxreal_0(esk5_1(X1),k6_numbers))).#trainneg
cnf(i_0_116, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_157, plain, (k3_xcmplx_0(X1,k6_numbers)=k6_numbers|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_142, plain, (k7_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_47, plain, (esk5_1(X1)=X1|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_80, plain, (m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_123, plain, (k4_xxreal_1(X1,k1_xxreal_0)=k3_limfunc1(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_99, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_159, plain, (v3_valued_0(k2_zfmisc_1(X2,X1))|~v3_membered(X1))).#trainneg
cnf(i_0_56, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_135, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_25, plain, (r2_relset_1(k1_numbers,k1_numbers,k3_taylor_1,k2_partfun2(k1_numbers,k1_numbers,k24_sin_cos)))).#trainneg
cnf(i_0_13, plain, (X1=k2_fdiff_1(X2,X3)|k1_relset_1(k1_numbers,X1)!=X3|k1_fdiff_1(X2,esk4_3(X2,X3,X1))!=k1_seq_1(X1,esk4_3(X2,X3,X1))|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_152, plain, (v1_funct_1(k2_funct_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_133, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainneg
cnf(i_0_18, plain, (~r1_xxreal_0(k1_fdiff_1(k3_taylor_1,X1),k6_numbers)|~m2_subset_1(X1,k1_numbers,k3_limfunc1(k6_numbers)))).#trainneg
cnf(i_0_72, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_150, plain, (v3_membered(k10_xtuple_0(X1))|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_153, plain, (v1_relat_1(k2_funct_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_48, plain, (m1_subset_1(esk5_1(X1),k1_numbers)|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_129, plain, (v1_int_1(X2)|~v5_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_88, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_137, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_139, plain, (v1_membered(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_141, plain, (v3_rcomp_1(k4_xxreal_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_126, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_83, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_138, plain, (v1_valued_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_144, plain, (v1_int_1(k6_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_118, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_62, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_140, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_143, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_67, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_90, plain, (m1_subset_1(k2_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).#trainneg
cnf(i_0_12, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_136, plain, (r2_relset_1(X2,X3,X1,X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_113, plain, (k1_funct_1(X1,X2)=k1_seq_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_128, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_74, plain, (m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_145, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_146, plain, (k2_partfun2(X2,X3,X1)=k2_funct_1(X1)|~v1_funct_1(X1)|~v2_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_11, plain, (r1_fdiff_1(X3,X1)|~v3_rcomp_1(X2)|~v1_funct_1(X3)|~r2_hidden(X1,X2)|~r2_fdiff_1(X3,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_148, plain, (v1_funct_1(k2_partfun2(X1,X2,X3))|~v1_funct_1(X3)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_147, plain, (m1_subset_1(k2_partfun2(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v1_funct_1(X3)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_156, plain, (r2_relset_1(X2,X3,X4,X1)|~r2_relset_1(X2,X3,X1,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_17, plain, (k1_relset_1(k1_numbers,X1)=X2|X1!=k2_fdiff_1(X3,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_15, plain, (X3=k2_fdiff_1(X1,X2)|m1_subset_1(esk4_3(X1,X2,X3),k1_numbers)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_14, plain, (X3=k2_fdiff_1(X1,X2)|r2_hidden(esk4_3(X1,X2,X3),X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_182, negated_conjecture, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_160, plain, (np__0=k6_numbers)).#trainneg
cnf(i_0_163, plain, (v1_xcmplx_0(np__1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_351, plain, (v4_relat_1(k24_sin_cos,k1_numbers))).#trainneg
cnf(i_0_213, plain, (v1_xreal_0(np__2))).#trainneg
cnf(i_0_218, plain, (v1_xreal_0(np__1))).#trainneg
cnf(i_0_274, plain, (v7_ordinal1(k6_numbers))).#trainneg
cnf(i_0_165, plain, (v1_xcmplx_0(np__1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_166, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_403, plain, (r2_fdiff_1(k3_taylor_1,X1)|r1_xxreal_0(esk3_2(X1,k3_taylor_1),k6_numbers)|~r1_tarski(X1,a_0_0_fdiff_7)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(esk3_2(X1,k3_taylor_1),k1_numbers))).#trainneg
cnf(i_0_402, plain, (r2_fdiff_1(k19_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xreal_0(esk3_2(X1,k19_sin_cos)))).#trainneg
cnf(i_0_318, plain, (r2_fdiff_1(k3_taylor_1,X1)|m1_subset_1(esk3_2(X1,k3_taylor_1),k1_numbers)|~r1_tarski(X1,a_0_0_fdiff_7)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_176, negated_conjecture, (k4_xcmplx_0(k10_real_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k10_real_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0),k1_numbers))).#trainneg
cnf(i_0_321, plain, (r2_fdiff_1(X1,X2)|m1_subset_1(esk3_2(X2,X1),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_183, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_208, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_263, plain, (v7_ordinal1(np__2))).#trainneg
cnf(i_0_175, plain, (~r1_xxreal_0(X1,k6_numbers)|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_168, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_352, plain, (v4_relat_1(k16_sin_cos,k1_numbers))).#trainneg
cnf(i_0_268, plain, (v7_ordinal1(np__1))).#trainneg
cnf(i_0_297, plain, (v1_relat_1(k24_sin_cos))).#trainneg
cnf(i_0_299, plain, (v1_relat_1(k19_sin_cos))).#trainneg
cnf(i_0_301, negated_conjecture, (v3_membered(esk1_0))).#trainneg
cnf(i_0_170, plain, (v1_int_1(np__1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_173, plain, (v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_337, plain, (v1_xreal_0(k10_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_336, plain, (r2_hidden(k10_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_330, plain, (v3_valued_0(k19_sin_cos))).#trainneg
cnf(i_0_184, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k3_limfunc1(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_396, plain, (r1_tarski(X1,a_0_0_fdiff_7)|~r2_fdiff_1(k3_taylor_1,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_199, plain, (r2_hidden(np__1,k1_numbers))).#trainneg
cnf(i_0_189, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_328, plain, (v3_valued_0(k24_sin_cos))).#trainneg
cnf(i_0_260, plain, (v7_ordinal1(esk6_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_186, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k16_sin_cos))).#trainneg
cnf(i_0_353, plain, (v4_relat_1(k19_sin_cos,k1_numbers))).#trainneg
cnf(i_0_308, negated_conjecture, (v1_membered(esk1_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_333, plain, (v1_valued_0(k16_sin_cos)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_181, plain, (m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_357, plain, (v5_relat_1(k19_sin_cos,k1_numbers))).#trainneg
cnf(i_0_363, plain, (v1_xreal_0(k8_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_294, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1))=k1_fdiff_1(k19_sin_cos,X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_344, plain, (r2_hidden(k1_seq_1(X1,X2),k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_295, plain, (k1_real_1(k1_fdiff_1(k19_sin_cos,X1))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_391, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k2_relset_1(X2,X3))|~v5_relat_1(X3,X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_354, plain, (v4_relat_1(k3_taylor_1,k1_numbers))).#trainneg
cnf(i_0_187, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k19_sin_cos))).#trainneg
cnf(i_0_185, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k24_sin_cos))).#trainneg
cnf(i_0_355, plain, (v5_relat_1(k24_sin_cos,k1_numbers))).#trainneg
cnf(i_0_229, plain, (v1_xcmplx_0(np__2)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_230, plain, (v1_xcmplx_0(np__2)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_334, plain, (v1_valued_0(k19_sin_cos)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_335, plain, (v1_valued_0(k3_taylor_1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_188, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_taylor_1))).#trainneg
cnf(i_0_190, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))))).#trainneg
cnf(i_0_362, plain, (r2_hidden(k8_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_369, plain, (m1_subset_1(k8_real_1(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_44, plain, (v1_xboole_0(k6_numbers))).#trainneg
cnf(i_0_151, plain, (k4_xcmplx_0(k6_numbers)=k6_numbers)).#trainneg
cnf(i_0_501, plain, (v1_int_1(k6_numbers))).#trainneg
cnf(i_0_158, plain, (~r1_xxreal_0(np__2,k6_numbers))).#trainneg
cnf(i_0_192, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|r2_hidden(esk1_0,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_93, plain, (m2_subset_1(k6_numbers,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_132, plain, (~r1_xxreal_0(np__1,k6_numbers))).#trainneg
cnf(i_0_149, plain, (r1_xxreal_0(k6_numbers,k6_numbers))).#trainneg
cnf(i_0_91, plain, (m1_subset_1(k6_numbers,k1_numbers))).#trainneg
cnf(i_0_193, plain, (v1_xboole_0(k5_numbers)|r2_hidden(np__2,k5_numbers))).#trainneg
cnf(i_0_195, plain, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|r2_hidden(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_476, plain, (r2_fdiff_1(X2,X1)|~v1_xboole_0(X1)|~v1_funct_1(X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_524, negated_conjecture, (k4_xcmplx_0(k7_xcmplx_0(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k7_xcmplx_0(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0),k1_numbers)|~v1_xreal_0(esk2_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_494, negated_conjecture, (m1_subset_1(esk3_2(esk1_0,X1),k1_numbers)|r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~r1_tarski(esk1_0,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_484, negated_conjecture, (X2=k2_fdiff_1(X1,esk1_0)|~r1_xxreal_0(esk4_3(X1,esk1_0,X2),k6_numbers)|~m1_subset_1(esk4_3(X1,esk1_0,X2),k1_numbers)|k1_relset_1(k1_numbers,X2)!=esk1_0|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_608, plain, (k4_xcmplx_0(k1_fdiff_1(k19_sin_cos,X1))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_477, plain, (m1_subset_1(esk3_2(X1,X2),X1)|r2_fdiff_1(X2,X1)|~v1_funct_1(X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_564, plain, (v1_int_1(np__2))).#trainneg
cnf(i_0_572, plain, (v1_int_1(np__1))).#trainneg
cnf(i_0_121, plain, (~r1_xxreal_0(k6_numbers,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_198, plain, (v1_xboole_0(k5_numbers)|r2_hidden(np__1,k5_numbers))).#trainneg
cnf(i_0_356, plain, (v5_relat_1(k16_sin_cos,k1_numbers))).#trainneg
cnf(i_0_666, plain, (v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_234, plain, (v1_xcmplx_0(np__1)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_255, plain, (v1_int_1(k1_real_1(X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_200, plain, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|r2_hidden(k3_limfunc1(X1),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_201, plain, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|r2_hidden(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_436, plain, (r1_fdiff_1(k3_taylor_1,X1)|~r2_hidden(X1,a_0_0_fdiff_7)|~m1_subset_1(a_0_0_fdiff_7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(a_0_0_fdiff_7))).#trainneg
cnf(i_0_618, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_566, plain, (r2_fdiff_1(X1,a_0_0_fdiff_7)|~r1_xxreal_0(esk3_2(a_0_0_fdiff_7,X1),k6_numbers)|~v1_funct_1(X1)|~r1_tarski(a_0_0_fdiff_7,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(a_0_0_fdiff_7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(a_0_0_fdiff_7))).#trainneg
cnf(i_0_616, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_518, plain, (r2_hidden(esk3_2(X1,k3_taylor_1),k1_numbers)|r2_fdiff_1(k3_taylor_1,X1)|~r1_tarski(X1,a_0_0_fdiff_7)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_596, plain, (r1_xxreal_0(esk6_0,esk6_0))).#trainneg
cnf(i_0_659, plain, (k9_xtuple_0(k3_taylor_1)=a_0_0_fdiff_7)).#trainneg
cnf(i_0_155, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),k6_numbers))).#trainneg
cnf(i_0_203, plain, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|r2_hidden(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_587, negated_conjecture, (r1_tarski(esk1_0,a_0_0_fdiff_7)|~r2_fdiff_1(k3_taylor_1,esk1_0))).#trainneg
cnf(i_0_235, plain, (v1_xcmplx_0(np__1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_240, plain, (v1_xcmplx_0(k6_numbers)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_272, plain, (v7_ordinal1(k1_real_1(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_289, plain, (v1_xxreal_0(k1_real_1(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_204, plain, (v1_xboole_0(k5_numbers)|r2_hidden(k6_numbers,k5_numbers))).#trainneg
cnf(i_0_205, plain, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|r2_hidden(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_533, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_519, plain, (v1_xreal_0(esk3_2(X1,k3_taylor_1))|r2_fdiff_1(k3_taylor_1,X1)|~r1_tarski(X1,a_0_0_fdiff_7)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_490, plain, (k8_real_1(k1_fdiff_1(X2,X3),k1_fdiff_1(X1,k1_seq_1(X2,X3)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3)|~v1_xreal_0(k1_fdiff_1(X2,X3))|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(X2,X3)),k1_numbers)|~r1_fdiff_1(X1,k1_seq_1(X2,X3))|~r1_fdiff_1(X2,X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_455, plain, (v1_relat_1(k1_partfun1(X3,X4,X5,X6,X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_532, plain, (r2_hidden(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_624, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_656, plain, (m1_subset_1(X1,X2)|~v5_relat_1(X3,X2)|~v1_relat_1(X3)|~r2_hidden(X1,k10_xtuple_0(X3)))).#trainneg
cnf(i_0_509, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_154, plain, (r1_xxreal_0(k6_numbers,k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_206, plain, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|r2_hidden(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_571, plain, (m1_subset_1(k1_relset_1(k1_numbers,k16_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_611, plain, (m1_subset_1(k1_relset_1(k1_numbers,k19_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_210, plain, (v1_xreal_0(esk5_1(X1))|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_211, negated_conjecture, (v1_xreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_626, plain, (k1_fdiff_1(k3_taylor_1,np__1)=np__1|~v1_xcmplx_0(np__1)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_489, plain, (k8_real_1(k1_fdiff_1(X1,X2),k1_fdiff_1(X3,k1_seq_1(X1,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X3),X2)|~v1_xreal_0(k1_fdiff_1(X3,k1_seq_1(X1,X2)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~r1_fdiff_1(X3,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_581, plain, (r2_hidden(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_526, plain, (r2_hidden(esk3_2(X1,X2),k1_numbers)|r2_fdiff_1(X2,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_591, plain, (k3_xcmplx_0(k1_fdiff_1(X2,X3),k1_fdiff_1(X1,k1_seq_1(X2,X3)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3)|~r1_fdiff_1(X1,k1_seq_1(X2,X3))|~r1_fdiff_1(X2,X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~v1_xreal_0(k1_fdiff_1(X1,k1_seq_1(X2,X3)))|~m1_subset_1(k1_fdiff_1(X2,X3),k1_numbers))).#trainneg
cnf(i_0_612, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_309, plain, (v1_membered(k5_numbers)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_332, plain, (v1_valued_0(k24_sin_cos)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_485, plain, (X3=k2_fdiff_1(X2,X1)|~v1_xboole_0(X1)|k1_relset_1(k1_numbers,X3)!=X1|~v1_funct_1(X3)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_214, plain, (v1_xreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_280, plain, (v1_xxreal_0(np__2)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_281, plain, (v1_xxreal_0(np__2)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_228, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_245, negated_conjecture, (v1_int_1(esk1_0)|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_219, plain, (v1_xreal_0(k3_limfunc1(X1))|~v3_membered(k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_220, plain, (v1_xreal_0(k24_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_527, plain, (v1_xreal_0(esk3_2(X1,X2))|r2_fdiff_1(X2,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_551, negated_conjecture, (r2_hidden(esk3_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k1_numbers)|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_926, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k16_sin_cos)))).#trainneg
cnf(i_0_746, plain, (v3_membered(a_0_0_fdiff_7))).#trainneg
cnf(i_0_843, negated_conjecture, (k4_xcmplx_0(k10_real_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~v1_xcmplx_0(k10_real_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0))|~m1_subset_1(k4_xcmplx_0(k10_real_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0)),k1_numbers))).#trainneg
cnf(i_0_222, plain, (v1_xreal_0(k16_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_770, plain, (k5_numbers=k6_numbers|r2_hidden(np__2,k5_numbers))).#trainneg
cnf(i_0_806, plain, (v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_869, plain, (v1_xcmplx_0(k6_numbers))).#trainneg
cnf(i_0_224, plain, (v1_xreal_0(k19_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_225, plain, (v1_xreal_0(k3_taylor_1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_713, plain, (r1_fdiff_1(k1_partfun1(X4,X5,X6,X7,X1,X2),X3)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~r1_fdiff_1(X1,X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_935, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k19_sin_cos)))).#trainneg
cnf(i_0_987, plain, (v1_xxreal_0(np__2))).#trainneg
cnf(i_0_747, plain, (v1_membered(a_0_0_fdiff_7)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_227, plain, (v1_xcmplx_0(esk5_1(X1))|~v1_membered(k1_numbers)|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_803, plain, (k5_numbers=k6_numbers|r2_hidden(np__1,k5_numbers))).#trainneg
cnf(i_0_285, plain, (v1_xxreal_0(np__1)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_286, plain, (v1_xxreal_0(np__1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_769, plain, (v1_xxreal_0(k6_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_231, plain, (v1_xcmplx_0(k5_numbers)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_236, plain, (v1_xcmplx_0(k3_limfunc1(X1))|~v1_membered(k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_739, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|r1_xxreal_0(k1_seq_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1))),k6_numbers)|~r1_fdiff_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k1_seq_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1))),k1_numbers))).#trainneg
cnf(i_0_963, plain, (v1_xcmplx_0(X1)|~v1_int_1(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_781, plain, (r2_fdiff_1(k19_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(esk3_2(X1,k19_sin_cos),k1_numbers))).#trainneg
cnf(i_0_757, plain, (k1_real_1(k6_numbers)=k6_numbers|~m1_subset_1(k6_numbers,k1_numbers))).#trainneg
cnf(i_0_913, plain, (m1_subset_1(X1,X2)|~v5_relat_1(k3_taylor_1,X2)|~r2_hidden(X1,k1_numbers))).#trainneg
cnf(i_0_761, negated_conjecture, (k1_zfmisc_1(k1_numbers)=k6_numbers|r2_hidden(esk1_0,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_237, plain, (v1_xcmplx_0(k24_sin_cos)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_885, plain, (k5_numbers=k6_numbers|r2_hidden(k6_numbers,k5_numbers))).#trainneg
cnf(i_0_291, plain, (v1_xxreal_0(k6_numbers)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_248, plain, (v1_int_1(k5_numbers)|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_924, plain, (v3_membered(k1_relset_1(k1_numbers,k16_sin_cos)))).#trainneg
cnf(i_0_933, plain, (v3_membered(k1_relset_1(k1_numbers,k19_sin_cos)))).#trainneg
cnf(i_0_239, plain, (v1_xcmplx_0(k16_sin_cos)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_241, plain, (v1_xcmplx_0(k19_sin_cos)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_795, plain, (r2_fdiff_1(X1,X2)|m1_subset_1(esk3_2(X2,X1),X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_716, negated_conjecture, (r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_867, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(a_0_0_fdiff_7))|~r2_fdiff_1(k3_taylor_1,esk1_0))).#trainneg
cnf(i_0_771, plain, (k1_zfmisc_1(k1_numbers)=k6_numbers|r2_hidden(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_191, plain, (r2_hidden(esk6_0,k1_numbers))).#trainneg
cnf(i_0_844, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1))=k1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1)),k1_numbers))).#trainneg
cnf(i_0_242, plain, (v1_xcmplx_0(k3_taylor_1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_262, negated_conjecture, (v7_ordinal1(esk1_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_265, plain, (v7_ordinal1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_925, plain, (v1_membered(k1_relset_1(k1_numbers,k16_sin_cos))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_934, plain, (v1_membered(k1_relset_1(k1_numbers,k19_sin_cos))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_244, plain, (v1_int_1(esk5_1(X1))|~v5_membered(k1_numbers)|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_253, plain, (v1_int_1(k3_limfunc1(X1))|~v5_membered(k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_852, plain, (v1_xreal_0(k1_fdiff_1(k3_taylor_1,X1))|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_970, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_793, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|m1_subset_1(esk3_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_778, negated_conjecture, (r2_fdiff_1(X1,esk1_0)|m1_subset_1(esk3_2(esk1_0,X1),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_904, plain, (r2_hidden(X1,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_964, plain, (v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_936, plain, (v1_xreal_0(X1)|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_510, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_194, plain, (r2_hidden(np__2,k1_numbers))).#trainneg
cnf(i_0_254, plain, (v1_int_1(k24_sin_cos)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_304, plain, (v3_membered(k24_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_279, negated_conjecture, (v1_xxreal_0(esk1_0)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_282, plain, (v1_xxreal_0(k5_numbers)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_870, plain, (v1_int_1(k1_real_1(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_256, plain, (v1_int_1(k16_sin_cos)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_258, plain, (v1_int_1(k19_sin_cos)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_749, negated_conjecture, (r2_fdiff_1(k3_taylor_1,esk1_0)|~r1_tarski(esk1_0,a_0_0_fdiff_7)|~m1_subset_1(esk3_2(esk1_0,k3_taylor_1),k1_numbers))).#trainneg
cnf(i_0_962, plain, (v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_832, plain, (v1_xreal_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_775, plain, (r2_fdiff_1(X1,X2)|~v1_xboole_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_824, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3),X4)|~r1_fdiff_1(X3,k1_fdiff_1(X1,esk3_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3))))|~r1_fdiff_1(k2_fdiff_1(X1,X2),esk3_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3)))|~v1_funct_1(X3)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r1_tarski(X4,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X4)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(esk3_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3)),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3)),k1_numbers))).#trainneg
cnf(i_0_504, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_513, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_570, plain, (k9_xtuple_0(k16_sin_cos)=k1_relset_1(k1_numbers,k16_sin_cos))).#trainneg
cnf(i_0_259, plain, (v1_int_1(k3_taylor_1)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_305, plain, (v3_membered(k16_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_741, plain, (v1_xreal_0(a_0_0_fdiff_7)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_742, plain, (v1_xcmplx_0(a_0_0_fdiff_7)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_919, plain, (v1_xreal_0(k1_relset_1(k1_numbers,k16_sin_cos))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_922, plain, (v7_ordinal1(k1_relset_1(k1_numbers,k16_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_261, plain, (v7_ordinal1(esk5_1(X1))|~v6_membered(k1_numbers)|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_270, plain, (v7_ordinal1(k3_limfunc1(X1))|~v6_membered(k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_996, plain, (v1_xcmplx_0(k1_real_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_628, plain, (k3_xcmplx_0(X1,k1_fdiff_1(k3_taylor_1,X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(X2,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_864, plain, (m1_subset_1(np__1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_979, plain, (k3_xcmplx_0(k1_fdiff_1(X2,X3),k1_fdiff_1(X1,k1_seq_1(X2,X3)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3)|~v1_xcmplx_0(k1_fdiff_1(X2,X3))|~v1_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(X2,X3)))|~v1_xreal_0(k1_fdiff_1(X2,X3))|~r1_fdiff_1(X1,k1_seq_1(X2,X3))|~r1_fdiff_1(X2,X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_505, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1051, plain, (v1_xxreal_0(np__1))).#trainneg
cnf(i_0_626, plain, (k1_fdiff_1(k3_taylor_1,np__1)=np__1|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_271, plain, (v7_ordinal1(k24_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1045, plain, (r1_xxreal_0(np__2,np__2))).#trainneg
cnf(i_0_1053, plain, (v1_xxreal_0(k6_numbers))).#trainneg
cnf(i_0_273, plain, (v7_ordinal1(k16_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_275, plain, (v7_ordinal1(k19_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1037, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,X5,X6),X7)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~r1_tarski(X7,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,X6)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X7)|~r1_fdiff_1(X6,k1_seq_1(X5,esk3_2(X7,k1_partfun1(X1,X2,X3,X4,X5,X6))))|~r1_fdiff_1(X5,esk3_2(X7,k1_partfun1(X1,X2,X3,X4,X5,X6)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(esk3_2(X7,k1_partfun1(X1,X2,X3,X4,X5,X6)),k1_numbers))).#trainneg
cnf(i_0_1032, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),X2)|~r1_fdiff_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k1_seq_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_1251, plain, (v3_valued_0(k2_zfmisc_1(X1,k1_numbers)))).#trainneg
cnf(i_0_1288, plain, (v1_xboole_0(a_0_0_fdiff_7)|r2_hidden(np__1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_565, plain, (X2=k2_fdiff_1(X1,a_0_0_fdiff_7)|~r1_xxreal_0(esk4_3(X1,a_0_0_fdiff_7,X2),k6_numbers)|k1_relset_1(k1_numbers,X2)!=a_0_0_fdiff_7|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,a_0_0_fdiff_7)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_276, plain, (v7_ordinal1(k3_taylor_1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_743, plain, (v1_int_1(a_0_0_fdiff_7)|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1203, negated_conjecture, (v1_xxreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1070, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_278, plain, (v1_xxreal_0(esk5_1(X1))|~v2_membered(k1_numbers)|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_287, plain, (v1_xxreal_0(k3_limfunc1(X1))|~v2_membered(k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1107, plain, (v1_xboole_0(X1)|r2_hidden(esk3_2(X1,X2),X1)|r2_fdiff_1(X2,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1031, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos),X3)|~v1_xreal_0(k1_fdiff_1(X1,esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos))))|~r1_fdiff_1(k2_fdiff_1(X1,X2),esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos)))|~v1_funct_1(k2_fdiff_1(X1,X2))|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos)),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_1012, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(X6,k1_numbers)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1008, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X6,k1_numbers)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_503, plain, (v1_xboole_0(k1_zfmisc_1(X1))|r2_hidden(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_610, plain, (k9_xtuple_0(k19_sin_cos)=k1_relset_1(k1_numbers,k19_sin_cos))).#trainneg
cnf(i_0_1216, negated_conjecture, (r2_fdiff_1(k3_taylor_1,esk1_0)|~r1_tarski(esk1_0,a_0_0_fdiff_7))).#trainneg
cnf(i_0_288, plain, (v1_xxreal_0(k24_sin_cos)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1204, plain, (v1_xxreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_744, plain, (v7_ordinal1(a_0_0_fdiff_7)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_745, plain, (v1_xxreal_0(a_0_0_fdiff_7)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_923, plain, (v1_xxreal_0(k1_relset_1(k1_numbers,k16_sin_cos))|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_290, plain, (v1_xxreal_0(k16_sin_cos)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_292, plain, (v1_xxreal_0(k19_sin_cos)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1285, plain, (k3_xcmplx_0(k1_fdiff_1(k3_taylor_1,X2),X1)=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m2_subset_1(X2,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_1284, plain, (k3_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2)=k7_xcmplx_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_1295, plain, (k7_xcmplx_0(k1_fdiff_1(X1,X2),k1_seq_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_xcmplx_0(k1_fdiff_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,X2))|~m2_subset_1(k1_seq_1(X1,X2),k1_numbers,a_0_0_fdiff_7)|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2)))|~v1_xreal_0(k1_fdiff_1(X1,X2))|~r1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1130, plain, (k7_xcmplx_0(k1_real_1(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~m1_subset_1(k7_xcmplx_0(k1_real_1(X1),X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1124, plain, (k7_xcmplx_0(k1_real_1(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(k7_xcmplx_0(k1_real_1(X1),X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1282, plain, (k7_xcmplx_0(k1_fdiff_1(X1,X2),k1_seq_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~v1_xreal_0(k1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2)))|~r1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_xcmplx_0(k1_fdiff_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,X2))|~m2_subset_1(k1_seq_1(X1,X2),k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_506, plain, (v1_int_1(X1)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_303, plain, (v3_membered(k3_limfunc1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_764, plain, (r2_hidden(k6_numbers,k1_numbers))).#trainneg
cnf(i_0_1026, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k3_relat_1(k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_293, plain, (v1_xxreal_0(k3_taylor_1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1151, plain, (v1_int_1(X1)|~v5_membered(k1_numbers)|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_989, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_990, negated_conjecture, (v1_int_1(esk1_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_928, plain, (v1_xreal_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_931, plain, (v7_ordinal1(k1_relset_1(k1_numbers,k19_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_306, plain, (v3_membered(k19_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1274, plain, (m1_subset_1(k3_xcmplx_0(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1163, plain, (k4_xcmplx_0(k7_xcmplx_0(k7_xcmplx_0(k1_real_1(X1),X2),X3))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~m1_subset_1(k7_xcmplx_0(k1_real_1(X1),X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1281, plain, (k7_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2)),X2)=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),X2)|~v1_xreal_0(k1_fdiff_1(k3_taylor_1,X2))|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2))|~r1_fdiff_1(k3_taylor_1,X2)|~v1_funct_1(X1)|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2)))|~v1_xreal_0(X2)|~m2_subset_1(X2,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_1177, negated_conjecture, (v1_xreal_0(esk3_2(esk1_0,X1))|r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_507, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_620, plain, (v1_int_1(k4_xcmplx_0(X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_863, plain, (m1_subset_1(np__2,a_0_0_fdiff_7))).#trainneg
cnf(i_0_1034, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k3_relat_1(k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_307, plain, (v3_membered(k3_taylor_1)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_311, plain, (v1_membered(k24_sin_cos)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1056, plain, (v1_xcmplx_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_932, plain, (v1_xxreal_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_335, plain, (v1_valued_0(k3_taylor_1))).#trainneg
cnf(i_0_310, plain, (v1_membered(k3_limfunc1(X1))|~v1_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_312, plain, (v1_membered(k16_sin_cos)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1305, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos),X3)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~r1_fdiff_1(X2,k1_seq_1(X1,esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos))))|~r1_fdiff_1(X1,esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_1356, negated_conjecture, (r2_fdiff_1(X1,esk1_0)|~r1_xxreal_0(esk3_2(esk1_0,X1),k6_numbers)|~m1_subset_1(esk3_2(esk1_0,X1),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_1307, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xreal_0(esk3_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_747, plain, (v1_membered(a_0_0_fdiff_7))).#trainneg
cnf(i_0_334, plain, (v1_valued_0(k19_sin_cos))).#trainneg
cnf(i_0_333, plain, (v1_valued_0(k16_sin_cos))).#trainneg
cnf(i_0_308, negated_conjecture, (v1_membered(esk1_0))).#trainneg
cnf(i_0_313, plain, (v1_membered(k19_sin_cos)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1299, plain, (r1_xxreal_0(np__1,np__1))).#trainneg
cnf(i_0_332, plain, (v1_valued_0(k24_sin_cos))).#trainneg
cnf(i_0_309, plain, (v1_membered(k5_numbers))).#trainneg
cnf(i_0_934, plain, (v1_membered(k1_relset_1(k1_numbers,k19_sin_cos)))).#trainneg
cnf(i_0_925, plain, (v1_membered(k1_relset_1(k1_numbers,k16_sin_cos)))).#trainneg
cnf(i_0_314, plain, (v1_membered(k3_taylor_1)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_315, plain, (v3_rcomp_1(k3_limfunc1(X1))|~v1_xreal_0(k1_xxreal_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1398, negated_conjecture, (r2_fdiff_1(k3_taylor_1,esk1_0)|~m1_subset_1(esk1_0,k1_zfmisc_1(a_0_0_fdiff_7)))).#trainneg
cnf(i_0_1300, plain, (k1_fdiff_1(k3_taylor_1,np__1)=np__1|v1_xboole_0(a_0_0_fdiff_7))).#trainneg
cnf(i_0_1319, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos),X3)|~r1_fdiff_1(k2_fdiff_1(X1,X2),esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos)))|~v1_funct_1(k2_fdiff_1(X1,X2))|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos)))|~m1_subset_1(k1_fdiff_1(X1,esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos))),k1_numbers)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos)),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_1397, plain, (k1_zfmisc_1(X1)=k6_numbers|r2_hidden(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_227, plain, (v1_xcmplx_0(esk5_1(X1))|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_1551, plain, (v1_xboole_0(a_0_0_fdiff_7)|r2_hidden(np__2,a_0_0_fdiff_7))).#trainneg
cnf(i_0_1535, negated_conjecture, (r2_fdiff_1(k19_sin_cos,esk1_0)|~r1_tarski(esk1_0,k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k19_sin_cos))))).#trainneg
cnf(i_0_317, plain, (r1_xxreal_0(X1,esk6_0)|r1_xxreal_0(esk6_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1403, plain, (v1_xxreal_0(a_0_0_fdiff_7)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1094, plain, (v1_int_1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1301, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1539, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_320, plain, (r2_fdiff_1(X1,k1_relset_1(k1_numbers,X1))|m1_subset_1(esk3_2(k1_relset_1(k1_numbers,X1),X1),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,X1)))).#trainneg
cnf(i_0_322, plain, (k3_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_1439, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1383, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1313, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,X7)),X8)|~r1_fdiff_1(X5,esk3_2(X8,k1_partfun1(X1,X2,X3,X4,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,X7))))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,X7))|~v1_funct_1(X5)|~r1_tarski(X8,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,X7))))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,X7),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,X7),k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X8,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X8)|~r1_fdiff_1(X7,k1_seq_1(X6,k1_seq_1(X5,esk3_2(X8,k1_partfun1(X1,X2,X3,X4,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,X7))))))|~r1_fdiff_1(X6,k1_seq_1(X5,esk3_2(X8,k1_partfun1(X1,X2,X3,X4,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,X7)))))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X5,esk3_2(X8,k1_partfun1(X1,X2,X3,X4,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,X7)))),k1_numbers))).#trainneg
cnf(i_0_1315, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,X5,k19_sin_cos),X6)|~r1_fdiff_1(X5,esk3_2(X6,k1_partfun1(X1,X2,X3,X4,X5,k19_sin_cos)))|~v1_funct_1(X5)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,k19_sin_cos)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~v1_xreal_0(k1_seq_1(X5,esk3_2(X6,k1_partfun1(X1,X2,X3,X4,X5,k19_sin_cos)))))).#trainneg
cnf(i_0_1462, plain, (k4_xcmplx_0(k7_xcmplx_0(X1,np__1))=k1_real_1(X1)|~v1_xcmplx_0(k1_real_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1317, plain, (r2_fdiff_1(k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos),X2)|~r1_fdiff_1(X1,esk3_2(X2,k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos)))|~m1_subset_1(k1_seq_1(X1,esk3_2(X2,k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos))),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_621, plain, (v7_ordinal1(k4_xcmplx_0(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1071, plain, (r2_fdiff_1(k19_sin_cos,k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(esk3_2(k1_relset_1(k1_numbers,k19_sin_cos),k19_sin_cos),k1_numbers)|~v3_rcomp_1(k1_relset_1(k1_numbers,k19_sin_cos)))).#trainneg
cnf(i_0_323, plain, (k3_xcmplx_0(k6_numbers,X1)=k6_numbers|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k6_numbers))).#trainneg
cnf(i_0_1352, plain, (v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_1625, plain, (v3_rcomp_1(a_0_0_fdiff_7)|~v1_xreal_0(k1_xxreal_0))).#trainneg
cnf(i_0_1256, plain, (v1_int_1(k1_relset_1(k1_numbers,k16_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1576, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),k19_sin_cos),X2)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos))|~r1_fdiff_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),k19_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_xreal_0(k1_seq_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),k19_sin_cos)))))).#trainneg
cnf(i_0_1672, plain, (v1_xcmplx_0(X1)|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_1585, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1694, plain, (r1_xxreal_0(k6_numbers,np__2))).#trainneg
cnf(i_0_1675, negated_conjecture, (r2_fdiff_1(k19_sin_cos,esk1_0)|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k19_sin_cos))))).#trainneg
cnf(i_0_338, plain, (v1_xcmplx_0(k10_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1680, plain, (r1_xxreal_0(esk6_0,np__2)|r1_xxreal_0(np__2,esk6_0))).#trainneg
cnf(i_0_1255, plain, (v1_xcmplx_0(a_0_0_fdiff_7)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1475, plain, (v1_int_1(k1_relset_1(k1_numbers,k19_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_508, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_339, plain, (v1_int_1(k10_real_1(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_340, plain, (v7_ordinal1(k10_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1619, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(esk3_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_1828, plain, (k7_xcmplx_0(k6_numbers,X1)=k6_numbers|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))).#trainneg
cnf(i_0_1567, plain, (r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos),X1)|r1_xxreal_0(esk3_2(X1,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos)),k6_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos)))|~m1_subset_1(esk3_2(X1,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1648, negated_conjecture, (r1_fdiff_1(k3_taylor_1,X1)|~r2_hidden(X1,esk1_0)|~m1_subset_1(esk1_0,k1_zfmisc_1(a_0_0_fdiff_7)))).#trainneg
cnf(i_0_1572, plain, (r2_fdiff_1(k1_partfun1(X4,X5,X6,X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos),X3)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~r1_fdiff_1(X2,k1_seq_1(X1,esk3_2(X3,k1_partfun1(X4,X5,X6,X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos))))|~r1_fdiff_1(X1,esk3_2(X3,k1_partfun1(X4,X5,X6,X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(X4,X5,X6,X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos)))|~m1_subset_1(esk3_2(X3,k1_partfun1(X4,X5,X6,X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1616, plain, (k7_xcmplx_0(k1_real_1(X1),np__1)=k1_real_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1577, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),k19_sin_cos),X2)|r1_xxreal_0(k1_seq_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),k19_sin_cos))),k6_numbers)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1))|~r1_fdiff_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),k19_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),k19_sin_cos)))|~m1_subset_1(esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),k19_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k1_seq_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),k19_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_1835, plain, (v1_xxreal_0(X1)|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_622, plain, (v1_xxreal_0(k4_xcmplx_0(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1697, plain, (r1_xxreal_0(k6_numbers,np__1))).#trainneg
cnf(i_0_1825, plain, (r2_fdiff_1(k19_sin_cos,k1_relset_1(k1_numbers,k19_sin_cos))|~v3_rcomp_1(k1_relset_1(k1_numbers,k19_sin_cos)))).#trainneg
cnf(i_0_341, plain, (v1_xxreal_0(k10_real_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1682, plain, (r1_xxreal_0(esk6_0,np__1)|r1_xxreal_0(np__1,esk6_0))).#trainneg
cnf(i_0_1337, plain, (v1_int_1(a_0_0_fdiff_7)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1297, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1467, plain, (v1_int_1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_346, plain, (v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1617, plain, (r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k19_sin_cos),X1)|~v1_xreal_0(esk3_2(X1,k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k19_sin_cos)))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1801, plain, (k1_real_1(k1_real_1(X1))=k7_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers))).#trainneg
cnf(i_0_1569, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X4,X5,X6,X7,X1,X2),k19_sin_cos),X3)|~v3_valued_0(k1_partfun1(X4,X5,X6,X7,X1,X2))|~r1_fdiff_1(X2,k1_seq_1(X1,esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X4,X5,X6,X7,X1,X2),k19_sin_cos))))|~r1_fdiff_1(X1,esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X4,X5,X6,X7,X1,X2),k19_sin_cos)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X4,X5,X6,X7,X1,X2),k19_sin_cos)))|~m1_subset_1(esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X4,X5,X6,X7,X1,X2),k19_sin_cos)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_1779, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,X5,k19_sin_cos),X6)|~r1_fdiff_1(X5,esk3_2(X6,k1_partfun1(X1,X2,X3,X4,X5,k19_sin_cos)))|~v1_funct_1(X5)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,k19_sin_cos)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~v3_valued_0(X5)|~v1_relat_1(X5))).#trainneg
cnf(i_0_1574, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),k19_sin_cos),X4)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)))|~r1_fdiff_1(X1,esk3_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),k19_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))|~v1_funct_1(X1)|~r1_tarski(X4,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),k19_sin_cos)))|~m1_subset_1(esk3_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),k19_sin_cos)),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X4)|~r1_fdiff_1(X3,k1_seq_1(X2,k1_seq_1(X1,esk3_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),k19_sin_cos)))))|~r1_fdiff_1(X2,k1_seq_1(X1,esk3_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),k19_sin_cos))))|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X1,esk3_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),k19_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_347, plain, (v1_int_1(k1_seq_1(X1,X2))|~v5_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1685, plain, (r1_xxreal_0(esk6_0,k6_numbers)|r1_xxreal_0(k6_numbers,esk6_0))).#trainneg
cnf(i_0_1559, plain, (v1_membered(k24_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1564, plain, (v1_membered(k16_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1339, negated_conjecture, (r1_xxreal_0(esk1_0,esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_348, plain, (v7_ordinal1(k1_seq_1(X1,X2))|~v6_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_349, plain, (v1_xxreal_0(k1_seq_1(X1,X2))|~v2_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1799, plain, (k1_real_1(k4_xcmplx_0(X1))=k7_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_1593, plain, (k7_xcmplx_0(k1_real_1(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1772, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,k19_sin_cos)),X7)|~r1_fdiff_1(X6,k1_seq_1(X5,esk3_2(X7,k1_partfun1(X1,X2,X3,X4,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,k19_sin_cos)))))|~r1_fdiff_1(X5,esk3_2(X7,k1_partfun1(X1,X2,X3,X4,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,k19_sin_cos))))|~v1_funct_1(X5)|~v1_funct_1(X6)|~r1_tarski(X7,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,k19_sin_cos))))|~m1_subset_1(k1_seq_1(X5,esk3_2(X7,k1_partfun1(X1,X2,X3,X4,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,k19_sin_cos)))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X7)|~v1_xreal_0(k1_seq_1(X6,k1_seq_1(X5,esk3_2(X7,k1_partfun1(X1,X2,X3,X4,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,k19_sin_cos)))))))).#trainneg
cnf(i_0_1845, negated_conjecture, (k7_xcmplx_0(k4_xcmplx_0(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0))),esk2_0)!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)|~v1_xcmplx_0(esk2_0)|~v1_xreal_0(esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_1855, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),k19_sin_cos),X2)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos))|~r1_fdiff_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),k19_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1854, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),k19_sin_cos),X2)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos))|~r1_fdiff_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),k19_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k1_seq_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),k19_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_2095, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1941, plain, (r1_xxreal_0(X1,X1)|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_1931, plain, (k7_xcmplx_0(k6_numbers,np__1)=k6_numbers)).#trainneg
cnf(i_0_350, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_1622, plain, (v1_membered(k19_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1623, plain, (v1_membered(k3_taylor_1)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1877, plain, (v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_364, plain, (v1_xcmplx_0(k8_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_365, plain, (v1_int_1(k8_real_1(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1884, plain, (r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k19_sin_cos),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k19_sin_cos)))|~m1_subset_1(esk3_2(X1,k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k19_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1999, plain, (k1_real_1(k4_xcmplx_0(X1))=k7_xcmplx_0(X1,np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1874, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1992, plain, (k7_xcmplx_0(X1,np__1)=X1|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1940, plain, (r1_xxreal_0(k6_numbers,X1)|r1_xxreal_0(X1,k6_numbers)|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_1334, plain, (a_0_0_fdiff_7=k6_numbers|r2_hidden(np__1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_366, plain, (v7_ordinal1(k8_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1473, plain, (v1_int_1(X1)|~r2_hidden(X1,a_0_0_fdiff_7)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_937, negated_conjecture, (v1_xreal_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_986, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2037, negated_conjecture, (r1_xxreal_0(esk1_0,esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_367, plain, (v1_xxreal_0(k8_real_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_371, plain, (k4_xcmplx_0(k7_xcmplx_0(X1,np__1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1594, plain, (k1_real_1(k7_xcmplx_0(k1_real_1(X1),X2))=k7_xcmplx_0(X1,X2)|~m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1912, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k19_sin_cos),k19_sin_cos),X6)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k19_sin_cos))|~r1_fdiff_1(X5,esk3_2(X6,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k19_sin_cos),k19_sin_cos)))|~v1_funct_1(X5)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X6,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~v1_xreal_0(k1_seq_1(X5,esk3_2(X6,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k19_sin_cos),k19_sin_cos)))))).#trainneg
cnf(i_0_1773, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,k3_taylor_1)),X7)|r1_xxreal_0(k1_seq_1(X6,k1_seq_1(X5,esk3_2(X7,k1_partfun1(X1,X2,X3,X4,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,k3_taylor_1))))),k6_numbers)|~r1_fdiff_1(X6,k1_seq_1(X5,esk3_2(X7,k1_partfun1(X1,X2,X3,X4,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,k3_taylor_1)))))|~r1_fdiff_1(X5,esk3_2(X7,k1_partfun1(X1,X2,X3,X4,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,k3_taylor_1))))|~v1_funct_1(X5)|~v1_funct_1(X6)|~r1_tarski(X7,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,k3_taylor_1))))|~m1_subset_1(k1_seq_1(X5,esk3_2(X7,k1_partfun1(X1,X2,X3,X4,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,k3_taylor_1)))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X7)|~m1_subset_1(k1_seq_1(X6,k1_seq_1(X5,esk3_2(X7,k1_partfun1(X1,X2,X3,X4,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X6,k3_taylor_1))))),k1_numbers))).#trainneg
cnf(i_0_1993, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),np__1)=k1_real_1(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2106, negated_conjecture, (r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2153, plain, (r1_xxreal_0(X1,X1)|r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_373, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1))=k7_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1089, plain, (v1_xcmplx_0(k24_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1095, plain, (v1_xcmplx_0(k16_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2166, plain, (r1_xxreal_0(X1,X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_375, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,X1),X2)=k7_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2123, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),X1)|r1_xxreal_0(esk3_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos)),k6_numbers)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_2161, negated_conjecture, (k7_xcmplx_0(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk2_0)|~v1_xcmplx_0(esk2_0)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~v1_xreal_0(esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~r2_hidden(esk2_0,esk1_0)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk2_0,k1_numbers))).#trainneg
cnf(i_0_2366, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k19_sin_cos),k19_sin_cos),X6)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k19_sin_cos))|~r1_fdiff_1(X5,esk3_2(X6,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k19_sin_cos),k19_sin_cos)))|~v1_funct_1(X5)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X6,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~m1_subset_1(k1_seq_1(X5,esk3_2(X6,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k19_sin_cos),k19_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_2367, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k19_sin_cos),k19_sin_cos),X6)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k19_sin_cos))|~r1_fdiff_1(X5,esk3_2(X6,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k19_sin_cos),k19_sin_cos)))|~v1_funct_1(X5)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X6,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~v3_valued_0(X5)|~v1_relat_1(X5))).#trainneg
cnf(i_0_2212, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),np__1)=k1_real_1(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1402, plain, (r1_xxreal_0(k5_numbers,k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1673, plain, (a_0_0_fdiff_7=k6_numbers|r2_hidden(np__2,a_0_0_fdiff_7))).#trainneg
cnf(i_0_377, plain, (k3_xcmplx_0(X1,np__1)=k7_xcmplx_0(X1,np__1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1837, plain, (r1_fdiff_1(k3_taylor_1,X1)|~r2_hidden(X1,a_0_0_fdiff_7)|~v1_xreal_0(k1_xxreal_0))).#trainneg
cnf(i_0_1096, plain, (v1_xcmplx_0(k19_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1149, plain, (v1_xcmplx_0(k3_taylor_1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_378, plain, (v1_xboole_0(k1_zfmisc_1(X1))|r2_hidden(k2_relset_1(X1,X2),k1_zfmisc_1(X1))|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_379, plain, (v1_xreal_0(k2_relset_1(X1,X2))|~v3_membered(k1_zfmisc_1(X1))|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_2122, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos),k19_sin_cos),X1)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xreal_0(esk3_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos),k19_sin_cos))))).#trainneg
cnf(i_0_2193, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=k7_xcmplx_0(X1,np__1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2120, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos),k19_sin_cos),X3)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos))|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~r1_fdiff_1(X2,k1_seq_1(X1,esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos),k19_sin_cos))))|~r1_fdiff_1(X1,esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos),k19_sin_cos)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2311, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=k7_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2474, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_380, plain, (v1_xcmplx_0(k2_relset_1(X1,X2))|~v1_membered(k1_zfmisc_1(X1))|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_2543, plain, (r1_xxreal_0(k5_numbers,k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1201, plain, (v1_int_1(k24_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1338, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|r1_xxreal_0(esk1_0,X1)|~v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_381, plain, (v1_int_1(k2_relset_1(X1,X2))|~v5_membered(k1_zfmisc_1(X1))|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_382, plain, (v7_ordinal1(k2_relset_1(X1,X2))|~v6_membered(k1_zfmisc_1(X1))|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_2424, plain, (r1_xxreal_0(esk3_2(X1,X2),k6_numbers)|r1_xxreal_0(esk3_2(X1,X2),esk3_2(X1,X2))|r2_fdiff_1(X2,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_2415, negated_conjecture, (r1_fdiff_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos),X1)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2421, negated_conjecture, (r1_xxreal_0(esk3_2(esk1_0,X1),k6_numbers)|r1_xxreal_0(esk3_2(esk1_0,X1),esk3_2(esk1_0,X1))|r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_2406, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~r2_hidden(esk3_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),esk1_0))).#trainneg
cnf(i_0_2513, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos),k19_sin_cos),X5)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos))|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X5,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_xreal_0(esk3_2(X5,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos),k19_sin_cos))))).#trainneg
cnf(i_0_1689, plain, (r1_xxreal_0(a_0_0_fdiff_7,a_0_0_fdiff_7)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2553, plain, (k7_xcmplx_0(np__1,np__1)=np__1)).#trainneg
cnf(i_0_383, plain, (v1_xxreal_0(k2_relset_1(X1,X2))|~v2_membered(k1_zfmisc_1(X1))|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1213, plain, (v1_int_1(k16_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1214, plain, (v1_int_1(k19_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_918, plain, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|r2_hidden(k1_relset_1(k1_numbers,k16_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_927, plain, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|r2_hidden(k1_relset_1(k1_numbers,k19_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_384, plain, (v1_relat_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_385, plain, (v4_relat_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_2638, plain, (r2_fdiff_1(k1_partfun1(X4,X5,X6,X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos),k19_sin_cos),X3)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos))|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~r1_fdiff_1(X2,k1_seq_1(X1,esk3_2(X3,k1_partfun1(X4,X5,X6,X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos),k19_sin_cos))))|~r1_fdiff_1(X1,esk3_2(X3,k1_partfun1(X4,X5,X6,X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos),k19_sin_cos)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(X4,X5,X6,X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X3,k1_partfun1(X4,X5,X6,X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_2505, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(X7,X8,X9,X10,X5,k19_sin_cos),k19_sin_cos),X6)|~v3_valued_0(k1_partfun1(X7,X8,X9,X10,X5,k19_sin_cos))|~v3_valued_0(X5)|~r1_fdiff_1(X5,esk3_2(X6,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X7,X8,X9,X10,X5,k19_sin_cos),k19_sin_cos)))|~v1_funct_1(X5)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X7,X8,X9,X10,X5,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X6,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X7,X8,X9,X10,X5,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(k1_partfun1(X7,X8,X9,X10,X5,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_2408, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),X2)|~r1_fdiff_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~r2_hidden(k1_seq_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))),esk1_0))).#trainneg
cnf(i_0_2353, plain, (k1_real_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2349, plain, (k1_real_1(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k1_real_1(X1),X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2124, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),X1)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~r2_hidden(esk3_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos)),esk1_0)|~m1_subset_1(esk1_0,k1_zfmisc_1(a_0_0_fdiff_7)))).#trainneg
cnf(i_0_2295, plain, (k1_real_1(k4_xcmplx_0(X1))=k7_xcmplx_0(X1,np__1)|~m1_subset_1(k7_xcmplx_0(X1,np__1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1258, plain, (v7_ordinal1(X1)|~v6_membered(k1_numbers)|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_386, plain, (v3_valued_0(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v3_membered(X2)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_1202, plain, (v3_membered(k24_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1252, plain, (v1_int_1(k3_taylor_1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1254, plain, (v1_xreal_0(a_0_0_fdiff_7)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_920, plain, (v1_xcmplx_0(k1_relset_1(k1_numbers,k16_sin_cos))|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_921, plain, (v1_int_1(k1_relset_1(k1_numbers,k16_sin_cos))|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_387, plain, (v1_valued_0(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_membered(X2)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_388, plain, (v5_relat_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X2)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_2696, negated_conjecture, (r2_hidden(esk2_0,k1_numbers))).#trainneg
cnf(i_0_2896, plain, (k7_xcmplx_0(k1_real_1(np__1),np__1)=k1_real_1(np__1))).#trainneg
cnf(i_0_2769, plain, (k1_real_1(np__1)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_2738, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),X5)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos))|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~r2_hidden(esk3_2(X5,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos)),esk1_0)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2700, negated_conjecture, (v7_ordinal1(esk2_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_389, plain, (v3_membered(k2_relset_1(X1,X2))|~v3_membered(X1)|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_2692, negated_conjecture, (~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_2699, negated_conjecture, (v1_int_1(esk2_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_2956, plain, (v1_xcmplx_0(k1_relset_1(k1_numbers,k16_sin_cos))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_390, plain, (v1_membered(k2_relset_1(X1,X2))|~v1_membered(X1)|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_393, plain, (m1_subset_1(k10_xtuple_0(X2),k1_zfmisc_1(X1))|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_2826, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(X1),X2))=k7_xcmplx_0(X1,X2)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2846, plain, (k1_real_1(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2498, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(X7,X8,X9,X10,X5,k19_sin_cos),k19_sin_cos),X6)|~v3_valued_0(k1_partfun1(X7,X8,X9,X10,X5,k19_sin_cos))|~r1_fdiff_1(X5,esk3_2(X6,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X7,X8,X9,X10,X5,k19_sin_cos),k19_sin_cos)))|~v1_funct_1(X5)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X7,X8,X9,X10,X5,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(k1_seq_1(X5,esk3_2(X6,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X7,X8,X9,X10,X5,k19_sin_cos),k19_sin_cos))),k1_numbers)|~m1_subset_1(esk3_2(X6,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X7,X8,X9,X10,X5,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(k1_partfun1(X7,X8,X9,X10,X5,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_2636, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X4,X5,X6,X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos),k19_sin_cos),X3)|~v3_valued_0(k1_partfun1(X4,X5,X6,X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos))|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~r1_fdiff_1(X2,k1_seq_1(X1,esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X4,X5,X6,X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos),k19_sin_cos))))|~r1_fdiff_1(X1,esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X4,X5,X6,X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos),k19_sin_cos)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X4,X5,X6,X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X4,X5,X6,X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_2411, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k19_sin_cos),X6)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))|~r1_fdiff_1(X5,esk3_2(X6,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k19_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~v1_funct_1(X5)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k19_sin_cos)))|~m1_subset_1(esk3_2(X6,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k19_sin_cos)),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~r2_hidden(k1_seq_1(X5,esk3_2(X6,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k19_sin_cos))),esk1_0))).#trainneg
cnf(i_0_2585, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k19_sin_cos),k19_sin_cos),X1)|~v3_valued_0(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k19_sin_cos))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1211, plain, (v1_int_1(k4_xcmplx_0(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2693, negated_conjecture, (m1_subset_1(esk2_0,esk1_0))).#trainneg
cnf(i_0_1890, plain, (k7_xcmplx_0(k6_numbers,X1)=k6_numbers|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers))).#trainneg
cnf(i_0_394, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_2701, negated_conjecture, (v1_xxreal_0(esk2_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_397, plain, (r2_relset_1(X1,X2,X3,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v5_relat_1(X4,k2_zfmisc_1(X1,X2))|~v1_relat_1(X4))).#trainneg
cnf(i_0_398, plain, (r2_relset_1(k1_numbers,k1_numbers,X1,X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2125, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos),k19_sin_cos),X2)|~v3_valued_0(k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos))|~r1_fdiff_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos),k19_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(k1_seq_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos),k19_sin_cos))),k1_numbers)|~m1_subset_1(esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_2510, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k19_sin_cos),k19_sin_cos),X7)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k19_sin_cos))|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6))|~r1_tarski(X7,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X7,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X7)|~r1_fdiff_1(X6,k1_seq_1(X5,esk3_2(X7,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k19_sin_cos),k19_sin_cos))))|~r1_fdiff_1(X5,esk3_2(X7,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k19_sin_cos),k19_sin_cos)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1845, negated_conjecture, (k7_xcmplx_0(k4_xcmplx_0(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0))),esk2_0)!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)|~v1_xcmplx_0(esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_1608, negated_conjecture, (k7_xcmplx_0(k1_real_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0))),esk2_0)!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)|~v1_xcmplx_0(esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_3211, negated_conjecture, (v1_xxreal_0(esk2_0))).#trainneg
cnf(i_0_524, negated_conjecture, (k4_xcmplx_0(k7_xcmplx_0(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_327, negated_conjecture, (k1_real_1(k7_xcmplx_0(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_2985, plain, (m1_subset_1(k4_xcmplx_0(np__1),k1_numbers))).#trainneg
cnf(i_0_2984, plain, (k1_real_1(k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_3210, plain, (m1_subset_1(X1,k1_zfmisc_1(a_0_0_fdiff_7))|~r2_fdiff_1(k3_taylor_1,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_2987, plain, (v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_3190, negated_conjecture, (v7_ordinal1(esk2_0)|~v6_membered(esk1_0))).#trainneg
cnf(i_0_1847, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)!=k7_xcmplx_0(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0)|~v1_xcmplx_0(esk2_0)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_3169, plain, (v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_3005, negated_conjecture, (v1_int_1(esk2_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_3032, plain, (m1_subset_1(k1_numbers,k1_zfmisc_1(X1))|~v5_relat_1(k3_taylor_1,X1))).#trainneg
cnf(i_0_3082, plain, (k4_xcmplx_0(k1_real_1(X1))=k7_xcmplx_0(X1,np__1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3162, plain, (r2_fdiff_1(k1_partfun1(X6,X7,X8,X9,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos),k19_sin_cos),X5)|~v3_valued_0(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos))|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X6,X7,X8,X9,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X5,k1_partfun1(X6,X7,X8,X9,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X8,X9)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X6,X7))))).#trainneg
cnf(i_0_2847, plain, (k1_real_1(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2820, negated_conjecture, (r2_fdiff_1(k1_partfun1(X3,X4,X5,X6,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),X2)|~r1_fdiff_1(X1,esk3_2(X2,k1_partfun1(X3,X4,X5,X6,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~v1_funct_1(X1)|~r2_hidden(k1_seq_1(X1,esk3_2(X2,k1_partfun1(X3,X4,X5,X6,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))),esk1_0)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_2736, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k19_sin_cos),X9)|~v3_valued_0(k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos))|~r1_tarski(X9,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X9,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X9)|~r2_hidden(esk3_2(X9,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k19_sin_cos)),esk1_0)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_2493, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos),k19_sin_cos),X1)|r1_xxreal_0(esk3_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos),k19_sin_cos)),k6_numbers)|~v3_valued_0(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2991, plain, (v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_2161, negated_conjecture, (k7_xcmplx_0(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk2_0)|~v1_xcmplx_0(esk2_0)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_2763, plain, (r1_xxreal_0(a_0_0_fdiff_7,a_0_0_fdiff_7)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3189, negated_conjecture, (v1_int_1(esk2_0)|~v5_membered(esk1_0))).#trainneg
cnf(i_0_929, plain, (v1_xcmplx_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_930, plain, (v1_int_1(k1_relset_1(k1_numbers,k19_sin_cos))|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_409, plain, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|r2_hidden(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_410, plain, (v1_xreal_0(k2_fdiff_1(X1,X2))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2360, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(X7,X8,X9,X10,X5,k19_sin_cos),k19_sin_cos),X6)|~v3_valued_0(k1_partfun1(X7,X8,X9,X10,X5,k19_sin_cos))|~v1_xreal_0(k1_seq_1(X5,esk3_2(X6,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X7,X8,X9,X10,X5,k19_sin_cos),k19_sin_cos))))|~r1_fdiff_1(X5,esk3_2(X6,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X7,X8,X9,X10,X5,k19_sin_cos),k19_sin_cos)))|~v1_funct_1(X5)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X7,X8,X9,X10,X5,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X6,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X7,X8,X9,X10,X5,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(k1_partfun1(X7,X8,X9,X10,X5,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_2325, plain, (k4_xcmplx_0(k7_xcmplx_0(k1_real_1(X1),X2))=k7_xcmplx_0(X1,X2)|~m1_subset_1(k7_xcmplx_0(k1_real_1(X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2988, plain, (v1_int_1(k4_xcmplx_0(np__1))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_499, plain, (k9_xtuple_0(k24_sin_cos)=k1_relset_1(k1_numbers,k24_sin_cos)|~v1_relat_1(k24_sin_cos))).#trainneg
cnf(i_0_411, plain, (v1_xcmplx_0(k2_fdiff_1(X1,X2))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1253, plain, (v3_membered(k16_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1400, plain, (v1_xxreal_0(k24_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1406, plain, (v1_xxreal_0(k16_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1840, plain, (v1_xreal_0(k1_relset_1(k1_numbers,k16_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1876, plain, (v1_xreal_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_412, plain, (v1_int_1(k2_fdiff_1(X1,X2))|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_413, plain, (v7_ordinal1(k2_fdiff_1(X1,X2))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3117, plain, (k1_real_1(k1_fdiff_1(k3_taylor_1,X1))=k7_xcmplx_0(k4_xcmplx_0(np__1),X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(k4_xcmplx_0(np__1),k1_numbers)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_3248, negated_conjecture, (r1_xxreal_0(k6_numbers,esk2_0))).#trainneg
cnf(i_0_3249, negated_conjecture, (r1_xxreal_0(esk2_0,esk2_0))).#trainneg
cnf(i_0_3421, negated_conjecture, (v1_xboole_0(a_0_0_fdiff_7)|r2_hidden(esk2_0,a_0_0_fdiff_7))).#trainneg
cnf(i_0_3425, negated_conjecture, (v7_ordinal1(esk2_0)|~v6_membered(a_0_0_fdiff_7))).#trainneg
cnf(i_0_2896, plain, (k7_xcmplx_0(k4_xcmplx_0(np__1),np__1)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_414, plain, (v1_xxreal_0(k2_fdiff_1(X1,X2))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1482, plain, (v3_membered(k19_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3292, negated_conjecture, (v1_int_1(esk2_0)|~v6_membered(esk1_0))).#trainneg
cnf(i_0_3424, negated_conjecture, (v1_int_1(esk2_0)|~v5_membered(a_0_0_fdiff_7))).#trainneg
cnf(i_0_3429, plain, (v1_xcmplx_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_415, plain, (v3_membered(k2_fdiff_1(X1,X2))|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_416, plain, (v1_membered(k2_fdiff_1(X1,X2))|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3247, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|r1_xxreal_0(esk2_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_2116, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos),k19_sin_cos),X2)|~v3_valued_0(k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos))|~v3_valued_0(X1)|~r1_fdiff_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos),k19_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_2254, plain, (r2_hidden(X1,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2740, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos)),X6)|~r1_fdiff_1(X1,esk3_2(X6,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos))))|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos))|~v1_funct_1(X1)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos))))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~r2_hidden(k1_seq_1(X1,esk3_2(X6,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos)))),esk1_0)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_602, plain, (v1_int_1(X1)|~v1_int_1(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_373, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1))=k7_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_417, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k2_fdiff_1(X2,X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1558, plain, (v3_membered(k3_taylor_1)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1407, plain, (v1_xxreal_0(k19_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1472, plain, (v1_xxreal_0(k3_taylor_1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1693, negated_conjecture, (r1_xxreal_0(k6_numbers,esk1_0)|r1_xxreal_0(esk1_0,k6_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2165, plain, (r1_xxreal_0(k6_numbers,X1)|r1_xxreal_0(X1,k6_numbers)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_418, plain, (v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_419, plain, (v4_relat_1(k2_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2409, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k19_sin_cos),X2)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))|~r1_fdiff_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k19_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k19_sin_cos)))|~m1_subset_1(esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k19_sin_cos)),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~r2_hidden(k1_seq_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k19_sin_cos))),esk1_0))).#trainneg
cnf(i_0_2523, plain, (k1_real_1(k4_xcmplx_0(X1))=k4_xcmplx_0(k1_real_1(X1))|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2129, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos),k19_sin_cos),X3)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos))|~r1_fdiff_1(k2_fdiff_1(X1,X2),esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos),k19_sin_cos)))|~v1_funct_1(k2_fdiff_1(X1,X2))|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos),k19_sin_cos)))|~m1_subset_1(k1_fdiff_1(X1,esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos),k19_sin_cos))),k1_numbers)|~m1_subset_1(esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(esk3_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos),k19_sin_cos)),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2490, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos))|~r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos)),k1_numbers)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))).#trainneg
cnf(i_0_2814, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k19_sin_cos),k19_sin_cos),X9)|~v3_valued_0(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k19_sin_cos))|~r1_tarski(X9,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X9,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X9,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X9)|~v1_xreal_0(esk3_2(X9,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k19_sin_cos),k19_sin_cos))))).#trainneg
cnf(i_0_2989, plain, (v7_ordinal1(k4_xcmplx_0(np__1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2990, plain, (v1_xxreal_0(k4_xcmplx_0(np__1))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_3501, negated_conjecture, (k1_seq_1(k2_fdiff_1(k3_relat_1(k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)!=k7_xcmplx_0(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~r2_fdiff_1(k3_relat_1(k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_420, plain, (v3_valued_0(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2035, plain, (v1_membered(k24_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1474, negated_conjecture, (v1_xcmplx_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1560, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1699, plain, (r1_xxreal_0(k6_numbers,k1_real_1(X1))|r1_xxreal_0(k1_real_1(X1),k6_numbers)|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_421, plain, (v1_valued_0(k2_fdiff_1(X1,X2))|~v1_membered(k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_422, plain, (v5_relat_1(k2_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2405, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),X5)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~r2_hidden(esk3_2(X5,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos)),esk1_0))).#trainneg
cnf(i_0_2815, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k19_sin_cos),X9)|r1_xxreal_0(esk3_2(X9,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k19_sin_cos)),k6_numbers)|~v3_valued_0(k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos))|~r1_tarski(X9,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X9,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X9,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X9))).#trainneg
cnf(i_0_2566, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),X5)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_xreal_0(k1_xxreal_0)|~r2_hidden(esk3_2(X5,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos)),a_0_0_fdiff_7))).#trainneg
cnf(i_0_3380, plain, (k7_xcmplx_0(k1_real_1(X1),X2)=k7_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3602, plain, (v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_3536, negated_conjecture, (v1_int_1(esk2_0)|~v6_membered(a_0_0_fdiff_7))).#trainneg
cnf(i_0_1875, plain, (v1_xcmplx_0(a_0_0_fdiff_7)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3669, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|r1_xxreal_0(k6_numbers,esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_878, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_425, plain, (v1_xreal_0(k1_fdiff_1(X1,X2))|~v1_xreal_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_426, plain, (v1_xcmplx_0(k1_fdiff_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3685, plain, (k1_real_1(k4_xcmplx_0(k1_real_1(X1)))=k4_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3735, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),esk1_0)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos)))))).#trainneg
cnf(i_0_3534, negated_conjecture, (a_0_0_fdiff_7=k6_numbers|r2_hidden(esk2_0,a_0_0_fdiff_7))).#trainneg
cnf(i_0_427, plain, (v1_int_1(k1_fdiff_1(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2036, plain, (v1_membered(k16_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2162, plain, (v1_membered(k19_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2163, plain, (v1_membered(k3_taylor_1)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_606, plain, (v1_int_1(k1_real_1(X1))|~v1_int_1(X1)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2436, plain, (r1_xxreal_0(k1_real_1(X1),k6_numbers)|r1_xxreal_0(k1_real_1(X1),k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_428, plain, (v7_ordinal1(k1_fdiff_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_429, plain, (v1_xxreal_0(k1_fdiff_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3687, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_real_1(X1)))=k4_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3599, negated_conjecture, (r2_fdiff_1(k1_partfun1(X7,X8,X9,X10,X1,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos)),X6)|~r1_fdiff_1(X1,esk3_2(X6,k1_partfun1(X7,X8,X9,X10,X1,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos))))|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos))|~v1_funct_1(X1)|~r2_hidden(k1_seq_1(X1,esk3_2(X6,k1_partfun1(X7,X8,X9,X10,X1,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos)))),esk1_0)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(X7,X8,X9,X10,X1,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_2514, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),X5)|r1_xxreal_0(esk3_2(X5,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos)),k6_numbers)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X5,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5))).#trainneg
cnf(i_0_3514, plain, (k7_xcmplx_0(k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_3710, plain, (k4_xcmplx_0(k1_real_1(k4_xcmplx_0(X1)))=k1_real_1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4128, plain, (k1_real_1(k1_real_1(X1))=k1_real_1(k4_xcmplx_0(X1))|~m1_subset_1(k1_real_1(k4_xcmplx_0(X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3754, negated_conjecture, (k7_xcmplx_0(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0)!=k1_fdiff_1(k3_relat_1(k3_taylor_1,k19_sin_cos),esk2_0)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~r2_fdiff_1(k3_relat_1(k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(k3_relat_1(k3_taylor_1,k19_sin_cos))|~m1_subset_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_434, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_1679, negated_conjecture, (r1_xxreal_0(esk6_0,esk1_0)|r1_xxreal_0(esk1_0,esk6_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_435, plain, (r2_relset_1(k1_numbers,k1_numbers,k3_taylor_1,k2_funct_1(k24_sin_cos))|~v2_funct_1(k24_sin_cos))).#trainneg
cnf(i_0_437, plain, (v1_funct_1(k2_funct_1(X3))|~v2_funct_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4099, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1)))=np__1|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_4117, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))=k7_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_438, plain, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|r2_hidden(k2_partfun2(X2,X1,X3),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v2_funct_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1))))).#trainneg
cnf(i_0_1681, plain, (r1_xxreal_0(esk6_0,k5_numbers)|r1_xxreal_0(k5_numbers,esk6_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_992, plain, (v1_xreal_0(k24_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1018, plain, (v1_xreal_0(k16_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_439, plain, (v1_xreal_0(k2_partfun2(X1,X2,X3))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v2_funct_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_440, plain, (v1_xcmplx_0(k2_partfun2(X1,X2,X3))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v2_funct_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2260, plain, (r1_xxreal_0(k6_numbers,X1)|r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2403, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos),X1)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos))|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~r2_hidden(esk3_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos)),esk1_0))).#trainneg
cnf(i_0_2589, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos),k19_sin_cos),k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos),k19_sin_cos)))|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos))|~m1_subset_1(esk3_2(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos),k19_sin_cos)),k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos),k19_sin_cos)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos),k19_sin_cos))))).#trainneg
cnf(i_0_2590, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos),k19_sin_cos),X1)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos))|~m1_subset_1(esk3_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos),k19_sin_cos)))))).#trainneg
cnf(i_0_3250, negated_conjecture, (r1_xxreal_0(esk6_0,esk2_0)|r1_xxreal_0(esk2_0,esk6_0))).#trainneg
cnf(i_0_441, plain, (v1_int_1(k2_partfun2(X1,X2,X3))|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v2_funct_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1690, plain, (r1_xxreal_0(esk6_0,a_0_0_fdiff_7)|r1_xxreal_0(a_0_0_fdiff_7,esk6_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1024, plain, (v1_xreal_0(k19_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1025, plain, (v1_xreal_0(k3_taylor_1)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_442, plain, (v7_ordinal1(k2_partfun2(X1,X2,X3))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v2_funct_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_443, plain, (v1_xxreal_0(k2_partfun2(X1,X2,X3))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v2_funct_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4113, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))=k7_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_3589, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos),X9)|~v3_valued_0(k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k19_sin_cos))|~v3_valued_0(k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos))|~r1_tarski(X9,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X9,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X9,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X9)|~r2_hidden(esk3_2(X9,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos)),esk1_0)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_4391, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos),esk1_0)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos))|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4388, negated_conjecture, (r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos),X1)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos))|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~r2_hidden(esk3_2(X1,k1_partfun1(X2,X3,X4,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos)),esk1_0)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_4457, plain, (k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1)=k7_xcmplx_0(np__1,X1)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_4273, plain, (v1_funct_1(k2_funct_1(k3_taylor_1)))).#trainneg
cnf(i_0_4265, plain, (v1_funct_1(k2_funct_1(k24_sin_cos))|~v2_funct_1(k24_sin_cos))).#trainneg
cnf(i_0_4236, negated_conjecture, (k1_fdiff_1(k3_relat_1(k3_taylor_1,k19_sin_cos),esk2_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk2_0)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~v1_funct_1(k3_relat_1(k3_taylor_1,k19_sin_cos))|~r2_fdiff_1(k3_relat_1(k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_xcmplx_0(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)))|~m2_subset_1(esk2_0,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0))|~r1_fdiff_1(k3_taylor_1,esk2_0)|~m1_subset_1(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_444, plain, (v3_membered(k2_partfun2(X1,X2,X3))|~v3_membered(k2_zfmisc_1(X2,X1))|~v2_funct_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4259, negated_conjecture, (r1_xxreal_0(esk1_0,esk6_0)|r1_xxreal_0(esk6_0,esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4266, plain, (v1_funct_1(k2_funct_1(k16_sin_cos))|~v2_funct_1(k16_sin_cos))).#trainneg
cnf(i_0_4272, plain, (v1_funct_1(k2_funct_1(k19_sin_cos))|~v2_funct_1(k19_sin_cos))).#trainneg
cnf(i_0_445, plain, (v1_membered(k2_partfun2(X1,X2,X3))|~v1_membered(k2_zfmisc_1(X2,X1))|~v2_funct_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_446, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X1,k2_partfun2(X3,X2,X4))|~v2_funct_1(X4)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_4467, negated_conjecture, (r2_fdiff_1(k1_partfun1(X10,X11,X12,X13,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos),X9)|~v3_valued_0(k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k19_sin_cos))|~v3_valued_0(k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos))|~r2_hidden(esk3_2(X9,k1_partfun1(X10,X11,X12,X13,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos)),esk1_0)|~r1_tarski(X9,k1_relset_1(k1_numbers,k1_partfun1(X10,X11,X12,X13,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos)))|~m1_subset_1(k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X9,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X9)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k19_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X12,X13)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k3_taylor_1,k19_sin_cos),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X10,X11))))).#trainneg
cnf(i_0_4434, plain, (k4_xcmplx_0(k7_xcmplx_0(k7_xcmplx_0(np__1,X1),X2))=k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_2516, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3930, plain, (k1_real_1(k4_xcmplx_0(X1))=k4_xcmplx_0(k1_real_1(X1))|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(k4_xcmplx_0(k1_real_1(X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2515, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),X5)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos)))|~m1_subset_1(esk3_2(X5,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos)),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~r2_hidden(esk3_2(X5,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos)),esk1_0)|~m1_subset_1(esk1_0,k1_zfmisc_1(a_0_0_fdiff_7)))).#trainneg
cnf(i_0_3841, plain, (r1_xxreal_0(X1,X1)|~v2_membered(k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_447, plain, (v1_relat_1(k2_partfun2(X1,X2,X3))|~v2_funct_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4355, plain, (r1_xxreal_0(k5_numbers,esk6_0)|r1_xxreal_0(esk6_0,k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2472, plain, (v1_xcmplx_0(k24_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1881, plain, (v1_int_1(k10_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_448, plain, (v4_relat_1(k2_partfun2(X1,X2,X3),X2)|~v2_funct_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_449, plain, (v3_valued_0(k2_partfun2(X1,X2,X3))|~v3_membered(X1)|~v2_funct_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2522, plain, (k4_xcmplx_0(k1_real_1(X1))=k4_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_4153, plain, (k1_real_1(k1_real_1(X1))=k4_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4614, plain, (k1_real_1(k4_xcmplx_0(X1))=k4_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers))).#trainneg
cnf(i_0_4748, plain, (k1_real_1(k4_xcmplx_0(X1))=k4_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4622, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4481, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(X6,X7,X8,X9,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos),X5)|~v3_valued_0(k1_partfun1(X6,X7,X8,X9,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos))|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~v1_funct_1(k1_partfun1(X6,X7,X8,X9,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~r2_hidden(esk3_2(X5,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X6,X7,X8,X9,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos)),esk1_0)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X6,X7,X8,X9,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos)))|~m1_subset_1(k1_partfun1(X6,X7,X8,X9,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(X6,X7,X8,X9,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X8,X9)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X6,X7))))).#trainneg
cnf(i_0_3564, negated_conjecture, (r1_xxreal_0(esk2_0,np__2)|r1_xxreal_0(np__2,esk2_0))).#trainneg
cnf(i_0_450, plain, (v1_valued_0(k2_partfun2(X1,X2,X3))|~v1_membered(X1)|~v2_funct_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4412, plain, (r1_xxreal_0(a_0_0_fdiff_7,esk6_0)|r1_xxreal_0(esk6_0,a_0_0_fdiff_7)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2473, plain, (v1_xcmplx_0(k16_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2577, plain, (v1_xcmplx_0(k19_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_451, plain, (v5_relat_1(k2_partfun2(X1,X2,X3),X1)|~v2_funct_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_452, plain, (r2_relset_1(X1,X2,X3,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v2_funct_1(X4)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X1))))).#trainneg
cnf(i_0_4517, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k1_partfun1(X9,X10,X11,X12,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos),esk1_0)|~v3_valued_0(k1_partfun1(X5,X6,X7,X8,k1_partfun1(X9,X10,X11,X12,k3_taylor_1,k19_sin_cos),k19_sin_cos))|~v3_valued_0(k1_partfun1(X9,X10,X11,X12,k3_taylor_1,k19_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k1_partfun1(X9,X10,X11,X12,k3_taylor_1,k19_sin_cos),k19_sin_cos))|~v1_funct_1(k1_partfun1(X9,X10,X11,X12,k3_taylor_1,k19_sin_cos))|~r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k1_partfun1(X9,X10,X11,X12,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos)))|~m1_subset_1(k1_partfun1(X5,X6,X7,X8,k1_partfun1(X9,X10,X11,X12,k3_taylor_1,k19_sin_cos),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(X5,X6,X7,X8,k1_partfun1(X9,X10,X11,X12,k3_taylor_1,k19_sin_cos),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k1_partfun1(X9,X10,X11,X12,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X11,X12)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k1_partfun1(X9,X10,X11,X12,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,k1_partfun1(X9,X10,X11,X12,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4479, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X6,X7,X8,X9,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos),X5)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X6,X7,X8,X9,k3_taylor_1,k19_sin_cos),k19_sin_cos))|~v3_valued_0(k1_partfun1(X6,X7,X8,X9,k3_taylor_1,k19_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X6,X7,X8,X9,k3_taylor_1,k19_sin_cos),k19_sin_cos))|~v1_funct_1(k1_partfun1(X6,X7,X8,X9,k3_taylor_1,k19_sin_cos))|~r2_hidden(esk3_2(X5,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X6,X7,X8,X9,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos)),esk1_0)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X6,X7,X8,X9,k3_taylor_1,k19_sin_cos),k19_sin_cos),k19_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X6,X7,X8,X9,k3_taylor_1,k19_sin_cos),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X6,X7,X8,X9,k3_taylor_1,k19_sin_cos),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k1_partfun1(X6,X7,X8,X9,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X8,X9)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))).#trainneg
cnf(i_0_3675, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k19_sin_cos),X2)|~v3_valued_0(k1_partfun1(X3,X4,X5,X6,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))|~r1_fdiff_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k19_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~v1_funct_1(X1)|~r2_hidden(k1_seq_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k19_sin_cos))),esk1_0)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k19_sin_cos)))|~m1_subset_1(esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)),k19_sin_cos)),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_3340, plain, (k7_xcmplx_0(k1_real_1(X1),np__1)=k4_xcmplx_0(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3569, negated_conjecture, (r1_xxreal_0(esk2_0,np__1)|r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_453, plain, (m1_subset_1(k2_funct_1(X3),k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v2_funct_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1695, plain, (r1_xxreal_0(k6_numbers,k5_numbers)|r1_xxreal_0(k5_numbers,k6_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2578, plain, (v1_xcmplx_0(k3_taylor_1)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3506, plain, (v1_xxreal_0(k24_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_454, plain, (r2_relset_1(k1_numbers,k1_numbers,k2_partfun2(k1_numbers,k1_numbers,k24_sin_cos),k3_taylor_1)|~m1_subset_1(k2_partfun2(k1_numbers,k1_numbers,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_5027, plain, (v1_funct_1(k3_relat_1(X1,k3_taylor_1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_4951, plain, (k4_xcmplx_0(k7_xcmplx_0(X1,np__1))=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4787, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_real_1(X1)))=k1_real_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_458, plain, (k1_relset_1(k1_numbers,k2_fdiff_1(X1,X2))=X2|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_5005, plain, (r1_xxreal_0(k5_numbers,k6_numbers)|r1_xxreal_0(k6_numbers,k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3507, plain, (v1_xxreal_0(k16_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3667, plain, (v1_xxreal_0(k19_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_877, plain, (r1_xxreal_0(X1,k1_real_1(X2))|r1_xxreal_0(k1_real_1(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2164, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_459, plain, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|r2_hidden(k1_partfun1(X1,X3,X4,X2,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_460, plain, (v1_xreal_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3305, plain, (k1_real_1(k7_xcmplx_0(X1,np__1))=k1_real_1(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4104, plain, (k1_real_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))=k7_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_5034, plain, (v1_funct_1(k3_relat_1(k24_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_740, plain, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|r2_hidden(a_0_0_fdiff_7,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_461, plain, (v1_xcmplx_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1696, plain, (r1_xxreal_0(k6_numbers,a_0_0_fdiff_7)|r1_xxreal_0(a_0_0_fdiff_7,k6_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3668, plain, (v1_xxreal_0(k3_taylor_1)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5203, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X2)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5202, plain, (r1_xxreal_0(k1_real_1(X1),X2)|r1_xxreal_0(X2,k1_real_1(X1))|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_462, plain, (v1_int_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_463, plain, (v7_ordinal1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5259, plain, (k1_real_1(k1_real_1(X1))=k1_real_1(k4_xcmplx_0(X1))|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5305, plain, (k4_xcmplx_0(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)))=k7_xcmplx_0(np__1,X1)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_5266, plain, (k1_real_1(k1_real_1(X1))=k1_real_1(k4_xcmplx_0(X1))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5187, plain, (r2_fdiff_1(k2_fdiff_1(X1,X2),X3)|m1_subset_1(esk3_2(X3,k2_fdiff_1(X1,X2)),X3)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r1_tarski(X3,X2)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_5314, plain, (k1_zfmisc_1(k1_numbers)=k6_numbers|r2_hidden(a_0_0_fdiff_7,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_464, plain, (v1_xxreal_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5316, plain, (r1_xxreal_0(a_0_0_fdiff_7,k6_numbers)|r1_xxreal_0(k6_numbers,a_0_0_fdiff_7)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3256, plain, (r2_hidden(k4_xcmplx_0(np__1),k1_numbers))).#trainneg
cnf(i_0_2038, plain, (v1_int_1(k1_seq_1(X1,X2))|~v6_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_465, plain, (v3_membered(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v3_membered(k2_zfmisc_1(X1,X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_466, plain, (v1_membered(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_membered(k2_zfmisc_1(X1,X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5254, plain, (k1_real_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1)))=np__1|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_5117, plain, (k1_real_1(k1_real_1(X1))=k4_xcmplx_0(k1_real_1(X1))|~m1_subset_1(k4_xcmplx_0(k1_real_1(X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5035, plain, (v1_funct_1(k3_relat_1(k16_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_467, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X1,k1_partfun1(X2,X4,X5,X3,X6,X7))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X3)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X4))))).#trainneg
cnf(i_0_5041, plain, (v1_funct_1(k3_relat_1(k19_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_1701, plain, (r1_xxreal_0(k6_numbers,k1_relset_1(k1_numbers,k16_sin_cos))|r1_xxreal_0(k1_relset_1(k1_numbers,k16_sin_cos),k6_numbers)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_469, plain, (v4_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6),X1)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_470, plain, (v3_valued_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v3_membered(X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5507, plain, (k1_real_1(k1_real_1(k4_xcmplx_0(X1)))=k1_real_1(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5580, plain, (v1_xboole_0(X1)|r2_hidden(esk3_2(X1,k2_fdiff_1(X2,X3)),X1)|r2_fdiff_1(k2_fdiff_1(X2,X3),X1)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~r1_tarski(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_3673, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X3,X4,X5,X6,k3_taylor_1,k19_sin_cos)),k19_sin_cos),X2)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X3,X4,X5,X6,k3_taylor_1,k19_sin_cos)))|~r1_fdiff_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X3,X4,X5,X6,k3_taylor_1,k19_sin_cos)),k19_sin_cos)))|~v1_funct_1(k1_partfun1(X3,X4,X5,X6,k3_taylor_1,k19_sin_cos))|~v1_funct_1(X1)|~r2_hidden(k1_seq_1(X1,esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X3,X4,X5,X6,k3_taylor_1,k19_sin_cos)),k19_sin_cos))),esk1_0)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X3,X4,X5,X6,k3_taylor_1,k19_sin_cos)),k19_sin_cos)))|~m1_subset_1(esk3_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X3,X4,X5,X6,k3_taylor_1,k19_sin_cos)),k19_sin_cos)),k1_numbers)|~m1_subset_1(k1_partfun1(X3,X4,X5,X6,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_1250, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5049, plain, (v1_funct_1(k3_relat_1(k24_sin_cos,k3_taylor_1)))).#trainneg
cnf(i_0_5050, plain, (v1_funct_1(k3_relat_1(k16_sin_cos,k3_taylor_1)))).#trainneg
cnf(i_0_471, plain, (v1_valued_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_membered(X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2713, negated_conjecture, (r1_xxreal_0(esk1_0,np__2)|r1_xxreal_0(np__2,esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5056, plain, (v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1)))).#trainneg
cnf(i_0_5057, plain, (v1_funct_1(k3_relat_1(k3_taylor_1,k3_taylor_1)))).#trainneg
cnf(i_0_1702, plain, (r1_xxreal_0(k6_numbers,k1_relset_1(k1_numbers,k19_sin_cos))|r1_xxreal_0(k1_relset_1(k1_numbers,k19_sin_cos),k6_numbers)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_472, plain, (v5_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6),X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_473, plain, (r2_relset_1(X1,X2,X3,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,X4))))).#trainneg
cnf(i_0_5821, plain, (k1_real_1(k4_xcmplx_0(k1_real_1(X1)))=k1_real_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1964, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4236, negated_conjecture, (k1_fdiff_1(k3_relat_1(k3_taylor_1,k19_sin_cos),esk2_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk2_0)|~v1_xcmplx_0(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)))|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~m2_subset_1(esk2_0,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(k3_taylor_1,esk2_0)|~r2_fdiff_1(k3_relat_1(k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_474, plain, (m1_subset_1(k3_relat_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5863, negated_conjecture, (r1_xxreal_0(np__2,esk1_0)|r1_xxreal_0(esk1_0,np__2)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5736, plain, (r1_xxreal_0(k1_relset_1(k1_numbers,k16_sin_cos),k6_numbers)|r1_xxreal_0(k6_numbers,k1_relset_1(k1_numbers,k16_sin_cos))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5866, plain, (r1_xxreal_0(k1_relset_1(k1_numbers,k19_sin_cos),k6_numbers)|r1_xxreal_0(k6_numbers,k1_relset_1(k1_numbers,k19_sin_cos))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_479, plain, (v1_xreal_0(esk4_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_480, plain, (v1_xcmplx_0(esk4_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|~v1_membered(k1_numbers)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_5865, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k3_taylor_1))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5689, plain, (k4_xcmplx_0(k1_real_1(X1))=k4_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k4_xcmplx_0(k1_real_1(X1)),k1_numbers))).#trainneg
cnf(i_0_5856, negated_conjecture, (r2_fdiff_1(k1_partfun1(X7,X8,X9,X10,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos)),k19_sin_cos),X6)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos)))|~r1_fdiff_1(X1,esk3_2(X6,k1_partfun1(X7,X8,X9,X10,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos)),k19_sin_cos)))|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos))|~v1_funct_1(X1)|~r2_hidden(k1_seq_1(X1,esk3_2(X6,k1_partfun1(X7,X8,X9,X10,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos)),k19_sin_cos))),esk1_0)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(X7,X8,X9,X10,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos)),k19_sin_cos)))|~m1_subset_1(esk3_2(X6,k1_partfun1(X7,X8,X9,X10,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos)),k19_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3515, plain, (k1_real_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1))=k1_fdiff_1(k3_taylor_1,X1)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_3670, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3754, negated_conjecture, (k7_xcmplx_0(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0)!=k1_fdiff_1(k3_relat_1(k3_taylor_1,k19_sin_cos),esk2_0)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~r2_fdiff_1(k3_relat_1(k3_taylor_1,k19_sin_cos),esk1_0)|~m1_subset_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_481, plain, (v1_int_1(esk4_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|~v5_membered(k1_numbers)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2714, negated_conjecture, (r1_xxreal_0(esk1_0,k5_numbers)|r1_xxreal_0(k5_numbers,esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_671, negated_conjecture, (v1_xreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))))).#trainneg
cnf(i_0_672, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v1_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))))).#trainneg
cnf(i_0_5372, plain, (r1_xxreal_0(X1,k1_real_1(k1_real_1(X2)))|r1_xxreal_0(k1_real_1(k1_real_1(X2)),X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1405, plain, (r1_xxreal_0(k1_relset_1(k1_numbers,k16_sin_cos),k1_relset_1(k1_numbers,k16_sin_cos))|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_482, plain, (v7_ordinal1(esk4_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|~v6_membered(k1_numbers)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_483, plain, (v1_xxreal_0(esk4_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|~v2_membered(k1_numbers)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_5868, plain, (v3_valued_0(k1_partfun1(X1,X2,X3,k1_numbers,X4,X5))|~v1_relat_1(k1_partfun1(X1,X2,X3,k1_numbers,X4,X5))|~v1_funct_1(X5)|~v1_funct_1(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6049, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1))=k1_fdiff_1(k3_taylor_1,X1)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_512, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~v1_membered(X2))).#trainneg
cnf(i_0_486, plain, (m1_subset_1(esk4_3(X1,X2,X3),X2)|X3=k2_fdiff_1(X1,X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_673, negated_conjecture, (v1_int_1(esk1_0)|~v5_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))))).#trainneg
cnf(i_0_674, negated_conjecture, (v7_ordinal1(esk1_0)|~v6_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))))).#trainneg
cnf(i_0_1562, plain, (r1_xxreal_0(k1_relset_1(k1_numbers,k19_sin_cos),k1_relset_1(k1_numbers,k19_sin_cos))|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2727, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|r1_xxreal_0(X1,esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers))|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_492, plain, (m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3),k1_numbers)|~v1_xreal_0(k1_fdiff_1(X2,X3))|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(X2,X3)),k1_numbers)|~r1_fdiff_1(X1,k1_seq_1(X2,X3))|~r1_fdiff_1(X2,X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_6088, plain, (k1_real_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))=k1_fdiff_1(k3_taylor_1,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_6085, plain, (k1_real_1(k4_xcmplx_0(k7_xcmplx_0(np__1,X1)))=k1_fdiff_1(k3_taylor_1,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_6067, plain, (k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),np__1)=k1_real_1(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_3750, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_5725, negated_conjecture, (k7_xcmplx_0(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0)!=k1_fdiff_1(k3_relat_1(k3_taylor_1,k19_sin_cos),esk2_0)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~v1_funct_1(k3_relat_1(k3_taylor_1,k19_sin_cos))|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_500, plain, (m1_subset_1(k1_relset_1(k1_numbers,k24_sin_cos),k1_zfmisc_1(k1_numbers))|~v1_relat_1(k24_sin_cos))).#trainneg
cnf(i_0_6091, negated_conjecture, (r1_xxreal_0(k5_numbers,esk1_0)|r1_xxreal_0(esk1_0,k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_675, negated_conjecture, (v1_xxreal_0(esk1_0)|~v2_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))))).#trainneg
cnf(i_0_2720, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(X1))|r1_xxreal_0(k1_real_1(X1),esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_514, plain, (r2_relset_1(X1,X2,X3,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_520, plain, (v1_xcmplx_0(esk3_2(X1,k3_taylor_1))|r2_fdiff_1(k3_taylor_1,X1)|~v1_membered(k1_numbers)|~r1_tarski(X1,a_0_0_fdiff_7)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_6331, plain, (k7_xcmplx_0(k7_xcmplx_0(np__1,X1),np__1)=k1_fdiff_1(k3_taylor_1,X1)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_6241, plain, (k1_real_1(k1_fdiff_1(k3_taylor_1,X1))=k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_6373, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k24_sin_cos)))).#trainneg
cnf(i_0_6370, plain, (v3_membered(k1_relset_1(k1_numbers,k24_sin_cos)))).#trainneg
cnf(i_0_6371, plain, (v1_membered(k1_relset_1(k1_numbers,k24_sin_cos)))).#trainneg
cnf(i_0_6363, negated_conjecture, (k1_fdiff_1(k3_relat_1(k3_taylor_1,k19_sin_cos),esk2_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk2_0)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_xcmplx_0(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)))|~m2_subset_1(esk2_0,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0))|~r1_fdiff_1(k3_taylor_1,esk2_0)|~m1_subset_1(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_521, plain, (v1_int_1(esk3_2(X1,k3_taylor_1))|r2_fdiff_1(k3_taylor_1,X1)|~v5_membered(k1_numbers)|~r1_tarski(X1,a_0_0_fdiff_7)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_6120, plain, (r1_xxreal_0(k1_real_1(k1_real_1(k1_real_1(X1))),X2)|r1_xxreal_0(X2,k1_real_1(k1_real_1(k1_real_1(X1))))|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1404, plain, (r1_xxreal_0(X1,k1_relset_1(k1_numbers,k16_sin_cos))|r1_xxreal_0(k1_relset_1(k1_numbers,k16_sin_cos),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_522, plain, (v7_ordinal1(esk3_2(X1,k3_taylor_1))|r2_fdiff_1(k3_taylor_1,X1)|~v6_membered(k1_numbers)|~r1_tarski(X1,a_0_0_fdiff_7)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_523, plain, (v1_xxreal_0(esk3_2(X1,k3_taylor_1))|r2_fdiff_1(k3_taylor_1,X1)|~v2_membered(k1_numbers)|~r1_tarski(X1,a_0_0_fdiff_7)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_6163, plain, (k1_real_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1))=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_6169, plain, (k1_real_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2))=k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_6266, plain, (k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),np__1)=k1_real_1(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_6183, plain, (k4_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(np__1,X1)))=k1_fdiff_1(k3_taylor_1,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(np__1,X1)),k1_numbers))).#trainneg
cnf(i_0_6586, plain, (k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),X2)=k4_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2))|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_528, plain, (v1_xcmplx_0(esk3_2(X1,X2))|r2_fdiff_1(X2,X1)|~v1_membered(k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_5009, plain, (r1_xxreal_0(k24_sin_cos,k24_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5193, plain, (r1_xxreal_0(k16_sin_cos,k16_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1561, plain, (r1_xxreal_0(X1,k1_relset_1(k1_numbers,k19_sin_cos))|r1_xxreal_0(k1_relset_1(k1_numbers,k19_sin_cos),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6505, plain, (r1_xxreal_0(k1_relset_1(k1_numbers,k16_sin_cos),X1)|r1_xxreal_0(X1,k1_relset_1(k1_numbers,k16_sin_cos))|~v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_529, plain, (v1_int_1(esk3_2(X1,X2))|r2_fdiff_1(X2,X1)|~v5_membered(k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_530, plain, (v7_ordinal1(esk3_2(X1,X2))|r2_fdiff_1(X2,X1)|~v6_membered(k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_6530, plain, (k1_real_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1))=k7_xcmplx_0(k4_xcmplx_0(np__1),X1)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_6532, plain, (k4_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1))=k7_xcmplx_0(k4_xcmplx_0(np__1),X1)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_6529, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1))=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_6514, plain, (k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),np__1)=k4_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1))|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_6561, plain, (k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1)=k1_fdiff_1(k3_taylor_1,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_6538, plain, (k4_xcmplx_0(k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),X2))=k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_6572, plain, (k7_xcmplx_0(k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),X2),np__1)=k4_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2))|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_6589, plain, (k4_xcmplx_0(k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),X2))=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_569, plain, (v1_xreal_0(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3748, negated_conjecture, (r1_xxreal_0(esk2_0,k4_xcmplx_0(np__1))|r1_xxreal_0(k4_xcmplx_0(np__1),esk2_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_531, plain, (v1_xxreal_0(esk3_2(X1,X2))|r2_fdiff_1(X2,X1)|~v2_membered(k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_5199, plain, (r1_xxreal_0(k19_sin_cos,k19_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5320, plain, (r1_xxreal_0(k3_taylor_1,k3_taylor_1)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6715, plain, (r1_xxreal_0(k1_relset_1(k1_numbers,k19_sin_cos),X1)|r1_xxreal_0(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1959, plain, (r1_xxreal_0(k6_numbers,k10_real_1(X1,X2))|r1_xxreal_0(k10_real_1(X1,X2),k6_numbers)|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_535, plain, (v1_int_1(k7_xcmplx_0(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6571, plain, (k4_xcmplx_0(k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),X2))=k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),np__1)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_6603, plain, (k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),X2)=k4_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2))|~v1_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2))|~m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2)),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_6555, plain, (k4_xcmplx_0(k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),X2))=k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),X2)|~v1_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1))|~v1_xcmplx_0(X2)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_6677, plain, (k1_real_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2))=k7_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(np__1,X1)),X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(np__1,X1)),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_6515, plain, (k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),np__1)=k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1))|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_6587, plain, (k1_real_1(k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),X2))=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_6441, plain, (k1_real_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2))=k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),X2)|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~v1_xcmplx_0(X2)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),X2),k1_numbers)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers))).#trainneg
cnf(i_0_541, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3751, plain, (r1_xxreal_0(esk6_0,k4_xcmplx_0(np__1))|r1_xxreal_0(k4_xcmplx_0(np__1),esk6_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_4284, plain, (k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1))=k4_xcmplx_0(np__1)|~m1_subset_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1)),k1_numbers)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_536, plain, (v7_ordinal1(k7_xcmplx_0(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6097, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))))).#trainneg
cnf(i_0_537, plain, (v1_xxreal_0(k7_xcmplx_0(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_538, plain, (m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6570, plain, (k1_real_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers))).#trainneg
cnf(i_0_6326, plain, (k4_xcmplx_0(k1_real_1(k7_xcmplx_0(np__1,X1)))=k1_fdiff_1(k3_taylor_1,X1)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_542, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6220, negated_conjecture, (v1_int_1(esk1_0)|~v6_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))))).#trainneg
cnf(i_0_6365, plain, (v1_xreal_0(k1_relset_1(k1_numbers,k24_sin_cos))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_543, plain, (v1_int_1(k3_xcmplx_0(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_544, plain, (v7_ordinal1(k3_xcmplx_0(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6528, plain, (k4_xcmplx_0(k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),X2))=k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),X2)|~v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1))|~v1_xcmplx_0(X2)|~m1_subset_1(k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),X2),k1_numbers)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_7551, plain, (m1_subset_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_7687, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1))=k1_fdiff_1(k3_taylor_1,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers))).#trainneg
cnf(i_0_7632, plain, (k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1)=k1_fdiff_1(k3_taylor_1,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers))).#trainneg
cnf(i_0_6596, plain, (k1_real_1(k7_xcmplx_0(k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),X2),X3))=k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),X3)|~v1_xcmplx_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_632, plain, (v1_xreal_0(k8_real_1(X2,X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7075, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk2_0)|r1_xxreal_0(esk2_0,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_545, plain, (v1_xxreal_0(k3_xcmplx_0(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6366, plain, (v1_xcmplx_0(k1_relset_1(k1_numbers,k24_sin_cos))|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6367, plain, (v1_int_1(k1_relset_1(k1_numbers,k24_sin_cos))|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7732, plain, (v7_ordinal1(X1)|~v6_membered(k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_547, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(k7_xcmplx_0(np__1,X2))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_549, plain, (m1_subset_1(k3_xcmplx_0(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_6547, plain, (k1_real_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1))=k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),np__1)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_7808, plain, (v1_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_7994, plain, (m1_subset_1(k7_xcmplx_0(X1,np__1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7450, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),esk6_0)|r1_xxreal_0(esk6_0,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_6368, plain, (v7_ordinal1(k1_relset_1(k1_numbers,k24_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6369, plain, (v1_xxreal_0(k1_relset_1(k1_numbers,k24_sin_cos))|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7713, plain, (v1_int_1(X1)|~v5_membered(k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7974, plain, (v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_560, plain, (m1_subset_1(esk3_2(k5_numbers,X1),k1_numbers)|r2_fdiff_1(X1,k5_numbers)|~v1_funct_1(X1)|~r1_tarski(k5_numbers,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(k5_numbers))).#trainneg
cnf(i_0_7870, plain, (k4_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(np__1,X1)))=k1_fdiff_1(k3_taylor_1,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers))).#trainneg
cnf(i_0_7807, plain, (v1_xreal_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_6376, negated_conjecture, (v1_xxreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))))).#trainneg
cnf(i_0_7991, plain, (v1_xcmplx_0(k1_relset_1(k1_numbers,k24_sin_cos))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7993, plain, (v1_int_1(X1)|~v6_membered(k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_562, plain, (m1_subset_1(esk3_2(a_0_0_fdiff_7,X1),k1_numbers)|r2_fdiff_1(X1,a_0_0_fdiff_7)|~v1_funct_1(X1)|~r1_tarski(a_0_0_fdiff_7,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(a_0_0_fdiff_7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(a_0_0_fdiff_7))).#trainneg
cnf(i_0_7836, plain, (v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers))).#trainneg
cnf(i_0_7851, plain, (k1_real_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1))=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1)|~v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1))|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers))).#trainneg
cnf(i_0_8072, plain, (k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),np__1)=k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),np__1)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_1939, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~r2_hidden(X2,a_0_0_fdiff_7))).#trainneg
cnf(i_0_2173, plain, (v1_xcmplx_0(k8_real_1(X2,X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7992, plain, (v1_int_1(k1_relset_1(k1_numbers,k24_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5652, plain, (k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1),np__1)=np__1|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1))|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1),np__1),k1_numbers)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_6094, negated_conjecture, (v1_xreal_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7485, negated_conjecture, (v1_xcmplx_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6521, plain, (k1_real_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1))=k4_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1))|~m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1)),k1_numbers)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_8249, plain, (k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1),np__1)=np__1|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_8281, plain, (k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),np__1)=k7_xcmplx_0(k4_xcmplx_0(np__1),X1)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_8134, plain, (v1_xxreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8086, plain, (v1_xreal_0(k7_xcmplx_0(X1,np__1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_588, plain, (r1_tarski(k4_xxreal_1(X1,X2),a_0_0_fdiff_7)|~r2_fdiff_1(k3_taylor_1,k4_xxreal_1(X1,X2))|~m1_subset_1(k4_xxreal_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8317, plain, (v1_xreal_0(k1_relset_1(k1_numbers,k24_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8196, plain, (v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1))|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_8308, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,k6_numbers)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_598, plain, (m1_subset_1(esk4_3(X1,k16_sin_cos,X2),k2_zfmisc_1(k1_numbers,k1_numbers))|X2=k2_fdiff_1(X1,k16_sin_cos)|k1_relset_1(k1_numbers,X2)!=k16_sin_cos|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k16_sin_cos)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_8154, plain, (k7_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(np__1,X1)),X2)=k4_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(np__1,X1)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers))).#trainneg
cnf(i_0_8291, plain, (k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),np__1)=k1_real_1(k1_fdiff_1(k3_taylor_1,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_8170, plain, (k1_real_1(k4_xcmplx_0(k7_xcmplx_0(np__1,X1)))=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1)|~v1_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(np__1,X1)))|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers))).#trainneg
cnf(i_0_6602, plain, (k4_xcmplx_0(k7_xcmplx_0(k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),X2),X3))=k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),X3)|~v1_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2))|~v1_xcmplx_0(X3)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_8087, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,np__1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_599, plain, (m1_subset_1(esk3_2(k16_sin_cos,X1),k2_zfmisc_1(k1_numbers,k1_numbers))|r2_fdiff_1(X1,k16_sin_cos)|~v1_funct_1(X1)|~r1_tarski(k16_sin_cos,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k16_sin_cos))).#trainneg
cnf(i_0_8440, plain, (r1_xxreal_0(X1,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_604, plain, (k4_xcmplx_0(k7_xcmplx_0(k1_real_1(X1),X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8460, plain, (v1_xreal_0(k1_fdiff_1(k3_taylor_1,np__1))|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_6378, negated_conjecture, (r1_xxreal_0(np__1,esk1_0)|r1_xxreal_0(esk1_0,np__1)|~v2_membered(k1_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8439, plain, (r1_xxreal_0(k6_numbers,X1)|r1_xxreal_0(X1,k6_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_630, plain, (v1_xreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3))|~v1_xreal_0(k1_fdiff_1(X2,X3))|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(X2,X3)),k1_numbers)|~r1_fdiff_1(X1,k1_seq_1(X2,X3))|~r1_fdiff_1(X2,X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_8383, plain, (k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1)),np__1)=k4_xcmplx_0(np__1)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,np__1),k1_numbers)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_8630, plain, (k7_xcmplx_0(k7_xcmplx_0(np__1,X1),np__1)=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_8692, plain, (k4_xcmplx_0(k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),X2))=k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_635, plain, (v1_xcmplx_0(k1_fdiff_1(k19_sin_cos,X1))|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_8436, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_637, plain, (v1_int_1(k1_fdiff_1(k19_sin_cos,X1))|~v1_int_1(k1_seq_1(k16_sin_cos,X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_8834, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(X1),X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8828, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(X1),X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8753, plain, (k1_real_1(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k1_real_1(X1),X2)|~m1_subset_1(k7_xcmplx_0(k1_real_1(X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8755, plain, (k7_xcmplx_0(k1_real_1(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~m1_subset_1(k7_xcmplx_0(k1_real_1(X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8513, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X2,a_0_0_fdiff_7))).#trainneg
cnf(i_0_2715, negated_conjecture, (r1_xxreal_0(esk1_0,a_0_0_fdiff_7)|r1_xxreal_0(a_0_0_fdiff_7,esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_639, plain, (k4_xcmplx_0(k7_xcmplx_0(k1_seq_1(k16_sin_cos,X1),X2))=k7_xcmplx_0(k1_fdiff_1(k19_sin_cos,X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_8190, negated_conjecture, (v1_xxreal_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_916, plain, (r1_xxreal_0(k6_numbers,k1_fdiff_1(k3_taylor_1,np__2))|~m2_subset_1(np__2,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_640, plain, (v1_xreal_0(k1_fdiff_1(k19_sin_cos,X1))|~v1_int_1(k1_seq_1(k16_sin_cos,X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_641, plain, (v1_int_1(k1_seq_1(k16_sin_cos,X1))|~v1_int_1(k1_fdiff_1(k19_sin_cos,X1))|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_8995, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|r1_xxreal_0(esk2_0,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_9023, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8551, plain, (k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),X2)=k4_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers))).#trainneg
cnf(i_0_8514, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v3_membered(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_8088, plain, (v1_int_1(k7_xcmplx_0(X1,np__1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2718, negated_conjecture, (r1_xxreal_0(esk1_0,np__1)|r1_xxreal_0(np__1,esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5613, plain, (k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1))))=k4_xcmplx_0(np__1)|~m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1)))),k1_numbers)|~m1_subset_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1)),k1_numbers)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_8863, plain, (r1_xxreal_0(k1_real_1(X1),k6_numbers)|r1_xxreal_0(k6_numbers,k1_real_1(X1))|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_648, plain, (v1_xreal_0(k1_fdiff_1(k19_sin_cos,X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_9396, negated_conjecture, (r1_xxreal_0(esk2_0,esk3_2(X1,X2))|r1_xxreal_0(esk3_2(X1,X2),esk2_0)|r2_fdiff_1(X2,X1)|~v1_xreal_0(esk3_2(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_9393, negated_conjecture, (r1_xxreal_0(esk2_0,esk5_1(X1))|r1_xxreal_0(esk5_1(X1),esk2_0)|~v1_xreal_0(esk5_1(X1))|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_9138, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8824, plain, (k7_xcmplx_0(k1_real_1(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8412, plain, (k7_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))),np__1)=k1_real_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1))|~v1_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_9218, plain, (k4_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_658, plain, (m1_subset_1(esk3_2(k2_relset_1(X1,X2),X3),X1)|r2_fdiff_1(X3,k2_relset_1(X1,X2))|~v5_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X3)|~r1_tarski(k2_relset_1(X1,X2),k1_relset_1(k1_numbers,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_relset_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_relset_1(X1,X2)))).#trainneg
cnf(i_0_9511, plain, (v1_int_1(k1_fdiff_1(k3_taylor_1,np__1))|~v5_membered(k1_numbers)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_657, plain, (m1_subset_1(esk4_3(X1,k2_relset_1(X2,X3),X4),X2)|X4=k2_fdiff_1(X1,k2_relset_1(X2,X3))|~v5_relat_1(X3,X2)|~v1_relat_1(X3)|k1_relset_1(k1_numbers,X4)!=k2_relset_1(X2,X3)|~v1_funct_1(X4)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k2_relset_1(X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_661, plain, (m1_subset_1(esk4_3(X1,k19_sin_cos,X2),k2_zfmisc_1(k1_numbers,k1_numbers))|X2=k2_fdiff_1(X1,k19_sin_cos)|k1_relset_1(k1_numbers,X2)!=k19_sin_cos|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k19_sin_cos)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_9716, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8842, plain, (r1_xxreal_0(esk3_2(X1,k3_taylor_1),k6_numbers)|r1_xxreal_0(k6_numbers,esk3_2(X1,k3_taylor_1))|r2_fdiff_1(k3_taylor_1,X1)|~v1_xreal_0(esk3_2(X1,k3_taylor_1))|~r1_tarski(X1,a_0_0_fdiff_7)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_8089, plain, (v7_ordinal1(k7_xcmplx_0(X1,np__1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_662, plain, (m1_subset_1(esk3_2(k19_sin_cos,X1),k2_zfmisc_1(k1_numbers,k1_numbers))|r2_fdiff_1(X1,k19_sin_cos)|~v1_funct_1(X1)|~r1_tarski(k19_sin_cos,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k19_sin_cos))).#trainneg
cnf(i_0_8512, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_663, plain, (m1_subset_1(esk4_3(X1,k24_sin_cos,X2),k2_zfmisc_1(k1_numbers,k1_numbers))|X2=k2_fdiff_1(X1,k24_sin_cos)|k1_relset_1(k1_numbers,X2)!=k24_sin_cos|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k24_sin_cos)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_664, plain, (m1_subset_1(esk3_2(k24_sin_cos,X1),k2_zfmisc_1(k1_numbers,k1_numbers))|r2_fdiff_1(X1,k24_sin_cos)|~v1_funct_1(X1)|~r1_tarski(k24_sin_cos,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k24_sin_cos))).#trainneg
cnf(i_0_8693, plain, (k4_xcmplx_0(k7_xcmplx_0(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)),X2))=k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_9586, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|r1_xxreal_0(esk2_0,X1)|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_8044, plain, (k1_real_1(k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),np__1))=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_8398, plain, (k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1)),np__1)=k4_xcmplx_0(np__1)|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1))|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1)),np__1),k1_numbers)|~m1_subset_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1)),k1_numbers))).#trainneg
cnf(i_0_7825, plain, (v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers))).#trainneg
cnf(i_0_668, plain, (m1_subset_1(esk4_3(X1,k3_taylor_1,X2),k2_zfmisc_1(k1_numbers,k1_numbers))|X2=k2_fdiff_1(X1,k3_taylor_1)|k1_relset_1(k1_numbers,X2)!=k3_taylor_1|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k3_taylor_1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_10024, plain, (v7_ordinal1(k1_fdiff_1(k3_taylor_1,np__1))|~v6_membered(k1_numbers)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_669, plain, (m1_subset_1(esk3_2(k3_taylor_1,X1),k2_zfmisc_1(k1_numbers,k1_numbers))|r2_fdiff_1(X1,k3_taylor_1)|~v1_funct_1(X1)|~r1_tarski(k3_taylor_1,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k3_taylor_1))).#trainneg
cnf(i_0_670, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))))|r2_hidden(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))))).#trainneg
cnf(i_0_7155, plain, (k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),np__1)=k1_real_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1))|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers))).#trainneg
cnf(i_0_6162, plain, (k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),np__1)=k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_8090, plain, (v1_xxreal_0(k7_xcmplx_0(X1,np__1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_678, negated_conjecture, (m1_subset_1(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_9340, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,k6_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_10394, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))|r2_hidden(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_10335, plain, (k1_real_1(k1_fdiff_1(k3_taylor_1,X1))=k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_10327, plain, (k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),np__1)=k4_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_10326, plain, (k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),np__1)=k4_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~v1_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_10317, plain, (k1_real_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1))=k4_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_10325, plain, (k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),np__1)=k4_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_689, plain, (v1_int_1(k8_real_1(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_690, plain, (v7_ordinal1(k8_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_691, plain, (v1_xxreal_0(k8_real_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_10313, plain, (k1_real_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1))=k1_real_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_10440, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))|r2_hidden(esk2_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))))).#trainneg
cnf(i_0_10439, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))|r2_hidden(esk3_2(esk1_0,X1),k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))|r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~r1_tarski(esk1_0,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_10438, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))|r2_hidden(esk3_2(esk1_0,X1),k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos)))|r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_10328, plain, (k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),np__1)=k4_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_10330, plain, (k1_real_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_10306, plain, (k7_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1)),np__1)=k1_real_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_10331, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_10314, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1))=k4_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_9011, plain, (r1_xxreal_0(X1,k1_real_1(X2))|r1_xxreal_0(k1_real_1(X2),X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_692, plain, (m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3),k1_numbers)|~v1_xreal_0(k1_fdiff_1(X1,k1_seq_1(X2,X3)))|~m1_subset_1(k1_fdiff_1(X2,X3),k1_numbers)|~r1_fdiff_1(X1,k1_seq_1(X2,X3))|~r1_fdiff_1(X2,X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_10304, plain, (k7_xcmplx_0(k1_real_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1)),np__1)=k1_real_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_10715, negated_conjecture, (k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))=k6_numbers|r2_hidden(esk2_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))))).#trainneg
cnf(i_0_10429, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k1_fdiff_1(k3_taylor_1,X2),k6_numbers)|r1_xxreal_0(X1,k1_fdiff_1(k3_taylor_1,X2))|r1_xxreal_0(k1_fdiff_1(k3_taylor_1,X2),X1)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)|~m2_subset_1(X2,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_7360, plain, (k7_xcmplx_0(k1_real_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1)),X2)=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),k1_numbers))).#trainneg
cnf(i_0_10002, plain, (v1_int_1(k7_xcmplx_0(X1,np__1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_697, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(X1),X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10259, plain, (v1_int_1(k1_fdiff_1(k3_taylor_1,np__1))|~v6_membered(k1_numbers)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_10921, plain, (r1_xxreal_0(k1_real_1(X1),X2)|r1_xxreal_0(X2,k1_real_1(X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8442, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|r1_xxreal_0(X1,esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_698, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_int_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_699, plain, (v1_int_1(k7_xcmplx_0(X1,X2))|~v1_int_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6604, plain, (k1_real_1(k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),X2))=k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),np__1)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_8575, plain, (k4_xcmplx_0(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)))=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1)|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_8263, plain, (k1_real_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1)|~v1_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers))).#trainneg
cnf(i_0_9341, negated_conjecture, (r1_xxreal_0(a_0_0_fdiff_7,esk1_0)|r1_xxreal_0(esk1_0,a_0_0_fdiff_7)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_10380, plain, (v1_xxreal_0(k1_fdiff_1(k3_taylor_1,np__1))|~v2_membered(k1_numbers)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_701, plain, (k4_xcmplx_0(k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(X1),X2),X3))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(X1),X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_702, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_int_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6258, plain, (k1_real_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2))=k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_1945, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9526, negated_conjecture, (r1_xxreal_0(np__1,esk1_0)|r1_xxreal_0(esk1_0,np__1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_703, plain, (v1_int_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2))|~v1_int_1(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(X1),X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_704, plain, (k1_real_1(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k4_xcmplx_0(X1),X2)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_705, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_11730, plain, (k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),np__1)=k1_real_1(k1_fdiff_1(k3_taylor_1,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_11693, plain, (k1_real_1(k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),X2))=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_11907, plain, (k1_real_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2))=k7_xcmplx_0(X1,X2)|~m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2),k1_numbers))).#trainneg
cnf(i_0_11997, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),np__1)=k1_real_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),np__1),k1_numbers))).#trainneg
cnf(i_0_12080, plain, (k1_real_1(k1_fdiff_1(k3_taylor_1,X1))=k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~v1_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_3251, negated_conjecture, (r1_xxreal_0(esk1_0,esk2_0)|r1_xxreal_0(esk2_0,esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_706, plain, (k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(X1),X2),X3)=k4_xcmplx_0(k7_xcmplx_0(k7_xcmplx_0(X1,X2),X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(X1),X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_707, plain, (k7_xcmplx_0(k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_12223, plain, (k1_real_1(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_12224, plain, (k1_real_1(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_710, plain, (k4_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k19_sin_cos,X1),X2))=k7_xcmplx_0(k1_seq_1(k16_sin_cos,X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_8111, plain, (m1_subset_1(k1_fdiff_1(k3_taylor_1,np__1),k1_numbers)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_9025, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_9457, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_712, plain, (k4_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_11894, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(X1),X2))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),np__1)|~m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2),k1_numbers))).#trainneg
cnf(i_0_11742, plain, (r1_xxreal_0(X1,X1)|~v2_membered(k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_717, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos))))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1942, plain, (r1_xxreal_0(esk6_0,X1)|r1_xxreal_0(X1,esk6_0)|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_12161, plain, (k1_real_1(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)))=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers))).#trainneg
cnf(i_0_12257, plain, (k1_real_1(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k1_real_1(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12376, negated_conjecture, (r1_xxreal_0(esk2_0,esk1_0)|r1_xxreal_0(esk1_0,esk2_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8500, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k1_real_1(X2),X1)|r1_xxreal_0(X1,k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~v2_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_8485, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(esk1_0,X1)|r1_xxreal_0(X1,esk1_0)|~m1_subset_1(X1,k1_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_11924, plain, (k1_real_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),np__1)|~m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2),k1_numbers))).#trainneg
cnf(i_0_13215, plain, (k1_real_1(k7_xcmplx_0(k1_real_1(X1),X2))=k7_xcmplx_0(X1,X2)|~m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12329, plain, (k4_xcmplx_0(k7_xcmplx_0(X1,np__1))=k1_real_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(X1,np__1)),k1_numbers))).#trainneg
cnf(i_0_12572, plain, (k1_real_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2))=k7_xcmplx_0(X1,X2)|~m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13131, plain, (r1_xxreal_0(X1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_754, plain, (k4_xcmplx_0(k7_xcmplx_0(k6_numbers,X1))=k7_xcmplx_0(k6_numbers,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k6_numbers))).#trainneg
cnf(i_0_8444, plain, (v1_xreal_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1)))|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_8703, plain, (v1_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1)))|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_13315, plain, (r1_xxreal_0(X1,k1_real_1(X2))|r1_xxreal_0(k1_real_1(X2),X1)|r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_11186, plain, (r1_xxreal_0(X1,k1_real_1(k1_real_1(X2)))|r1_xxreal_0(k1_real_1(k1_real_1(X2)),X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_13547, plain, (k1_real_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13754, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1044, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_13795, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X1))|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13626, plain, (k4_xcmplx_0(k7_xcmplx_0(X1,np__1))=k1_real_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(X1,np__1),k1_numbers))).#trainneg
cnf(i_0_14322, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(X1,np__1),k1_numbers))).#trainneg
cnf(i_0_13500, plain, (k1_real_1(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2),k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14168, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14083, plain, (r1_xxreal_0(np__2,np__1)|r1_xxreal_0(np__1,np__2))).#trainneg
cnf(i_0_13627, plain, (k4_xcmplx_0(k7_xcmplx_0(X1,np__1))=k1_real_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_12832, plain, (r2_hidden(k1_fdiff_1(k3_taylor_1,np__1),k1_numbers)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_792, plain, (r2_fdiff_1(k3_taylor_1,X1)|m1_subset_1(esk3_2(X1,k3_taylor_1),X1)|~r1_tarski(X1,a_0_0_fdiff_7)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_794, plain, (r2_fdiff_1(X1,k1_relset_1(k1_numbers,X1))|m1_subset_1(esk3_2(k1_relset_1(k1_numbers,X1),X1),k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,X1)))).#trainneg
cnf(i_0_798, plain, (m1_subset_1(k1_funct_1(X1,X2),k1_numbers)|~v1_xreal_0(k1_funct_1(X1,X2))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1958, plain, (r1_xxreal_0(X1,k10_real_1(X2,X3))|r1_xxreal_0(k10_real_1(X2,X3),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_799, plain, (m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_800, plain, (m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_1484, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14077, plain, (r1_xxreal_0(np__2,k4_xcmplx_0(np__1))|r1_xxreal_0(k4_xcmplx_0(np__1),np__2)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_5339, plain, (r1_xxreal_0(k1_real_1(X1),X2)|r1_xxreal_0(X2,k1_real_1(X1))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_numbers))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5326, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|r1_xxreal_0(X1,esk1_0)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_817, plain, (v1_int_1(k1_fdiff_1(k19_sin_cos,X1))|~v5_membered(k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_12961, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))=k1_fdiff_1(k3_taylor_1,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_1485, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_822, plain, (k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))=k6_numbers|r2_hidden(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_9491, plain, (v1_int_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1)))|~v5_membered(k1_numbers)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_5350, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_numbers))|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_6223, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|r1_xxreal_0(esk1_0,X1)|~v3_membered(k1_zfmisc_1(X1))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_823, plain, (k2_fdiff_1(X1,X2)=k2_fdiff_1(X3,X4)|k1_fdiff_1(X3,esk4_3(X3,X4,k2_fdiff_1(X1,X2)))!=k1_fdiff_1(X1,esk4_3(X3,X4,k2_fdiff_1(X1,X2)))|k1_relset_1(k1_numbers,k2_fdiff_1(X1,X2))!=X4|~v1_funct_1(X3)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r2_fdiff_1(X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(esk4_3(X3,X4,k2_fdiff_1(X1,X2)),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_3(X3,X4,k2_fdiff_1(X1,X2)),k1_numbers))).#trainneg
cnf(i_0_825, plain, (m1_subset_1(k1_fdiff_1(X1,X3),k1_numbers)|~v3_valued_0(k2_fdiff_1(X1,X2))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_13720, plain, (k1_real_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2))=k7_xcmplx_0(k7_xcmplx_0(X1,np__1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k7_xcmplx_0(X1,np__1))|~m1_subset_1(k7_xcmplx_0(k7_xcmplx_0(X1,np__1),X2),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_15027, plain, (m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~r2_hidden(X1,a_0_0_fdiff_7)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15024, negated_conjecture, (m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),X1),k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14869, plain, (m1_subset_1(k4_xcmplx_0(esk6_0),k1_numbers))).#trainneg
cnf(i_0_14875, plain, (m1_subset_1(k4_xcmplx_0(np__2),k1_numbers))).#trainneg
cnf(i_0_827, plain, (v1_xreal_0(k1_fdiff_1(X1,X3))|~v3_valued_0(k2_fdiff_1(X1,X2))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_10004, plain, (v7_ordinal1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1)))|~v6_membered(k1_numbers)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_14810, plain, (r1_xxreal_0(k10_real_1(X1,X2),X3)|r1_xxreal_0(X3,k10_real_1(X1,X2))|~v1_xxreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_828, plain, (k8_real_1(k1_fdiff_1(X1,k1_fdiff_1(X2,X4)),k1_fdiff_1(k2_fdiff_1(X2,X3),X4))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X2,X3),X1),X4)|~r1_fdiff_1(X1,k1_fdiff_1(X2,X4))|~r1_fdiff_1(k2_fdiff_1(X2,X3),X4)|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_numbers)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~r2_hidden(X4,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_833, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_14948, plain, (k1_real_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2))=k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_14951, plain, (k1_real_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2))=k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_15187, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),X2),k6_numbers)|r1_xxreal_0(X1,k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),X2))|r1_xxreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),X2),X1)|~m1_subset_1(X1,k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~r2_hidden(X2,esk1_0)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_15184, negated_conjecture, (r1_xxreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),X1),k6_numbers)|r1_xxreal_0(k6_numbers,k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),X1))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_15193, plain, (v1_xreal_0(k4_xcmplx_0(esk6_0)))).#trainneg
cnf(i_0_15194, plain, (v1_xcmplx_0(k4_xcmplx_0(esk6_0)))).#trainneg
cnf(i_0_15211, plain, (v1_xreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_14074, plain, (r1_xxreal_0(np__2,k5_numbers)|r1_xxreal_0(k5_numbers,np__2)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_15212, plain, (v1_xcmplx_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_15195, plain, (v1_int_1(k4_xcmplx_0(esk6_0))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_15274, plain, (r1_xxreal_0(X1,k10_real_1(k6_numbers,X2))|r1_xxreal_0(k10_real_1(k6_numbers,X2),X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_6500, plain, (r1_xxreal_0(X1,k1_real_1(k1_real_1(k1_real_1(k1_real_1(X2)))))|r1_xxreal_0(k1_real_1(k1_real_1(k1_real_1(k1_real_1(X2)))),X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_15177, negated_conjecture, (v1_xcmplx_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),X1))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_15470, plain, (k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1)),np__1)=k4_xcmplx_0(np__1)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_15077, plain, (k7_xcmplx_0(k7_xcmplx_0(X1,np__1),X2)=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2),k1_numbers)|~v1_xcmplx_0(k7_xcmplx_0(X1,np__1))|~m1_subset_1(k7_xcmplx_0(k7_xcmplx_0(X1,np__1),X2),k1_numbers))).#trainneg
cnf(i_0_14936, plain, (v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~v1_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_15160, plain, (v1_xreal_0(k1_fdiff_1(k3_taylor_1,X1))|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_15196, plain, (v7_ordinal1(k4_xcmplx_0(esk6_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_14076, plain, (r1_xxreal_0(np__2,a_0_0_fdiff_7)|r1_xxreal_0(a_0_0_fdiff_7,np__2)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_15197, plain, (v1_xxreal_0(k4_xcmplx_0(esk6_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_15213, plain, (v1_int_1(k4_xcmplx_0(np__2))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_5375, plain, (r1_xxreal_0(X1,k1_real_1(k1_seq_1(X2,X3)))|r1_xxreal_0(k1_real_1(k1_seq_1(X2,X3)),X1)|~v1_xxreal_0(X1)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_853, plain, (v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~v1_membered(k1_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_854, plain, (v1_int_1(k1_fdiff_1(k3_taylor_1,X1))|~v5_membered(k1_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_15579, negated_conjecture, (r1_xxreal_0(esk2_0,k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),X1))|r1_xxreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),X1),esk2_0)|~v1_xreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),X1))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_14939, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),X2))=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_14934, plain, (k1_real_1(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)))=k1_real_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~m1_subset_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_14952, plain, (k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),X2),np__1)=k1_real_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_15161, plain, (v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_15214, plain, (v7_ordinal1(k4_xcmplx_0(np__2))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_855, plain, (v7_ordinal1(k1_fdiff_1(k3_taylor_1,X1))|~v6_membered(k1_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_15488, plain, (r1_xxreal_0(k5_numbers,np__2)|r1_xxreal_0(np__2,k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_15215, plain, (v1_xxreal_0(k4_xcmplx_0(np__2))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_15703, plain, (v1_int_1(k4_xcmplx_0(esk6_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_15527, plain, (r1_xxreal_0(k10_real_1(k6_numbers,X1),k6_numbers)|r1_xxreal_0(k6_numbers,k10_real_1(k6_numbers,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_15534, plain, (r1_xxreal_0(k10_real_1(k6_numbers,X1),X2)|r1_xxreal_0(X2,k10_real_1(k6_numbers,X1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_856, plain, (v1_xxreal_0(k1_fdiff_1(k3_taylor_1,X1))|~v2_membered(k1_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_866, plain, (k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))=k6_numbers|r2_hidden(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_15150, plain, (k1_real_1(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)))=k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),np__1)|~v1_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1))|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),np__1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_14791, plain, (v1_xboole_0(X1)|r2_hidden(esk3_2(X1,k3_taylor_1),X1)|r2_fdiff_1(k3_taylor_1,X1)|~r1_tarski(X1,a_0_0_fdiff_7)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_11710, plain, (k1_real_1(k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),X2))=k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),np__1)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_15949, plain, (v1_int_1(k4_xcmplx_0(np__2))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_15683, plain, (k1_real_1(k1_fdiff_1(k3_taylor_1,np__1))=k4_xcmplx_0(np__1)|~m1_subset_1(k1_real_1(k1_fdiff_1(k3_taylor_1,np__1)),k1_numbers)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_15704, plain, (r1_xxreal_0(a_0_0_fdiff_7,np__2)|r1_xxreal_0(np__2,a_0_0_fdiff_7)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_15192, plain, (r2_hidden(k4_xcmplx_0(esk6_0),k1_numbers))).#trainneg
cnf(i_0_15210, plain, (r2_hidden(k4_xcmplx_0(np__2),k1_numbers))).#trainneg
cnf(i_0_15504, negated_conjecture, (r1_xxreal_0(k10_real_1(k6_numbers,X1),esk1_0)|r1_xxreal_0(esk1_0,k10_real_1(k6_numbers,X1))|~v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_874, plain, (v7_ordinal1(k1_fdiff_1(k19_sin_cos,X1))|~v6_membered(k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_12092, plain, (k4_xcmplx_0(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)))=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1)|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_15176, negated_conjecture, (v1_xreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),X1))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_13290, plain, (k7_xcmplx_0(k1_real_1(k7_xcmplx_0(np__1,X1)),np__1)=k1_real_1(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_14879, plain, (m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k1_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_14103, plain, (r1_xxreal_0(np__2,X1)|r1_xxreal_0(X1,np__2)|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_15494, plain, (m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(esk6_0)),k1_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_16366, plain, (m1_subset_1(X1,k1_numbers)|~v1_int_1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_882, plain, (v1_xxreal_0(k1_fdiff_1(k19_sin_cos,X1))|~v2_membered(k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_16124, plain, (k1_real_1(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)))=k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),np__1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers))).#trainneg
cnf(i_0_993, plain, (r2_fdiff_1(k19_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k19_sin_cos))))).#trainneg
cnf(i_0_16282, plain, (v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_15709, negated_conjecture, (r1_xxreal_0(esk2_0,k4_xcmplx_0(esk6_0))|r1_xxreal_0(k4_xcmplx_0(esk6_0),esk2_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_886, plain, (k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))=k6_numbers|r2_hidden(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1298, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_15712, plain, (r1_xxreal_0(k4_xcmplx_0(esk6_0),k4_xcmplx_0(esk6_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_15717, plain, (m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(np__2)),k1_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_15522, plain, (r1_xxreal_0(k10_real_1(k6_numbers,X1),k1_real_1(X2))|r1_xxreal_0(k1_real_1(X2),k10_real_1(k6_numbers,X1))|~v1_xreal_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_15536, plain, (r1_xxreal_0(k10_real_1(k6_numbers,X1),X2)|r1_xxreal_0(X2,k10_real_1(k6_numbers,X1))|~v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_887, plain, (v1_xreal_0(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_16512, plain, (r2_fdiff_1(k19_sin_cos,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k19_sin_cos))))).#trainneg
cnf(i_0_16165, plain, (k4_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1))=k1_real_1(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_16163, plain, (k1_real_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1))=k1_real_1(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_12153, plain, (k1_real_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers)|~v1_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))))).#trainneg
cnf(i_0_16627, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__1)|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_16283, plain, (v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_16380, plain, (v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(esk6_0)))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_16381, plain, (v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(esk6_0)))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_8479, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(esk6_0,X1)|r1_xxreal_0(X1,esk6_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_16382, plain, (v1_int_1(k4_xcmplx_0(k4_xcmplx_0(esk6_0)))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_907, plain, (m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2),k1_numbers)|~v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(X1),X2))|~m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_16401, plain, (m1_subset_1(k7_xcmplx_0(esk6_0,np__1),k1_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_16630, plain, (v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(np__2)))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_16631, plain, (v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(np__2)))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_896, plain, (k8_real_1(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_fdiff_1(X4,k1_fdiff_1(X1,X3)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X4),X3)|~v1_xreal_0(k1_fdiff_1(k2_fdiff_1(X1,X2),X3))|~r1_fdiff_1(X4,k1_fdiff_1(X1,X3))|~r1_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(X4)|~v1_funct_1(k2_fdiff_1(X1,X2))|~m1_subset_1(k1_fdiff_1(X4,k1_fdiff_1(X1,X3)),k1_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_16887, plain, (k4_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1))=k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers))).#trainneg
cnf(i_0_16873, plain, (k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),np__1)=k4_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~m1_subset_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers))).#trainneg
cnf(i_0_17028, plain, (r1_xxreal_0(esk3_2(X1,X2),esk6_0)|r1_xxreal_0(esk6_0,esk3_2(X1,X2))|r1_xxreal_0(esk3_2(X1,X2),k6_numbers)|r2_fdiff_1(X2,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_17221, plain, (k1_real_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1))=k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),np__1)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_914, plain, (m1_subset_1(esk4_3(X1,k10_xtuple_0(X2),X3),X4)|X3=k2_fdiff_1(X1,k10_xtuple_0(X2))|~v5_relat_1(X2,X4)|~v1_relat_1(X2)|k1_relset_1(k1_numbers,X3)!=k10_xtuple_0(X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k10_xtuple_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_16632, plain, (v1_int_1(k4_xcmplx_0(k4_xcmplx_0(np__2)))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_17006, plain, (v1_xreal_0(k7_xcmplx_0(esk6_0,np__1))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_915, plain, (m1_subset_1(esk3_2(k10_xtuple_0(X1),X2),X3)|r2_fdiff_1(X2,k10_xtuple_0(X1))|~v5_relat_1(X1,X3)|~v1_relat_1(X1)|~v1_funct_1(X2)|~r1_tarski(k10_xtuple_0(X1),k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k10_xtuple_0(X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k10_xtuple_0(X1)))).#trainneg
cnf(i_0_917, plain, (k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))=k6_numbers|r2_hidden(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_16782, plain, (k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),np__1)=k4_xcmplx_0(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)))|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers))).#trainneg
cnf(i_0_16760, plain, (k4_xcmplx_0(k7_xcmplx_0(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)),X2))=k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),X2)|~v1_xcmplx_0(X2)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers))).#trainneg
cnf(i_0_16872, plain, (k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),np__1)=k4_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1))|~m1_subset_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers))).#trainneg
cnf(i_0_16651, plain, (m1_subset_1(k7_xcmplx_0(np__2,np__1),k1_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_17015, plain, (v1_xcmplx_0(k7_xcmplx_0(esk6_0,np__1))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_943, plain, (v1_xreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X3),X2))|~v1_xreal_0(k1_fdiff_1(X3,k1_seq_1(X1,X2)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~r1_fdiff_1(X3,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_16896, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),X2))=k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),X2)|~v1_xcmplx_0(X2)|~m1_subset_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers))).#trainneg
cnf(i_0_17222, plain, (k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),np__1)=k1_real_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_16787, plain, (k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),np__1)=k4_xcmplx_0(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)))|~v1_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers))).#trainneg
cnf(i_0_16757, plain, (k4_xcmplx_0(k7_xcmplx_0(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)),X2))=k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers))).#trainneg
cnf(i_0_17691, plain, (k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),X2)=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2)|~v1_xcmplx_0(X2)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers))).#trainneg
cnf(i_0_8486, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(np__2,X1)|r1_xxreal_0(X1,np__2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_17069, plain, (v1_int_1(k7_xcmplx_0(esk6_0,np__1))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_16653, plain, (r1_xxreal_0(k1_real_1(X1),k10_real_1(k6_numbers,X2))|r1_xxreal_0(k10_real_1(k6_numbers,X2),k1_real_1(X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_16445, plain, (m1_subset_1(k1_relset_1(k1_numbers,k16_sin_cos),k1_numbers)|~v1_xcmplx_0(k1_relset_1(k1_numbers,k16_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_945, plain, (k8_real_1(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_fdiff_1(X4,k1_fdiff_1(X1,X3)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X4),X3)|~v1_xreal_0(k1_fdiff_1(X4,k1_fdiff_1(X1,X3)))|~r1_fdiff_1(X4,k1_fdiff_1(X1,X3))|~r1_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(X4)|~v1_funct_1(k2_fdiff_1(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_numbers)|~m1_subset_1(X3,k1_numbers)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_17877, plain, (k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),X2)=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2)|~v1_xcmplx_0(X2)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~v1_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers))).#trainneg
cnf(i_0_16885, plain, (k1_real_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1))=k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers))).#trainneg
cnf(i_0_16884, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1)|~m1_subset_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers))).#trainneg
cnf(i_0_8494, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_16810, plain, (k1_real_1(k1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,np__1)))=k4_xcmplx_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),np__1))|~v1_xcmplx_0(k1_seq_1(k3_taylor_1,np__1))|~v1_xreal_0(k1_seq_1(k3_taylor_1,np__1))|~m2_subset_1(k1_seq_1(k3_taylor_1,np__1),k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k1_real_1(k1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,np__1))),k1_numbers)|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,np__1)))|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,np__1))|~r1_fdiff_1(k3_taylor_1,np__1)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,np__1)),k1_numbers))).#trainneg
cnf(i_0_17150, plain, (v1_xreal_0(k7_xcmplx_0(np__2,np__1))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_16446, plain, (m1_subset_1(k1_relset_1(k1_numbers,k19_sin_cos),k1_numbers)|~v1_xcmplx_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_961, plain, (k3_xcmplx_0(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_fdiff_1(X4,k1_fdiff_1(X1,X3)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X4),X3)|~v1_xreal_0(k1_fdiff_1(X4,k1_fdiff_1(X1,X3)))|~r1_fdiff_1(X4,k1_fdiff_1(X1,X3))|~r1_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(X4)|~v1_funct_1(k2_fdiff_1(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_numbers)|~m1_subset_1(X3,k1_numbers)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_18109, plain, (k1_real_1(k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),X2))=k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_17159, plain, (v1_xcmplx_0(k7_xcmplx_0(np__2,np__1))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_17146, plain, (m1_subset_1(k7_xcmplx_0(esk6_0,np__1),k1_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_16009, plain, (r1_xxreal_0(X1,k10_real_1(k6_numbers,X2))|r1_xxreal_0(k10_real_1(k6_numbers,X2),X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_18022, plain, (r1_xxreal_0(k10_real_1(k6_numbers,X1),k1_real_1(k1_real_1(X2)))|r1_xxreal_0(k1_real_1(k1_real_1(X2)),k10_real_1(k6_numbers,X1))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_981, plain, (k3_xcmplx_0(k1_fdiff_1(X1,k1_fdiff_1(X2,X4)),k1_fdiff_1(k2_fdiff_1(X2,X3),X4))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X2,X3),X1),X4)|~v1_xreal_0(k1_fdiff_1(k2_fdiff_1(X2,X3),X4))|~r1_fdiff_1(X1,k1_fdiff_1(X2,X4))|~r1_fdiff_1(k2_fdiff_1(X2,X3),X4)|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(X2,X3))|~m1_subset_1(k1_fdiff_1(X1,k1_fdiff_1(X2,X4)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_numbers)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~r2_hidden(X4,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_18114, plain, (k1_real_1(k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),X2))=k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),X2)|~v1_xcmplx_0(X2)|~m1_subset_1(k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),X2),k1_numbers)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_18173, plain, (r1_xxreal_0(esk3_2(X1,X2),np__1)|r1_xxreal_0(np__1,esk3_2(X1,X2))|r1_xxreal_0(esk3_2(X1,X2),k6_numbers)|r2_fdiff_1(X2,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_18170, negated_conjecture, (r1_xxreal_0(esk3_2(esk1_0,X1),np__1)|r1_xxreal_0(np__1,esk3_2(esk1_0,X1))|r1_xxreal_0(esk3_2(esk1_0,X1),k6_numbers)|r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_18195, negated_conjecture, (r1_xxreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),X1),np__1)|r1_xxreal_0(np__1,k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),X1))|r1_xxreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),X1),k6_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_17402, plain, (v1_int_1(k7_xcmplx_0(np__2,np__1))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1007, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X6,k1_numbers)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1005, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k2_relset_1(k2_zfmisc_1(X6,X7),X8)),k1_zfmisc_1(k2_zfmisc_1(X9,X7)))|~v1_funct_1(k2_relset_1(k2_zfmisc_1(X6,X7),X8))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(k2_relset_1(k2_zfmisc_1(X6,X7),X8),k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v5_relat_1(X8,k2_zfmisc_1(X6,X7))|~v1_relat_1(X8))).#trainneg
cnf(i_0_1006, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k2_partfun2(X6,X7,X8)),k1_zfmisc_1(k2_zfmisc_1(X9,X6)))|~v1_funct_1(k2_partfun2(X6,X7,X8))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(k2_partfun2(X6,X7,X8),k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v2_funct_1(X8)|~v1_funct_1(X8)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))).#trainneg
cnf(i_0_16871, plain, (k1_real_1(k7_xcmplx_0(np__1,X1))=k1_real_1(k1_fdiff_1(k3_taylor_1,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers))).#trainneg
cnf(i_0_16815, plain, (k7_xcmplx_0(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)),np__1)=k4_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1))|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers))).#trainneg
cnf(i_0_11729, plain, (k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),np__1)=k1_real_1(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_13423, plain, (k7_xcmplx_0(k7_xcmplx_0(X1,np__1),np__1)=k1_real_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_13743, plain, (k1_real_1(k1_real_1(X1))=k7_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(X1,np__1),k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_1013, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k2_zfmisc_1(X6,X7)),k1_zfmisc_1(k2_zfmisc_1(X8,X7)))|~v1_funct_1(k2_zfmisc_1(X6,X7))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X8,X9)))|~m1_subset_1(k2_zfmisc_1(X6,X7),k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_17611, plain, (m1_subset_1(k7_xcmplx_0(np__2,np__1),k1_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_16383, plain, (v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(esk6_0)))|~v6_membered(k1_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_16384, plain, (v1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(esk6_0)))|~v2_membered(k1_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_13902, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|r1_xxreal_0(k1_real_1(X2),k1_real_1(X1))|r1_xxreal_0(k1_real_1(X2),k6_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1009, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k2_fdiff_1(X6,X7)),k1_zfmisc_1(k2_zfmisc_1(X8,k1_numbers)))|~v1_funct_1(k2_fdiff_1(X6,X7))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X8,X9)))|~m1_subset_1(k2_fdiff_1(X6,X7),k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1010, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k1_partfun1(X6,X7,X8,X9,X10,X11)),k1_zfmisc_1(k2_zfmisc_1(X12,X9)))|~v1_funct_1(k1_partfun1(X6,X7,X8,X9,X10,X11))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X12,X13)))|~m1_subset_1(k1_partfun1(X6,X7,X8,X9,X10,X11),k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X11)|~v1_funct_1(X10)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X8,X9)))|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(X6,X7))))).#trainneg
cnf(i_0_12341, plain, (k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),np__1)=k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_19123, plain, (k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1))=k4_xcmplx_0(np__1)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_19115, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_18952, plain, (v1_int_1(k4_xcmplx_0(k4_xcmplx_0(esk6_0)))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_18955, plain, (v7_ordinal1(k7_xcmplx_0(esk6_0,np__1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_6406, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|r1_xxreal_0(esk1_0,X1)|~v2_membered(k1_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1015, plain, (m1_subset_1(esk3_2(k1_relset_1(k1_numbers,k16_sin_cos),X1),k1_numbers)|r2_fdiff_1(X1,k1_relset_1(k1_numbers,k16_sin_cos))|~v1_funct_1(X1)|~r1_tarski(k1_relset_1(k1_numbers,k16_sin_cos),k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(k1_relset_1(k1_numbers,k16_sin_cos)))).#trainneg
cnf(i_0_19096, plain, (k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),np__1)=k4_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)))|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_16594, plain, (r1_xxreal_0(np__1,k5_numbers)|r1_xxreal_0(k5_numbers,np__1)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_16633, plain, (v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(np__2)))|~v6_membered(k1_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_16634, plain, (v1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(np__2)))|~v2_membered(k1_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_2041, plain, (r1_xxreal_0(k6_numbers,k1_seq_1(X1,X2))|r1_xxreal_0(k1_seq_1(X1,X2),k6_numbers)|~v2_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2724, negated_conjecture, (r1_xxreal_0(esk1_0,k1_relset_1(k1_numbers,k16_sin_cos))|r1_xxreal_0(k1_relset_1(k1_numbers,k16_sin_cos),esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers))|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1039, plain, (r1_fdiff_1(k3_relat_1(X5,X6),X7)|~r1_fdiff_1(X6,k1_seq_1(X5,X7))|~r1_fdiff_1(X5,X7)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X7,k1_numbers))).#trainneg
cnf(i_0_18821, plain, (k1_real_1(k7_xcmplx_0(X1,np__1))=k1_real_1(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(X1,np__1),k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_1043, plain, (m1_subset_1(esk3_2(k1_relset_1(k1_numbers,k19_sin_cos),X1),k1_numbers)|r2_fdiff_1(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~v1_funct_1(X1)|~r1_tarski(k1_relset_1(k1_numbers,k19_sin_cos),k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(k1_relset_1(k1_numbers,k19_sin_cos)))).#trainneg
cnf(i_0_19303, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|r1_xxreal_0(X1,esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2725, negated_conjecture, (r1_xxreal_0(esk1_0,k1_relset_1(k1_numbers,k19_sin_cos))|r1_xxreal_0(k1_relset_1(k1_numbers,k19_sin_cos),esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers))|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1057, plain, (m1_subset_1(k3_limfunc1(X1),k1_numbers)|~v1_xreal_0(k3_limfunc1(X1))|~v1_membered(k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1059, plain, (r2_fdiff_1(k1_partfun1(X3,X4,X5,X6,X1,k3_taylor_1),X2)|r1_xxreal_0(k1_seq_1(X1,esk3_2(X2,k1_partfun1(X3,X4,X5,X6,X1,k3_taylor_1))),k6_numbers)|~r1_fdiff_1(X1,esk3_2(X2,k1_partfun1(X3,X4,X5,X6,X1,k3_taylor_1)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k3_taylor_1)))|~m1_subset_1(k1_seq_1(X1,esk3_2(X2,k1_partfun1(X3,X4,X5,X6,X1,k3_taylor_1))),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_16173, plain, (k1_real_1(k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),X2))=k7_xcmplx_0(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)))|~m1_subset_1(k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_16788, plain, (k1_real_1(k7_xcmplx_0(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)),X2))=k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)),k1_numbers))).#trainneg
cnf(i_0_19290, plain, (v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(esk6_0)))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1097, plain, (v1_relat_1(esk3_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3))|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_19298, plain, (v1_int_1(k7_xcmplx_0(esk6_0,np__1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_5374, plain, (r1_xxreal_0(X1,k1_real_1(k10_real_1(X2,X3)))|r1_xxreal_0(k1_real_1(k10_real_1(X2,X3)),X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_13938, plain, (r1_xxreal_0(k1_real_1(k1_real_1(k1_real_1(X1))),X2)|r1_xxreal_0(X2,k1_real_1(k1_real_1(k1_real_1(X1))))|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1098, plain, (v4_relat_1(esk3_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3),X1)|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_13114, plain, (k7_xcmplx_0(k7_xcmplx_0(X1,np__1),np__1)=k4_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_16270, plain, (k7_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(np__1,X1)),np__1)=k1_real_1(k1_fdiff_1(k3_taylor_1,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers))).#trainneg
cnf(i_0_19539, plain, (m1_subset_1(a_0_0_fdiff_7,k1_numbers)|~v1_membered(k1_zfmisc_1(k1_numbers))|~v1_xreal_0(a_0_0_fdiff_7))).#trainneg
cnf(i_0_1101, plain, (v5_relat_1(esk3_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3),X2)|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_19365, plain, (v1_int_1(k4_xcmplx_0(k4_xcmplx_0(np__2)))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_19368, plain, (v7_ordinal1(k7_xcmplx_0(np__2,np__1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_19730, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k6_numbers,X1)),X2)|r1_xxreal_0(X2,k1_real_1(k10_real_1(k6_numbers,X1)))|~v1_xxreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1099, plain, (v3_valued_0(esk3_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3))|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_membered(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1100, plain, (v1_valued_0(esk3_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3))|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_membered(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6623, plain, (k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),k1_seq_1(k3_taylor_1,X1)),X1)=k1_real_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k1_seq_1(k3_taylor_1,X1))|~v1_xreal_0(X1)|~v1_xreal_0(k1_seq_1(k3_taylor_1,X1))|~m2_subset_1(k1_seq_1(k3_taylor_1,X1),k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),k1_seq_1(k3_taylor_1,X1)),k1_numbers)|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1)))|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1)),k1_numbers))).#trainneg
cnf(i_0_5349, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_19679, plain, (v1_xreal_0(k7_xcmplx_0(esk6_0,np__1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1104, plain, (v3_membered(esk3_2(k1_zfmisc_1(X1),X2))|r2_fdiff_1(X2,k1_zfmisc_1(X1))|~v3_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_15270, plain, (r1_xxreal_0(X1,k10_real_1(k1_real_1(X2),X3))|r1_xxreal_0(k10_real_1(k1_real_1(X2),X3),X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1960, plain, (r1_xxreal_0(k10_real_1(X1,X2),k10_real_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1103, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,esk3_2(k1_zfmisc_1(k2_zfmisc_1(X6,X7)),X8)),k1_zfmisc_1(k2_zfmisc_1(X9,X7)))|r2_fdiff_1(X8,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~v1_funct_1(esk3_2(k1_zfmisc_1(k2_zfmisc_1(X6,X7)),X8))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(esk3_2(k1_zfmisc_1(k2_zfmisc_1(X6,X7)),X8),k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X8)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X6,X7)),k1_zfmisc_1(k1_relset_1(k1_numbers,X8)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X6,X7)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X6,X7))))).#trainneg
cnf(i_0_1105, plain, (v1_membered(esk3_2(k1_zfmisc_1(X1),X2))|r2_fdiff_1(X2,k1_zfmisc_1(X1))|~v1_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_19973, plain, (k1_real_1(k7_xcmplx_0(np__1,X1))=k1_real_1(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_19962, plain, (k1_real_1(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)))=k7_xcmplx_0(k7_xcmplx_0(np__1,X1),np__1)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_19601, plain, (k1_real_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2))=k7_xcmplx_0(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)),X2)|~v1_xcmplx_0(k1_real_1(k1_fdiff_1(k3_taylor_1,X1)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),k1_numbers))).#trainneg
cnf(i_0_10900, plain, (k1_real_1(k7_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))),X2))=k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_19984, plain, (m1_subset_1(a_0_0_fdiff_7,k1_numbers)|~v1_xreal_0(a_0_0_fdiff_7)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1106, plain, (m1_subset_1(X1,X2)|r2_fdiff_1(X3,k1_zfmisc_1(X2))|~r2_hidden(X1,esk3_2(k1_zfmisc_1(X2),X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_19986, plain, (v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(np__2)))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_19994, plain, (v1_int_1(k7_xcmplx_0(np__2,np__1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_20034, plain, (r1_xxreal_0(k6_numbers,k1_real_1(k10_real_1(k6_numbers,X1)))|r1_xxreal_0(k1_real_1(k10_real_1(k6_numbers,X1)),k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_20361, plain, (k7_xcmplx_0(k1_real_1(k1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))),X1)=k1_real_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1))|~v1_xcmplx_0(k1_real_1(k1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k1_seq_1(k3_taylor_1,X1))|~v1_xreal_0(k1_seq_1(k3_taylor_1,X1))|~m2_subset_1(k1_seq_1(k3_taylor_1,X1),k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1),k1_numbers)|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1)))|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1)),k1_numbers))).#trainneg
cnf(i_0_12064, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_13471, plain, (k7_xcmplx_0(k7_xcmplx_0(k1_real_1(X1),X2),np__1)=k1_real_1(k7_xcmplx_0(X1,X2))|~m1_subset_1(k7_xcmplx_0(k1_real_1(X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1110, plain, (v1_int_1(esk3_2(X1,X2))|r2_fdiff_1(X2,X1)|~v5_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_19364, plain, (r1_xxreal_0(k5_numbers,np__1)|r1_xxreal_0(np__1,k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_20434, plain, (m1_subset_1(a_0_0_fdiff_7,k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_20443, plain, (v1_xreal_0(k7_xcmplx_0(np__2,np__1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_20041, plain, (r1_xxreal_0(X1,k1_real_1(k10_real_1(k6_numbers,X2)))|r1_xxreal_0(k1_real_1(k10_real_1(k6_numbers,X2)),X1)|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_20029, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(k10_real_1(k6_numbers,X2)))|r1_xxreal_0(k1_real_1(k10_real_1(k6_numbers,X2)),k1_real_1(X1))|~v1_xreal_0(X2)|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1111, plain, (v7_ordinal1(esk3_2(X1,X2))|r2_fdiff_1(X2,X1)|~v6_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1112, plain, (v1_xxreal_0(esk3_2(X1,X2))|r2_fdiff_1(X2,X1)|~v2_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_20538, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),np__1)=k1_real_1(X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k1_numbers))).#trainneg
cnf(i_0_20733, plain, (k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(X1),X2),np__1)=k1_real_1(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2),k1_numbers)|~m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_20834, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k1_numbers))).#trainneg
cnf(i_0_20743, plain, (k7_xcmplx_0(k1_real_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)),np__1)=k1_real_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1)),X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k1_real_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)),k1_numbers)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1)),X1),k1_numbers)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1)),k1_numbers)|~v1_xcmplx_0(k1_real_1(k1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))))|~v1_xcmplx_0(k1_seq_1(k3_taylor_1,X1))|~v1_xreal_0(k1_seq_1(k3_taylor_1,X1))|~v1_xreal_0(X1)|~m2_subset_1(k1_seq_1(k3_taylor_1,X1),k1_numbers,a_0_0_fdiff_7)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1),k1_numbers))).#trainneg
cnf(i_0_9488, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_20920, plain, (k1_real_1(k1_fdiff_1(k3_taylor_1,np__1))=k4_xcmplx_0(np__1)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7)|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1))|~m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1))),k1_numbers))).#trainneg
cnf(i_0_18972, plain, (v1_xxreal_0(k7_xcmplx_0(esk6_0,np__1))|~v5_membered(k1_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_20796, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k6_numbers,X1)),k1_real_1(X2))|r1_xxreal_0(k1_real_1(X2),k1_real_1(k10_real_1(k6_numbers,X1)))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_21209, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_21143, plain, (k1_real_1(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8311, plain, (v1_xcmplx_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X3),X2))|~v1_xreal_0(k1_fdiff_1(X1,X2))|~m1_subset_1(k1_fdiff_1(X3,k1_seq_1(X1,X2)),k1_numbers)|~v1_xreal_0(k1_fdiff_1(X3,k1_seq_1(X1,X2)))|~r1_fdiff_1(X3,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_15959, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),k4_xcmplx_0(np__2))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_21210, plain, (k1_real_1(esk6_0)=k4_xcmplx_0(esk6_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_21331, plain, (k1_real_1(k1_fdiff_1(k3_taylor_1,np__1))=k4_xcmplx_0(np__1)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1),np__1),k1_numbers))).#trainneg
cnf(i_0_1118, negated_conjecture, (r1_tarski(esk1_0,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k19_sin_cos)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_19385, plain, (v1_xxreal_0(k7_xcmplx_0(np__2,np__1))|~v5_membered(k1_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_20043, plain, (r1_xxreal_0(X1,k1_real_1(k10_real_1(k6_numbers,X2)))|r1_xxreal_0(k1_real_1(k10_real_1(k6_numbers,X2)),X1)|~v1_xreal_0(X2)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_21586, plain, (k4_xcmplx_0(k1_real_1(X1))=X1|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_21495, negated_conjecture, (k4_xcmplx_0(k10_real_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~v1_int_1(k10_real_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0))|~v1_xcmplx_0(k10_real_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0)))).#trainneg
cnf(i_0_15965, plain, (m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(esk6_0)),k1_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_21211, plain, (k1_real_1(np__2)=k4_xcmplx_0(np__2)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_16134, plain, (m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(np__2)),k1_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_16379, plain, (r2_hidden(k4_xcmplx_0(k4_xcmplx_0(esk6_0)),k1_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_20795, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k6_numbers,X1)),X2)|r1_xxreal_0(X2,k1_real_1(k10_real_1(k6_numbers,X1)))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_21703, plain, (k1_real_1(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k1_real_1(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_21728, plain, (k1_real_1(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k1_real_1(X1),X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_22134, plain, (k1_real_1(k7_xcmplx_0(k1_real_1(X1),X2))=k7_xcmplx_0(X1,X2)|~m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_21742, plain, (k1_real_1(k1_real_1(X1))=X1|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_21655, plain, (k1_real_1(esk6_0)=k4_xcmplx_0(esk6_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1152, plain, (v1_xreal_0(k3_limfunc1(X1))|~v5_membered(k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_21840, plain, (v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(esk6_0)))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2040, plain, (r1_xxreal_0(X1,k1_seq_1(X2,X3))|r1_xxreal_0(k1_seq_1(X2,X3),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1178, negated_conjecture, (v1_xcmplx_0(esk3_2(esk1_0,X1))|r2_fdiff_1(X1,esk1_0)|~v1_membered(k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_1179, negated_conjecture, (v1_int_1(esk3_2(esk1_0,X1))|r2_fdiff_1(X1,esk1_0)|~v5_membered(k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_21740, plain, (k1_real_1(k1_real_1(X1))=k7_xcmplx_0(X1,np__1)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_22136, plain, (k4_xcmplx_0(k7_xcmplx_0(k1_real_1(X1),X2))=k7_xcmplx_0(X1,X2)|~m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_22818, plain, (k7_xcmplx_0(k1_real_1(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_real_1(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_22915, plain, (k4_xcmplx_0(k1_real_1(X1))=k7_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(k1_real_1(X1)))).#trainneg
cnf(i_0_22760, plain, (k7_xcmplx_0(k1_real_1(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~m1_subset_1(k7_xcmplx_0(k1_real_1(X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_21874, plain, (k1_real_1(np__2)=k4_xcmplx_0(np__2)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_21836, negated_conjecture, (k4_xcmplx_0(k10_real_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)|~v1_xcmplx_0(k10_real_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~v5_membered(k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_1180, negated_conjecture, (v7_ordinal1(esk3_2(esk1_0,X1))|r2_fdiff_1(X1,esk1_0)|~v6_membered(k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_21877, plain, (v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(np__2)))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_22589, plain, (r1_xxreal_0(k1_seq_1(X1,X2),X3)|r1_xxreal_0(X3,k1_seq_1(X1,X2))|~v1_xxreal_0(X3)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_12881, plain, (r1_xxreal_0(k1_seq_1(X1,X2),X3)|r1_xxreal_0(X3,k1_seq_1(X1,X2))|~v1_xcmplx_0(X3)|~v1_xreal_0(X3)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1181, negated_conjecture, (v1_xxreal_0(esk3_2(esk1_0,X1))|r2_fdiff_1(X1,esk1_0)|~v2_membered(k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_22815, plain, (k7_xcmplx_0(k1_real_1(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_21541, plain, (k1_real_1(k4_xcmplx_0(k4_xcmplx_0(X1)))=k1_real_1(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_22587, plain, (v1_xcmplx_0(k7_xcmplx_0(esk6_0,np__1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_21837, negated_conjecture, (k4_xcmplx_0(k10_real_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0),esk2_0)|~v1_xcmplx_0(k10_real_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),esk2_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk1_0)|~v6_membered(k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_23367, plain, (v1_xcmplx_0(k7_xcmplx_0(np__2,np__1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_23423, plain, (r1_xxreal_0(X1,k1_seq_1(k16_sin_cos,X2))|r1_xxreal_0(k1_seq_1(k16_sin_cos,X2),X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_23427, plain, (r1_xxreal_0(X1,k1_seq_1(k19_sin_cos,X2))|r1_xxreal_0(k1_seq_1(k19_sin_cos,X2),X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1187, plain, (v1_int_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2))|~v1_int_1(k7_xcmplx_0(X1,X2))|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_22591, negated_conjecture, (r1_xxreal_0(esk2_0,esk3_2(esk1_0,X1))|r1_xxreal_0(esk3_2(esk1_0,X1),esk2_0)|r2_fdiff_1(X1,esk1_0)|~v1_xreal_0(esk3_2(esk1_0,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_21843, plain, (v1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(esk6_0)))|~v2_membered(k1_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_23787, plain, (r1_xxreal_0(k1_seq_1(k16_sin_cos,X1),k6_numbers)|r1_xxreal_0(k6_numbers,k1_seq_1(k16_sin_cos,X1)))).#trainneg
cnf(i_0_23838, plain, (r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),k6_numbers)|r1_xxreal_0(k6_numbers,k1_seq_1(k19_sin_cos,X1)))).#trainneg
cnf(i_0_1195, plain, (v1_xreal_0(esk3_2(a_0_0_fdiff_7,X1))|r2_fdiff_1(X1,a_0_0_fdiff_7)|~v1_funct_1(X1)|~r1_tarski(a_0_0_fdiff_7,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(a_0_0_fdiff_7))).#trainneg
cnf(i_0_23096, plain, (k1_real_1(k7_xcmplx_0(X1,np__1))=k1_real_1(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~v1_xcmplx_0(k1_real_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_23723, plain, (k1_real_1(k7_xcmplx_0(X1,np__1))=k1_real_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_16396, plain, (r1_xxreal_0(k4_xcmplx_0(esk6_0),k4_xcmplx_0(esk6_0))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_15711, plain, (r1_xxreal_0(k6_numbers,k4_xcmplx_0(esk6_0))|r1_xxreal_0(k4_xcmplx_0(esk6_0),k6_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1197, plain, (k1_relset_1(X1,k2_zfmisc_1(X1,X2))=k9_xtuple_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_21880, plain, (v1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(np__2)))|~v2_membered(k1_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_23428, plain, (r1_xxreal_0(X1,k1_seq_1(k3_taylor_1,X2))|r1_xxreal_0(k1_seq_1(k3_taylor_1,X2),X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_23791, plain, (r1_xxreal_0(k1_seq_1(k16_sin_cos,X1),X2)|r1_xxreal_0(X2,k1_seq_1(k16_sin_cos,X1))|~v1_xreal_0(X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1198, plain, (m1_subset_1(k1_relset_1(X1,k2_zfmisc_1(X1,X2)),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_22590, negated_conjecture, (r1_xxreal_0(esk3_2(esk1_0,X1),k6_numbers)|r1_xxreal_0(k6_numbers,esk3_2(esk1_0,X1))|r2_fdiff_1(X1,esk1_0)|~v1_xreal_0(esk3_2(esk1_0,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_22916, plain, (k4_xcmplx_0(k1_real_1(X1))=k7_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers))).#trainneg
cnf(i_0_24409, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_24410, plain, (m1_subset_1(k9_xtuple_0(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_15713, plain, (r1_xxreal_0(esk6_0,k4_xcmplx_0(esk6_0))|r1_xxreal_0(k4_xcmplx_0(esk6_0),esk6_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1209, plain, (v1_int_1(k1_fdiff_1(k19_sin_cos,X1))|~v6_membered(k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_23951, plain, (v1_xxreal_0(k7_xcmplx_0(esk6_0,np__1))|~v2_membered(k1_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_24323, plain, (r1_xxreal_0(k1_seq_1(k3_taylor_1,X1),k6_numbers)|r1_xxreal_0(k6_numbers,k1_seq_1(k3_taylor_1,X1)))).#trainneg
cnf(i_0_1232, plain, (v1_xreal_0(k7_xcmplx_0(k4_xcmplx_0(X1),X2))|~v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(X1),X2))|~m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_24281, plain, (r1_xxreal_0(k1_seq_1(k3_taylor_1,X1),esk3_2(X2,k3_taylor_1))|r1_xxreal_0(esk3_2(X2,k3_taylor_1),k1_seq_1(k3_taylor_1,X1))|r2_fdiff_1(k3_taylor_1,X2)|~v1_xreal_0(esk3_2(X2,k3_taylor_1))|~r1_tarski(X2,a_0_0_fdiff_7)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2))).#trainneg
cnf(i_0_21700, plain, (k1_real_1(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k1_real_1(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(k7_xcmplx_0(k1_real_1(X1),X2),k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_22472, plain, (k7_xcmplx_0(k1_real_1(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~v1_xcmplx_0(k1_real_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_24581, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k9_xtuple_0(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_15714, plain, (r1_xxreal_0(np__2,k4_xcmplx_0(esk6_0))|r1_xxreal_0(k4_xcmplx_0(esk6_0),np__2)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_16596, plain, (r1_xxreal_0(np__1,a_0_0_fdiff_7)|r1_xxreal_0(a_0_0_fdiff_7,np__1)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_24275, plain, (v1_xxreal_0(k7_xcmplx_0(np__2,np__1))|~v2_membered(k1_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_24371, plain, (r1_xxreal_0(k1_real_1(X1),k1_seq_1(k16_sin_cos,X2))|r1_xxreal_0(k1_seq_1(k16_sin_cos,X2),k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_23842, plain, (r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),X2)|r1_xxreal_0(X2,k1_seq_1(k19_sin_cos,X1))|~v1_xreal_0(X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1257, plain, (v1_int_1(esk5_1(X1))|~v6_membered(k1_numbers)|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_22106, plain, (k7_xcmplx_0(k1_real_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2))),X2)=k1_real_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),X2))|~v1_xcmplx_0(k1_real_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2))))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2)))|~m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),X2),k1_numbers)|~m1_subset_1(k4_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2))),k1_numbers)|~v1_xreal_0(X2)|~m2_subset_1(X2,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2))|~r1_fdiff_1(k3_taylor_1,X2)|~v1_funct_1(X1)|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_23461, plain, (k1_real_1(k1_real_1(X1))=k4_xcmplx_0(k1_real_1(X1))|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(X1,np__1)),k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_23953, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos),a_0_0_fdiff_7)|~r1_tarski(a_0_0_fdiff_7,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos)))|~v3_rcomp_1(a_0_0_fdiff_7)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_16646, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),k4_xcmplx_0(np__2))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_15956, negated_conjecture, (r1_xxreal_0(esk2_0,k4_xcmplx_0(np__2))|r1_xxreal_0(k4_xcmplx_0(np__2),esk2_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1259, plain, (v1_int_1(k3_limfunc1(X1))|~v6_membered(k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_25026, plain, (r1_xxreal_0(a_0_0_fdiff_7,np__1)|r1_xxreal_0(np__1,a_0_0_fdiff_7)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_16629, plain, (r2_hidden(k4_xcmplx_0(k4_xcmplx_0(np__2)),k1_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_15958, plain, (r1_xxreal_0(k6_numbers,k4_xcmplx_0(np__2))|r1_xxreal_0(k4_xcmplx_0(np__2),k6_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_25135, plain, (r1_xxreal_0(k1_real_1(X1),k1_seq_1(k19_sin_cos,X2))|r1_xxreal_0(k1_seq_1(k19_sin_cos,X2),k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_25207, plain, (k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2))),X2)=k1_real_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),X2))|~v1_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2))))|~v1_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2)))|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m2_subset_1(X2,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2))|~r1_fdiff_1(k3_taylor_1,X2)|~v1_funct_1(X1)|~m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),X2),k1_numbers)|~m1_subset_1(k4_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2))),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2)),k1_numbers))).#trainneg
cnf(i_0_22811, plain, (k1_real_1(k1_real_1(X1))=k7_xcmplx_0(X1,np__1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_25455, plain, (k3_xcmplx_0(X1,k1_fdiff_1(k3_taylor_1,X2))=k7_xcmplx_0(X1,X2)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m2_subset_1(X2,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_25469, plain, (k7_xcmplx_0(k1_fdiff_1(X1,X2),k1_seq_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~v1_xreal_0(k1_fdiff_1(X1,X2))|~r1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,X2))|~m2_subset_1(k1_seq_1(X1,X2),k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers))).#trainneg
cnf(i_0_21866, plain, (r1_xxreal_0(k4_xcmplx_0(esk6_0),k4_xcmplx_0(esk6_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_15960, plain, (r1_xxreal_0(esk6_0,k4_xcmplx_0(np__2))|r1_xxreal_0(k4_xcmplx_0(np__2),esk6_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1350, plain, (r1_xxreal_0(X1,esk5_1(X2))|r1_xxreal_0(esk5_1(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~r2_hidden(X2,a_0_0_fdiff_7))).#trainneg
cnf(i_0_21838, plain, (r2_hidden(k4_xcmplx_0(k4_xcmplx_0(esk6_0)),k1_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_21875, plain, (r2_hidden(k4_xcmplx_0(k4_xcmplx_0(np__2)),k1_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_24327, plain, (r1_xxreal_0(k1_seq_1(k3_taylor_1,X1),X2)|r1_xxreal_0(X2,k1_seq_1(k3_taylor_1,X1))|~v1_xreal_0(X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_21903, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),k4_xcmplx_0(np__2))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_15961, plain, (r1_xxreal_0(np__2,k4_xcmplx_0(np__2))|r1_xxreal_0(k4_xcmplx_0(np__2),np__2)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1351, plain, (r1_xxreal_0(esk5_1(X1),esk5_1(X1))|~v2_membered(k1_numbers)|~r2_hidden(X1,a_0_0_fdiff_7))).#trainneg
cnf(i_0_21915, plain, (r2_hidden(k7_xcmplx_0(esk6_0,np__1),k1_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_25087, plain, (r1_xxreal_0(k1_seq_1(k16_sin_cos,X1),k1_real_1(k1_real_1(X2)))|r1_xxreal_0(k1_real_1(k1_real_1(X2)),k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_25444, plain, (r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),k1_real_1(k1_real_1(X2)))|r1_xxreal_0(k1_real_1(k1_real_1(X2)),k1_seq_1(k19_sin_cos,X1))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1309, plain, (r2_fdiff_1(k3_relat_1(X5,X6),X7)|~r1_fdiff_1(X6,k1_seq_1(X5,esk3_2(X7,k3_relat_1(X5,X6))))|~r1_fdiff_1(X5,esk3_2(X7,k3_relat_1(X5,X6)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~r1_tarski(X7,k1_relset_1(k1_numbers,k3_relat_1(X5,X6)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X7))).#trainneg
cnf(i_0_22422, plain, (k1_real_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_25798, negated_conjecture, (r1_xxreal_0(esk2_0,k1_seq_1(k3_taylor_1,X1))|r1_xxreal_0(k1_seq_1(k3_taylor_1,X1),esk2_0))).#trainneg
cnf(i_0_26085, plain, (k1_real_1(k4_xcmplx_0(X1))=k7_xcmplx_0(X1,np__1)|~v1_int_1(X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_26032, plain, (k1_real_1(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1353, plain, (r1_xxreal_0(X1,k3_limfunc1(X2))|r1_xxreal_0(k3_limfunc1(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1354, plain, (r1_xxreal_0(k3_limfunc1(X1),k3_limfunc1(X1))|~v2_membered(k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_25828, plain, (r1_xxreal_0(k1_real_1(X1),k1_seq_1(k3_taylor_1,X2))|r1_xxreal_0(k1_seq_1(k3_taylor_1,X2),k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_23407, plain, (r1_xxreal_0(k6_numbers,k1_seq_1(X1,X2))|r1_xxreal_0(k1_seq_1(X1,X2),k6_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_26400, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),np__1)=k1_real_1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_24857, plain, (k1_real_1(k7_xcmplx_0(X1,np__1))=k7_xcmplx_0(k1_real_1(X1),np__1)|~v1_xcmplx_0(k1_real_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers))).#trainneg
cnf(i_0_26403, plain, (r1_xxreal_0(a_0_0_fdiff_7,a_0_0_fdiff_7)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_16597, plain, (r1_xxreal_0(np__1,k4_xcmplx_0(esk6_0))|r1_xxreal_0(k4_xcmplx_0(esk6_0),np__1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1376, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_relset_1(k2_zfmisc_1(X5,X6),X7),k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(X5,k1_numbers)))|~v1_funct_1(k2_relset_1(k2_zfmisc_1(X5,X6),X7))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_relset_1(k2_zfmisc_1(X5,X6),X7),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v5_relat_1(X7,k2_zfmisc_1(X5,X6))|~v1_relat_1(X7))).#trainneg
cnf(i_0_25376, plain, (r2_hidden(k7_xcmplx_0(np__2,np__1),k1_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_26448, plain, (r1_xxreal_0(k1_seq_1(k3_taylor_1,X1),k1_real_1(k1_real_1(X2)))|r1_xxreal_0(k1_real_1(k1_real_1(X2)),k1_seq_1(k3_taylor_1,X1))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_25717, plain, (k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2)=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X2),X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m2_subset_1(X2,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers))).#trainneg
cnf(i_0_26544, plain, (k4_xcmplx_0(k1_real_1(X1))=k7_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1377, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_partfun2(X5,X6,X7),k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(X6,k1_numbers)))|~v1_funct_1(k2_partfun2(X5,X6,X7))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_partfun2(X5,X6,X7),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v2_funct_1(X7)|~v1_funct_1(X7)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_25883, plain, (r1_xxreal_0(k1_real_1(k1_real_1(k1_real_1(X1))),k1_seq_1(k16_sin_cos,X2))|r1_xxreal_0(k1_seq_1(k16_sin_cos,X2),k1_real_1(k1_real_1(k1_real_1(X1))))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_25931, plain, (r1_xxreal_0(k1_real_1(k1_real_1(k1_real_1(X1))),k1_seq_1(k19_sin_cos,X2))|r1_xxreal_0(k1_seq_1(k19_sin_cos,X2),k1_real_1(k1_real_1(k1_real_1(X1))))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1323, plain, (k3_xcmplx_0(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_fdiff_1(X4,k1_fdiff_1(X1,X3)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X4),X3)|~r1_fdiff_1(X4,k1_fdiff_1(X1,X3))|~r1_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(X4)|~v1_funct_1(k2_fdiff_1(X1,X2))|~m1_subset_1(k1_fdiff_1(X4,k1_fdiff_1(X1,X3)),k1_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_numbers)|~m1_subset_1(X3,k1_numbers)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1375, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,esk3_2(k1_zfmisc_1(k2_zfmisc_1(X5,X6)),X7),k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(X5,k1_numbers)))|r2_fdiff_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~v1_funct_1(esk3_2(k1_zfmisc_1(k2_zfmisc_1(X5,X6)),X7))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(esk3_2(k1_zfmisc_1(k2_zfmisc_1(X5,X6)),X7),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X7)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X5,X6)),k1_zfmisc_1(k1_relset_1(k1_numbers,X7)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X5,X6)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_26968, plain, (k1_real_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X2),X1))=k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m2_subset_1(X2,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_27427, plain, (k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X2)),X1)=k4_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2))|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(X2,k1_numbers,a_0_0_fdiff_7)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_27428, plain, (k1_real_1(k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X2)),X1))=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(X2,k1_numbers,a_0_0_fdiff_7)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_27498, plain, (k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1)),X1)=k1_real_1(k1_fdiff_1(k3_taylor_1,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_21556, plain, (r1_xxreal_0(k6_numbers,X1)|r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_14878, plain, (m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k1_numbers)|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_16598, plain, (r1_xxreal_0(np__1,k4_xcmplx_0(np__2))|r1_xxreal_0(k4_xcmplx_0(np__2),np__1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1378, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_25777, plain, (r2_hidden(k7_xcmplx_0(esk6_0,np__1),k1_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1379, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1380, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(k2_fdiff_1(X5,X6))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_fdiff_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_27662, plain, (k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),np__1)=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X2),X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),k1_numbers)|~m2_subset_1(X2,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X2),X1),k1_numbers))).#trainneg
cnf(i_0_27664, plain, (k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),X2)=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X2),X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),X2),k1_numbers)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers)|~v1_xreal_0(X2)|~m2_subset_1(X2,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X2),X1),k1_numbers))).#trainneg
cnf(i_0_27608, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),X2))=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X2),X1)|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m2_subset_1(X2,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X2),X1),k1_numbers))).#trainneg
cnf(i_0_27855, plain, (v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_27703, plain, (k1_real_1(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1)))=k4_xcmplx_0(k1_real_1(k1_fdiff_1(k3_taylor_1,np__1)))|~v1_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1)))|~m1_subset_1(k4_xcmplx_0(k1_real_1(k1_fdiff_1(k3_taylor_1,np__1))),k1_numbers)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_1382, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_25783, plain, (r2_hidden(k7_xcmplx_0(np__2,np__1),k1_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1381, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,X9,X10),k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(X5,k1_numbers)))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,X9,X10))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k1_partfun1(X5,X6,X7,X8,X9,X10),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X10)|~v1_funct_1(X9)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_1384, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_zfmisc_1(X5,X6),k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(X5,k1_numbers)))|~v1_funct_1(k2_zfmisc_1(X5,X6))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_zfmisc_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_27667, plain, (k7_xcmplx_0(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),np__1)=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X2),X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1)),X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m2_subset_1(X2,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X2),X1),k1_numbers))).#trainneg
cnf(i_0_27856, plain, (v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_28658, negated_conjecture, (k1_fdiff_1(k3_relat_1(k3_taylor_1,k19_sin_cos),esk2_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k19_sin_cos),esk2_0)|~v1_xreal_0(k1_seq_1(k3_taylor_1,esk2_0))|~m1_subset_1(k1_seq_1(k16_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m2_subset_1(esk2_0,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0))|~r1_fdiff_1(k3_taylor_1,esk2_0)|~m1_subset_1(k1_fdiff_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_1386, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_relset_1(k2_zfmisc_1(X5,X6),X7),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X5,k1_numbers)))|~v1_funct_1(k2_relset_1(k2_zfmisc_1(X5,X6),X7))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_relset_1(k2_zfmisc_1(X5,X6),X7),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v5_relat_1(X7,k2_zfmisc_1(X5,X6))|~v1_relat_1(X7))).#trainneg
cnf(i_0_1385, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,esk3_2(k1_zfmisc_1(k2_zfmisc_1(X5,X6)),X7),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X5,k1_numbers)))|r2_fdiff_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~v1_funct_1(esk3_2(k1_zfmisc_1(k2_zfmisc_1(X5,X6)),X7))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(esk3_2(k1_zfmisc_1(k2_zfmisc_1(X5,X6)),X7),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X7)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X5,X6)),k1_zfmisc_1(k1_relset_1(k1_numbers,X7)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X5,X6)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_1387, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_partfun2(X5,X6,X7),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X6,k1_numbers)))|~v1_funct_1(k2_partfun2(X5,X6,X7))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_partfun2(X5,X6,X7),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v2_funct_1(X7)|~v1_funct_1(X7)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_28637, plain, (k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2))),X2)=k1_real_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),X2))|~v1_xcmplx_0(X2)|~m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),X2),k1_numbers)|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2)),k1_numbers)|~v1_xreal_0(X2)|~m2_subset_1(X2,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2))|~r1_fdiff_1(k3_taylor_1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_28638, plain, (k7_xcmplx_0(k1_real_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2))),X2)=k1_real_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),X2))|~v1_xcmplx_0(X2)|~m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),X2),k1_numbers)|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2)),k1_numbers)|~v1_xreal_0(X2)|~m2_subset_1(X2,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2))|~r1_fdiff_1(k3_taylor_1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_27430, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X2)),X1))=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(X2,k1_numbers,a_0_0_fdiff_7)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_25681, plain, (k1_real_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),np__1)|~v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers))).#trainneg
cnf(i_0_29314, plain, (k7_xcmplx_0(k7_xcmplx_0(X1,np__1),np__1)=k1_real_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(k7_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(X1,np__1)),k1_numbers)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_27858, plain, (v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1388, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1389, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1390, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(k2_fdiff_1(X5,X6))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_fdiff_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_23447, plain, (k4_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(X1,np__1)))=k4_xcmplx_0(k1_real_1(X1))|~v1_xcmplx_0(k1_real_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_27484, plain, (k7_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1)),X1)=k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_numbers)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_16599, plain, (r1_xxreal_0(np__1,k4_xcmplx_0(np__1))|r1_xxreal_0(k4_xcmplx_0(np__1),np__1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1392, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1391, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,X9,X10),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X5,k1_numbers)))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,X9,X10))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k1_partfun1(X5,X6,X7,X8,X9,X10),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X10)|~v1_funct_1(X9)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_1393, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_28642, plain, (k7_xcmplx_0(k1_real_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2))),X2)=k1_real_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2)),k1_numbers)|~m2_subset_1(X2,k1_numbers,a_0_0_fdiff_7)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2))|~r1_fdiff_1(k3_taylor_1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_23613, plain, (k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))=k1_real_1(X1)|~m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_23444, plain, (k1_real_1(k4_xcmplx_0(k7_xcmplx_0(X1,np__1)))=k1_real_1(k1_real_1(X1))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_27846, plain, (r1_xxreal_0(k4_xcmplx_0(esk6_0),k6_numbers)|r1_xxreal_0(k6_numbers,k4_xcmplx_0(esk6_0))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1394, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_zfmisc_1(X5,X6),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X5,k1_numbers)))|~v1_funct_1(k2_zfmisc_1(X5,X6))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_zfmisc_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_26892, plain, (r1_xxreal_0(k1_real_1(k1_real_1(k1_real_1(X1))),k1_seq_1(k3_taylor_1,X2))|r1_xxreal_0(k1_seq_1(k3_taylor_1,X2),k1_real_1(k1_real_1(k1_real_1(X1))))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_24374, plain, (r1_xxreal_0(k1_seq_1(X1,X2),k1_seq_1(k16_sin_cos,X3))|r1_xxreal_0(k1_seq_1(k16_sin_cos,X3),k1_seq_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1401, plain, (r1_xxreal_0(X1,k5_numbers)|r1_xxreal_0(k5_numbers,X1)|~v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1408, plain, (k7_xcmplx_0(k7_xcmplx_0(np__1,X2),X1)=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~v1_xcmplx_0(k7_xcmplx_0(np__1,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_30278, plain, (k1_real_1(k4_xcmplx_0(X1))=k4_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_30790, plain, (k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))=k1_real_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_30792, plain, (k1_real_1(k4_xcmplx_0(k4_xcmplx_0(X1)))=k1_real_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_25138, plain, (r1_xxreal_0(k1_seq_1(X1,X2),k1_seq_1(k19_sin_cos,X3))|r1_xxreal_0(k1_seq_1(k19_sin_cos,X3),k1_seq_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_25831, plain, (r1_xxreal_0(k1_seq_1(X1,X2),k1_seq_1(k3_taylor_1,X3))|r1_xxreal_0(k1_seq_1(k3_taylor_1,X3),k1_seq_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1414, plain, (m1_subset_1(k7_xcmplx_0(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_30189, plain, (k1_real_1(k1_real_1(X1))=k4_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_27847, plain, (r1_xxreal_0(k4_xcmplx_0(esk6_0),k6_numbers)|r1_xxreal_0(k6_numbers,k4_xcmplx_0(esk6_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_29951, plain, (k1_real_1(k1_fdiff_1(k3_taylor_1,np__1))=k4_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1))|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7)|~m1_subset_1(k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1),np__1),k1_numbers))).#trainneg
cnf(i_0_1433, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_relset_1(k2_zfmisc_1(X5,X6),X7),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X5,k1_numbers)))|~v1_funct_1(k2_relset_1(k2_zfmisc_1(X5,X6),X7))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_relset_1(k2_zfmisc_1(X5,X6),X7),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v5_relat_1(X7,k2_zfmisc_1(X5,X6))|~v1_relat_1(X7))).#trainneg
cnf(i_0_19140, plain, (v1_int_1(k1_fdiff_1(k3_taylor_1,np__1))|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,np__1))|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_24373, plain, (r1_xxreal_0(k10_real_1(X1,X2),k1_seq_1(k16_sin_cos,X3))|r1_xxreal_0(k1_seq_1(k16_sin_cos,X3),k10_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_25137, plain, (r1_xxreal_0(k10_real_1(X1,X2),k1_seq_1(k19_sin_cos,X3))|r1_xxreal_0(k1_seq_1(k19_sin_cos,X3),k10_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1419, plain, (k7_xcmplx_0(k1_fdiff_1(X2,k1_seq_1(k3_taylor_1,X1)),X1)=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X2),X1)|~v1_xreal_0(k1_fdiff_1(X2,k1_seq_1(k3_taylor_1,X1)))|~r1_fdiff_1(X2,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~v1_xcmplx_0(k1_fdiff_1(X2,k1_seq_1(k3_taylor_1,X1)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_7))).#trainneg
cnf(i_0_30326, plain, (k1_real_1(k1_real_1(X1))=k4_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_31507, negated_conjecture, (a_0_0_fdiff_7=k6_numbers)).#trainneg
cnf(i_0_19, plain, (k10_real_1(np__1,X1)=k1_fdiff_1(k3_taylor_1,X1)|~m2_subset_1(X1,k1_numbers,k6_numbers))).#trainneg
cnf(i_0_47, plain, (esk5_1(X1)=X1|~r2_hidden(X1,k6_numbers))).#trainneg
cnf(i_0_2962, negated_conjecture, (m1_subset_1(esk2_0,k6_numbers))).#trainneg
cnf(i_0_18, plain, (~m2_subset_1(X1,k1_numbers,k6_numbers)|~r1_xxreal_0(k1_fdiff_1(k3_taylor_1,X1),k6_numbers))).#trainneg
cnf(i_0_627, plain, (m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,k6_numbers))).#trainneg
cnf(i_0_864, plain, (m1_subset_1(np__1,k6_numbers))).#trainneg
cnf(i_0_23, plain, (k1_relset_1(k1_numbers,k3_taylor_1)=k6_numbers)).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 104
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 159
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 157
# Processed clauses                    : 3611
# ...of these trivial                  : 44
# ...subsumed                          : 1855
# ...remaining for further processing  : 1712
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 136
# Backward-rewritten                   : 348
# Generated clauses                    : 31842
# ...of the previous two non-trivial   : 30734
# Contextual simplify-reflections      : 300
# Paramodulations                      : 31836
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1227
#    Positive orientable unit clauses  : 159
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 1061
# Current number of unprocessed clauses: 26985
# ...number of literals in the above   : 181900
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2196
# Proof object given clauses           : 149
# Proof search given clauses           : 1712
# Clause-clause subsumption calls (NU) : 708789
# Rec. Clause-clause subsumption calls : 71772
# Non-unit clause-clause subsumptions  : 2255
# Unit Clause-clause subsumption calls : 8525
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 49
# BW rewrite match successes           : 22
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 826122

# -------------------------------------------------
# User time                : 51.808 s
# System time              : 1.798 s
# Total time               : 53.606 s
# Maximum resident set size: 432716 pages
51.80user 1.89system 0:55.18elapsed 97%CPU (0avgtext+0avgdata 433296maxresident)k
0inputs+952outputs (0major+495712minor)pagefaults 0swaps
