2020-05-25 01:24:42.638986: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_37
2020-05-25 01:24:42.747476: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-05-25 01:24:43.589893: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-05-25 01:24:45.024050: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 2385064 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_37' loaded (query=256; context=768; weigths=real; len_mult=0.000000).
# Preprocessing time       : 3.109 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t33_mesfunc9, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,k4_partfun1(X1,k7_numbers)))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X1,k7_numbers)))))=>![X3]:(m1_subset_1(X3,X1)=>![X4]:((((v2_mesfunc9(X2,X1)&v1_mesfunc8(X2,X1,k7_numbers))&r1_tarski(X4,k1_relset_1(X1,k4_mesfunc5(X1,k7_numbers,X2,k6_numbers))))&r2_hidden(X3,X4))=>((((((((v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))=>v7_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3)))&(v7_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))=>v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))))&(v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))=>v8_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))))&(v8_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))=>v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))))&(v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))=>v9_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))))&(v9_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))=>v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))))&(v10_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))=>v10_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))))&(v10_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))=>v10_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3))))))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', t33_mesfunc9)).
fof(dt_k3_mesfunc5, axiom, ![X1, X2, X3]:(((((~(v1_xboole_0(X1))&v1_funct_1(X2))&v1_funct_2(X2,k5_numbers,k4_partfun1(X1,k7_numbers)))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X1,k7_numbers)))))&m1_subset_1(X3,X1))=>((v1_funct_1(k3_mesfunc5(X1,X2,X3))&v1_funct_2(k3_mesfunc5(X1,X2,X3),k5_numbers,k7_numbers))&m1_subset_1(k3_mesfunc5(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', dt_k3_mesfunc5)).
fof(dt_k4_mesfunc9, axiom, ![X1, X2]:((((~(v1_xboole_0(X1))&v1_funct_1(X2))&v1_funct_2(X2,k5_numbers,k4_partfun1(X1,k7_numbers)))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X1,k7_numbers)))))=>((v1_funct_1(k4_mesfunc9(X1,X2))&v1_funct_2(k4_mesfunc9(X1,X2),k5_numbers,k4_partfun1(X1,k7_numbers)))&m1_subset_1(k4_mesfunc9(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X1,k7_numbers)))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', dt_k4_mesfunc9)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', t4_subset)).
fof(dt_k1_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>m1_subset_1(k1_relset_1(X1,X2),k1_zfmisc_1(X1))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', dt_k1_relset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', t3_subset)).
fof(t7_boole, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', t7_boole)).
fof(t32_mesfunc9, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,k4_partfun1(X1,k7_numbers)))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X1,k7_numbers)))))=>![X3]:(v7_ordinal1(X3)=>![X4]:(m1_subset_1(X4,X1)=>![X5]:((((v2_mesfunc9(X2,X1)&v1_mesfunc8(X2,X1,k7_numbers))&r1_tarski(X5,k1_relset_1(X1,k4_mesfunc5(X1,k7_numbers,X2,k6_numbers))))&r2_hidden(X4,X5))=>k1_mesfunc9(k2_mesfunc9(k3_mesfunc5(X1,X2,X4)),X3)=k1_mesfunc9(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X4),X3)))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', t32_mesfunc9)).
fof(cc17_valued_0, axiom, ![X1, X2]:(v2_membered(X2)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v2_valued_0(X3))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', cc17_valued_0)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', cc2_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', fc6_relat_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', t2_subset)).
fof(cc1_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', cc1_subset_1)).
fof(dt_k4_mesfunc5, axiom, ![X1, X2, X3, X4]:((((v1_funct_1(X3)&v1_funct_2(X3,k5_numbers,k3_rfunct_3(X1,X2)))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k3_rfunct_3(X1,X2)))))&v7_ordinal1(X4))=>(v1_funct_1(k4_mesfunc5(X1,X2,X3,X4))&m1_subset_1(k4_mesfunc5(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', dt_k4_mesfunc5)).
fof(redefinition_k3_rfunct_3, axiom, ![X1, X2]:k3_rfunct_3(X1,X2)=k4_partfun1(X1,X2), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', redefinition_k3_rfunct_3)).
fof(redefinition_k1_mesfunc9, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v2_valued_0(X1))=>k1_mesfunc9(X1,X2)=k1_funct_1(X1,X2)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', redefinition_k1_mesfunc9)).
fof(dt_k2_mesfunc9, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v2_valued_0(X1))=>((v1_funct_1(k2_mesfunc9(X1))&v1_funct_2(k2_mesfunc9(X1),k5_numbers,k7_numbers))&m1_subset_1(k2_mesfunc9(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', dt_k2_mesfunc9)).
fof(fc2_membered, axiom, v2_membered(k7_numbers), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', fc2_membered)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', cc2_relset_1)).
fof(cc17_membered, axiom, ![X1]:(v6_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v7_ordinal1(X2))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', cc17_membered)).
fof(t1_numerals, axiom, m1_subset_1(k1_xboole_0,k4_ordinal1), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', t1_numerals)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', redefinition_k6_numbers)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', redefinition_k5_numbers)).
fof(fc6_membered, axiom, v6_membered(k4_ordinal1), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', fc6_membered)).
fof(redefinition_k8_nat_1, axiom, ![X1, X2, X3]:((((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,X1))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))&v7_ordinal1(X3))=>k8_nat_1(X1,X2,X3)=k1_funct_1(X2,X3)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', redefinition_k8_nat_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', cc1_relset_1)).
fof(d10_mesfunc5, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k7_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers))))=>(v9_mesfunc5(X1)<=>![X2]:(v1_xreal_0(X2)=>~((~(r1_xxreal_0(k6_numbers,X2))&![X3]:(v7_ordinal1(X3)=>?[X4]:((v7_ordinal1(X4)&r1_xxreal_0(X3,X4))&~(r1_xxreal_0(k8_nat_1(k7_numbers,X1,X4),X2))))))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', d10_mesfunc5)).
fof(d8_mesfunc5, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k7_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers))))=>(v7_mesfunc5(X1)<=>?[X2]:(v1_xreal_0(X2)&![X3]:(v1_xreal_0(X3)=>~((~(r1_xxreal_0(X3,k6_numbers))&![X4]:(v7_ordinal1(X4)=>?[X5]:((v7_ordinal1(X5)&r1_xxreal_0(X4,X5))&r1_xxreal_0(X3,k3_extreal1(k4_supinf_2(k8_nat_1(k7_numbers,X1,X5),k1_measure6(X2)))))))))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', d8_mesfunc5)).
fof(d9_mesfunc5, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k7_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers))))=>(v8_mesfunc5(X1)<=>![X2]:(v1_xreal_0(X2)=>~((~(r1_xxreal_0(X2,k6_numbers))&![X3]:(v7_ordinal1(X3)=>?[X4]:((v7_ordinal1(X4)&r1_xxreal_0(X3,X4))&~(r1_xxreal_0(X2,k8_nat_1(k7_numbers,X1,X4)))))))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', d9_mesfunc5)).
fof(d11_mesfunc5, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k7_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers))))=>(v10_mesfunc5(X1)<=>~(((~(v7_mesfunc5(X1))&~(v8_mesfunc5(X1)))&~(v9_mesfunc5(X1)))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t33_mesfunc9', d11_mesfunc5)).
fof(c_0_30, plain, ![X2, X1, X3]:(epred1_3(X3,X1,X2)<=>(((((v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))=>v7_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3)))&(v7_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))=>v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))))&(v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))=>v8_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))))&(v8_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))=>v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))))&(v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))=>v9_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))))), introduced(definition)).
fof(c_0_31, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,k4_partfun1(X1,k7_numbers)))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X1,k7_numbers)))))=>![X3]:(m1_subset_1(X3,X1)=>![X4]:((((v2_mesfunc9(X2,X1)&v1_mesfunc8(X2,X1,k7_numbers))&r1_tarski(X4,k1_relset_1(X1,k4_mesfunc5(X1,k7_numbers,X2,k6_numbers))))&r2_hidden(X3,X4))=>(((epred1_3(X3,X1,X2)&(v9_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))=>v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))))&(v10_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))=>v10_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))))&(v10_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))=>v10_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))))))))), inference(apply_def,[status(thm)],[inference(assume_negation,[status(cth)],[t33_mesfunc9]), c_0_30])).
fof(c_0_32, plain, ![X24, X25, X26]:(((v1_funct_1(k3_mesfunc5(X24,X25,X26))|(v1_xboole_0(X24)|~v1_funct_1(X25)|~v1_funct_2(X25,k5_numbers,k4_partfun1(X24,k7_numbers))|~m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X24,k7_numbers))))|~m1_subset_1(X26,X24)))&(v1_funct_2(k3_mesfunc5(X24,X25,X26),k5_numbers,k7_numbers)|(v1_xboole_0(X24)|~v1_funct_1(X25)|~v1_funct_2(X25,k5_numbers,k4_partfun1(X24,k7_numbers))|~m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X24,k7_numbers))))|~m1_subset_1(X26,X24))))&(m1_subset_1(k3_mesfunc5(X24,X25,X26),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|(v1_xboole_0(X24)|~v1_funct_1(X25)|~v1_funct_2(X25,k5_numbers,k4_partfun1(X24,k7_numbers))|~m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X24,k7_numbers))))|~m1_subset_1(X26,X24)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_mesfunc5])])])])).
fof(c_0_33, plain, ![X28, X29]:(((v1_funct_1(k4_mesfunc9(X28,X29))|(v1_xboole_0(X28)|~v1_funct_1(X29)|~v1_funct_2(X29,k5_numbers,k4_partfun1(X28,k7_numbers))|~m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X28,k7_numbers))))))&(v1_funct_2(k4_mesfunc9(X28,X29),k5_numbers,k4_partfun1(X28,k7_numbers))|(v1_xboole_0(X28)|~v1_funct_1(X29)|~v1_funct_2(X29,k5_numbers,k4_partfun1(X28,k7_numbers))|~m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X28,k7_numbers)))))))&(m1_subset_1(k4_mesfunc9(X28,X29),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X28,k7_numbers))))|(v1_xboole_0(X28)|~v1_funct_1(X29)|~v1_funct_2(X29,k5_numbers,k4_partfun1(X28,k7_numbers))|~m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X28,k7_numbers))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k4_mesfunc9])])])])).
fof(c_0_34, negated_conjecture, (~v1_xboole_0(esk1_0)&(((v1_funct_1(esk2_0)&v1_funct_2(esk2_0,k5_numbers,k4_partfun1(esk1_0,k7_numbers)))&m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(esk1_0,k7_numbers)))))&(m1_subset_1(esk3_0,esk1_0)&((((v2_mesfunc9(esk2_0,esk1_0)&v1_mesfunc8(esk2_0,esk1_0,k7_numbers))&r1_tarski(esk4_0,k1_relset_1(esk1_0,k4_mesfunc5(esk1_0,k7_numbers,esk2_0,k6_numbers))))&r2_hidden(esk3_0,esk4_0))&((((v10_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|(v10_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|(v9_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~epred1_3(esk3_0,esk1_0,esk2_0))))&(~v10_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|(v10_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|(v9_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~epred1_3(esk3_0,esk1_0,esk2_0)))))&((v10_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|(~v10_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|(v9_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~epred1_3(esk3_0,esk1_0,esk2_0))))&(~v10_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|(~v10_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|(v9_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~epred1_3(esk3_0,esk1_0,esk2_0))))))&(((v10_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|(v10_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|(~v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~epred1_3(esk3_0,esk1_0,esk2_0))))&(~v10_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|(v10_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|(~v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~epred1_3(esk3_0,esk1_0,esk2_0)))))&((v10_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|(~v10_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|(~v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~epred1_3(esk3_0,esk1_0,esk2_0))))&(~v10_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|(~v10_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|(~v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~epred1_3(esk3_0,esk1_0,esk2_0))))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_31])])])])])).
fof(c_0_35, plain, ![X91, X92, X93]:(~r2_hidden(X91,X92)|~m1_subset_1(X92,k1_zfmisc_1(X93))|m1_subset_1(X91,X93)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_36, plain, ![X56, X57]:(~v1_relat_1(X57)|~v4_relat_1(X57,X56)|m1_subset_1(k1_relset_1(X56,X57),k1_zfmisc_1(X56))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_relset_1])])).
fof(c_0_37, plain, ![X68, X69]:((~m1_subset_1(X68,k1_zfmisc_1(X69))|r1_tarski(X68,X69))&(~r1_tarski(X68,X69)|m1_subset_1(X68,k1_zfmisc_1(X69)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_38, plain, ![X70, X71]:(~r2_hidden(X70,X71)|~v1_xboole_0(X71)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
fof(c_0_39, plain, ![X10, X11, X12, X13, X14]:(v1_xboole_0(X10)|(~v1_funct_1(X11)|~v1_funct_2(X11,k5_numbers,k4_partfun1(X10,k7_numbers))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X10,k7_numbers))))|(~v7_ordinal1(X12)|(~m1_subset_1(X13,X10)|(~v2_mesfunc9(X11,X10)|~v1_mesfunc8(X11,X10,k7_numbers)|~r1_tarski(X14,k1_relset_1(X10,k4_mesfunc5(X10,k7_numbers,X11,k6_numbers)))|~r2_hidden(X13,X14)|k1_mesfunc9(k2_mesfunc9(k3_mesfunc5(X10,X11,X13)),X12)=k1_mesfunc9(k3_mesfunc5(X10,k4_mesfunc9(X10,X11),X13),X12)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t32_mesfunc9])])])])).
cnf(c_0_40, plain, (v1_funct_1(k3_mesfunc5(X1,X2,X3))|v1_xboole_0(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k4_partfun1(X1,k7_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X1,k7_numbers))))|~m1_subset_1(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_41, plain, (m1_subset_1(k4_mesfunc9(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X1,k7_numbers))))|v1_xboole_0(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k4_partfun1(X1,k7_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X1,k7_numbers))))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42, plain, (v1_funct_1(k4_mesfunc9(X1,X2))|v1_xboole_0(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k4_partfun1(X1,k7_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X1,k7_numbers))))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_43, plain, (v1_funct_2(k4_mesfunc9(X1,X2),k5_numbers,k4_partfun1(X1,k7_numbers))|v1_xboole_0(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k4_partfun1(X1,k7_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X1,k7_numbers))))), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_44, plain, ![X107, X108, X109]:(~v2_membered(X108)|(~m1_subset_1(X109,k1_zfmisc_1(k2_zfmisc_1(X107,X108)))|v2_valued_0(X109))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc17_valued_0])])])).
cnf(c_0_45, plain, (m1_subset_1(k3_mesfunc5(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|v1_xboole_0(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k4_partfun1(X1,k7_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X1,k7_numbers))))|~m1_subset_1(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_46, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(esk1_0,k7_numbers))))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_47, negated_conjecture, (v1_funct_2(esk2_0,k5_numbers,k4_partfun1(esk1_0,k7_numbers))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_48, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_49, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_50, plain, ![X110, X111]:(~v1_relat_1(X110)|(~m1_subset_1(X111,k1_zfmisc_1(X110))|v1_relat_1(X111))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_51, plain, ![X99, X100]:v1_relat_1(k2_zfmisc_1(X99,X100)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
cnf(c_0_52, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_53, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_54, plain, ![X72, X73]:(~m1_subset_1(X72,X73)|(v1_xboole_0(X73)|r2_hidden(X72,X73))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_55, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_56, negated_conjecture, (r1_tarski(esk4_0,k1_relset_1(esk1_0,k4_mesfunc5(esk1_0,k7_numbers,esk2_0,k6_numbers)))), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_57, plain, ![X151, X152]:(~v1_xboole_0(X151)|(~m1_subset_1(X152,k1_zfmisc_1(X151))|v1_xboole_0(X152))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])).
cnf(c_0_58, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_59, negated_conjecture, (r2_hidden(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_60, plain, ![X74, X75, X76, X77]:((v1_funct_1(k4_mesfunc5(X74,X75,X76,X77))|(~v1_funct_1(X76)|~v1_funct_2(X76,k5_numbers,k3_rfunct_3(X74,X75))|~m1_subset_1(X76,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k3_rfunct_3(X74,X75))))|~v7_ordinal1(X77)))&(m1_subset_1(k4_mesfunc5(X74,X75,X76,X77),k1_zfmisc_1(k2_zfmisc_1(X74,X75)))|(~v1_funct_1(X76)|~v1_funct_2(X76,k5_numbers,k3_rfunct_3(X74,X75))|~m1_subset_1(X76,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k3_rfunct_3(X74,X75))))|~v7_ordinal1(X77)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_mesfunc5])])])).
fof(c_0_61, plain, ![X41, X42]:k3_rfunct_3(X41,X42)=k4_partfun1(X41,X42), inference(variable_rename,[status(thm)],[redefinition_k3_rfunct_3])).
fof(c_0_62, plain, ![X58, X59]:(~v1_relat_1(X58)|~v1_funct_1(X58)|~v2_valued_0(X58)|k1_mesfunc9(X58,X59)=k1_funct_1(X58,X59)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_mesfunc9])])).
cnf(c_0_63, plain, (v1_xboole_0(X1)|k1_mesfunc9(k2_mesfunc9(k3_mesfunc5(X1,X2,X4)),X3)=k1_mesfunc9(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X4),X3)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k4_partfun1(X1,k7_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X1,k7_numbers))))|~v7_ordinal1(X3)|~m1_subset_1(X4,X1)|~v2_mesfunc9(X2,X1)|~v1_mesfunc8(X2,X1,k7_numbers)|~r1_tarski(X5,k1_relset_1(X1,k4_mesfunc5(X1,k7_numbers,X2,k6_numbers)))|~r2_hidden(X4,X5)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_64, negated_conjecture, (v1_mesfunc8(esk2_0,esk1_0,k7_numbers)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_65, negated_conjecture, (v2_mesfunc9(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_66, plain, (v1_funct_1(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X1,k7_numbers))))|~m1_subset_1(X3,X1)|~v1_funct_2(X2,k5_numbers,k4_partfun1(X1,k7_numbers))|~v1_funct_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42]), c_0_43])).
fof(c_0_67, plain, ![X27]:(((v1_funct_1(k2_mesfunc9(X27))|(~v1_relat_1(X27)|~v1_funct_1(X27)|~v2_valued_0(X27)))&(v1_funct_2(k2_mesfunc9(X27),k5_numbers,k7_numbers)|(~v1_relat_1(X27)|~v1_funct_1(X27)|~v2_valued_0(X27))))&(m1_subset_1(k2_mesfunc9(X27),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|(~v1_relat_1(X27)|~v1_funct_1(X27)|~v2_valued_0(X27)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_mesfunc9])])])).
cnf(c_0_68, plain, (v2_valued_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_69, negated_conjecture, (m1_subset_1(k3_mesfunc5(esk1_0,esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47]), c_0_48])]), c_0_49])).
cnf(c_0_70, plain, (v2_membered(k7_numbers)), inference(split_conjunct,[status(thm)],[fc2_membered])).
cnf(c_0_71, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_72, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_73, plain, (m1_subset_1(X1,X2)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~r2_hidden(X1,k1_relset_1(X2,X3))), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_74, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_75, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k1_relset_1(esk1_0,k4_mesfunc5(esk1_0,k7_numbers,esk2_0,k6_numbers))))), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_76, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_77, negated_conjecture, (~v1_xboole_0(esk4_0)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
fof(c_0_78, plain, ![X60, X61, X62]:((v4_relat_1(X62,X60)|~m1_subset_1(X62,k1_zfmisc_1(k2_zfmisc_1(X60,X61))))&(v5_relat_1(X62,X61)|~m1_subset_1(X62,k1_zfmisc_1(k2_zfmisc_1(X60,X61))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
cnf(c_0_79, plain, (m1_subset_1(k4_mesfunc5(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~v1_funct_2(X3,k5_numbers,k3_rfunct_3(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k3_rfunct_3(X1,X2))))|~v7_ordinal1(X4)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_80, plain, (k3_rfunct_3(X1,X2)=k4_partfun1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_61])).
fof(c_0_81, plain, ![X88, X89]:(~v6_membered(X88)|(~m1_subset_1(X89,X88)|v7_ordinal1(X89))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc17_membered])])])).
cnf(c_0_82, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1)), inference(split_conjunct,[status(thm)],[t1_numerals])).
cnf(c_0_83, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
cnf(c_0_84, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
cnf(c_0_85, plain, (v6_membered(k4_ordinal1)), inference(split_conjunct,[status(thm)],[fc6_membered])).
cnf(c_0_86, plain, (k1_mesfunc9(X1,X2)=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v2_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_87, negated_conjecture, (k1_mesfunc9(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),X2)=k1_mesfunc9(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),X2)|~v7_ordinal1(X2)|~r2_hidden(X1,esk4_0)|~m1_subset_1(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_56]), c_0_64]), c_0_65]), c_0_46]), c_0_47]), c_0_48])]), c_0_49])).
cnf(c_0_88, negated_conjecture, (v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~m1_subset_1(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_46]), c_0_47]), c_0_48])]), c_0_49])).
cnf(c_0_89, plain, (v1_funct_1(k2_mesfunc9(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v2_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_90, negated_conjecture, (v2_valued_0(k3_mesfunc5(esk1_0,esk2_0,X1))|~m1_subset_1(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70])])).
cnf(c_0_91, negated_conjecture, (v1_funct_1(k3_mesfunc5(esk1_0,esk2_0,X1))|~m1_subset_1(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_46]), c_0_47]), c_0_48])]), c_0_49])).
cnf(c_0_92, negated_conjecture, (v1_relat_1(k3_mesfunc5(esk1_0,esk2_0,X1))|~m1_subset_1(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_69]), c_0_72])])).
cnf(c_0_93, plain, (m1_subset_1(k2_mesfunc9(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v2_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_94, plain, (m1_subset_1(X1,X2)|v1_xboole_0(k1_relset_1(X2,X3))|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~m1_subset_1(X1,k1_relset_1(X2,X3))), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_95, negated_conjecture, (m1_subset_1(X1,k1_relset_1(esk1_0,k4_mesfunc5(esk1_0,k7_numbers,esk2_0,k6_numbers)))|~r2_hidden(X1,esk4_0)), inference(spm,[status(thm)],[c_0_52, c_0_75])).
cnf(c_0_96, negated_conjecture, (~v1_xboole_0(k1_relset_1(esk1_0,k4_mesfunc5(esk1_0,k7_numbers,esk2_0,k6_numbers)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_75]), c_0_77])).
cnf(c_0_97, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_98, plain, (m1_subset_1(k4_mesfunc5(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~v7_ordinal1(X4)|~v1_funct_2(X3,k5_numbers,k4_partfun1(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X1,X2))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_80]), c_0_80])).
cnf(c_0_99, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_100, plain, (m1_subset_1(k6_numbers,k5_numbers)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_82, c_0_83]), c_0_84])).
cnf(c_0_101, plain, (v6_membered(k5_numbers)), inference(rw,[status(thm)],[c_0_85, c_0_84])).
fof(c_0_102, plain, ![X65, X66, X67]:(~v1_funct_1(X66)|~v1_funct_2(X66,k5_numbers,X65)|~m1_subset_1(X66,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X65)))|~v7_ordinal1(X67)|k8_nat_1(X65,X66,X67)=k1_funct_1(X66,X67)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_nat_1])])).
cnf(c_0_103, negated_conjecture, (k1_mesfunc9(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),X2)=k1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),X2)|~v2_valued_0(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_relat_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v7_ordinal1(X2)|~r2_hidden(X1,esk4_0)|~m1_subset_1(X1,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_88])).
cnf(c_0_104, negated_conjecture, (v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~m1_subset_1(X1,esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_91]), c_0_92])).
cnf(c_0_105, negated_conjecture, (m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(X1,esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_90]), c_0_91]), c_0_92])).
fof(c_0_106, plain, ![X44, X45, X46]:(~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(X44,X45)))|v1_relat_1(X46)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_107, negated_conjecture, (m1_subset_1(X1,esk1_0)|~v4_relat_1(k4_mesfunc5(esk1_0,k7_numbers,esk2_0,k6_numbers),esk1_0)|~v1_relat_1(k4_mesfunc5(esk1_0,k7_numbers,esk2_0,k6_numbers))|~r2_hidden(X1,esk4_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_96])).
cnf(c_0_108, plain, (v4_relat_1(k4_mesfunc5(X1,X2,X3,X4),X1)|~v7_ordinal1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X1,X2))))|~v1_funct_2(X3,k5_numbers,k4_partfun1(X1,X2))|~v1_funct_1(X3)), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_109, plain, (v7_ordinal1(k6_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_101])])).
cnf(c_0_110, plain, (k8_nat_1(X2,X1,X3)=k1_funct_1(X1,X3)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v7_ordinal1(X3)), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_111, negated_conjecture, (k1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),X2)=k1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),X2)|~v2_valued_0(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v2_valued_0(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v1_relat_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_relat_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v7_ordinal1(X2)|~r2_hidden(X1,esk4_0)|~m1_subset_1(X1,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_103]), c_0_104])).
cnf(c_0_112, negated_conjecture, (v1_relat_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~m1_subset_1(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_105]), c_0_72])])).
cnf(c_0_113, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_106])).
cnf(c_0_114, negated_conjecture, (v2_valued_0(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~m1_subset_1(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_105]), c_0_70])])).
cnf(c_0_115, negated_conjecture, (m1_subset_1(X1,esk1_0)|~v1_relat_1(k4_mesfunc5(esk1_0,k7_numbers,esk2_0,k6_numbers))|~r2_hidden(X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_109]), c_0_46]), c_0_47]), c_0_48])])).
cnf(c_0_116, plain, (v1_relat_1(k4_mesfunc5(X1,X2,X3,X4))|~v7_ordinal1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X1,X2))))|~v1_funct_2(X3,k5_numbers,k4_partfun1(X1,X2))|~v1_funct_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_98]), c_0_72])])).
cnf(c_0_117, negated_conjecture, (k8_nat_1(X1,k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X2),X3)=k1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),X3)|~v2_valued_0(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X2))|~v7_ordinal1(X3)|~r2_hidden(X2,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X2,esk1_0)|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X2),k5_numbers,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_88]), c_0_112]), c_0_113]), c_0_114])).
cnf(c_0_118, negated_conjecture, (m1_subset_1(X1,esk1_0)|~r2_hidden(X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_109]), c_0_46]), c_0_47]), c_0_48])])).
cnf(c_0_119, plain, (m1_subset_1(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X1,k7_numbers))))|~m1_subset_1(X3,X1)|~v1_funct_2(X2,k5_numbers,k4_partfun1(X1,k7_numbers))|~v1_funct_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_41]), c_0_42]), c_0_43])).
fof(c_0_120, plain, ![X34, X35, X37, X39]:(((v7_ordinal1(esk9_2(X34,X35))|r1_xxreal_0(k6_numbers,X35)|~v1_xreal_0(X35)|~v9_mesfunc5(X34)|(~v1_funct_1(X34)|~v1_funct_2(X34,k5_numbers,k7_numbers)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))))&(~v7_ordinal1(X37)|~r1_xxreal_0(esk9_2(X34,X35),X37)|r1_xxreal_0(k8_nat_1(k7_numbers,X34,X37),X35)|r1_xxreal_0(k6_numbers,X35)|~v1_xreal_0(X35)|~v9_mesfunc5(X34)|(~v1_funct_1(X34)|~v1_funct_2(X34,k5_numbers,k7_numbers)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers))))))&((v1_xreal_0(esk10_1(X34))|v9_mesfunc5(X34)|(~v1_funct_1(X34)|~v1_funct_2(X34,k5_numbers,k7_numbers)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))))&((~r1_xxreal_0(k6_numbers,esk10_1(X34))|v9_mesfunc5(X34)|(~v1_funct_1(X34)|~v1_funct_2(X34,k5_numbers,k7_numbers)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))))&(((v7_ordinal1(esk11_2(X34,X39))|~v7_ordinal1(X39)|v9_mesfunc5(X34)|(~v1_funct_1(X34)|~v1_funct_2(X34,k5_numbers,k7_numbers)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))))&(r1_xxreal_0(X39,esk11_2(X34,X39))|~v7_ordinal1(X39)|v9_mesfunc5(X34)|(~v1_funct_1(X34)|~v1_funct_2(X34,k5_numbers,k7_numbers)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers))))))&(~r1_xxreal_0(k8_nat_1(k7_numbers,X34,esk11_2(X34,X39)),esk10_1(X34))|~v7_ordinal1(X39)|v9_mesfunc5(X34)|(~v1_funct_1(X34)|~v1_funct_2(X34,k5_numbers,k7_numbers)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d10_mesfunc5])])])])])])).
cnf(c_0_121, negated_conjecture, (k8_nat_1(X1,k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),X3)=k8_nat_1(X4,k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X2),X3)|~v2_valued_0(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X2))|~v7_ordinal1(X3)|~r2_hidden(X2,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X2),k5_numbers,X4)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),k5_numbers,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_117]), c_0_104]), c_0_118])).
cnf(c_0_122, plain, (v2_valued_0(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X1,k7_numbers))))|~m1_subset_1(X3,X1)|~v1_funct_2(X2,k5_numbers,k4_partfun1(X1,k7_numbers))|~v1_funct_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_119]), c_0_70])])).
cnf(c_0_123, plain, (v1_funct_2(k3_mesfunc5(X1,X2,X3),k5_numbers,k7_numbers)|v1_xboole_0(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k4_partfun1(X1,k7_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X1,k7_numbers))))|~m1_subset_1(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_124, plain, (v9_mesfunc5(X1)|~r1_xxreal_0(k8_nat_1(k7_numbers,X1,esk11_2(X1,X2)),esk10_1(X1))|~v7_ordinal1(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_125, plain, (k8_nat_1(X1,X2,X3)=k8_nat_1(X4,X2,X3)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_2(X2,k5_numbers,X4)|~v1_funct_1(X2)), inference(spm,[status(thm)],[c_0_110, c_0_110])).
cnf(c_0_126, plain, (v7_ordinal1(esk11_2(X1,X2))|v9_mesfunc5(X1)|~v7_ordinal1(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_127, negated_conjecture, (k8_nat_1(X1,k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),X3)=k8_nat_1(X4,k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X2),X3)|~v7_ordinal1(X3)|~r2_hidden(X2,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X2),k5_numbers,X4)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),k5_numbers,X1)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122]), c_0_46]), c_0_47]), c_0_48])]), c_0_49]), c_0_118])).
cnf(c_0_128, plain, (v1_funct_2(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3),k5_numbers,k7_numbers)|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k4_partfun1(X1,k7_numbers))))|~m1_subset_1(X3,X1)|~v1_funct_2(X2,k5_numbers,k4_partfun1(X1,k7_numbers))|~v1_funct_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_41]), c_0_42]), c_0_43])).
cnf(c_0_129, plain, (v9_mesfunc5(X1)|~r1_xxreal_0(k8_nat_1(X2,X1,esk11_2(X1,X3)),esk10_1(X1))|~v7_ordinal1(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(X1,k5_numbers,k7_numbers)|~v1_funct_2(X1,k5_numbers,X2)|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_125]), c_0_126])).
cnf(c_0_130, plain, (r1_xxreal_0(k8_nat_1(k7_numbers,X2,X1),X3)|r1_xxreal_0(k6_numbers,X3)|~v7_ordinal1(X1)|~r1_xxreal_0(esk9_2(X2,X3),X1)|~v1_xreal_0(X3)|~v9_mesfunc5(X2)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_131, plain, (r1_xxreal_0(X1,esk11_2(X2,X1))|v9_mesfunc5(X2)|~v7_ordinal1(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_132, plain, (v7_ordinal1(esk9_2(X1,X2))|r1_xxreal_0(k6_numbers,X2)|~v1_xreal_0(X2)|~v9_mesfunc5(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_120])).
fof(c_0_133, plain, ![X15, X17, X19, X20, X22]:(((v1_xreal_0(esk5_1(X15))|~v7_mesfunc5(X15)|(~v1_funct_1(X15)|~v1_funct_2(X15,k5_numbers,k7_numbers)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))))&((v7_ordinal1(esk6_2(X15,X17))|r1_xxreal_0(X17,k6_numbers)|~v1_xreal_0(X17)|~v7_mesfunc5(X15)|(~v1_funct_1(X15)|~v1_funct_2(X15,k5_numbers,k7_numbers)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))))&(~v7_ordinal1(X19)|~r1_xxreal_0(esk6_2(X15,X17),X19)|~r1_xxreal_0(X17,k3_extreal1(k4_supinf_2(k8_nat_1(k7_numbers,X15,X19),k1_measure6(esk5_1(X15)))))|r1_xxreal_0(X17,k6_numbers)|~v1_xreal_0(X17)|~v7_mesfunc5(X15)|(~v1_funct_1(X15)|~v1_funct_2(X15,k5_numbers,k7_numbers)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))))))&((v1_xreal_0(esk7_2(X15,X20))|~v1_xreal_0(X20)|v7_mesfunc5(X15)|(~v1_funct_1(X15)|~v1_funct_2(X15,k5_numbers,k7_numbers)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))))&((~r1_xxreal_0(esk7_2(X15,X20),k6_numbers)|~v1_xreal_0(X20)|v7_mesfunc5(X15)|(~v1_funct_1(X15)|~v1_funct_2(X15,k5_numbers,k7_numbers)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))))&(((v7_ordinal1(esk8_3(X15,X20,X22))|~v7_ordinal1(X22)|~v1_xreal_0(X20)|v7_mesfunc5(X15)|(~v1_funct_1(X15)|~v1_funct_2(X15,k5_numbers,k7_numbers)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))))&(r1_xxreal_0(X22,esk8_3(X15,X20,X22))|~v7_ordinal1(X22)|~v1_xreal_0(X20)|v7_mesfunc5(X15)|(~v1_funct_1(X15)|~v1_funct_2(X15,k5_numbers,k7_numbers)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers))))))&(r1_xxreal_0(esk7_2(X15,X20),k3_extreal1(k4_supinf_2(k8_nat_1(k7_numbers,X15,esk8_3(X15,X20,X22)),k1_measure6(X20))))|~v7_ordinal1(X22)|~v1_xreal_0(X20)|v7_mesfunc5(X15)|(~v1_funct_1(X15)|~v1_funct_2(X15,k5_numbers,k7_numbers)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_mesfunc5])])])])])])).
fof(c_0_134, plain, ![X47, X48, X50, X52]:(((v7_ordinal1(esk12_2(X47,X48))|r1_xxreal_0(X48,k6_numbers)|~v1_xreal_0(X48)|~v8_mesfunc5(X47)|(~v1_funct_1(X47)|~v1_funct_2(X47,k5_numbers,k7_numbers)|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))))&(~v7_ordinal1(X50)|~r1_xxreal_0(esk12_2(X47,X48),X50)|r1_xxreal_0(X48,k8_nat_1(k7_numbers,X47,X50))|r1_xxreal_0(X48,k6_numbers)|~v1_xreal_0(X48)|~v8_mesfunc5(X47)|(~v1_funct_1(X47)|~v1_funct_2(X47,k5_numbers,k7_numbers)|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers))))))&((v1_xreal_0(esk13_1(X47))|v8_mesfunc5(X47)|(~v1_funct_1(X47)|~v1_funct_2(X47,k5_numbers,k7_numbers)|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))))&((~r1_xxreal_0(esk13_1(X47),k6_numbers)|v8_mesfunc5(X47)|(~v1_funct_1(X47)|~v1_funct_2(X47,k5_numbers,k7_numbers)|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))))&(((v7_ordinal1(esk14_2(X47,X52))|~v7_ordinal1(X52)|v8_mesfunc5(X47)|(~v1_funct_1(X47)|~v1_funct_2(X47,k5_numbers,k7_numbers)|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))))&(r1_xxreal_0(X52,esk14_2(X47,X52))|~v7_ordinal1(X52)|v8_mesfunc5(X47)|(~v1_funct_1(X47)|~v1_funct_2(X47,k5_numbers,k7_numbers)|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers))))))&(~r1_xxreal_0(esk13_1(X47),k8_nat_1(k7_numbers,X47,esk14_2(X47,X52)))|~v7_ordinal1(X52)|v8_mesfunc5(X47)|(~v1_funct_1(X47)|~v1_funct_2(X47,k5_numbers,k7_numbers)|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d9_mesfunc5])])])])])])).
cnf(c_0_135, negated_conjecture, (k8_nat_1(X1,k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),X3)=k8_nat_1(X4,k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),X3)|~v7_ordinal1(X3)|~r2_hidden(X2,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X2),k5_numbers,X5)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),k5_numbers,X1)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),k5_numbers,X4)), inference(spm,[status(thm)],[c_0_127, c_0_127])).
cnf(c_0_136, negated_conjecture, (v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~m1_subset_1(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_46]), c_0_47]), c_0_48])]), c_0_49])).
cnf(c_0_137, negated_conjecture, (v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r1_xxreal_0(k8_nat_1(X2,k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk11_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),X3)),esk10_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1))))|~v7_ordinal1(X3)|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,X2)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,X4)|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_127]), c_0_126])).
cnf(c_0_138, plain, (r1_xxreal_0(k8_nat_1(k7_numbers,X1,esk11_2(X2,esk9_2(X1,X3))),X3)|r1_xxreal_0(k6_numbers,X3)|v9_mesfunc5(X2)|~v1_xreal_0(X3)|~v7_ordinal1(esk11_2(X2,esk9_2(X1,X3)))|~v9_mesfunc5(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(X1,k5_numbers,k7_numbers)|~v1_funct_2(X2,k5_numbers,k7_numbers)|~v1_funct_1(X1)|~v1_funct_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_131]), c_0_132])).
cnf(c_0_139, plain, (v1_xreal_0(esk10_1(X1))|v9_mesfunc5(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_140, plain, (v9_mesfunc5(X1)|~r1_xxreal_0(k6_numbers,esk10_1(X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_120])).
fof(c_0_141, plain, ![X43]:((~v10_mesfunc5(X43)|(v7_mesfunc5(X43)|v8_mesfunc5(X43)|v9_mesfunc5(X43))|(~v1_funct_1(X43)|~v1_funct_2(X43,k5_numbers,k7_numbers)|~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))))&(((~v7_mesfunc5(X43)|v10_mesfunc5(X43)|(~v1_funct_1(X43)|~v1_funct_2(X43,k5_numbers,k7_numbers)|~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))))&(~v8_mesfunc5(X43)|v10_mesfunc5(X43)|(~v1_funct_1(X43)|~v1_funct_2(X43,k5_numbers,k7_numbers)|~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers))))))&(~v9_mesfunc5(X43)|v10_mesfunc5(X43)|(~v1_funct_1(X43)|~v1_funct_2(X43,k5_numbers,k7_numbers)|~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d11_mesfunc5])])])])).
cnf(c_0_142, negated_conjecture, (v9_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~r1_xxreal_0(k8_nat_1(X2,k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),esk11_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),X3)),esk10_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1)))|~v7_ordinal1(X3)|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,X2)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_127]), c_0_126])).
cnf(c_0_143, plain, (r1_xxreal_0(esk7_2(X1,X2),k3_extreal1(k4_supinf_2(k8_nat_1(k7_numbers,X1,esk8_3(X1,X2,X3)),k1_measure6(X2))))|v7_mesfunc5(X1)|~v7_ordinal1(X3)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_133])).
cnf(c_0_144, plain, (v7_ordinal1(esk8_3(X1,X2,X3))|v7_mesfunc5(X1)|~v7_ordinal1(X3)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_133])).
cnf(c_0_145, negated_conjecture, (k8_nat_1(X1,k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),X3)=k1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X2),X3)|~v7_ordinal1(X3)|~r2_hidden(X2,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X2),k5_numbers,X4)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),k5_numbers,X1)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X2))), inference(spm,[status(thm)],[c_0_110, c_0_127])).
cnf(c_0_146, plain, (v8_mesfunc5(X1)|~r1_xxreal_0(esk13_1(X1),k8_nat_1(k7_numbers,X1,esk14_2(X1,X2)))|~v7_ordinal1(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_147, negated_conjecture, (k8_nat_1(X1,k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),X3)=k8_nat_1(X4,k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),X3)|~v7_ordinal1(X3)|~r2_hidden(X2,esk4_0)|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),k5_numbers,X1)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),k5_numbers,X4)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_119]), c_0_46]), c_0_47]), c_0_48])]), c_0_49]), c_0_136]), c_0_118])).
cnf(c_0_148, plain, (v7_ordinal1(esk14_2(X1,X2))|v8_mesfunc5(X1)|~v7_ordinal1(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_149, plain, (r1_xxreal_0(X3,k6_numbers)|~v7_ordinal1(X1)|~r1_xxreal_0(esk6_2(X2,X3),X1)|~r1_xxreal_0(X3,k3_extreal1(k4_supinf_2(k8_nat_1(k7_numbers,X2,X1),k1_measure6(esk5_1(X2)))))|~v1_xreal_0(X3)|~v7_mesfunc5(X2)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_133])).
cnf(c_0_150, negated_conjecture, (v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v7_ordinal1(esk11_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),esk9_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk10_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1))))))|~v9_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,X2)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_138]), c_0_132]), c_0_139]), c_0_140])).
cnf(c_0_151, negated_conjecture, (~v10_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~v10_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~epred1_3(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_152, plain, (v10_mesfunc5(X1)|~v9_mesfunc5(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_141])).
cnf(c_0_153, negated_conjecture, (v9_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v7_ordinal1(esk11_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk9_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),esk10_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1)))))|~v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_138]), c_0_132]), c_0_139]), c_0_140])).
cnf(c_0_154, plain, (r1_xxreal_0(X3,k8_nat_1(k7_numbers,X2,X1))|r1_xxreal_0(X3,k6_numbers)|~v7_ordinal1(X1)|~r1_xxreal_0(esk12_2(X2,X3),X1)|~v1_xreal_0(X3)|~v8_mesfunc5(X2)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_155, plain, (r1_xxreal_0(X1,esk14_2(X2,X1))|v8_mesfunc5(X2)|~v7_ordinal1(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_156, plain, (v7_ordinal1(esk12_2(X1,X2))|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X2)|~v8_mesfunc5(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_157, negated_conjecture, (r1_xxreal_0(esk7_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),X2),k3_extreal1(k4_supinf_2(k8_nat_1(X3,k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),esk8_3(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),X2,X4)),k1_measure6(X2))))|v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_xreal_0(X2)|~v7_ordinal1(X4)|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,X3)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_127]), c_0_144])).
cnf(c_0_158, negated_conjecture, (k8_nat_1(X1,k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),X3)=k1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X2),X3)|~v7_ordinal1(X3)|~r2_hidden(X2,esk4_0)|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X2)),k5_numbers,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_119]), c_0_46]), c_0_47]), c_0_48])]), c_0_49]), c_0_88]), c_0_136]), c_0_118])).
cnf(c_0_159, plain, (v1_funct_2(k2_mesfunc9(X1),k5_numbers,k7_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v2_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_160, negated_conjecture, (v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r1_xxreal_0(esk13_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1))),k8_nat_1(X2,k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),esk14_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),X3)))|~v7_ordinal1(X3)|~r2_hidden(X1,esk4_0)|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,X2)|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_147]), c_0_148])).
cnf(c_0_161, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(X1,k3_extreal1(k4_supinf_2(k8_nat_1(X2,k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X3)),X4),k1_measure6(esk5_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X3))))))|~r1_xxreal_0(esk6_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X3),X1),X4)|~v1_xreal_0(X1)|~v7_ordinal1(X4)|~v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X3))|~r2_hidden(X3,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X3)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X3),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X3)),k5_numbers,X2)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X3))), inference(spm,[status(thm)],[c_0_149, c_0_127])).
cnf(c_0_162, plain, (v1_xreal_0(esk5_1(X1))|~v7_mesfunc5(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_133])).
fof(c_0_163, plain, ![X2, X1, X3]:((((((v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))=>v7_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3)))&(v7_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))=>v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))))&(v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))=>v8_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))))&(v8_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))=>v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))))&(v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))=>v9_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))))=>epred1_3(X3,X1,X2)), inference(split_equiv,[status(thm)],[c_0_30])).
cnf(c_0_164, negated_conjecture, (v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v7_ordinal1(esk9_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk10_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))))|~v9_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,X2)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(spm,[status(thm)],[c_0_150, c_0_126])).
cnf(c_0_165, negated_conjecture, (v9_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v10_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~v10_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~epred1_3(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_166, plain, (v10_mesfunc5(X1)|~v8_mesfunc5(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_141])).
cnf(c_0_167, negated_conjecture, (~epred1_3(esk3_0,esk1_0,esk2_0)|~v10_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~v9_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_151, c_0_152])).
cnf(c_0_168, negated_conjecture, (v9_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v7_ordinal1(esk9_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),esk10_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))))|~v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(spm,[status(thm)],[c_0_153, c_0_126])).
cnf(c_0_169, negated_conjecture, (v8_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~r1_xxreal_0(esk13_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1)),k8_nat_1(X2,k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),esk14_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),X3)))|~v7_ordinal1(X3)|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,X2)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_127]), c_0_148])).
cnf(c_0_170, plain, (r1_xxreal_0(X1,k8_nat_1(k7_numbers,X2,esk14_2(X3,esk12_2(X2,X1))))|r1_xxreal_0(X1,k6_numbers)|v8_mesfunc5(X3)|~v1_xreal_0(X1)|~v7_ordinal1(esk14_2(X3,esk12_2(X2,X1)))|~v8_mesfunc5(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(X2,k5_numbers,k7_numbers)|~v1_funct_2(X3,k5_numbers,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_155]), c_0_156])).
cnf(c_0_171, plain, (v1_xreal_0(esk13_1(X1))|v8_mesfunc5(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_172, plain, (v8_mesfunc5(X1)|~r1_xxreal_0(esk13_1(X1),k6_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_173, negated_conjecture, (r1_xxreal_0(esk7_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),X2),k3_extreal1(k4_supinf_2(k1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk8_3(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),X2,X3)),k1_measure6(X2))))|v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_xreal_0(X2)|~v7_ordinal1(X3)|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,X4)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_158]), c_0_144])).
cnf(c_0_174, negated_conjecture, (v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~m1_subset_1(X1,esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_90]), c_0_91]), c_0_92])).
cnf(c_0_175, negated_conjecture, (v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r1_xxreal_0(esk13_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1))),k8_nat_1(X2,k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk14_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),X3)))|~v7_ordinal1(X3)|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,X2)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,X4)|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_127]), c_0_148])).
cnf(c_0_176, negated_conjecture, (r1_xxreal_0(esk7_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),esk5_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))),k6_numbers)|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r1_xxreal_0(esk6_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk7_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),esk5_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1)))),esk8_3(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),esk5_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1)),X2))|~v1_xreal_0(esk7_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),esk5_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))))|~v7_ordinal1(X2)|~v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_143]), c_0_144]), c_0_162])).
cnf(c_0_177, plain, (r1_xxreal_0(X1,esk8_3(X2,X3,X1))|v7_mesfunc5(X2)|~v7_ordinal1(X1)|~v1_xreal_0(X3)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_133])).
cnf(c_0_178, plain, (v7_ordinal1(esk6_2(X1,X2))|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X2)|~v7_mesfunc5(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_133])).
cnf(c_0_179, plain, (v7_mesfunc5(X1)|v8_mesfunc5(X1)|v9_mesfunc5(X1)|~v10_mesfunc5(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_141])).
cnf(c_0_180, negated_conjecture, (v10_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|v10_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|v9_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~epred1_3(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_181, plain, ![X157, X158, X159]:((((((v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159))))))|epred1_3(X159,X158,X157))&(~v9_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159))))))|epred1_3(X159,X158,X157)))&((v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159))))))|epred1_3(X159,X158,X157))&(~v9_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159))))))|epred1_3(X159,X158,X157))))&(((v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159))))))|epred1_3(X159,X158,X157))&(~v9_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159))))))|epred1_3(X159,X158,X157)))&((v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159))))))|epred1_3(X159,X158,X157))&(~v9_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159))))))|epred1_3(X159,X158,X157)))))&((((v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159))))))|epred1_3(X159,X158,X157))&(~v9_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159))))))|epred1_3(X159,X158,X157)))&((v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159))))))|epred1_3(X159,X158,X157))&(~v9_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159))))))|epred1_3(X159,X158,X157))))&(((v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159))))))|epred1_3(X159,X158,X157))&(~v9_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159))))))|epred1_3(X159,X158,X157)))&((v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159))))))|epred1_3(X159,X158,X157))&(~v9_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159))))))|epred1_3(X159,X158,X157))))))&(((((v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|~v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159)))))|epred1_3(X159,X158,X157))&(~v9_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|~v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159)))))|epred1_3(X159,X158,X157)))&((v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|~v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159)))))|epred1_3(X159,X158,X157))&(~v9_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|~v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159)))))|epred1_3(X159,X158,X157))))&(((v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|~v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159)))))|epred1_3(X159,X158,X157))&(~v9_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|~v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159)))))|epred1_3(X159,X158,X157)))&((v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|~v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159)))))|epred1_3(X159,X158,X157))&(~v9_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|~v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159)))))|epred1_3(X159,X158,X157)))))&((((v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|~v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159)))))|epred1_3(X159,X158,X157))&(~v9_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|~v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159)))))|epred1_3(X159,X158,X157)))&((v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|~v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159)))))|epred1_3(X159,X158,X157))&(~v9_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|~v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159)))))|epred1_3(X159,X158,X157))))&(((v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|~v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159)))))|epred1_3(X159,X158,X157))&(~v9_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|~v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159)))))|epred1_3(X159,X158,X157)))&((v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|~v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159)))))|epred1_3(X159,X158,X157))&(~v9_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|(~v8_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159))|(~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X158,X157,X159)))|~v7_mesfunc5(k3_mesfunc5(X158,k4_mesfunc9(X158,X157),X159)))))|epred1_3(X159,X158,X157))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_163])])])).
cnf(c_0_182, negated_conjecture, (v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v9_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,X2)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_132]), c_0_139]), c_0_140])).
cnf(c_0_183, negated_conjecture, (v9_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~epred1_3(esk3_0,esk1_0,esk2_0)|~v10_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~v8_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_165, c_0_166])).
cnf(c_0_184, negated_conjecture, (~epred1_3(esk3_0,esk1_0,esk2_0)|~v9_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))), inference(spm,[status(thm)],[c_0_167, c_0_152])).
cnf(c_0_185, negated_conjecture, (v9_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_132]), c_0_139]), c_0_140])).
cnf(c_0_186, negated_conjecture, (v8_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v7_ordinal1(esk14_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk12_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),esk13_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1)))))|~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_170]), c_0_156]), c_0_171]), c_0_172])).
cnf(c_0_187, negated_conjecture, (r1_xxreal_0(esk7_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),X2),k3_extreal1(k4_supinf_2(k1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk8_3(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),X2,X3)),k1_measure6(X2))))|v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_xreal_0(X2)|~v7_ordinal1(X3)|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_105]), c_0_88]), c_0_174]), c_0_136]), c_0_118])).
cnf(c_0_188, negated_conjecture, (v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v7_ordinal1(esk14_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),esk12_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk13_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1))))))|~v8_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,X2)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_170]), c_0_156]), c_0_171]), c_0_172])).
cnf(c_0_189, plain, (v7_mesfunc5(X1)|~r1_xxreal_0(esk7_2(X1,X2),k6_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_133])).
cnf(c_0_190, negated_conjecture, (r1_xxreal_0(esk7_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),esk5_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))),k6_numbers)|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v1_xreal_0(esk7_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),esk5_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))))|~v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_177]), c_0_178]), c_0_162])).
cnf(c_0_191, negated_conjecture, (v10_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|v9_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|v8_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~epred1_3(esk3_0,esk1_0,esk2_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_179, c_0_180])).
cnf(c_0_192, plain, (v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))|v8_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))|v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))|v7_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))|epred1_3(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_181])).
cnf(c_0_193, negated_conjecture, (v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v9_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_105]), c_0_104]), c_0_88]), c_0_174]), c_0_136]), c_0_118])).
cnf(c_0_194, negated_conjecture, (v9_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~epred1_3(esk3_0,esk1_0,esk2_0)|~v8_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))), inference(spm,[status(thm)],[c_0_183, c_0_166])).
cnf(c_0_195, negated_conjecture, (~epred1_3(esk3_0,esk1_0,esk2_0)|~v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_185]), c_0_59])])).
cnf(c_0_196, negated_conjecture, (v8_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v7_ordinal1(esk12_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),esk13_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))))|~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(spm,[status(thm)],[c_0_186, c_0_148])).
cnf(c_0_197, negated_conjecture, (r1_xxreal_0(esk7_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),X2),k3_extreal1(k4_supinf_2(k8_nat_1(X3,k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),esk8_3(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),X2,X4)),k1_measure6(X2))))|v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_xreal_0(X2)|~v7_ordinal1(esk8_3(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),X2,X4))|~v7_ordinal1(X4)|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,X3)), inference(spm,[status(thm)],[c_0_187, c_0_158])).
cnf(c_0_198, negated_conjecture, (v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v7_ordinal1(esk12_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk13_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))))|~v8_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,X2)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(spm,[status(thm)],[c_0_188, c_0_148])).
cnf(c_0_199, negated_conjecture, (v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v1_xreal_0(esk7_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),esk5_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))))|~v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_190]), c_0_162])).
cnf(c_0_200, plain, (v1_xreal_0(esk7_2(X1,X2))|v7_mesfunc5(X1)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_133])).
cnf(c_0_201, negated_conjecture, (v9_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|v8_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_191]), c_0_192])).
cnf(c_0_202, negated_conjecture, (~epred1_3(esk3_0,esk1_0,esk2_0)|~v8_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_194]), c_0_59])]), c_0_195])).
cnf(c_0_203, negated_conjecture, (v8_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_156]), c_0_171]), c_0_172])).
cnf(c_0_204, negated_conjecture, (r1_xxreal_0(esk7_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk5_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))),k6_numbers)|v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~r1_xxreal_0(esk6_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),esk7_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk5_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1))))),esk8_3(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk5_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1))),X2))|~v1_xreal_0(esk7_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk5_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))))|~v7_ordinal1(esk8_3(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk5_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1))),X2))|~v7_ordinal1(X2)|~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_197]), c_0_162])).
cnf(c_0_205, plain, (v7_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))|epred1_3(X3,X1,X2)|~v9_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))|~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))|~v8_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))), inference(split_conjunct,[status(thm)],[c_0_181])).
cnf(c_0_206, negated_conjecture, (v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v8_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,X2)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_156]), c_0_171]), c_0_172])).
cnf(c_0_207, plain, (v8_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))|v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))|v7_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))|epred1_3(X3,X1,X2)|~v9_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))), inference(split_conjunct,[status(thm)],[c_0_181])).
cnf(c_0_208, negated_conjecture, (v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_200]), c_0_162])).
cnf(c_0_209, negated_conjecture, (v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|v8_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_201]), c_0_59])])).
cnf(c_0_210, negated_conjecture, (~epred1_3(esk3_0,esk1_0,esk2_0)|~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_203]), c_0_59])])).
cnf(c_0_211, negated_conjecture, (m1_subset_1(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_212, plain, (v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))|epred1_3(X3,X1,X2)|~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))|~v8_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))|~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))|~v7_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))), inference(split_conjunct,[status(thm)],[c_0_181])).
cnf(c_0_213, plain, (v8_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))|v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))|epred1_3(X3,X1,X2)|~v9_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))|~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))|~v7_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))), inference(split_conjunct,[status(thm)],[c_0_181])).
cnf(c_0_214, plain, (v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))|v8_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))|v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))|epred1_3(X3,X1,X2)|~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))|~v7_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))), inference(split_conjunct,[status(thm)],[c_0_181])).
cnf(c_0_215, negated_conjecture, (r1_xxreal_0(esk7_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk5_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))),k6_numbers)|v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_xreal_0(esk7_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk5_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))))|~v7_ordinal1(esk8_3(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk5_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1))),esk6_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),esk7_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk5_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))))))|~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_177]), c_0_178]), c_0_162])).
cnf(c_0_216, plain, (v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))|v7_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))|epred1_3(X3,X1,X2)|~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))|~v8_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))), inference(split_conjunct,[status(thm)],[c_0_181])).
cnf(c_0_217, plain, (epred1_3(X1,esk1_0,esk2_0)|v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_185]), c_0_203])).
cnf(c_0_218, negated_conjecture, (v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v8_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_105]), c_0_104]), c_0_88]), c_0_174]), c_0_136]), c_0_118])).
cnf(c_0_219, plain, (epred1_3(X1,esk1_0,esk2_0)|v8_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_185]), c_0_192]), c_0_203])).
cnf(c_0_220, negated_conjecture, (v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_119]), c_0_46]), c_0_47]), c_0_48])]), c_0_49]), c_0_104]), c_0_88]), c_0_174]), c_0_136]), c_0_118])).
cnf(c_0_221, plain, (v10_mesfunc5(X1)|~v7_mesfunc5(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_141])).
cnf(c_0_222, negated_conjecture, (v8_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~epred1_3(esk3_0,esk1_0,esk2_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_209]), c_0_210])).
cnf(c_0_223, negated_conjecture, (~epred1_3(esk3_0,esk1_0,esk2_0)|~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_119]), c_0_46]), c_0_211]), c_0_47]), c_0_48])]), c_0_49])).
cnf(c_0_224, plain, (epred1_3(X3,X1,X2)|~v9_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))|~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))|~v8_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))|~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(X1,X2,X3)))|~v7_mesfunc5(k3_mesfunc5(X1,k4_mesfunc9(X1,X2),X3))), inference(split_conjunct,[status(thm)],[c_0_181])).
cnf(c_0_225, plain, (epred1_3(X1,esk1_0,esk2_0)|v9_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(spm,[status(thm)],[c_0_212, c_0_203])).
cnf(c_0_226, plain, (epred1_3(X1,esk1_0,esk2_0)|v8_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_185]), c_0_214]), c_0_203])).
cnf(c_0_227, negated_conjecture, (r1_xxreal_0(esk7_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk5_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))),k6_numbers)|v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_xreal_0(esk7_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk5_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))))|~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_144]), c_0_178]), c_0_162])).
cnf(c_0_228, plain, (epred1_3(X1,esk1_0,esk2_0)|v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_203]), c_0_217])).
cnf(c_0_229, plain, (epred1_3(X1,esk1_0,esk2_0)|v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_219]), c_0_220])).
cnf(c_0_230, negated_conjecture, (v9_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~epred1_3(esk3_0,esk1_0,esk2_0)|~v10_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_165, c_0_221])).
cnf(c_0_231, negated_conjecture, (v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~epred1_3(esk3_0,esk1_0,esk2_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_222]), c_0_59])]), c_0_223])).
cnf(c_0_232, plain, (epred1_3(X1,esk1_0,esk2_0)|~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_224, c_0_185]), c_0_203]), c_0_225])).
cnf(c_0_233, plain, (epred1_3(X1,esk1_0,esk2_0)|v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_226]), c_0_220])).
cnf(c_0_234, negated_conjecture, (v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_xreal_0(esk7_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),esk5_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))))|~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_227]), c_0_162])).
cnf(c_0_235, plain, (epred1_3(X1,esk1_0,esk2_0)|v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_119]), c_0_46]), c_0_47]), c_0_48])]), c_0_49]), c_0_104]), c_0_88]), c_0_174]), c_0_136]), c_0_118])).
cnf(c_0_236, plain, (epred1_3(X1,esk1_0,esk2_0)|v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_119]), c_0_46]), c_0_47]), c_0_48])]), c_0_49]), c_0_104]), c_0_88]), c_0_174]), c_0_136]), c_0_118])).
cnf(c_0_237, negated_conjecture, (v9_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~epred1_3(esk3_0,esk1_0,esk2_0)|~v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))), inference(spm,[status(thm)],[c_0_230, c_0_221])).
cnf(c_0_238, negated_conjecture, (v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~epred1_3(esk3_0,esk1_0,esk2_0)|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_119]), c_0_46]), c_0_211]), c_0_47]), c_0_48])]), c_0_49])).
cnf(c_0_239, plain, (epred1_3(X1,esk1_0,esk2_0)|~v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_119]), c_0_46]), c_0_47]), c_0_48])]), c_0_49]), c_0_104]), c_0_88]), c_0_174]), c_0_136]), c_0_118])).
cnf(c_0_240, plain, (epred1_3(X1,esk1_0,esk2_0)|v8_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_119]), c_0_46]), c_0_47]), c_0_48])]), c_0_49]), c_0_104]), c_0_88]), c_0_174]), c_0_136]), c_0_118])).
cnf(c_0_241, negated_conjecture, (v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_200]), c_0_162])).
cnf(c_0_242, plain, (epred1_3(X1,esk1_0,esk2_0)|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_236]), c_0_220])).
cnf(c_0_243, negated_conjecture, (~epred1_3(esk3_0,esk1_0,esk2_0)|~v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~m1_subset_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_237]), c_0_59])]), c_0_195])).
cnf(c_0_244, negated_conjecture, (v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~epred1_3(esk3_0,esk1_0,esk2_0)|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_238]), c_0_59])])).
cnf(c_0_245, plain, (epred1_3(X1,esk1_0,esk2_0)|~v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_240]), c_0_220])).
cnf(c_0_246, negated_conjecture, (v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),X1))|~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_119]), c_0_46]), c_0_47]), c_0_48])]), c_0_49]), c_0_104]), c_0_88]), c_0_174]), c_0_136]), c_0_118])).
cnf(c_0_247, negated_conjecture, (epred1_3(X1,esk1_0,esk2_0)|v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)))|~r2_hidden(X1,esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_105]), c_0_118])).
cnf(c_0_248, negated_conjecture, (~epred1_3(esk3_0,esk1_0,esk2_0)|~v7_mesfunc5(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_119]), c_0_46]), c_0_211]), c_0_47]), c_0_48])]), c_0_49])).
cnf(c_0_249, negated_conjecture, (v7_mesfunc5(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))|~epred1_3(esk3_0,esk1_0,esk2_0)|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_244, c_0_105]), c_0_211])])).
cnf(c_0_250, plain, (epred1_3(X1,esk1_0,esk2_0)|~r2_hidden(X1,esk4_0)|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_246]), c_0_247])).
cnf(c_0_251, negated_conjecture, (~epred1_3(esk3_0,esk1_0,esk2_0)|~m1_subset_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k7_numbers)))|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_248, c_0_246]), c_0_59])]), c_0_249])).
cnf(c_0_252, negated_conjecture, (epred1_3(X1,esk1_0,esk2_0)|~r2_hidden(X1,esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_105]), c_0_118])).
cnf(c_0_253, negated_conjecture, (~epred1_3(esk3_0,esk1_0,esk2_0)|~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_105]), c_0_211])])).
cnf(c_0_254, negated_conjecture, (epred1_3(esk3_0,esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_252, c_0_59])).
cnf(c_0_255, negated_conjecture, (~v1_funct_2(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0),k5_numbers,k7_numbers)|~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_253, c_0_254])])).
cnf(c_0_256, negated_conjecture, (~v1_funct_2(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)),k5_numbers,k7_numbers)|~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_136]), c_0_211])])).
cnf(c_0_257, negated_conjecture, (~v1_funct_1(k3_mesfunc5(esk1_0,k4_mesfunc9(esk1_0,esk2_0),esk3_0))|~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_256, c_0_174]), c_0_211])])).
cnf(c_0_258, negated_conjecture, (~v1_funct_1(k2_mesfunc9(k3_mesfunc5(esk1_0,esk2_0,esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_88]), c_0_211])])).
cnf(c_0_259, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_104]), c_0_211])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 260
# Proof object clause steps            : 201
# Proof object formula steps           : 59
# Proof object conjectures             : 108
# Proof object clause conjectures      : 105
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 77
# Proof object initial formulas used   : 30
# Proof object generating inferences   : 120
# Proof object simplifying inferences  : 281
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 70
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 161
# Removed in clause preprocessing      : 32
# Initial clauses in saturation        : 129
# Processed clauses                    : 2081
# ...of these trivial                  : 2
# ...subsumed                          : 1111
# ...remaining for further processing  : 968
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 139
# Backward-rewritten                   : 12
# Generated clauses                    : 6522
# ...of the previous two non-trivial   : 6381
# Contextual simplify-reflections      : 551
# Paramodulations                      : 6521
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 816
#    Positive orientable unit clauses  : 31
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 10
#    Non-unit-clauses                  : 775
# Current number of unprocessed clauses: 4311
# ...number of literals in the above   : 52651
# Current number of archived formulas  : 0
# Current number of archived clauses   : 152
# Clause-clause subsumption calls (NU) : 766766
# Rec. Clause-clause subsumption calls : 16147
# Non-unit clause-clause subsumptions  : 1757
# Unit Clause-clause subsumption calls : 949
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 28
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 476039

# -------------------------------------------------
# User time                : 49.017 s
# System time              : 1.268 s
# Total time               : 50.285 s
# Maximum resident set size: 470860 pages
49.01user 1.33system 0:50.36elapsed 99%CPU (0avgtext+0avgdata 471260maxresident)k
0inputs+280outputs (0major+458688minor)pagefaults 0swaps
