2020-06-03 04:47:26.269217: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_35
2020-06-03 04:47:26.368951: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-06-03 04:47:26.502887: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-06-03 04:47:26.997599: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-06-03 04:47:29.753683: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 3484476 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_35' loaded (query=512; context=768; weigths=real; len_mult=0.000000).
# Preprocessing time       : 4.509 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', t4_subset)).
fof(dt_k1_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>m1_subset_1(k1_relset_1(X1,X2),k1_zfmisc_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', dt_k1_relset_1)).
fof(t31_goboard7, axiom, ![X1]:((((((((~(v3_funct_1(X1))&~(v1_xboole_0(X1)))&v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))&v1_topreal1(X1))&v2_topreal1(X1))&v1_goboard5(X1))&v2_goboard5(X1))&m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))=>?[X2]:((m2_subset_1(X2,k1_numbers,k5_numbers)&r2_hidden(X2,k4_finseq_1(X1)))&k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1)))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', t31_goboard7)).
fof(redefinition_k4_finseq_1, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))=>k4_finseq_1(X1)=k9_xtuple_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', redefinition_k4_finseq_1)).
fof(dt_m1_finseq_1, axiom, ![X1, X2]:(m1_finseq_1(X2,X1)=>((v1_relat_1(X2)&v1_funct_1(X2))&v1_finseq_1(X2))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', dt_m1_finseq_1)).
fof(redefinition_m2_finseq_1, axiom, ![X1, X2]:(m2_finseq_1(X2,X1)<=>m1_finseq_1(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', redefinition_m2_finseq_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>k1_relset_1(X1,X2)=k9_xtuple_0(X2)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', redefinition_k1_relset_1)).
fof(dt_m2_finseq_1, axiom, ![X1, X2]:(m2_finseq_1(X2,X1)=>((v1_funct_1(X2)&v1_finseq_1(X2))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', dt_m2_finseq_1)).
fof(t13_nat_1, axiom, ![X1]:(v7_ordinal1(X1)=>![X2]:(v7_ordinal1(X2)=>(~(r1_xxreal_0(k1_nat_1(X2,np__1),X1))<=>r1_xxreal_0(X1,X2)))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', t13_nat_1)).
fof(commutativity_k1_nat_1, axiom, ![X1, X2]:((v7_ordinal1(X1)&m1_subset_1(X2,k5_numbers))=>k1_nat_1(X1,X2)=k1_nat_1(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', commutativity_k1_nat_1)).
fof(cc17_membered, axiom, ![X1]:(v6_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v7_ordinal1(X2))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', cc17_membered)).
fof(fc6_membered, axiom, v6_membered(k4_ordinal1), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', fc6_membered)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', redefinition_k5_numbers)).
fof(t25_finseq_3, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))=>![X2]:(v7_ordinal1(X2)=>(r2_hidden(X2,k1_relset_1(k5_numbers,X1))<=>(r1_xxreal_0(np__1,X2)&r1_xxreal_0(X2,k3_finseq_1(X1)))))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', t25_finseq_3)).
fof(t30_goboard7, axiom, ![X1]:((((((((~(v3_funct_1(X1))&~(v1_xboole_0(X1)))&v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))&v1_topreal1(X1))&v2_topreal1(X1))&v1_goboard5(X1))&v2_goboard5(X1))&m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))=>?[X2]:((m2_subset_1(X2,k1_numbers,k5_numbers)&r2_hidden(X2,k4_finseq_1(X1)))&k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1)))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', t30_goboard7)).
fof(cc20_membered, axiom, ![X1]:(v2_membered(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v2_membered(X2))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', cc20_membered)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', spc1_numerals)).
fof(redefinition_k1_nat_1, axiom, ![X1, X2]:((v7_ordinal1(X1)&m1_subset_1(X2,k5_numbers))=>k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', redefinition_k1_nat_1)).
fof(cc8_ordinal1, axiom, ![X1]:(m1_subset_1(X1,k4_ordinal1)=>v7_ordinal1(X1)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', cc8_ordinal1)).
fof(fc6_ordinal1, axiom, (~(v1_xboole_0(k4_ordinal1))&v3_ordinal1(k4_ordinal1)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', fc6_ordinal1)).
fof(cc5_finseq_1, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))=>(((v1_relat_1(X1)&v4_relat_1(X1,k5_numbers))&v1_funct_1(X1))&v1_finseq_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', cc5_finseq_1)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', t6_boole)).
fof(t1_xxreal_0, axiom, ![X1]:(v1_xxreal_0(X1)=>![X2]:(v1_xxreal_0(X2)=>((r1_xxreal_0(X1,X2)&r1_xxreal_0(X2,X1))=>X1=X2))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', t1_xxreal_0)).
fof(cc13_membered, axiom, ![X1]:(v2_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xxreal_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', cc13_membered)).
fof(cc21_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v3_membered(X2))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', cc21_membered)).
fof(dt_k5_numbers, axiom, m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', dt_k5_numbers)).
fof(cc4_membered, axiom, ![X1]:(v3_membered(X1)=>v2_membered(X1)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', cc4_membered)).
fof(redefinition_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', redefinition_m2_subset_1)).
fof(fc1_numbers, axiom, ~(v1_xboole_0(k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', fc1_numbers)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', redefinition_k6_numbers)).
fof(dt_k3_finseq_1, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))=>m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', dt_k3_finseq_1)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', fc3_membered)).
fof(spc2_numerals, axiom, (((v2_xxreal_0(np__2)&m2_subset_1(np__2,k1_numbers,k5_numbers))&m1_subset_1(np__2,k5_numbers))&m1_subset_1(np__2,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', spc2_numerals)).
fof(t28_nat_1, axiom, ![X1]:(v7_ordinal1(X1)=>~((((((r1_xxreal_0(X1,np__4)&X1!=k6_numbers)&X1!=np__1)&X1!=np__2)&X1!=np__3)&X1!=np__4))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', t28_nat_1)).
fof(spc0_boole, axiom, v1_xboole_0(np__0), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', spc0_boole)).
fof(t34_goboard7, conjecture, ![X1]:((((((((~(v3_funct_1(X1))&~(v1_xboole_0(X1)))&v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))&v1_topreal1(X1))&v2_topreal1(X1))&v1_goboard5(X1))&v2_goboard5(X1))&m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))=>~(r1_xxreal_0(k3_finseq_1(X1),np__4))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', t34_goboard7)).
fof(spc0_numerals, axiom, ((m2_subset_1(np__0,k1_numbers,k5_numbers)&m1_subset_1(np__0,k5_numbers))&m1_subset_1(np__0,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', spc0_numerals)).
fof(rqRealAdd__k2_xcmplx_0__r1_r0_r1, axiom, k2_xcmplx_0(np__1,np__0)=np__1, file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', rqRealAdd__k2_xcmplx_0__r1_r0_r1)).
fof(spc3_numerals, axiom, (((v2_xxreal_0(np__3)&m2_subset_1(np__3,k1_numbers,k5_numbers))&m1_subset_1(np__3,k5_numbers))&m1_subset_1(np__3,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', spc3_numerals)).
fof(rqRealAdd__k2_xcmplx_0__r1_r1_r2, axiom, k2_xcmplx_0(np__1,np__1)=np__2, file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', rqRealAdd__k2_xcmplx_0__r1_r1_r2)).
fof(rqRealAdd__k2_xcmplx_0__r0_r1_r1, axiom, k2_xcmplx_0(np__0,np__1)=np__1, file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', rqRealAdd__k2_xcmplx_0__r0_r1_r1)).
fof(t2_xxreal_0, axiom, ![X1]:(v1_xxreal_0(X1)=>![X2]:(v1_xxreal_0(X2)=>![X3]:(v1_xxreal_0(X3)=>((r1_xxreal_0(X1,X2)&r1_xxreal_0(X2,X3))=>r1_xxreal_0(X1,X3))))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', t2_xxreal_0)).
fof(rqRealAdd__k2_xcmplx_0__r2_r1_r3, axiom, k2_xcmplx_0(np__2,np__1)=np__3, file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', rqRealAdd__k2_xcmplx_0__r2_r1_r3)).
fof(redefinition_k3_finseq_1, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))=>k3_finseq_1(X1)=k1_card_1(X1)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', redefinition_k3_finseq_1)).
fof(reflexivity_r1_xxreal_0, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>r1_xxreal_0(X1,X1)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', reflexivity_r1_xxreal_0)).
fof(rc1_int_1, axiom, ?[X1]:((((m1_subset_1(X1,k1_numbers)&v1_xxreal_0(X1))&v1_xcmplx_0(X1))&v1_xreal_0(X1))&v1_int_1(X1)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', rc1_int_1)).
fof(cc4_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>v1_xboole_0(X3))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', cc4_relset_1)).
fof(l30_goboard7, axiom, ![X1]:((((((((~(v3_funct_1(X1))&~(v1_xboole_0(X1)))&v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))&v1_topreal1(X1))&v2_topreal1(X1))&v1_goboard5(X1))&v2_goboard5(X1))&m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))=>~(r1_xxreal_0(k3_finseq_1(X1),np__1))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', l30_goboard7)).
fof(rqLessOrEqual__r1_xxreal_0__r1_r0, axiom, ~(r1_xxreal_0(np__1,np__0)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', rqLessOrEqual__r1_xxreal_0__r1_r0)).
fof(d1_finseq_6, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m2_finseq_1(X2,X1)=>(v1_finseq_6(X2,X1)<=>k7_partfun1(X1,X2,np__1)=k7_partfun1(X1,X2,k3_finseq_1(X2))))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', d1_finseq_6)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', cc2_relset_1)).
fof(cc19_membered, axiom, ![X1]:(v1_membered(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_membered(X2))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', cc19_membered)).
fof(spc4_numerals, axiom, (((v2_xxreal_0(np__4)&m2_subset_1(np__4,k1_numbers,k5_numbers))&m1_subset_1(np__4,k5_numbers))&m1_subset_1(np__4,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', spc4_numerals)).
fof(rqLessOrEqual__r1_xxreal_0__r1_r1, axiom, r1_xxreal_0(np__1,np__1), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', rqLessOrEqual__r1_xxreal_0__r1_r1)).
fof(rqLessOrEqual__r1_xxreal_0__r3_r2, axiom, ~(r1_xxreal_0(np__3,np__2)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', rqLessOrEqual__r1_xxreal_0__r3_r2)).
fof(cc12_membered, axiom, ![X1]:(v1_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xcmplx_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', cc12_membered)).
fof(cc5_membered, axiom, ![X1]:(v3_membered(X1)=>v1_membered(X1)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', cc5_membered)).
fof(d5_topreal1, axiom, ![X1]:(m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))=>(v1_topreal1(X1)<=>![X2]:(v7_ordinal1(X2)=>~((((r1_xxreal_0(np__1,X2)&r1_xxreal_0(k1_nat_1(X2,np__1),k3_finseq_1(X1)))&k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1))))&k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)))))))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', d5_topreal1)).
fof(dt_k1_nat_1, axiom, ![X1, X2]:((v7_ordinal1(X1)&m1_subset_1(X2,k5_numbers))=>m2_subset_1(k1_nat_1(X1,X2),k1_numbers,k5_numbers)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', dt_k1_nat_1)).
fof(t1_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k2_xcmplx_0(X1,k6_numbers)=X1), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', t1_arithm)).
fof(rqRealAdd__k2_xcmplx_0__r3_r1_r4, axiom, k2_xcmplx_0(np__3,np__1)=np__4, file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', rqRealAdd__k2_xcmplx_0__r3_r1_r4)).
fof(rqLessOrEqual__r1_xxreal_0__r2_r2, axiom, r1_xxreal_0(np__2,np__2), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', rqLessOrEqual__r1_xxreal_0__r2_r2)).
fof(rqLessOrEqual__r1_xxreal_0__r1_r3, axiom, r1_xxreal_0(np__1,np__3), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', rqLessOrEqual__r1_xxreal_0__r1_r3)).
fof(t29_goboard7, axiom, ![X1]:(m2_subset_1(X1,k1_numbers,k5_numbers)=>![X2]:(((~(v1_xboole_0(X2))&v2_goboard5(X2))&m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2))))=>~(((r2_hidden(X1,k4_finseq_1(X2))&r2_hidden(k2_nat_1(X1,np__1),k4_finseq_1(X2)))&k7_partfun1(u1_struct_0(k15_euclid(np__2)),X2,X1)=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X2,k2_nat_1(X1,np__1)))))), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', t29_goboard7)).
fof(rqLessOrEqual__r1_xxreal_0__r1_r2, axiom, r1_xxreal_0(np__1,np__2), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', rqLessOrEqual__r1_xxreal_0__r1_r2)).
fof(redefinition_k2_nat_1, axiom, ![X1, X2]:((m1_subset_1(X1,k5_numbers)&v7_ordinal1(X2))=>k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', redefinition_k2_nat_1)).
fof(rqLessOrEqual__r1_xxreal_0__r3_r4, axiom, r1_xxreal_0(np__3,np__4), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', rqLessOrEqual__r1_xxreal_0__r3_r4)).
fof(rqLessOrEqual__r1_xxreal_0__r4_r4, axiom, r1_xxreal_0(np__4,np__4), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', rqLessOrEqual__r1_xxreal_0__r4_r4)).
fof(rqLessOrEqual__r1_xxreal_0__r2_r3, axiom, r1_xxreal_0(np__2,np__3), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', rqLessOrEqual__r1_xxreal_0__r2_r3)).
fof(rqLessOrEqual__r1_xxreal_0__r2_r4, axiom, r1_xxreal_0(np__2,np__4), file('/scratch/mptp/allpred/preds__0.005/t34_goboard7', rqLessOrEqual__r1_xxreal_0__r2_r4)).
fof(c_0_70, plain, ![X74, X75, X76]:(~r2_hidden(X74,X75)|~m1_subset_1(X75,k1_zfmisc_1(X76))|m1_subset_1(X74,X76)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_71, plain, ![X60, X61]:(~v1_relat_1(X61)|~v4_relat_1(X61,X60)|m1_subset_1(k1_relset_1(X60,X61),k1_zfmisc_1(X60))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_relset_1])])).
fof(c_0_72, plain, ![X10]:(((m2_subset_1(esk3_1(X10),k1_numbers,k5_numbers)|(v3_funct_1(X10)|v1_xboole_0(X10)|~v1_finseq_6(X10,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X10)|~v2_topreal1(X10)|~v1_goboard5(X10)|~v2_goboard5(X10)|~m2_finseq_1(X10,u1_struct_0(k15_euclid(np__2)))))&(r2_hidden(esk3_1(X10),k4_finseq_1(X10))|(v3_funct_1(X10)|v1_xboole_0(X10)|~v1_finseq_6(X10,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X10)|~v2_topreal1(X10)|~v1_goboard5(X10)|~v2_goboard5(X10)|~m2_finseq_1(X10,u1_struct_0(k15_euclid(np__2))))))&(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X10,esk3_1(X10)))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X10,np__1))|(v3_funct_1(X10)|v1_xboole_0(X10)|~v1_finseq_6(X10,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X10)|~v2_topreal1(X10)|~v1_goboard5(X10)|~v2_goboard5(X10)|~m2_finseq_1(X10,u1_struct_0(k15_euclid(np__2)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t31_goboard7])])])])])).
fof(c_0_73, plain, ![X45]:(~v1_relat_1(X45)|~v1_funct_1(X45)|~v1_finseq_1(X45)|k4_finseq_1(X45)=k9_xtuple_0(X45)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_finseq_1])])).
fof(c_0_74, plain, ![X34, X35]:(((v1_relat_1(X35)|~m1_finseq_1(X35,X34))&(v1_funct_1(X35)|~m1_finseq_1(X35,X34)))&(v1_finseq_1(X35)|~m1_finseq_1(X35,X34))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_finseq_1])])])).
fof(c_0_75, plain, ![X12, X13]:((~m2_finseq_1(X13,X12)|m1_finseq_1(X13,X12))&(~m1_finseq_1(X13,X12)|m2_finseq_1(X13,X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_m2_finseq_1])])).
cnf(c_0_76, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_77, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_71])).
fof(c_0_78, plain, ![X37, X38]:(~v1_relat_1(X38)|~v4_relat_1(X38,X37)|k1_relset_1(X37,X38)=k9_xtuple_0(X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_79, plain, (r2_hidden(esk3_1(X1),k4_finseq_1(X1))|v3_funct_1(X1)|v1_xboole_0(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)|~v2_topreal1(X1)|~v1_goboard5(X1)|~v2_goboard5(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_80, plain, (k4_finseq_1(X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_81, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_82, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_75])).
fof(c_0_83, plain, ![X28, X29]:(((v1_funct_1(X29)|~m2_finseq_1(X29,X28))&(v1_finseq_1(X29)|~m2_finseq_1(X29,X28)))&(m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X28)))|~m2_finseq_1(X29,X28))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m2_finseq_1])])])).
cnf(c_0_84, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_76])).
cnf(c_0_85, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_77])).
cnf(c_0_86, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_87, plain, (v3_funct_1(X1)|v1_xboole_0(X1)|r2_hidden(esk3_1(X1),k4_finseq_1(X1))|~v1_topreal1(X1)|~v2_topreal1(X1)|~v1_goboard5(X1)|~v2_goboard5(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[c_0_79])).
cnf(c_0_88, plain, (k4_finseq_1(X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_80])).
cnf(c_0_89, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_81])).
cnf(c_0_90, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_82])).
cnf(c_0_91, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_92, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_83])).
fof(c_0_93, plain, ![X24, X25]:((r1_xxreal_0(k1_nat_1(X25,np__1),X24)|r1_xxreal_0(X24,X25)|~v7_ordinal1(X25)|~v7_ordinal1(X24))&(~r1_xxreal_0(X24,X25)|~r1_xxreal_0(k1_nat_1(X25,np__1),X24)|~v7_ordinal1(X25)|~v7_ordinal1(X24))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t13_nat_1])])])])])).
fof(c_0_94, plain, ![X57, X58]:(~v7_ordinal1(X57)|~m1_subset_1(X58,k5_numbers)|k1_nat_1(X57,X58)=k1_nat_1(X58,X57)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k1_nat_1])])).
fof(c_0_95, plain, ![X32, X33]:(~v6_membered(X32)|(~m1_subset_1(X33,X32)|v7_ordinal1(X33))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc17_membered])])])).
cnf(c_0_96, plain, (v6_membered(k4_ordinal1)), inference(split_conjunct,[status(thm)],[fc6_membered])).
cnf(c_0_97, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
fof(c_0_98, plain, ![X18, X19]:(((r1_xxreal_0(np__1,X19)|~r2_hidden(X19,k1_relset_1(k5_numbers,X18))|~v7_ordinal1(X19)|(~v1_relat_1(X18)|~v1_funct_1(X18)|~v1_finseq_1(X18)))&(r1_xxreal_0(X19,k3_finseq_1(X18))|~r2_hidden(X19,k1_relset_1(k5_numbers,X18))|~v7_ordinal1(X19)|(~v1_relat_1(X18)|~v1_funct_1(X18)|~v1_finseq_1(X18))))&(~r1_xxreal_0(np__1,X19)|~r1_xxreal_0(X19,k3_finseq_1(X18))|r2_hidden(X19,k1_relset_1(k5_numbers,X18))|~v7_ordinal1(X19)|(~v1_relat_1(X18)|~v1_funct_1(X18)|~v1_finseq_1(X18)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t25_finseq_3])])])])).
fof(c_0_99, plain, ![X16]:(((m2_subset_1(esk4_1(X16),k1_numbers,k5_numbers)|(v3_funct_1(X16)|v1_xboole_0(X16)|~v1_finseq_6(X16,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X16)|~v2_topreal1(X16)|~v1_goboard5(X16)|~v2_goboard5(X16)|~m2_finseq_1(X16,u1_struct_0(k15_euclid(np__2)))))&(r2_hidden(esk4_1(X16),k4_finseq_1(X16))|(v3_funct_1(X16)|v1_xboole_0(X16)|~v1_finseq_6(X16,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X16)|~v2_topreal1(X16)|~v1_goboard5(X16)|~v2_goboard5(X16)|~m2_finseq_1(X16,u1_struct_0(k15_euclid(np__2))))))&(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X16,esk4_1(X16)))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X16,np__1))|(v3_funct_1(X16)|v1_xboole_0(X16)|~v1_finseq_6(X16,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X16)|~v2_topreal1(X16)|~v1_goboard5(X16)|~v2_goboard5(X16)|~m2_finseq_1(X16,u1_struct_0(k15_euclid(np__2)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t30_goboard7])])])])])).
cnf(c_0_100, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,X3))|~v4_relat_1(X3,X2)|~v1_relat_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85])])).
cnf(c_0_101, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_86])).
cnf(c_0_102, plain, (r2_hidden(esk3_1(X1),k9_xtuple_0(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88])])).
cnf(c_0_103, plain, (v1_relat_1(X1)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90])])).
cnf(c_0_104, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_91])).
cnf(c_0_105, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_92])).
fof(c_0_106, plain, ![X98, X99]:(~v2_membered(X98)|(~m1_subset_1(X99,k1_zfmisc_1(X98))|v2_membered(X99))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc20_membered])])])).
cnf(c_0_107, plain, (~r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_nat_1(X2,np__1),X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_108, plain, (k1_nat_1(X1,X2)=k1_nat_1(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_109, plain, (m1_subset_1(np__1,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
fof(c_0_110, plain, ![X84, X85]:(~v7_ordinal1(X84)|~m1_subset_1(X85,k5_numbers)|k1_nat_1(X84,X85)=k2_xcmplx_0(X84,X85)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_nat_1])])).
cnf(c_0_111, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_112, plain, (v6_membered(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_96])).
cnf(c_0_113, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_97])).
fof(c_0_114, plain, ![X36]:(~m1_subset_1(X36,k4_ordinal1)|v7_ordinal1(X36)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_ordinal1])])).
fof(c_0_115, plain, (~v1_xboole_0(k4_ordinal1)&v3_ordinal1(k4_ordinal1)), inference(fof_simplification,[status(thm)],[fc6_ordinal1])).
cnf(c_0_116, plain, (r1_xxreal_0(np__1,X1)|~r2_hidden(X1,k1_relset_1(k5_numbers,X2))|~v7_ordinal1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)), inference(split_conjunct,[status(thm)],[c_0_98])).
fof(c_0_117, plain, ![X101]:((((v1_relat_1(X101)|(~v1_relat_1(X101)|~v1_funct_1(X101)|~v1_finseq_1(X101)))&(v4_relat_1(X101,k5_numbers)|(~v1_relat_1(X101)|~v1_funct_1(X101)|~v1_finseq_1(X101))))&(v1_funct_1(X101)|(~v1_relat_1(X101)|~v1_funct_1(X101)|~v1_finseq_1(X101))))&(v1_finseq_1(X101)|(~v1_relat_1(X101)|~v1_funct_1(X101)|~v1_finseq_1(X101)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc5_finseq_1])])])).
fof(c_0_118, plain, ![X59]:(~v1_xboole_0(X59)|X59=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
cnf(c_0_119, plain, (r2_hidden(esk4_1(X1),k4_finseq_1(X1))|v3_funct_1(X1)|v1_xboole_0(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)|~v2_topreal1(X1)|~v1_goboard5(X1)|~v2_goboard5(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(split_conjunct,[status(thm)],[c_0_99])).
fof(c_0_120, plain, ![X14, X15]:(~v1_xxreal_0(X14)|(~v1_xxreal_0(X15)|(~r1_xxreal_0(X14,X15)|~r1_xxreal_0(X15,X14)|X14=X15))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xxreal_0])])])).
fof(c_0_121, plain, ![X46, X47]:(~v2_membered(X46)|(~m1_subset_1(X47,X46)|v1_xxreal_0(X47))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc13_membered])])])).
fof(c_0_122, plain, ![X67, X68]:(~v3_membered(X67)|(~m1_subset_1(X68,k1_zfmisc_1(X67))|v3_membered(X68))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc21_membered])])])).
cnf(c_0_123, plain, (m1_subset_1(X1,X2)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~r2_hidden(X1,k9_xtuple_0(X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101])])).
cnf(c_0_124, plain, (r2_hidden(esk3_1(X1),k9_xtuple_0(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_102, c_0_103]), c_0_104]), c_0_105])).
cnf(c_0_125, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_106])).
cnf(c_0_126, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[dt_k5_numbers])).
fof(c_0_127, plain, ![X44]:(~v3_membered(X44)|v2_membered(X44)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_membered])])).
cnf(c_0_128, plain, (~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_nat_1(X2,np__1),X1)), inference(evalgc,[status(thm)],[c_0_107])).
cnf(c_0_129, plain, (k1_nat_1(X1,X2)=k1_nat_1(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_108])).
cnf(c_0_130, plain, (m1_subset_1(np__1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_109])).
cnf(c_0_131, plain, (k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_132, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_111])).
cnf(c_0_133, plain, (v6_membered(k5_numbers)), inference(rw,[status(thm)],[c_0_112, c_0_113])).
cnf(c_0_134, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_114])).
fof(c_0_135, plain, ![X50, X51, X52]:((~m2_subset_1(X52,X50,X51)|m1_subset_1(X52,X51)|(v1_xboole_0(X50)|v1_xboole_0(X51)|~m1_subset_1(X51,k1_zfmisc_1(X50))))&(~m1_subset_1(X52,X51)|m2_subset_1(X52,X50,X51)|(v1_xboole_0(X50)|v1_xboole_0(X51)|~m1_subset_1(X51,k1_zfmisc_1(X50))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m2_subset_1])])])])])).
fof(c_0_136, plain, ~v1_xboole_0(k1_numbers), inference(fof_simplification,[status(thm)],[fc1_numbers])).
cnf(c_0_137, plain, (~v1_xboole_0(k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_115])).
cnf(c_0_138, plain, (r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)|~r2_hidden(X1,k1_relset_1(k5_numbers,X2))), inference(evalgc,[status(thm)],[c_0_116])).
cnf(c_0_139, plain, (v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_117])).
cnf(c_0_140, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|~r2_hidden(X1,k1_relset_1(k5_numbers,X2))|~v7_ordinal1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_141, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_142, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
fof(c_0_143, plain, ![X53]:(~v1_relat_1(X53)|~v1_funct_1(X53)|~v1_finseq_1(X53)|m2_subset_1(k3_finseq_1(X53),k1_numbers,k5_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_finseq_1])])).
cnf(c_0_144, plain, (v3_funct_1(X1)|v1_xboole_0(X1)|r2_hidden(esk4_1(X1),k4_finseq_1(X1))|~v1_topreal1(X1)|~v2_topreal1(X1)|~v1_goboard5(X1)|~v2_goboard5(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[c_0_119])).
cnf(c_0_145, plain, (X1=X2|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_146, plain, (r1_xxreal_0(k1_nat_1(X1,np__1),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_147, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_148, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_122])).
cnf(c_0_149, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
cnf(c_0_150, plain, (m1_subset_1(esk3_1(X1),X2)|v1_xboole_0(X1)|v3_funct_1(X1)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_124])])).
cnf(c_0_151, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_125])).
cnf(c_0_152, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_126])).
cnf(c_0_153, plain, (v2_membered(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_154, plain, (~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(k1_nat_1(np__1,X1),X2)|~r1_xxreal_0(X2,X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_129]), c_0_130])])])).
cnf(c_0_155, plain, (k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_131])).
cnf(c_0_156, plain, (v7_ordinal1(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_130]), c_0_133])])])).
cnf(c_0_157, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_134])).
cnf(c_0_158, plain, (m1_subset_1(X1,X3)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_135])).
cnf(c_0_159, plain, (m2_subset_1(esk3_1(X1),k1_numbers,k5_numbers)|v3_funct_1(X1)|v1_xboole_0(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)|~v2_topreal1(X1)|~v1_goboard5(X1)|~v2_goboard5(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_160, plain, (~v1_xboole_0(k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_136])).
cnf(c_0_161, plain, (~v1_xboole_0(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_137])).
cnf(c_0_162, plain, (r1_xxreal_0(np__1,X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k9_xtuple_0(X2))|~v7_ordinal1(X1)|~v4_relat_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_101])])).
cnf(c_0_163, plain, (v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_139])).
cnf(c_0_164, plain, (m1_subset_1(np__2,k5_numbers)), inference(split_conjunct,[status(thm)],[spc2_numerals])).
cnf(c_0_165, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|~v7_ordinal1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)|~r2_hidden(X1,k1_relset_1(k5_numbers,X2))), inference(evalgc,[status(thm)],[c_0_140])).
fof(c_0_166, plain, ![X20]:(~v7_ordinal1(X20)|(~r1_xxreal_0(X20,np__4)|X20=k6_numbers|X20=np__1|X20=np__2|X20=np__3|X20=np__4)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_nat_1])])).
cnf(c_0_167, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_141])).
cnf(c_0_168, plain, (k1_xboole_0=k6_numbers), inference(evalgc,[status(thm)],[c_0_142])).
cnf(c_0_169, plain, (v1_xboole_0(np__0)), inference(split_conjunct,[status(thm)],[spc0_boole])).
cnf(c_0_170, plain, (m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_143])).
fof(c_0_171, negated_conjecture, ~(![X1]:((((((((~(v3_funct_1(X1))&~(v1_xboole_0(X1)))&v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))&v1_topreal1(X1))&v2_topreal1(X1))&v1_goboard5(X1))&v2_goboard5(X1))&m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))=>~(r1_xxreal_0(k3_finseq_1(X1),np__4)))), inference(assume_negation,[status(cth)],[t34_goboard7])).
cnf(c_0_172, plain, (m2_subset_1(esk4_1(X1),k1_numbers,k5_numbers)|v3_funct_1(X1)|v1_xboole_0(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)|~v2_topreal1(X1)|~v1_goboard5(X1)|~v2_goboard5(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_173, plain, (r2_hidden(esk4_1(X1),k9_xtuple_0(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_88])])).
cnf(c_0_174, plain, (X1=X2|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_145])).
cnf(c_0_175, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(k1_nat_1(X1,np__1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_146])).
cnf(c_0_176, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_147])).
cnf(c_0_177, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_148])).
cnf(c_0_178, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_149])).
cnf(c_0_179, plain, (m1_subset_1(np__0,k5_numbers)), inference(split_conjunct,[status(thm)],[spc0_numerals])).
cnf(c_0_180, plain, (m1_subset_1(esk3_1(X1),X2)|v1_xboole_0(X1)|v3_funct_1(X1)|~v4_relat_1(X1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[c_0_150, c_0_103])).
cnf(c_0_181, plain, (v2_membered(k5_numbers)|~v2_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152])])).
cnf(c_0_182, plain, (v2_membered(X1)|~v3_membered(X1)), inference(evalgc,[status(thm)],[c_0_153])).
cnf(c_0_183, plain, (~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(k2_xcmplx_0(np__1,X1),X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_155]), c_0_156])])])).
cnf(c_0_184, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)), inference(rw,[status(thm)],[c_0_157, c_0_113])).
cnf(c_0_185, plain, (k2_xcmplx_0(np__1,np__0)=np__1), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r0_r1])).
cnf(c_0_186, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_158])).
cnf(c_0_187, plain, (v3_funct_1(X1)|v1_xboole_0(X1)|m2_subset_1(esk3_1(X1),k1_numbers,k5_numbers)|~v1_topreal1(X1)|~v2_topreal1(X1)|~v1_goboard5(X1)|~v2_goboard5(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[c_0_159])).
cnf(c_0_188, plain, (~v1_xboole_0(k1_numbers)), inference(evalgc,[status(thm)],[c_0_160])).
cnf(c_0_189, plain, (~v1_xboole_0(k5_numbers)), inference(rw,[status(thm)],[c_0_161, c_0_113])).
cnf(c_0_190, plain, (r1_xxreal_0(np__1,X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k9_xtuple_0(X2))|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_162, c_0_163])).
cnf(c_0_191, plain, (m1_subset_1(np__2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_164])).
cnf(c_0_192, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k9_xtuple_0(X2))|~v7_ordinal1(X1)|~v4_relat_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_101])])).
cnf(c_0_193, plain, (m1_subset_1(np__3,k5_numbers)), inference(split_conjunct,[status(thm)],[spc3_numerals])).
cnf(c_0_194, plain, (X1=k6_numbers|X1=np__1|X1=np__2|X1=np__3|X1=np__4|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__4)), inference(split_conjunct,[status(thm)],[c_0_166])).
cnf(c_0_195, plain, (X1=k6_numbers|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_167, c_0_168])).
cnf(c_0_196, plain, (v1_xboole_0(np__0)), inference(evalgc,[status(thm)],[c_0_169])).
cnf(c_0_197, plain, (m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_170])).
fof(c_0_198, negated_conjecture, ((((((((~v3_funct_1(esk1_0)&~v1_xboole_0(esk1_0))&v1_finseq_6(esk1_0,u1_struct_0(k15_euclid(np__2))))&v1_topreal1(esk1_0))&v2_topreal1(esk1_0))&v1_goboard5(esk1_0))&v2_goboard5(esk1_0))&m2_finseq_1(esk1_0,u1_struct_0(k15_euclid(np__2))))&r1_xxreal_0(k3_finseq_1(esk1_0),np__4)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_171])])])])).
cnf(c_0_199, plain, (k2_xcmplx_0(np__1,np__1)=np__2), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r1_r2])).
cnf(c_0_200, plain, (v3_funct_1(X1)|v1_xboole_0(X1)|m2_subset_1(esk4_1(X1),k1_numbers,k5_numbers)|~v1_topreal1(X1)|~v2_topreal1(X1)|~v1_goboard5(X1)|~v2_goboard5(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[c_0_172])).
cnf(c_0_201, plain, (r2_hidden(esk4_1(X1),k9_xtuple_0(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_173, c_0_103]), c_0_104]), c_0_105])).
cnf(c_0_202, plain, (X1=k1_nat_1(X2,np__1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(k1_nat_1(X2,np__1))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k1_nat_1(X2,np__1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_175])])).
cnf(c_0_203, plain, (k2_xcmplx_0(np__0,np__1)=np__1), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r0_r1_r1])).
cnf(c_0_204, plain, (v1_xxreal_0(np__1)|~v2_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_130])])).
cnf(c_0_205, plain, (v3_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_152]), c_0_178])])])).
cnf(c_0_206, plain, (m1_subset_1(np__0,k5_numbers)), inference(evalgc,[status(thm)],[c_0_179])).
cnf(c_0_207, plain, (v1_xxreal_0(esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v2_membered(X2)|~v4_relat_1(X1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_180])])).
cnf(c_0_208, plain, (v2_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_182]), c_0_178])])])).
cnf(c_0_209, plain, (~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)|~r1_xxreal_0(k2_xcmplx_0(np__1,X1),X2)|~r1_xxreal_0(X2,X1)), inference(csr,[status(thm)],[c_0_183, c_0_184])).
cnf(c_0_210, plain, (k2_xcmplx_0(np__1,np__0)=np__1), inference(evalgc,[status(thm)],[c_0_185])).
cnf(c_0_211, plain, (m1_subset_1(esk3_1(X1),k5_numbers)|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_187]), c_0_152])]), c_0_188]), c_0_189])])).
cnf(c_0_212, plain, (r1_xxreal_0(np__1,esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_190, c_0_124])])).
cnf(c_0_213, plain, (v1_xxreal_0(np__2)|~v2_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_191])])).
fof(c_0_214, plain, ![X21, X22, X23]:(~v1_xxreal_0(X21)|(~v1_xxreal_0(X22)|(~v1_xxreal_0(X23)|(~r1_xxreal_0(X21,X22)|~r1_xxreal_0(X22,X23)|r1_xxreal_0(X21,X23))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_xxreal_0])])])).
cnf(c_0_215, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k9_xtuple_0(X2))|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_192, c_0_163])).
cnf(c_0_216, plain, (m1_subset_1(np__3,k5_numbers)), inference(evalgc,[status(thm)],[c_0_193])).
cnf(c_0_217, plain, (X1=np__2|X1=np__4|X1=np__1|X1=k6_numbers|X1=np__3|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__4)), inference(evalgc,[status(thm)],[c_0_194])).
cnf(c_0_218, plain, (k6_numbers=np__0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_196])])).
cnf(c_0_219, plain, (m1_subset_1(k3_finseq_1(X1),k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_197]), c_0_152])]), c_0_188]), c_0_189])])).
cnf(c_0_220, negated_conjecture, (m2_finseq_1(esk1_0,u1_struct_0(k15_euclid(np__2)))), inference(split_conjunct,[status(thm)],[c_0_198])).
cnf(c_0_221, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__1),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_155]), c_0_130])])])).
cnf(c_0_222, plain, (k2_xcmplx_0(np__1,np__1)=np__2), inference(evalgc,[status(thm)],[c_0_199])).
cnf(c_0_223, plain, (m1_subset_1(esk4_1(X1),k5_numbers)|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_200]), c_0_152])]), c_0_188]), c_0_189])])).
cnf(c_0_224, plain, (m1_subset_1(esk4_1(X1),X2)|v1_xboole_0(X1)|v3_funct_1(X1)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_201])])).
cnf(c_0_225, plain, (X1=k2_xcmplx_0(X2,np__1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(k2_xcmplx_0(X2,np__1))|~v1_xxreal_0(X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(X1,k2_xcmplx_0(X2,np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_155]), c_0_130])])])).
cnf(c_0_226, plain, (k2_xcmplx_0(np__0,np__1)=np__1), inference(evalgc,[status(thm)],[c_0_203])).
cnf(c_0_227, plain, (v1_xxreal_0(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_182]), c_0_205])])])).
cnf(c_0_228, plain, (v7_ordinal1(np__0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_206]), c_0_133])])])).
cnf(c_0_229, plain, (v1_xxreal_0(esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_163]), c_0_208])])])).
cnf(c_0_230, plain, (~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_210]), c_0_206])])])).
cnf(c_0_231, plain, (v7_ordinal1(esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_211]), c_0_133])])])).
cnf(c_0_232, plain, (r1_xxreal_0(np__1,esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_212, c_0_103]), c_0_104]), c_0_105])).
cnf(c_0_233, plain, (v1_xxreal_0(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_182]), c_0_205])])])).
cnf(c_0_234, plain, (r1_xxreal_0(X1,X3)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_214])).
cnf(c_0_235, plain, (r1_xxreal_0(esk3_1(X1),k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_124])])).
cnf(c_0_236, plain, (k2_xcmplx_0(np__2,np__1)=np__3), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r2_r1_r3])).
cnf(c_0_237, plain, (v1_xxreal_0(np__3)|~v2_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_216])])).
cnf(c_0_238, plain, (X1=np__3|X1=np__0|X1=np__1|X1=np__4|X1=np__2|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__4)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_217, c_0_218])])).
cnf(c_0_239, plain, (v7_ordinal1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_219]), c_0_133])])])).
cnf(c_0_240, negated_conjecture, (r1_xxreal_0(k3_finseq_1(esk1_0),np__4)), inference(split_conjunct,[status(thm)],[c_0_198])).
cnf(c_0_241, negated_conjecture, (m2_finseq_1(esk1_0,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[c_0_220])).
cnf(c_0_242, plain, (r1_xxreal_0(np__2,X1)|r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_222]), c_0_156])])])).
cnf(c_0_243, plain, (v7_ordinal1(esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_223]), c_0_133])])])).
cnf(c_0_244, plain, (r1_xxreal_0(np__1,esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_190, c_0_201])])).
cnf(c_0_245, plain, (m1_subset_1(esk4_1(X1),X2)|v1_xboole_0(X1)|v3_funct_1(X1)|~v4_relat_1(X1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[c_0_224, c_0_103])).
cnf(c_0_246, plain, (X1=np__1|r1_xxreal_0(X1,np__0)|~v1_xxreal_0(X1)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_226]), c_0_227]), c_0_228])])])).
cnf(c_0_247, plain, (v1_xxreal_0(esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_229, c_0_103]), c_0_104]), c_0_105])).
cnf(c_0_248, plain, (v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(np__1,esk3_1(X1))|~r1_xxreal_0(esk3_1(X1),np__0)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_231])])).
cnf(c_0_249, plain, (r1_xxreal_0(np__1,esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_231])])).
cnf(c_0_250, plain, (X1=np__2|r1_xxreal_0(X1,np__1)|~v1_xxreal_0(X1)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_222]), c_0_233]), c_0_156])])])).
cnf(c_0_251, plain, (r1_xxreal_0(X1,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_234])).
cnf(c_0_252, plain, (r1_xxreal_0(esk3_1(X1),k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_235, c_0_103]), c_0_104]), c_0_105])).
fof(c_0_253, plain, ![X43]:(~v1_relat_1(X43)|~v1_funct_1(X43)|~v1_finseq_1(X43)|k3_finseq_1(X43)=k1_card_1(X43)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k3_finseq_1])])).
fof(c_0_254, plain, ![X77, X78]:(~v1_xxreal_0(X77)|~v1_xxreal_0(X78)|r1_xxreal_0(X77,X77)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[reflexivity_r1_xxreal_0])])).
fof(c_0_255, plain, ((((m1_subset_1(esk8_0,k1_numbers)&v1_xxreal_0(esk8_0))&v1_xcmplx_0(esk8_0))&v1_xreal_0(esk8_0))&v1_int_1(esk8_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc1_int_1])])).
cnf(c_0_256, plain, (k2_xcmplx_0(np__2,np__1)=np__3), inference(evalgc,[status(thm)],[c_0_236])).
cnf(c_0_257, plain, (v1_xxreal_0(np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_182]), c_0_205])])])).
cnf(c_0_258, plain, (v7_ordinal1(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_191]), c_0_133])])])).
cnf(c_0_259, plain, (k3_finseq_1(X1)=np__2|k3_finseq_1(X1)=np__4|k3_finseq_1(X1)=np__1|k3_finseq_1(X1)=np__0|k3_finseq_1(X1)=np__3|~r1_xxreal_0(k3_finseq_1(X1),np__4)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_239])])).
cnf(c_0_260, negated_conjecture, (r1_xxreal_0(k3_finseq_1(esk1_0),np__4)), inference(evalgc,[status(thm)],[c_0_240])).
cnf(c_0_261, negated_conjecture, (v1_finseq_1(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_241])])).
cnf(c_0_262, negated_conjecture, (v1_funct_1(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_241])])).
cnf(c_0_263, negated_conjecture, (v1_relat_1(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_241])])).
cnf(c_0_264, negated_conjecture, (v2_goboard5(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_198])).
cnf(c_0_265, negated_conjecture, (v1_goboard5(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_198])).
cnf(c_0_266, negated_conjecture, (v2_topreal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_198])).
cnf(c_0_267, negated_conjecture, (v1_topreal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_198])).
cnf(c_0_268, negated_conjecture, (v1_finseq_6(esk1_0,u1_struct_0(k15_euclid(np__2)))), inference(split_conjunct,[status(thm)],[c_0_198])).
cnf(c_0_269, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_198])).
cnf(c_0_270, negated_conjecture, (~v3_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_198])).
cnf(c_0_271, plain, (r1_xxreal_0(esk4_1(X1),np__1)|r1_xxreal_0(np__2,esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_243])])).
cnf(c_0_272, plain, (r1_xxreal_0(np__1,esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_244, c_0_103]), c_0_104]), c_0_105])).
cnf(c_0_273, plain, (v1_xxreal_0(esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v2_membered(X2)|~v4_relat_1(X1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_245])])).
cnf(c_0_274, plain, (esk3_1(X1)=np__1|r1_xxreal_0(esk3_1(X1),np__0)|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(esk3_1(X1))|~r1_xxreal_0(esk3_1(X1),np__1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_247])])).
cnf(c_0_275, plain, (v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(esk3_1(X1),np__0)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[c_0_248, c_0_249])).
cnf(c_0_276, plain, (esk3_1(X1)=np__2|r1_xxreal_0(esk3_1(X1),np__1)|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(esk3_1(X1))|~r1_xxreal_0(esk3_1(X1),np__2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_247])])).
cnf(c_0_277, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|v1_xboole_0(X2)|v3_funct_1(X2)|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(esk3_1(X2))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk3_1(X2))|~v7_ordinal1(esk3_1(X2))|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_252])])).
cnf(c_0_278, plain, (k3_finseq_1(X1)=k1_card_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_253])).
cnf(c_0_279, plain, (r1_xxreal_0(esk4_1(X1),k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_201])])).
cnf(c_0_280, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_254])).
cnf(c_0_281, plain, (v1_xxreal_0(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_255])).
cnf(c_0_282, plain, (X1=np__3|r1_xxreal_0(X1,np__2)|~v1_xxreal_0(X1)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_256]), c_0_257]), c_0_258])])])).
cnf(c_0_283, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_259, c_0_260]), c_0_261]), c_0_262]), c_0_263])])])).
cnf(c_0_284, negated_conjecture, (v2_goboard5(esk1_0)), inference(evalgc,[status(thm)],[c_0_264])).
cnf(c_0_285, negated_conjecture, (v1_goboard5(esk1_0)), inference(evalgc,[status(thm)],[c_0_265])).
cnf(c_0_286, negated_conjecture, (v2_topreal1(esk1_0)), inference(evalgc,[status(thm)],[c_0_266])).
cnf(c_0_287, negated_conjecture, (v1_topreal1(esk1_0)), inference(evalgc,[status(thm)],[c_0_267])).
cnf(c_0_288, negated_conjecture, (v1_finseq_6(esk1_0,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[c_0_268])).
cnf(c_0_289, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_269])).
cnf(c_0_290, negated_conjecture, (~v3_funct_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_270])).
fof(c_0_291, plain, ![X109, X110, X111]:(~v1_xboole_0(X109)|(~m1_subset_1(X111,k1_zfmisc_1(k2_zfmisc_1(X110,X109)))|v1_xboole_0(X111))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_relset_1])])])).
cnf(c_0_292, plain, (esk4_1(X1)=np__1|r1_xxreal_0(np__2,esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(esk4_1(X1))|~r1_xxreal_0(np__1,esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_271]), c_0_227])])])).
cnf(c_0_293, plain, (r1_xxreal_0(np__1,esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_272, c_0_243])])).
cnf(c_0_294, plain, (v1_xxreal_0(esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_163]), c_0_208])])])).
cnf(c_0_295, plain, (esk3_1(X1)=np__1|v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(esk3_1(X1),np__1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_274, c_0_231]), c_0_275])).
cnf(c_0_296, plain, (esk3_1(X1)=np__2|r1_xxreal_0(esk3_1(X1),np__1)|v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(esk3_1(X1),np__2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[c_0_276, c_0_231])).
cnf(c_0_297, plain, (r1_xxreal_0(np__3,X1)|r1_xxreal_0(X1,np__2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_256]), c_0_258])])])).
fof(c_0_298, plain, ![X6]:(v3_funct_1(X6)|v1_xboole_0(X6)|~v1_finseq_6(X6,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X6)|~v2_topreal1(X6)|~v1_goboard5(X6)|~v2_goboard5(X6)|~m2_finseq_1(X6,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(k3_finseq_1(X6),np__1)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l30_goboard7])])])).
cnf(c_0_299, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|v1_xboole_0(X2)|v3_funct_1(X2)|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(esk3_1(X2))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk3_1(X2))|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[c_0_277, c_0_231])).
cnf(c_0_300, plain, (k3_finseq_1(X1)=k1_card_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_278])).
fof(c_0_301, plain, ~r1_xxreal_0(np__1,np__0), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r1_r0])).
cnf(c_0_302, plain, (r1_xxreal_0(esk4_1(X1),k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_279, c_0_103]), c_0_104]), c_0_105])).
cnf(c_0_303, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_280])).
cnf(c_0_304, plain, (v1_xxreal_0(esk8_0)), inference(evalgc,[status(thm)],[c_0_281])).
cnf(c_0_305, plain, (esk3_1(X1)=np__3|r1_xxreal_0(esk3_1(X1),np__2)|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(esk3_1(X1))|~r1_xxreal_0(esk3_1(X1),np__3)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_247])])).
cnf(c_0_306, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),np__3)|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|~v7_ordinal1(esk3_1(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_283]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
fof(c_0_307, plain, ![X26, X27]:((~v1_finseq_6(X27,X26)|k7_partfun1(X26,X27,np__1)=k7_partfun1(X26,X27,k3_finseq_1(X27))|~m2_finseq_1(X27,X26)|v1_xboole_0(X26))&(k7_partfun1(X26,X27,np__1)!=k7_partfun1(X26,X27,k3_finseq_1(X27))|v1_finseq_6(X27,X26)|~m2_finseq_1(X27,X26)|v1_xboole_0(X26))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_finseq_6])])])])])).
cnf(c_0_308, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_291])).
cnf(c_0_309, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_310, plain, (esk4_1(X1)=np__1|r1_xxreal_0(np__2,esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[c_0_292, c_0_293])).
cnf(c_0_311, plain, (v1_xxreal_0(esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_294, c_0_103]), c_0_104]), c_0_105])).
cnf(c_0_312, plain, (esk3_1(X1)=np__1|v1_xboole_0(X1)|v3_funct_1(X1)|esk3_1(X1)=np__2|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(esk3_1(X1),np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_295, c_0_296])])).
cnf(c_0_313, plain, (r1_xxreal_0(esk3_1(X1),np__2)|r1_xxreal_0(np__3,esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_297, c_0_231])])).
cnf(c_0_314, plain, (v3_funct_1(X1)|v1_xboole_0(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)|~v2_topreal1(X1)|~v1_goboard5(X1)|~v2_goboard5(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(k3_finseq_1(X1),np__1)), inference(split_conjunct,[status(thm)],[c_0_298])).
cnf(c_0_315, plain, (r1_xxreal_0(np__1,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(k3_finseq_1(X1))|~v1_xxreal_0(esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_299, c_0_249]), c_0_227])])])).
cnf(c_0_316, negated_conjecture, (k1_card_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_300, c_0_283]), c_0_261]), c_0_262]), c_0_263])])])).
cnf(c_0_317, plain, (v1_xxreal_0(np__0)|~v2_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_206])])).
cnf(c_0_318, plain, (~r1_xxreal_0(np__1,np__0)), inference(split_conjunct,[status(thm)],[c_0_301])).
cnf(c_0_319, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|v1_xboole_0(X2)|v3_funct_1(X2)|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(esk4_1(X2))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk4_1(X2))|~v7_ordinal1(esk4_1(X2))|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_302])])).
cnf(c_0_320, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_303, c_0_304])])).
cnf(c_0_321, plain, (esk3_1(X1)=np__3|r1_xxreal_0(esk3_1(X1),np__2)|v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(esk3_1(X1),np__3)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[c_0_305, c_0_231])).
cnf(c_0_322, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2|r1_xxreal_0(esk3_1(esk1_0),np__3)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_306, c_0_231]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_323, plain, (k7_partfun1(X2,X1,np__1)=k7_partfun1(X2,X1,k3_finseq_1(X1))|v1_xboole_0(X2)|~v1_finseq_6(X1,X2)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_307])).
cnf(c_0_324, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_308])).
cnf(c_0_325, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_309])).
cnf(c_0_326, plain, (esk4_1(X1)=np__1|r1_xxreal_0(np__2,esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_311])])).
cnf(c_0_327, plain, (esk4_1(X1)=np__3|r1_xxreal_0(esk4_1(X1),np__2)|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(esk4_1(X1))|~r1_xxreal_0(esk4_1(X1),np__3)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_311])])).
cnf(c_0_328, negated_conjecture, (r1_xxreal_0(esk4_1(esk1_0),np__3)|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|~v7_ordinal1(esk4_1(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_283]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
fof(c_0_329, plain, ![X91, X92, X93]:((v4_relat_1(X93,X91)|~m1_subset_1(X93,k1_zfmisc_1(k2_zfmisc_1(X91,X92))))&(v5_relat_1(X93,X92)|~m1_subset_1(X93,k1_zfmisc_1(k2_zfmisc_1(X91,X92))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
fof(c_0_330, plain, ![X227, X228]:(~v1_membered(X227)|(~m1_subset_1(X228,k1_zfmisc_1(X227))|v1_membered(X228))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc19_membered])])])).
cnf(c_0_331, plain, (m1_subset_1(np__4,k5_numbers)), inference(split_conjunct,[status(thm)],[spc4_numerals])).
cnf(c_0_332, plain, (esk3_1(X1)=np__2|esk3_1(X1)=np__1|v1_xboole_0(X1)|v3_funct_1(X1)|r1_xxreal_0(np__3,esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_312, c_0_313])])).
cnf(c_0_333, plain, (v3_funct_1(X1)|v1_xboole_0(X1)|~v1_topreal1(X1)|~v2_topreal1(X1)|~v1_goboard5(X1)|~v2_goboard5(X1)|~r1_xxreal_0(k3_finseq_1(X1),np__1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[c_0_314])).
cnf(c_0_334, plain, (r1_xxreal_0(np__1,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_247])])).
cnf(c_0_335, negated_conjecture, (k1_card_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__3), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_300, c_0_316]), c_0_261]), c_0_262]), c_0_263])])])).
cnf(c_0_336, plain, (v1_xxreal_0(np__0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_317, c_0_182]), c_0_205])])])).
cnf(c_0_337, plain, (~r1_xxreal_0(np__1,np__0)), inference(evalgc,[status(thm)],[c_0_318])).
cnf(c_0_338, plain, (r1_xxreal_0(np__1,np__1)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r1_r1])).
fof(c_0_339, plain, ~r1_xxreal_0(np__3,np__2), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r3_r2])).
cnf(c_0_340, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|v1_xboole_0(X2)|v3_funct_1(X2)|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(esk4_1(X2))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk4_1(X2))|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[c_0_319, c_0_243])).
cnf(c_0_341, plain, (r1_xxreal_0(esk4_1(X1),esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_320, c_0_311])])).
cnf(c_0_342, plain, (v3_funct_1(X1)|v1_xboole_0(X1)|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,esk3_1(X1)))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1))|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)|~v2_topreal1(X1)|~v1_goboard5(X1)|~v2_goboard5(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_343, negated_conjecture, (esk3_1(esk1_0)=np__3|r1_xxreal_0(esk3_1(esk1_0),np__2)|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_321, c_0_322]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_344, plain, (k7_partfun1(X2,X1,k3_finseq_1(X1))=k7_partfun1(X2,X1,np__1)|v1_xboole_0(X2)|~v1_finseq_6(X1,X2)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_323])).
cnf(c_0_345, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_324, c_0_325])])).
cnf(c_0_346, plain, (esk4_1(X1)=np__2|esk4_1(X1)=np__1|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(esk4_1(X1))|~r1_xxreal_0(esk4_1(X1),np__2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_326]), c_0_233])])])).
cnf(c_0_347, plain, (esk4_1(X1)=np__3|r1_xxreal_0(esk4_1(X1),np__2)|v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(esk4_1(X1),np__3)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[c_0_327, c_0_243])).
cnf(c_0_348, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2|r1_xxreal_0(esk4_1(esk1_0),np__3)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_328, c_0_243]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
fof(c_0_349, plain, ![X174, X175]:(~v1_membered(X174)|(~m1_subset_1(X175,X174)|v1_xcmplx_0(X175))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc12_membered])])])).
cnf(c_0_350, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_329])).
cnf(c_0_351, plain, (v1_membered(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_330])).
fof(c_0_352, plain, ![X154]:(~v3_membered(X154)|v1_membered(X154)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc5_membered])])).
cnf(c_0_353, plain, (m1_subset_1(np__4,k5_numbers)), inference(evalgc,[status(thm)],[c_0_331])).
cnf(c_0_354, plain, (r1_xxreal_0(np__3,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|esk3_1(X1)=np__1|esk3_1(X1)=np__2|~v1_xxreal_0(k3_finseq_1(X1))|~v1_xxreal_0(esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_299, c_0_332]), c_0_257])])])).
cnf(c_0_355, negated_conjecture, (~r1_xxreal_0(k3_finseq_1(esk1_0),np__1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_333, c_0_288]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287])]), c_0_289]), c_0_290])])).
cnf(c_0_356, negated_conjecture, (k1_card_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k1_card_1(esk1_0)=np__4), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_334, c_0_335]), c_0_336]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_337]), c_0_289]), c_0_290])])).
cnf(c_0_357, plain, (r1_xxreal_0(np__1,np__1)), inference(evalgc,[status(thm)],[c_0_338])).
cnf(c_0_358, plain, (~r1_xxreal_0(np__3,np__2)), inference(split_conjunct,[status(thm)],[c_0_339])).
fof(c_0_359, plain, ![X7, X8]:((~v1_topreal1(X7)|(~v7_ordinal1(X8)|(~r1_xxreal_0(np__1,X8)|~r1_xxreal_0(k1_nat_1(X8,np__1),k3_finseq_1(X7))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X7,X8))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X7,k1_nat_1(X8,np__1)))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X7,X8))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X7,k1_nat_1(X8,np__1)))))|~m2_finseq_1(X7,u1_struct_0(k15_euclid(np__2))))&((v7_ordinal1(esk2_1(X7))|v1_topreal1(X7)|~m2_finseq_1(X7,u1_struct_0(k15_euclid(np__2))))&((((r1_xxreal_0(np__1,esk2_1(X7))|v1_topreal1(X7)|~m2_finseq_1(X7,u1_struct_0(k15_euclid(np__2))))&(r1_xxreal_0(k1_nat_1(esk2_1(X7),np__1),k3_finseq_1(X7))|v1_topreal1(X7)|~m2_finseq_1(X7,u1_struct_0(k15_euclid(np__2)))))&(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X7,esk2_1(X7)))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X7,k1_nat_1(esk2_1(X7),np__1)))|v1_topreal1(X7)|~m2_finseq_1(X7,u1_struct_0(k15_euclid(np__2)))))&(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X7,esk2_1(X7)))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X7,k1_nat_1(esk2_1(X7),np__1)))|v1_topreal1(X7)|~m2_finseq_1(X7,u1_struct_0(k15_euclid(np__2))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_topreal1])])])])])).
cnf(c_0_360, plain, (r1_xxreal_0(esk4_1(X1),k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(k3_finseq_1(X1))|~v1_xxreal_0(esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_340, c_0_341])])).
cnf(c_0_361, plain, (v3_funct_1(X1)|v1_xboole_0(X1)|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,esk3_1(X1)))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1))|~v1_topreal1(X1)|~v2_topreal1(X1)|~v1_goboard5(X1)|~v2_goboard5(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[c_0_342])).
cnf(c_0_362, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2|esk3_1(esk1_0)=np__3), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_312, c_0_343]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_363, negated_conjecture, (k7_partfun1(X1,esk1_0,np__3)=k7_partfun1(X1,esk1_0,np__1)|v1_xboole_0(X1)|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|~m2_finseq_1(esk1_0,X1)|~v1_finseq_6(esk1_0,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_283])])).
cnf(c_0_364, negated_conjecture, (~v1_xboole_0(u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_345, c_0_241]), c_0_289])])).
cnf(c_0_365, plain, (v3_funct_1(X1)|v1_xboole_0(X1)|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,esk4_1(X1)))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1))|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)|~v2_topreal1(X1)|~v1_goboard5(X1)|~v2_goboard5(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_366, plain, (esk4_1(X1)=np__1|esk4_1(X1)=np__2|v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(esk4_1(X1),np__2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[c_0_346, c_0_311])).
cnf(c_0_367, negated_conjecture, (esk4_1(esk1_0)=np__3|r1_xxreal_0(esk4_1(esk1_0),np__2)|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_347, c_0_348]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
fof(c_0_368, plain, ![X141, X142]:(~v7_ordinal1(X141)|~m1_subset_1(X142,k5_numbers)|m2_subset_1(k1_nat_1(X141,X142),k1_numbers,k5_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_nat_1])])).
fof(c_0_369, plain, ![X138]:(~v1_xcmplx_0(X138)|k2_xcmplx_0(X138,k6_numbers)=X138), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_arithm])])).
cnf(c_0_370, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_349])).
cnf(c_0_371, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_350])).
cnf(c_0_372, plain, (v1_membered(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_351])).
cnf(c_0_373, plain, (v1_membered(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_352])).
cnf(c_0_374, plain, (k2_xcmplx_0(np__3,np__1)=np__4), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r3_r1_r4])).
cnf(c_0_375, plain, (v1_xxreal_0(np__4)|~v2_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_353])])).
cnf(c_0_376, plain, (esk3_1(X1)=np__2|esk3_1(X1)=np__1|r1_xxreal_0(np__3,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[c_0_354, c_0_247])).
cnf(c_0_377, negated_conjecture, (k1_card_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__3), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_355, c_0_356]), c_0_357])])])).
cnf(c_0_378, plain, (~r1_xxreal_0(np__3,np__2)), inference(evalgc,[status(thm)],[c_0_358])).
cnf(c_0_379, plain, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)))|~v1_topreal1(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k1_nat_1(X2,np__1),k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(split_conjunct,[status(thm)],[c_0_359])).
cnf(c_0_380, plain, (r1_xxreal_0(esk4_1(X1),k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[c_0_360, c_0_311])).
cnf(c_0_381, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_362]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_382, negated_conjecture, (k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3)=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_363, c_0_288]), c_0_241])]), c_0_364])])).
cnf(c_0_383, plain, (v3_funct_1(X1)|v1_xboole_0(X1)|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,esk4_1(X1)))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1))|~v1_topreal1(X1)|~v2_topreal1(X1)|~v1_goboard5(X1)|~v2_goboard5(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[c_0_365])).
cnf(c_0_384, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2|esk4_1(esk1_0)=np__3), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_367]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_385, plain, (r1_xxreal_0(np__2,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|esk4_1(X1)=np__1|~v1_xxreal_0(k3_finseq_1(X1))|~v1_xxreal_0(esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_340, c_0_326]), c_0_233])])])).
cnf(c_0_386, plain, (v1_xxreal_0(k3_finseq_1(X1))|~v2_membered(k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_219])])).
cnf(c_0_387, plain, (m2_subset_1(k1_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_368])).
cnf(c_0_388, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_369])).
cnf(c_0_389, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_370])).
cnf(c_0_390, plain, (v4_relat_1(X1,k5_numbers)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_371, c_0_325])])).
cnf(c_0_391, plain, (v1_membered(k5_numbers)|~v1_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_372, c_0_152])])).
cnf(c_0_392, plain, (v1_membered(X1)|~v3_membered(X1)), inference(evalgc,[status(thm)],[c_0_373])).
cnf(c_0_393, plain, (k2_xcmplx_0(np__3,np__1)=np__4), inference(evalgc,[status(thm)],[c_0_374])).
cnf(c_0_394, plain, (v1_xxreal_0(np__4)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_375, c_0_182]), c_0_205])])])).
cnf(c_0_395, plain, (v7_ordinal1(np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_216]), c_0_133])])])).
cnf(c_0_396, negated_conjecture, (r1_xxreal_0(X1,np__4)|~v1_xxreal_0(np__4)|~v1_xxreal_0(k3_finseq_1(esk1_0))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k3_finseq_1(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_260])])).
cnf(c_0_397, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|k1_card_1(esk1_0)=np__3|k1_card_1(esk1_0)=np__4), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_376, c_0_377]), c_0_233]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_378]), c_0_289]), c_0_290])])).
cnf(c_0_398, plain, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~v1_topreal1(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(k1_nat_1(X2,np__1),k3_finseq_1(X1))), inference(evalgc,[status(thm)],[c_0_379])).
cnf(c_0_399, negated_conjecture, (r1_xxreal_0(esk4_1(esk1_0),np__2)|k1_card_1(esk1_0)=np__3|k1_card_1(esk1_0)=np__4), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_380, c_0_377]), c_0_233]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_400, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_381, c_0_382])])).
cnf(c_0_401, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_384]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_402, plain, (esk4_1(X1)=np__1|r1_xxreal_0(np__2,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[c_0_385, c_0_311])).
cnf(c_0_403, plain, (v1_xxreal_0(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_386, c_0_208])])])).
cnf(c_0_404, plain, (m2_subset_1(k1_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_387])).
cnf(c_0_405, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_388])).
cnf(c_0_406, plain, (v1_xcmplx_0(esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_membered(X2)|~v4_relat_1(X1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_389, c_0_180])])).
cnf(c_0_407, negated_conjecture, (v4_relat_1(esk1_0,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_390, c_0_241])])).
cnf(c_0_408, plain, (v1_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_392]), c_0_178])])])).
cnf(c_0_409, plain, (X1=np__4|r1_xxreal_0(X1,np__3)|~v1_xxreal_0(X1)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__4)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_393]), c_0_394]), c_0_395])])])).
cnf(c_0_410, negated_conjecture, (r1_xxreal_0(X1,np__4)|~v1_xxreal_0(k3_finseq_1(esk1_0))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k3_finseq_1(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_396, c_0_394])])])).
cnf(c_0_411, negated_conjecture, (k1_card_1(esk1_0)=np__4|k1_card_1(esk1_0)=np__3|esk3_1(esk1_0)=np__2), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_397]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_412, plain, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1)))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~v7_ordinal1(X2)|~r1_xxreal_0(k2_xcmplx_0(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_398, c_0_155]), c_0_130])])])).
cnf(c_0_413, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__4|k1_card_1(esk1_0)=np__3), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_399]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_414, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|esk3_1(esk1_0)=np__2), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_400]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_415, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_401, c_0_382])])).
cnf(c_0_416, plain, (esk4_1(X1)=np__1|r1_xxreal_0(np__2,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_403])])).
cnf(c_0_417, plain, (m1_subset_1(k1_nat_1(X1,X2),k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_404]), c_0_152])]), c_0_188]), c_0_189])])).
cnf(c_0_418, plain, (k2_xcmplx_0(X1,np__0)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_405, c_0_218])])).
cnf(c_0_419, negated_conjecture, (v1_xcmplx_0(esk3_1(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_406, c_0_407]), c_0_408]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_420, plain, (esk3_1(X1)=np__4|r1_xxreal_0(esk3_1(X1),np__3)|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(esk3_1(X1))|~r1_xxreal_0(esk3_1(X1),np__4)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_409, c_0_247])])).
cnf(c_0_421, negated_conjecture, (r1_xxreal_0(X1,np__4)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k3_finseq_1(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_410, c_0_403]), c_0_261]), c_0_262]), c_0_263])])])).
cnf(c_0_422, negated_conjecture, (v1_xxreal_0(esk3_1(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_407]), c_0_208]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_423, negated_conjecture, (k1_card_1(esk1_0)=np__3|k1_card_1(esk1_0)=np__4|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_411]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_424, plain, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|~r1_xxreal_0(np__2,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_412, c_0_222]), c_0_156]), c_0_357])])])).
cnf(c_0_425, negated_conjecture, (k1_card_1(esk1_0)=np__3|k1_card_1(esk1_0)=np__4|esk4_1(esk1_0)=np__2), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_413]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_426, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_414]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_427, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|esk4_1(esk1_0)=np__2), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_415]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_428, plain, (esk4_1(X1)=np__1|r1_xxreal_0(np__2,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_416, c_0_103]), c_0_104]), c_0_105])).
cnf(c_0_429, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_417, c_0_155])])).
cnf(c_0_430, negated_conjecture, (k2_xcmplx_0(esk3_1(esk1_0),np__0)=esk3_1(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_418, c_0_419])])).
cnf(c_0_431, plain, (r1_xxreal_0(X1,esk3_1(X2))|v1_xboole_0(X2)|v3_funct_1(X2)|~v1_xxreal_0(esk3_1(X2))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_249]), c_0_227])])])).
cnf(c_0_432, plain, (esk3_1(X1)=np__4|r1_xxreal_0(esk3_1(X1),np__3)|v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(esk3_1(X1),np__4)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[c_0_420, c_0_231])).
cnf(c_0_433, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),np__4)|~v7_ordinal1(esk3_1(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_421, c_0_252]), c_0_422]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_434, negated_conjecture, (k1_card_1(esk1_0)=np__4|k1_card_1(esk1_0)=np__3|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(np__2,k3_finseq_1(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_423, c_0_424]), c_0_241]), c_0_287])])])).
cnf(c_0_435, negated_conjecture, (k1_card_1(esk1_0)=np__4|k1_card_1(esk1_0)=np__3|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_425]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_436, plain, (r1_xxreal_0(np__2,np__2)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r2_r2])).
cnf(c_0_437, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(np__2,k3_finseq_1(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_424]), c_0_241]), c_0_287])])])).
cnf(c_0_438, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_427]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_439, negated_conjecture, (esk4_1(esk1_0)=np__1|r1_xxreal_0(np__2,k3_finseq_1(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_428, c_0_288]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287])]), c_0_289]), c_0_290])])).
cnf(c_0_440, negated_conjecture, (m1_subset_1(esk3_1(esk1_0),k5_numbers)|~v7_ordinal1(esk3_1(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_429, c_0_430]), c_0_206])])])).
cnf(c_0_441, negated_conjecture, (r1_xxreal_0(X1,esk3_1(esk1_0))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_431, c_0_422]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_442, plain, (esk3_1(X1)=np__3|r1_xxreal_0(esk3_1(X1),np__2)|v1_xboole_0(X1)|v3_funct_1(X1)|esk3_1(X1)=np__4|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(esk3_1(X1),np__4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_321, c_0_432])])).
cnf(c_0_443, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),np__4)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_433, c_0_231]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_444, plain, (k7_partfun1(X1,X2,k1_card_1(X2))=k7_partfun1(X1,X2,np__1)|v1_xboole_0(X1)|~m2_finseq_1(X2,X1)|~v1_finseq_6(X2,X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_300])])).
cnf(c_0_445, negated_conjecture, (k1_card_1(esk1_0)=np__3|k1_card_1(esk1_0)=np__4|~r1_xxreal_0(np__2,k3_finseq_1(esk1_0))), inference(csr,[status(thm)],[c_0_434, c_0_435])).
cnf(c_0_446, plain, (r1_xxreal_0(np__2,np__2)), inference(evalgc,[status(thm)],[c_0_436])).
cnf(c_0_447, plain, (r1_xxreal_0(esk3_1(X1),esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_320, c_0_247])])).
cnf(c_0_448, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2|~r1_xxreal_0(np__2,k3_finseq_1(esk1_0))), inference(csr,[status(thm)],[c_0_437, c_0_438])).
cnf(c_0_449, negated_conjecture, (r1_xxreal_0(np__2,k3_finseq_1(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_439]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_450, negated_conjecture, (m1_subset_1(esk3_1(esk1_0),k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_440, c_0_231]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_451, negated_conjecture, (r1_xxreal_0(X1,esk3_1(esk1_0))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_441]), c_0_422])])])).
cnf(c_0_452, negated_conjecture, (esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__3|r1_xxreal_0(esk3_1(esk1_0),np__2)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_442, c_0_443]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_453, plain, (k7_partfun1(X1,X2,k1_card_1(X2))=k7_partfun1(X1,X2,np__1)|v1_xboole_0(X1)|~m2_finseq_1(X2,X1)|~v1_finseq_6(X2,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_444, c_0_103]), c_0_104]), c_0_105])).
cnf(c_0_454, negated_conjecture, (k1_card_1(esk1_0)=np__3|k1_card_1(esk1_0)=np__4), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_445, c_0_377]), c_0_446])])])).
cnf(c_0_455, plain, (r1_xxreal_0(esk3_1(X1),k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(k3_finseq_1(X1))|~v1_xxreal_0(esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_299, c_0_447])])).
cnf(c_0_456, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_448, c_0_449])])])).
cnf(c_0_457, negated_conjecture, (v7_ordinal1(esk3_1(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_450]), c_0_133])])])).
cnf(c_0_458, negated_conjecture, (r1_xxreal_0(np__1,esk3_1(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_451, c_0_357]), c_0_227]), c_0_357])])])).
cnf(c_0_459, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__3|esk3_1(esk1_0)=np__4), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_312, c_0_452]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_460, negated_conjecture, (k7_partfun1(X1,esk1_0,np__3)=k7_partfun1(X1,esk1_0,np__1)|v1_xboole_0(X1)|k1_card_1(esk1_0)=np__4|~m2_finseq_1(esk1_0,X1)|~v1_finseq_6(esk1_0,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_453, c_0_454])])).
cnf(c_0_461, plain, (r1_xxreal_0(esk3_1(X1),k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[c_0_455, c_0_247])).
cnf(c_0_462, negated_conjecture, (k1_card_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_300, c_0_456]), c_0_261]), c_0_262]), c_0_263])])])).
cnf(c_0_463, negated_conjecture, (~r1_xxreal_0(esk3_1(esk1_0),np__0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_457]), c_0_458])])])).
cnf(c_0_464, plain, (esk3_1(X1)=k3_finseq_1(X1)|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(esk3_1(X1))|~v1_xxreal_0(k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),esk3_1(X1))|~v7_ordinal1(esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_252])])).
cnf(c_0_465, negated_conjecture, (esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_459]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_466, negated_conjecture, (k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3)=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)|k1_card_1(esk1_0)=np__4), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_460, c_0_288]), c_0_241])]), c_0_364])])).
cnf(c_0_467, negated_conjecture, (k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__4), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_461, c_0_462]), c_0_336]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_463]), c_0_289]), c_0_290])])).
cnf(c_0_468, plain, (esk3_1(X1)=k3_finseq_1(X1)|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(esk3_1(X1))|~v1_xxreal_0(k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[c_0_464, c_0_231])).
cnf(c_0_469, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4|k1_card_1(esk1_0)=np__4), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_465, c_0_466])])).
cnf(c_0_470, negated_conjecture, (k1_card_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_355, c_0_467]), c_0_357])])])).
cnf(c_0_471, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|k1_card_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|~v1_xxreal_0(k3_finseq_1(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_468, c_0_469]), c_0_394]), c_0_260]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_472, negated_conjecture, (r1_xxreal_0(esk4_1(esk1_0),np__2)|k1_card_1(esk1_0)=np__4), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_380, c_0_470]), c_0_233]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_473, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|k1_card_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__4), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_471, c_0_403]), c_0_261]), c_0_262]), c_0_263])])])).
cnf(c_0_474, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__4), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_472]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_475, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|k1_card_1(esk1_0)=np__4|esk3_1(esk1_0)=np__2), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_473]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_476, plain, (esk4_1(X1)=np__4|r1_xxreal_0(esk4_1(X1),np__3)|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(esk4_1(X1))|~r1_xxreal_0(esk4_1(X1),np__4)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_409, c_0_311])])).
cnf(c_0_477, negated_conjecture, (v1_xxreal_0(esk4_1(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_407]), c_0_208]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_478, negated_conjecture, (k1_card_1(esk1_0)=np__4|esk4_1(esk1_0)=np__2), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_474]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_479, negated_conjecture, (k1_card_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__4|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_475]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_480, plain, (esk4_1(X1)=np__4|r1_xxreal_0(esk4_1(X1),np__3)|v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(esk4_1(X1),np__4)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[c_0_476, c_0_243])).
cnf(c_0_481, negated_conjecture, (r1_xxreal_0(esk4_1(esk1_0),np__4)|~v1_xxreal_0(k3_finseq_1(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_421, c_0_380]), c_0_477]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_482, negated_conjecture, (k1_card_1(esk1_0)=np__4|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_478]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_483, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|k1_card_1(esk1_0)=np__4|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_479, c_0_424]), c_0_449]), c_0_241]), c_0_287])])])).
cnf(c_0_484, plain, (esk4_1(X1)=np__3|r1_xxreal_0(esk4_1(X1),np__2)|v1_xboole_0(X1)|v3_funct_1(X1)|esk4_1(X1)=np__4|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(esk4_1(X1),np__4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_347, c_0_480])])).
cnf(c_0_485, negated_conjecture, (r1_xxreal_0(esk4_1(esk1_0),np__4)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_481, c_0_403]), c_0_261]), c_0_262]), c_0_263])])])).
cnf(c_0_486, negated_conjecture, (k1_card_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__4), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_482, c_0_483])])).
cnf(c_0_487, plain, (esk4_1(X1)=k3_finseq_1(X1)|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(esk4_1(X1))|~v1_xxreal_0(k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),esk4_1(X1))|~v7_ordinal1(esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_302])])).
cnf(c_0_488, negated_conjecture, (esk4_1(esk1_0)=np__4|esk4_1(esk1_0)=np__3|r1_xxreal_0(esk4_1(esk1_0),np__2)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_484, c_0_485]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_489, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|~v1_xxreal_0(k3_finseq_1(esk1_0))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_468, c_0_459]), c_0_257]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_490, negated_conjecture, (k1_card_1(esk1_0)=np__4), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_300, c_0_486]), c_0_261]), c_0_262]), c_0_263])])])).
cnf(c_0_491, plain, (esk4_1(X1)=k3_finseq_1(X1)|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(esk4_1(X1))|~v1_xxreal_0(k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[c_0_487, c_0_243])).
cnf(c_0_492, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__3|esk4_1(esk1_0)=np__4), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_488]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_493, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__3|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_489, c_0_403]), c_0_261]), c_0_262]), c_0_263])])])).
cnf(c_0_494, negated_conjecture, (k7_partfun1(X1,esk1_0,np__4)=k7_partfun1(X1,esk1_0,np__1)|v1_xboole_0(X1)|~m2_finseq_1(esk1_0,X1)|~v1_finseq_6(esk1_0,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_453, c_0_490])])).
cnf(c_0_495, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|esk4_1(esk1_0)=np__4|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|~v1_xxreal_0(k3_finseq_1(esk1_0))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_491, c_0_492]), c_0_257]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_496, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_493]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_497, negated_conjecture, (k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4)=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_494, c_0_288]), c_0_241])]), c_0_364])])).
cnf(c_0_498, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__3|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_495, c_0_403]), c_0_261]), c_0_262]), c_0_263])])])).
cnf(c_0_499, plain, (r1_xxreal_0(np__1,np__3)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r1_r3])).
cnf(c_0_500, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__3|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_496, c_0_497])])])).
cnf(c_0_501, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_498]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_502, plain, (r1_xxreal_0(np__1,np__3)), inference(evalgc,[status(thm)],[c_0_499])).
cnf(c_0_503, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|esk3_1(esk1_0)=np__2|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_500]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_504, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__3|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_501, c_0_497])])])).
cnf(c_0_505, plain, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__4))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__4))|~v7_ordinal1(np__3)|~r1_xxreal_0(np__4,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_412, c_0_393]), c_0_502])])])).
fof(c_0_506, plain, ![X30, X31]:(~m2_subset_1(X30,k1_numbers,k5_numbers)|(v1_xboole_0(X31)|~v2_goboard5(X31)|~m2_finseq_1(X31,u1_struct_0(k15_euclid(np__2)))|(~r2_hidden(X30,k4_finseq_1(X31))|~r2_hidden(k2_nat_1(X30,np__1),k4_finseq_1(X31))|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X31,X30)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X31,k2_nat_1(X30,np__1))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t29_goboard7])])])])).
cnf(c_0_507, plain, (r1_xxreal_0(np__1,np__2)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r1_r2])).
cnf(c_0_508, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_503]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_509, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|esk4_1(esk1_0)=np__2|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_504]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_510, plain, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__4))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__4))|~r1_xxreal_0(np__4,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_505, c_0_395])])).
cnf(c_0_511, plain, (v1_xboole_0(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v2_goboard5(X2)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k4_finseq_1(X2))|~r2_hidden(k2_nat_1(X1,np__1),k4_finseq_1(X2))|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X2,X1)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X2,k2_nat_1(X1,np__1))), inference(split_conjunct,[status(thm)],[c_0_506])).
cnf(c_0_512, plain, (r1_xxreal_0(np__1,np__2)), inference(evalgc,[status(thm)],[c_0_507])).
cnf(c_0_513, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_508, c_0_424]), c_0_449]), c_0_241]), c_0_287])])])).
cnf(c_0_514, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_509]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_515, plain, (k3_finseq_1(X1)=np__4|r1_xxreal_0(k3_finseq_1(X1),np__3)|~v7_ordinal1(k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),np__4)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_409, c_0_403])])).
cnf(c_0_516, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_465, c_0_510]), c_0_241]), c_0_287])])])).
cnf(c_0_517, plain, (v1_xboole_0(X2)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X2,k2_nat_1(X1,np__1))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X2,X1)|~v2_goboard5(X2)|~r2_hidden(X1,k4_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(k2_nat_1(X1,np__1),k4_finseq_1(X2))), inference(evalgc,[status(thm)],[c_0_511])).
fof(c_0_518, plain, ![X102, X103]:(~m1_subset_1(X102,k5_numbers)|~v7_ordinal1(X103)|k2_nat_1(X102,X103)=k2_xcmplx_0(X102,X103)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_nat_1])])).
cnf(c_0_519, plain, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|~r1_xxreal_0(np__3,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_412, c_0_256]), c_0_258]), c_0_512])])])).
cnf(c_0_520, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3)), inference(csr,[status(thm)],[c_0_513, c_0_514])).
cnf(c_0_521, plain, (k3_finseq_1(X1)=np__4|r1_xxreal_0(k3_finseq_1(X1),np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r1_xxreal_0(k3_finseq_1(X1),np__4)), inference(csr,[status(thm)],[c_0_515, c_0_239])).
cnf(c_0_522, negated_conjecture, (esk4_1(esk1_0)=np__4|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_492]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_523, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_516, c_0_497]), c_0_497])])])).
cnf(c_0_524, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~r2_hidden(k2_nat_1(X2,np__1),k9_xtuple_0(X1))|~r2_hidden(X2,k9_xtuple_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_88])])).
cnf(c_0_525, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_518])).
cnf(c_0_526, plain, (r2_hidden(X1,k1_relset_1(k5_numbers,X2))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,k3_finseq_1(X2))|~v7_ordinal1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_527, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(np__3,k3_finseq_1(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_465, c_0_519]), c_0_241]), c_0_287])])])).
cnf(c_0_528, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__4), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_520, c_0_521]), c_0_261]), c_0_262]), c_0_263]), c_0_260])])])).
cnf(c_0_529, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__4|esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_522, c_0_523])])).
cnf(c_0_530, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~r2_hidden(k2_nat_1(X2,np__1),k9_xtuple_0(X1))|~r2_hidden(X2,k9_xtuple_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_524, c_0_103]), c_0_104]), c_0_105])).
cnf(c_0_531, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_525])).
cnf(c_0_532, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_533, plain, (r2_hidden(X1,k1_relset_1(k5_numbers,X2))|~v7_ordinal1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,k3_finseq_1(X2))), inference(evalgc,[status(thm)],[c_0_526])).
cnf(c_0_534, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(np__3,k3_finseq_1(esk1_0))|~r1_xxreal_0(np__2,k3_finseq_1(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_527, c_0_424]), c_0_241]), c_0_287])])])).
cnf(c_0_535, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|np__3!=np__4), inference(evalgc,[status(thm)],[inference(ef,[status(thm)],[c_0_528])])).
cnf(c_0_536, negated_conjecture, (np__3=np__4|k3_finseq_1(esk1_0)=np__4), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_300, c_0_528]), c_0_490]), c_0_261]), c_0_262]), c_0_263])])])).
cnf(c_0_537, plain, (r1_xxreal_0(np__3,np__4)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r3_r4])).
cnf(c_0_538, plain, (r1_xxreal_0(np__4,np__4)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r4_r4])).
cnf(c_0_539, negated_conjecture, (esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_529]), c_0_497]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_540, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~r2_hidden(k2_xcmplx_0(X2,np__1),k9_xtuple_0(X1))|~r2_hidden(X2,k9_xtuple_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_530, c_0_531]), c_0_156])])])).
cnf(c_0_541, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_532])).
cnf(c_0_542, plain, (r2_hidden(X1,k9_xtuple_0(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X1)|~v4_relat_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_533, c_0_101])])).
cnf(c_0_543, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4|~r1_xxreal_0(np__3,k3_finseq_1(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_534, c_0_449])])])).
cnf(c_0_544, negated_conjecture, (k3_finseq_1(esk1_0)=np__4), inference(csr,[status(thm)],[c_0_535, c_0_536])).
cnf(c_0_545, plain, (r1_xxreal_0(np__3,np__4)), inference(evalgc,[status(thm)],[c_0_537])).
cnf(c_0_546, plain, (r1_xxreal_0(np__4,np__4)), inference(evalgc,[status(thm)],[c_0_538])).
cnf(c_0_547, negated_conjecture, (esk4_1(esk1_0)=np__2|esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_539]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_548, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)|~r2_hidden(np__2,k9_xtuple_0(X1))|~r2_hidden(np__1,k9_xtuple_0(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_540, c_0_222]), c_0_130]), c_0_541])])])).
cnf(c_0_549, plain, (r2_hidden(X1,k9_xtuple_0(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X1)), inference(csr,[status(thm)],[c_0_542, c_0_163])).
cnf(c_0_550, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|np__4!=np__1), inference(evalgc,[status(thm)],[inference(ef,[status(thm)],[c_0_456])])).
cnf(c_0_551, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|np__4!=np__2), inference(evalgc,[status(thm)],[inference(ef,[status(thm)],[c_0_456])])).
cnf(c_0_552, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_543, c_0_544]), c_0_545])])])).
cnf(c_0_553, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_523, c_0_544]), c_0_546])])])).
cnf(c_0_554, negated_conjecture, (esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_547]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_555, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|np__4!=np__0), inference(evalgc,[status(thm)],[inference(ef,[status(thm)],[c_0_456])])).
cnf(c_0_556, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)|~r2_hidden(np__2,k9_xtuple_0(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_548, c_0_549]), c_0_156]), c_0_357])])])).
cnf(c_0_557, negated_conjecture, (k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|np__4!=np__1), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_461, c_0_550]), c_0_336]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_463]), c_0_289]), c_0_290])])).
cnf(c_0_558, negated_conjecture, (k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|np__4!=np__2), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_461, c_0_551]), c_0_336]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_463]), c_0_289]), c_0_290])])).
cnf(c_0_559, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_552, c_0_553])])).
cnf(c_0_560, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_554, c_0_544]), c_0_546])])])).
cnf(c_0_561, negated_conjecture, (k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|np__4!=np__0), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_461, c_0_555]), c_0_336]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_463]), c_0_289]), c_0_290])])).
cnf(c_0_562, plain, (r1_xxreal_0(np__2,np__3)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r2_r3])).
cnf(c_0_563, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)|~r2_hidden(np__2,k9_xtuple_0(X1))|~r1_xxreal_0(np__1,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_556, c_0_103]), c_0_104]), c_0_105])).
cnf(c_0_564, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|np__4!=np__1), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_355, c_0_557]), c_0_357])])])).
cnf(c_0_565, plain, (r1_xxreal_0(np__1,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_334, c_0_403])])).
cnf(c_0_566, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|np__4!=np__2), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_355, c_0_558]), c_0_357])])])).
cnf(c_0_567, negated_conjecture, (esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2), inference(csr,[status(thm)],[c_0_559, c_0_560])).
cnf(c_0_568, negated_conjecture, (k7_partfun1(X1,esk1_0,np__4)=k7_partfun1(X1,esk1_0,np__1)|v1_xboole_0(X1)|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|~m2_finseq_1(esk1_0,X1)|~v1_finseq_6(esk1_0,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_456])])).
cnf(c_0_569, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|np__4!=np__0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_355, c_0_561]), c_0_357])])])).
cnf(c_0_570, plain, (r1_xxreal_0(np__2,np__3)), inference(evalgc,[status(thm)],[c_0_562])).
cnf(c_0_571, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)|~r1_xxreal_0(np__1,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r1_xxreal_0(np__2,k3_finseq_1(X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_563, c_0_549]), c_0_258]), c_0_512])])])).
cnf(c_0_572, negated_conjecture, (k7_partfun1(X1,esk1_0,np__1)=k7_partfun1(X1,esk1_0,np__2)|v1_xboole_0(X1)|~m2_finseq_1(esk1_0,X1)|~v1_finseq_6(esk1_0,X1)|np__4!=np__1), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_564])])).
cnf(c_0_573, plain, (r1_xxreal_0(np__1,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_565, c_0_103]), c_0_104]), c_0_105])).
cnf(c_0_574, negated_conjecture, (k7_partfun1(X1,esk1_0,np__1)=k7_partfun1(X1,esk1_0,np__2)|v1_xboole_0(X1)|~m2_finseq_1(esk1_0,X1)|~v1_finseq_6(esk1_0,X1)|np__4!=np__2), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_566])])).
cnf(c_0_575, negated_conjecture, (esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_567]), c_0_497]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_576, plain, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__4))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__4))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|~r1_xxreal_0(np__4,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)|~r1_xxreal_0(np__3,k3_finseq_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_510, c_0_519])])).
cnf(c_0_577, negated_conjecture, (k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4)=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_568, c_0_288]), c_0_241])]), c_0_364])])).
cnf(c_0_578, negated_conjecture, (np__3=np__2|np__4!=np__0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_520, c_0_569]), c_0_570])])])).
cnf(c_0_579, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)|~r1_xxreal_0(np__1,k3_finseq_1(X1))|~r1_xxreal_0(np__2,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_571, c_0_103]), c_0_104]), c_0_105])).
cnf(c_0_580, negated_conjecture, (k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)|np__4!=np__1), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_572, c_0_288]), c_0_241])]), c_0_364])])).
cnf(c_0_581, negated_conjecture, (r1_xxreal_0(np__1,k3_finseq_1(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_573, c_0_288]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287])]), c_0_289]), c_0_290])])).
cnf(c_0_582, negated_conjecture, (k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)|np__4!=np__2), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_574, c_0_288]), c_0_241])]), c_0_364])])).
cnf(c_0_583, negated_conjecture, (esk3_1(esk1_0)=np__2), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_575]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_584, plain, (r1_xxreal_0(np__2,np__4)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r2_r4])).
cnf(c_0_585, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0))|~r1_xxreal_0(np__3,k3_finseq_1(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_576, c_0_577]), c_0_241]), c_0_287])])])).
cnf(c_0_586, negated_conjecture, (np__4!=np__0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_378, c_0_578]), c_0_446])])])).
cnf(c_0_587, negated_conjecture, (np__4!=np__1), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_579, c_0_580]), c_0_581]), c_0_449]), c_0_241]), c_0_284])]), c_0_289])])).
cnf(c_0_588, negated_conjecture, (np__4!=np__2), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_579, c_0_582]), c_0_581]), c_0_449]), c_0_241]), c_0_284])]), c_0_289])])).
cnf(c_0_589, negated_conjecture, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_583]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_590, plain, (r1_xxreal_0(np__2,np__4)), inference(evalgc,[status(thm)],[c_0_584])).
cnf(c_0_591, negated_conjecture, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_585, c_0_544]), c_0_544]), c_0_544]), c_0_544]), c_0_546]), c_0_544]), c_0_545])]), c_0_586]), c_0_587]), c_0_588])])).
cnf(c_0_592, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_589, c_0_424]), c_0_544]), c_0_590]), c_0_241]), c_0_287])])])).
cnf(c_0_593, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[c_0_591, c_0_589])])).
cnf(c_0_594, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__4|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_522, c_0_592])])).
cnf(c_0_595, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_593, c_0_592])])).
cnf(c_0_596, negated_conjecture, (esk4_1(esk1_0)=np__4|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_594, c_0_595])])).
cnf(c_0_597, negated_conjecture, (esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_596]), c_0_497]), c_0_592]), c_0_592]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_598, negated_conjecture, (esk4_1(esk1_0)=np__2), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_597]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290])])).
cnf(c_0_599, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_598]), c_0_592]), c_0_241]), c_0_284]), c_0_285]), c_0_286]), c_0_287]), c_0_288])]), c_0_289]), c_0_290]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 600
# Proof object clause steps            : 483
# Proof object formula steps           : 117
# Proof object conjectures             : 173
# Proof object clause conjectures      : 170
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 88
# Proof object initial formulas used   : 70
# Proof object generating inferences   : 242
# Proof object simplifying inferences  : 1007
# Training examples: 344 positive, 1963 negative
# Training: Positive examples begin
cnf(c_0_113, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_168, plain, (k1_xboole_0=k6_numbers)).# trainpos
cnf(c_0_196, plain, (v1_xboole_0(np__0))).# trainpos
cnf(c_0_287, negated_conjecture, (v1_topreal1(esk1_0))).# trainpos
cnf(c_0_286, negated_conjecture, (v2_topreal1(esk1_0))).# trainpos
cnf(c_0_285, negated_conjecture, (v1_goboard5(esk1_0))).# trainpos
cnf(c_0_284, negated_conjecture, (v2_goboard5(esk1_0))).# trainpos
cnf(c_0_290, negated_conjecture, (~v3_funct_1(esk1_0))).# trainpos
cnf(c_0_289, negated_conjecture, (~v1_xboole_0(esk1_0))).# trainpos
cnf(c_0_304, plain, (v1_xxreal_0(esk8_0))).# trainpos
cnf(c_0_112, plain, (v6_membered(k4_ordinal1))).# trainpos
cnf(c_0_178, plain, (v3_membered(k1_numbers))).# trainpos
cnf(c_0_260, negated_conjecture, (r1_xxreal_0(k3_finseq_1(esk1_0),np__4))).# trainpos
cnf(c_0_288, negated_conjecture, (v1_finseq_6(esk1_0,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_188, plain, (~v1_xboole_0(k1_numbers))).# trainpos
cnf(c_0_446, plain, (r1_xxreal_0(np__2,np__2))).# trainpos
cnf(c_0_590, plain, (r1_xxreal_0(np__2,np__4))).# trainpos
cnf(c_0_161, plain, (~v1_xboole_0(k4_ordinal1))).# trainpos
cnf(c_0_256, plain, (k2_xcmplx_0(np__2,np__1)=np__3)).# trainpos
cnf(c_0_241, negated_conjecture, (m2_finseq_1(esk1_0,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_570, plain, (r1_xxreal_0(np__2,np__3))).# trainpos
cnf(c_0_167, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_541, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_546, plain, (r1_xxreal_0(np__4,np__4))).# trainpos
cnf(c_0_222, plain, (k2_xcmplx_0(np__1,np__1)=np__2)).# trainpos
cnf(c_0_210, plain, (k2_xcmplx_0(np__1,np__0)=np__1)).# trainpos
cnf(c_0_393, plain, (k2_xcmplx_0(np__3,np__1)=np__4)).# trainpos
cnf(c_0_226, plain, (k2_xcmplx_0(np__0,np__1)=np__1)).# trainpos
cnf(c_0_337, plain, (~r1_xxreal_0(np__1,np__0))).# trainpos
cnf(c_0_512, plain, (r1_xxreal_0(np__1,np__2))).# trainpos
cnf(c_0_378, plain, (~r1_xxreal_0(np__3,np__2))).# trainpos
cnf(c_0_357, plain, (r1_xxreal_0(np__1,np__1))).# trainpos
cnf(c_0_502, plain, (r1_xxreal_0(np__1,np__3))).# trainpos
cnf(c_0_545, plain, (r1_xxreal_0(np__3,np__4))).# trainpos
cnf(c_0_191, plain, (m1_subset_1(np__2,k5_numbers))).# trainpos
cnf(c_0_353, plain, (m1_subset_1(np__4,k5_numbers))).# trainpos
cnf(c_0_130, plain, (m1_subset_1(np__1,k5_numbers))).# trainpos
cnf(c_0_216, plain, (m1_subset_1(np__3,k5_numbers))).# trainpos
cnf(c_0_206, plain, (m1_subset_1(np__0,k5_numbers))).# trainpos
cnf(c_0_152, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_128, plain, (~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_nat_1(X2,np__1),X1))).# trainpos
cnf(c_0_182, plain, (v2_membered(X1)|~v3_membered(X1))).# trainpos
cnf(c_0_392, plain, (v1_membered(X1)|~v3_membered(X1))).# trainpos
cnf(c_0_333, plain, (v3_funct_1(X1)|v1_xboole_0(X1)|~v1_topreal1(X1)|~v2_topreal1(X1)|~v1_goboard5(X1)|~v2_goboard5(X1)|~r1_xxreal_0(k3_finseq_1(X1),np__1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_105, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_157, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).# trainpos
cnf(c_0_104, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_303, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_383, plain, (v3_funct_1(X1)|v1_xboole_0(X1)|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,esk4_1(X1)))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1))|~v1_topreal1(X1)|~v2_topreal1(X1)|~v1_goboard5(X1)|~v2_goboard5(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_405, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_324, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).# trainpos
cnf(c_0_251, plain, (r1_xxreal_0(X1,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_361, plain, (v3_funct_1(X1)|v1_xboole_0(X1)|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,esk3_1(X1)))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1))|~v1_topreal1(X1)|~v2_topreal1(X1)|~v1_goboard5(X1)|~v2_goboard5(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_89, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2))).# trainpos
cnf(c_0_90, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_88, plain, (k4_finseq_1(X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_132, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_176, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_389, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_177, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_151, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_372, plain, (v1_membered(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_217, plain, (X1=np__2|X1=np__4|X1=np__1|X1=k6_numbers|X1=np__3|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__4))).# trainpos
cnf(c_0_300, plain, (k3_finseq_1(X1)=k1_card_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_163, plain, (v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_325, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_371, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_101, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).# trainpos
cnf(c_0_129, plain, (k1_nat_1(X1,X2)=k1_nat_1(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_174, plain, (X1=X2|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_84, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_517, plain, (v1_xboole_0(X2)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X2,k2_nat_1(X1,np__1))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X2,X1)|~v2_goboard5(X2)|~r2_hidden(X1,k4_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(k2_nat_1(X1,np__1),k4_finseq_1(X2)))).# trainpos
cnf(c_0_197, plain, (m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_138, plain, (r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)|~r2_hidden(X1,k1_relset_1(k5_numbers,X2)))).# trainpos
cnf(c_0_186, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_155, plain, (k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_531, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_85, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).# trainpos
cnf(c_0_404, plain, (m2_subset_1(k1_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_175, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(k1_nat_1(X1,np__1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_165, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|~v7_ordinal1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)|~r2_hidden(X1,k1_relset_1(k5_numbers,X2)))).# trainpos
cnf(c_0_87, plain, (v3_funct_1(X1)|v1_xboole_0(X1)|r2_hidden(esk3_1(X1),k4_finseq_1(X1))|~v1_topreal1(X1)|~v2_topreal1(X1)|~v1_goboard5(X1)|~v2_goboard5(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_187, plain, (v3_funct_1(X1)|v1_xboole_0(X1)|m2_subset_1(esk3_1(X1),k1_numbers,k5_numbers)|~v1_topreal1(X1)|~v2_topreal1(X1)|~v1_goboard5(X1)|~v2_goboard5(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_533, plain, (r2_hidden(X1,k1_relset_1(k5_numbers,X2))|~v7_ordinal1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,k3_finseq_1(X2)))).# trainpos
cnf(c_0_344, plain, (k7_partfun1(X2,X1,k3_finseq_1(X1))=k7_partfun1(X2,X1,np__1)|v1_xboole_0(X2)|~v1_finseq_6(X1,X2)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_200, plain, (v3_funct_1(X1)|v1_xboole_0(X1)|m2_subset_1(esk4_1(X1),k1_numbers,k5_numbers)|~v1_topreal1(X1)|~v2_topreal1(X1)|~v1_goboard5(X1)|~v2_goboard5(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_144, plain, (v3_funct_1(X1)|v1_xboole_0(X1)|r2_hidden(esk4_1(X1),k4_finseq_1(X1))|~v1_topreal1(X1)|~v2_topreal1(X1)|~v1_goboard5(X1)|~v2_goboard5(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_398, plain, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~v1_topreal1(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(k1_nat_1(X2,np__1),k3_finseq_1(X1)))).# trainpos
cnf(c_0_261, negated_conjecture, (v1_finseq_1(esk1_0))).# trainpos
cnf(c_0_262, negated_conjecture, (v1_funct_1(esk1_0))).# trainpos
cnf(c_0_218, plain, (k6_numbers=np__0)).# trainpos
cnf(c_0_103, plain, (v1_relat_1(X1)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_355, negated_conjecture, (~r1_xxreal_0(k3_finseq_1(esk1_0),np__1))).# trainpos
cnf(c_0_156, plain, (v7_ordinal1(np__1))).# trainpos
cnf(c_0_390, plain, (v4_relat_1(X1,k5_numbers)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_345, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_228, plain, (v7_ordinal1(np__0))).# trainpos
cnf(c_0_258, plain, (v7_ordinal1(np__2))).# trainpos
cnf(c_0_412, plain, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1)))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~v7_ordinal1(X2)|~r1_xxreal_0(k2_xcmplx_0(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).# trainpos
cnf(c_0_542, plain, (r2_hidden(X1,k9_xtuple_0(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X1)|~v4_relat_1(X2,k5_numbers))).# trainpos
cnf(c_0_219, plain, (m1_subset_1(k3_finseq_1(X1),k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_395, plain, (v7_ordinal1(np__3))).# trainpos
cnf(c_0_524, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~r2_hidden(k2_nat_1(X2,np__1),k9_xtuple_0(X1))|~r2_hidden(X2,k9_xtuple_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_154, plain, (~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(k1_nat_1(np__1,X1),X2)|~r1_xxreal_0(X2,X1))).# trainpos
cnf(c_0_320, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_173, plain, (r2_hidden(esk4_1(X1),k9_xtuple_0(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_221, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__1),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_202, plain, (X1=k1_nat_1(X2,np__1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(k1_nat_1(X2,np__1))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k1_nat_1(X2,np__1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_102, plain, (r2_hidden(esk3_1(X1),k9_xtuple_0(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_205, plain, (v3_membered(k5_numbers))).# trainpos
cnf(c_0_396, negated_conjecture, (r1_xxreal_0(X1,np__4)|~v1_xxreal_0(np__4)|~v1_xxreal_0(k3_finseq_1(esk1_0))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k3_finseq_1(esk1_0)))).# trainpos
cnf(c_0_162, plain, (r1_xxreal_0(np__1,X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k9_xtuple_0(X2))|~v7_ordinal1(X1)|~v4_relat_1(X2,k5_numbers))).# trainpos
cnf(c_0_100, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,X3))|~v4_relat_1(X3,X2)|~v1_relat_1(X3))).# trainpos
cnf(c_0_192, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k9_xtuple_0(X2))|~v7_ordinal1(X1)|~v4_relat_1(X2,k5_numbers))).# trainpos
cnf(c_0_317, plain, (v1_xxreal_0(np__0)|~v2_membered(k5_numbers))).# trainpos
cnf(c_0_417, plain, (m1_subset_1(k1_nat_1(X1,X2),k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_444, plain, (k7_partfun1(X1,X2,k1_card_1(X2))=k7_partfun1(X1,X2,np__1)|v1_xboole_0(X1)|~m2_finseq_1(X2,X1)|~v1_finseq_6(X2,X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).# trainpos
cnf(c_0_237, plain, (v1_xxreal_0(np__3)|~v2_membered(k5_numbers))).# trainpos
cnf(c_0_204, plain, (v1_xxreal_0(np__1)|~v2_membered(k5_numbers))).# trainpos
cnf(c_0_375, plain, (v1_xxreal_0(np__4)|~v2_membered(k5_numbers))).# trainpos
cnf(c_0_213, plain, (v1_xxreal_0(np__2)|~v2_membered(k5_numbers))).# trainpos
cnf(c_0_223, plain, (m1_subset_1(esk4_1(X1),k5_numbers)|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_211, plain, (m1_subset_1(esk3_1(X1),k5_numbers)|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_239, plain, (v7_ordinal1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_263, negated_conjecture, (v1_relat_1(esk1_0))).# trainpos
cnf(c_0_394, plain, (v1_xxreal_0(np__4))).# trainpos
cnf(c_0_233, plain, (v1_xxreal_0(np__2))).# trainpos
cnf(c_0_227, plain, (v1_xxreal_0(np__1))).# trainpos
cnf(c_0_235, plain, (r1_xxreal_0(esk3_1(X1),k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_279, plain, (r1_xxreal_0(esk4_1(X1),k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_212, plain, (r1_xxreal_0(np__1,esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_244, plain, (r1_xxreal_0(np__1,esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_418, plain, (k2_xcmplx_0(X1,np__0)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_238, plain, (X1=np__3|X1=np__0|X1=np__1|X1=np__4|X1=np__2|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__4))).# trainpos
cnf(c_0_424, plain, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|~r1_xxreal_0(np__2,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).# trainpos
cnf(c_0_336, plain, (v1_xxreal_0(np__0))).# trainpos
cnf(c_0_231, plain, (v7_ordinal1(esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_243, plain, (v7_ordinal1(esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_364, negated_conjecture, (~v1_xboole_0(u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_257, plain, (v1_xxreal_0(np__3))).# trainpos
cnf(c_0_407, negated_conjecture, (v4_relat_1(esk1_0,k5_numbers))).# trainpos
cnf(c_0_505, plain, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__4))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__4))|~v7_ordinal1(np__3)|~r1_xxreal_0(np__4,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).# trainpos
cnf(c_0_225, plain, (X1=k2_xcmplx_0(X2,np__1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(k2_xcmplx_0(X2,np__1))|~v1_xxreal_0(X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(X1,k2_xcmplx_0(X2,np__1)))).# trainpos
cnf(c_0_519, plain, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|~r1_xxreal_0(np__3,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).# trainpos
cnf(c_0_183, plain, (~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(k2_xcmplx_0(np__1,X1),X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_540, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~r2_hidden(k2_xcmplx_0(X2,np__1),k9_xtuple_0(X1))|~r2_hidden(X2,k9_xtuple_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_386, plain, (v1_xxreal_0(k3_finseq_1(X1))|~v2_membered(k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_429, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_123, plain, (m1_subset_1(X1,X2)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~r2_hidden(X1,k9_xtuple_0(X3)))).# trainpos
cnf(c_0_181, plain, (v2_membered(k5_numbers)|~v2_membered(k1_numbers))).# trainpos
cnf(c_0_242, plain, (r1_xxreal_0(np__2,X1)|r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_297, plain, (r1_xxreal_0(np__3,X1)|r1_xxreal_0(X1,np__2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_293, plain, (r1_xxreal_0(np__1,esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_249, plain, (r1_xxreal_0(np__1,esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_246, plain, (X1=np__1|r1_xxreal_0(X1,np__0)|~v1_xxreal_0(X1)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__1))).# trainpos
cnf(c_0_224, plain, (m1_subset_1(esk4_1(X1),X2)|v1_xboole_0(X1)|v3_funct_1(X1)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_319, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|v1_xboole_0(X2)|v3_funct_1(X2)|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(esk4_1(X2))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk4_1(X2))|~v7_ordinal1(esk4_1(X2))|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_277, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|v1_xboole_0(X2)|v3_funct_1(X2)|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(esk3_1(X2))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk3_1(X2))|~v7_ordinal1(esk3_1(X2))|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_150, plain, (m1_subset_1(esk3_1(X1),X2)|v1_xboole_0(X1)|v3_funct_1(X1)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_487, plain, (esk4_1(X1)=k3_finseq_1(X1)|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(esk4_1(X1))|~v1_xxreal_0(k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),esk4_1(X1))|~v7_ordinal1(esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_410, negated_conjecture, (r1_xxreal_0(X1,np__4)|~v1_xxreal_0(k3_finseq_1(esk1_0))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k3_finseq_1(esk1_0)))).# trainpos
cnf(c_0_464, plain, (esk3_1(X1)=k3_finseq_1(X1)|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(esk3_1(X1))|~v1_xxreal_0(k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),esk3_1(X1))|~v7_ordinal1(esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_576, plain, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__4))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__4))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|~r1_xxreal_0(np__4,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)|~r1_xxreal_0(np__3,k3_finseq_1(X1)))).# trainpos
cnf(c_0_409, plain, (X1=np__4|r1_xxreal_0(X1,np__3)|~v1_xxreal_0(X1)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__4))).# trainpos
cnf(c_0_282, plain, (X1=np__3|r1_xxreal_0(X1,np__2)|~v1_xxreal_0(X1)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__3))).# trainpos
cnf(c_0_548, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)|~r2_hidden(np__2,k9_xtuple_0(X1))|~r2_hidden(np__1,k9_xtuple_0(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1))).# trainpos
cnf(c_0_208, plain, (v2_membered(k5_numbers))).# trainpos
cnf(c_0_391, plain, (v1_membered(k5_numbers)|~v1_membered(k1_numbers))).# trainpos
cnf(c_0_313, plain, (r1_xxreal_0(esk3_1(X1),np__2)|r1_xxreal_0(np__3,esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_230, plain, (~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__0))).# trainpos
cnf(c_0_250, plain, (X1=np__2|r1_xxreal_0(X1,np__1)|~v1_xxreal_0(X1)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__2))).# trainpos
cnf(c_0_259, plain, (k3_finseq_1(X1)=np__2|k3_finseq_1(X1)=np__4|k3_finseq_1(X1)=np__1|k3_finseq_1(X1)=np__0|k3_finseq_1(X1)=np__3|~r1_xxreal_0(k3_finseq_1(X1),np__4)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_271, plain, (r1_xxreal_0(esk4_1(X1),np__1)|r1_xxreal_0(np__2,esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_403, plain, (v1_xxreal_0(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_283, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2)).# trainpos
cnf(c_0_315, plain, (r1_xxreal_0(np__1,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(k3_finseq_1(X1))|~v1_xxreal_0(esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_408, plain, (v1_membered(k5_numbers))).# trainpos
cnf(c_0_273, plain, (v1_xxreal_0(esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v2_membered(X2)|~v4_relat_1(X1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_556, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)|~r2_hidden(np__2,k9_xtuple_0(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(X1)))).# trainpos
cnf(c_0_207, plain, (v1_xxreal_0(esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v2_membered(X2)|~v4_relat_1(X1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_431, plain, (r1_xxreal_0(X1,esk3_1(X2))|v1_xboole_0(X2)|v3_funct_1(X2)|~v1_xxreal_0(esk3_1(X2))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_248, plain, (v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(np__1,esk3_1(X1))|~r1_xxreal_0(esk3_1(X1),np__0)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_292, plain, (esk4_1(X1)=np__1|r1_xxreal_0(np__2,esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(esk4_1(X1))|~r1_xxreal_0(np__1,esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_406, plain, (v1_xcmplx_0(esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_membered(X2)|~v4_relat_1(X1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_422, negated_conjecture, (v1_xxreal_0(esk3_1(esk1_0)))).# trainpos
cnf(c_0_477, negated_conjecture, (v1_xxreal_0(esk4_1(esk1_0)))).# trainpos
cnf(c_0_229, plain, (v1_xxreal_0(esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_294, plain, (v1_xxreal_0(esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_316, negated_conjecture, (k1_card_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0)).# trainpos
cnf(c_0_421, negated_conjecture, (r1_xxreal_0(X1,np__4)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k3_finseq_1(esk1_0)))).# trainpos
cnf(c_0_363, negated_conjecture, (k7_partfun1(X1,esk1_0,np__3)=k7_partfun1(X1,esk1_0,np__1)|v1_xboole_0(X1)|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|~m2_finseq_1(esk1_0,X1)|~v1_finseq_6(esk1_0,X1))).# trainpos
cnf(c_0_306, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),np__3)|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|~v7_ordinal1(esk3_1(esk1_0)))).# trainpos
cnf(c_0_328, negated_conjecture, (r1_xxreal_0(esk4_1(esk1_0),np__3)|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|~v7_ordinal1(esk4_1(esk1_0)))).# trainpos
cnf(c_0_571, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)|~r1_xxreal_0(np__1,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r1_xxreal_0(np__2,k3_finseq_1(X1)))).# trainpos
cnf(c_0_334, plain, (r1_xxreal_0(np__1,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_322, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2|r1_xxreal_0(esk3_1(esk1_0),np__3))).# trainpos
cnf(c_0_335, negated_conjecture, (k1_card_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__3)).# trainpos
cnf(c_0_274, plain, (esk3_1(X1)=np__1|r1_xxreal_0(esk3_1(X1),np__0)|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(esk3_1(X1))|~r1_xxreal_0(esk3_1(X1),np__1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_382, negated_conjecture, (k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3)=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2)).# trainpos
cnf(c_0_419, negated_conjecture, (v1_xcmplx_0(esk3_1(esk1_0)))).# trainpos
cnf(c_0_476, plain, (esk4_1(X1)=np__4|r1_xxreal_0(esk4_1(X1),np__3)|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(esk4_1(X1))|~r1_xxreal_0(esk4_1(X1),np__4)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_441, negated_conjecture, (r1_xxreal_0(X1,esk3_1(esk1_0))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).# trainpos
cnf(c_0_420, plain, (esk3_1(X1)=np__4|r1_xxreal_0(esk3_1(X1),np__3)|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(esk3_1(X1))|~r1_xxreal_0(esk3_1(X1),np__4)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_276, plain, (esk3_1(X1)=np__2|r1_xxreal_0(esk3_1(X1),np__1)|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(esk3_1(X1))|~r1_xxreal_0(esk3_1(X1),np__2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_305, plain, (esk3_1(X1)=np__3|r1_xxreal_0(esk3_1(X1),np__2)|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(esk3_1(X1))|~r1_xxreal_0(esk3_1(X1),np__3)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_327, plain, (esk4_1(X1)=np__3|r1_xxreal_0(esk4_1(X1),np__2)|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(esk4_1(X1))|~r1_xxreal_0(esk4_1(X1),np__3)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_433, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),np__4)|~v7_ordinal1(esk3_1(esk1_0)))).# trainpos
cnf(c_0_348, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2|r1_xxreal_0(esk4_1(esk1_0),np__3))).# trainpos
cnf(c_0_326, plain, (esk4_1(X1)=np__1|r1_xxreal_0(np__2,esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_443, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),np__4))).# trainpos
cnf(c_0_346, plain, (esk4_1(X1)=np__2|esk4_1(X1)=np__1|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(esk4_1(X1))|~r1_xxreal_0(esk4_1(X1),np__2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_356, negated_conjecture, (k1_card_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k1_card_1(esk1_0)=np__4)).# trainpos
cnf(c_0_343, negated_conjecture, (esk3_1(esk1_0)=np__3|r1_xxreal_0(esk3_1(esk1_0),np__2)|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0)).# trainpos
cnf(c_0_484, plain, (esk4_1(X1)=np__3|r1_xxreal_0(esk4_1(X1),np__2)|v1_xboole_0(X1)|v3_funct_1(X1)|esk4_1(X1)=np__4|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(esk4_1(X1),np__4))).# trainpos
cnf(c_0_367, negated_conjecture, (esk4_1(esk1_0)=np__3|r1_xxreal_0(esk4_1(esk1_0),np__2)|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0)).# trainpos
cnf(c_0_312, plain, (esk3_1(X1)=np__1|v1_xboole_0(X1)|v3_funct_1(X1)|esk3_1(X1)=np__2|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(esk3_1(X1),np__2))).# trainpos
cnf(c_0_430, negated_conjecture, (k2_xcmplx_0(esk3_1(esk1_0),np__0)=esk3_1(esk1_0))).# trainpos
cnf(c_0_442, plain, (esk3_1(X1)=np__3|r1_xxreal_0(esk3_1(X1),np__2)|v1_xboole_0(X1)|v3_funct_1(X1)|esk3_1(X1)=np__4|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(esk3_1(X1),np__4))).# trainpos
cnf(c_0_565, plain, (r1_xxreal_0(np__1,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_385, plain, (r1_xxreal_0(np__2,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|esk4_1(X1)=np__1|~v1_xxreal_0(k3_finseq_1(X1))|~v1_xxreal_0(esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_451, negated_conjecture, (r1_xxreal_0(X1,esk3_1(esk1_0))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,np__1))).# trainpos
cnf(c_0_384, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2|esk4_1(esk1_0)=np__3)).# trainpos
cnf(c_0_362, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2|esk3_1(esk1_0)=np__3)).# trainpos
cnf(c_0_440, negated_conjecture, (m1_subset_1(esk3_1(esk1_0),k5_numbers)|~v7_ordinal1(esk3_1(esk1_0)))).# trainpos
cnf(c_0_458, negated_conjecture, (r1_xxreal_0(np__1,esk3_1(esk1_0)))).# trainpos
cnf(c_0_377, negated_conjecture, (k1_card_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__3)).# trainpos
cnf(c_0_581, negated_conjecture, (r1_xxreal_0(np__1,k3_finseq_1(esk1_0)))).# trainpos
cnf(c_0_452, negated_conjecture, (esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__3|r1_xxreal_0(esk3_1(esk1_0),np__2))).# trainpos
cnf(c_0_447, plain, (r1_xxreal_0(esk3_1(X1),esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_341, plain, (r1_xxreal_0(esk4_1(X1),esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_332, plain, (esk3_1(X1)=np__2|esk3_1(X1)=np__1|v1_xboole_0(X1)|v3_funct_1(X1)|r1_xxreal_0(np__3,esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_450, negated_conjecture, (m1_subset_1(esk3_1(esk1_0),k5_numbers))).# trainpos
cnf(c_0_459, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__3|esk3_1(esk1_0)=np__4)).# trainpos
cnf(c_0_401, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)))).# trainpos
cnf(c_0_360, plain, (r1_xxreal_0(esk4_1(X1),k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(k3_finseq_1(X1))|~v1_xxreal_0(esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_455, plain, (r1_xxreal_0(esk3_1(X1),k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(k3_finseq_1(X1))|~v1_xxreal_0(esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_381, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)))).# trainpos
cnf(c_0_354, plain, (r1_xxreal_0(np__3,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|esk3_1(X1)=np__1|esk3_1(X1)=np__2|~v1_xxreal_0(k3_finseq_1(X1))|~v1_xxreal_0(esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_400, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2)).# trainpos
cnf(c_0_415, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2)).# trainpos
cnf(c_0_457, negated_conjecture, (v7_ordinal1(esk3_1(esk1_0)))).# trainpos
cnf(c_0_397, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|k1_card_1(esk1_0)=np__3|k1_card_1(esk1_0)=np__4)).# trainpos
cnf(c_0_465, negated_conjecture, (esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)))).# trainpos
cnf(c_0_399, negated_conjecture, (r1_xxreal_0(esk4_1(esk1_0),np__2)|k1_card_1(esk1_0)=np__3|k1_card_1(esk1_0)=np__4)).# trainpos
cnf(c_0_489, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|~v1_xxreal_0(k3_finseq_1(esk1_0))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).# trainpos
cnf(c_0_481, negated_conjecture, (r1_xxreal_0(esk4_1(esk1_0),np__4)|~v1_xxreal_0(k3_finseq_1(esk1_0)))).# trainpos
cnf(c_0_414, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|esk3_1(esk1_0)=np__2)).# trainpos
cnf(c_0_427, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|esk4_1(esk1_0)=np__2)).# trainpos
cnf(c_0_485, negated_conjecture, (r1_xxreal_0(esk4_1(esk1_0),np__4))).# trainpos
cnf(c_0_463, negated_conjecture, (~r1_xxreal_0(esk3_1(esk1_0),np__0))).# trainpos
cnf(c_0_411, negated_conjecture, (k1_card_1(esk1_0)=np__4|k1_card_1(esk1_0)=np__3|esk3_1(esk1_0)=np__2)).# trainpos
cnf(c_0_413, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__4|k1_card_1(esk1_0)=np__3)).# trainpos
cnf(c_0_493, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__3|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).# trainpos
cnf(c_0_488, negated_conjecture, (esk4_1(esk1_0)=np__4|esk4_1(esk1_0)=np__3|r1_xxreal_0(esk4_1(esk1_0),np__2))).# trainpos
cnf(c_0_426, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)))).# trainpos
cnf(c_0_425, negated_conjecture, (k1_card_1(esk1_0)=np__3|k1_card_1(esk1_0)=np__4|esk4_1(esk1_0)=np__2)).# trainpos
cnf(c_0_496, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).# trainpos
cnf(c_0_438, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)))).# trainpos
cnf(c_0_423, negated_conjecture, (k1_card_1(esk1_0)=np__3|k1_card_1(esk1_0)=np__4|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)))).# trainpos
cnf(c_0_492, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__3|esk4_1(esk1_0)=np__4)).# trainpos
cnf(c_0_437, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(np__2,k3_finseq_1(esk1_0)))).# trainpos
cnf(c_0_435, negated_conjecture, (k1_card_1(esk1_0)=np__4|k1_card_1(esk1_0)=np__3|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)))).# trainpos
cnf(c_0_434, negated_conjecture, (k1_card_1(esk1_0)=np__4|k1_card_1(esk1_0)=np__3|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(np__2,k3_finseq_1(esk1_0)))).# trainpos
cnf(c_0_516, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0)))).# trainpos
cnf(c_0_527, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(np__3,k3_finseq_1(esk1_0)))).# trainpos
cnf(c_0_416, plain, (esk4_1(X1)=np__1|r1_xxreal_0(np__2,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_454, negated_conjecture, (k1_card_1(esk1_0)=np__3|k1_card_1(esk1_0)=np__4)).# trainpos
cnf(c_0_495, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|esk4_1(esk1_0)=np__4|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|~v1_xxreal_0(k3_finseq_1(esk1_0))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).# trainpos
cnf(c_0_534, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(np__3,k3_finseq_1(esk1_0))|~r1_xxreal_0(np__2,k3_finseq_1(esk1_0)))).# trainpos
cnf(c_0_522, negated_conjecture, (esk4_1(esk1_0)=np__4|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)))).# trainpos
cnf(c_0_439, negated_conjecture, (esk4_1(esk1_0)=np__1|r1_xxreal_0(np__2,k3_finseq_1(esk1_0)))).# trainpos
cnf(c_0_498, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__3|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).# trainpos
cnf(c_0_460, negated_conjecture, (k7_partfun1(X1,esk1_0,np__3)=k7_partfun1(X1,esk1_0,np__1)|v1_xboole_0(X1)|k1_card_1(esk1_0)=np__4|~m2_finseq_1(esk1_0,X1)|~v1_finseq_6(esk1_0,X1))).# trainpos
cnf(c_0_449, negated_conjecture, (r1_xxreal_0(np__2,k3_finseq_1(esk1_0)))).# trainpos
cnf(c_0_456, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2)).# trainpos
cnf(c_0_466, negated_conjecture, (k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3)=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)|k1_card_1(esk1_0)=np__4)).# trainpos
cnf(c_0_501, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).# trainpos
cnf(c_0_462, negated_conjecture, (k1_card_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0)).# trainpos
cnf(c_0_568, negated_conjecture, (k7_partfun1(X1,esk1_0,np__4)=k7_partfun1(X1,esk1_0,np__1)|v1_xboole_0(X1)|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|~m2_finseq_1(esk1_0,X1)|~v1_finseq_6(esk1_0,X1))).# trainpos
cnf(c_0_543, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4|~r1_xxreal_0(np__3,k3_finseq_1(esk1_0)))).# trainpos
cnf(c_0_550, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|np__4!=np__1)).# trainpos
cnf(c_0_551, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|np__4!=np__2)).# trainpos
cnf(c_0_555, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|np__4!=np__0)).# trainpos
cnf(c_0_469, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4|k1_card_1(esk1_0)=np__4)).# trainpos
cnf(c_0_467, negated_conjecture, (k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__4)).# trainpos
cnf(c_0_577, negated_conjecture, (k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4)=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2)).# trainpos
cnf(c_0_471, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|k1_card_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|~v1_xxreal_0(k3_finseq_1(esk1_0)))).# trainpos
cnf(c_0_557, negated_conjecture, (k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|np__4!=np__1)).# trainpos
cnf(c_0_561, negated_conjecture, (k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|np__4!=np__0)).# trainpos
cnf(c_0_558, negated_conjecture, (k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|np__4!=np__2)).# trainpos
cnf(c_0_470, negated_conjecture, (k1_card_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2)).# trainpos
cnf(c_0_473, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|k1_card_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__4)).# trainpos
cnf(c_0_569, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|np__4!=np__0)).# trainpos
cnf(c_0_566, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|np__4!=np__2)).# trainpos
cnf(c_0_564, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|np__4!=np__1)).# trainpos
cnf(c_0_475, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|k1_card_1(esk1_0)=np__4|esk3_1(esk1_0)=np__2)).# trainpos
cnf(c_0_472, negated_conjecture, (r1_xxreal_0(esk4_1(esk1_0),np__2)|k1_card_1(esk1_0)=np__4)).# trainpos
cnf(c_0_474, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__4)).# trainpos
cnf(c_0_479, negated_conjecture, (k1_card_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__4|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)))).# trainpos
cnf(c_0_478, negated_conjecture, (k1_card_1(esk1_0)=np__4|esk4_1(esk1_0)=np__2)).# trainpos
cnf(c_0_483, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|k1_card_1(esk1_0)=np__4|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)))).# trainpos
cnf(c_0_572, negated_conjecture, (k7_partfun1(X1,esk1_0,np__1)=k7_partfun1(X1,esk1_0,np__2)|v1_xboole_0(X1)|~m2_finseq_1(esk1_0,X1)|~v1_finseq_6(esk1_0,X1)|np__4!=np__1)).# trainpos
cnf(c_0_580, negated_conjecture, (k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)|np__4!=np__1)).# trainpos
cnf(c_0_482, negated_conjecture, (k1_card_1(esk1_0)=np__4|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)))).# trainpos
cnf(c_0_585, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0))|~r1_xxreal_0(np__3,k3_finseq_1(esk1_0)))).# trainpos
cnf(c_0_486, negated_conjecture, (k1_card_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__4)).# trainpos
cnf(c_0_587, negated_conjecture, (np__4!=np__1)).# trainpos
cnf(c_0_490, negated_conjecture, (k1_card_1(esk1_0)=np__4)).# trainpos
cnf(c_0_515, plain, (k3_finseq_1(X1)=np__4|r1_xxreal_0(k3_finseq_1(X1),np__3)|~v7_ordinal1(k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),np__4)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_494, negated_conjecture, (k7_partfun1(X1,esk1_0,np__4)=k7_partfun1(X1,esk1_0,np__1)|v1_xboole_0(X1)|~m2_finseq_1(esk1_0,X1)|~v1_finseq_6(esk1_0,X1))).# trainpos
cnf(c_0_497, negated_conjecture, (k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4)=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))).# trainpos
cnf(c_0_500, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__3|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).# trainpos
cnf(c_0_523, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0)))).# trainpos
cnf(c_0_504, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__3|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).# trainpos
cnf(c_0_503, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|esk3_1(esk1_0)=np__2|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).# trainpos
cnf(c_0_529, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__4|esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0)))).# trainpos
cnf(c_0_509, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|esk4_1(esk1_0)=np__2|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).# trainpos
cnf(c_0_514, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).# trainpos
cnf(c_0_508, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).# trainpos
cnf(c_0_539, negated_conjecture, (esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0)))).# trainpos
cnf(c_0_574, negated_conjecture, (k7_partfun1(X1,esk1_0,np__1)=k7_partfun1(X1,esk1_0,np__2)|v1_xboole_0(X1)|~m2_finseq_1(esk1_0,X1)|~v1_finseq_6(esk1_0,X1)|np__4!=np__2)).# trainpos
cnf(c_0_582, negated_conjecture, (k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)|np__4!=np__2)).# trainpos
cnf(c_0_513, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).# trainpos
cnf(c_0_547, negated_conjecture, (esk4_1(esk1_0)=np__2|esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0)))).# trainpos
cnf(c_0_528, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__4)).# trainpos
cnf(c_0_578, negated_conjecture, (np__3=np__2|np__4!=np__0)).# trainpos
cnf(c_0_588, negated_conjecture, (np__4!=np__2)).# trainpos
cnf(c_0_554, negated_conjecture, (esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0)))).# trainpos
cnf(c_0_586, negated_conjecture, (np__4!=np__0)).# trainpos
cnf(c_0_536, negated_conjecture, (np__3=np__4|k3_finseq_1(esk1_0)=np__4)).# trainpos
cnf(c_0_535, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|np__3!=np__4)).# trainpos
cnf(c_0_552, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2)).# trainpos
cnf(c_0_553, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4)).# trainpos
cnf(c_0_560, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)))).# trainpos
cnf(c_0_591, negated_conjecture, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)))).# trainpos
cnf(c_0_559, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4)).# trainpos
cnf(c_0_575, negated_conjecture, (esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2)).# trainpos
cnf(c_0_583, negated_conjecture, (esk3_1(esk1_0)=np__2)).# trainpos
cnf(c_0_589, negated_conjecture, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)))).# trainpos
cnf(c_0_592, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)))).# trainpos
cnf(c_0_593, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)))).# trainpos
cnf(c_0_594, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__4|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)))).# trainpos
cnf(c_0_597, negated_conjecture, (esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1)).# trainpos
cnf(c_0_598, negated_conjecture, (esk4_1(esk1_0)=np__2)).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_75, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_278, plain, (v1_xboole_0(esk10_0))).#trainneg
cnf(i_0_228, plain, (v1_topreal1(esk7_0))).#trainneg
cnf(i_0_227, plain, (v2_topreal1(esk7_0))).#trainneg
cnf(i_0_226, plain, (v1_goboard5(esk7_0))).#trainneg
cnf(i_0_225, plain, (v2_goboard5(esk7_0))).#trainneg
cnf(i_0_216, plain, (~v1_xboole_0(np__2))).#trainneg
cnf(i_0_324, plain, (v7_ordinal1(esk13_0))).#trainneg
cnf(i_0_55, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_234, plain, (~v3_funct_1(esk7_0))).#trainneg
cnf(i_0_221, plain, (~v1_xboole_0(np__4))).#trainneg
cnf(i_0_56, plain, (v2_xxreal_0(np__2))).#trainneg
cnf(i_0_79, plain, (v2_xxreal_0(np__4))).#trainneg
cnf(i_0_238, plain, (v1_relat_1(esk7_0))).#trainneg
cnf(i_0_235, plain, (v1_funct_1(esk7_0))).#trainneg
cnf(i_0_231, plain, (v1_finseq_1(esk7_0))).#trainneg
cnf(i_0_48, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_232, plain, (v1_finset_1(esk7_0))).#trainneg
cnf(i_0_100, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_117, plain, (v2_xxreal_0(np__3))).#trainneg
cnf(i_0_341, plain, (v1_card_1(k4_ordinal1))).#trainneg
cnf(i_0_250, plain, (v1_xreal_0(esk8_0))).#trainneg
cnf(i_0_251, plain, (v1_xcmplx_0(esk8_0))).#trainneg
cnf(i_0_249, plain, (v1_int_1(esk8_0))).#trainneg
cnf(i_0_230, plain, (v2_finseq_1(esk7_0))).#trainneg
cnf(i_0_133, plain, (~r1_xxreal_0(np__4,np__2))).#trainneg
cnf(i_0_257, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_78, plain, (m2_subset_1(np__4,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_191, plain, (~v1_xboole_0(np__3))).#trainneg
cnf(i_0_233, plain, (~v1_xboole_0(esk7_0))).#trainneg
cnf(i_0_325, plain, (~v1_xboole_0(esk13_0))).#trainneg
cnf(i_0_155, plain, (~v1_finset_1(k4_ordinal1))).#trainneg
cnf(i_0_89, plain, (v1_xboole_0(esk5_1(X1)))).#trainneg
cnf(i_0_93, plain, (v1_relat_1(esk5_1(X1)))).#trainneg
cnf(i_0_90, plain, (v1_funct_1(esk5_1(X1)))).#trainneg
cnf(i_0_87, plain, (v1_finseq_1(esk5_1(X1)))).#trainneg
cnf(i_0_88, plain, (v1_finset_1(esk5_1(X1)))).#trainneg
cnf(i_0_120, plain, (v1_card_1(k1_card_1(X1)))).#trainneg
cnf(i_0_197, plain, (k2_xcmplx_0(np__2,np__2)=np__4)).#trainneg
cnf(i_0_242, plain, (k2_xcmplx_0(np__2,np__0)=np__2)).#trainneg
cnf(i_0_201, plain, (k2_xcmplx_0(np__1,np__2)=np__3)).#trainneg
cnf(i_0_174, plain, (~r1_xxreal_0(np__2,np__1))).#trainneg
cnf(i_0_179, plain, (~r1_xxreal_0(np__2,np__0))).#trainneg
cnf(i_0_170, plain, (~r1_xxreal_0(np__4,np__1))).#trainneg
cnf(i_0_121, plain, (~r1_xxreal_0(np__4,np__3))).#trainneg
cnf(i_0_152, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_209, plain, (k2_xcmplx_0(np__1,np__3)=np__4)).#trainneg
cnf(i_0_207, plain, (k2_xcmplx_0(np__0,np__2)=np__2)).#trainneg
cnf(i_0_264, plain, (k2_xcmplx_0(np__0,np__0)=np__0)).#trainneg
cnf(i_0_203, plain, (k3_xcmplx_0(np__2,np__2)=np__4)).#trainneg
cnf(i_0_266, plain, (k3_xcmplx_0(np__2,np__1)=np__2)).#trainneg
cnf(i_0_241, plain, (k3_xcmplx_0(np__2,np__0)=np__0)).#trainneg
cnf(i_0_337, plain, (k3_xcmplx_0(np__4,np__1)=np__4)).#trainneg
cnf(i_0_287, plain, (k3_xcmplx_0(np__1,np__2)=np__2)).#trainneg
cnf(i_0_350, plain, (k3_xcmplx_0(np__1,np__1)=np__1)).#trainneg
cnf(i_0_356, plain, (k3_xcmplx_0(np__1,np__0)=np__0)).#trainneg
cnf(i_0_129, plain, (~r1_xxreal_0(np__4,np__0))).#trainneg
cnf(i_0_338, plain, (k3_xcmplx_0(np__3,np__1)=np__3)).#trainneg
cnf(i_0_62, plain, (m2_subset_1(k6_numbers,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_196, plain, (~r1_xxreal_0(np__3,np__1))).#trainneg
cnf(i_0_208, plain, (~r1_xxreal_0(np__3,np__0))).#trainneg
cnf(i_0_305, plain, (v6_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_282, plain, (v3_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_248, plain, (k3_xcmplx_0(np__0,np__2)=np__0)).#trainneg
cnf(i_0_352, plain, (k3_xcmplx_0(np__0,np__1)=np__0)).#trainneg
cnf(i_0_291, plain, (k3_xcmplx_0(np__0,np__0)=np__0)).#trainneg
cnf(i_0_128, plain, (r1_xxreal_0(np__1,np__4))).#trainneg
cnf(i_0_175, plain, (r1_xxreal_0(np__3,np__3))).#trainneg
cnf(i_0_222, plain, (r1_xxreal_0(np__0,np__2))).#trainneg
cnf(i_0_153, plain, (r1_xxreal_0(np__0,np__4))).#trainneg
cnf(i_0_165, plain, (r1_xxreal_0(np__0,np__1))).#trainneg
cnf(i_0_211, plain, (r1_xxreal_0(np__0,np__3))).#trainneg
cnf(i_0_243, plain, (r1_xxreal_0(np__0,np__0))).#trainneg
cnf(i_0_53, plain, (m1_subset_1(np__2,k1_numbers))).#trainneg
cnf(i_0_147, plain, (v1_xboole_0(X1)|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_317, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk12_1(X1)))).#trainneg
cnf(i_0_265, plain, (~v1_xboole_0(k10_finseq_1(X1,X2)))).#trainneg
cnf(i_0_76, plain, (m1_subset_1(np__4,k1_numbers))).#trainneg
cnf(i_0_116, plain, (m2_subset_1(np__3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_97, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_163, plain, (v1_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_286, plain, (v1_card_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_114, plain, (m1_subset_1(np__3,k1_numbers))).#trainneg
cnf(i_0_66, plain, (m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_101, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_253, plain, (m1_subset_1(esk8_0,k1_numbers))).#trainneg
cnf(i_0_237, plain, (v4_relat_1(esk7_0,k5_numbers))).#trainneg
cnf(i_0_127, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_195, plain, (k1_card_1(k1_card_1(X1))=k1_card_1(X1))).#trainneg
cnf(i_0_92, plain, (v4_relat_1(esk5_1(X1),k5_numbers))).#trainneg
cnf(i_0_292, plain, (m2_finseq_1(esk11_1(X1),X1))).#trainneg
cnf(i_0_94, plain, (m1_finseq_1(esk5_1(X1),X1))).#trainneg
cnf(i_0_358, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk16_1(X1)))).#trainneg
cnf(i_0_159, plain, (~v1_xboole_0(X1)|~v3_card_1(X1,np__1))).#trainneg
cnf(i_0_68, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_323, plain, (v3_ordinal1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_288, plain, (v1_xboole_0(X1)|~v3_card_1(X1,k1_xboole_0))).#trainneg
cnf(i_0_229, plain, (v1_finseq_6(esk7_0,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_239, plain, (m1_finseq_1(esk7_0,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_168, plain, (v5_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_112, plain, (v3_ordinal1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_162, plain, (v1_finset_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_149, plain, (v1_finset_1(X1)|~v1_zfmisc_1(X1))).#trainneg
cnf(i_0_166, plain, (m1_subset_1(esk6_1(X1),X1))).#trainneg
cnf(i_0_91, plain, (v5_relat_1(esk5_1(X1),X1))).#trainneg
cnf(i_0_270, plain, (v1_xboole_0(esk9_2(X1,X2)))).#trainneg
cnf(i_0_173, plain, (v1_relat_1(k10_finseq_1(X1,X2)))).#trainneg
cnf(i_0_269, plain, (v1_relat_1(esk9_2(X1,X2)))).#trainneg
cnf(i_0_172, plain, (v1_funct_1(k10_finseq_1(X1,X2)))).#trainneg
cnf(i_0_258, plain, (v1_finseq_1(k10_finseq_1(X1,X2)))).#trainneg
cnf(i_0_210, plain, (v3_membered(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_310, plain, (v1_card_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_256, plain, (v1_xreal_0(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_20, plain, (m2_finseq_1(X1,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_360, plain, (~r2_hidden(X2,X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_236, plain, (v5_relat_1(esk7_0,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_12, plain, (v1_topreal1(X1)|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(esk2_1(X1),np__1)))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,esk2_1(X1)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_193, plain, (l1_rlvect_1(X1)|~l1_rltopsp1(X1))).#trainneg
cnf(i_0_247, plain, (l1_struct_0(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_178, plain, (~v7_ordinal1(X1)|~v2_struct_0(k15_euclid(X1)))).#trainneg
cnf(i_0_107, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_192, plain, (l1_pre_topc(X1)|~l1_rltopsp1(X1))).#trainneg
cnf(i_0_202, plain, (v1_int_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_280, plain, (v4_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_213, plain, (v5_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_272, plain, (l1_struct_0(X1)|~l1_algstr_0(X1))).#trainneg
cnf(i_0_336, plain, (l1_struct_0(X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_274, plain, (l1_algstr_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_345, plain, (l2_algstr_0(X1)|~l1_rlvect_1(X1))).#trainneg
cnf(i_0_275, plain, (l2_struct_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_244, plain, (v3_card_1(k10_finseq_1(X1,X2),np__2))).#trainneg
cnf(i_0_268, plain, (v4_relat_1(esk9_2(X1,X2),X1))).#trainneg
cnf(i_0_267, plain, (v5_relat_1(esk9_2(X1,X2),X2))).#trainneg
cnf(i_0_321, plain, (v1_xboole_0(X1)|v1_relat_1(esk12_1(X1)))).#trainneg
cnf(i_0_318, plain, (v1_xboole_0(X1)|v1_funct_1(esk12_1(X1)))).#trainneg
cnf(i_0_315, plain, (v1_xboole_0(X1)|v1_finseq_1(esk12_1(X1)))).#trainneg
cnf(i_0_316, plain, (v1_xboole_0(X1)|v1_finset_1(esk12_1(X1)))).#trainneg
cnf(i_0_11, plain, (v1_topreal1(X1)|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(esk2_1(X1),np__1)))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,esk2_1(X1)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_357, plain, (v1_xboole_0(X1)|v1_finset_1(esk16_1(X1)))).#trainneg
cnf(i_0_143, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_200, plain, (v1_finset_1(X1)|~v1_finset_1(k1_card_1(X1)))).#trainneg
cnf(i_0_217, plain, (v1_finseq_1(X1)|~v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_314, plain, (v1_xboole_0(X1)|v2_finseq_1(esk12_1(X1)))).#trainneg
cnf(i_0_103, plain, (v1_xboole_0(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_335, plain, (v1_relat_1(esk14_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_333, plain, (v1_funct_1(esk14_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_330, plain, (v1_finseq_1(esk14_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_71, plain, (v1_finset_1(k1_card_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_332, plain, (v1_finset_1(esk14_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_177, plain, (v5_rltopsp1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_320, plain, (v1_xboole_0(X1)|v4_relat_1(esk12_1(X1),k5_numbers))).#trainneg
cnf(i_0_169, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_131, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_158, plain, (v1_zfmisc_1(X1)|~v3_card_1(X1,np__1))).#trainneg
cnf(i_0_255, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_342, plain, (v7_struct_0(X1)|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_190, plain, (v2_pre_topc(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_189, plain, (v13_algstr_0(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_188, plain, (v2_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_187, plain, (v3_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_186, plain, (v4_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_185, plain, (v5_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_184, plain, (v6_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_183, plain, (v7_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_182, plain, (v8_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_205, plain, (l1_rltopsp1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_329, plain, (v2_finseq_1(esk14_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_284, plain, (v5_finset_1(k1_zfmisc_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_281, plain, (k3_xcmplx_0(X1,k6_numbers)=k6_numbers|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_334, plain, (v4_relat_1(esk14_1(X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_51, plain, (v7_ordinal1(X1)|~v3_ordinal1(X1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_132, plain, (m1_subset_1(k17_euclid(X1),k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_135, plain, (m1_subset_1(k18_euclid(X1),k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_223, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_260, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_262, plain, (v1_xboole_0(X1)|v3_card_1(X1,np__1)|~v1_zfmisc_1(X1))).#trainneg
cnf(i_0_44, plain, (v1_funct_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_43, plain, (v1_finseq_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_322, plain, (v1_xboole_0(X1)|m1_finseq_1(esk12_1(X1),X1))).#trainneg
cnf(i_0_319, plain, (v1_xboole_0(X1)|v5_relat_1(esk12_1(X1),X1))).#trainneg
cnf(i_0_331, plain, (v3_card_1(esk14_1(X1),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_118, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_124, plain, (v5_relat_1(X1,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_359, plain, (v1_xboole_0(X1)|m1_subset_1(esk16_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_156, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_157, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_271, plain, (m1_subset_1(esk9_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_145, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_83, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_15, plain, (v1_topreal1(X1)|v7_ordinal1(esk2_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_34, plain, (v1_xboole_0(X1)|v1_finseq_6(X2,X1)|k7_partfun1(X1,X2,k3_finseq_1(X2))!=k7_partfun1(X1,X2,np__1)|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_254, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_308, plain, (v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v2_finseq_1(X1))).#trainneg
cnf(i_0_285, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v1_zfmisc_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_161, plain, (v1_finset_1(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_353, plain, (v2_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_113, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_80, plain, (v1_finset_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_142, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_289, plain, (v1_finset_1(X2)|~v5_finset_1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_96, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_340, plain, (v1_int_1(X2)|~v5_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_57, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_312, plain, (v6_membered(X2)|~v6_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_215, plain, (v1_finset_1(X2)|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_14, plain, (v1_topreal1(X1)|r1_xxreal_0(np__1,esk2_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_167, plain, (v3_ordinal1(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_ordinal1(X1))).#trainneg
cnf(i_0_327, plain, (v1_finseq_1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1))))).#trainneg
cnf(i_0_290, plain, (v3_valued_0(X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1))))).#trainneg
cnf(i_0_279, plain, (k3_finseq_1(k3_finseq_1(X1))=k3_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_276, plain, (v2_xxreal_0(X2)|~v2_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_151, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_245, plain, (k19_euclid(X1,X2)=k10_finseq_1(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_346, plain, (v1_finset_1(k2_zfmisc_1(X1,X2))|~v1_finset_1(X2)|~v1_finset_1(X1))).#trainneg
cnf(i_0_263, plain, (v1_int_1(k2_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_328, plain, (v1_int_1(k3_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_313, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_125, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_204, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_65, plain, (k19_euclid(k17_euclid(X1),k18_euclid(X1))=X1|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_123, plain, (k2_nat_1(X1,X2)=k2_nat_1(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_130, plain, (m1_subset_1(k4_finseq_1(X1),k1_zfmisc_1(k5_numbers))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_351, plain, (v1_relat_1(k9_xtuple_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4))))).#trainneg
cnf(i_0_72, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_148, plain, (m1_subset_1(k19_euclid(X1,X2),u1_struct_0(k15_euclid(np__2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_224, plain, (m2_subset_1(k2_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_13, plain, (v1_topreal1(X1)|r1_xxreal_0(k1_nat_1(esk2_1(X1),np__1),k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_154, plain, (m1_subset_1(k7_partfun1(X2,X1,X3),X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_relat_1(X1,X2))).#trainneg
cnf(i_0_273, plain, (v3_card_1(k9_xtuple_0(X2),X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_card_1(X1)|~v3_card_1(X2,X1))).#trainneg
cnf(i_0_339, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m2_subset_1(esk15_2(X1,X2),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_261, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_311, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_105, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_347, plain, (v1_finset_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X2)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_283, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,X3),k3_xcmplx_0(X2,X3))=k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_361, plain, (esk10_0=k6_numbers)).#trainneg
cnf(i_0_362, plain, (esk5_1(X1)=k6_numbers)).#trainneg
cnf(i_0_940, plain, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,X2)))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,X2)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~v7_ordinal1(X2)|~r1_xxreal_0(k1_nat_1(np__1,X2),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_486, plain, (v7_ordinal1(np__4))).#trainneg
cnf(i_0_371, plain, (m2_finseq_1(esk7_0,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_368, plain, (~v7_ordinal1(k5_numbers))).#trainneg
cnf(i_0_608, plain, (~v1_finset_1(k1_numbers))).#trainneg
cnf(i_0_522, plain, (v3_ordinal1(np__4))).#trainneg
cnf(i_0_373, plain, (l1_struct_0(X1)|~l1_rltopsp1(X1))).#trainneg
cnf(i_0_374, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_713, plain, (v5_relat_1(X1,X2)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_938, plain, (k19_euclid(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1))),k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)))=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k1_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_512, plain, (v3_ordinal1(np__0))).#trainneg
cnf(i_0_514, plain, (v3_ordinal1(np__3))).#trainneg
cnf(i_0_379, plain, (~r1_xxreal_0(k3_finseq_1(esk7_0),np__1)|~m2_finseq_1(esk7_0,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_375, plain, (v3_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_403, plain, (~m1_subset_1(k5_numbers,k5_numbers))).#trainneg
cnf(i_0_535, plain, (~v5_finset_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_524, plain, (v3_ordinal1(np__2))).#trainneg
cnf(i_0_559, plain, (v1_xreal_0(np__4))).#trainneg
cnf(i_0_376, plain, (l1_struct_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_377, plain, (l2_algstr_0(X1)|~l1_rltopsp1(X1))).#trainneg
cnf(i_0_776, plain, (~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(k2_xcmplx_0(X1,np__1),X2)|~r1_xxreal_0(X2,X1))).#trainneg
cnf(i_0_756, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(np__1,X2))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~r2_hidden(k2_nat_1(np__1,X2),k4_finseq_1(X1))|~r2_hidden(X2,k4_finseq_1(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_842, plain, (m1_subset_1(k7_partfun1(X1,X2,X3),X1)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_731, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|~v1_xxreal_0(k3_finseq_1(esk1_0))|~v1_xxreal_0(np__4)|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_627, negated_conjecture, (r1_xxreal_0(k1_card_1(esk1_0),np__4)|~v1_finseq_1(esk1_0)|~v1_funct_1(esk1_0)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_757, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(np__1,X2))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~r2_hidden(k2_nat_1(np__1,X2),k4_finseq_1(X1))|~r2_hidden(X2,k4_finseq_1(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_520, plain, (v3_ordinal1(np__1))).#trainneg
cnf(i_0_549, plain, (v1_xreal_0(np__0))).#trainneg
cnf(i_0_442, plain, (~v1_xboole_0(esk14_1(np__1))|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_551, plain, (v1_xreal_0(np__3))).#trainneg
cnf(i_0_557, plain, (v1_xreal_0(np__1))).#trainneg
cnf(i_0_561, plain, (v1_xreal_0(np__2))).#trainneg
cnf(i_0_664, plain, (v1_int_1(np__4)|~v1_int_1(np__2))).#trainneg
cnf(i_0_381, plain, (v1_finseq_1(esk11_1(X1)))).#trainneg
cnf(i_0_386, plain, (v1_finset_1(X1)|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_427, plain, (v7_ordinal1(X1)|~v1_finset_1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_777, plain, (k1_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_779, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~r2_hidden(k2_xcmplx_0(X2,np__1),k4_finseq_1(X1))|~r2_hidden(X2,k4_finseq_1(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_903, plain, (k7_partfun1(X1,k3_finseq_1(X2),k3_finseq_1(X2))=k7_partfun1(X1,k3_finseq_1(X2),np__1)|v1_xboole_0(X1)|~m2_finseq_1(k3_finseq_1(X2),X1)|~v1_finseq_6(k3_finseq_1(X2),X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_780, plain, (k2_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_826, plain, (r1_xxreal_0(k1_nat_1(np__1,X1),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_747, plain, (v6_membered(k4_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_400, plain, (X1=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_393, plain, (v1_funct_1(esk11_1(X1)))).#trainneg
cnf(i_0_510, plain, (v3_ordinal1(esk8_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_702, plain, (~r2_hidden(X3,X1)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_387, plain, (v1_finset_1(X1)|~v7_ordinal1(k1_card_1(X1)))).#trainneg
cnf(i_0_474, plain, (v7_ordinal1(esk8_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_397, plain, (X1=esk9_2(X2,X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_399, plain, (X1=k1_card_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_885, plain, (v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xboole_0(k4_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_824, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(k1_nat_1(X3,np__1))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k1_nat_1(X3,np__1))|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_898, plain, (~r2_hidden(k1_relset_1(k5_numbers,X2),X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_818, plain, (m1_subset_1(k2_nat_1(X1,X2),k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_411, plain, (r1_xxreal_0(X1,np__4)|~v1_xxreal_0(np__4)|~v1_xxreal_0(np__0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__0))).#trainneg
cnf(i_0_896, plain, (~v1_xboole_0(k1_relset_1(k5_numbers,X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_402, plain, (v1_finset_1(X1)|~m1_subset_1(k1_card_1(X1),k5_numbers))).#trainneg
cnf(i_0_495, plain, (v1_xxreal_0(np__0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_422, plain, (r1_xxreal_0(X1,np__4)|~v1_xxreal_0(np__4)|~v1_xxreal_0(np__2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__2))).#trainneg
cnf(i_0_404, plain, (v1_finset_1(X1)|~v3_card_1(X1,np__1))).#trainneg
cnf(i_0_407, plain, (v1_xboole_0(esk6_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_417, plain, (r1_xxreal_0(X1,np__4)|~v1_xxreal_0(np__4)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_815, plain, (m2_subset_1(k1_nat_1(X2,X1),k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_758, plain, (v1_xboole_0(X1)|v1_xboole_0(k4_finseq_1(X1))|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~r2_hidden(X2,k4_finseq_1(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~m1_subset_1(k2_nat_1(X2,np__1),k4_finseq_1(X1)))).#trainneg
cnf(i_0_880, plain, (m1_subset_1(k3_finseq_1(X1),k1_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_415, plain, (r1_xxreal_0(X1,np__3)|~v1_xxreal_0(np__3)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_418, plain, (r1_xxreal_0(X1,np__2)|~v1_xxreal_0(np__2)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_414, plain, (r1_xxreal_0(X1,np__4)|~v1_xxreal_0(np__4)|~v1_xxreal_0(np__3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__3))).#trainneg
cnf(i_0_497, plain, (v1_xxreal_0(np__3)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_844, plain, (~v1_xboole_0(k9_xtuple_0(X1))|~v3_card_1(X1,np__1)|~v1_card_1(np__1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_409, plain, (r1_xxreal_0(X1,np__3)|~v1_xxreal_0(np__3)|~v1_xxreal_0(np__0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__0))).#trainneg
cnf(i_0_503, plain, (v1_xxreal_0(np__1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_785, plain, (v1_finset_1(k19_euclid(X1,X2))|~v5_finset_1(u1_struct_0(k15_euclid(np__2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_456, plain, (r1_xxreal_0(X1,esk8_0)|r1_xxreal_0(esk8_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_410, plain, (r1_xxreal_0(X1,np__1)|~v1_xxreal_0(np__1)|~v1_xxreal_0(np__0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__0))).#trainneg
cnf(i_0_412, plain, (r1_xxreal_0(X1,np__2)|~v1_xxreal_0(np__2)|~v1_xxreal_0(np__0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__0))).#trainneg
cnf(i_0_430, plain, (~v7_ordinal1(X1)|~l1_struct_0(k15_euclid(X1))|~v1_xboole_0(u1_struct_0(k15_euclid(X1))))).#trainneg
cnf(i_0_819, plain, (m2_subset_1(k2_nat_1(X2,X1),k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_483, plain, (v7_ordinal1(k17_euclid(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_837, plain, (r1_xxreal_0(k1_nat_1(np__1,esk2_1(X1)),k3_finseq_1(X1))|v1_topreal1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk2_1(X1),k5_numbers)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_482, plain, (v7_ordinal1(k18_euclid(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_421, plain, (r1_xxreal_0(X1,np__3)|~v1_xxreal_0(np__3)|~v1_xxreal_0(np__2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__2))).#trainneg
cnf(i_0_835, plain, (r1_xxreal_0(k1_nat_1(np__1,esk2_1(X1)),k3_finseq_1(X1))|v1_topreal1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(esk2_1(X1)))).#trainneg
cnf(i_0_753, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k4_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_555, plain, (v1_xreal_0(k17_euclid(X1))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_505, plain, (v1_xxreal_0(np__4)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_719, plain, (v1_topreal1(X1)|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,esk2_1(X1))))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,esk2_1(X1)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(esk2_1(X1)))).#trainneg
cnf(i_0_431, plain, (k3_xcmplx_0(np__1,esk8_0)=esk8_0)).#trainneg
cnf(i_0_507, plain, (v1_xxreal_0(np__2)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_739, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|~r2_hidden(X1,X3)|~m2_finseq_1(X3,X2))).#trainneg
cnf(i_0_440, plain, (m2_finseq_1(esk12_1(X1),X1)|v1_xboole_0(X1))).#trainneg
cnf(i_0_447, plain, (v1_xboole_0(esk16_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_449, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_717, plain, (v1_topreal1(X1)|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,esk2_1(X1))))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,esk2_1(X1)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk2_1(X1),k5_numbers)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_833, plain, (r1_xxreal_0(k1_nat_1(esk2_1(k3_finseq_1(X1)),np__1),k3_finseq_1(X1))|v1_topreal1(k3_finseq_1(X1))|~m2_finseq_1(k3_finseq_1(X1),u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_554, plain, (v1_xreal_0(k18_euclid(X1))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_806, plain, (v6_membered(k1_relset_1(X1,X2))|~v6_membered(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_830, plain, (k1_nat_1(esk2_1(X1),np__1)=k3_finseq_1(X1)|v1_topreal1(X1)|~v1_xxreal_0(k1_nat_1(esk2_1(X1),np__1))|~v1_xxreal_0(k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),k1_nat_1(esk2_1(X1),np__1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_816, plain, (m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_897, plain, (m1_subset_1(X1,k1_relset_1(k5_numbers,X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_470, plain, (v1_finset_1(k10_finseq_1(X1,X2)))).#trainneg
cnf(i_0_720, plain, (v1_topreal1(X1)|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,esk2_1(X1))))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,esk2_1(X1)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(esk2_1(X1)))).#trainneg
cnf(i_0_453, plain, (v1_xboole_0(esk16_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_528, plain, (v1_finset_1(esk8_0)|~v5_finset_1(k1_numbers))).#trainneg
cnf(i_0_530, plain, (v1_finset_1(np__0)|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_475, plain, (v7_ordinal1(esk6_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_493, plain, (v1_xxreal_0(esk6_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_455, plain, (v1_xboole_0(esk6_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_715, plain, (k10_finseq_1(k17_euclid(X1),k18_euclid(X1))=X1|~v1_xreal_0(k18_euclid(X1))|~v1_xreal_0(k17_euclid(X1))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_754, plain, (m1_subset_1(k9_xtuple_0(X1),k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_832, plain, (v1_topreal1(X1)|~v7_ordinal1(esk2_1(X1))|~v7_ordinal1(k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),esk2_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_789, plain, (v1_finseq_1(k19_euclid(X1,X2))|~v7_ordinal1(np__2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_760, plain, (m2_subset_1(k1_card_1(X1),k1_numbers,k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_631, plain, (k1_card_1(k3_finseq_1(X1))=k3_finseq_1(X1)|~v1_finseq_1(k3_finseq_1(X1))|~v1_funct_1(k3_finseq_1(X1))|~v1_relat_1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_937, plain, (m1_subset_1(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)),k1_numbers)|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k1_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_820, plain, (m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_531, plain, (v1_finset_1(np__0)|~v5_finset_1(k1_numbers))).#trainneg
cnf(i_0_511, plain, (v3_ordinal1(esk6_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_532, plain, (v1_finset_1(np__3)|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_774, plain, (v1_topreal1(X1)|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(esk2_1(X1),np__1)))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,esk2_1(X1)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(esk2_1(X1)))).#trainneg
cnf(i_0_533, plain, (v1_finset_1(np__3)|~v5_finset_1(k1_numbers))).#trainneg
cnf(i_0_538, plain, (v1_finset_1(np__1)|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_539, plain, (v1_finset_1(np__1)|~v5_finset_1(k1_numbers))).#trainneg
cnf(i_0_529, plain, (v1_finset_1(esk6_1(X1))|~v5_finset_1(X1))).#trainneg
cnf(i_0_547, plain, (v1_xreal_0(esk6_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_461, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v3_card_1(u1_struct_0(X1),np__1))).#trainneg
cnf(i_0_633, plain, (v1_finseq_6(k3_finseq_1(X1),X2)|v1_xboole_0(X2)|k7_partfun1(X2,k3_finseq_1(X1),k3_finseq_1(X1))!=k7_partfun1(X2,k3_finseq_1(X1),np__1)|~m2_finseq_1(k3_finseq_1(X1),X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_884, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|v1_xboole_0(k1_relset_1(k5_numbers,X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_relset_1(k5_numbers,X2)))).#trainneg
cnf(i_0_873, plain, (m1_subset_1(k2_nat_1(X1,X2),k1_numbers)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_877, plain, (m1_subset_1(k1_nat_1(X1,X2),k1_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_836, plain, (r1_xxreal_0(k2_xcmplx_0(esk2_1(X1),np__1),k3_finseq_1(X1))|v1_topreal1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(esk2_1(X1)))).#trainneg
cnf(i_0_882, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_925, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)),k3_xcmplx_0(X4,X3))=k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),X4),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_975, plain, (m1_subset_1(k7_partfun1(X1,X2,X3),X1)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_1361, plain, (v1_relat_1(X1)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_288, plain, (v1_xboole_0(X1)|~v3_card_1(X1,np__0))).#trainneg
cnf(i_0_88, plain, (v1_finset_1(k6_numbers))).#trainneg
cnf(i_0_361, plain, (esk10_0=np__0)).#trainneg
cnf(i_0_52, plain, (k1_xboole_0=np__0)).#trainneg
cnf(i_0_472, plain, (v7_ordinal1(esk16_1(X1))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_473, plain, (v7_ordinal1(esk9_2(X1,X2))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1362, plain, (v1_funct_1(X1)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_87, plain, (v1_finseq_1(k6_numbers))).#trainneg
cnf(i_0_90, plain, (v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_93, plain, (v1_relat_1(k6_numbers))).#trainneg
cnf(i_0_281, plain, (k3_xcmplx_0(X1,np__0)=np__0|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_481, plain, (v7_ordinal1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_92, plain, (v4_relat_1(k6_numbers,k5_numbers))).#trainneg
cnf(i_0_964, plain, (~v7_ordinal1(k1_numbers))).#trainneg
cnf(i_0_1010, plain, (v3_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_490, plain, (v1_xxreal_0(esk16_1(X1))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_491, plain, (v1_xxreal_0(esk9_2(X1,X2))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1363, plain, (v1_finseq_1(X1)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_946, negated_conjecture, (~r1_xxreal_0(k1_card_1(esk1_0),np__1)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_1006, plain, (r1_xxreal_0(esk8_0,esk8_0))).#trainneg
cnf(i_0_362, plain, (esk5_1(X1)=np__0)).#trainneg
cnf(i_0_499, plain, (v1_xxreal_0(k5_numbers)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_967, plain, (l1_struct_0(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_500, plain, (v1_xxreal_0(k18_euclid(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_501, plain, (v1_xxreal_0(k17_euclid(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1081, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__0)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1)|~r2_hidden(np__1,k4_finseq_1(X1))|~r2_hidden(np__0,k4_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1))).#trainneg
cnf(i_0_1083, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)|~r2_hidden(np__2,k4_finseq_1(X1))|~r2_hidden(np__1,k4_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1))).#trainneg
cnf(i_0_954, plain, (v1_xboole_0(esk11_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1163, plain, (v1_finset_1(esk14_1(np__1)))).#trainneg
cnf(i_0_91, plain, (v5_relat_1(k6_numbers,X1))).#trainneg
cnf(i_0_508, plain, (v3_ordinal1(esk16_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_540, plain, (v1_finset_1(np__4)|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_965, plain, (~m1_subset_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1135, plain, (~v1_xboole_0(k4_finseq_1(esk7_0)))).#trainneg
cnf(i_0_943, negated_conjecture, (v1_finset_1(esk1_0)|~v1_funct_1(esk1_0)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_509, plain, (v3_ordinal1(esk9_2(X1,X2))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_518, plain, (v3_ordinal1(k18_euclid(X1))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1014, plain, (v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xboole_0(k9_xtuple_0(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1365, plain, (~v1_xboole_0(X1)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1195, plain, (v1_xreal_0(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_541, plain, (v1_finset_1(np__4)|~v5_finset_1(k1_numbers))).#trainneg
cnf(i_0_94, plain, (m1_finseq_1(k6_numbers,X1))).#trainneg
cnf(i_0_542, plain, (v1_finset_1(np__2)|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_1136, negated_conjecture, (~v1_xboole_0(k4_finseq_1(esk1_0)))).#trainneg
cnf(i_0_519, plain, (v3_ordinal1(k17_euclid(X1))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_543, plain, (v1_finset_1(np__2)|~v5_finset_1(k1_numbers))).#trainneg
cnf(i_0_1009, plain, (~r1_xxreal_0(k1_card_1(esk7_0),np__1))).#trainneg
cnf(i_0_566, plain, (v1_xcmplx_0(np__0)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_1364, plain, (v3_card_1(X1,np__2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1366, plain, (v1_finset_1(X1)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_527, plain, (v1_finset_1(esk9_2(X1,X2))|~v5_finset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_950, plain, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(np__1,X2)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(np__1,X2)))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~v7_ordinal1(X2)|~r1_xxreal_0(k2_xcmplx_0(np__1,X2),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1076, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1021, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1110, plain, (v6_membered(k9_xtuple_0(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1192, plain, (v1_xxreal_0(k3_finseq_1(X1))|~v2_membered(k1_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1102, plain, (X1=k1_nat_1(np__1,X2)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(k1_nat_1(np__1,X2))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k1_nat_1(np__1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_567, plain, (v1_xcmplx_0(np__0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_944, plain, (v1_relat_1(esk11_1(X1)))).#trainneg
cnf(i_0_568, plain, (v1_xcmplx_0(np__3)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_976, plain, (~v1_xboole_0(k9_xtuple_0(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_536, plain, (v1_finset_1(k18_euclid(X1))|~v5_finset_1(k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_569, plain, (v1_xcmplx_0(np__3)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_574, plain, (v1_xcmplx_0(np__1)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_575, plain, (v1_xcmplx_0(np__1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1398, plain, (v1_card_1(k3_finseq_1(X1))|~v1_finseq_1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(k3_finseq_1(X1))|~v1_funct_1(X1)|~v1_relat_1(k3_finseq_1(X1))|~v1_relat_1(X1))).#trainneg
cnf(i_0_537, plain, (v1_finset_1(k17_euclid(X1))|~v5_finset_1(k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_544, plain, (v1_xreal_0(esk16_1(X1))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1384, plain, (v1_topreal1(k3_finseq_1(X1))|~v7_ordinal1(k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),esk2_1(k3_finseq_1(X1)))|~m2_finseq_1(k3_finseq_1(X1),u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1180, plain, (v7_ordinal1(k1_nat_1(X1,X2))|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1175, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~m1_subset_1(k2_nat_1(X2,np__1),k9_xtuple_0(X1))|~r2_hidden(X2,k9_xtuple_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_997, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(np__1,X2))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~r2_hidden(k2_nat_1(np__1,X2),k9_xtuple_0(X1))|~r2_hidden(X2,k9_xtuple_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1082, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__4)|~r2_hidden(np__4,k4_finseq_1(X1))|~r2_hidden(np__3,k4_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1))).#trainneg
cnf(i_0_576, plain, (v1_xcmplx_0(np__4)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_577, plain, (v1_xcmplx_0(np__4)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_545, plain, (v1_xreal_0(esk9_2(X1,X2))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_578, plain, (v1_xcmplx_0(np__2)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_579, plain, (v1_xcmplx_0(np__2)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_584, plain, (v1_int_1(np__0)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1261, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_553, plain, (v1_xreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_562, plain, (v1_xcmplx_0(esk16_1(X1))|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1053, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(np__1,X2))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~r2_hidden(k2_xcmplx_0(np__1,X2),k4_finseq_1(X1))|~r2_hidden(X2,k4_finseq_1(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1))).#trainneg
cnf(i_0_999, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(np__1,X2))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~r2_hidden(k2_nat_1(np__1,X2),k9_xtuple_0(X1))|~r2_hidden(X2,k9_xtuple_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1156, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1367, plain, (k10_finseq_1(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1))),k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)))=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k1_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_585, plain, (v1_int_1(np__0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_586, plain, (v1_int_1(np__3)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1146, plain, (~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(k1_relset_1(k5_numbers,X2))|~r1_xxreal_0(k1_relset_1(k5_numbers,X2),k3_finseq_1(X1))|~r1_xxreal_0(np__1,k1_relset_1(k5_numbers,X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(k1_relset_1(k5_numbers,X1))|~r1_xxreal_0(k1_relset_1(k5_numbers,X1),k3_finseq_1(X2))|~r1_xxreal_0(np__1,k1_relset_1(k5_numbers,X1)))).#trainneg
cnf(i_0_563, plain, (v1_xcmplx_0(esk9_2(X1,X2))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_587, plain, (v1_int_1(np__3)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_592, plain, (v1_int_1(np__1)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_593, plain, (v1_int_1(np__1)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1345, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1022, plain, (r1_xxreal_0(np__4,X1)|r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_565, plain, (v1_xcmplx_0(esk6_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_571, plain, (v1_xcmplx_0(k5_numbers)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1084, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)|~r2_hidden(np__3,k4_finseq_1(X1))|~r2_hidden(np__2,k4_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1))).#trainneg
cnf(i_0_1095, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(np__1,X2))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~r2_hidden(k2_xcmplx_0(np__1,X2),k9_xtuple_0(X1))|~r2_hidden(X2,k9_xtuple_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1027, plain, (~v7_ordinal1(X1)|~r1_xxreal_0(np__4,X1)|~r1_xxreal_0(X1,np__3))).#trainneg
cnf(i_0_1169, plain, (m1_subset_1(k1_nat_1(X1,X2),k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1176, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(np__1,X2))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~m1_subset_1(k2_nat_1(np__1,X2),k4_finseq_1(X1))|~r2_hidden(X2,k4_finseq_1(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_594, plain, (v1_int_1(np__4)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_595, plain, (v1_int_1(np__4)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1145, plain, (~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(k9_xtuple_0(X2))|~r1_xxreal_0(k9_xtuple_0(X2),k3_finseq_1(X1))|~r1_xxreal_0(np__1,k9_xtuple_0(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(k1_relset_1(k5_numbers,X1))|~r1_xxreal_0(k1_relset_1(k5_numbers,X1),k3_finseq_1(X2))|~r1_xxreal_0(np__1,k1_relset_1(k5_numbers,X1)))).#trainneg
cnf(i_0_572, plain, (v1_xcmplx_0(k18_euclid(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_596, plain, (v1_int_1(np__2)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_597, plain, (v1_int_1(np__2)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_573, plain, (v1_xcmplx_0(k17_euclid(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_580, plain, (v1_int_1(esk16_1(X1))|v1_xboole_0(X1)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1606, plain, (esk4_1(X1)=np__2|esk4_1(X1)=np__4|esk4_1(X1)=np__1|esk4_1(X1)=np__0|esk4_1(X1)=np__3|v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(esk4_1(X1),np__4)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1604, plain, (esk3_1(X1)=np__2|esk3_1(X1)=np__4|esk3_1(X1)=np__1|esk3_1(X1)=np__0|esk3_1(X1)=np__3|v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(esk3_1(X1),np__4)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1686, plain, (m2_finseq_1(np__0,X1))).#trainneg
cnf(i_0_1664, plain, (~v1_xboole_0(k19_euclid(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_422, plain, (r1_xxreal_0(X1,np__4)|~v1_xxreal_0(np__2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__2))).#trainneg
cnf(i_0_1690, plain, (v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_1791, plain, (v1_xcmplx_0(np__3))).#trainneg
cnf(i_0_731, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|~v1_xxreal_0(k3_finseq_1(esk1_0))|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_581, plain, (v1_int_1(esk9_2(X1,X2))|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1573, plain, (esk13_0=np__2|esk13_0=np__4|esk13_0=np__1|esk13_0=np__0|esk13_0=np__3|~r1_xxreal_0(esk13_0,np__4))).#trainneg
cnf(i_0_1583, plain, (~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1867, plain, (v1_xcmplx_0(np__4))).#trainneg
cnf(i_0_1870, plain, (v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_583, plain, (v1_int_1(esk6_1(X1))|~v5_membered(X1))).#trainneg
cnf(i_0_589, plain, (v1_int_1(k5_numbers)|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1770, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__0)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1)|~r2_hidden(np__1,k9_xtuple_0(X1))|~r2_hidden(np__0,k9_xtuple_0(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1))).#trainneg
cnf(i_0_1580, plain, (X1=np__2|X1=np__4|X1=np__1|X1=np__0|X1=np__3|~r1_xxreal_0(X1,np__4)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_417, plain, (r1_xxreal_0(X1,np__4)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_1811, plain, (v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_1872, plain, (v1_int_1(np__0))).#trainneg
cnf(i_0_590, plain, (v1_int_1(k18_euclid(X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1545, plain, (r1_xxreal_0(esk8_0,np__4)|r1_xxreal_0(np__4,esk8_0))).#trainneg
cnf(i_0_1651, plain, (~v1_xboole_0(k9_xtuple_0(esk7_0)))).#trainneg
cnf(i_0_2029, plain, (v1_int_1(np__4))).#trainneg
cnf(i_0_2046, plain, (v1_int_1(np__2))).#trainneg
cnf(i_0_591, plain, (v1_int_1(k17_euclid(X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_598, plain, (v1_relat_1(esk16_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_414, plain, (r1_xxreal_0(X1,np__4)|~v1_xxreal_0(np__3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__3))).#trainneg
cnf(i_0_1945, plain, (v1_int_1(np__3))).#trainneg
cnf(i_0_1951, plain, (v1_int_1(np__1))).#trainneg
cnf(i_0_411, plain, (r1_xxreal_0(X1,np__4)|~v1_xxreal_0(np__0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__0))).#trainneg
cnf(i_0_600, plain, (v1_relat_1(esk6_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_1548, plain, (r1_xxreal_0(esk8_0,np__2)|r1_xxreal_0(np__2,esk8_0))).#trainneg
cnf(i_0_1662, negated_conjecture, (~v1_xboole_0(k9_xtuple_0(esk1_0)))).#trainneg
cnf(i_0_1589, plain, (v3_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1546, plain, (r1_xxreal_0(X1,np__4)|r1_xxreal_0(np__4,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_601, plain, (v6_membered(esk16_1(X1))|v1_xboole_0(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_602, plain, (v6_membered(esk9_2(X1,X2))|~v6_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1963, plain, (r1_xxreal_0(esk4_1(X1),np__3)|r1_xxreal_0(np__4,esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1964, plain, (r1_xxreal_0(esk3_1(X1),np__3)|r1_xxreal_0(np__4,esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1530, plain, (v1_xxreal_0(k7_partfun1(X1,X2,X3))|~v2_membered(X1)|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_1719, plain, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1)))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~m1_subset_1(X2,k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_1764, plain, (r1_xxreal_0(k3_finseq_1(X1),np__0)|r1_xxreal_0(np__1,k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1542, plain, (v1_xboole_0(esk14_1(np__0)))).#trainneg
cnf(i_0_418, plain, (r1_xxreal_0(X1,np__2)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_603, plain, (v6_membered(esk6_1(k1_zfmisc_1(X1)))|~v6_membered(X1))).#trainneg
cnf(i_0_1553, plain, (r1_xxreal_0(esk8_0,np__1)|r1_xxreal_0(np__1,esk8_0))).#trainneg
cnf(i_0_1549, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_606, plain, (v1_finset_1(esk9_2(X1,X2))|~v1_finset_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2067, plain, (r1_xxreal_0(esk4_1(X1),np__2)|r1_xxreal_0(np__3,esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1771, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__4)|~r2_hidden(np__4,k9_xtuple_0(X1))|~r2_hidden(np__3,k9_xtuple_0(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1))).#trainneg
cnf(i_0_1614, plain, (~v7_ordinal1(X1)|~v1_xboole_0(u1_struct_0(k15_euclid(X1))))).#trainneg
cnf(i_0_410, plain, (r1_xxreal_0(X1,np__1)|~v1_xxreal_0(np__0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__0))).#trainneg
cnf(i_0_1995, plain, (v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(np__4,esk4_1(X1))|~r1_xxreal_0(esk4_1(X1),np__3)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1767, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_421, plain, (r1_xxreal_0(X1,np__3)|~v1_xxreal_0(np__3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__2))).#trainneg
cnf(i_0_627, negated_conjecture, (r1_xxreal_0(k1_card_1(esk1_0),np__4))).#trainneg
cnf(i_0_412, plain, (r1_xxreal_0(X1,np__2)|~v1_xxreal_0(np__0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__0))).#trainneg
cnf(i_0_607, plain, (v1_finset_1(esk6_1(k1_zfmisc_1(X1)))|~v1_finset_1(X1))).#trainneg
cnf(i_0_1602, plain, (r1_xxreal_0(esk8_0,np__0)|r1_xxreal_0(np__0,esk8_0))).#trainneg
cnf(i_0_609, plain, (v3_membered(esk16_1(X1))|v1_xboole_0(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_610, plain, (v3_membered(esk9_2(X1,X2))|~v3_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1996, plain, (v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(np__4,esk3_1(X1))|~r1_xxreal_0(esk3_1(X1),np__3)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1773, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)|~r2_hidden(np__3,k9_xtuple_0(X1))|~r2_hidden(np__2,k9_xtuple_0(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1))).#trainneg
cnf(i_0_1718, plain, (~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)|~r1_xxreal_0(k2_xcmplx_0(X1,np__1),X2)|~r1_xxreal_0(X2,X1))).#trainneg
cnf(i_0_1973, plain, (r1_xxreal_0(k3_finseq_1(X1),np__3)|r1_xxreal_0(np__4,k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_415, plain, (r1_xxreal_0(X1,np__3)|~v1_xxreal_0(np__3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_409, plain, (r1_xxreal_0(X1,np__3)|~v1_xxreal_0(np__3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__0))).#trainneg
cnf(i_0_611, plain, (v3_membered(esk6_1(k1_zfmisc_1(X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_1609, plain, (r1_xxreal_0(esk8_0,np__3)|r1_xxreal_0(np__3,esk8_0))).#trainneg
cnf(i_0_1663, plain, (~v1_xboole_0(esk6_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_1975, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__4,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2063, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__2,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_613, plain, (v2_membered(esk16_1(X1))|v1_xboole_0(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_614, plain, (v2_membered(esk9_2(X1,X2))|~v2_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1723, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(np__1,X2))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~m1_subset_1(X2,k5_numbers)|~r2_hidden(k2_xcmplx_0(np__1,X2),k4_finseq_1(X1))|~r2_hidden(X2,k4_finseq_1(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1))).#trainneg
cnf(i_0_2023, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~m1_subset_1(k2_xcmplx_0(X2,np__1),k4_finseq_1(X1))|~m1_subset_1(X2,k5_numbers)|~r2_hidden(X2,k4_finseq_1(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1))).#trainneg
cnf(i_0_952, plain, (v4_relat_1(esk11_1(X1),k5_numbers))).#trainneg
cnf(i_0_1993, plain, (~r1_xxreal_0(np__4,esk13_0)|~r1_xxreal_0(esk13_0,np__3))).#trainneg
cnf(i_0_615, plain, (v2_membered(esk6_1(k1_zfmisc_1(X1)))|~v2_membered(X1))).#trainneg
cnf(i_0_1753, plain, (r1_xxreal_0(esk13_0,np__0)|r1_xxreal_0(np__1,esk13_0))).#trainneg
cnf(i_0_1613, plain, (v1_xxreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1982, plain, (v1_xcmplx_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2079, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__3,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1648, plain, (X1=esk11_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_617, plain, (r1_xxreal_0(X1,esk2_1(X2))|v1_topreal1(X2)|~v1_xxreal_0(esk2_1(X2))|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_618, plain, (v1_membered(esk16_1(X1))|v1_xboole_0(X1)|~v1_membered(X1))).#trainneg
cnf(i_0_1774, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(np__1,X2))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~m1_subset_1(X2,k5_numbers)|~r2_hidden(k2_xcmplx_0(np__1,X2),k9_xtuple_0(X1))|~r2_hidden(X2,k9_xtuple_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1))).#trainneg
cnf(i_0_1554, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1833, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~m1_subset_1(k2_xcmplx_0(X2,np__1),k9_xtuple_0(X1))|~r2_hidden(X2,k9_xtuple_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2052, plain, (r1_xxreal_0(esk3_1(X1),np__1)|r1_xxreal_0(np__2,esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1826, plain, (k1_nat_1(X1,X2)=np__2|k1_nat_1(X1,X2)=np__4|k1_nat_1(X1,X2)=np__1|k1_nat_1(X1,X2)=np__0|k1_nat_1(X1,X2)=np__3|~r1_xxreal_0(k1_nat_1(X1,X2),np__4)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1263, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1271, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1572, plain, (k2_xcmplx_0(esk8_0,np__0)=esk8_0)).#trainneg
cnf(i_0_2007, plain, (~r1_xxreal_0(np__4,X1)|~r1_xxreal_0(X1,np__3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_619, plain, (v1_membered(esk9_2(X1,X2))|~v1_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1961, plain, (r1_xxreal_0(esk13_0,np__3)|r1_xxreal_0(np__4,esk13_0))).#trainneg
cnf(i_0_1529, plain, (v7_ordinal1(k7_partfun1(X1,X2,X3))|~v6_membered(X1)|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_1532, plain, (v1_finset_1(k7_partfun1(X1,X2,X3))|~v5_finset_1(X1)|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_620, plain, (v1_membered(esk6_1(k1_zfmisc_1(X1)))|~v1_membered(X1))).#trainneg
cnf(i_0_622, plain, (v7_ordinal1(k9_xtuple_0(X1))|~v1_finset_1(k9_xtuple_0(X1))|~v5_ordinal1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2005, plain, (~r1_xxreal_0(np__4,k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1610, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__3,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1832, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(np__1,X2))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~m1_subset_1(k2_nat_1(np__1,X2),k9_xtuple_0(X1))|~r2_hidden(X2,k9_xtuple_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1680, plain, (~v7_ordinal1(X1)|~r1_xxreal_0(np__3,X1)|~r1_xxreal_0(X1,np__2))).#trainneg
cnf(i_0_2243, negated_conjecture, (r1_xxreal_0(X1,np__4)|~v1_xxreal_0(k1_card_1(esk1_0))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k1_card_1(esk1_0)))).#trainneg
cnf(i_0_2391, plain, (esk14_1(np__0)=np__0)).#trainneg
cnf(i_0_628, plain, (v1_finseq_6(X1,X2)|v1_xboole_0(X2)|k7_partfun1(X2,X1,k1_card_1(X1))!=k7_partfun1(X2,X1,np__1)|~m2_finseq_1(X1,X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2257, plain, (esk8_0=np__3|esk8_0=np__0|esk8_0=np__1|esk8_0=np__4|esk8_0=np__2|r1_xxreal_0(np__4,esk8_0)|~m1_subset_1(esk8_0,k5_numbers))).#trainneg
cnf(i_0_1302, plain, (v6_membered(k9_xtuple_0(X2))|~v4_relat_1(X2,X1)|~v6_membered(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_629, plain, (v1_finseq_1(esk6_1(u1_struct_0(k15_euclid(X1))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_630, plain, (v3_valued_0(esk6_1(u1_struct_0(k15_euclid(X1))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2142, plain, (r1_xxreal_0(np__1,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(k3_finseq_1(X1))|~v1_xxreal_0(esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2173, plain, (k2_xcmplx_0(np__4,np__0)=np__4)).#trainneg
cnf(i_0_2164, plain, (k2_xcmplx_0(np__3,np__0)=np__3)).#trainneg
cnf(i_0_2206, plain, (v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xboole_0(esk3_1(X1))|~v4_relat_1(X1,u1_struct_0(k15_euclid(np__2)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2128, plain, (v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xboole_0(esk4_1(X1))|~v4_relat_1(X1,u1_struct_0(k15_euclid(np__2)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2172, plain, (k3_xcmplx_0(np__1,np__4)=np__4)).#trainneg
cnf(i_0_634, plain, (k3_finseq_1(k1_card_1(X1))=k1_card_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2662, plain, (m2_subset_1(esk8_0,k1_numbers,k5_numbers)|~m1_subset_1(esk8_0,k5_numbers))).#trainneg
cnf(i_0_2134, plain, (v3_ordinal1(esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v3_ordinal1(X2)|~v4_relat_1(X1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_642, plain, (v2_xxreal_0(esk2_1(X1))|v1_topreal1(X1)|~v1_xreal_0(esk2_1(X1))|~v1_xreal_0(np__1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_652, plain, (k2_xcmplx_0(X1,esk8_0)=k2_xcmplx_0(esk8_0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_653, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_2661, plain, (m2_subset_1(esk8_0,k1_numbers,k5_numbers)|~v7_ordinal1(esk8_0))).#trainneg
cnf(i_0_2135, plain, (v1_finset_1(esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v5_finset_1(X2)|~v4_relat_1(X1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2136, plain, (v1_xreal_0(esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v3_membered(X2)|~v4_relat_1(X1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_665, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_677, plain, (v7_ordinal1(u1_pre_topc(X1))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2657, plain, (k1_nat_1(X1,np__1)=np__3|k1_nat_1(X1,np__1)=np__0|k1_nat_1(X1,np__1)=np__1|k1_nat_1(X1,np__1)=np__4|k1_nat_1(X1,np__1)=np__2|r1_xxreal_0(np__4,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2653, plain, (k2_xcmplx_0(X1,X2)=np__3|k2_xcmplx_0(X1,X2)=np__0|k2_xcmplx_0(X1,X2)=np__1|k2_xcmplx_0(X1,X2)=np__4|k2_xcmplx_0(X1,X2)=np__2|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__4))).#trainneg
cnf(i_0_2716, plain, (~r1_xxreal_0(np__3,k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),np__2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2652, plain, (k1_nat_1(X2,X1)=np__3|k1_nat_1(X2,X1)=np__0|k1_nat_1(X2,X1)=np__1|k1_nat_1(X2,X1)=np__4|k1_nat_1(X2,X1)=np__2|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)|~r1_xxreal_0(k1_nat_1(X2,X1),np__4))).#trainneg
cnf(i_0_2353, plain, (k7_partfun1(X1,X2,X3)=np__4|r1_xxreal_0(k7_partfun1(X1,X2,X3),np__3)|~v7_ordinal1(k7_partfun1(X1,X2,X3))|~r1_xxreal_0(k7_partfun1(X1,X2,X3),np__4)|~v2_membered(X1)|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_2658, plain, (v7_ordinal1(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1582, plain, (k3_xcmplx_0(esk8_0,np__0)=np__0)).#trainneg
cnf(i_0_678, plain, (v1_xxreal_0(u1_pre_topc(X1))|~v2_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2299, plain, (r1_xxreal_0(esk8_0,np__4)|~r1_xxreal_0(esk8_0,np__3))).#trainneg
cnf(i_0_2137, plain, (v1_xcmplx_0(esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_membered(X2)|~v4_relat_1(X1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2138, plain, (v1_int_1(esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v5_membered(X2)|~v4_relat_1(X1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_679, plain, (v3_ordinal1(u1_pre_topc(X1))|~v3_ordinal1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_680, plain, (v1_finset_1(u1_pre_topc(X1))|~v5_finset_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2706, plain, (v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(np__3,esk4_1(X1))|~r1_xxreal_0(esk4_1(X1),np__2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2707, plain, (v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(np__3,esk3_1(X1))|~r1_xxreal_0(esk3_1(X1),np__2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2526, plain, (~v7_ordinal1(X1)|~r1_xxreal_0(np__2,X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_2656, plain, (k1_nat_1(np__1,X1)=np__3|k1_nat_1(np__1,X1)=np__0|k1_nat_1(np__1,X1)=np__1|k1_nat_1(np__1,X1)=np__4|k1_nat_1(np__1,X1)=np__2|r1_xxreal_0(np__4,X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2089, plain, (esk4_1(X1)=np__1|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(esk4_1(X1))|~r1_xxreal_0(esk4_1(X1),np__1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2147, plain, (r1_xxreal_0(esk8_0,np__4)|~r1_xxreal_0(esk8_0,np__2))).#trainneg
cnf(i_0_2226, plain, (r1_xxreal_0(esk8_0,np__4)|~r1_xxreal_0(esk8_0,np__1))).#trainneg
cnf(i_0_2163, plain, (k3_xcmplx_0(np__1,np__3)=np__3)).#trainneg
cnf(i_0_681, plain, (v1_xreal_0(u1_pre_topc(X1))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2397, plain, (r1_xxreal_0(esk8_0,np__2)|~r1_xxreal_0(esk8_0,np__1))).#trainneg
cnf(i_0_2496, plain, (~r1_xxreal_0(np__1,esk13_0)|~r1_xxreal_0(esk13_0,np__0))).#trainneg
cnf(i_0_682, plain, (v1_xcmplx_0(u1_pre_topc(X1))|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_683, plain, (v1_int_1(u1_pre_topc(X1))|~v5_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2092, plain, (esk3_1(X1)=np__1|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(esk3_1(X1))|~r1_xxreal_0(esk3_1(X1),np__1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2126, plain, (v1_relat_1(esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v4_relat_1(X1,u1_struct_0(k15_euclid(np__2)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2354, plain, (k7_partfun1(X1,X2,X3)=np__3|r1_xxreal_0(k7_partfun1(X1,X2,X3),np__2)|~v7_ordinal1(k7_partfun1(X1,X2,X3))|~r1_xxreal_0(k7_partfun1(X1,X2,X3),np__3)|~v2_membered(X1)|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_2204, plain, (v1_relat_1(esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v4_relat_1(X1,u1_struct_0(k15_euclid(np__2)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2519, plain, (k7_partfun1(X1,X2,X3)=np__2|r1_xxreal_0(k7_partfun1(X1,X2,X3),np__1)|~v7_ordinal1(k7_partfun1(X1,X2,X3))|~r1_xxreal_0(k7_partfun1(X1,X2,X3),np__2)|~v2_membered(X1)|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_2186, plain, (r1_xxreal_0(k1_nat_1(X1,np__1),k3_finseq_1(X2))|v1_xboole_0(X2)|v3_funct_1(X2)|r1_xxreal_0(esk3_1(X2),X1)|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(esk3_1(X2))|~v1_xxreal_0(k1_nat_1(X1,np__1))|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(esk3_1(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2678, plain, (v7_ordinal1(k7_partfun1(k5_numbers,X1,X2))|~m2_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_2141, plain, (r1_xxreal_0(k1_nat_1(X1,np__1),k3_finseq_1(X2))|v1_xboole_0(X2)|v3_funct_1(X2)|r1_xxreal_0(esk4_1(X2),X1)|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(esk4_1(X2))|~v1_xxreal_0(k1_nat_1(X1,np__1))|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(esk4_1(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2313, plain, (r1_xxreal_0(esk8_0,np__4)|~r1_xxreal_0(esk8_0,np__0))).#trainneg
cnf(i_0_2463, plain, (r1_xxreal_0(esk8_0,np__3)|~r1_xxreal_0(esk8_0,np__2))).#trainneg
cnf(i_0_2165, plain, (k3_xcmplx_0(np__3,np__0)=np__0)).#trainneg
cnf(i_0_2704, plain, (~r1_xxreal_0(np__3,esk13_0)|~r1_xxreal_0(esk13_0,np__2))).#trainneg
cnf(i_0_684, plain, (v6_membered(u1_pre_topc(X1))|~v6_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2433, plain, (r1_xxreal_0(esk8_0,np__1)|~r1_xxreal_0(esk8_0,np__0))).#trainneg
cnf(i_0_1117, plain, (v7_ordinal1(esk8_0)|~v1_finset_1(esk8_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_2212, plain, (v3_ordinal1(esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v3_ordinal1(X2)|~v4_relat_1(X1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2213, plain, (v1_finset_1(esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v5_finset_1(X2)|~v4_relat_1(X1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_685, plain, (v1_finset_1(u1_pre_topc(X1))|~v1_finset_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_686, plain, (v3_membered(u1_pre_topc(X1))|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2090, plain, (r1_xxreal_0(X1,esk4_1(X2))|v1_xboole_0(X2)|v3_funct_1(X2)|~v1_xxreal_0(esk4_1(X2))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2383, plain, (k3_finseq_1(X1)=np__0|r1_xxreal_0(np__1,k3_finseq_1(X1))|~v1_xxreal_0(k3_finseq_1(X1))|~r1_xxreal_0(np__0,k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2498, plain, (v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(np__1,esk4_1(X1))|~r1_xxreal_0(esk4_1(X1),np__0)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2650, plain, (esk3_1(X1)=np__1|r1_xxreal_0(np__2,esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(esk3_1(X1))|~r1_xxreal_0(np__1,esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2479, plain, (r1_xxreal_0(esk8_0,np__2)|~r1_xxreal_0(esk8_0,np__0))).#trainneg
cnf(i_0_2538, plain, (r1_xxreal_0(esk8_0,np__3)|~r1_xxreal_0(esk8_0,np__1))).#trainneg
cnf(i_0_2174, plain, (k3_xcmplx_0(np__4,np__0)=np__0)).#trainneg
cnf(i_0_687, plain, (v2_membered(u1_pre_topc(X1))|~v2_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2550, plain, (r1_xxreal_0(esk8_0,np__3)|~r1_xxreal_0(esk8_0,np__0))).#trainneg
cnf(i_0_1881, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2600, plain, (v1_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2214, plain, (v1_xreal_0(esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v3_membered(X2)|~v4_relat_1(X1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_688, plain, (v1_membered(u1_pre_topc(X1))|~v1_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_690, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_2217, plain, (v1_xboole_0(X1)|v1_xboole_0(k9_xtuple_0(X1))|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__0)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1)|~r2_hidden(np__1,k9_xtuple_0(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~m1_subset_1(np__0,k9_xtuple_0(X1)))).#trainneg
cnf(i_0_2533, plain, (k3_finseq_1(X1)=np__3|r1_xxreal_0(np__4,k3_finseq_1(X1))|~v1_xxreal_0(k3_finseq_1(X1))|~r1_xxreal_0(np__3,k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2719, plain, (~r1_xxreal_0(np__3,X1)|~r1_xxreal_0(X1,np__2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2428, plain, (esk3_1(X1)=np__2|r1_xxreal_0(np__3,esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(esk3_1(X1))|~r1_xxreal_0(np__2,esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2508, plain, (~r1_xxreal_0(np__1,k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),np__0)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2425, plain, (esk4_1(X1)=np__2|r1_xxreal_0(np__3,esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(esk4_1(X1))|~r1_xxreal_0(np__2,esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2140, plain, (r1_xxreal_0(k1_nat_1(np__1,X1),k3_finseq_1(X2))|v1_xboole_0(X2)|v3_funct_1(X2)|r1_xxreal_0(esk4_1(X2),X1)|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(esk4_1(X2))|~v1_xxreal_0(k1_nat_1(np__1,X1))|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(esk4_1(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2352, plain, (k7_partfun1(X1,X2,X3)=np__1|r1_xxreal_0(k7_partfun1(X1,X2,X3),np__0)|~v7_ordinal1(k7_partfun1(X1,X2,X3))|~r1_xxreal_0(k7_partfun1(X1,X2,X3),np__1)|~v2_membered(X1)|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_2185, plain, (r1_xxreal_0(k1_nat_1(np__1,X1),k3_finseq_1(X2))|v1_xboole_0(X2)|v3_funct_1(X2)|r1_xxreal_0(esk3_1(X2),X1)|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(esk3_1(X2))|~v1_xxreal_0(k1_nat_1(np__1,X1))|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(esk3_1(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1264, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1270, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_2665, plain, (m1_subset_1(esk8_0,k5_numbers)|~v7_ordinal1(esk8_0))).#trainneg
cnf(i_0_691, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_2049, plain, (r1_xxreal_0(esk13_0,np__1)|r1_xxreal_0(np__2,esk13_0))).#trainneg
cnf(i_0_2602, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2216, plain, (v1_int_1(esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v5_membered(X2)|~v4_relat_1(X1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2130, plain, (v1_finseq_1(esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(X2)|~v4_relat_1(X1,u1_struct_0(k15_euclid(X2)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_692, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_693, plain, (v1_finset_1(X1)|~v5_finset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_1722, plain, (r1_xxreal_0(k2_xcmplx_0(np__1,X1),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2677, plain, (v7_ordinal1(k7_partfun1(k4_finseq_1(X1),X2,X3))|~m2_finseq_1(X2,k4_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2495, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)|~r2_hidden(np__2,k9_xtuple_0(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r1_xxreal_0(np__3,k3_finseq_1(X1)))).#trainneg
cnf(i_0_1265, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1274, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_2065, plain, (r1_xxreal_0(esk13_0,np__2)|r1_xxreal_0(np__3,esk13_0))).#trainneg
cnf(i_0_2008, plain, (~r1_xxreal_0(np__4,X1)|~r1_xxreal_0(X1,np__3)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_694, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_1539, plain, (v1_relat_1(esk6_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_1569, plain, (v1_funct_1(esk6_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_2131, plain, (v3_valued_0(esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(X2)|~v4_relat_1(X1,u1_struct_0(k15_euclid(X2)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2208, plain, (v1_finseq_1(esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(X2)|~v4_relat_1(X1,u1_struct_0(k15_euclid(X2)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_695, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_696, plain, (v1_int_1(X1)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_969, plain, (k19_euclid(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1))),k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)))=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k2_xcmplx_0(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_1369, plain, (k10_finseq_1(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1))),k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)))=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k2_xcmplx_0(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_1835, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(np__1,X2))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~m1_subset_1(k2_nat_1(np__1,X2),k9_xtuple_0(X1))|~r2_hidden(X2,k9_xtuple_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3075, plain, (r1_xxreal_0(X1,np__4)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk8_0)|~r1_xxreal_0(esk8_0,np__2))).#trainneg
cnf(i_0_2788, plain, (v2_finseq_1(np__0))).#trainneg
cnf(i_0_2790, plain, (v3_card_1(np__0,np__0))).#trainneg
cnf(i_0_697, plain, (v6_membered(X1)|~v6_membered(k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_698, plain, (v1_finset_1(X1)|~v1_finset_1(k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_699, plain, (v3_membered(X1)|~v3_membered(k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_2742, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_3079, plain, (r1_xxreal_0(X1,np__4)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk8_0)|~r1_xxreal_0(esk8_0,np__1))).#trainneg
cnf(i_0_2802, plain, (v6_membered(k9_xtuple_0(esk7_0)))).#trainneg
cnf(i_0_2803, negated_conjecture, (v6_membered(k9_xtuple_0(esk1_0)))).#trainneg
cnf(i_0_700, plain, (v2_membered(X1)|~v2_membered(k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_3076, plain, (esk8_0=np__3|esk8_0=np__0|esk8_0=np__1|esk8_0=np__4|esk8_0=np__2|~m1_subset_1(esk8_0,k5_numbers)|~r1_xxreal_0(esk8_0,np__2))).#trainneg
cnf(i_0_701, plain, (v1_membered(X1)|~v1_membered(k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_706, plain, (v4_relat_1(esk16_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2996, plain, (k7_partfun1(X1,X2,X3)=np__4|r1_xxreal_0(k7_partfun1(X1,X2,X3),np__3)|~v2_membered(X1)|~r1_xxreal_0(k7_partfun1(X1,X2,X3),np__4)|~m2_finseq_1(X2,X1)|~m1_subset_1(k7_partfun1(X1,X2,X3),k5_numbers))).#trainneg
cnf(i_0_3253, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__1),X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3189, plain, (v1_xreal_0(esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2928, plain, (v1_xreal_0(esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3127, plain, (r1_xxreal_0(X1,np__4)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk8_0)|~r1_xxreal_0(esk8_0,np__0))).#trainneg
cnf(i_0_2804, plain, (v6_membered(k9_xtuple_0(np__0)))).#trainneg
cnf(i_0_2879, plain, (v3_ordinal1(esk4_1(esk7_0)))).#trainneg
cnf(i_0_708, plain, (v4_relat_1(esk6_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_3080, plain, (esk8_0=np__3|esk8_0=np__0|esk8_0=np__1|esk8_0=np__4|esk8_0=np__2|~m1_subset_1(esk8_0,k5_numbers)|~r1_xxreal_0(esk8_0,np__1))).#trainneg
cnf(i_0_3026, plain, (~r1_xxreal_0(np__2,esk13_0)|~r1_xxreal_0(esk13_0,np__1))).#trainneg
cnf(i_0_2880, negated_conjecture, (v3_ordinal1(esk4_1(esk1_0)))).#trainneg
cnf(i_0_2209, plain, (v3_valued_0(esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(X2)|~v4_relat_1(X1,u1_struct_0(k15_euclid(X2)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3278, plain, (v3_ordinal1(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_710, plain, (v5_relat_1(esk16_1(k2_zfmisc_1(X1,X2)),X2)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_712, plain, (v5_relat_1(esk6_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_2819, plain, (k2_xcmplx_0(np__0,np__3)=np__3)).#trainneg
cnf(i_0_3104, plain, (k7_partfun1(X1,X2,X3)=np__3|r1_xxreal_0(k7_partfun1(X1,X2,X3),np__2)|~v2_membered(X1)|~r1_xxreal_0(k7_partfun1(X1,X2,X3),np__3)|~m2_finseq_1(X2,X1)|~m1_subset_1(k7_partfun1(X1,X2,X3),k5_numbers))).#trainneg
cnf(i_0_3215, plain, (k7_partfun1(X1,X2,X3)=np__1|r1_xxreal_0(k7_partfun1(X1,X2,X3),np__0)|~v2_membered(X1)|~r1_xxreal_0(k7_partfun1(X1,X2,X3),np__1)|~m2_finseq_1(X2,X1)|~m1_subset_1(k7_partfun1(X1,X2,X3),k5_numbers))).#trainneg
cnf(i_0_3119, plain, (k7_partfun1(k5_numbers,X1,X2)=np__4|r1_xxreal_0(k7_partfun1(k5_numbers,X1,X2),np__3)|~r1_xxreal_0(k7_partfun1(k5_numbers,X1,X2),np__4)|~m2_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_2811, plain, (k2_xcmplx_0(np__0,np__4)=np__4)).#trainneg
cnf(i_0_3109, plain, (k7_partfun1(X1,X2,X3)=np__2|r1_xxreal_0(k7_partfun1(X1,X2,X3),np__1)|~v2_membered(X1)|~r1_xxreal_0(k7_partfun1(X1,X2,X3),np__2)|~m2_finseq_1(X2,X1)|~m1_subset_1(k7_partfun1(X1,X2,X3),k5_numbers))).#trainneg
cnf(i_0_2969, plain, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k1_nat_1(X2,np__1)=np__2|k1_nat_1(X2,np__1)=np__4|k1_nat_1(X2,np__1)=np__1|k1_nat_1(X2,np__1)=np__0|r1_xxreal_0(np__4,X2)|~v7_ordinal1(X2)|~r1_xxreal_0(np__3,k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_3074, plain, (esk8_0=np__4|~r1_xxreal_0(np__4,esk8_0)|~r1_xxreal_0(esk8_0,np__2))).#trainneg
cnf(i_0_2800, plain, (v6_membered(k9_xtuple_0(esk11_1(X1))))).#trainneg
cnf(i_0_3010, plain, (v1_xcmplx_0(esk4_1(esk7_0)))).#trainneg
cnf(i_0_3130, plain, (esk8_0=np__3|~r1_xxreal_0(np__3,esk8_0)|~r1_xxreal_0(esk8_0,np__2))).#trainneg
cnf(i_0_714, plain, (k3_xcmplx_0(X1,esk8_0)=k3_xcmplx_0(esk8_0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3128, plain, (esk8_0=np__3|esk8_0=np__0|esk8_0=np__1|esk8_0=np__4|esk8_0=np__2|~m1_subset_1(esk8_0,k5_numbers)|~r1_xxreal_0(esk8_0,np__0))).#trainneg
cnf(i_0_3011, negated_conjecture, (v1_xcmplx_0(esk4_1(esk1_0)))).#trainneg
cnf(i_0_3279, plain, (v3_ordinal1(X1)|~v7_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1766, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__1,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_727, plain, (esk2_1(X1)=np__1|v1_topreal1(X1)|~v1_xxreal_0(np__1)|~v1_xxreal_0(esk2_1(X1))|~r1_xxreal_0(esk2_1(X1),np__1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_734, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk16_1(X2)))).#trainneg
cnf(i_0_3029, plain, (v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(np__2,esk3_1(X1))|~r1_xxreal_0(esk3_1(X1),np__1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2990, plain, (k1_nat_1(np__1,X1)=np__2|k1_nat_1(np__1,X1)=np__4|k1_nat_1(np__1,X1)=np__1|k1_nat_1(np__1,X1)=np__0|k1_nat_1(np__1,X1)=np__3|r1_xxreal_0(np__4,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3028, plain, (v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(np__2,esk4_1(X1))|~r1_xxreal_0(esk4_1(X1),np__1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3277, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)|~r1_xxreal_0(np__3,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r1_xxreal_0(np__2,k3_finseq_1(X1)))).#trainneg
cnf(i_0_2744, plain, (v1_xxreal_0(k1_card_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3069, plain, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k1_nat_1(np__1,X2)=np__2|k1_nat_1(np__1,X2)=np__4|k1_nat_1(np__1,X2)=np__1|k1_nat_1(np__1,X2)=np__0|r1_xxreal_0(np__4,X2)|~v7_ordinal1(X2)|~r1_xxreal_0(np__3,k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3111, plain, (k7_partfun1(k5_numbers,X1,X2)=np__2|k7_partfun1(k5_numbers,X1,X2)=np__4|k7_partfun1(k5_numbers,X1,X2)=np__1|k7_partfun1(k5_numbers,X1,X2)=np__0|k7_partfun1(k5_numbers,X1,X2)=np__3|~r1_xxreal_0(k7_partfun1(k5_numbers,X1,X2),np__4)|~m2_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_2753, negated_conjecture, (k3_finseq_1(np__3)=np__3|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0)).#trainneg
cnf(i_0_3249, plain, (k2_xcmplx_0(np__1,X1)=np__2|k2_xcmplx_0(np__1,X1)=np__4|k2_xcmplx_0(np__1,X1)=np__1|k2_xcmplx_0(np__1,X1)=np__0|k2_xcmplx_0(np__1,X1)=np__3|r1_xxreal_0(np__4,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3272, plain, (v7_ordinal1(k7_partfun1(k9_xtuple_0(X1),X2,X3))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m2_finseq_1(X2,k9_xtuple_0(X1)))).#trainneg
cnf(i_0_3078, plain, (esk8_0=np__4|~r1_xxreal_0(np__4,esk8_0)|~r1_xxreal_0(esk8_0,np__1))).#trainneg
cnf(i_0_3559, plain, (v3_membered(k9_xtuple_0(esk7_0)))).#trainneg
cnf(i_0_3150, plain, (v3_ordinal1(esk3_1(esk7_0)))).#trainneg
cnf(i_0_3500, negated_conjecture, (k1_card_1(esk1_0)=np__4|~r1_xxreal_0(np__4,k1_card_1(esk1_0)))).#trainneg
cnf(i_0_735, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X1,esk9_2(X2,X3)))).#trainneg
cnf(i_0_3564, plain, (esk8_0=np__2|esk8_0=np__4|esk8_0=np__1|esk8_0=np__0|esk8_0=np__3|r1_xxreal_0(np__2,esk8_0)|~m1_subset_1(esk8_0,k5_numbers))).#trainneg
cnf(i_0_3145, plain, (v7_ordinal1(esk8_0)|~v1_finset_1(esk8_0)|~v1_card_1(k1_numbers))).#trainneg
cnf(i_0_736, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk6_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_737, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,u1_pre_topc(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_3690, plain, (k7_partfun1(X1,X2,X3)=np__1|r1_xxreal_0(k7_partfun1(X1,X2,X3),np__0)|~m1_subset_1(k7_partfun1(X1,X2,X3),k5_numbers)|~r1_xxreal_0(k7_partfun1(X1,X2,X3),np__1)|~m2_finseq_1(X2,X1)|~v3_membered(X1))).#trainneg
cnf(i_0_3713, plain, (k7_partfun1(X1,X2,X3)=np__2|r1_xxreal_0(k7_partfun1(X1,X2,X3),np__1)|~m1_subset_1(k7_partfun1(X1,X2,X3),k5_numbers)|~r1_xxreal_0(k7_partfun1(X1,X2,X3),np__2)|~m2_finseq_1(X2,X1)|~v3_membered(X1))).#trainneg
cnf(i_0_3587, plain, (k7_partfun1(X1,X2,X3)=np__4|r1_xxreal_0(k7_partfun1(X1,X2,X3),np__3)|~m1_subset_1(k7_partfun1(X1,X2,X3),k5_numbers)|~r1_xxreal_0(k7_partfun1(X1,X2,X3),np__4)|~m2_finseq_1(X2,X1)|~v3_membered(X1))).#trainneg
cnf(i_0_3363, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),esk3_1(esk1_0)))).#trainneg
cnf(i_0_3477, plain, (esk4_1(X1)=np__1|r1_xxreal_0(esk4_1(X1),np__0)|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(esk4_1(X1))|~r1_xxreal_0(esk4_1(X1),np__1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3126, plain, (esk8_0=np__4|~r1_xxreal_0(np__4,esk8_0)|~r1_xxreal_0(esk8_0,np__0))).#trainneg
cnf(i_0_3561, negated_conjecture, (v3_membered(k9_xtuple_0(esk1_0)))).#trainneg
cnf(i_0_3151, negated_conjecture, (v3_ordinal1(esk3_1(esk1_0)))).#trainneg
cnf(i_0_3168, plain, (esk8_0=np__3|~r1_xxreal_0(np__3,esk8_0)|~r1_xxreal_0(esk8_0,np__1))).#trainneg
cnf(i_0_738, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_3658, plain, (esk8_0=np__2|esk8_0=np__4|esk8_0=np__1|esk8_0=np__0|esk8_0=np__3|r1_xxreal_0(np__1,esk8_0)|~m1_subset_1(esk8_0,k5_numbers))).#trainneg
cnf(i_0_3194, plain, (v1_xcmplx_0(esk3_1(esk7_0)))).#trainneg
cnf(i_0_1603, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_740, plain, (v7_ordinal1(k4_finseq_1(X1))|~v6_membered(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_741, plain, (v1_xxreal_0(k4_finseq_1(X1))|~v2_membered(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3505, negated_conjecture, (k7_partfun1(X1,esk1_0,np__4)=k7_partfun1(X1,esk1_0,np__1)|v1_xboole_0(X1)|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__3|~m2_finseq_1(esk1_0,X1)|~v1_finseq_6(esk1_0,X1))).#trainneg
cnf(i_0_3479, plain, (esk4_1(X1)=np__2|r1_xxreal_0(esk4_1(X1),np__1)|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(esk4_1(X1))|~r1_xxreal_0(esk4_1(X1),np__2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3824, plain, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k1_nat_1(np__1,X2)=np__0|k1_nat_1(np__1,X2)=np__1|k1_nat_1(np__1,X2)=np__4|k1_nat_1(np__1,X2)=np__2|r1_xxreal_0(np__4,X2)|~v7_ordinal1(X2)|~r1_xxreal_0(np__3,k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_3371, negated_conjecture, (esk3_1(esk1_0)=np__3|r1_xxreal_0(esk3_1(esk1_0),np__2)|~v7_ordinal1(esk3_1(esk1_0))|~r1_xxreal_0(esk3_1(esk1_0),np__3))).#trainneg
cnf(i_0_3610, negated_conjecture, (r1_xxreal_0(X1,np__4)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k1_card_1(esk1_0)))).#trainneg
cnf(i_0_3375, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),np__3)|~r1_xxreal_0(esk3_1(esk1_0),np__2))).#trainneg
cnf(i_0_3503, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__3|np__4!=np__1)).#trainneg
cnf(i_0_3131, plain, (r1_xxreal_0(X1,np__3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk8_0)|~r1_xxreal_0(esk8_0,np__2))).#trainneg
cnf(i_0_3641, plain, (v3_membered(k9_xtuple_0(np__0)))).#trainneg
cnf(i_0_2885, plain, (k2_xcmplx_0(np__0,esk8_0)=esk8_0)).#trainneg
cnf(i_0_3179, plain, (esk8_0=np__3|~r1_xxreal_0(np__3,esk8_0)|~r1_xxreal_0(esk8_0,np__0))).#trainneg
cnf(i_0_742, plain, (v3_ordinal1(k4_finseq_1(X1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3776, plain, (esk8_0=np__2|esk8_0=np__4|esk8_0=np__1|esk8_0=np__0|esk8_0=np__3|r1_xxreal_0(np__0,esk8_0)|~m1_subset_1(esk8_0,k5_numbers))).#trainneg
cnf(i_0_3787, plain, (v7_ordinal1(esk8_0)|~v1_finset_1(esk8_0)|~v7_ordinal1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2175, plain, (k2_xcmplx_0(X1,np__4)=k2_xcmplx_0(np__4,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_743, plain, (v1_finset_1(k4_finseq_1(X1))|~v5_finset_1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_744, plain, (v1_xreal_0(k4_finseq_1(X1))|~v3_membered(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3504, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__3|np__4!=np__2)).#trainneg
cnf(i_0_3468, plain, (r1_xxreal_0(X1,esk3_1(X2))|v1_xboole_0(X2)|v3_funct_1(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3502, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__3|np__4!=np__0)).#trainneg
cnf(i_0_3697, plain, (k7_partfun1(k5_numbers,X1,X2)=np__3|r1_xxreal_0(k7_partfun1(k5_numbers,X1,X2),np__2)|k7_partfun1(k5_numbers,X1,X2)=np__4|~m1_subset_1(k7_partfun1(k5_numbers,X1,X2),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~r1_xxreal_0(k7_partfun1(k5_numbers,X1,X2),np__4))).#trainneg
cnf(i_0_3467, plain, (esk3_1(X1)=np__1|r1_xxreal_0(np__2,esk3_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3378, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),np__2)|~r1_xxreal_0(esk3_1(esk1_0),np__0))).#trainneg
cnf(i_0_3460, plain, (r1_xxreal_0(esk3_1(X1),np__3)|v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(esk3_1(X1),np__1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3536, negated_conjecture, (esk3_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__3|~r1_xxreal_0(np__4,esk3_1(esk1_0)))).#trainneg
cnf(i_0_3169, plain, (r1_xxreal_0(X1,np__3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk8_0)|~r1_xxreal_0(esk8_0,np__1))).#trainneg
cnf(i_0_3757, plain, (v3_membered(k9_xtuple_0(esk11_1(X1))))).#trainneg
cnf(i_0_3769, plain, (k3_xcmplx_0(np__0,esk8_0)=np__0)).#trainneg
cnf(i_0_745, plain, (v1_xcmplx_0(k4_finseq_1(X1))|~v1_membered(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2794, plain, (v2_xxreal_0(esk8_0)|esk8_0=np__2|esk8_0=np__4|esk8_0=np__1|esk8_0=np__0|esk8_0=np__3|~m1_subset_1(esk8_0,k5_numbers))).#trainneg
cnf(i_0_2917, plain, (v1_finset_1(esk4_1(esk7_0))|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_2176, plain, (k3_xcmplx_0(X1,np__4)=k3_xcmplx_0(np__4,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2180, plain, (k2_xcmplx_0(X1,np__2)=k2_xcmplx_0(np__2,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_746, plain, (v1_int_1(k4_finseq_1(X1))|~v5_membered(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3439, negated_conjecture, (r1_xxreal_0(X1,esk4_1(esk1_0))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_3369, negated_conjecture, (esk3_1(esk1_0)=np__4|r1_xxreal_0(esk3_1(esk1_0),np__3)|~v7_ordinal1(esk3_1(esk1_0))|~r1_xxreal_0(esk3_1(esk1_0),np__4))).#trainneg
cnf(i_0_3712, plain, (k7_partfun1(k5_numbers,X1,X2)=np__2|r1_xxreal_0(k7_partfun1(k5_numbers,X1,X2),np__1)|~m1_subset_1(k7_partfun1(k5_numbers,X1,X2),k5_numbers)|~r1_xxreal_0(k7_partfun1(k5_numbers,X1,X2),np__2)|~m2_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_3382, negated_conjecture, (r1_xxreal_0(np__1,esk3_1(esk1_0))|r1_xxreal_0(esk3_1(esk1_0),np__1))).#trainneg
cnf(i_0_3497, plain, (r1_xxreal_0(X1,esk4_1(X2))|v1_xboole_0(X2)|v3_funct_1(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3373, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),np__4)|~r1_xxreal_0(esk3_1(esk1_0),np__3))).#trainneg
cnf(i_0_3457, plain, (r1_xxreal_0(esk3_1(X1),np__2)|v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(esk3_1(X1),np__1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3597, plain, (X1=k2_xcmplx_0(X2,np__1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(k2_xcmplx_0(X2,np__1))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k2_xcmplx_0(X2,np__1))|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3180, plain, (r1_xxreal_0(X1,np__3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk8_0)|~r1_xxreal_0(esk8_0,np__0))).#trainneg
cnf(i_0_3773, plain, (k3_xcmplx_0(esk8_0,np__1)=esk8_0)).#trainneg
cnf(i_0_4102, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)|~r1_xxreal_0(np__3,k3_finseq_1(esk1_0))|~r1_xxreal_0(np__2,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_749, plain, (v3_membered(k4_finseq_1(X1))|~v3_membered(k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4319, plain, (~r1_xxreal_0(esk8_0,np__2)|~r1_xxreal_0(np__4,esk8_0))).#trainneg
cnf(i_0_4489, plain, (~r1_xxreal_0(esk8_0,np__1)|~r1_xxreal_0(np__4,esk8_0))).#trainneg
cnf(i_0_750, plain, (v2_membered(k4_finseq_1(X1))|~v2_membered(k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_751, plain, (v1_membered(k4_finseq_1(X1))|~v1_membered(k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4173, negated_conjecture, (k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4)=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)|k1_card_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0)).#trainneg
cnf(i_0_4585, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|~r1_xxreal_0(np__4,esk4_1(esk1_0)))).#trainneg
cnf(i_0_4582, negated_conjecture, (esk4_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|~r1_xxreal_0(np__3,esk4_1(esk1_0)))).#trainneg
cnf(i_0_4518, plain, (k3_xcmplx_0(np__0,np__4)=np__0)).#trainneg
cnf(i_0_4257, negated_conjecture, (~r1_xxreal_0(np__4,esk3_1(esk1_0))|~r1_xxreal_0(esk3_1(esk1_0),np__2))).#trainneg
cnf(i_0_761, plain, (v1_relat_1(k9_xtuple_0(esk16_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))|v1_xboole_0(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)))).#trainneg
cnf(i_0_2181, plain, (k3_xcmplx_0(X1,np__2)=k3_xcmplx_0(np__2,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1349, plain, (r1_xxreal_0(X1,esk6_1(X2))|r1_xxreal_0(esk6_1(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(X2))).#trainneg
cnf(i_0_762, plain, (v1_relat_1(k9_xtuple_0(esk9_2(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_763, plain, (v1_relat_1(k9_xtuple_0(esk6_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))))).#trainneg
cnf(i_0_4471, plain, (v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(np__4,esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(esk3_1(X1),np__1))).#trainneg
cnf(i_0_4402, plain, (esk4_1(X1)=np__3|r1_xxreal_0(esk4_1(X1),np__2)|v1_xboole_0(X1)|v3_funct_1(X1)|r1_xxreal_0(np__4,esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4133, plain, (r1_xxreal_0(np__0,esk4_1(X1))|r1_xxreal_0(esk4_1(X1),np__0)|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4578, plain, (esk4_1(X1)=X2|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(esk4_1(X1))|~r1_xxreal_0(esk4_1(X1),X2)|~r1_xxreal_0(X2,np__1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4400, plain, (esk3_1(X1)=np__3|r1_xxreal_0(esk3_1(X1),np__2)|v1_xboole_0(X1)|v3_funct_1(X1)|r1_xxreal_0(np__4,esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4104, negated_conjecture, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|~r1_xxreal_0(np__3,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_4030, plain, (k7_partfun1(X1,X2,X3)=np__2|r1_xxreal_0(k7_partfun1(X1,X2,X3),np__1)|~m1_subset_1(k7_partfun1(X1,X2,X3),k5_numbers)|~r1_xxreal_0(k7_partfun1(X1,X2,X3),np__2)|~m2_finseq_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_765, plain, (r1_xxreal_0(np__1,X1)|v1_xboole_0(k1_relset_1(k5_numbers,X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_relset_1(k5_numbers,X2)))).#trainneg
cnf(i_0_1252, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(X3)|~m2_finseq_1(X3,X2)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_775, plain, (v1_topreal1(X1)|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(esk2_1(X1),np__1)))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,esk2_1(X1)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(esk2_1(X1)))).#trainneg
cnf(i_0_782, plain, (v7_ordinal1(k19_euclid(X1,X2))|~v6_membered(u1_struct_0(k15_euclid(np__2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4109, negated_conjecture, (v7_ordinal1(esk3_1(esk1_0))|~v1_finset_1(esk3_1(esk1_0)))).#trainneg
cnf(i_0_4122, negated_conjecture, (k3_xcmplx_0(np__1,esk3_1(esk1_0))=esk3_1(esk1_0))).#trainneg
cnf(i_0_4021, plain, (k7_partfun1(X1,X2,X3)=np__1|r1_xxreal_0(k7_partfun1(X1,X2,X3),np__0)|~m1_subset_1(k7_partfun1(X1,X2,X3),k5_numbers)|~r1_xxreal_0(k7_partfun1(X1,X2,X3),np__1)|~m2_finseq_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4180, plain, (esk4_1(X1)=np__3|esk4_1(X1)=np__4|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(esk4_1(X1))|~r1_xxreal_0(np__3,esk4_1(X1))|~r1_xxreal_0(esk4_1(X1),np__4)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4556, negated_conjecture, (esk3_1(esk1_0)=np__4|r1_xxreal_0(esk3_1(esk1_0),np__3))).#trainneg
cnf(i_0_4135, plain, (r1_xxreal_0(np__0,esk3_1(X1))|r1_xxreal_0(esk3_1(X1),np__0)|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4221, plain, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X3))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k1_nat_1(np__1,X2)=np__2|k1_nat_1(np__1,X2)=np__4|k1_nat_1(np__1,X2)=np__1|k1_nat_1(np__1,X2)=np__0|r1_xxreal_0(np__4,X2)|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X3))|k1_nat_1(np__1,X3)=np__2|k1_nat_1(np__1,X3)=np__4|k1_nat_1(np__1,X3)=np__1|k1_nat_1(np__1,X3)=np__0|r1_xxreal_0(np__4,X3)|~v7_ordinal1(X2)|~r1_xxreal_0(np__3,k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)|~v7_ordinal1(X3)|~r1_xxreal_0(np__1,X3))).#trainneg
cnf(i_0_2918, negated_conjecture, (v1_finset_1(esk4_1(esk1_0))|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_783, plain, (v1_xxreal_0(k19_euclid(X1,X2))|~v2_membered(u1_struct_0(k15_euclid(np__2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2358, plain, (r1_xxreal_0(np__4,k7_partfun1(X1,X2,X3))|r1_xxreal_0(k7_partfun1(X1,X2,X3),np__4)|~v2_membered(X1)|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_2419, plain, (r1_xxreal_0(np__2,k7_partfun1(X1,X2,X3))|r1_xxreal_0(k7_partfun1(X1,X2,X3),np__2)|~v2_membered(X1)|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_784, plain, (v3_ordinal1(k19_euclid(X1,X2))|~v3_ordinal1(u1_struct_0(k15_euclid(np__2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_786, plain, (v1_xreal_0(k19_euclid(X1,X2))|~v3_membered(u1_struct_0(k15_euclid(np__2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4039, plain, (k7_partfun1(X1,X2,X3)=np__4|r1_xxreal_0(k7_partfun1(X1,X2,X3),np__3)|~m1_subset_1(k7_partfun1(X1,X2,X3),k5_numbers)|~r1_xxreal_0(k7_partfun1(X1,X2,X3),np__4)|~m2_finseq_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4552, negated_conjecture, (r1_xxreal_0(X1,esk4_1(esk1_0))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,np__1))).#trainneg
cnf(i_0_5090, negated_conjecture, (r1_xxreal_0(np__1,esk4_1(esk1_0)))).#trainneg
cnf(i_0_4787, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__0|esk3_1(esk1_0)=np__3|~v7_ordinal1(esk3_1(esk1_0)))).#trainneg
cnf(i_0_3018, plain, (v1_int_1(esk4_1(esk7_0))|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_4648, negated_conjecture, (esk3_1(esk1_0)=np__4|~r1_xxreal_0(np__4,esk3_1(esk1_0)))).#trainneg
cnf(i_0_787, plain, (v1_xcmplx_0(k19_euclid(X1,X2))|~v1_membered(u1_struct_0(k15_euclid(np__2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4626, plain, (~r1_xxreal_0(esk8_0,np__0)|~r1_xxreal_0(np__4,esk8_0))).#trainneg
cnf(i_0_4842, plain, (v7_ordinal1(k10_finseq_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v6_membered(u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4741, plain, (r1_xxreal_0(esk6_1(X1),X2)|r1_xxreal_0(X2,esk6_1(X1))|~v1_xxreal_0(X2)|~v3_membered(X1))).#trainneg
cnf(i_0_788, plain, (v1_int_1(k19_euclid(X1,X2))|~v5_membered(u1_struct_0(k15_euclid(np__2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_790, plain, (v3_valued_0(k19_euclid(X1,X2))|~v7_ordinal1(np__2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4743, negated_conjecture, (esk4_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2|esk4_1(esk1_0)=np__3|~r1_xxreal_0(np__2,esk4_1(esk1_0)))).#trainneg
cnf(i_0_4651, negated_conjecture, (esk3_1(esk1_0)=np__3|esk3_1(esk1_0)=np__0|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__2)).#trainneg
cnf(i_0_4709, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2|esk3_1(esk1_0)=np__3|~r1_xxreal_0(np__3,esk3_1(esk1_0)))).#trainneg
cnf(i_0_5156, negated_conjecture, (r1_xxreal_0(np__0,esk3_1(esk1_0)))).#trainneg
cnf(i_0_4707, negated_conjecture, (esk3_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2|esk3_1(esk1_0)=np__3|~r1_xxreal_0(np__2,esk3_1(esk1_0)))).#trainneg
cnf(i_0_4876, negated_conjecture, (esk3_1(esk1_0)=np__3|esk3_1(esk1_0)=np__4|~r1_xxreal_0(np__3,esk3_1(esk1_0)))).#trainneg
cnf(i_0_1266, plain, (v1_finset_1(X1)|~v5_finset_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3019, negated_conjecture, (v1_int_1(esk4_1(esk1_0))|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_4255, negated_conjecture, (esk3_1(esk1_0)=np__3|~r1_xxreal_0(np__3,esk3_1(esk1_0))|~r1_xxreal_0(esk3_1(esk1_0),np__2))).#trainneg
cnf(i_0_792, plain, (m1_subset_1(k10_finseq_1(X1,X2),u1_struct_0(k15_euclid(np__2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3158, plain, (v1_finset_1(esk3_1(esk7_0))|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_3159, negated_conjecture, (v1_finset_1(esk3_1(esk1_0))|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_4994, plain, (r1_xxreal_0(k7_partfun1(X1,X2,X3),np__4)|r1_xxreal_0(np__4,k7_partfun1(X1,X2,X3))|~m2_finseq_1(X2,X1)|~v3_membered(X1))).#trainneg
cnf(i_0_5001, plain, (r1_xxreal_0(k7_partfun1(X1,X2,X3),np__2)|r1_xxreal_0(np__2,k7_partfun1(X1,X2,X3))|~m2_finseq_1(X2,X1)|~v3_membered(X1))).#trainneg
cnf(i_0_793, plain, (v7_ordinal1(k1_relset_1(X1,X2))|~v6_membered(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_794, plain, (v1_xxreal_0(k1_relset_1(X1,X2))|~v2_membered(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_5068, negated_conjecture, (r1_xxreal_0(np__0,esk4_1(esk1_0)))).#trainneg
cnf(i_0_4720, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2|esk3_1(esk1_0)=np__3|~r1_xxreal_0(np__4,esk3_1(esk1_0)))).#trainneg
cnf(i_0_4751, plain, (v1_xboole_0(X1)|v3_funct_1(X1)|r1_xxreal_0(np__2,esk3_1(X1))|~r1_xxreal_0(np__4,esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4649, negated_conjecture, (r1_xxreal_0(X1,np__4)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk3_1(esk1_0)))).#trainneg
cnf(i_0_4719, negated_conjecture, (r1_xxreal_0(np__3,esk3_1(esk1_0))|~r1_xxreal_0(np__4,esk3_1(esk1_0)))).#trainneg
cnf(i_0_1267, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3219, plain, (m1_subset_1(esk8_0,k5_numbers)|~v1_xboole_0(esk8_0))).#trainneg
cnf(i_0_4013, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|~r1_xxreal_0(np__4,esk3_1(esk1_0)))).#trainneg
cnf(i_0_795, plain, (v3_ordinal1(k1_relset_1(X1,X2))|~v3_ordinal1(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_3229, plain, (v1_int_1(esk3_1(esk7_0))|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_3230, negated_conjecture, (v1_int_1(esk3_1(esk1_0))|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_2359, plain, (r1_xxreal_0(X1,k7_partfun1(X2,X3,X4))|r1_xxreal_0(k7_partfun1(X2,X3,X4),X1)|~v1_xxreal_0(X1)|~v2_membered(X2)|~m2_finseq_1(X3,X2))).#trainneg
cnf(i_0_4976, plain, (r1_xxreal_0(np__4,k19_euclid(X1,X2))|r1_xxreal_0(k19_euclid(X1,X2),np__4)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_membered(u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_796, plain, (v1_finset_1(k1_relset_1(X1,X2))|~v5_finset_1(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_797, plain, (v1_xreal_0(k1_relset_1(X1,X2))|~v3_membered(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_4785, negated_conjecture, (k2_xcmplx_0(np__0,esk3_1(esk1_0))=esk3_1(esk1_0)|~m1_subset_1(esk3_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_4551, negated_conjecture, (esk4_1(esk1_0)=X1|~v1_xxreal_0(X1)|~r1_xxreal_0(esk4_1(esk1_0),X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_4579, plain, (r1_xxreal_0(X1,esk4_1(X2))|v1_xboole_0(X2)|v3_funct_1(X2)|~v1_xxreal_0(esk4_1(X2))|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X3,np__1)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1268, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1691, plain, (v1_finset_1(esk6_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_798, plain, (v1_xcmplx_0(k1_relset_1(X1,X2))|~v1_membered(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_4340, plain, (v7_ordinal1(esk8_0)|~v1_finset_1(esk8_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_3643, plain, (v7_ordinal1(esk4_1(esk7_0))|~v1_finset_1(esk4_1(esk7_0)))).#trainneg
cnf(i_0_4979, plain, (r1_xxreal_0(np__2,k19_euclid(X1,X2))|r1_xxreal_0(k19_euclid(X1,X2),np__2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_membered(u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4983, plain, (r1_xxreal_0(X1,k19_euclid(X2,X3))|r1_xxreal_0(k19_euclid(X2,X3),X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v2_membered(u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_799, plain, (v1_int_1(k1_relset_1(X1,X2))|~v5_membered(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_800, plain, (v5_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_4463, plain, (r1_xxreal_0(np__2,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|esk3_1(X1)=np__1|~v1_xxreal_0(k3_finseq_1(X1))|~v1_xxreal_0(esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4177, plain, (v1_xboole_0(X1)|v3_funct_1(X1)|esk4_1(X1)=np__2|~r1_xxreal_0(np__2,esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(esk4_1(X1),np__2))).#trainneg
cnf(i_0_4883, plain, (r1_xxreal_0(np__0,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(k3_finseq_1(X1))|~v1_xxreal_0(esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4247, negated_conjecture, (esk3_1(esk1_0)=X1|~v1_xxreal_0(X1)|~r1_xxreal_0(esk3_1(esk1_0),X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_4867, negated_conjecture, (esk4_1(esk1_0)=np__1|k1_card_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|np__4!=np__1)).#trainneg
cnf(i_0_5622, plain, (v1_xreal_0(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1123, plain, (k1_card_1(X1)=esk9_2(X2,X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_801, plain, (v1_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_3765, plain, (esk8_0=np__3|r1_xxreal_0(np__2,esk8_0)|~r1_xxreal_0(np__3,esk8_0))).#trainneg
cnf(i_0_5383, plain, (r1_xxreal_0(X1,esk6_1(X2))|r1_xxreal_0(esk6_1(X2),X1)|~v1_xxreal_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_802, plain, (v1_xboole_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_xboole_0(X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_803, plain, (v1_xboole_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_xboole_0(X1)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_5838, plain, (r1_xxreal_0(np__0,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1269, plain, (v1_int_1(X1)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1645, plain, (esk11_1(X1)=esk9_2(X2,X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_804, plain, (v4_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_4110, plain, (esk8_0=np__3|r1_xxreal_0(np__1,esk8_0)|~r1_xxreal_0(np__3,esk8_0))).#trainneg
cnf(i_0_5654, plain, (r1_xxreal_0(k7_partfun1(X1,X2,X3),X4)|r1_xxreal_0(X4,k7_partfun1(X1,X2,X3))|~v1_xxreal_0(X4)|~m2_finseq_1(X2,X1)|~v3_membered(X1))).#trainneg
cnf(i_0_5808, plain, (r1_xxreal_0(k19_euclid(X1,X2),X3)|r1_xxreal_0(X3,k19_euclid(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xxreal_0(X3)|~v3_membered(u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_805, plain, (v1_relat_1(k9_xtuple_0(k1_relset_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)))|~v4_relat_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))|~v1_relat_1(X4))).#trainneg
cnf(i_0_807, plain, (v1_finset_1(k1_relset_1(X1,X2))|~v1_finset_1(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_5671, plain, (esk3_1(X1)=np__3|esk3_1(X1)=np__1|esk3_1(X1)=np__2|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(esk3_1(X1))|~r1_xxreal_0(esk3_1(X1),np__3)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5343, negated_conjecture, (esk4_1(esk1_0)=np__1|~r1_xxreal_0(esk4_1(esk1_0),np__1))).#trainneg
cnf(i_0_5687, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|~r1_xxreal_0(esk4_1(esk1_0),np__2))).#trainneg
cnf(i_0_4413, plain, (esk3_1(X1)=X2|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(esk3_1(X1))|~r1_xxreal_0(esk3_1(X1),X2)|~r1_xxreal_0(X2,np__1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4141, plain, (r1_xxreal_0(np__0,k7_partfun1(X1,X2,X3))|r1_xxreal_0(k7_partfun1(X1,X2,X3),np__0)|~v2_membered(X1)|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_3217, plain, (v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1126, plain, (k1_card_1(X1)=k1_card_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5350, plain, (v1_xreal_0(esk4_1(esk7_0))|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_808, plain, (v3_membered(k1_relset_1(X1,X2))|~v3_membered(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_4333, plain, (esk8_0=np__3|r1_xxreal_0(np__0,esk8_0)|~r1_xxreal_0(np__3,esk8_0))).#trainneg
cnf(i_0_5499, negated_conjecture, (v1_xreal_0(esk4_1(esk1_0))|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_5521, negated_conjecture, (v7_ordinal1(esk3_1(esk1_0))|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_2028, plain, (m1_subset_1(X1,X2)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,k3_finseq_1(X3))|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_809, plain, (v2_membered(k1_relset_1(X1,X2))|~v2_membered(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_810, plain, (v1_membered(k1_relset_1(X1,X2))|~v1_membered(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_4296, negated_conjecture, (k1_card_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|np__4!=np__1)).#trainneg
cnf(i_0_5196, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|np__3!=np__2)).#trainneg
cnf(i_0_5197, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|np__3!=np__1)).#trainneg
cnf(i_0_4126, negated_conjecture, (k3_xcmplx_0(esk3_1(esk1_0),np__0)=np__0)).#trainneg
cnf(i_0_3280, plain, (v3_ordinal1(X1)|~v1_card_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1646, plain, (esk11_1(X1)=k1_card_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5644, plain, (v1_xreal_0(esk3_1(esk7_0))|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_5296, negated_conjecture, (esk3_1(esk1_0)=np__1|~r1_xxreal_0(esk3_1(esk1_0),np__1))).#trainneg
cnf(i_0_813, plain, (m1_subset_1(k9_xtuple_0(X2),k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_5645, negated_conjecture, (v1_xreal_0(esk3_1(esk1_0))|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_5806, plain, (v7_ordinal1(esk4_1(esk7_0))|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_3463, plain, (r1_xxreal_0(X1,esk3_1(X2))|r1_xxreal_0(esk3_1(X2),X1)|v1_xboole_0(X2)|v3_funct_1(X2)|~v1_xxreal_0(X1)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_822, plain, (v2_xxreal_0(X1)|r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(k1_nat_1(X2,np__1))|~v2_xxreal_0(k1_nat_1(X2,np__1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_829, plain, (v2_xxreal_0(k3_finseq_1(X1))|v1_topreal1(X1)|~v1_xreal_0(k3_finseq_1(X1))|~v1_xreal_0(k1_nat_1(esk2_1(X1),np__1))|~v2_xxreal_0(k1_nat_1(esk2_1(X1),np__1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4785, negated_conjecture, (k2_xcmplx_0(np__0,esk3_1(esk1_0))=esk3_1(esk1_0))).#trainneg
cnf(i_0_6516, negated_conjecture, (esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|~r1_xxreal_0(esk3_1(esk1_0),np__2))).#trainneg
cnf(i_0_5911, negated_conjecture, (v1_xreal_0(esk3_1(esk1_0)))).#trainneg
cnf(i_0_6550, plain, (~r1_xxreal_0(np__3,esk4_1(esk7_0))|~r1_xxreal_0(esk4_1(esk7_0),np__2)|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_831, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|v1_topreal1(X2)|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(k1_nat_1(esk2_1(X2),np__1))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k1_nat_1(esk2_1(X2),np__1))|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_6542, plain, (~r1_xxreal_0(np__1,esk4_1(esk7_0))|~r1_xxreal_0(esk4_1(esk7_0),np__0)|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_6548, plain, (~r1_xxreal_0(np__4,esk4_1(esk7_0))|~r1_xxreal_0(esk4_1(esk7_0),np__3)|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_6001, plain, (v1_card_1(esk9_2(X2,X3))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6117, plain, (v1_funct_1(esk9_2(X2,X3))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_834, plain, (r1_xxreal_0(k1_nat_1(esk2_1(X1),np__1),k1_card_1(X1))|v1_topreal1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_838, plain, (m1_subset_1(k7_partfun1(X1,esk12_1(X1),X2),X1)|v1_xboole_0(X1)|~v1_funct_1(esk12_1(X1))|~v1_relat_1(esk12_1(X1)))).#trainneg
cnf(i_0_6101, negated_conjecture, (r1_xxreal_0(np__0,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_6067, negated_conjecture, (r1_xxreal_0(esk4_1(esk1_0),np__4)|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__3)).#trainneg
cnf(i_0_5919, negated_conjecture, (esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|np__3!=np__2)).#trainneg
cnf(i_0_6168, plain, (esk3_1(X1)=np__2|esk3_1(X1)=np__1|esk3_1(X1)=np__3|v1_xboole_0(X1)|v3_funct_1(X1)|esk3_1(X1)=np__4|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(esk3_1(X1),np__4))).#trainneg
cnf(i_0_6175, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|k1_card_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|np__4!=np__2)).#trainneg
cnf(i_0_5917, negated_conjecture, (esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|np__3!=np__4)).#trainneg
cnf(i_0_6549, plain, (~r1_xxreal_0(np__2,esk4_1(esk7_0))|~r1_xxreal_0(esk4_1(esk7_0),np__1)|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_839, plain, (m1_subset_1(k7_partfun1(X1,esk9_2(X2,X1),X3),X1)|~v1_funct_1(esk9_2(X2,X1)))).#trainneg
cnf(i_0_3005, plain, (r1_xxreal_0(esk8_0,np__4)|r1_xxreal_0(np__3,esk8_0))).#trainneg
cnf(i_0_6119, plain, (v1_finseq_1(esk9_2(X2,X3))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_840, plain, (m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk7_0,X1),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_845, plain, (m1_subset_1(esk15_2(X1,X2),X2)|v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_6364, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__3|np__4!=np__1)).#trainneg
cnf(i_0_6170, plain, (esk3_1(X1)=np__2|esk3_1(X1)=np__1|esk3_1(X1)=np__3|v1_xboole_0(X1)|v3_funct_1(X1)|r1_xxreal_0(np__4,esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_6173, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|k1_card_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|np__4!=np__0)).#trainneg
cnf(i_0_6181, plain, (esk3_1(X1)=np__2|esk3_1(X1)=np__1|r1_xxreal_0(np__3,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6815, plain, (v1_card_1(esk9_2(X1,X2)))).#trainneg
cnf(i_0_7267, negated_conjecture, (k1_card_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__3|np__4!=np__1)).#trainneg
cnf(i_0_3371, negated_conjecture, (esk3_1(esk1_0)=np__3|r1_xxreal_0(esk3_1(esk1_0),np__2)|~r1_xxreal_0(esk3_1(esk1_0),np__3))).#trainneg
cnf(i_0_846, plain, (v1_int_1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_int_1(X3)|~v1_int_1(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7207, plain, (esk8_0=np__3|esk8_0=np__0|esk8_0=np__1|esk8_0=np__4|esk8_0=np__2|r1_xxreal_0(np__3,esk8_0)|~m1_subset_1(esk8_0,k5_numbers))).#trainneg
cnf(i_0_847, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),X4)=k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_848, plain, (k2_xcmplx_0(np__0,k2_xcmplx_0(np__0,X1))=k2_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_7283, negated_conjecture, (r1_xxreal_0(esk4_1(esk1_0),np__2)|k1_card_1(esk1_0)=np__3|np__4!=np__1)).#trainneg
cnf(i_0_6931, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|np__4!=np__1|np__3!=np__2)).#trainneg
cnf(i_0_6825, plain, (v1_funct_1(esk9_2(X1,X2)))).#trainneg
cnf(i_0_7268, negated_conjecture, (k3_finseq_1(np__2)=np__2|k1_card_1(esk1_0)=np__3|np__4!=np__1)).#trainneg
cnf(i_0_6934, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|np__3!=np__2)).#trainneg
cnf(i_0_849, plain, (k2_xcmplx_0(np__0,k2_xcmplx_0(np__1,X1))=k2_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_7205, plain, (esk8_0=np__4|r1_xxreal_0(np__3,esk8_0)|~r1_xxreal_0(np__4,esk8_0))).#trainneg
cnf(i_0_6699, negated_conjecture, (~r1_xxreal_0(np__4,esk3_1(esk1_0))|~r1_xxreal_0(esk3_1(esk1_0),np__3))).#trainneg
cnf(i_0_6700, negated_conjecture, (~r1_xxreal_0(np__2,esk3_1(esk1_0))|~r1_xxreal_0(esk3_1(esk1_0),np__1))).#trainneg
cnf(i_0_850, plain, (k2_xcmplx_0(np__0,k2_xcmplx_0(np__2,X1))=k2_xcmplx_0(np__2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_851, plain, (k2_xcmplx_0(np__3,k2_xcmplx_0(np__1,X1))=k2_xcmplx_0(np__4,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__3))).#trainneg
cnf(i_0_6930, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|np__4!=np__2|np__3!=np__2)).#trainneg
cnf(i_0_6968, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|~v1_xxreal_0(k3_finseq_1(esk1_0))|np__3!=np__2)).#trainneg
cnf(i_0_7282, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),np__2)|k1_card_1(esk1_0)=np__3|np__4!=np__1)).#trainneg
cnf(i_0_6697, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),np__1)|r1_xxreal_0(np__2,esk3_1(esk1_0)))).#trainneg
cnf(i_0_7538, negated_conjecture, (r1_xxreal_0(X1,np__4)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk4_1(esk1_0)))).#trainneg
cnf(i_0_7544, negated_conjecture, (esk4_1(esk1_0)=np__3|esk4_1(esk1_0)=np__0|esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__4|esk4_1(esk1_0)=np__2)).#trainneg
cnf(i_0_7209, plain, (v1_finseq_1(esk9_2(X1,X2)))).#trainneg
cnf(i_0_7555, negated_conjecture, (v1_card_1(np__3)|k1_card_1(esk1_0)=np__2|np__4!=np__1)).#trainneg
cnf(i_0_839, plain, (m1_subset_1(k7_partfun1(X1,esk9_2(X2,X1),X3),X1))).#trainneg
cnf(i_0_852, plain, (k2_xcmplx_0(np__1,k2_xcmplx_0(np__0,X1))=k2_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_7617, plain, (esk8_0=np__3|esk8_0=np__2|esk8_0=np__4|esk8_0=np__1|esk8_0=np__0|~r1_xxreal_0(esk8_0,np__3)|~m1_subset_1(esk8_0,k5_numbers))).#trainneg
cnf(i_0_6701, negated_conjecture, (~r1_xxreal_0(np__3,esk3_1(esk1_0))|~r1_xxreal_0(esk3_1(esk1_0),np__2))).#trainneg
cnf(i_0_8009, negated_conjecture, (~r1_xxreal_0(np__4,esk3_1(esk1_0))|~r1_xxreal_0(esk3_1(esk1_0),np__1))).#trainneg
cnf(i_0_6491, plain, (v1_card_1(esk11_1(X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_853, plain, (k2_xcmplx_0(np__1,k2_xcmplx_0(np__3,X1))=k2_xcmplx_0(np__4,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__3)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_854, plain, (k2_xcmplx_0(np__1,k2_xcmplx_0(np__1,X1))=k2_xcmplx_0(np__2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_8100, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__3|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|~v1_xxreal_0(k3_finseq_1(esk1_0))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_7443, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|~v1_xxreal_0(k3_finseq_1(esk1_0))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__2))).#trainneg
cnf(i_0_7542, negated_conjecture, (esk4_1(esk1_0)=np__4|esk4_1(esk1_0)=np__3|~r1_xxreal_0(np__3,esk4_1(esk1_0)))).#trainneg
cnf(i_0_6218, plain, (v1_xxreal_0(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_6118, plain, (m2_finseq_1(esk9_2(X2,X3),X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7993, plain, (v2_xxreal_0(esk8_0)|esk8_0=np__4|~r1_xxreal_0(np__4,esk8_0))).#trainneg
cnf(i_0_7091, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|np__3!=np__4)).#trainneg
cnf(i_0_855, plain, (k2_xcmplx_0(np__1,k2_xcmplx_0(np__2,X1))=k2_xcmplx_0(np__3,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_3660, negated_conjecture, (v7_ordinal1(esk4_1(esk1_0))|~v1_finset_1(esk4_1(esk1_0)))).#trainneg
cnf(i_0_6504, plain, (v1_funct_1(k1_card_1(X2))|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6506, plain, (v1_finseq_1(k1_card_1(X2))|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_856, plain, (k2_xcmplx_0(np__2,k2_xcmplx_0(np__0,X1))=k2_xcmplx_0(np__2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_857, plain, (k2_xcmplx_0(np__2,k2_xcmplx_0(np__1,X1))=k2_xcmplx_0(np__3,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_8061, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|np__4!=np__1|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_7852, negated_conjecture, (esk3_1(esk1_0)=np__2|np__4!=np__1|np__3!=np__2)).#trainneg
cnf(i_0_7537, negated_conjecture, (esk4_1(esk1_0)=np__4|~r1_xxreal_0(np__4,esk4_1(esk1_0)))).#trainneg
cnf(i_0_7264, negated_conjecture, (k7_partfun1(X1,esk1_0,np__1)=k7_partfun1(X1,esk1_0,np__2)|v1_xboole_0(X1)|k1_card_1(esk1_0)=np__3|~m2_finseq_1(esk1_0,X1)|~v1_finseq_6(esk1_0,X1)|np__4!=np__1)).#trainneg
cnf(i_0_6903, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|~v1_xxreal_0(k3_finseq_1(esk1_0))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__0))).#trainneg
cnf(i_0_7319, plain, (r1_xxreal_0(np__4,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|esk3_1(X1)=np__3|esk3_1(X1)=np__1|esk3_1(X1)=np__2|~v1_xxreal_0(k3_finseq_1(X1))|~v1_xxreal_0(esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5983, plain, (v1_relat_1(k1_card_1(X3))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_7552, negated_conjecture, (v1_finset_1(np__3)|k1_card_1(esk1_0)=np__2|np__4!=np__1)).#trainneg
cnf(i_0_858, plain, (k2_xcmplx_0(np__2,k2_xcmplx_0(np__2,X1))=k2_xcmplx_0(np__4,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_859, plain, (v1_int_1(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)))|~v1_int_1(X3)|~v1_int_1(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_860, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)),X4)=k3_xcmplx_0(k3_xcmplx_0(X1,X2),k3_xcmplx_0(X3,X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8539, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(np__3,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_8353, negated_conjecture, (esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__3|esk4_1(esk1_0)=np__4|~r1_xxreal_0(np__2,esk4_1(esk1_0)))).#trainneg
cnf(i_0_8535, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__2|~r1_xxreal_0(k3_finseq_1(esk1_0),np__2))).#trainneg
cnf(i_0_8532, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__4|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_8613, plain, (r1_xxreal_0(X1,np__4)|~r1_xxreal_0(X1,np__2)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_8628, plain, (r1_xxreal_0(X1,np__4)|~r1_xxreal_0(X1,np__1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_8948, negated_conjecture, (v1_finset_1(np__2)|v1_finset_1(np__3)|np__4!=np__1)).#trainneg
cnf(i_0_8836, negated_conjecture, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__4!=np__1|np__3!=np__2)).#trainneg
cnf(i_0_861, plain, (k3_xcmplx_0(np__0,k3_xcmplx_0(np__0,X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_862, plain, (k3_xcmplx_0(np__0,k3_xcmplx_0(np__1,X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_863, plain, (k3_xcmplx_0(np__0,k3_xcmplx_0(np__2,X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_8354, negated_conjecture, (r1_xxreal_0(X1,np__2)|esk4_1(esk1_0)=np__3|esk4_1(esk1_0)=np__4|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk4_1(esk1_0)))).#trainneg
cnf(i_0_8925, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|~r1_xxreal_0(k3_finseq_1(esk1_0),np__0))).#trainneg
cnf(i_0_8283, negated_conjecture, (r1_xxreal_0(k1_nat_1(np__1,X1),np__4)|r1_xxreal_0(esk4_1(esk1_0),X1)|~v1_xxreal_0(k1_nat_1(np__1,X1))|~v7_ordinal1(esk4_1(esk1_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8938, plain, (esk3_1(X1)=np__2|esk3_1(X1)=np__1|esk3_1(X1)=np__3|r1_xxreal_0(np__4,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8284, negated_conjecture, (r1_xxreal_0(k1_nat_1(X1,np__1),np__4)|r1_xxreal_0(esk4_1(esk1_0),X1)|~v1_xxreal_0(k1_nat_1(X1,np__1))|~v7_ordinal1(esk4_1(esk1_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8279, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,np__1),np__4)|r1_xxreal_0(esk4_1(esk1_0),X1)|~v1_xxreal_0(k2_xcmplx_0(X1,np__1))|~v7_ordinal1(esk4_1(esk1_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8753, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__3|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3)|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0))|~r1_xxreal_0(np__3,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_8629, plain, (r1_xxreal_0(X1,np__4)|~r1_xxreal_0(X1,np__3)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_8630, plain, (r1_xxreal_0(X1,np__4)|~r1_xxreal_0(X1,np__0)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_8951, negated_conjecture, (v1_card_1(np__2)|v1_finset_1(np__3)|np__4!=np__1)).#trainneg
cnf(i_0_864, plain, (k3_xcmplx_0(np__3,k3_xcmplx_0(np__1,X1))=k3_xcmplx_0(np__3,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__3))).#trainneg
cnf(i_0_865, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(np__0,X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_866, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(np__1,X1))=k3_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_8278, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(np__1,X1),np__4)|r1_xxreal_0(esk4_1(esk1_0),X1)|~v1_xxreal_0(k2_xcmplx_0(np__1,X1))|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(esk4_1(esk1_0)))).#trainneg
cnf(i_0_8280, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,np__1),np__4)|r1_xxreal_0(esk4_1(esk1_0),X1)|~v1_xxreal_0(k2_xcmplx_0(X1,np__1))|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(esk4_1(esk1_0)))).#trainneg
cnf(i_0_4645, plain, (v2_membered(k9_xtuple_0(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3455, plain, (r1_xxreal_0(esk3_1(X1),np__3)|v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(esk3_1(X1),np__2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_8917, negated_conjecture, (k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)|k1_card_1(esk1_0)=np__3|np__4!=np__1)).#trainneg
cnf(i_0_8615, plain, (r1_xxreal_0(X1,np__4)|~r1_xxreal_0(esk8_0,np__2)|~r1_xxreal_0(X1,esk8_0)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_8459, plain, (v1_card_1(esk11_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_867, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(np__2,X1))=k3_xcmplx_0(np__2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_868, plain, (k3_xcmplx_0(np__4,k3_xcmplx_0(np__1,X1))=k3_xcmplx_0(np__4,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__4))).#trainneg
cnf(i_0_869, plain, (k3_xcmplx_0(np__2,k3_xcmplx_0(np__0,X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_3383, negated_conjecture, (r1_xxreal_0(X1,esk3_1(esk1_0))|r1_xxreal_0(esk3_1(esk1_0),X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_5484, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__3|esk3_1(esk1_0)=np__4|~r1_xxreal_0(np__2,esk3_1(esk1_0)))).#trainneg
cnf(i_0_4223, plain, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k1_nat_1(np__1,X2)=np__2|k1_nat_1(np__1,X2)=np__4|k1_nat_1(np__1,X2)=np__1|k1_nat_1(np__1,X2)=np__0|r1_xxreal_0(np__4,X2)|~r1_xxreal_0(np__3,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(np__1,X2))).#trainneg
cnf(i_0_3975, plain, (v1_xboole_0(X1)|k2_xcmplx_0(np__1,X2)=np__0|k2_xcmplx_0(np__1,X2)=np__1|k2_xcmplx_0(np__1,X2)=np__4|k2_xcmplx_0(np__1,X2)=np__2|r1_xxreal_0(np__4,X2)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~m1_subset_1(X2,k5_numbers)|~r2_hidden(np__3,k4_finseq_1(X1))|~r2_hidden(X2,k4_finseq_1(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1))).#trainneg
cnf(i_0_7010, negated_conjecture, (k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__3|esk3_1(esk1_0)=np__2|np__4!=np__2)).#trainneg
cnf(i_0_8616, plain, (r1_xxreal_0(X1,np__4)|~r1_xxreal_0(esk8_0,np__1)|~r1_xxreal_0(X1,esk8_0)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_8690, plain, (v1_funct_1(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_870, plain, (k3_xcmplx_0(np__2,k3_xcmplx_0(np__1,X1))=k3_xcmplx_0(np__2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_871, plain, (k3_xcmplx_0(np__2,k3_xcmplx_0(np__2,X1))=k3_xcmplx_0(np__4,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_872, plain, (m1_subset_1(esk15_2(X1,X2),X1)|v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_9126, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__3|esk3_1(esk1_0)=np__2|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_9430, negated_conjecture, (k1_card_1(esk1_0)=np__3|~r1_xxreal_0(np__2,k3_finseq_1(esk1_0))|np__4!=np__1)).#trainneg
cnf(i_0_9313, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__3|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_9100, negated_conjecture, (k1_card_1(esk1_0)=np__3|esk4_1(esk1_0)=np__4|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|np__4!=np__1)).#trainneg
cnf(i_0_9117, negated_conjecture, (np__3=np__0|np__3=np__1|np__3=np__4|np__3=np__2|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0)).#trainneg
cnf(i_0_8617, plain, (r1_xxreal_0(X1,np__4)|~r1_xxreal_0(esk8_0,np__0)|~r1_xxreal_0(X1,esk8_0)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_8700, plain, (v1_finseq_1(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_9523, plain, (esk4_1(esk7_0)=np__1|r1_xxreal_0(np__2,k3_finseq_1(esk7_0)))).#trainneg
cnf(i_0_9590, negated_conjecture, (k1_card_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__4!=np__2)).#trainneg
cnf(i_0_886, plain, (m1_subset_1(esk3_1(X1),k4_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_892, plain, (v1_finset_1(esk9_2(X1,X2))|~v1_funct_2(esk9_2(X1,X2),X1,X2)|~v1_finset_1(X1)|~v1_funct_1(esk9_2(X1,X2)))).#trainneg
cnf(i_0_9437, negated_conjecture, (k1_card_1(esk1_0)=np__3|np__4!=np__1|np__3!=np__2)).#trainneg
cnf(i_0_9058, negated_conjecture, (esk4_1(esk1_0)=np__4|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|np__3!=np__1)).#trainneg
cnf(i_0_9207, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__4!=np__1|np__3!=np__2|~r1_xxreal_0(np__2,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_9513, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),np__2)|r1_xxreal_0(np__2,esk3_1(esk1_0)))).#trainneg
cnf(i_0_9509, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),k7_partfun1(X1,X2,X3))|r1_xxreal_0(k7_partfun1(X1,X2,X3),esk3_1(esk1_0))|~v2_membered(X1)|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_9622, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__1|~v1_xxreal_0(k3_finseq_1(esk1_0))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__2)|np__4!=np__2)).#trainneg
cnf(i_0_9056, negated_conjecture, (esk4_1(esk1_0)=np__4|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|np__3!=np__4)).#trainneg
cnf(i_0_8612, plain, (r1_xxreal_0(X1,X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_8363, plain, (v1_finset_1(esk9_2(X1,X2)))).#trainneg
cnf(i_0_893, plain, (v1_finset_1(esk6_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_funct_2(esk6_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1,X2)|~v1_finset_1(X1)|~v1_funct_1(esk6_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_8688, negated_conjecture, (v7_ordinal1(esk4_1(esk1_0))|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_894, plain, (v1_finset_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_funct_2(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v1_finset_1(X1)|~v1_funct_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_9057, negated_conjecture, (esk4_1(esk1_0)=np__4|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|np__3!=np__2)).#trainneg
cnf(i_0_9531, plain, (k10_finseq_1(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)),k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)))=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k1_nat_1(np__1,X2)=np__0|k1_nat_1(np__1,X2)=np__1|k1_nat_1(np__1,X2)=np__4|k1_nat_1(np__1,X2)=np__2|r1_xxreal_0(np__4,X2)|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(np__3,k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_8147, negated_conjecture, (esk3_1(esk1_0)=np__2|np__4!=np__2|np__3!=np__2)).#trainneg
cnf(i_0_3893, plain, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X3))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k1_nat_1(np__1,X2)=np__0|k1_nat_1(np__1,X2)=np__1|k1_nat_1(np__1,X2)=np__4|k1_nat_1(np__1,X2)=np__2|r1_xxreal_0(np__4,X2)|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X3))|k1_nat_1(np__1,X3)=np__0|k1_nat_1(np__1,X3)=np__1|k1_nat_1(np__1,X3)=np__4|k1_nat_1(np__1,X3)=np__2|r1_xxreal_0(np__4,X3)|~m1_subset_1(X2,k5_numbers)|~r1_xxreal_0(np__3,k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)|~m1_subset_1(X3,k5_numbers)|~r1_xxreal_0(np__1,X3))).#trainneg
cnf(i_0_9937, negated_conjecture, (k1_card_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|np__3=np__2|np__3=np__4|np__3=np__1|np__3=np__0)).#trainneg
cnf(i_0_10208, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__4|k1_card_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_8621, plain, (r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X1,np__0)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_9746, negated_conjecture, (v1_card_1(np__3)|k1_card_1(esk1_0)=np__4)).#trainneg
cnf(i_0_10458, negated_conjecture, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__4!=np__2|np__3!=np__2)).#trainneg
cnf(i_0_908, plain, (m1_subset_1(esk4_1(X1),k4_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_7994, plain, (esk8_0=np__3|esk8_0=np__4|~r1_xxreal_0(esk8_0,np__3)|~r1_xxreal_0(np__4,esk8_0))).#trainneg
cnf(i_0_10346, negated_conjecture, (~r1_xxreal_0(esk4_1(esk1_0),np__0)|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_911, plain, (v1_int_1(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2))|~v1_int_1(k3_xcmplx_0(X3,X2))|~v1_int_1(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_912, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,X2),k2_xcmplx_0(k3_xcmplx_0(X3,X2),X4))=k2_xcmplx_0(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X4)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k3_xcmplx_0(X3,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9932, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|np__3=np__2|np__3=np__4|np__3=np__1|np__3=np__0|np__4!=np__1)).#trainneg
cnf(i_0_9845, negated_conjecture, (k1_card_1(esk1_0)=np__3|np__4!=np__1)).#trainneg
cnf(i_0_10092, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|np__4!=np__1|np__3!=np__1)).#trainneg
cnf(i_0_8631, plain, (r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X1,np__2)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_9743, negated_conjecture, (v1_finset_1(np__3)|k1_card_1(esk1_0)=np__4)).#trainneg
cnf(i_0_9767, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__4|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_913, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,X2),k3_xcmplx_0(X3,k3_xcmplx_0(X4,X2)))=k3_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(X3,X4)),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k3_xcmplx_0(X3,X4))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_7997, plain, (esk8_0=np__3|r1_xxreal_0(np__1,esk8_0)|esk8_0=np__4|~r1_xxreal_0(np__4,esk8_0))).#trainneg
cnf(i_0_6228, plain, (v1_finset_1(k1_card_1(X2))|~v1_finset_1(X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_914, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,np__0),np__0)=k2_xcmplx_0(k3_xcmplx_0(X1,np__0),np__0)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_915, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,np__0),np__1)=k2_xcmplx_0(k3_xcmplx_0(X1,np__1),np__0)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9799, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__4|~v1_xxreal_0(k3_finseq_1(esk1_0))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__2)|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_10657, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k1_card_1(esk1_0)=np__4|np__3=np__0|np__3=np__1|np__3=np__4|np__3=np__2)).#trainneg
cnf(i_0_11256, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__2|~r1_xxreal_0(k3_finseq_1(esk1_0),np__2))).#trainneg
cnf(i_0_10365, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|np__4!=np__1|np__3!=np__2)).#trainneg
cnf(i_0_8633, plain, (r1_xxreal_0(X1,np__2)|~r1_xxreal_0(X1,np__1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_11189, plain, (v1_finset_1(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_11061, negated_conjecture, (v1_finset_1(np__4)|v1_finset_1(np__3))).#trainneg
cnf(i_0_10223, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|np__3!=np__4)).#trainneg
cnf(i_0_916, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,np__0),np__2)=k2_xcmplx_0(k3_xcmplx_0(X1,np__2),np__0)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6488, plain, (v1_finset_1(esk11_1(X2))|~v1_finset_1(X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_917, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,np__3),np__1)=k2_xcmplx_0(k3_xcmplx_0(X1,np__1),np__3)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_918, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,np__1),np__0)=k2_xcmplx_0(k3_xcmplx_0(X1,np__0),np__0)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_11290, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__3|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|~v1_xxreal_0(k3_finseq_1(esk1_0))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_10124, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|~v1_xxreal_0(k3_finseq_1(esk1_0))|np__3!=np__1)).#trainneg
cnf(i_0_10397, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|~v1_xxreal_0(k3_finseq_1(esk1_0))|np__3!=np__2)).#trainneg
cnf(i_0_10937, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|np__3=np__0|np__3=np__1|np__3=np__4|np__3=np__2|np__4!=np__1)).#trainneg
cnf(i_0_9934, negated_conjecture, (k7_partfun1(X1,esk1_0,np__4)=k7_partfun1(X1,esk1_0,np__1)|v1_xboole_0(X1)|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|np__3=np__2|np__3=np__4|np__3=np__1|np__3=np__0|~m2_finseq_1(esk1_0,X1)|~v1_finseq_6(esk1_0,X1))).#trainneg
cnf(i_0_10219, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|np__4!=np__1|np__3!=np__4)).#trainneg
cnf(i_0_10996, negated_conjecture, (esk4_1(esk1_0)=np__2|np__4!=np__1|np__3!=np__1)).#trainneg
cnf(i_0_10975, negated_conjecture, (k3_finseq_1(np__3)=np__3|np__4!=np__1)).#trainneg
cnf(i_0_10093, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|np__4!=np__2|np__3!=np__1)).#trainneg
cnf(i_0_8634, plain, (r1_xxreal_0(X1,np__2)|~r1_xxreal_0(X1,np__0)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_11064, negated_conjecture, (v1_card_1(np__4)|v1_finset_1(np__3))).#trainneg
cnf(i_0_10369, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|np__3!=np__2)).#trainneg
cnf(i_0_919, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,np__1),np__1)=k2_xcmplx_0(k3_xcmplx_0(X1,np__1),np__1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6164, plain, (v1_finset_1(k9_xtuple_0(X2))|~v1_finset_1(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6274, plain, (v3_membered(k9_xtuple_0(X2))|~v3_membered(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_920, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,np__1),np__2)=k2_xcmplx_0(k3_xcmplx_0(X1,np__2),np__2)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_921, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,np__4),np__1)=k2_xcmplx_0(k3_xcmplx_0(X1,np__1),np__4)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__4)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9933, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|np__3=np__2|np__3=np__4|np__3=np__1|np__3=np__0|np__4!=np__2)).#trainneg
cnf(i_0_9931, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|np__3=np__2|np__3=np__4|np__3=np__1|np__3=np__0|np__4!=np__0)).#trainneg
cnf(i_0_10251, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|~v1_xxreal_0(k3_finseq_1(esk1_0))|np__3!=np__4)).#trainneg
cnf(i_0_10366, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|np__4!=np__2|np__3!=np__2)).#trainneg
cnf(i_0_8636, plain, (r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X1,np__1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_10970, negated_conjecture, (v1_finset_1(np__3)|np__4!=np__1)).#trainneg
cnf(i_0_10096, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|np__3!=np__1)).#trainneg
cnf(i_0_922, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,np__2),np__0)=k2_xcmplx_0(k3_xcmplx_0(X1,np__0),np__0)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_923, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,np__2),np__1)=k2_xcmplx_0(k3_xcmplx_0(X1,np__1),np__2)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_924, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,np__2),np__2)=k2_xcmplx_0(k3_xcmplx_0(X1,np__2),np__4)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_11939, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|~v1_xxreal_0(k3_finseq_1(esk1_0))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__2))).#trainneg
cnf(i_0_11362, negated_conjecture, (k1_card_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__2|np__3=np__2|np__3=np__4|np__3=np__1|np__3=np__0)).#trainneg
cnf(i_0_11606, negated_conjecture, (k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4)=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|np__3=np__0|np__3=np__1|np__3=np__4|np__3=np__2)).#trainneg
cnf(i_0_11539, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__4|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_12076, negated_conjecture, (esk4_1(esk1_0)=np__2|np__4!=np__2|np__3!=np__2)).#trainneg
cnf(i_0_12028, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|np__3=np__0|np__3=np__1|np__3=np__4|np__3=np__2|np__4!=np__0)).#trainneg
cnf(i_0_11973, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|np__3=np__0|np__3=np__1|np__3=np__4|np__3=np__2|np__4!=np__2)).#trainneg
cnf(i_0_11555, negated_conjecture, (esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|np__3!=np__2)).#trainneg
cnf(i_0_11950, negated_conjecture, (k1_card_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|np__4!=np__0)).#trainneg
cnf(i_0_11378, negated_conjecture, (esk4_1(esk1_0)=np__2|np__4!=np__1|np__3!=np__2)).#trainneg
cnf(i_0_8637, plain, (r1_xxreal_0(X1,np__1)|~r1_xxreal_0(X1,np__0)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_11480, plain, (v1_finset_1(esk11_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_10973, negated_conjecture, (v1_card_1(np__3)|np__4!=np__1)).#trainneg
cnf(i_0_9207, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__4!=np__1|np__3!=np__2)).#trainneg
cnf(i_0_926, plain, (k3_xcmplx_0(k2_xcmplx_0(np__0,X1),np__0)=k2_xcmplx_0(np__0,k3_xcmplx_0(X1,np__0))|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_927, plain, (k3_xcmplx_0(k2_xcmplx_0(np__0,X1),np__1)=k2_xcmplx_0(np__0,k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_928, plain, (k3_xcmplx_0(k2_xcmplx_0(np__0,X1),np__2)=k2_xcmplx_0(np__0,k3_xcmplx_0(X1,np__2))|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_11595, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|np__3=np__2|np__3=np__4|np__3=np__1|np__3=np__0|np__4!=np__1)).#trainneg
cnf(i_0_11778, negated_conjecture, (esk4_1(esk1_0)=np__2|np__4!=np__2|np__3!=np__1)).#trainneg
cnf(i_0_11938, negated_conjecture, (v2_xxreal_0(k3_finseq_1(esk1_0))|~v1_xreal_0(k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_10220, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|np__4!=np__2|np__3!=np__4)).#trainneg
cnf(i_0_12272, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|~r1_xxreal_0(k3_finseq_1(esk1_0),np__2))).#trainneg
cnf(i_0_8618, plain, (r1_xxreal_0(X1,np__3)|~r1_xxreal_0(esk8_0,np__2)|~r1_xxreal_0(X1,esk8_0)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_929, plain, (k3_xcmplx_0(k2_xcmplx_0(np__3,X1),np__1)=k2_xcmplx_0(np__3,k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__3))).#trainneg
cnf(i_0_7998, plain, (esk8_0=np__3|r1_xxreal_0(np__2,esk8_0)|esk8_0=np__4|~r1_xxreal_0(np__4,esk8_0))).#trainneg
cnf(i_0_930, plain, (k3_xcmplx_0(k2_xcmplx_0(np__1,X1),np__0)=k2_xcmplx_0(np__0,k3_xcmplx_0(X1,np__0))|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_931, plain, (k3_xcmplx_0(k2_xcmplx_0(np__1,X1),np__1)=k2_xcmplx_0(np__1,k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_12356, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__3|esk4_1(esk1_0)=np__2|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_12905, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),np__2)|np__3=np__0|np__3=np__1|np__3=np__4|np__3=np__2|np__4!=np__1)).#trainneg
cnf(i_0_12349, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|np__3=np__2|np__3=np__4|np__3=np__1|np__3=np__0|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_12214, negated_conjecture, (k3_finseq_1(np__4)=np__4|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0)).#trainneg
cnf(i_0_12906, negated_conjecture, (r1_xxreal_0(esk4_1(esk1_0),np__2)|np__3=np__0|np__3=np__1|np__3=np__4|np__3=np__2|np__4!=np__1)).#trainneg
cnf(i_0_13004, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|np__3=np__2|np__3=np__4|np__3=np__1|np__3=np__0|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_8619, plain, (r1_xxreal_0(X1,np__3)|~r1_xxreal_0(esk8_0,np__1)|~r1_xxreal_0(X1,esk8_0)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_6505, plain, (m2_finseq_1(k1_card_1(X2),X1)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_12448, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__4!=np__2|np__3!=np__2)).#trainneg
cnf(i_0_932, plain, (k3_xcmplx_0(k2_xcmplx_0(np__1,X1),np__2)=k2_xcmplx_0(np__2,k3_xcmplx_0(X1,np__2))|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_7999, plain, (esk8_0=np__3|r1_xxreal_0(np__0,esk8_0)|esk8_0=np__4|~r1_xxreal_0(np__4,esk8_0))).#trainneg
cnf(i_0_6321, plain, (v2_membered(k9_xtuple_0(X2))|~v2_membered(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6322, plain, (v1_membered(k9_xtuple_0(X2))|~v1_membered(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_933, plain, (k3_xcmplx_0(k2_xcmplx_0(np__4,X1),np__1)=k2_xcmplx_0(np__4,k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__4))).#trainneg
cnf(i_0_934, plain, (k3_xcmplx_0(k2_xcmplx_0(np__2,X1),np__0)=k2_xcmplx_0(np__0,k3_xcmplx_0(X1,np__0))|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_13820, negated_conjecture, (np__4!=np__2|np__3!=np__2|np__4!=np__1)).#trainneg
cnf(i_0_13587, negated_conjecture, (k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(np__4)=np__4)).#trainneg
cnf(i_0_13654, negated_conjecture, (k1_card_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)))).#trainneg
cnf(i_0_13634, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|np__3=np__2|np__3=np__4|np__3=np__1|np__3=np__0|np__4!=np__1)).#trainneg
cnf(i_0_12614, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|esk4_1(esk1_0)=np__2|np__3!=np__2)).#trainneg
cnf(i_0_13855, plain, (v2_membered(k9_xtuple_0(esk7_0)))).#trainneg
cnf(i_0_12728, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__4!=np__1|np__3!=np__2)).#trainneg
cnf(i_0_935, plain, (k3_xcmplx_0(k2_xcmplx_0(np__2,X1),np__1)=k2_xcmplx_0(np__2,k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_13815, plain, (m2_finseq_1(k3_finseq_1(X1),X2)|~v1_xboole_0(k3_finseq_1(X1))|~v1_xboole_0(X2)|~v1_finseq_1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(k3_finseq_1(X1))|~v1_funct_1(X1)|~v1_relat_1(k3_finseq_1(X1))|~v1_relat_1(X1))).#trainneg
cnf(i_0_936, plain, (k3_xcmplx_0(k2_xcmplx_0(np__2,X1),np__2)=k2_xcmplx_0(np__4,k3_xcmplx_0(X1,np__2))|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_947, plain, (m1_subset_1(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)),k1_numbers)|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,X2)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,X2)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k1_nat_1(np__1,X2),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_13472, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|np__3=np__2|np__3=np__4|np__3=np__1|np__3=np__0|np__4!=np__1)).#trainneg
cnf(i_0_13650, negated_conjecture, (k1_card_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|np__4!=np__1)).#trainneg
cnf(i_0_12445, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__4|k1_card_1(esk1_0)=np__4)).#trainneg
cnf(i_0_12939, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__4!=np__2|np__3!=np__1)).#trainneg
cnf(i_0_13857, negated_conjecture, (v2_membered(k9_xtuple_0(esk1_0)))).#trainneg
cnf(i_0_8620, plain, (r1_xxreal_0(X1,np__3)|~r1_xxreal_0(esk8_0,np__0)|~r1_xxreal_0(X1,esk8_0)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_13859, plain, (v2_membered(k9_xtuple_0(np__0)))).#trainneg
cnf(i_0_11059, negated_conjecture, (k1_card_1(np__4)=np__4|v1_finset_1(np__3))).#trainneg
cnf(i_0_11675, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__4!=np__1|np__3!=np__1)).#trainneg
cnf(i_0_948, plain, (k19_euclid(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,X2))),k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)))=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,X2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,X2)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,X2)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k1_nat_1(np__1,X2),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_968, plain, (m1_subset_1(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)),k1_numbers)|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k2_xcmplx_0(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_14657, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|k1_card_1(esk1_0)=np__4|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|~v1_xxreal_0(k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_14002, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_13998, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_12887, negated_conjecture, (k7_partfun1(X1,esk1_0,np__1)=k7_partfun1(X1,esk1_0,np__2)|v1_xboole_0(X1)|np__3=np__0|np__3=np__1|np__3=np__4|np__3=np__2|~m2_finseq_1(esk1_0,X1)|~v1_finseq_6(esk1_0,X1)|np__4!=np__1)).#trainneg
cnf(i_0_13467, negated_conjecture, (esk3_1(esk1_0)=np__2|np__3=np__2|np__3=np__4|np__3=np__1|np__3=np__0|~r1_xxreal_0(np__2,esk3_1(esk1_0))|np__4!=np__1)).#trainneg
cnf(i_0_5790, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_13867, plain, (v1_membered(k9_xtuple_0(esk7_0)))).#trainneg
cnf(i_0_12440, negated_conjecture, (k1_card_1(esk1_0)=np__4|k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)|~r1_xxreal_0(np__3,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_977, plain, (m1_subset_1(X1,k9_xtuple_0(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_970, plain, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,np__1))))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,X3)))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,np__1))))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,X3)))|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~r1_xxreal_0(k2_xcmplx_0(X2,k2_xcmplx_0(X3,np__1)),k3_finseq_1(X1))|~r1_xxreal_0(np__1,k2_xcmplx_0(X2,X3))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_13651, negated_conjecture, (k1_card_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|np__4!=np__2)).#trainneg
cnf(i_0_14001, negated_conjecture, (esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|np__3!=np__2)).#trainneg
cnf(i_0_13628, negated_conjecture, (esk4_1(esk1_0)=np__2|np__3=np__2|np__3=np__4|np__3=np__1|np__3=np__0|~r1_xxreal_0(np__2,esk4_1(esk1_0))|np__4!=np__1)).#trainneg
cnf(i_0_13471, negated_conjecture, (np__3=np__2|np__3=np__4|np__3=np__1|np__3=np__0|~r1_xxreal_0(np__3,esk3_1(esk1_0))|np__4!=np__1)).#trainneg
cnf(i_0_13995, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_6102, plain, (v1_int_1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_13869, negated_conjecture, (v1_membered(k9_xtuple_0(esk1_0)))).#trainneg
cnf(i_0_14308, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__3!=np__2)).#trainneg
cnf(i_0_1015, plain, (m1_subset_1(esk4_1(X1),k9_xtuple_0(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3088, plain, (r1_xxreal_0(esk8_0,np__2)|r1_xxreal_0(np__1,esk8_0))).#trainneg
cnf(i_0_990, plain, (k19_euclid(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(k3_finseq_1(X1),np__1))),k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1)))=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(k3_finseq_1(X1),np__1))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(k3_finseq_1(X1),np__1)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1))|v1_xboole_0(u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(k3_finseq_1(X1),np__1)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(k3_finseq_1(X1))|~r1_xxreal_0(k1_nat_1(k3_finseq_1(X1),np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_14936, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__4)).#trainneg
cnf(i_0_14940, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__3|esk4_1(esk1_0)=np__2|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_15007, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__3|esk3_1(esk1_0)=np__2|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_15223, negated_conjecture, (r1_xxreal_0(esk4_1(esk1_0),np__2)|np__4!=np__1)).#trainneg
cnf(i_0_15258, plain, (v1_xcmplx_0(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_13871, plain, (v1_membered(k9_xtuple_0(np__0)))).#trainneg
cnf(i_0_13292, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__4|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_1017, plain, (v2_xxreal_0(X1)|r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,np__1))|~v2_xxreal_0(k2_xcmplx_0(X2,np__1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13819, plain, (m2_finseq_1(esk11_1(X3),X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1001, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~r2_hidden(X2,k9_xtuple_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(k2_nat_1(X2,np__1))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,k2_nat_1(X2,np__1)))).#trainneg
cnf(i_0_15156, negated_conjecture, (r1_xxreal_0(esk4_1(esk1_0),np__2)|np__4!=np__2)).#trainneg
cnf(i_0_15510, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_15404, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|esk4_1(esk1_0)=np__2|np__3!=np__2)).#trainneg
cnf(i_0_14798, negated_conjecture, (k1_card_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__4)).#trainneg
cnf(i_0_15222, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),np__2)|np__4!=np__1)).#trainneg
cnf(i_0_15103, negated_conjecture, (r1_xxreal_0(esk4_1(esk1_0),np__2)|np__4!=np__0)).#trainneg
cnf(i_0_15102, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),np__2)|np__4!=np__0)).#trainneg
cnf(i_0_15512, plain, (v1_int_1(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_13853, plain, (v2_membered(k9_xtuple_0(esk11_1(X1))))).#trainneg
cnf(i_0_11066, negated_conjecture, (k3_finseq_1(np__4)=np__4|v1_finset_1(np__3))).#trainneg
cnf(i_0_14628, negated_conjecture, (k1_card_1(esk1_0)=np__4|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)))).#trainneg
cnf(i_0_1019, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(k2_xcmplx_0(X3,np__1))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k2_xcmplx_0(X3,np__1))|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1020, plain, (r1_xxreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,np__1)),X3)|r1_xxreal_0(X3,k2_xcmplx_0(X1,X2))|~v7_ordinal1(X3)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1034, plain, (k2_xcmplx_0(X1,np__1)=k1_nat_1(X2,np__1)|r1_xxreal_0(k2_xcmplx_0(X1,np__1),X2)|r1_xxreal_0(k1_nat_1(X2,np__1),X1)|~v1_xxreal_0(k1_nat_1(X2,np__1))|~v1_xxreal_0(k2_xcmplx_0(X1,np__1))|~v7_ordinal1(k2_xcmplx_0(X1,np__1))|~v7_ordinal1(X2)|~v7_ordinal1(k1_nat_1(X2,np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15155, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),np__2)|np__4!=np__2)).#trainneg
cnf(i_0_16295, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X2,np__2))).#trainneg
cnf(i_0_16292, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X2,np__1))).#trainneg
cnf(i_0_16293, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X2,np__4))).#trainneg
cnf(i_0_16191, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|np__4!=np__0)).#trainneg
cnf(i_0_13865, plain, (v1_membered(k9_xtuple_0(esk11_1(X1))))).#trainneg
cnf(i_0_16134, negated_conjecture, (v1_card_1(np__4)|k1_card_1(esk1_0)=np__2)).#trainneg
cnf(i_0_1035, plain, (k1_nat_1(X1,np__1)=k1_nat_1(X2,np__1)|r1_xxreal_0(k1_nat_1(X1,np__1),X2)|r1_xxreal_0(k1_nat_1(X2,np__1),X1)|~v1_xxreal_0(k1_nat_1(X2,np__1))|~v1_xxreal_0(k1_nat_1(X1,np__1))|~v7_ordinal1(k1_nat_1(X1,np__1))|~v7_ordinal1(X2)|~v7_ordinal1(k1_nat_1(X2,np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15542, plain, (esk8_0=np__2|r1_xxreal_0(np__1,esk8_0)|~r1_xxreal_0(np__2,esk8_0))).#trainneg
cnf(i_0_15943, plain, (v2_xxreal_0(X1)|r1_xxreal_0(X1,np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15939, plain, (v2_xxreal_0(X1)|r1_xxreal_0(X1,np__0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1037, plain, (m1_subset_1(esk3_1(X1),k9_xtuple_0(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_15906, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|np__4!=np__1)).#trainneg
cnf(i_0_16296, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X2,np__3))).#trainneg
cnf(i_0_15870, negated_conjecture, (esk4_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__4|~r1_xxreal_0(np__2,esk4_1(esk1_0)))).#trainneg
cnf(i_0_15871, negated_conjecture, (r1_xxreal_0(X1,np__2)|k1_card_1(esk1_0)=np__4|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk4_1(esk1_0)))).#trainneg
cnf(i_0_16149, negated_conjecture, (k7_partfun1(X1,esk1_0,np__4)=k7_partfun1(X1,esk1_0,np__1)|v1_xboole_0(X1)|k1_card_1(esk1_0)=np__2|~m2_finseq_1(esk1_0,X1)|~v1_finseq_6(esk1_0,X1))).#trainneg
cnf(i_0_15970, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|np__4!=np__2)).#trainneg
cnf(i_0_15947, plain, (m2_finseq_1(esk11_1(X1),X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_16067, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__3!=np__2)).#trainneg
cnf(i_0_3142, plain, (r1_xxreal_0(esk8_0,np__1)|r1_xxreal_0(np__0,esk8_0))).#trainneg
cnf(i_0_15940, plain, (v2_xxreal_0(X1)|r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15942, plain, (v2_xxreal_0(X1)|r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1055, plain, (v1_xboole_0(X1)|v1_xboole_0(k4_finseq_1(X1))|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(np__1,X2))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~r2_hidden(X2,k4_finseq_1(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~m1_subset_1(k2_nat_1(np__1,X2),k4_finseq_1(X1)))).#trainneg
cnf(i_0_16164, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|np__4!=np__1)).#trainneg
cnf(i_0_15900, negated_conjecture, (esk4_1(esk1_0)=np__2|~r1_xxreal_0(np__2,esk4_1(esk1_0))|np__4!=np__1)).#trainneg
cnf(i_0_16220, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|np__4!=np__0)).#trainneg
cnf(i_0_15790, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_3667, plain, (v3_ordinal1(X1)|~v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_5981, plain, (v1_relat_1(k9_xtuple_0(k1_card_1(X4)))|~v1_xboole_0(X4))).#trainneg
cnf(i_0_16131, negated_conjecture, (v1_finset_1(np__4)|k1_card_1(esk1_0)=np__2)).#trainneg
cnf(i_0_15651, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_1058, plain, (v1_finset_1(esk11_1(X1))|~v1_funct_1(esk11_1(X1))|~v1_relat_1(esk11_1(X1)))).#trainneg
cnf(i_0_3166, plain, (r1_xxreal_0(esk8_0,np__2)|r1_xxreal_0(np__0,esk8_0))).#trainneg
cnf(i_0_15904, negated_conjecture, (~r1_xxreal_0(np__3,esk4_1(esk1_0))|np__4!=np__1)).#trainneg
cnf(i_0_1061, plain, (v7_ordinal1(k1_card_1(X1))|~v1_finset_1(k1_card_1(X1)))).#trainneg
cnf(i_0_15874, negated_conjecture, (k1_card_1(esk1_0)=np__4|~r1_xxreal_0(np__3,esk4_1(esk1_0)))).#trainneg
cnf(i_0_16652, negated_conjecture, (esk4_1(esk1_0)=np__2|np__4!=np__1)).#trainneg
cnf(i_0_17086, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__3|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_17181, plain, (v7_ordinal1(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_17107, negated_conjecture, (v1_finset_1(np__2)|v1_finset_1(np__4))).#trainneg
cnf(i_0_15994, negated_conjecture, (esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_1065, plain, (r1_xxreal_0(k2_xcmplx_0(np__1,esk2_1(X1)),k3_finseq_1(X1))|v1_topreal1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk2_1(X1),k5_numbers))).#trainneg
cnf(i_0_17173, plain, (v7_ordinal1(esk9_2(X2,X3))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1066, plain, (v1_topreal1(X1)|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(np__1,esk2_1(X1))))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,esk2_1(X1)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk2_1(X1),k5_numbers))).#trainneg
cnf(i_0_1067, plain, (v1_topreal1(X1)|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(np__1,esk2_1(X1))))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,esk2_1(X1)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk2_1(X1),k5_numbers))).#trainneg
cnf(i_0_16792, negated_conjecture, (esk4_1(esk1_0)=np__2|np__4!=np__2)).#trainneg
cnf(i_0_16529, negated_conjecture, (esk4_1(esk1_0)=np__2|np__4!=np__0)).#trainneg
cnf(i_0_16899, negated_conjecture, (esk3_1(esk1_0)=np__2|np__4!=np__1)).#trainneg
cnf(i_0_16611, plain, (k1_nat_1(X1,np__1)=k1_nat_1(np__1,X2)|r1_xxreal_0(k1_nat_1(np__1,X2),X1)|r1_xxreal_0(k1_nat_1(X1,np__1),X2)|~v1_xxreal_0(k1_nat_1(np__1,X2))|~v1_xxreal_0(k1_nat_1(X1,np__1))|~v7_ordinal1(k1_nat_1(X1,np__1))|~v7_ordinal1(k1_nat_1(np__1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8632, plain, (r1_xxreal_0(np__4,X1)|r1_xxreal_0(X1,np__4)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5984, plain, (v4_relat_1(k1_card_1(X3),X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_17110, negated_conjecture, (v1_card_1(np__2)|v1_finset_1(np__4))).#trainneg
cnf(i_0_1073, plain, (X1=k2_xcmplx_0(np__1,X2)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(k2_xcmplx_0(np__1,X2))|~v1_xxreal_0(X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(X1,k2_xcmplx_0(np__1,X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_15968, negated_conjecture, (~r1_xxreal_0(np__3,esk4_1(esk1_0))|np__4!=np__2)).#trainneg
cnf(i_0_16163, negated_conjecture, (~r1_xxreal_0(np__3,esk3_1(esk1_0))|np__4!=np__1)).#trainneg
cnf(i_0_17176, plain, (v7_ordinal1(esk11_1(X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1074, plain, (k19_euclid(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(np__1,X2))),k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)))=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(np__1,X2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(np__1,X2)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(np__1,X2)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k2_xcmplx_0(np__1,X2),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_16788, negated_conjecture, (k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4)=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)|k1_card_1(esk1_0)=np__2)).#trainneg
cnf(i_0_15263, plain, (v1_xxreal_0(X1)|~v2_membered(k9_xtuple_0(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_3406, plain, (v6_membered(X1)|~m2_finseq_1(X1,X2)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,X2)))).#trainneg
cnf(i_0_5985, plain, (v5_relat_1(k1_card_1(X3),X2)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_14238, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|k1_card_1(esk1_0)=np__4|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0))|~r1_xxreal_0(np__3,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_16621, plain, (esk8_0=np__1|esk8_0=np__2|~r1_xxreal_0(esk8_0,np__1)|~r1_xxreal_0(np__2,esk8_0))).#trainneg
cnf(i_0_16167, negated_conjecture, (~r1_xxreal_0(np__4,esk3_1(esk1_0))|np__4!=np__1)).#trainneg
cnf(i_0_1080, plain, (v1_xboole_0(X1)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,np__1)))!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,X3))|~m1_subset_1(k2_xcmplx_0(X2,X3),k5_numbers)|~r2_hidden(k2_xcmplx_0(X2,k2_xcmplx_0(X3,np__1)),k4_finseq_1(X1))|~r2_hidden(k2_xcmplx_0(X2,X3),k4_finseq_1(X1))|~m2_subset_1(k2_xcmplx_0(X2,X3),k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1089, plain, (k7_partfun1(X1,k1_card_1(X2),k1_card_1(X2))=k7_partfun1(X1,k1_card_1(X2),np__1)|v1_xboole_0(X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m2_finseq_1(k1_card_1(X2),X1)|~v1_finseq_6(k1_card_1(X2),X1))).#trainneg
cnf(i_0_17723, plain, (v7_ordinal1(esk11_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_17406, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__4!=np__2)).#trainneg
cnf(i_0_17865, plain, (esk8_0=np__2|esk8_0=np__1|r1_xxreal_0(np__0,esk8_0)|~r1_xxreal_0(np__2,esk8_0))).#trainneg
cnf(i_0_17844, plain, (v6_membered(X1)|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_17669, plain, (v4_relat_1(k3_finseq_1(X1),X2)|~v1_xboole_0(k3_finseq_1(X1))|~v1_finseq_1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(k3_finseq_1(X1))|~v1_funct_1(X1)|~v1_relat_1(k3_finseq_1(X1))|~v1_relat_1(X1))).#trainneg
cnf(i_0_17640, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__3|k1_card_1(esk1_0)=np__4|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_17540, negated_conjecture, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__4!=np__1)).#trainneg
cnf(i_0_17270, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__4!=np__1)).#trainneg
cnf(i_0_17384, plain, (v7_ordinal1(esk9_2(X1,X2)))).#trainneg
cnf(i_0_17129, negated_conjecture, (np__4=np__2|np__3=np__2|v1_finset_1(np__4))).#trainneg
cnf(i_0_17473, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__4!=np__0)).#trainneg
cnf(i_0_1101, plain, (v2_xxreal_0(X1)|r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(k1_nat_1(np__1,X2))|~v2_xxreal_0(k1_nat_1(np__1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_17854, plain, (v5_relat_1(k3_finseq_1(X1),X2)|~v1_xboole_0(k3_finseq_1(X1))|~v1_finseq_1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(k3_finseq_1(X1))|~v1_funct_1(X1)|~v1_relat_1(k3_finseq_1(X1))|~v1_relat_1(X1))).#trainneg
cnf(i_0_1103, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(k1_nat_1(np__1,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k1_nat_1(np__1,X3))|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1113, plain, (r1_xxreal_0(np__1,X1)|v1_xboole_0(k9_xtuple_0(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k9_xtuple_0(X2)))).#trainneg
cnf(i_0_3407, plain, (v1_finset_1(X1)|~m2_finseq_1(X1,X2)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,X2)))).#trainneg
cnf(i_0_6103, plain, (v1_relat_1(k9_xtuple_0(esk11_1(X4)))|~v1_xboole_0(X4))).#trainneg
cnf(i_0_1115, plain, (esk9_2(X1,X2)=np__0)).#trainneg
cnf(i_0_1116, plain, (k1_card_1(X1)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_17807, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__3|k1_card_1(esk1_0)=np__2|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_17641, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|k1_card_1(esk1_0)=np__4|np__3!=np__2)).#trainneg
cnf(i_0_3412, plain, (v3_membered(X1)|~m2_finseq_1(X1,X2)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,X2)))).#trainneg
cnf(i_0_2614, plain, (esk11_1(X1)=esk11_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_17131, negated_conjecture, (np__3=np__2|v1_finset_1(np__4)|np__4!=np__1)).#trainneg
cnf(i_0_1129, plain, (m1_subset_1(X1,X2)|v1_xboole_0(k1_relset_1(X2,X3))|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~m1_subset_1(X1,k1_relset_1(X2,X3)))).#trainneg
cnf(i_0_14813, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),np__2)|k1_card_1(esk1_0)=np__4)).#trainneg
cnf(i_0_17105, negated_conjecture, (k1_card_1(np__2)=np__2|v1_finset_1(np__4))).#trainneg
cnf(i_0_16879, plain, (esk8_0=np__1|r1_xxreal_0(np__0,esk8_0)|~r1_xxreal_0(np__1,esk8_0))).#trainneg
cnf(i_0_17673, plain, (v4_relat_1(esk11_1(X3),X2)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1133, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|v1_xboole_0(k9_xtuple_0(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k9_xtuple_0(X2)))).#trainneg
cnf(i_0_13996, negated_conjecture, (esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|np__3!=np__2)).#trainneg
cnf(i_0_2182, plain, (v1_xreal_0(esk6_1(X1))|~v5_membered(X1))).#trainneg
cnf(i_0_17112, negated_conjecture, (k3_finseq_1(np__2)=np__2|v1_finset_1(np__4))).#trainneg
cnf(i_0_17141, plain, (esk8_0=np__2|r1_xxreal_0(np__0,esk8_0)|~r1_xxreal_0(np__2,esk8_0))).#trainneg
cnf(i_0_1139, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X3,X1)|~v1_xxreal_0(k2_xcmplx_0(np__1,X2))|~v1_xxreal_0(X1)|~v1_xxreal_0(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(X3,k2_xcmplx_0(np__1,X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_7317, plain, (r1_xxreal_0(X1,esk3_1(X2))|esk3_1(X2)=np__3|esk3_1(X2)=np__1|esk3_1(X2)=np__2|v1_xboole_0(X2)|v3_funct_1(X2)|~v1_xxreal_0(esk3_1(X2))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__4)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_9504, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),np__3)|r1_xxreal_0(np__3,esk3_1(esk1_0)))).#trainneg
cnf(i_0_18384, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|k1_card_1(esk1_0)=np__4)).#trainneg
cnf(i_0_6815, plain, (v1_card_1(np__0))).#trainneg
cnf(i_0_602, plain, (v6_membered(np__0)|~v6_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1141, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(k2_xcmplx_0(X3,np__1),X1)|r1_xxreal_0(k1_nat_1(X2,np__1),X3)|~v1_xxreal_0(k1_nat_1(X2,np__1))|~v1_xxreal_0(X1)|~v1_xxreal_0(k2_xcmplx_0(X3,np__1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(k1_nat_1(X2,np__1))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_18433, plain, (esk8_0=np__0|esk8_0=np__1|~r1_xxreal_0(esk8_0,np__0)|~r1_xxreal_0(np__1,esk8_0))).#trainneg
cnf(i_0_1142, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(k1_nat_1(np__1,X3),X1)|r1_xxreal_0(k1_nat_1(X2,np__1),X3)|~v1_xxreal_0(k1_nat_1(X2,np__1))|~v1_xxreal_0(X1)|~v1_xxreal_0(k1_nat_1(np__1,X3))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(k1_nat_1(X2,np__1))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1143, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(k1_nat_1(X3,np__1),X1)|r1_xxreal_0(k1_nat_1(X2,np__1),X3)|~v1_xxreal_0(k1_nat_1(X2,np__1))|~v1_xxreal_0(X1)|~v1_xxreal_0(k1_nat_1(X3,np__1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(k1_nat_1(X2,np__1))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_18486, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|esk3_1(esk1_0)=np__2|np__3!=np__2)).#trainneg
cnf(i_0_18184, negated_conjecture, (k1_card_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__3|esk3_1(esk1_0)=np__2|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_17990, plain, (v6_membered(np__0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_619, plain, (v1_membered(np__0)|~v1_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_18620, plain, (esk8_0=np__0|esk8_0=np__2|~r1_xxreal_0(esk8_0,np__0)|~r1_xxreal_0(np__2,esk8_0))).#trainneg
cnf(i_0_1164, plain, (v1_finset_1(k9_xtuple_0(X1))|~v3_card_1(X1,np__1)|~v1_card_1(np__1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1166, plain, (esk6_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=k1_card_1(X3)|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_18678, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__3|r1_xxreal_0(np__3,esk3_1(esk1_0)))).#trainneg
cnf(i_0_18379, negated_conjecture, (esk3_1(esk1_0)=np__2|k1_card_1(esk1_0)=np__4|~r1_xxreal_0(np__2,esk3_1(esk1_0)))).#trainneg
cnf(i_0_268, plain, (v4_relat_1(np__0,X1))).#trainneg
cnf(i_0_614, plain, (v2_membered(np__0)|~v2_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1167, plain, (esk6_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=np__0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_16189, negated_conjecture, (~r1_xxreal_0(np__3,esk4_1(esk1_0))|np__4!=np__0)).#trainneg
cnf(i_0_18135, plain, (v1_finset_1(X1)|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_18328, plain, (v3_membered(X1)|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1168, plain, (X1=esk6_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_18380, negated_conjecture, (r1_xxreal_0(X1,np__2)|k1_card_1(esk1_0)=np__4|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk3_1(esk1_0)))).#trainneg
cnf(i_0_18679, negated_conjecture, (esk3_1(esk1_0)=np__3|r1_xxreal_0(esk3_1(esk1_0),np__2)|r1_xxreal_0(np__3,esk3_1(esk1_0)))).#trainneg
cnf(i_0_18383, negated_conjecture, (k1_card_1(esk1_0)=np__4|~r1_xxreal_0(np__3,esk3_1(esk1_0)))).#trainneg
cnf(i_0_18452, plain, (r1_xxreal_0(k7_partfun1(k9_xtuple_0(X1),X2,X3),k3_finseq_1(X1))|v1_xboole_0(k9_xtuple_0(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(k7_partfun1(k9_xtuple_0(X1),X2,X3))|~m2_finseq_1(X2,k9_xtuple_0(X1)))).#trainneg
cnf(i_0_18662, plain, (esk3_1(X1)=X2|esk3_1(X1)=np__2|esk3_1(X1)=np__1|esk3_1(X1)=np__3|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(esk3_1(X1))|~r1_xxreal_0(esk3_1(X1),X2)|~r1_xxreal_0(X2,np__4)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_762, plain, (v1_relat_1(k9_xtuple_0(np__0)))).#trainneg
cnf(i_0_610, plain, (v3_membered(np__0)|~v3_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_16219, negated_conjecture, (~r1_xxreal_0(np__3,esk3_1(esk1_0))|np__4!=np__0)).#trainneg
cnf(i_0_1200, plain, (k19_euclid(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(k1_card_1(X1),np__1))),k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1)))=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(k1_card_1(X1),np__1))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(k1_card_1(X1),np__1)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1))|v1_xboole_0(u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(k1_card_1(X1),np__1)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(k1_card_1(X1))|~r1_xxreal_0(k1_nat_1(k1_card_1(X1),np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,k1_card_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1201, plain, (k7_partfun1(X1,k1_card_1(X2),k1_card_1(X2))=k7_partfun1(X1,k1_card_1(X2),np__1)|v1_xboole_0(X1)|~m2_finseq_1(k1_card_1(X2),X1)|~v1_finseq_6(k1_card_1(X2),X1))).#trainneg
cnf(i_0_19112, negated_conjecture, (esk3_1(esk1_0)=np__3|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|~r1_xxreal_0(esk3_1(esk1_0),np__3))).#trainneg
cnf(i_0_19337, negated_conjecture, (esk3_1(esk1_0)=np__3|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|np__4!=np__2)).#trainneg
cnf(i_0_19046, plain, (v6_membered(np__0))).#trainneg
cnf(i_0_17984, plain, (v6_membered(esk11_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1645, plain, (esk11_1(X1)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_19111, negated_conjecture, (v2_xxreal_0(esk3_1(esk1_0))|esk3_1(esk1_0)=np__3|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2)).#trainneg
cnf(i_0_18133, plain, (r1_xxreal_0(np__1,k7_partfun1(k9_xtuple_0(X1),X2,X3))|v1_xboole_0(k9_xtuple_0(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(k7_partfun1(k9_xtuple_0(X1),X2,X3))|~m2_finseq_1(X2,k9_xtuple_0(X1)))).#trainneg
cnf(i_0_19292, negated_conjecture, (r1_xxreal_0(k7_partfun1(k9_xtuple_0(esk1_0),X1,X2),np__4)|~v1_xxreal_0(k7_partfun1(k9_xtuple_0(esk1_0),X1,X2))|~m2_finseq_1(X1,k9_xtuple_0(esk1_0)))).#trainneg
cnf(i_0_18874, negated_conjecture, (k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__3!=np__2)).#trainneg
cnf(i_0_19204, plain, (v3_membered(np__0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_19199, plain, (v3_membered(esk11_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_271, plain, (m1_subset_1(np__0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1224, plain, (v2_xxreal_0(k3_finseq_1(X1))|v1_topreal1(X1)|~v1_xreal_0(k3_finseq_1(X1))|~v1_xreal_0(k1_nat_1(np__1,esk2_1(X1)))|~v2_xxreal_0(k1_nat_1(np__1,esk2_1(X1)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1221, plain, (v7_ordinal1(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~v6_membered(k1_numbers)|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k1_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_1222, plain, (v7_ordinal1(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,X2)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~v6_membered(k1_numbers)|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,X2)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k1_nat_1(np__1,X2),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_18441, plain, (v4_relat_1(esk11_1(X1),X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_15087, negated_conjecture, (k1_card_1(esk1_0)=np__2|np__4!=np__0)).#trainneg
cnf(i_0_735, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X1,np__0))).#trainneg
cnf(i_0_1225, plain, (k1_nat_1(np__1,esk2_1(X1))=k3_finseq_1(X1)|v1_topreal1(X1)|~v1_xxreal_0(k1_nat_1(np__1,esk2_1(X1)))|~v1_xxreal_0(k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),k1_nat_1(np__1,esk2_1(X1)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2100, plain, (esk8_0=np__1|r1_xxreal_0(esk8_0,np__0)|~v7_ordinal1(esk8_0)|~r1_xxreal_0(esk8_0,np__1))).#trainneg
cnf(i_0_16223, negated_conjecture, (~r1_xxreal_0(np__4,esk3_1(esk1_0))|np__4!=np__0)).#trainneg
cnf(i_0_1223, plain, (v7_ordinal1(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~v6_membered(k1_numbers)|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k2_xcmplx_0(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_1226, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|v1_topreal1(X2)|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(k1_nat_1(np__1,esk2_1(X2)))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k1_nat_1(np__1,esk2_1(X2)))|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4596, plain, (r1_xxreal_0(X1,esk4_1(X2))|esk4_1(X2)=np__1|v1_xboole_0(X2)|v3_funct_1(X2)|~v1_xxreal_0(esk4_1(X2))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__2)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_8526, plain, (k2_xcmplx_0(np__2,X1)=np__4|k2_xcmplx_0(np__1,X1)=np__0|k2_xcmplx_0(np__1,X1)=np__1|k2_xcmplx_0(np__1,X1)=np__4|k2_xcmplx_0(np__1,X1)=np__2|r1_xxreal_0(np__4,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2732, plain, (r1_xxreal_0(k3_finseq_1(X1),np__3)|~r1_xxreal_0(k3_finseq_1(X1),np__2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3976, plain, (v1_xboole_0(X1)|k2_xcmplx_0(np__1,X2)=np__0|k2_xcmplx_0(np__1,X2)=np__1|k2_xcmplx_0(np__1,X2)=np__4|k2_xcmplx_0(np__1,X2)=np__2|r1_xxreal_0(np__4,X2)|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|~m1_subset_1(X2,k5_numbers)|~r2_hidden(np__3,k9_xtuple_0(X1))|~r2_hidden(X2,k9_xtuple_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1))).#trainneg
cnf(i_0_4222, plain, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__4))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__4))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k1_nat_1(np__1,X2)=np__2|k1_nat_1(np__1,X2)=np__4|k1_nat_1(np__1,X2)=np__1|k1_nat_1(np__1,X2)=np__0|r1_xxreal_0(np__4,X2)|~r1_xxreal_0(np__4,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(np__3,k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2))).#trainneg
cnf(i_0_15088, negated_conjecture, (k3_finseq_1(np__2)=np__2|np__4!=np__0)).#trainneg
cnf(i_0_839, plain, (m1_subset_1(k7_partfun1(X1,np__0,X3),X1))).#trainneg
cnf(i_0_1228, plain, (r1_xxreal_0(k1_nat_1(np__1,esk2_1(k3_finseq_1(X1))),k3_finseq_1(X1))|v1_topreal1(k3_finseq_1(X1))|~m2_finseq_1(k3_finseq_1(X1),u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2277, plain, (esk8_0=np__4|r1_xxreal_0(esk8_0,np__3)|~v7_ordinal1(esk8_0)|~r1_xxreal_0(esk8_0,np__4))).#trainneg
cnf(i_0_16371, negated_conjecture, (~r1_xxreal_0(np__3,esk3_1(esk1_0))|np__4!=np__2)).#trainneg
cnf(i_0_1229, plain, (r1_xxreal_0(k1_nat_1(np__1,esk2_1(X1)),k1_card_1(X1))|v1_topreal1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16788, negated_conjecture, (k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4)=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)|np__4=np__2)).#trainneg
cnf(i_0_14798, negated_conjecture, (k1_card_1(esk1_0)=np__4|np__4=np__2)).#trainneg
cnf(i_0_19853, plain, (esk4_1(X1)=X2|esk4_1(X1)=np__1|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(esk4_1(X1))|~r1_xxreal_0(esk4_1(X1),X2)|~r1_xxreal_0(X2,np__2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_19597, plain, (v3_membered(np__0))).#trainneg
cnf(i_0_19464, negated_conjecture, (v1_finset_1(np__4))).#trainneg
cnf(i_0_19784, plain, (esk8_0=np__1|r1_xxreal_0(esk8_0,np__0)|~r1_xxreal_0(esk8_0,np__1)|~m1_subset_1(esk8_0,k5_numbers))).#trainneg
cnf(i_0_1232, plain, (m1_subset_1(X1,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k9_xtuple_0(X2)))).#trainneg
cnf(i_0_18184, negated_conjecture, (esk3_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__3|np__4=np__2|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_19469, negated_conjecture, (k3_finseq_1(np__4)=np__4)).#trainneg
cnf(i_0_19462, negated_conjecture, (k1_card_1(np__4)=np__4)).#trainneg
cnf(i_0_19842, plain, (esk4_1(X2)=np__1|v1_xboole_0(X2)|v3_funct_1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(esk4_1(X2))|~r1_xxreal_0(esk4_1(X2),X1)|~v1_xxreal_0(k1_nat_1(np__1,X1))|~r1_xxreal_0(k1_nat_1(np__1,X1),np__2)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_19467, negated_conjecture, (v1_card_1(np__4))).#trainneg
cnf(i_0_16134, negated_conjecture, (np__4=np__2|v1_card_1(np__4))).#trainneg
cnf(i_0_20128, plain, (esk8_0=np__4|r1_xxreal_0(esk8_0,np__3)|~r1_xxreal_0(esk8_0,np__4)|~m1_subset_1(esk8_0,k5_numbers))).#trainneg
cnf(i_0_1235, plain, (m1_subset_1(X1,k5_numbers)|v1_xboole_0(k4_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(X1,k4_finseq_1(X2)))).#trainneg
cnf(i_0_19892, plain, (k3_finseq_1(X1)=np__3|k3_finseq_1(X1)=np__4|~v1_xxreal_0(k3_finseq_1(X1))|~r1_xxreal_0(np__3,k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r1_xxreal_0(k3_finseq_1(X1),np__4))).#trainneg
cnf(i_0_19912, plain, (~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r1_xxreal_0(np__4,k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),np__2))).#trainneg
cnf(i_0_19847, plain, (esk4_1(X1)=k3_finseq_1(X1)|v1_xboole_0(X1)|v3_funct_1(X1)|esk4_1(X1)=np__1|~v1_xxreal_0(esk4_1(X1))|~v1_xxreal_0(k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(k3_finseq_1(X1),np__2))).#trainneg
cnf(i_0_19843, plain, (esk4_1(X2)=np__1|v1_xboole_0(X2)|v3_funct_1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(esk4_1(X2))|~r1_xxreal_0(esk4_1(X2),X1)|~v1_xxreal_0(k1_nat_1(X1,np__1))|~r1_xxreal_0(k1_nat_1(X1,np__1),np__2)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_18161, plain, (r1_xxreal_0(k1_nat_1(esk2_1(X1),np__1),np__0)|v1_topreal1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_19508, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|np__4!=np__2)).#trainneg
cnf(i_0_19783, plain, (esk8_0=np__1|r1_xxreal_0(esk8_0,np__0)|~r1_xxreal_0(esk8_0,np__1)|~v1_xboole_0(esk8_0))).#trainneg
cnf(i_0_16375, negated_conjecture, (~r1_xxreal_0(np__4,esk3_1(esk1_0))|np__4!=np__2)).#trainneg
cnf(i_0_17989, plain, (v6_membered(k1_card_1(X1))|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1241, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(esk8_0,X1))=k3_xcmplx_0(esk8_0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_1242, plain, (k3_xcmplx_0(k2_xcmplx_0(np__1,X1),esk8_0)=k2_xcmplx_0(esk8_0,k3_xcmplx_0(X1,esk8_0))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_20007, negated_conjecture, (k1_card_1(np__2)=np__2|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2)|np__4!=np__0)).#trainneg
cnf(i_0_1243, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,np__1),esk8_0)=k2_xcmplx_0(k3_xcmplx_0(X1,esk8_0),esk8_0)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_20127, plain, (esk8_0=np__4|r1_xxreal_0(esk8_0,np__3)|~r1_xxreal_0(esk8_0,np__4)|~v1_xboole_0(esk8_0))).#trainneg
cnf(i_0_1246, plain, (m1_subset_1(esk4_1(X1),k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(k9_xtuple_0(X1),X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1247, plain, (m1_subset_1(esk4_1(X1),k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(k4_finseq_1(X1),X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3689, plain, (k7_partfun1(k5_numbers,X1,X2)=np__1|r1_xxreal_0(k7_partfun1(k5_numbers,X1,X2),np__0)|~m1_subset_1(k7_partfun1(k5_numbers,X1,X2),k5_numbers)|~r1_xxreal_0(k7_partfun1(k5_numbers,X1,X2),np__1)|~m2_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_17360, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_20495, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__3|~r1_xxreal_0(np__3,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_20742, negated_conjecture, (esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|esk4_1(esk1_0)=np__2|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_20577, plain, (v6_membered(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_20293, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|np__4=np__2|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(k3_finseq_1(esk1_0),np__3))).#trainneg
cnf(i_0_1248, plain, (m1_subset_1(esk3_1(X1),k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(k9_xtuple_0(X1),X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1249, plain, (m1_subset_1(esk3_1(X1),k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(k4_finseq_1(X1),X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1250, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(k9_xtuple_0(X3),X2)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,k3_finseq_1(X3))|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_20159, plain, (r1_xxreal_0(k1_nat_1(np__1,esk2_1(X1)),np__0)|v1_topreal1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_20541, negated_conjecture, (esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__4))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|np__4!=np__2|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_1251, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(k1_relset_1(k5_numbers,X3),X2)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,k3_finseq_1(X3))|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_1259, plain, (k2_zfmisc_1(X1,X2)=k1_card_1(X3)|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1260, plain, (k2_zfmisc_1(X1,X2)=np__0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_17337, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__0|r1_xxreal_0(np__4,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_4580, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|v1_xboole_0(X2)|v3_funct_1(X2)|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(esk4_1(X2))|~v1_xxreal_0(X1)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_3486, plain, (r1_xxreal_0(esk4_1(X1),np__2)|v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(esk4_1(X1),np__1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3780, negated_conjecture, (k2_xcmplx_0(esk4_1(esk1_0),np__0)=esk4_1(esk1_0))).#trainneg
cnf(i_0_9498, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),esk4_1(X1))|r1_xxreal_0(esk4_1(X1),esk3_1(esk1_0))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_9402, plain, (v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(np__4,esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(esk3_1(X1),np__2))).#trainneg
cnf(i_0_3452, plain, (r1_xxreal_0(esk3_1(X1),np__4)|v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(esk3_1(X1),np__1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3786, plain, (v3_ordinal1(X1)|~v7_ordinal1(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_20542, negated_conjecture, (esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__4!=np__2|~r1_xxreal_0(np__3,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_1275, plain, (v1_finset_1(k2_zfmisc_1(X1,X2))|~v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v1_finset_1(X1)|~v1_funct_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_19203, plain, (v3_membered(k1_card_1(X1))|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1276, plain, (v1_relat_1(k9_xtuple_0(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_1287, plain, (v2_xxreal_0(k3_finseq_1(X1))|v1_topreal1(k3_finseq_1(X1))|~v1_xreal_0(k3_finseq_1(X1))|~v1_xreal_0(k1_nat_1(esk2_1(k3_finseq_1(X1)),np__1))|~v2_xxreal_0(k1_nat_1(esk2_1(k3_finseq_1(X1)),np__1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m2_finseq_1(k3_finseq_1(X1),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_9385, plain, (r1_xxreal_0(k7_partfun1(k9_xtuple_0(X1),X2,X3),np__2)|r1_xxreal_0(np__2,k7_partfun1(k9_xtuple_0(X1),X2,X3))|~m2_finseq_1(X2,k9_xtuple_0(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3120, plain, (k7_partfun1(k5_numbers,X1,X2)=np__3|r1_xxreal_0(k7_partfun1(k5_numbers,X1,X2),np__2)|~r1_xxreal_0(k7_partfun1(k5_numbers,X1,X2),np__3)|~m2_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_11547, negated_conjecture, (esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__4|np__3!=np__1)).#trainneg
cnf(i_0_21267, negated_conjecture, (r1_xxreal_0(esk4_1(X1),np__4)|r1_xxreal_0(esk3_1(esk1_0),esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_20823, negated_conjecture, (m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_21182, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(np__0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_20914, plain, (v3_membered(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1288, plain, (k1_nat_1(esk2_1(k3_finseq_1(X1)),np__1)=k3_finseq_1(X1)|v1_topreal1(k3_finseq_1(X1))|~v1_xxreal_0(k1_nat_1(esk2_1(k3_finseq_1(X1)),np__1))|~v1_xxreal_0(k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),k1_nat_1(esk2_1(k3_finseq_1(X1)),np__1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m2_finseq_1(k3_finseq_1(X1),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1289, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|v1_topreal1(k3_finseq_1(X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(k1_nat_1(esk2_1(k3_finseq_1(X2)),np__1))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k1_nat_1(esk2_1(k3_finseq_1(X2)),np__1))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m2_finseq_1(k3_finseq_1(X2),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_21216, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|v1_xboole_0(X2)|v3_funct_1(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_21119, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_card_1(X4)))|~v1_xboole_0(X3)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_10369, negated_conjecture, (esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|np__3!=np__2)).#trainneg
cnf(i_0_1292, plain, (r1_xxreal_0(k1_nat_1(esk2_1(k1_card_1(X1)),np__1),k1_card_1(X1))|v1_topreal1(k1_card_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m2_finseq_1(k1_card_1(X1),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_21266, negated_conjecture, (r1_xxreal_0(esk4_1(X1),k3_finseq_1(esk1_0))|r1_xxreal_0(esk3_1(esk1_0),esk4_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(k3_finseq_1(esk1_0))|~v1_xxreal_0(esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_21231, plain, (r1_xxreal_0(esk4_1(X1),np__2)|v1_xboole_0(X1)|v3_funct_1(X1)|r1_xxreal_0(np__2,esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_6934, negated_conjecture, (esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|np__3!=np__2)).#trainneg
cnf(i_0_21269, negated_conjecture, (esk4_1(X1)=np__1|esk4_1(X1)=esk3_1(esk1_0)|v1_xboole_0(X1)|v3_funct_1(X1)|r1_xxreal_0(esk3_1(esk1_0),esk4_1(X1))|~r1_xxreal_0(esk3_1(esk1_0),np__2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_20846, negated_conjecture, (k3_finseq_1(esk1_0)=np__0|k3_finseq_1(esk1_0)=np__1|k3_finseq_1(esk1_0)=np__2|esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~r1_xxreal_0(np__4,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_9526, plain, (k19_euclid(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)),k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)))=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k1_nat_1(np__1,X2)=np__0|k1_nat_1(np__1,X2)=np__1|k1_nat_1(np__1,X2)=np__4|k1_nat_1(np__1,X2)=np__2|r1_xxreal_0(np__4,X2)|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(np__3,k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_4572, plain, (v1_xboole_0(X2)|v3_funct_1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(esk4_1(X2))|~r1_xxreal_0(esk4_1(X2),X1)|~v1_xxreal_0(k1_nat_1(np__1,X1))|~r1_xxreal_0(k1_nat_1(np__1,X1),np__1)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_17180, plain, (v7_ordinal1(k1_card_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_10223, negated_conjecture, (esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|np__3!=np__4)).#trainneg
cnf(i_0_21622, negated_conjecture, (~v1_xboole_0(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)))).#trainneg
cnf(i_0_1296, plain, (r1_xxreal_0(k2_xcmplx_0(esk2_1(k3_finseq_1(X1)),np__1),k3_finseq_1(X1))|v1_topreal1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m2_finseq_1(k3_finseq_1(X1),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(esk2_1(k3_finseq_1(X1))))).#trainneg
cnf(i_0_1297, plain, (r1_xxreal_0(k2_xcmplx_0(np__1,esk2_1(k3_finseq_1(X1))),k3_finseq_1(X1))|v1_topreal1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m2_finseq_1(k3_finseq_1(X1),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk2_1(k3_finseq_1(X1)),k5_numbers))).#trainneg
cnf(i_0_21992, plain, (v1_xboole_0(X1)|v3_funct_1(X1)|r1_xxreal_0(np__2,esk4_1(X1))|~r1_xxreal_0(np__3,esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_21584, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|r1_xxreal_0(esk4_1(X1),np__4)|~v1_xxreal_0(k3_finseq_1(X1))|~v1_xxreal_0(esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_22185, plain, (v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(k1_nat_1(X2,np__1))|~v7_ordinal1(X2)|~r1_xxreal_0(k1_nat_1(X2,np__1),np__1)|~r1_xxreal_0(esk4_1(X1),X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_21803, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_card_1(X2)))|~v1_xboole_0(X3)|~v1_xboole_0(k1_card_1(X2)))).#trainneg
cnf(i_0_10096, negated_conjecture, (esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1|np__3!=np__1)).#trainneg
cnf(i_0_20278, plain, (esk8_0=np__0|esk8_0=np__1|~r1_xxreal_0(np__0,esk8_0)|~m1_subset_1(esk8_0,k5_numbers)|~r1_xxreal_0(esk8_0,np__1))).#trainneg
cnf(i_0_1299, plain, (v1_xreal_0(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k1_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_1300, plain, (v1_xreal_0(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,X2)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,X2)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k1_nat_1(np__1,X2),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_16372, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|np__4!=np__2)).#trainneg
cnf(i_0_19491, negated_conjecture, (k7_partfun1(X1,np__4,np__4)=k7_partfun1(X1,np__4,np__1)|v1_xboole_0(X1)|~m2_finseq_1(np__4,X1)|~v1_finseq_6(np__4,X1))).#trainneg
cnf(i_0_19308, plain, (r1_xxreal_0(k7_partfun1(k9_xtuple_0(X1),X2,X3),k1_card_1(X1))|v1_xboole_0(k9_xtuple_0(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m2_finseq_1(X2,k9_xtuple_0(X1)))).#trainneg
cnf(i_0_5141, negated_conjecture, (r1_xxreal_0(esk3_1(X1),esk3_1(esk1_0))|esk3_1(X1)=np__2|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(esk3_1(X1))|~r1_xxreal_0(esk3_1(X1),np__2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_7091, negated_conjecture, (esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|np__3!=np__4)).#trainneg
cnf(i_0_20281, plain, (esk8_0=np__1|esk8_0=np__0|~r1_xxreal_0(np__1,esk8_0)|~m1_subset_1(esk8_0,k5_numbers)|~r1_xxreal_0(esk8_0,np__1))).#trainneg
cnf(i_0_1301, plain, (v1_xreal_0(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k2_xcmplx_0(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_1304, plain, (k2_xcmplx_0(esk2_1(X1),np__1)=k3_finseq_1(X1)|v1_topreal1(X1)|~v1_xxreal_0(k2_xcmplx_0(esk2_1(X1),np__1))|~v1_xxreal_0(k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),k2_xcmplx_0(esk2_1(X1),np__1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(esk2_1(X1)))).#trainneg
cnf(i_0_22837, negated_conjecture, (esk3_1(esk1_0)=np__2|np__4!=np__2)).#trainneg
cnf(i_0_1305, plain, (k2_xcmplx_0(np__1,esk2_1(X1))=k3_finseq_1(X1)|v1_topreal1(X1)|~v1_xxreal_0(k2_xcmplx_0(np__1,esk2_1(X1)))|~v1_xxreal_0(k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),k2_xcmplx_0(np__1,esk2_1(X1)))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk2_1(X1),k5_numbers))).#trainneg
cnf(i_0_22570, negated_conjecture, (k3_finseq_1(X1)=esk3_1(esk1_0)|r1_xxreal_0(esk4_1(X1),np__4)|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),esk3_1(esk1_0))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4666, plain, (esk4_1(X1)=np__1|esk4_1(X1)=np__2|v1_xboole_0(X1)|v3_funct_1(X1)|r1_xxreal_0(np__3,esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_21952, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|r1_xxreal_0(esk4_1(X1),k3_finseq_1(esk1_0))|~v1_xxreal_0(k3_finseq_1(X1))|~v1_xxreal_0(esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_xxreal_0(k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_21810, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk11_1(X4)))|~v1_xboole_0(X3)|~v1_xboole_0(X2)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_1308, plain, (k1_nat_1(esk2_1(X1),np__1)=k1_card_1(X1)|v1_topreal1(X1)|~v1_xxreal_0(k1_nat_1(esk2_1(X1),np__1))|~v1_xxreal_0(k1_card_1(X1))|~r1_xxreal_0(k1_card_1(X1),k1_nat_1(esk2_1(X1),np__1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_20472, plain, (esk8_0=np__3|esk8_0=np__4|~r1_xxreal_0(np__3,esk8_0)|~m1_subset_1(esk8_0,k5_numbers)|~r1_xxreal_0(esk8_0,np__4))).#trainneg
cnf(i_0_1311, plain, (m2_subset_1(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),k1_numbers,k5_numbers)|~m1_subset_1(k3_xcmplx_0(X3,X2),k5_numbers)|~v7_ordinal1(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1312, plain, (m2_subset_1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),k1_numbers,k5_numbers)|~m1_subset_1(X3,k5_numbers)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4302, plain, (v1_xboole_0(X1)|v3_funct_1(X1)|esk3_1(X1)=np__2|~r1_xxreal_0(np__2,esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(esk3_1(X1),np__2))).#trainneg
cnf(i_0_21799, negated_conjecture, (r1_xxreal_0(X1,k3_finseq_1(esk1_0))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_21241, negated_conjecture, (m1_subset_1(esk4_1(esk1_0),k5_numbers)|~v7_ordinal1(esk4_1(esk1_0)))).#trainneg
cnf(i_0_19630, plain, (k7_partfun1(k9_xtuple_0(X1),X2,X3)=np__1|v1_xboole_0(k9_xtuple_0(X1))|~v1_xxreal_0(k7_partfun1(k9_xtuple_0(X1),X2,X3))|~r1_xxreal_0(k7_partfun1(k9_xtuple_0(X1),X2,X3),np__1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m2_finseq_1(X2,k9_xtuple_0(X1)))).#trainneg
cnf(i_0_3377, negated_conjecture, (r1_xxreal_0(esk3_1(esk1_0),np__2)|~r1_xxreal_0(esk3_1(esk1_0),np__1))).#trainneg
cnf(i_0_8940, plain, (v1_relat_1(k1_card_1(X1))|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_20475, plain, (esk8_0=np__4|esk8_0=np__3|~r1_xxreal_0(np__4,esk8_0)|~m1_subset_1(esk8_0,k5_numbers)|~r1_xxreal_0(esk8_0,np__4))).#trainneg
cnf(i_0_7579, negated_conjecture, (esk3_1(esk1_0)=np__3|r1_xxreal_0(esk3_1(esk1_0),np__2)|r1_xxreal_0(np__4,esk3_1(esk1_0)))).#trainneg
cnf(i_0_4182, plain, (v1_xboole_0(X1)|v3_funct_1(X1)|esk4_1(X1)=np__4|~r1_xxreal_0(np__4,esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(esk4_1(X1),np__4))).#trainneg
cnf(i_0_3695, plain, (k7_partfun1(k5_numbers,X1,X2)=np__3|k7_partfun1(k5_numbers,X1,X2)=np__4|~v1_xxreal_0(k7_partfun1(k5_numbers,X1,X2))|~r1_xxreal_0(np__3,k7_partfun1(k5_numbers,X1,X2))|~r1_xxreal_0(k7_partfun1(k5_numbers,X1,X2),np__4)|~m2_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_9669, plain, (v1_funct_1(k1_card_1(X1))|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_22795, plain, (esk8_0=np__1|esk8_0=np__0|r1_xxreal_0(np__1,esk8_0)|~m1_subset_1(esk8_0,k5_numbers)|~r1_xxreal_0(np__0,esk8_0))).#trainneg
cnf(i_0_1336, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_relset_1(k5_numbers,X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_4559, plain, (r1_xxreal_0(X1,np__1)|k7_partfun1(k5_numbers,X2,X3)=np__2|~v1_xxreal_0(k7_partfun1(k5_numbers,X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k7_partfun1(k5_numbers,X2,X3))|~r1_xxreal_0(k7_partfun1(k5_numbers,X2,X3),np__2)|~m2_finseq_1(X2,k5_numbers))).#trainneg
cnf(i_0_5014, plain, (k7_partfun1(X1,X2,X3)=np__3|r1_xxreal_0(k7_partfun1(X1,X2,X3),np__2)|k7_partfun1(X1,X2,X3)=np__4|~v2_membered(X1)|~m1_subset_1(k7_partfun1(X1,X2,X3),k5_numbers)|~m2_finseq_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_xxreal_0(k7_partfun1(X1,X2,X3),np__4))).#trainneg
cnf(i_0_3489, plain, (r1_xxreal_0(esk4_1(X1),np__3)|v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(esk4_1(X1),np__1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_22571, negated_conjecture, (r1_xxreal_0(X1,k3_finseq_1(X2))|r1_xxreal_0(esk4_1(X2),np__4)|v1_xboole_0(X2)|v3_funct_1(X2)|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk3_1(esk1_0))|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2739, plain, (r1_xxreal_0(np__1,k3_finseq_1(X1))|r1_xxreal_0(k3_finseq_1(X1),np__1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4593, plain, (r1_xxreal_0(esk3_1(X1),np__2)|v1_xboole_0(X1)|v3_funct_1(X1)|r1_xxreal_0(np__2,esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_9988, plain, (v1_finseq_1(k1_card_1(X1))|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_23135, negated_conjecture, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__4!=np__2)).#trainneg
cnf(i_0_23630, negated_conjecture, (m1_subset_1(esk4_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1343, plain, (k2_zfmisc_1(X1,X2)=k1_card_1(X3)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1344, plain, (k2_zfmisc_1(X1,X2)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1347, plain, (r1_xxreal_0(esk6_1(X1),esk6_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_23363, plain, (r1_xxreal_0(np__3,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|esk4_1(X1)=np__2|esk4_1(X1)=np__1|~v1_xxreal_0(k3_finseq_1(X1))|~v1_xxreal_0(esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_23359, plain, (esk4_1(X1)=np__3|esk4_1(X1)=np__2|esk4_1(X1)=np__1|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(esk4_1(X1))|~r1_xxreal_0(esk4_1(X1),np__3)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1348, plain, (r1_xxreal_0(esk8_0,esk6_1(X1))|r1_xxreal_0(esk6_1(X1),esk8_0)|~v2_membered(X1))).#trainneg
cnf(i_0_17858, plain, (v5_relat_1(esk11_1(X3),X2)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1351, plain, (esk6_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=k1_card_1(X3)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1352, plain, (esk6_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_23137, negated_conjecture, (r2_hidden(np__2,k9_xtuple_0(esk1_0))|np__4!=np__2)).#trainneg
cnf(i_0_1353, plain, (X1=esk6_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1368, plain, (k10_finseq_1(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,X2))),k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)))=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,X2))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,X2)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,X2)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k1_nat_1(np__1,X2),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_1370, plain, (v7_ordinal1(k9_xtuple_0(X1))|~v6_membered(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4882, plain, (r1_xxreal_0(X1,esk3_1(X2))|v1_xboole_0(X2)|v3_funct_1(X2)|~v1_xxreal_0(esk3_1(X2))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__0)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3057, plain, (k2_xcmplx_0(np__1,X1)=np__3|k1_nat_1(np__1,X1)=np__2|k1_nat_1(np__1,X1)=np__4|k1_nat_1(np__1,X1)=np__1|k1_nat_1(np__1,X1)=np__0|r1_xxreal_0(np__4,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1371, plain, (v1_xxreal_0(k9_xtuple_0(X1))|~v2_membered(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1372, plain, (v3_ordinal1(k9_xtuple_0(X1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1373, plain, (v1_finset_1(k9_xtuple_0(X1))|~v5_finset_1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_24123, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|~r1_xxreal_0(k3_finseq_1(esk1_0),np__2))).#trainneg
cnf(i_0_24219, plain, (esk4_1(X1)=np__1|esk4_1(X1)=np__2|esk4_1(X1)=np__3|v1_xboole_0(X1)|v3_funct_1(X1)|esk4_1(X1)=np__4|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(esk4_1(X1),np__4))).#trainneg
cnf(i_0_24236, plain, (v5_relat_1(esk11_1(X1),X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1374, plain, (v1_xreal_0(k9_xtuple_0(X1))|~v3_membered(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1375, plain, (v1_xcmplx_0(k9_xtuple_0(X1))|~v1_membered(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1376, plain, (v1_int_1(k9_xtuple_0(X1))|~v5_membered(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_24218, plain, (esk4_1(X1)=np__1|esk4_1(X1)=np__2|r1_xxreal_0(np__3,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_24221, plain, (esk4_1(X1)=np__1|esk4_1(X1)=np__2|esk4_1(X1)=np__3|v1_xboole_0(X1)|v3_funct_1(X1)|r1_xxreal_0(np__4,esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_24031, plain, (v4_relat_1(esk6_1(k1_zfmisc_1(np__0)),X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1379, plain, (v3_membered(k9_xtuple_0(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_24124, negated_conjecture, (k3_finseq_1(esk1_0)=np__3|r1_xxreal_0(np__4,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_24446, plain, (esk3_1(X1)=X2|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(esk3_1(X1))|~r1_xxreal_0(esk3_1(X1),X2)|~r1_xxreal_0(X2,np__0)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_23864, negated_conjecture, (v7_ordinal1(esk4_1(esk1_0)))).#trainneg
cnf(i_0_24481, plain, (k2_xcmplx_0(np__1,X1)=np__4|k1_nat_1(np__1,X1)=np__0|k1_nat_1(np__1,X1)=np__1|k1_nat_1(np__1,X1)=np__2|k2_xcmplx_0(np__1,X1)=np__3|r1_xxreal_0(np__4,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_24447, plain, (r1_xxreal_0(X1,esk3_1(X2))|v1_xboole_0(X2)|v3_funct_1(X2)|~v1_xxreal_0(esk3_1(X2))|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X3,np__0)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_10543, plain, (k10_finseq_1(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)),k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X3)))=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X3))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k1_nat_1(np__1,X3)=np__2|k1_nat_1(np__1,X3)=np__4|k1_nat_1(np__1,X3)=np__1|k1_nat_1(np__1,X3)=np__0|k1_nat_1(np__1,X2)=np__2|k1_nat_1(np__1,X2)=np__4|k1_nat_1(np__1,X2)=np__1|k1_nat_1(np__1,X2)=np__0|r1_xxreal_0(np__4,X3)|r1_xxreal_0(np__4,X2)|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X3,k5_numbers)|~r1_xxreal_0(np__3,k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X3)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_24609, negated_conjecture, (np__3=np__4|np__4!=np__0)).#trainneg
cnf(i_0_8283, negated_conjecture, (r1_xxreal_0(k1_nat_1(np__1,X1),np__4)|r1_xxreal_0(esk4_1(esk1_0),X1)|~v1_xxreal_0(k1_nat_1(np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1381, plain, (v1_membered(k9_xtuple_0(X1))|~v1_membered(k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_24869, negated_conjecture, (~r1_xxreal_0(esk4_1(esk1_0),np__0))).#trainneg
cnf(i_0_21033, plain, (v1_card_1(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1386, plain, (v1_finset_1(k19_euclid(X1,X2))|~v1_funct_1(k19_euclid(X1,X2))|~v1_relat_1(k19_euclid(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_24876, negated_conjecture, (~r1_xxreal_0(np__2,esk4_1(esk1_0))|~r1_xxreal_0(esk4_1(esk1_0),np__1))).#trainneg
cnf(i_0_8284, negated_conjecture, (r1_xxreal_0(k1_nat_1(X1,np__1),np__4)|r1_xxreal_0(esk4_1(esk1_0),X1)|~v1_xxreal_0(k1_nat_1(X1,np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_23442, plain, (esk8_0=np__4|esk8_0=np__3|r1_xxreal_0(np__4,esk8_0)|~m1_subset_1(esk8_0,k5_numbers)|~r1_xxreal_0(np__3,esk8_0))).#trainneg
cnf(i_0_21047, plain, (v1_funct_1(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_21048, plain, (v1_finseq_1(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1392, plain, (m2_subset_1(k1_card_1(X1),k1_numbers,k5_numbers)|~v1_finseq_1(k1_card_1(X1))|~v1_funct_1(k1_card_1(X1))|~v1_relat_1(k1_card_1(X1)))).#trainneg
cnf(i_0_24573, negated_conjecture, (esk4_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__4|~r1_xxreal_0(np__3,esk4_1(esk1_0)))).#trainneg
cnf(i_0_24877, negated_conjecture, (~r1_xxreal_0(np__3,esk4_1(esk1_0))|~r1_xxreal_0(esk4_1(esk1_0),np__2))).#trainneg
cnf(i_0_24515, negated_conjecture, (k7_partfun1(X1,esk1_0,np__3)=k7_partfun1(X1,esk1_0,np__1)|v1_xboole_0(X1)|k3_finseq_1(esk1_0)=np__4|~m2_finseq_1(esk1_0,X1)|~v1_finseq_6(esk1_0,X1))).#trainneg
cnf(i_0_8278, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(np__1,X1),np__4)|r1_xxreal_0(esk4_1(esk1_0),X1)|~m1_subset_1(X1,k5_numbers)|~v1_xxreal_0(k2_xcmplx_0(np__1,X1)))).#trainneg
cnf(i_0_21049, plain, (v1_finset_1(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_21050, plain, (v7_ordinal1(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_24774, plain, (r1_xxreal_0(np__4,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|esk4_1(X1)=np__3|esk4_1(X1)=np__2|esk4_1(X1)=np__1|~v1_xxreal_0(k3_finseq_1(X1))|~v1_xxreal_0(esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_24572, negated_conjecture, (esk3_1(esk1_0)=np__3|k3_finseq_1(esk1_0)=np__4|~r1_xxreal_0(np__3,esk3_1(esk1_0)))).#trainneg
cnf(i_0_24521, negated_conjecture, (k3_finseq_1(np__3)=np__3|k3_finseq_1(esk1_0)=np__4)).#trainneg
cnf(i_0_24875, negated_conjecture, (~r1_xxreal_0(np__4,esk4_1(esk1_0))|~r1_xxreal_0(esk4_1(esk1_0),np__3))).#trainneg
cnf(i_0_8279, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,np__1),np__4)|r1_xxreal_0(esk4_1(esk1_0),X1)|~v1_xxreal_0(k2_xcmplx_0(X1,np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1418, plain, (m2_subset_1(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),k1_numbers,k5_numbers)|~m1_subset_1(k3_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(k3_xcmplx_0(X3,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_21051, plain, (v6_membered(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1406, plain, (m1_subset_1(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)),k1_numbers)|m1_subset_1(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)),k1_numbers)|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k1_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_1407, plain, (k19_euclid(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)),k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1))))=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1))|m1_subset_1(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)),k1_numbers)|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k1_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_25458, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|esk4_1(esk1_0)=np__3|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1)).#trainneg
cnf(i_0_25476, negated_conjecture, (k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3)=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)|k3_finseq_1(esk1_0)=np__4)).#trainneg
cnf(i_0_25582, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|esk3_1(esk1_0)=np__3|esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1)).#trainneg
cnf(i_0_25293, plain, (v1_card_1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_8280, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,np__1),np__4)|r1_xxreal_0(esk4_1(esk1_0),X1)|~m1_subset_1(X1,k5_numbers)|~v1_xxreal_0(k2_xcmplx_0(X1,np__1)))).#trainneg
cnf(i_0_1419, plain, (m2_subset_1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),k1_numbers,k5_numbers)|~m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_21658, plain, (v3_membered(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_23891, plain, (v1_card_1(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1408, plain, (v1_xreal_0(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)))|m1_subset_1(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)),k1_numbers)|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k1_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_1409, plain, (k10_finseq_1(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)),k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1))))=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1))|m1_subset_1(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)),k1_numbers)|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k1_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_25573, plain, (esk4_1(X1)=np__1|esk4_1(X1)=np__2|esk4_1(X1)=np__3|r1_xxreal_0(np__4,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_25416, plain, (v1_funct_1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1431, plain, (v7_ordinal1(esk6_1(X1))|~v1_finset_1(esk6_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_23906, plain, (v1_funct_1(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_23907, plain, (v1_finseq_1(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1410, plain, (v7_ordinal1(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)))|m1_subset_1(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)),k1_numbers)|~v6_membered(k1_numbers)|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k1_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_25425, plain, (v1_finseq_1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_23868, negated_conjecture, (v1_xreal_0(esk4_1(esk1_0)))).#trainneg
cnf(i_0_23910, plain, (v1_finset_1(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_23912, plain, (v7_ordinal1(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1414, plain, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(np__1,X2)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|m1_subset_1(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)),k1_numbers)|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(np__1,X2)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k2_xcmplx_0(np__1,X2),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_24858, plain, (esk3_1(X1)=k3_finseq_1(X1)|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),np__0)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(esk3_1(X1)))).#trainneg
cnf(i_0_26312, plain, (v1_card_1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_26045, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)|~r1_xxreal_0(np__3,k3_finseq_1(esk1_0)))).#trainneg
cnf(i_0_1433, plain, (v1_finseq_6(k1_card_1(X1),X2)|v1_xboole_0(X2)|k7_partfun1(X2,k1_card_1(X1),k1_card_1(X1))!=k7_partfun1(X2,k1_card_1(X1),np__1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m2_finseq_1(k1_card_1(X1),X2))).#trainneg
cnf(i_0_26049, negated_conjecture, (esk3_1(esk1_0)=np__2|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__4)).#trainneg
cnf(i_0_26050, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__4|k3_finseq_1(esk1_0)=np__4)).#trainneg
cnf(i_0_26402, plain, (v1_funct_1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_25919, negated_conjecture, (esk4_1(esk1_0)=np__1|esk4_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)))).#trainneg
cnf(i_0_20684, plain, (esk8_0=np__3|esk8_0=np__4|~r1_xxreal_0(np__3,esk8_0)|~r1_xxreal_0(esk8_0,np__4)|~v1_xboole_0(esk8_0))).#trainneg
cnf(i_0_23913, plain, (v3_membered(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_23915, plain, (v6_membered(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1438, plain, (r1_xxreal_0(esk6_1(k1_relset_1(k5_numbers,X1)),k3_finseq_1(X1))|v1_xboole_0(k1_relset_1(k5_numbers,X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(esk6_1(k1_relset_1(k5_numbers,X1))))).#trainneg
cnf(i_0_26054, negated_conjecture, (esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|np__3!=np__1)).#trainneg
cnf(i_0_26058, negated_conjecture, (esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)))).#trainneg
cnf(i_0_26411, plain, (v1_finseq_1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1462, plain, (v2_xxreal_0(k3_finseq_1(X1))|v1_topreal1(X1)|~v1_xreal_0(k3_finseq_1(X1))|~v1_xreal_0(k2_xcmplx_0(esk2_1(X1),np__1))|~v2_xxreal_0(k2_xcmplx_0(esk2_1(X1),np__1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_20687, plain, (esk8_0=np__4|esk8_0=np__3|~r1_xxreal_0(np__4,esk8_0)|~r1_xxreal_0(esk8_0,np__4)|~v1_xboole_0(esk8_0))).#trainneg
cnf(i_0_1464, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|v1_topreal1(X2)|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(k2_xcmplx_0(esk2_1(X2),np__1))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k2_xcmplx_0(esk2_1(X2),np__1))|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_26516, plain, (esk3_1(X1)=k3_finseq_1(X1)|v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(k3_finseq_1(X1),np__0)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_26463, plain, (v1_finset_1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1467, plain, (r1_xxreal_0(k2_xcmplx_0(esk2_1(X1),np__1),k1_card_1(X1))|v1_topreal1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1468, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk16_1(X2))|v1_xboole_0(X2)|~m1_subset_1(X1,esk16_1(X2)))).#trainneg
cnf(i_0_26472, plain, (v7_ordinal1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_21828, negated_conjecture, (esk4_1(esk1_0)=np__2|np__3!=np__2)).#trainneg
cnf(i_0_1470, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk6_1(k1_zfmisc_1(X2)))|v1_xboole_0(X2)|~m1_subset_1(X1,esk6_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_25683, negated_conjecture, (~r1_xxreal_0(np__4,esk4_1(esk1_0))|~r1_xxreal_0(esk4_1(esk1_0),np__1))).#trainneg
cnf(i_0_1471, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|v1_xboole_0(u1_pre_topc(X2))|v1_xboole_0(k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_pre_topc(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_21141, plain, (v2_membered(X1)|~v2_membered(np__0)|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_25548, plain, (v1_finset_1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_22012, negated_conjecture, (esk3_1(esk1_0)=np__2|np__3!=np__2)).#trainneg
cnf(i_0_2288, plain, (esk8_0=np__3|r1_xxreal_0(esk8_0,np__2)|~v7_ordinal1(esk8_0)|~r1_xxreal_0(esk8_0,np__3))).#trainneg
cnf(i_0_1478, plain, (v1_int_1(k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),X4),X3))|~v1_int_1(k3_xcmplx_0(X4,X3))|~v1_int_1(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1479, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)),k2_xcmplx_0(k3_xcmplx_0(X4,X3),X5))=k2_xcmplx_0(k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),X4),X3),X5)|~v1_xcmplx_0(X5)|~v1_xcmplx_0(k3_xcmplx_0(X4,X3))|~v1_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_26557, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2)).#trainneg
cnf(i_0_26680, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|esk4_1(esk1_0)=np__2|esk4_1(esk1_0)=np__1)).#trainneg
cnf(i_0_26794, plain, (v3_membered(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_22215, negated_conjecture, (esk4_1(esk1_0)=np__2|np__3!=np__4)).#trainneg
cnf(i_0_27034, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__3!=np__2)).#trainneg
cnf(i_0_1480, plain, (m2_subset_1(k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),X4),X3),k1_numbers,k5_numbers)|~m1_subset_1(k3_xcmplx_0(X4,X3),k5_numbers)|~v7_ordinal1(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_27153, plain, (v2_membered(X1)|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1481, plain, (m2_subset_1(k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),X4),X3),k1_numbers,k5_numbers)|~m1_subset_1(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)),k5_numbers)|~v7_ordinal1(k3_xcmplx_0(X4,X3))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1482, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)),k3_xcmplx_0(X4,k3_xcmplx_0(X5,X3)))=k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),k3_xcmplx_0(X4,X5)),X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X4,X5))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X5)|~v1_xcmplx_0(X4))).#trainneg
cnf(i_0_27516, negated_conjecture, (esk3_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4)).#trainneg
cnf(i_0_27628, negated_conjecture, (esk4_1(esk1_0)=np__2|k3_finseq_1(esk1_0)=np__4)).#trainneg
cnf(i_0_27904, plain, (v2_membered(np__0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_27899, plain, (v2_membered(esk11_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_22713, negated_conjecture, (esk4_1(esk1_0)=np__2|np__3!=np__1)).#trainneg
cnf(i_0_27161, negated_conjecture, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__3!=np__2)).#trainneg
cnf(i_0_1483, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),np__0),np__0)=k2_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__0)),np__0)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_27905, plain, (v2_membered(k3_finseq_1(X1))|~v1_xboole_0(X2)|~v1_finseq_1(X1)|~v1_funct_1(k3_finseq_1(X1))|~v1_funct_1(X1)|~v1_relat_1(k3_finseq_1(X1))|~v1_relat_1(X1)|~v1_xboole_0(k3_finseq_1(X1)))).#trainneg
cnf(i_0_1484, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),np__0),np__1)=k2_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__1)),np__0)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1485, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),np__0),np__2)=k2_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__2)),np__0)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_26803, plain, (v6_membered(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_22986, negated_conjecture, (esk3_1(esk1_0)=np__2|np__3!=np__4)).#trainneg
cnf(i_0_27730, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__3!=np__4)).#trainneg
cnf(i_0_1486, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),np__3),np__1)=k2_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__1)),np__3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_28564, plain, (v2_membered(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(k3_finseq_1(X1))|~v1_funct_1(X1)|~v1_relat_1(k3_finseq_1(X1))|~v1_relat_1(X1)|~v1_xboole_0(k3_finseq_1(X1)))).#trainneg
cnf(i_0_1487, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),np__1),esk8_0)=k2_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,esk8_0)),esk8_0)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1488, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),np__1),np__0)=k2_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__0)),np__0)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_28266, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)))).#trainneg
cnf(i_0_28160, negated_conjecture, (k3_finseq_1(esk1_0)=np__4|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)))).#trainneg
cnf(i_0_28163, negated_conjecture, (r2_hidden(np__2,k9_xtuple_0(esk1_0))|k3_finseq_1(esk1_0)=np__4)).#trainneg
cnf(i_0_28238, negated_conjecture, (k3_finseq_1(X1)=np__2|r1_xxreal_0(esk4_1(X1),np__4)|v1_xboole_0(X1)|v3_funct_1(X1)|k3_finseq_1(esk1_0)=np__4|~v1_xxreal_0(k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),np__2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_28357, plain, (v2_membered(np__0))).#trainneg
cnf(i_0_28762, negated_conjecture, (k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__3!=np__4)).#trainneg
cnf(i_0_1489, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),np__1),np__1)=k2_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__1)),np__1)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1490, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),np__1),np__2)=k2_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__2)),np__2)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1491, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),np__4),np__1)=k2_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__1)),np__4)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__4)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_23274, negated_conjecture, (esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|esk4_1(esk1_0)=np__2)).#trainneg
cnf(i_0_27013, plain, (m1_subset_1(esk6_1(esk16_1(X1)),X1)|v1_xboole_0(X1))).#trainneg
cnf(i_0_1492, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),np__2),np__0)=k2_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__0)),np__0)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1493, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),np__2),np__1)=k2_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__1)),np__2)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1494, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),np__2),np__2)=k2_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__2)),np__4)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6791, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|esk3_1(esk1_0)=np__4|esk3_1(esk1_0)=np__1|esk3_1(esk1_0)=np__2|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2)))).#trainneg
cnf(i_0_29298, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__3!=np__4)).#trainneg
cnf(i_0_29164, negated_conjecture, (k3_xcmplx_0(k2_xcmplx_0(esk3_1(esk1_0),np__1),np__0)=np__0)).#trainneg
cnf(i_0_27001, plain, (r1_xxreal_0(k2_xcmplx_0(esk2_1(X1),np__1),np__0)|v1_topreal1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1495, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,k3_xcmplx_0(X3,X4))),k3_xcmplx_0(X5,X4))=k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)),X5),X4)|~v1_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X5)|~v1_xcmplx_0(k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1496, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,np__0),X2),np__0)=k2_xcmplx_0(k3_xcmplx_0(X1,np__0),k3_xcmplx_0(X2,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X1,np__0))|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1497, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,np__0),X2),np__1)=k2_xcmplx_0(k3_xcmplx_0(X1,np__0),k3_xcmplx_0(X2,np__1))|~v1_xcmplx_0(k3_xcmplx_0(X1,np__0))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_30093, plain, (k3_xcmplx_0(np__3,np__2)=k2_xcmplx_0(np__2,np__4))).#trainneg
cnf(i_0_30095, plain, (k3_xcmplx_0(np__4,np__2)=k2_xcmplx_0(np__4,np__4))).#trainneg
cnf(i_0_29228, negated_conjecture, (v1_xxreal_0(X1)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__4)|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_21142, plain, (v1_membered(X1)|~v1_membered(np__0)|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_25557, plain, (v7_ordinal1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_28374, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))!=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|np__3!=np__1)).#trainneg
cnf(i_0_1498, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,np__0),X2),np__2)=k2_xcmplx_0(k3_xcmplx_0(X1,np__0),k3_xcmplx_0(X2,np__2))|~v1_xcmplx_0(k3_xcmplx_0(X1,np__0))|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_26787, plain, (esk8_0=np__4|esk8_0=np__3|r1_xxreal_0(np__4,esk8_0)|~r1_xxreal_0(np__3,esk8_0)|~v1_xboole_0(esk8_0))).#trainneg
cnf(i_0_1499, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,np__3),X2),np__1)=k2_xcmplx_0(k3_xcmplx_0(X1,np__3),k3_xcmplx_0(X2,np__1))|~v1_xcmplx_0(k3_xcmplx_0(X1,np__3))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(np__3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1500, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,np__1),X2),esk8_0)=k2_xcmplx_0(k3_xcmplx_0(X1,esk8_0),k3_xcmplx_0(X2,esk8_0))|~v1_xcmplx_0(k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_30780, negated_conjecture, (X1=np__4|r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__4)|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_30788, negated_conjecture, (r1_xxreal_0(X1,np__2)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__4)|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_30782, negated_conjecture, (X1=np__3|r1_xxreal_0(X1,np__2)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X1,np__4)|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_30786, negated_conjecture, (r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X1,np__2)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__4)|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_30781, negated_conjecture, (X1=np__2|r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__2)|~r1_xxreal_0(X1,np__4)|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_30779, negated_conjecture, (X1=np__1|r1_xxreal_0(X1,np__0)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(X1,np__4)|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_30791, negated_conjecture, (r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__4)|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_30796, plain, (v1_membered(X1)|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_30667, plain, (v1_int_1(k2_xcmplx_0(np__2,np__4)))).#trainneg
cnf(i_0_1501, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,np__1),X2),np__0)=k2_xcmplx_0(k3_xcmplx_0(X1,np__0),k3_xcmplx_0(X2,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_27903, plain, (v2_membered(k1_card_1(X1))|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1502, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,np__1),X2),np__1)=k2_xcmplx_0(k3_xcmplx_0(X1,np__1),k3_xcmplx_0(X2,np__1))|~v1_xcmplx_0(k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1503, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,np__1),X2),np__2)=k2_xcmplx_0(k3_xcmplx_0(X1,np__2),k3_xcmplx_0(X2,np__2))|~v1_xcmplx_0(k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_30144, negated_conjecture, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,X1))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,X1))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|k1_nat_1(np__1,X1)=np__0|k1_nat_1(np__1,X1)=np__1|k1_nat_1(np__1,X1)=np__4|k1_nat_1(np__1,X1)=np__2|k1_nat_1(np__1,np__1)=np__0|k1_nat_1(np__1,np__1)=np__1|k1_nat_1(np__1,np__1)=np__4|k1_nat_1(np__1,np__1)=np__2|r1_xxreal_0(np__4,X1)|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_31135, negated_conjecture, (esk4_1(esk1_0)=np__4|r1_xxreal_0(esk4_1(esk1_0),np__3))).#trainneg
cnf(i_0_31163, negated_conjecture, (r1_xxreal_0(esk4_1(esk1_0),np__2)|~r1_xxreal_0(esk4_1(esk1_0),np__1))).#trainneg
cnf(i_0_31192, negated_conjecture, (esk4_1(esk1_0)=np__3|r1_xxreal_0(esk4_1(esk1_0),np__2)|~r1_xxreal_0(esk4_1(esk1_0),np__3))).#trainneg
cnf(i_0_31249, negated_conjecture, (esk4_1(esk1_0)=np__2|r1_xxreal_0(esk4_1(esk1_0),np__1)|~r1_xxreal_0(esk4_1(esk1_0),np__2))).#trainneg
cnf(i_0_31220, negated_conjecture, (r1_xxreal_0(esk4_1(esk1_0),np__3)|~r1_xxreal_0(esk4_1(esk1_0),np__2))).#trainneg
cnf(i_0_31336, plain, (v1_membered(np__0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_31331, plain, (v1_membered(esk11_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1504, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,np__4),X2),np__1)=k2_xcmplx_0(k3_xcmplx_0(X1,np__4),k3_xcmplx_0(X2,np__1))|~v1_xcmplx_0(k3_xcmplx_0(X1,np__4))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(np__4)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_31337, plain, (v1_membered(k3_finseq_1(X1))|~v1_xboole_0(X2)|~v1_finseq_1(X1)|~v1_funct_1(k3_finseq_1(X1))|~v1_funct_1(X1)|~v1_relat_1(k3_finseq_1(X1))|~v1_relat_1(X1)|~v1_xboole_0(k3_finseq_1(X1)))).#trainneg
cnf(i_0_1505, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,np__2),X2),np__0)=k2_xcmplx_0(k3_xcmplx_0(X1,np__0),k3_xcmplx_0(X2,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X1,np__2))|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1506, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,np__2),X2),np__1)=k2_xcmplx_0(k3_xcmplx_0(X1,np__2),k3_xcmplx_0(X2,np__1))|~v1_xcmplx_0(k3_xcmplx_0(X1,np__2))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_31305, negated_conjecture, (r1_xxreal_0(esk4_1(esk1_0),np__3)|~r1_xxreal_0(esk4_1(esk1_0),np__1))).#trainneg
cnf(i_0_31300, negated_conjecture, (X1=np__1|r1_xxreal_0(X1,np__0)|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(X1,np__4)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_31424, plain, (v2_membered(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1507, plain, (k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,np__2),X2),np__2)=k2_xcmplx_0(k3_xcmplx_0(X1,np__4),k3_xcmplx_0(X2,np__2))|~v1_xcmplx_0(k3_xcmplx_0(X1,np__2))|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_27266, plain, (esk8_0=np__3|r1_xxreal_0(esk8_0,np__2)|~r1_xxreal_0(esk8_0,np__3)|~m1_subset_1(esk8_0,k5_numbers))).#trainneg
cnf(i_0_1508, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,k3_xcmplx_0(X3,X4))),k3_xcmplx_0(X5,X4))=k3_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3),X5),X4)|~v1_xcmplx_0(k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X5)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X3,X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1511, plain, (v5_relat_1(k7_partfun1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),X2)|~m2_finseq_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_31815, plain, (v1_membered(np__0))).#trainneg
cnf(i_0_30670, plain, (v1_xreal_0(k2_xcmplx_0(np__2,np__4)))).#trainneg
cnf(i_0_1512, plain, (v1_relat_1(k7_partfun1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4))|~m2_finseq_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_32184, plain, (esk8_0=np__3|r1_xxreal_0(esk8_0,np__2)|esk8_0=np__4|~m1_subset_1(esk8_0,k5_numbers)|~r1_xxreal_0(esk8_0,np__4))).#trainneg
cnf(i_0_31891, plain, (v1_membered(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(k3_finseq_1(X1))|~v1_funct_1(X1)|~v1_relat_1(k3_finseq_1(X1))|~v1_relat_1(X1)|~v1_xboole_0(k3_finseq_1(X1)))).#trainneg
cnf(i_0_1513, plain, (v1_xboole_0(k7_partfun1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4))|~v1_xboole_0(X2)|~m2_finseq_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1514, plain, (v1_xboole_0(k7_partfun1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4))|~v1_xboole_0(X1)|~m2_finseq_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_30703, plain, (v1_int_1(k2_xcmplx_0(np__4,np__4)))).#trainneg
cnf(i_0_1515, plain, (v4_relat_1(k7_partfun1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),X1)|~m2_finseq_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1516, plain, (v1_finset_1(k7_partfun1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4))|~v1_funct_2(k7_partfun1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),X1,X2)|~v1_finset_1(X1)|~v1_funct_1(k7_partfun1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4))|~m2_finseq_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1517, plain, (v1_relat_1(k9_xtuple_0(k7_partfun1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)),X4,X5)))|~m2_finseq_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_31767, negated_conjecture, (esk4_1(esk1_0)=np__3|r1_xxreal_0(esk4_1(esk1_0),np__2)|r1_xxreal_0(np__4,esk4_1(esk1_0)))).#trainneg
cnf(i_0_31677, negated_conjecture, (r2_hidden(np__2,k9_xtuple_0(esk1_0)))).#trainneg
cnf(i_0_31753, negated_conjecture, (r1_xxreal_0(esk4_1(esk1_0),np__2)|r1_xxreal_0(np__2,esk4_1(esk1_0)))).#trainneg
cnf(i_0_4248, negated_conjecture, (r1_xxreal_0(X1,np__2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_22570, negated_conjecture, (k3_finseq_1(X1)=np__2|r1_xxreal_0(esk4_1(X1),np__4)|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),np__2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_30706, plain, (v1_xreal_0(k2_xcmplx_0(np__4,np__4)))).#trainneg
cnf(i_0_1518, plain, (v6_membered(k7_partfun1(k1_zfmisc_1(X1),X2,X3))|~v6_membered(X1)|~m2_finseq_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1519, plain, (v1_finset_1(k7_partfun1(k1_zfmisc_1(X1),X2,X3))|~v1_finset_1(X1)|~m2_finseq_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1520, plain, (v3_membered(k7_partfun1(k1_zfmisc_1(X1),X2,X3))|~v3_membered(X1)|~m2_finseq_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_31672, negated_conjecture, (r2_hidden(np__2,k4_finseq_1(esk1_0)))).#trainneg
cnf(i_0_31778, negated_conjecture, (esk4_1(esk1_0)=np__2|~r1_xxreal_0(np__2,esk4_1(esk1_0))|~r1_xxreal_0(esk4_1(esk1_0),np__2))).#trainneg
cnf(i_0_31725, negated_conjecture, (r1_xxreal_0(X1,np__3)|esk4_1(esk1_0)=np__4|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk4_1(esk1_0)))).#trainneg
cnf(i_0_1521, plain, (v2_membered(k7_partfun1(k1_zfmisc_1(X1),X2,X3))|~v2_membered(X1)|~m2_finseq_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_31335, plain, (v1_membered(k1_card_1(X1))|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1522, plain, (v1_membered(k7_partfun1(k1_zfmisc_1(X1),X2,X3))|~v1_membered(X1)|~m2_finseq_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1524, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_partfun1(k1_zfmisc_1(X2),X3,X4))|~m2_finseq_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_32048, negated_conjecture, (r1_xxreal_0(X1,np__3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk4_1(esk1_0))|~r1_xxreal_0(esk4_1(esk1_0),np__1))).#trainneg
cnf(i_0_31686, negated_conjecture, (m1_subset_1(np__2,k9_xtuple_0(esk1_0)))).#trainneg
cnf(i_0_21811, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk6_1(k1_zfmisc_1(k2_zfmisc_1(X4,X5)))))|~v1_xboole_0(X3)|~v1_xboole_0(X2)|~v1_xboole_0(X5))).#trainneg
cnf(i_0_25823, plain, (v6_membered(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1525, plain, (m1_subset_1(X1,X2)|v1_xboole_0(k7_partfun1(k1_zfmisc_1(X2),X3,X4))|v1_xboole_0(X2)|~m1_subset_1(X1,k7_partfun1(k1_zfmisc_1(X2),X3,X4))|~m2_finseq_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_32187, plain, (esk8_0=np__3|r1_xxreal_0(esk8_0,np__2)|r1_xxreal_0(np__3,esk8_0)|~m1_subset_1(esk8_0,k5_numbers))).#trainneg
cnf(i_0_1526, plain, (r1_xxreal_0(k7_partfun1(k1_relset_1(k5_numbers,X1),X2,X3),k3_finseq_1(X1))|v1_xboole_0(k1_relset_1(k5_numbers,X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(k7_partfun1(k1_relset_1(k5_numbers,X1),X2,X3))|~m2_finseq_1(X2,k1_relset_1(k5_numbers,X1)))).#trainneg
cnf(i_0_1527, plain, (v1_finseq_1(k7_partfun1(u1_struct_0(k15_euclid(X1)),X2,X3))|~v7_ordinal1(X1)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(X1))))).#trainneg
cnf(i_0_31791, negated_conjecture, (r1_xxreal_0(X1,np__3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk4_1(esk1_0))|~r1_xxreal_0(esk4_1(esk1_0),np__2))).#trainneg
cnf(i_0_31158, negated_conjecture, (X1=np__4|r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X1,np__4)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_31782, negated_conjecture, (esk4_1(esk1_0)=np__2|~r1_xxreal_0(np__4,esk4_1(esk1_0))|~r1_xxreal_0(esk4_1(esk1_0),np__2))).#trainneg
cnf(i_0_32047, negated_conjecture, (esk4_1(esk1_0)=np__3|~r1_xxreal_0(np__3,esk4_1(esk1_0))|~r1_xxreal_0(esk4_1(esk1_0),np__1))).#trainneg
cnf(i_0_31771, negated_conjecture, (r1_xxreal_0(esk4_1(esk1_0),esk4_1(esk1_0))|esk4_1(esk1_0)=np__2|~r1_xxreal_0(esk4_1(esk1_0),np__2))).#trainneg
cnf(i_0_31215, negated_conjecture, (X1=np__3|r1_xxreal_0(X1,np__2)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X1,np__4)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_32835, plain, (v1_membered(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1528, plain, (v3_valued_0(k7_partfun1(u1_struct_0(k15_euclid(X1)),X2,X3))|~v7_ordinal1(X1)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(X1))))).#trainneg
cnf(i_0_1531, plain, (v3_ordinal1(k7_partfun1(X1,X2,X3))|~v3_ordinal1(X1)|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_1533, plain, (v1_xreal_0(k7_partfun1(X1,X2,X3))|~v3_membered(X1)|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_32571, negated_conjecture, (k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))).#trainneg
cnf(i_0_32572, negated_conjecture, (k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__3)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__2))).#trainneg
cnf(i_0_1534, plain, (v1_xcmplx_0(k7_partfun1(X1,X2,X3))|~v1_membered(X1)|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_1535, plain, (v1_int_1(k7_partfun1(X1,X2,X3))|~v5_membered(X1)|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_1540, plain, (v1_relat_1(k19_euclid(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_32771, negated_conjecture, (m1_subset_1(np__2,k4_finseq_1(esk1_0)))).#trainneg
cnf(i_0_32874, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk6_1(k1_zfmisc_1(np__0))))|~v1_xboole_0(X4)|~v1_xboole_0(X5)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_26303, plain, (v3_membered(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1541, plain, (v1_relat_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_32564, negated_conjecture, (~r2_hidden(k9_xtuple_0(esk1_0),np__2))).#trainneg
cnf(i_0_32772, negated_conjecture, (~r2_hidden(k4_finseq_1(esk1_0),np__2))).#trainneg
cnf(i_0_1543, plain, (v1_xboole_0(k9_xtuple_0(X1))|~v3_card_1(X1,np__0)|~v1_card_1(np__0)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1550, plain, (v7_ordinal1(esk16_1(X1))|v1_xboole_0(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_32721, negated_conjecture, (r1_xxreal_0(k3_finseq_1(X1),np__2)|r1_xxreal_0(np__1,k3_finseq_1(X1))|~v1_xxreal_0(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_32955, negated_conjecture, (v1_finseq_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),esk1_0,np__1)))).#trainneg
cnf(i_0_4247, negated_conjecture, (np__2=X1|~v1_xxreal_0(X1)|~r1_xxreal_0(np__2,X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_1555, plain, (v2_xxreal_0(k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xreal_0(k3_finseq_1(X1))|~v1_xreal_0(esk3_1(X1))|~v2_xxreal_0(esk3_1(X1))|~v7_ordinal1(esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_32917, plain, (esk8_0=np__2|esk8_0=np__3|r1_xxreal_0(np__3,esk8_0)|~r1_xxreal_0(np__2,esk8_0)|~m1_subset_1(esk8_0,k5_numbers))).#trainneg
cnf(i_0_32105, plain, (v2_membered(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_32107, plain, (v2_membered(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1558, plain, (r1_xxreal_0(esk3_1(k3_finseq_1(X1)),k3_finseq_1(X1))|v1_xboole_0(k3_finseq_1(X1))|v3_funct_1(k3_finseq_1(X1))|~v7_ordinal1(esk3_1(k3_finseq_1(X1)))|~m2_finseq_1(k3_finseq_1(X1),u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(k3_finseq_1(X1))|~v1_goboard5(k3_finseq_1(X1))|~v2_topreal1(k3_finseq_1(X1))|~v1_topreal1(k3_finseq_1(X1))|~v1_finseq_6(k3_finseq_1(X1),u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1559, plain, (r1_xxreal_0(esk3_1(X1),k1_card_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(esk3_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_32746, negated_conjecture, (k3_finseq_1(X1)=np__2|r1_xxreal_0(esk4_1(X1),np__4)|v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(k3_finseq_1(X1),np__2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_31797, negated_conjecture, (~r1_xxreal_0(np__4,esk4_1(esk1_0))|~r1_xxreal_0(esk4_1(esk1_0),np__2))).#trainneg
cnf(i_0_31774, negated_conjecture, (r1_xxreal_0(X1,np__1)|esk4_1(esk1_0)=np__2|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk4_1(esk1_0))|~r1_xxreal_0(esk4_1(esk1_0),np__2))).#trainneg
cnf(i_0_31297, negated_conjecture, (k3_finseq_1(X1)=np__1|r1_xxreal_0(k3_finseq_1(X1),np__0)|~r1_xxreal_0(k3_finseq_1(X1),np__1)|~r1_xxreal_0(k3_finseq_1(X1),np__4)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_21584, negated_conjecture, (r1_xxreal_0(np__2,k3_finseq_1(X1))|r1_xxreal_0(esk4_1(X1),np__4)|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1560, plain, (v2_xxreal_0(k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xreal_0(k3_finseq_1(X1))|~v1_xreal_0(esk4_1(X1))|~v2_xxreal_0(esk4_1(X1))|~v7_ordinal1(esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1563, plain, (r1_xxreal_0(esk4_1(k3_finseq_1(X1)),k3_finseq_1(X1))|v1_xboole_0(k3_finseq_1(X1))|v3_funct_1(k3_finseq_1(X1))|~v7_ordinal1(esk4_1(k3_finseq_1(X1)))|~m2_finseq_1(k3_finseq_1(X1),u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(k3_finseq_1(X1))|~v1_goboard5(k3_finseq_1(X1))|~v2_topreal1(k3_finseq_1(X1))|~v1_topreal1(k3_finseq_1(X1))|~v1_finseq_6(k3_finseq_1(X1),u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1564, plain, (r1_xxreal_0(esk4_1(X1),k1_card_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v7_ordinal1(esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_31272, negated_conjecture, (X1=np__2|r1_xxreal_0(X1,np__1)|~r1_xxreal_0(X1,np__2)|~r1_xxreal_0(X1,np__4)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_31184, negated_conjecture, (r1_xxreal_0(k3_finseq_1(X1),np__2)|~r1_xxreal_0(k3_finseq_1(X1),np__1)|~r1_xxreal_0(k3_finseq_1(X1),np__4)|~r1_xxreal_0(np__1,k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_31270, negated_conjecture, (k3_finseq_1(X1)=np__2|r1_xxreal_0(k3_finseq_1(X1),np__1)|~r1_xxreal_0(k3_finseq_1(X1),np__2)|~r1_xxreal_0(k3_finseq_1(X1),np__4)|~r1_xxreal_0(np__1,k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16620, plain, (v2_xxreal_0(esk8_0)|esk8_0=np__2|~r1_xxreal_0(np__2,esk8_0))).#trainneg
cnf(i_0_22571, negated_conjecture, (r1_xxreal_0(esk4_1(X1),np__4)|r1_xxreal_0(X2,k3_finseq_1(X1))|v1_xboole_0(X1)|v3_funct_1(X1)|~v1_xxreal_0(k3_finseq_1(X1))|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,np__2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1570, plain, (v1_funct_1(k19_euclid(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_32437, plain, (esk8_0=np__4|esk8_0=np__3|r1_xxreal_0(esk8_0,np__2)|r1_xxreal_0(np__4,esk8_0)|~m1_subset_1(esk8_0,k5_numbers))).#trainneg
cnf(i_0_1571, plain, (v1_funct_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1574, plain, (esk6_1(X1)=np__2|esk6_1(X1)=np__4|esk6_1(X1)=np__1|esk6_1(X1)=np__0|esk6_1(X1)=np__3|~r1_xxreal_0(esk6_1(X1),np__4)|~v6_membered(X1))).#trainneg
cnf(i_0_31244, negated_conjecture, (r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X1,np__2)|~r1_xxreal_0(X1,np__4)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_31326, negated_conjecture, (r1_xxreal_0(k3_finseq_1(X1),np__3)|~r1_xxreal_0(k3_finseq_1(X1),np__1)|~r1_xxreal_0(k3_finseq_1(X1),np__4)|~r1_xxreal_0(np__1,k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_24279, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk6_1(k1_zfmisc_1(k2_zfmisc_1(X4,X5)))))|~v1_xboole_0(X3)|~v1_xboole_0(X2)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_33317, plain, (v2_membered(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1575, plain, (esk2_1(X1)=np__2|esk2_1(X1)=np__4|esk2_1(X1)=np__1|esk2_1(X1)=np__0|esk2_1(X1)=np__3|v1_topreal1(X1)|~r1_xxreal_0(esk2_1(X1),np__4)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1576, plain, (k18_euclid(X1)=np__2|k18_euclid(X1)=np__4|k18_euclid(X1)=np__1|k18_euclid(X1)=np__0|k18_euclid(X1)=np__3|~r1_xxreal_0(k18_euclid(X1),np__4)|~v6_membered(k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1577, plain, (k17_euclid(X1)=np__2|k17_euclid(X1)=np__4|k17_euclid(X1)=np__1|k17_euclid(X1)=np__0|k17_euclid(X1)=np__3|~r1_xxreal_0(k17_euclid(X1),np__4)|~v6_membered(k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_27034, negated_conjecture, (np__3!=np__2)).#trainneg
cnf(i_0_28374, negated_conjecture, (np__3!=np__1)).#trainneg
cnf(i_0_33329, plain, (v2_membered(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1586, plain, (v3_membered(k1_relset_1(X1,X2))|~v4_relat_1(X2,X1)|~v6_membered(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_33304, plain, (esk8_0=np__3|esk8_0=np__4|r1_xxreal_0(np__4,esk8_0)|esk8_0=np__2|~m1_subset_1(esk8_0,k5_numbers)|~r1_xxreal_0(np__2,esk8_0))).#trainneg
cnf(i_0_33212, plain, (~v3_card_1(esk7_0,np__0))).#trainneg
cnf(i_0_33213, negated_conjecture, (~v3_card_1(esk1_0,np__0))).#trainneg
cnf(i_0_1590, plain, (r1_xxreal_0(esk16_1(X1),esk16_1(X1))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_33400, negated_conjecture, (esk4_1(X1)=np__3|esk4_1(X1)=np__0|esk4_1(X1)=np__1|esk4_1(X1)=np__4|esk4_1(X1)=np__2|k3_finseq_1(X1)=np__2|v1_xboole_0(X1)|v3_funct_1(X1)|~m1_subset_1(esk4_1(X1),k5_numbers)|~r1_xxreal_0(k3_finseq_1(X1),np__2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_33405, negated_conjecture, (esk4_1(X1)=np__4|esk4_1(X1)=np__3|esk4_1(X1)=np__2|esk4_1(X1)=np__1|v1_xboole_0(X1)|v3_funct_1(X1)|k3_finseq_1(X1)=np__2|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(k3_finseq_1(X1),np__2))).#trainneg
cnf(i_0_1591, plain, (r1_xxreal_0(esk8_0,esk16_1(X1))|r1_xxreal_0(esk16_1(X1),esk8_0)|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_33197, negated_conjecture, (~m1_subset_1(k9_xtuple_0(esk1_0),np__2))).#trainneg
cnf(i_0_1592, plain, (r1_xxreal_0(X1,esk16_1(X2))|r1_xxreal_0(esk16_1(X2),X1)|v1_xboole_0(X2)|~v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1595, plain, (m1_subset_1(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)),k1_numbers)|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1),u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(np__2,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_33201, plain, (k9_xtuple_0(X1)=np__0|~v3_card_1(X1,np__0)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_33239, negated_conjecture, (r1_xxreal_0(np__1,k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r1_xxreal_0(np__3,k3_finseq_1(X1)))).#trainneg
cnf(i_0_33214, plain, (~v3_card_1(X1,np__1)|~v1_card_1(np__1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_card_1(X1,np__0))).#trainneg
cnf(i_0_33238, negated_conjecture, (r1_xxreal_0(np__1,k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r1_xxreal_0(np__4,k3_finseq_1(X1)))).#trainneg
cnf(i_0_1597, plain, (v1_xreal_0(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1),u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(np__2,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_33300, plain, (esk8_0=np__3|esk8_0=np__2|~r1_xxreal_0(esk8_0,np__3)|~m1_subset_1(esk8_0,k5_numbers)|~r1_xxreal_0(np__2,esk8_0))).#trainneg
cnf(i_0_33199, negated_conjecture, (~m1_subset_1(k4_finseq_1(esk1_0),np__2))).#trainneg
cnf(i_0_29764, plain, (~v1_xboole_0(esk6_1(esk16_1(u1_struct_0(k15_euclid(np__2))))))).#trainneg
cnf(i_0_33075, plain, (v1_membered(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_33077, plain, (v1_membered(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1596, plain, (k19_euclid(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1)),k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)))=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1)|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1),u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(np__2,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_1598, plain, (k10_finseq_1(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1)),k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)))=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1)|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1),u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(np__2,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_33408, negated_conjecture, (esk4_1(X1)=np__4|esk4_1(X1)=np__3|r1_xxreal_0(esk4_1(X1),np__2)|v1_xboole_0(X1)|v3_funct_1(X1)|k3_finseq_1(X1)=np__2|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(k3_finseq_1(X1),np__2))).#trainneg
cnf(i_0_21267, negated_conjecture, (r1_xxreal_0(np__2,esk4_1(X1))|r1_xxreal_0(esk4_1(X1),np__4)|v1_xboole_0(X1)|v3_funct_1(X1)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1599, plain, (v7_ordinal1(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|~v6_membered(k1_numbers)|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1),u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(np__2,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_1600, plain, (k19_euclid(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,np__1))),k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)))=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,np__1))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,np__1)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,np__1)),u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(k1_nat_1(np__1,np__1),k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)|~r1_xxreal_0(np__2,k3_finseq_1(X1)))).#trainneg
cnf(i_0_1615, plain, (r1_xxreal_0(k18_euclid(X1),k18_euclid(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_32927, plain, (k7_partfun1(k1_relset_1(k5_numbers,X1),X2,X3)=k3_finseq_1(X1)|v1_xboole_0(k1_relset_1(k5_numbers,X1))|~v1_xxreal_0(k7_partfun1(k1_relset_1(k5_numbers,X1),X2,X3))|~v1_xxreal_0(k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),k7_partfun1(k1_relset_1(k5_numbers,X1),X2,X3))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(k7_partfun1(k1_relset_1(k5_numbers,X1),X2,X3))|~m2_finseq_1(X2,k1_relset_1(k5_numbers,X1)))).#trainneg
cnf(i_0_31221, negated_conjecture, (r1_xxreal_0(esk4_1(X1),np__3)|v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(esk4_1(X1),np__2)|~r1_xxreal_0(esk4_1(X1),np__4)|~r1_xxreal_0(np__1,esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_32096, plain, (v2_membered(k1_card_1(X1))|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_1616, plain, (r1_xxreal_0(esk8_0,k18_euclid(X1))|r1_xxreal_0(k18_euclid(X1),esk8_0)|~v2_membered(k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1617, plain, (r1_xxreal_0(X1,k18_euclid(X2))|r1_xxreal_0(k18_euclid(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1618, plain, (v1_xxreal_0(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~v2_membered(k1_numbers)|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k1_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_34036, plain, (v1_membered(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1619, plain, (v1_xxreal_0(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2))|~v2_membered(k1_numbers)|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1),u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(np__2,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_33872, plain, (esk8_0=np__4|esk8_0=np__2|esk8_0=np__3|~r1_xxreal_0(esk8_0,np__4)|~m1_subset_1(esk8_0,k5_numbers)|~r1_xxreal_0(np__2,esk8_0))).#trainneg
cnf(i_0_1620, plain, (v1_xxreal_0(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,X2)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~v2_membered(k1_numbers)|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(np__1,X2)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k1_nat_1(np__1,X2),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_1621, plain, (v1_xxreal_0(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1)))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|~v2_membered(k1_numbers)|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k2_xcmplx_0(X2,np__1)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k2_xcmplx_0(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_34274, negated_conjecture, (v1_xboole_0(X1)|v3_funct_1(X1)|~r1_xxreal_0(np__4,esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(esk4_1(X1),np__2)|~r1_xxreal_0(esk4_1(X1),np__4))).#trainneg
cnf(i_0_34128, negated_conjecture, (esk4_1(X1)=np__4|v1_xboole_0(X1)|v3_funct_1(X1)|r1_xxreal_0(np__2,esk4_1(X1))|~r1_xxreal_0(np__4,esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_34104, negated_conjecture, (esk4_1(X1)=np__2|v1_xboole_0(X1)|v3_funct_1(X1)|k3_finseq_1(X1)=np__2|esk4_1(X1)=np__3|esk4_1(X1)=np__4|~r1_xxreal_0(np__2,esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(k3_finseq_1(X1),np__2))).#trainneg
cnf(i_0_34048, plain, (v1_membered(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1622, plain, (r1_xxreal_0(k17_euclid(X1),k17_euclid(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_34001, plain, (~v3_card_1(np__0,np__1)|~v1_card_1(np__1))).#trainneg
cnf(i_0_1623, plain, (r1_xxreal_0(esk8_0,k17_euclid(X1))|r1_xxreal_0(k17_euclid(X1),esk8_0)|~v2_membered(k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1624, plain, (r1_xxreal_0(X1,k17_euclid(X2))|r1_xxreal_0(k17_euclid(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_34103, negated_conjecture, (v1_xboole_0(X1)|v3_funct_1(X1)|k3_finseq_1(X1)=np__2|esk4_1(X1)=np__3|esk4_1(X1)=np__4|~r1_xxreal_0(np__3,esk4_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~v1_goboard5(X1)|~v2_topreal1(X1)|~v1_topreal1(X1)|~v1_finseq_6(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(k3_finseq_1(X1),np__2))).#trainneg
cnf(i_0_34248, negated_conjecture, (k7_partfun1(k1_relset_1(k5_numbers,esk1_0),X1,X2)=np__4|v1_xboole_0(k1_relset_1(k5_numbers,esk1_0))|~v1_xxreal_0(k7_partfun1(k1_relset_1(k5_numbers,esk1_0),X1,X2))|~v7_ordinal1(k7_partfun1(k1_relset_1(k5_numbers,esk1_0),X1,X2))|~r1_xxreal_0(np__4,k7_partfun1(k1_relset_1(k5_numbers,esk1_0),X1,X2))|~m2_finseq_1(X1,k1_relset_1(k5_numbers,esk1_0)))).#trainneg
cnf(i_0_34092, negated_conjecture, (r1_xxreal_0(X1,np__2)|k3_finseq_1(X2)=np__2|esk4_1(X2)=np__3|esk4_1(X2)=np__4|v1_xboole_0(X2)|v3_funct_1(X2)|~v1_xxreal_0(esk4_1(X2))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk4_1(X2))|~r1_xxreal_0(k3_finseq_1(X2),np__2)|~m2_finseq_1(X2,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X2)|~v1_goboard5(X2)|~v2_topreal1(X2)|~v1_topreal1(X2)|~v1_finseq_6(X2,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_33066, plain, (v1_membered(k1_card_1(X1))|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_1627, plain, (v1_xboole_0(X1)|v1_xboole_0(k4_finseq_1(X1))|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__0)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1)|~r2_hidden(np__1,k4_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~m1_subset_1(np__0,k4_finseq_1(X1)))).#trainneg
cnf(i_0_1625, plain, (v1_xxreal_0(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)))|m1_subset_1(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2)),k1_numbers)|~v2_membered(k1_numbers)|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,k1_nat_1(X2,np__1)),u1_struct_0(k15_euclid(np__2)))|~v7_ordinal1(X2)|~r1_xxreal_0(k1_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_1629, plain, (v1_xboole_0(X1)|v1_xboole_0(k4_finseq_1(X1))|k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__1)!=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__2)|~r2_hidden(np__2,k4_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_goboard5(X1)|~m1_subset_1(np__1,k4_finseq_1(X1)))).#trainneg
cnf(i_0_33287, negated_conjecture, (k3_finseq_1(X1)=np__2|~r1_xxreal_0(np__2,k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),np__1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_24929, plain, (k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__0))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__0))=k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,X2))|k2_xcmplx_0(np__1,X2)=np__3|k1_nat_1(np__1,X2)=np__2|k1_nat_1(np__1,X2)=np__1|k2_xcmplx_0(np__1,X2)=np__4|r1_xxreal_0(np__4,X2)|~v7_ordinal1(X2)|~r1_xxreal_0(np__0,k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_29253, negated_conjecture, (r1_xxreal_0(k7_partfun1(k9_xtuple_0(esk1_0),X1,X2),np__4)|~m2_finseq_1(X1,k9_xtuple_0(esk1_0)))).#trainneg
cnf(i_0_31212, negated_conjecture, (k3_finseq_1(X1)=np__3|r1_xxreal_0(k3_finseq_1(X1),np__2)|~r1_xxreal_0(k3_finseq_1(X1),np__3)|~r1_xxreal_0(k3_finseq_1(X1),np__4)|~r1_xxreal_0(np__1,k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1630, plain, (m1_subset_1(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__4)),k1_numbers)|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__4))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3),u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(np__4,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_1631, plain, (k19_euclid(k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3)),k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__4)))=k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3)|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__4))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3),u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(np__4,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
cnf(i_0_1632, plain, (v1_xreal_0(k18_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__4)))|k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3))=k17_euclid(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__4))|~m1_subset_1(k7_partfun1(u1_struct_0(k15_euclid(np__2)),X1,np__3),u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(np__4,k3_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_topreal1(X1))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 237
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 360
# Removed in clause preprocessing      : 26
# Initial clauses in saturation        : 334
# Processed clauses                    : 4393
# ...of these trivial                  : 64
# ...subsumed                          : 2022
# ...remaining for further processing  : 2307
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 362
# Backward-rewritten                   : 335
# Generated clauses                    : 34712
# ...of the previous two non-trivial   : 25672
# Contextual simplify-reflections      : 305
# Paramodulations                      : 34516
# Factorizations                       : 193
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1607
#    Positive orientable unit clauses  : 239
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 57
#    Non-unit-clauses                  : 1311
# Current number of unprocessed clauses: 20259
# ...number of literals in the above   : 143427
# Current number of archived formulas  : 0
# Current number of archived clauses   : 3007
# Proof object given clauses           : 344
# Proof search given clauses           : 2307
# Clause-clause subsumption calls (NU) : 1602250
# Rec. Clause-clause subsumption calls : 125031
# Non-unit clause-clause subsumptions  : 1845
# Unit Clause-clause subsumption calls : 20914
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 83
# BW rewrite match successes           : 60
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 960012

# -------------------------------------------------
# User time                : 52.273 s
# System time              : 1.893 s
# Total time               : 54.166 s
# Maximum resident set size: 456080 pages
52.27user 1.96system 0:55.55elapsed 97%CPU (0avgtext+0avgdata 456080maxresident)k
0inputs+1096outputs (0major+507129minor)pagefaults 0swaps
