2020-08-16 22:43:40.023817: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_karel_mizar1/epoch1_mizar_8_96ffb007719f4afdb4a1e67b65c3d8cb_24
2020-08-16 22:43:40.108864: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-08-16 22:43:40.252059: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-08-16 22:43:40.787100: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-08-16 22:43:42.249613: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 2225819 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_karel_mizar1/epoch1_mizar_8_96ffb007719f4afdb4a1e67b65c3d8cb_24' loaded (query=1024; context=1024; weigths=real; len_mult=0.000000).
# Preprocessing time       : 2.970 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t128_xxreal_1, axiom, ![X1]:(v1_xxreal_0(X1)=>![X2]:(v1_xxreal_0(X2)=>(r1_xxreal_0(X1,X2)=>k1_xxreal_1(X1,X2)=k2_xboole_0(k4_xxreal_1(X1,X2),k2_tarski(X1,X2))))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t128_xxreal_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', commutativity_k2_xboole_0)).
fof(redefinition_k2_rcomp_1, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>k2_rcomp_1(X1,X2)=k4_xxreal_1(X1,X2)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', redefinition_k2_rcomp_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', commutativity_k3_xboole_0)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t28_xboole_1)).
fof(t24_xreal_1, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>(r1_xxreal_0(X1,X2)<=>r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t24_xreal_1)).
fof(redefinition_k1_real_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>k1_real_1(X1)=k4_xcmplx_0(X1)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', redefinition_k1_real_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t1_subset)).
fof(l19_sincos10, axiom, (r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))&r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', l19_sincos10)).
fof(t27_sincos10, axiom, k2_relset_1(k1_numbers,k3_sincos10)=k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t27_sincos10)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t1_xboole_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t3_subset)).
fof(d3_xboole_0, axiom, ![X1, X2, X3]:(X3=k2_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)|r2_hidden(X4,X2)))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', d3_xboole_0)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t6_boole)).
fof(t21_rfunct_2, axiom, ![X1, X2]:((v1_funct_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X2,X1))<=>![X3]:(m1_subset_1(X3,k1_numbers)=>![X4]:(m1_subset_1(X4,k1_numbers)=>~((r2_hidden(X3,k9_subset_1(k1_numbers,X1,k1_relset_1(k1_numbers,X2)))&(r2_hidden(X4,k9_subset_1(k1_numbers,X1,k1_relset_1(k1_numbers,X2)))&(~(r1_xxreal_0(X4,X3))&r1_xxreal_0(k1_seq_1(X2,X3),k1_seq_1(X2,X4)))))))))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t21_rfunct_2)).
fof(commutativity_k9_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>k9_subset_1(X1,X2,X3)=k9_subset_1(X1,X3,X2)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', commutativity_k9_subset_1)).
fof(redefinition_k9_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>k9_subset_1(X1,X2,X3)=k3_xboole_0(X2,X3)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', redefinition_k9_subset_1)).
fof(l63_sincos10, axiom, k2_funct_1(k3_sincos10)=k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', l63_sincos10)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', cc2_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', cc1_relset_1)).
fof(cc1_xreal_0, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>v1_xreal_0(X1)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', cc1_xreal_0)).
fof(cc13_membered, axiom, ![X1]:(v2_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xxreal_0(X2))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', cc13_membered)).
fof(dt_k1_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>m1_subset_1(k1_relset_1(X1,X2),k1_zfmisc_1(X1))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', dt_k1_relset_1)).
fof(fc6_ordinal1, axiom, (~(v1_xboole_0(k4_ordinal1))&v3_ordinal1(k4_ordinal1)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', fc6_ordinal1)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', redefinition_k6_numbers)).
fof(rc4_xreal_0, axiom, ?[X1]:(v1_xboole_0(X1)&(v1_xcmplx_0(X1)&(v1_xxreal_0(X1)&v1_xreal_0(X1)))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', rc4_xreal_0)).
fof(t29_rfunct_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>((r1_tarski(X1,X2)&v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,X2)))=>v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,X1)))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t29_rfunct_2)).
fof(t19_sincos10, axiom, v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t19_sincos10)).
fof(connectedness_r1_xxreal_0, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>(r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', connectedness_r1_xxreal_0)).
fof(cc4_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xxreal_0(X1)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', cc4_xreal_0)).
fof(fc63_valued_0, axiom, ![X1, X2]:((v1_relat_1(X1)&(v1_funct_1(X1)&v3_valued_0(X1)))=>v1_xreal_0(k1_funct_1(X1,X2))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', fc63_valued_0)).
fof(redefinition_k1_seq_1, axiom, ![X1, X2]:((v1_relat_1(X1)&(v1_funct_1(X1)&v3_valued_0(X1)))=>k1_seq_1(X1,X2)=k1_funct_1(X1,X2)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', redefinition_k1_seq_1)).
fof(cc30_valued_0, axiom, ![X1]:((v1_relat_1(X1)&v5_relat_1(X1,k1_numbers))=>(v1_relat_1(X1)&v3_valued_0(X1))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', cc30_valued_0)).
fof(dt_k2_fdiff_9, axiom, (v1_funct_1(k2_fdiff_9)&m1_subset_1(k2_fdiff_9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', dt_k2_fdiff_9)).
fof(fraenkel_a_2_0_rcomp_1, axiom, ![X1, X2, X3]:((v1_xreal_0(X2)&v1_xreal_0(X3))=>(r2_hidden(X1,a_2_0_rcomp_1(X2,X3))<=>?[X4]:(m1_subset_1(X4,k1_numbers)&(X1=X4&(r1_xxreal_0(X2,X4)&r1_xxreal_0(X4,X3)))))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', fraenkel_a_2_0_rcomp_1)).
fof(d1_rcomp_1, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>k1_rcomp_1(X1,X2)=a_2_0_rcomp_1(X1,X2))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', d1_rcomp_1)).
fof(cc4_membered, axiom, ![X1]:(v3_membered(X1)=>v2_membered(X1)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', cc4_membered)).
fof(t3_sincos10, axiom, r1_tarski(k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers),k1_relset_1(k1_numbers,k2_fdiff_9)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t3_sincos10)).
fof(redefinition_k1_rcomp_1, axiom, ![X1, X2]:((v1_xreal_0(X1)&v1_xreal_0(X2))=>k1_rcomp_1(X1,X2)=k1_xxreal_1(X1,X2)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', redefinition_k1_rcomp_1)).
fof(dt_m2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&(~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(X1))))=>![X3]:(m2_subset_1(X3,X1,X2)=>m1_subset_1(X3,X1))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', dt_m2_subset_1)).
fof(fc1_numbers, axiom, ~(v1_xboole_0(k1_numbers)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', fc1_numbers)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', redefinition_k5_numbers)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t4_subset)).
fof(t1_xxreal_0, axiom, ![X1]:(v1_xxreal_0(X1)=>![X2]:(v1_xxreal_0(X2)=>((r1_xxreal_0(X1,X2)&r1_xxreal_0(X2,X1))=>X1=X2))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t1_xxreal_0)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', d2_tarski)).
fof(cc8_membered, axiom, ![X1]:(m1_subset_1(X1,k1_zfmisc_1(k1_numbers))=>v3_membered(X1)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', cc8_membered)).
fof(cc14_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xreal_0(X2))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', cc14_membered)).
fof(dt_k6_numbers, axiom, m2_subset_1(k6_numbers,k1_numbers,k5_numbers), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', dt_k6_numbers)).
fof(dt_k5_numbers, axiom, m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', dt_k5_numbers)).
fof(l31_sincos10, axiom, r1_tarski(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))),k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', l31_sincos10)).
fof(fraenkel_a_2_2_rcomp_1, axiom, ![X1, X2, X3]:((v1_xreal_0(X2)&v1_xxreal_0(X3))=>(r2_hidden(X1,a_2_2_rcomp_1(X2,X3))<=>?[X4]:(m1_subset_1(X4,k1_numbers)&(X1=X4&(r1_xxreal_0(X2,X4)&~(r1_xxreal_0(X3,X4))))))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', fraenkel_a_2_2_rcomp_1)).
fof(d7_rcomp_1, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xxreal_0(X2)=>k3_rcomp_1(X1,X2)=a_2_2_rcomp_1(X1,X2))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', d7_rcomp_1)).
fof(involutiveness_k1_real_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>k1_real_1(k1_real_1(X1))=X1), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', involutiveness_k1_real_1)).
fof(reflexivity_r1_xxreal_0, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>r1_xxreal_0(X1,X1)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', reflexivity_r1_xxreal_0)).
fof(rc3_xxreal_0, axiom, ?[X1]:(v1_xxreal_0(X1)&v3_xxreal_0(X1)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', rc3_xxreal_0)).
fof(dt_k31_sin_cos, axiom, v1_xreal_0(k31_sin_cos), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', dt_k31_sin_cos)).
fof(redefinition_k32_sin_cos, axiom, k32_sin_cos=k31_sin_cos, file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', redefinition_k32_sin_cos)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', fc3_membered)).
fof(t5_comptrig, axiom, (~(r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k6_numbers))&(~(r1_xxreal_0(k32_sin_cos,k10_real_1(k32_sin_cos,np__2)))&(~(r1_xxreal_0(k32_sin_cos,k6_numbers))&(~(r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k1_real_1(k10_real_1(k32_sin_cos,np__2))))&(~(r1_xxreal_0(k8_real_1(np__2,k32_sin_cos),k32_sin_cos))&(~(r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos),k10_real_1(k32_sin_cos,np__2)))&(~(r1_xxreal_0(k6_numbers,k1_real_1(k10_real_1(k32_sin_cos,np__2))))&(~(r1_xxreal_0(k8_real_1(np__2,k32_sin_cos),k6_numbers))&(~(r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos),k32_sin_cos))&(~(r1_xxreal_0(k8_real_1(np__2,k32_sin_cos),k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos)))&~(r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos),k6_numbers)))))))))))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t5_comptrig)).
fof(l20_sincos10, axiom, (r2_hidden(k10_real_1(k32_sin_cos,np__4),k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))&r2_hidden(k10_real_1(k32_sin_cos,np__2),k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', l20_sincos10)).
fof(fraenkel_a_2_1_rcomp_1, axiom, ![X1, X2, X3]:((v1_xxreal_0(X2)&v1_xxreal_0(X3))=>(r2_hidden(X1,a_2_1_rcomp_1(X2,X3))<=>?[X4]:(m1_subset_1(X4,k1_numbers)&(X1=X4&(~(r1_xxreal_0(X4,X2))&~(r1_xxreal_0(X3,X4))))))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', fraenkel_a_2_1_rcomp_1)).
fof(d2_rcomp_1, axiom, ![X1]:(v1_xxreal_0(X1)=>![X2]:(v1_xxreal_0(X2)=>k2_rcomp_1(X1,X2)=a_2_1_rcomp_1(X1,X2))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', d2_rcomp_1)).
fof(dt_k1_real_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>m1_subset_1(k1_real_1(X1),k1_numbers)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', dt_k1_real_1)).
fof(dt_k1_rcomp_1, axiom, ![X1, X2]:((v1_xreal_0(X1)&v1_xreal_0(X2))=>m1_subset_1(k1_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', dt_k1_rcomp_1)).
fof(l25_sin_cos9, axiom, r2_hidden(k10_real_1(k32_sin_cos,np__4),k2_rcomp_1(k6_numbers,k32_sin_cos)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', l25_sin_cos9)).
fof(t35_sincos10, conjecture, ![X1]:(r2_hidden(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))=>r2_hidden(k1_seq_1(k2_fdiff_9,X1),k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1)))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t35_sincos10)).
fof(t1_xxreal_1, axiom, ![X1]:(v1_xxreal_0(X1)=>![X2]:(v1_xxreal_0(X2)=>![X3]:(v1_xxreal_0(X3)=>(r2_hidden(X1,k1_xxreal_1(X2,X3))<=>(r1_xxreal_0(X2,X1)&r1_xxreal_0(X1,X3)))))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t1_xxreal_1)).
fof(t4_sincos10, axiom, r1_tarski(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_relset_1(k1_numbers,k2_fdiff_9)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t4_sincos10)).
fof(dt_k1_seq_1, axiom, ![X1, X2]:((v1_relat_1(X1)&(v1_funct_1(X1)&v3_valued_0(X1)))=>m1_subset_1(k1_seq_1(X1,X2),k1_numbers)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', dt_k1_seq_1)).
fof(cc21_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v3_membered(X2))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', cc21_membered)).
fof(t32_sincos10, axiom, (k1_seq_1(k2_fdiff_9,k1_real_1(k10_real_1(k32_sin_cos,np__2)))=k1_real_1(np__1)&(k1_seq_1(k2_fdiff_9,k1_real_1(k10_real_1(k32_sin_cos,np__4)))=k1_real_1(k7_square_1(np__2))&(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))=k7_square_1(np__2)&k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__2))=np__1))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t32_sincos10)).
fof(t7_boole, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t7_boole)).
fof(t11_sin_cos9, axiom, k2_relset_1(k1_numbers,k1_sin_cos9)=k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t11_sin_cos9)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t2_subset)).
fof(l24_sin_cos9, axiom, r2_hidden(k10_real_1(k32_sin_cos,np__4),k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', l24_sin_cos9)).
fof(t27_comptrig, axiom, ![X1]:(v1_xreal_0(X1)=>(r2_hidden(k1_seq_1(k16_sin_cos,X1),k1_rcomp_1(k1_real_1(np__1),np__1))&r2_hidden(k1_seq_1(k19_sin_cos,X1),k1_rcomp_1(k1_real_1(np__1),np__1)))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t27_comptrig)).
fof(t76_sin_cos, axiom, (k1_seq_1(k19_sin_cos,k10_real_1(k32_sin_cos,np__2))=k6_numbers&(k1_seq_1(k16_sin_cos,k10_real_1(k32_sin_cos,np__2))=np__1&(k1_seq_1(k19_sin_cos,k32_sin_cos)=k1_real_1(np__1)&(k1_seq_1(k16_sin_cos,k32_sin_cos)=k6_numbers&(k1_seq_1(k19_sin_cos,k7_real_1(k32_sin_cos,k10_real_1(k32_sin_cos,np__2)))=k6_numbers&(k1_seq_1(k16_sin_cos,k7_real_1(k32_sin_cos,k10_real_1(k32_sin_cos,np__2)))=k1_real_1(np__1)&(k1_seq_1(k19_sin_cos,k8_real_1(np__2,k32_sin_cos))=np__1&k1_seq_1(k16_sin_cos,k8_real_1(np__2,k32_sin_cos))=k6_numbers))))))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t76_sin_cos)).
fof(spc1_numerals, axiom, ((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&(m1_subset_1(np__1,k5_numbers)&m1_subset_1(np__1,k1_numbers))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', spc1_numerals)).
fof(t63_sin_cos9, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>((r1_xxreal_0(k1_real_1(np__1),X1)&r1_xxreal_0(X1,np__1))=>(r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k5_sin_cos9(X1))&r1_xxreal_0(k5_sin_cos9(X1),k10_real_1(k32_sin_cos,np__4))))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t63_sin_cos9)).
fof(l23_sin_cos9, axiom, r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', l23_sin_cos9)).
fof(fraenkel_a_0_2_sin_cos9, axiom, ![X1]:(r2_hidden(X1,a_0_2_sin_cos9)<=>?[X2]:(m1_subset_1(X2,k1_numbers)&(X1=X2&(r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X2)&r1_xxreal_0(X2,k10_real_1(k32_sin_cos,np__4)))))), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', fraenkel_a_0_2_sin_cos9)).
fof(t39_sin_cos9, axiom, (k5_sin_cos9(np__1)=k10_real_1(k32_sin_cos,np__4)&k1_seq_1(k1_sin_cos9,np__1)=k10_real_1(k32_sin_cos,np__4)), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', t39_sin_cos9)).
fof(rqLessOrEqual__r1_xxreal_0__r1_r1, axiom, r1_xxreal_0(np__1,np__1), file('/local1/mptp/loop2/hardpred_knn256_miz1/preds__0.5/t35_sincos10', rqLessOrEqual__r1_xxreal_0__r1_r1)).
fof(c_0_83, plain, ![X53, X54]:(~v1_xxreal_0(X53)|(~v1_xxreal_0(X54)|(~r1_xxreal_0(X53,X54)|k1_xxreal_1(X53,X54)=k2_xboole_0(k4_xxreal_1(X53,X54),k2_tarski(X53,X54))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t128_xxreal_1])])])).
fof(c_0_84, plain, ![X253, X254]:k2_xboole_0(X253,X254)=k2_xboole_0(X254,X253), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
cnf(c_0_85, plain, (k1_xxreal_1(X1,X2)=k2_xboole_0(k4_xxreal_1(X1,X2),k2_tarski(X1,X2))|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_86, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
fof(c_0_87, plain, ![X148, X149]:(~v1_xxreal_0(X148)|~v1_xxreal_0(X149)|k2_rcomp_1(X148,X149)=k4_xxreal_1(X148,X149)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_rcomp_1])])).
fof(c_0_88, plain, ![X228, X229]:k3_xboole_0(X228,X229)=k3_xboole_0(X229,X228), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_89, plain, ![X109, X110]:(~r1_tarski(X109,X110)|k3_xboole_0(X109,X110)=X109), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_90, plain, ![X47, X48]:((~r1_xxreal_0(X47,X48)|r1_xxreal_0(k4_xcmplx_0(X48),k4_xcmplx_0(X47))|~v1_xreal_0(X48)|~v1_xreal_0(X47))&(~r1_xxreal_0(k4_xcmplx_0(X48),k4_xcmplx_0(X47))|r1_xxreal_0(X47,X48)|~v1_xreal_0(X48)|~v1_xreal_0(X47))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t24_xreal_1])])])])).
fof(c_0_91, plain, ![X30]:(~m1_subset_1(X30,k1_numbers)|k1_real_1(X30)=k4_xcmplx_0(X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_real_1])])).
fof(c_0_92, plain, ![X34, X35]:(~r2_hidden(X34,X35)|m1_subset_1(X34,X35)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_93, plain, (r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))), inference(split_conjunct,[status(thm)],[l19_sincos10])).
cnf(c_0_94, plain, (k2_relset_1(k1_numbers,k3_sincos10)=k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)), inference(split_conjunct,[status(thm)],[t27_sincos10])).
fof(c_0_95, plain, ![X143, X144, X145]:(~r1_tarski(X143,X144)|~r1_tarski(X144,X145)|r1_tarski(X143,X145)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_96, plain, ![X130, X131]:((~m1_subset_1(X130,k1_zfmisc_1(X131))|r1_tarski(X130,X131))&(~r1_tarski(X130,X131)|m1_subset_1(X130,k1_zfmisc_1(X131)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_97, plain, ![X38, X39, X40, X41, X42, X43, X44, X45]:(((~r2_hidden(X41,X40)|(r2_hidden(X41,X38)|r2_hidden(X41,X39))|X40!=k2_xboole_0(X38,X39))&((~r2_hidden(X42,X38)|r2_hidden(X42,X40)|X40!=k2_xboole_0(X38,X39))&(~r2_hidden(X42,X39)|r2_hidden(X42,X40)|X40!=k2_xboole_0(X38,X39))))&(((~r2_hidden(esk6_3(X43,X44,X45),X43)|~r2_hidden(esk6_3(X43,X44,X45),X45)|X45=k2_xboole_0(X43,X44))&(~r2_hidden(esk6_3(X43,X44,X45),X44)|~r2_hidden(esk6_3(X43,X44,X45),X45)|X45=k2_xboole_0(X43,X44)))&(r2_hidden(esk6_3(X43,X44,X45),X45)|(r2_hidden(esk6_3(X43,X44,X45),X43)|r2_hidden(esk6_3(X43,X44,X45),X44))|X45=k2_xboole_0(X43,X44)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_xboole_0])])])])])])).
cnf(c_0_98, plain, (k1_xxreal_1(X1,X2)=k2_xboole_0(k4_xxreal_1(X1,X2),k2_tarski(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_85])).
cnf(c_0_99, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(evalgc,[status(thm)],[c_0_86])).
cnf(c_0_100, plain, (k2_rcomp_1(X1,X2)=k4_xxreal_1(X1,X2)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_87])).
fof(c_0_101, plain, ![X114]:(~v1_xboole_0(X114)|X114=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
fof(c_0_102, plain, ![X8, X9, X10, X11]:((~v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X9,X8))|(~m1_subset_1(X10,k1_numbers)|(~m1_subset_1(X11,k1_numbers)|(~r2_hidden(X10,k9_subset_1(k1_numbers,X8,k1_relset_1(k1_numbers,X9)))|(~r2_hidden(X11,k9_subset_1(k1_numbers,X8,k1_relset_1(k1_numbers,X9)))|(r1_xxreal_0(X11,X10)|~r1_xxreal_0(k1_seq_1(X9,X10),k1_seq_1(X9,X11)))))))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&((m1_subset_1(esk2_2(X8,X9),k1_numbers)|v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X9,X8))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&((m1_subset_1(esk3_2(X8,X9),k1_numbers)|v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X9,X8))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&((r2_hidden(esk2_2(X8,X9),k9_subset_1(k1_numbers,X8,k1_relset_1(k1_numbers,X9)))|v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X9,X8))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&((r2_hidden(esk3_2(X8,X9),k9_subset_1(k1_numbers,X8,k1_relset_1(k1_numbers,X9)))|v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X9,X8))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&((~r1_xxreal_0(esk3_2(X8,X9),esk2_2(X8,X9))|v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X9,X8))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(r1_xxreal_0(k1_seq_1(X9,esk2_2(X8,X9)),k1_seq_1(X9,esk3_2(X8,X9)))|v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X9,X8))|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t21_rfunct_2])])])])])])).
fof(c_0_103, plain, ![X119, X120, X121]:(~m1_subset_1(X121,k1_zfmisc_1(X119))|k9_subset_1(X119,X120,X121)=k9_subset_1(X119,X121,X120)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k9_subset_1])])).
fof(c_0_104, plain, ![X184, X185, X186]:(~m1_subset_1(X186,k1_zfmisc_1(X184))|k9_subset_1(X184,X185,X186)=k3_xboole_0(X185,X186)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k9_subset_1])])).
cnf(c_0_105, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_106, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_107, plain, (k2_funct_1(k3_sincos10)=k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))), inference(split_conjunct,[status(thm)],[l63_sincos10])).
fof(c_0_108, plain, ![X94, X95, X96]:((v4_relat_1(X96,X94)|~m1_subset_1(X96,k1_zfmisc_1(k2_zfmisc_1(X94,X95))))&(v5_relat_1(X96,X95)|~m1_subset_1(X96,k1_zfmisc_1(k2_zfmisc_1(X94,X95))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
fof(c_0_109, plain, ![X20, X21, X22]:(~m1_subset_1(X22,k1_zfmisc_1(k2_zfmisc_1(X20,X21)))|v1_relat_1(X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_110, plain, (r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_111, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_91])).
fof(c_0_112, plain, ![X88]:(~m1_subset_1(X88,k1_numbers)|v1_xreal_0(X88)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_xreal_0])])).
fof(c_0_113, plain, ![X111, X112]:(~v2_membered(X111)|(~m1_subset_1(X112,X111)|v1_xxreal_0(X112))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc13_membered])])])).
cnf(c_0_114, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_115, plain, (r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))), inference(evalgc,[status(thm)],[c_0_93])).
cnf(c_0_116, plain, (k2_relset_1(k1_numbers,k3_sincos10)=k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)), inference(evalgc,[status(thm)],[c_0_94])).
cnf(c_0_117, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_118, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_96])).
fof(c_0_119, plain, ![X162, X163]:(~v1_relat_1(X163)|~v4_relat_1(X163,X162)|m1_subset_1(k1_relset_1(X162,X163),k1_zfmisc_1(X162))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_relset_1])])).
cnf(c_0_120, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_121, plain, (k2_xboole_0(k2_tarski(X1,X2),k4_xxreal_1(X1,X2))=k1_xxreal_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)), inference(rw,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_122, plain, (k2_rcomp_1(X1,X2)=k4_xxreal_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_100])).
fof(c_0_123, plain, (~v1_xboole_0(k4_ordinal1)&v3_ordinal1(k4_ordinal1)), inference(fof_simplification,[status(thm)],[fc6_ordinal1])).
cnf(c_0_124, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_125, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
fof(c_0_126, plain, (v1_xboole_0(esk8_0)&(v1_xcmplx_0(esk8_0)&(v1_xxreal_0(esk8_0)&v1_xreal_0(esk8_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc4_xreal_0])])).
cnf(c_0_127, plain, (r1_xxreal_0(X4,X3)|~v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X1,X2))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X4,k1_numbers)|~r2_hidden(X3,k9_subset_1(k1_numbers,X2,k1_relset_1(k1_numbers,X1)))|~r2_hidden(X4,k9_subset_1(k1_numbers,X2,k1_relset_1(k1_numbers,X1)))|~r1_xxreal_0(k1_seq_1(X1,X3),k1_seq_1(X1,X4))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_128, plain, (k9_subset_1(X2,X3,X1)=k9_subset_1(X2,X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_129, plain, (k9_subset_1(X2,X3,X1)=k3_xboole_0(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_130, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_131, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_106])).
fof(c_0_132, plain, ![X82, X83, X84]:(~v1_funct_1(X84)|~m1_subset_1(X84,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|(~r1_tarski(X82,X83)|~v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X84,X83))|v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X84,X82)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t29_rfunct_2])])).
cnf(c_0_133, plain, (v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))), inference(split_conjunct,[status(thm)],[t19_sincos10])).
cnf(c_0_134, plain, (k2_funct_1(k3_sincos10)=k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))), inference(evalgc,[status(thm)],[c_0_107])).
fof(c_0_135, plain, ![X338, X339]:(~v1_xxreal_0(X338)|~v1_xxreal_0(X339)|(r1_xxreal_0(X338,X339)|r1_xxreal_0(X339,X338))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[connectedness_r1_xxreal_0])])).
fof(c_0_136, plain, ![X76]:(~v1_xreal_0(X76)|v1_xxreal_0(X76)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_xreal_0])])).
fof(c_0_137, plain, ![X92, X93]:(~v1_relat_1(X92)|(~v1_funct_1(X92)|~v3_valued_0(X92))|v1_xreal_0(k1_funct_1(X92,X93))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc63_valued_0])])).
fof(c_0_138, plain, ![X97, X98]:(~v1_relat_1(X97)|(~v1_funct_1(X97)|~v3_valued_0(X97))|k1_seq_1(X97,X98)=k1_funct_1(X97,X98)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_seq_1])])).
fof(c_0_139, plain, ![X233]:((v1_relat_1(X233)|(~v1_relat_1(X233)|~v5_relat_1(X233,k1_numbers)))&(v3_valued_0(X233)|(~v1_relat_1(X233)|~v5_relat_1(X233,k1_numbers)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc30_valued_0])])])).
cnf(c_0_140, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_108])).
cnf(c_0_141, plain, (m1_subset_1(k2_fdiff_9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[dt_k2_fdiff_9])).
cnf(c_0_142, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_109])).
fof(c_0_143, plain, ![X284, X285, X286, X288]:(((m1_subset_1(esk15_3(X284,X285,X286),k1_numbers)|~r2_hidden(X284,a_2_0_rcomp_1(X285,X286))|(~v1_xreal_0(X285)|~v1_xreal_0(X286)))&((X284=esk15_3(X284,X285,X286)|~r2_hidden(X284,a_2_0_rcomp_1(X285,X286))|(~v1_xreal_0(X285)|~v1_xreal_0(X286)))&((r1_xxreal_0(X285,esk15_3(X284,X285,X286))|~r2_hidden(X284,a_2_0_rcomp_1(X285,X286))|(~v1_xreal_0(X285)|~v1_xreal_0(X286)))&(r1_xxreal_0(esk15_3(X284,X285,X286),X286)|~r2_hidden(X284,a_2_0_rcomp_1(X285,X286))|(~v1_xreal_0(X285)|~v1_xreal_0(X286))))))&(~m1_subset_1(X288,k1_numbers)|(X284!=X288|(~r1_xxreal_0(X285,X288)|~r1_xxreal_0(X288,X286)))|r2_hidden(X284,a_2_0_rcomp_1(X285,X286))|(~v1_xreal_0(X285)|~v1_xreal_0(X286)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fraenkel_a_2_0_rcomp_1])])])])])).
fof(c_0_144, plain, ![X65, X66]:(~v1_xreal_0(X65)|(~v1_xreal_0(X66)|k1_rcomp_1(X65,X66)=a_2_0_rcomp_1(X65,X66))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_rcomp_1])])])).
cnf(c_0_145, plain, (r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_110])).
cnf(c_0_146, plain, (k4_xcmplx_0(X1)=k1_real_1(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_111])).
cnf(c_0_147, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_148, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_149, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_114])).
cnf(c_0_150, plain, (r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k2_relset_1(k1_numbers,k3_sincos10))), inference(rw,[status(thm)],[c_0_115, c_0_116])).
fof(c_0_151, plain, ![X122]:(~v3_membered(X122)|v2_membered(X122)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_membered])])).
cnf(c_0_152, plain, (r1_tarski(X1,X3)|~r1_tarski(X2,X3)|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_117])).
cnf(c_0_153, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_118])).
cnf(c_0_154, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_155, plain, (r1_tarski(k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers),k1_relset_1(k1_numbers,k2_fdiff_9))), inference(split_conjunct,[status(thm)],[t3_sincos10])).
cnf(c_0_156, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_108])).
cnf(c_0_157, plain, (r2_hidden(X1,X3)|X3!=k2_xboole_0(X2,X4)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_120])).
cnf(c_0_158, plain, (k2_xboole_0(k2_tarski(X1,X2),k2_rcomp_1(X1,X2))=k1_xxreal_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122])])).
fof(c_0_159, plain, ![X55, X56]:(~v1_xreal_0(X55)|~v1_xreal_0(X56)|k1_rcomp_1(X55,X56)=k1_xxreal_1(X55,X56)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_rcomp_1])])).
fof(c_0_160, plain, ![X205, X206, X207]:(v1_xboole_0(X205)|(v1_xboole_0(X206)|~m1_subset_1(X206,k1_zfmisc_1(X205)))|(~m2_subset_1(X207,X205,X206)|m1_subset_1(X207,X205))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m2_subset_1])])])])).
fof(c_0_161, plain, ~v1_xboole_0(k1_numbers), inference(fof_simplification,[status(thm)],[fc1_numbers])).
cnf(c_0_162, plain, (~v1_xboole_0(k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_123])).
cnf(c_0_163, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
cnf(c_0_164, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_124])).
cnf(c_0_165, plain, (k1_xboole_0=k6_numbers), inference(evalgc,[status(thm)],[c_0_125])).
cnf(c_0_166, plain, (v1_xboole_0(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_126])).
cnf(c_0_167, plain, (r1_xxreal_0(X2,X1)|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_168, plain, (r1_xxreal_0(X4,X3)|~v1_funct_1(X1)|~m1_subset_1(X4,k1_numbers)|~m1_subset_1(X3,k1_numbers)|~v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X1,X2))|~r1_xxreal_0(k1_seq_1(X1,X3),k1_seq_1(X1,X4))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_hidden(X4,k9_subset_1(k1_numbers,X2,k1_relset_1(k1_numbers,X1)))|~r2_hidden(X3,k9_subset_1(k1_numbers,X2,k1_relset_1(k1_numbers,X1)))), inference(evalgc,[status(thm)],[c_0_127])).
cnf(c_0_169, plain, (k9_subset_1(X2,X1,X3)=k9_subset_1(X2,X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_128])).
cnf(c_0_170, plain, (k9_subset_1(X2,X3,X1)=k3_xboole_0(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_129])).
cnf(c_0_171, plain, (k3_xboole_0(X2,X1)=X1|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_131])])).
fof(c_0_172, plain, ![X79, X80, X81]:(~r2_hidden(X79,X80)|~m1_subset_1(X80,k1_zfmisc_1(X81))|m1_subset_1(X79,X81)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_173, plain, (v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r1_tarski(X2,X3)|~v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X1,X3))), inference(split_conjunct,[status(thm)],[c_0_132])).
cnf(c_0_174, plain, (v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))), inference(evalgc,[status(thm)],[c_0_133])).
cnf(c_0_175, plain, (k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k2_relset_1(k1_numbers,k3_sincos10))=k2_funct_1(k3_sincos10)), inference(rw,[status(thm)],[c_0_134, c_0_116])).
cnf(c_0_176, plain, (v1_funct_1(k2_fdiff_9)), inference(split_conjunct,[status(thm)],[dt_k2_fdiff_9])).
cnf(c_0_177, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_135])).
cnf(c_0_178, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_136])).
cnf(c_0_179, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_137])).
cnf(c_0_180, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_138])).
cnf(c_0_181, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_139])).
cnf(c_0_182, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(evalgc,[status(thm)],[c_0_140])).
cnf(c_0_183, plain, (m1_subset_1(k2_fdiff_9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_141])).
cnf(c_0_184, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_142])).
fof(c_0_185, plain, ![X200, X201]:(~v1_xxreal_0(X200)|(~v1_xxreal_0(X201)|(~r1_xxreal_0(X200,X201)|~r1_xxreal_0(X201,X200)|X200=X201))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xxreal_0])])])).
cnf(c_0_186, plain, (r1_xxreal_0(X1,esk15_3(X2,X1,X3))|~r2_hidden(X2,a_2_0_rcomp_1(X1,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X3)), inference(split_conjunct,[status(thm)],[c_0_143])).
cnf(c_0_187, plain, (k1_rcomp_1(X1,X2)=a_2_0_rcomp_1(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_144])).
fof(c_0_188, plain, ![X67, X68, X69, X70, X71, X72, X73, X74]:(((~r2_hidden(X70,X69)|(X70=X67|X70=X68)|X69!=k2_tarski(X67,X68))&((X71!=X67|r2_hidden(X71,X69)|X69!=k2_tarski(X67,X68))&(X71!=X68|r2_hidden(X71,X69)|X69!=k2_tarski(X67,X68))))&(((esk7_3(X72,X73,X74)!=X72|~r2_hidden(esk7_3(X72,X73,X74),X74)|X74=k2_tarski(X72,X73))&(esk7_3(X72,X73,X74)!=X73|~r2_hidden(esk7_3(X72,X73,X74),X74)|X74=k2_tarski(X72,X73)))&(r2_hidden(esk7_3(X72,X73,X74),X74)|(esk7_3(X72,X73,X74)=X72|esk7_3(X72,X73,X74)=X73)|X74=k2_tarski(X72,X73)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
cnf(c_0_189, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_146])])).
cnf(c_0_190, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_147])).
cnf(c_0_191, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_148])).
cnf(c_0_192, plain, (m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k2_relset_1(k1_numbers,k3_sincos10))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_150])])).
cnf(c_0_193, plain, (v2_membered(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_151])).
fof(c_0_194, plain, ![X128]:(~m1_subset_1(X128,k1_zfmisc_1(k1_numbers))|v3_membered(X128)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_membered])])).
cnf(c_0_195, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_153])])).
cnf(c_0_196, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_154])).
cnf(c_0_197, plain, (r1_tarski(k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers),k1_relset_1(k1_numbers,k2_fdiff_9))), inference(evalgc,[status(thm)],[c_0_155])).
cnf(c_0_198, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_156])).
cnf(c_0_199, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_157])])).
cnf(c_0_200, plain, (k2_xboole_0(k2_rcomp_1(X1,X2),k2_tarski(X1,X2))=k1_xxreal_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)), inference(rw,[status(thm)],[c_0_158, c_0_99])).
cnf(c_0_201, plain, (k1_rcomp_1(X1,X2)=k1_xxreal_1(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_159])).
fof(c_0_202, plain, ![X36, X37]:(~v3_membered(X36)|(~m1_subset_1(X37,X36)|v1_xreal_0(X37))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc14_membered])])])).
cnf(c_0_203, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m2_subset_1(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_160])).
cnf(c_0_204, plain, (m2_subset_1(k6_numbers,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[dt_k6_numbers])).
cnf(c_0_205, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[dt_k5_numbers])).
cnf(c_0_206, plain, (~v1_xboole_0(k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_161])).
cnf(c_0_207, plain, (~v1_xboole_0(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_162])).
cnf(c_0_208, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_163])).
cnf(c_0_209, plain, (v1_xxreal_0(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_126])).
cnf(c_0_210, plain, (X1=k6_numbers|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_164, c_0_165])).
cnf(c_0_211, plain, (v1_xboole_0(esk8_0)), inference(evalgc,[status(thm)],[c_0_166])).
cnf(c_0_212, plain, (r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))), inference(evalgc,[status(thm)],[c_0_167])).
cnf(c_0_213, plain, (r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))), inference(split_conjunct,[status(thm)],[l19_sincos10])).
cnf(c_0_214, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(X3,X2),k1_seq_1(X3,X1))|~v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_1(X3)|~r2_hidden(X1,k9_subset_1(k1_numbers,k1_relset_1(k1_numbers,X3),X4))|~r2_hidden(X2,k9_subset_1(k1_numbers,k1_relset_1(k1_numbers,X3),X4))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_169])])).
cnf(c_0_215, plain, (k9_subset_1(X1,X2,X3)=X3|~m1_subset_1(X3,k1_zfmisc_1(X1))|~r1_tarski(X3,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_171])])).
cnf(c_0_216, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_172])).
cnf(c_0_217, plain, (v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X1,X2))|~v1_funct_1(X1)|~r1_tarski(X2,X3)|~v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X1,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_173])).
cnf(c_0_218, plain, (v6_valued_0(k2_funct_1(k3_sincos10))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_174, c_0_116]), c_0_175])).
cnf(c_0_219, plain, (v1_funct_1(k2_fdiff_9)), inference(evalgc,[status(thm)],[c_0_176])).
cnf(c_0_220, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_177])).
cnf(c_0_221, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_178])).
cnf(c_0_222, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_179])).
cnf(c_0_223, plain, (k1_funct_1(X1,X2)=k1_seq_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_180])).
cnf(c_0_224, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_181])).
cnf(c_0_225, plain, (v5_relat_1(k2_fdiff_9,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_183])])).
cnf(c_0_226, plain, (v1_relat_1(k2_fdiff_9)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_183])])).
cnf(c_0_227, plain, (r1_tarski(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))),k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))), inference(split_conjunct,[status(thm)],[l31_sincos10])).
fof(c_0_228, plain, ![X279, X280, X281, X283]:(((m1_subset_1(esk14_3(X279,X280,X281),k1_numbers)|~r2_hidden(X279,a_2_2_rcomp_1(X280,X281))|(~v1_xreal_0(X280)|~v1_xxreal_0(X281)))&((X279=esk14_3(X279,X280,X281)|~r2_hidden(X279,a_2_2_rcomp_1(X280,X281))|(~v1_xreal_0(X280)|~v1_xxreal_0(X281)))&((r1_xxreal_0(X280,esk14_3(X279,X280,X281))|~r2_hidden(X279,a_2_2_rcomp_1(X280,X281))|(~v1_xreal_0(X280)|~v1_xxreal_0(X281)))&(~r1_xxreal_0(X281,esk14_3(X279,X280,X281))|~r2_hidden(X279,a_2_2_rcomp_1(X280,X281))|(~v1_xreal_0(X280)|~v1_xxreal_0(X281))))))&(~m1_subset_1(X283,k1_numbers)|(X279!=X283|(~r1_xxreal_0(X280,X283)|r1_xxreal_0(X281,X283)))|r2_hidden(X279,a_2_2_rcomp_1(X280,X281))|(~v1_xreal_0(X280)|~v1_xxreal_0(X281)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_2_rcomp_1])])])])])])).
fof(c_0_229, plain, ![X107, X108]:(~v1_xreal_0(X107)|(~v1_xxreal_0(X108)|k3_rcomp_1(X107,X108)=a_2_2_rcomp_1(X107,X108))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_rcomp_1])])])).
cnf(c_0_230, plain, (X1=X2|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_185])).
cnf(c_0_231, plain, (r1_xxreal_0(X1,esk15_3(X2,X1,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r2_hidden(X2,a_2_0_rcomp_1(X1,X3))), inference(evalgc,[status(thm)],[c_0_186])).
cnf(c_0_232, plain, (a_2_0_rcomp_1(X1,X2)=k1_rcomp_1(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_187])).
cnf(c_0_233, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_188])).
cnf(c_0_234, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X2))|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_189, c_0_190])).
fof(c_0_235, plain, ![X57]:(~m1_subset_1(X57,k1_numbers)|k1_real_1(k1_real_1(X57))=X57), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k1_real_1])])).
fof(c_0_236, plain, ![X178, X179]:(~v1_xxreal_0(X178)|~v1_xxreal_0(X179)|r1_xxreal_0(X178,X178)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[reflexivity_r1_xxreal_0])])).
fof(c_0_237, plain, (v1_xxreal_0(esk24_0)&v3_xxreal_0(esk24_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc3_xxreal_0])])).
cnf(c_0_238, plain, (v1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v2_membered(k2_relset_1(k1_numbers,k3_sincos10))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_192])])).
cnf(c_0_239, plain, (v2_membered(X1)|~v3_membered(X1)), inference(evalgc,[status(thm)],[c_0_193])).
cnf(c_0_240, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_194])).
cnf(c_0_241, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_242, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(X2,X3))|~v4_relat_1(X3,X2)|~v1_relat_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_196])])).
cnf(c_0_243, plain, (r1_tarski(k2_relset_1(k1_numbers,k3_sincos10),k1_relset_1(k1_numbers,k2_fdiff_9))), inference(rw,[status(thm)],[c_0_197, c_0_116])).
cnf(c_0_244, plain, (v4_relat_1(k2_fdiff_9,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_183])])).
cnf(c_0_245, plain, (r2_hidden(X1,k1_xxreal_1(X2,X3))|~r2_hidden(X1,k2_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_200])])).
cnf(c_0_246, plain, (k1_xxreal_1(X1,X2)=k1_rcomp_1(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_201])).
cnf(c_0_247, plain, (v1_xreal_0(k31_sin_cos)), inference(split_conjunct,[status(thm)],[dt_k31_sin_cos])).
cnf(c_0_248, plain, (k32_sin_cos=k31_sin_cos), inference(split_conjunct,[status(thm)],[redefinition_k32_sin_cos])).
cnf(c_0_249, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_202])).
cnf(c_0_250, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_203])).
cnf(c_0_251, plain, (m2_subset_1(k6_numbers,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_204])).
cnf(c_0_252, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_205])).
cnf(c_0_253, plain, (~v1_xboole_0(k1_numbers)), inference(evalgc,[status(thm)],[c_0_206])).
cnf(c_0_254, plain, (~v1_xboole_0(k5_numbers)), inference(rw,[status(thm)],[c_0_207, c_0_208])).
cnf(c_0_255, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
cnf(c_0_256, plain, (v1_xxreal_0(esk8_0)), inference(evalgc,[status(thm)],[c_0_209])).
cnf(c_0_257, plain, (esk8_0=k6_numbers), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_211])])).
fof(c_0_258, plain, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k6_numbers)&(~r1_xxreal_0(k32_sin_cos,k10_real_1(k32_sin_cos,np__2))&(~r1_xxreal_0(k32_sin_cos,k6_numbers)&(~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k1_real_1(k10_real_1(k32_sin_cos,np__2)))&(~r1_xxreal_0(k8_real_1(np__2,k32_sin_cos),k32_sin_cos)&(~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos),k10_real_1(k32_sin_cos,np__2))&(~r1_xxreal_0(k6_numbers,k1_real_1(k10_real_1(k32_sin_cos,np__2)))&(~r1_xxreal_0(k8_real_1(np__2,k32_sin_cos),k6_numbers)&(~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos),k32_sin_cos)&(~r1_xxreal_0(k8_real_1(np__2,k32_sin_cos),k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos))&~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos),k6_numbers))))))))))), inference(fof_simplification,[status(thm)],[t5_comptrig])).
cnf(c_0_259, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_146])])).
cnf(c_0_260, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))), inference(split_conjunct,[status(thm)],[l20_sincos10])).
fof(c_0_261, plain, ![X300, X301, X302, X304]:(((m1_subset_1(esk16_3(X300,X301,X302),k1_numbers)|~r2_hidden(X300,a_2_1_rcomp_1(X301,X302))|(~v1_xxreal_0(X301)|~v1_xxreal_0(X302)))&((X300=esk16_3(X300,X301,X302)|~r2_hidden(X300,a_2_1_rcomp_1(X301,X302))|(~v1_xxreal_0(X301)|~v1_xxreal_0(X302)))&((~r1_xxreal_0(esk16_3(X300,X301,X302),X301)|~r2_hidden(X300,a_2_1_rcomp_1(X301,X302))|(~v1_xxreal_0(X301)|~v1_xxreal_0(X302)))&(~r1_xxreal_0(X302,esk16_3(X300,X301,X302))|~r2_hidden(X300,a_2_1_rcomp_1(X301,X302))|(~v1_xxreal_0(X301)|~v1_xxreal_0(X302))))))&(~m1_subset_1(X304,k1_numbers)|(X300!=X304|(r1_xxreal_0(X304,X301)|r1_xxreal_0(X302,X304)))|r2_hidden(X300,a_2_1_rcomp_1(X301,X302))|(~v1_xxreal_0(X301)|~v1_xxreal_0(X302)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_1_rcomp_1])])])])])])).
fof(c_0_262, plain, ![X117, X118]:(~v1_xxreal_0(X117)|(~v1_xxreal_0(X118)|k2_rcomp_1(X117,X118)=a_2_1_rcomp_1(X117,X118))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_rcomp_1])])])).
cnf(c_0_263, plain, (r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))), inference(evalgc,[status(thm)],[c_0_213])).
cnf(c_0_264, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(X3,X2),k1_seq_1(X3,X1))|~v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_1(X3)|~r2_hidden(X1,X4)|~r2_hidden(X2,X4)|~r1_tarski(X4,k1_relset_1(k1_numbers,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_215])])).
cnf(c_0_265, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_216])).
cnf(c_0_266, plain, (v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))|~r1_tarski(X1,k2_relset_1(k1_numbers,k3_sincos10))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_175]), c_0_218]), c_0_183]), c_0_219])])])).
cnf(c_0_267, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_221])])).
cnf(c_0_268, plain, (v1_xreal_0(k1_seq_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_223])])).
cnf(c_0_269, plain, (v3_valued_0(k2_fdiff_9)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_224, c_0_225]), c_0_226])])])).
cnf(c_0_270, plain, (r1_tarski(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))),k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))), inference(evalgc,[status(thm)],[c_0_227])).
cnf(c_0_271, plain, (r1_xxreal_0(X1,esk14_3(X2,X1,X3))|~r2_hidden(X2,a_2_2_rcomp_1(X1,X3))|~v1_xreal_0(X1)|~v1_xxreal_0(X3)), inference(split_conjunct,[status(thm)],[c_0_228])).
cnf(c_0_272, plain, (k3_rcomp_1(X1,X2)=a_2_2_rcomp_1(X1,X2)|~v1_xreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_229])).
cnf(c_0_273, plain, (X1=esk14_3(X1,X2,X3)|~r2_hidden(X1,a_2_2_rcomp_1(X2,X3))|~v1_xreal_0(X2)|~v1_xxreal_0(X3)), inference(split_conjunct,[status(thm)],[c_0_228])).
cnf(c_0_274, plain, (X1=X2|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_230])).
cnf(c_0_275, plain, (r1_xxreal_0(X1,esk15_3(X2,X1,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r2_hidden(X2,k1_rcomp_1(X1,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_232])])).
cnf(c_0_276, plain, (X1=esk15_3(X1,X2,X3)|~r2_hidden(X1,a_2_0_rcomp_1(X2,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3)), inference(split_conjunct,[status(thm)],[c_0_143])).
cnf(c_0_277, plain, (r2_hidden(X1,k1_xxreal_1(X2,X3))|~r2_hidden(X1,k2_tarski(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_121])])).
cnf(c_0_278, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(evalgc,[status(thm)],[c_0_233])).
cnf(c_0_279, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_146])])).
cnf(c_0_280, plain, (k1_real_1(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_235])).
fof(c_0_281, plain, ![X51]:(~m1_subset_1(X51,k1_numbers)|m1_subset_1(k1_real_1(X51),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_real_1])])).
cnf(c_0_282, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_236])).
cnf(c_0_283, plain, (v1_xxreal_0(esk24_0)), inference(split_conjunct,[status(thm)],[c_0_237])).
cnf(c_0_284, plain, (v1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v3_membered(k2_relset_1(k1_numbers,k3_sincos10))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_239])])).
cnf(c_0_285, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_240])).
cnf(c_0_286, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_241])).
cnf(c_0_287, plain, (r1_tarski(k2_relset_1(k1_numbers,k3_sincos10),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_243]), c_0_244]), c_0_226])])])).
fof(c_0_288, plain, ![X115, X116]:(~v1_xreal_0(X115)|~v1_xreal_0(X116)|m1_subset_1(k1_rcomp_1(X115,X116),k1_zfmisc_1(k1_numbers))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_rcomp_1])])).
cnf(c_0_289, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3)|~r2_hidden(X1,k2_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_246])])).
cnf(c_0_290, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),k2_rcomp_1(k6_numbers,k32_sin_cos))), inference(split_conjunct,[status(thm)],[l25_sin_cos9])).
cnf(c_0_291, plain, (v1_xreal_0(k31_sin_cos)), inference(evalgc,[status(thm)],[c_0_247])).
cnf(c_0_292, plain, (k31_sin_cos=k32_sin_cos), inference(evalgc,[status(thm)],[c_0_248])).
cnf(c_0_293, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_249])).
cnf(c_0_294, plain, (m1_subset_1(k6_numbers,k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_251]), c_0_252])]), c_0_253]), c_0_254])])).
cnf(c_0_295, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_255])).
cnf(c_0_296, plain, (v1_xxreal_0(k6_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_256, c_0_257])])).
cnf(c_0_297, plain, (~r1_xxreal_0(k32_sin_cos,k6_numbers)), inference(split_conjunct,[status(thm)],[c_0_258])).
cnf(c_0_298, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_259, c_0_190])).
cnf(c_0_299, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[c_0_260])).
cnf(c_0_300, plain, (X1=esk16_3(X1,X2,X3)|~r2_hidden(X1,a_2_1_rcomp_1(X2,X3))|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)), inference(split_conjunct,[status(thm)],[c_0_261])).
cnf(c_0_301, plain, (k2_rcomp_1(X1,X2)=a_2_1_rcomp_1(X1,X2)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_262])).
cnf(c_0_302, plain, (r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k2_relset_1(k1_numbers,k3_sincos10))), inference(rw,[status(thm)],[c_0_263, c_0_116])).
fof(c_0_303, negated_conjecture, ~(![X1]:(r2_hidden(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))=>r2_hidden(k1_seq_1(k2_fdiff_9,X1),k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1))))), inference(assume_negation,[status(cth)],[t35_sincos10])).
fof(c_0_304, plain, ![X17, X18, X19]:(((r1_xxreal_0(X18,X17)|~r2_hidden(X17,k1_xxreal_1(X18,X19))|~v1_xxreal_0(X19)|~v1_xxreal_0(X18)|~v1_xxreal_0(X17))&(r1_xxreal_0(X17,X19)|~r2_hidden(X17,k1_xxreal_1(X18,X19))|~v1_xxreal_0(X19)|~v1_xxreal_0(X18)|~v1_xxreal_0(X17)))&(~r1_xxreal_0(X18,X17)|~r1_xxreal_0(X17,X19)|r2_hidden(X17,k1_xxreal_1(X18,X19))|~v1_xxreal_0(X19)|~v1_xxreal_0(X18)|~v1_xxreal_0(X17))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xxreal_1])])])])).
cnf(c_0_305, plain, (r1_xxreal_0(X1,X2)|~r1_tarski(X3,k1_relset_1(k1_numbers,X4))|~r1_xxreal_0(k1_seq_1(X4,X2),k1_seq_1(X4,X1))|~v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X4,X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v1_funct_1(X4)|~r2_hidden(X1,X3)|~r2_hidden(X2,X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_264, c_0_265]), c_0_265])).
cnf(c_0_306, plain, (v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_153])])).
cnf(c_0_307, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k2_fdiff_9))|~r1_tarski(X1,k2_relset_1(k1_numbers,k3_sincos10))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_243])])).
cnf(c_0_308, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_221])])).
cnf(c_0_309, plain, (v1_xreal_0(k1_seq_1(k2_fdiff_9,X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_269]), c_0_226]), c_0_219])])])).
cnf(c_0_310, plain, (r1_tarski(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))),k2_relset_1(k1_numbers,k3_sincos10))), inference(rw,[status(thm)],[c_0_270, c_0_116])).
cnf(c_0_311, plain, (r1_xxreal_0(X1,esk14_3(X2,X1,X3))|~v1_xreal_0(X1)|~v1_xxreal_0(X3)|~r2_hidden(X2,a_2_2_rcomp_1(X1,X3))), inference(evalgc,[status(thm)],[c_0_271])).
cnf(c_0_312, plain, (a_2_2_rcomp_1(X1,X2)=k3_rcomp_1(X1,X2)|~v1_xreal_0(X1)|~v1_xxreal_0(X2)), inference(evalgc,[status(thm)],[c_0_272])).
cnf(c_0_313, plain, (X1=esk14_3(X1,X2,X3)|~v1_xreal_0(X2)|~v1_xxreal_0(X3)|~r2_hidden(X1,a_2_2_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[c_0_273])).
cnf(c_0_314, plain, (esk15_3(X1,X2,X3)=X2|~v1_xxreal_0(X2)|~v1_xxreal_0(esk15_3(X1,X2,X3))|~r1_xxreal_0(esk15_3(X1,X2,X3),X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_274, c_0_275])])).
cnf(c_0_315, plain, (X1=esk15_3(X1,X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,a_2_0_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[c_0_276])).
cnf(c_0_316, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3)|~r2_hidden(X1,k2_tarski(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_277, c_0_246])])).
cnf(c_0_317, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X3,X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_278])])).
cnf(c_0_318, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_279, c_0_190])).
cnf(c_0_319, plain, (k1_real_1(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_280])).
cnf(c_0_320, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_281])).
cnf(c_0_321, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_282])).
cnf(c_0_322, plain, (v1_xxreal_0(esk24_0)), inference(evalgc,[status(thm)],[c_0_283])).
cnf(c_0_323, plain, (v1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(k2_relset_1(k1_numbers,k3_sincos10),k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_284, c_0_285])])).
cnf(c_0_324, plain, (m1_subset_1(k2_relset_1(k1_numbers,k3_sincos10),k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_287])])).
cnf(c_0_325, plain, (m1_subset_1(k1_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_288])).
cnf(c_0_326, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X3)|~r2_hidden(X1,k2_rcomp_1(X2,X3))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_289, c_0_221]), c_0_221])).
cnf(c_0_327, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),k2_rcomp_1(k6_numbers,k32_sin_cos))), inference(evalgc,[status(thm)],[c_0_290])).
cnf(c_0_328, plain, (v1_xreal_0(k32_sin_cos)), inference(rw,[status(thm)],[c_0_291, c_0_292])).
cnf(c_0_329, plain, (v1_xreal_0(k6_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_293, c_0_294]), c_0_295])])])).
cnf(c_0_330, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_296])])).
cnf(c_0_331, plain, (~r1_xxreal_0(k32_sin_cos,k6_numbers)), inference(evalgc,[status(thm)],[c_0_297])).
cnf(c_0_332, plain, (r1_tarski(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_relset_1(k1_numbers,k2_fdiff_9))), inference(split_conjunct,[status(thm)],[t4_sincos10])).
cnf(c_0_333, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k1_real_1(X2),k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_146])])).
cnf(c_0_334, plain, (m1_subset_1(k10_real_1(k32_sin_cos,np__4),k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_299])])).
cnf(c_0_335, plain, (~r1_xxreal_0(esk16_3(X1,X2,X3),X2)|~r2_hidden(X1,a_2_1_rcomp_1(X2,X3))|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)), inference(split_conjunct,[status(thm)],[c_0_261])).
cnf(c_0_336, plain, (X1=esk16_3(X1,X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,a_2_1_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[c_0_300])).
cnf(c_0_337, plain, (a_2_1_rcomp_1(X1,X2)=k2_rcomp_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_301])).
cnf(c_0_338, plain, (m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k2_relset_1(k1_numbers,k3_sincos10))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_302])])).
cnf(c_0_339, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__2),k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))), inference(split_conjunct,[status(thm)],[l20_sincos10])).
fof(c_0_340, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))&~r2_hidden(k1_seq_1(k2_fdiff_9,esk1_0),k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_303])])])).
cnf(c_0_341, plain, (r2_hidden(X2,k1_xxreal_1(X1,X3))|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_304])).
fof(c_0_342, plain, ![X90, X91]:(~v1_relat_1(X90)|(~v1_funct_1(X90)|~v3_valued_0(X90))|m1_subset_1(k1_seq_1(X90,X91),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_seq_1])])).
cnf(c_0_343, plain, (r1_xxreal_0(X1,X2)|~r1_tarski(X3,k1_relset_1(k1_numbers,k2_fdiff_9))|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,X3)|~r2_hidden(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_306]), c_0_183]), c_0_219])])])).
cnf(c_0_344, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k2_fdiff_9))|~m1_subset_1(X1,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_307, c_0_153])])).
cnf(c_0_345, plain, (r1_xxreal_0(X1,k1_seq_1(k2_fdiff_9,X2))|r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_308, c_0_309])])).
cnf(c_0_346, plain, (r1_tarski(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))),k1_relset_1(k1_numbers,k2_fdiff_9))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_307, c_0_310])])).
cnf(c_0_347, plain, (r1_xxreal_0(X1,esk14_3(X2,X1,X3))|~v1_xxreal_0(X3)|~v1_xreal_0(X1)|~r2_hidden(X2,k3_rcomp_1(X1,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_312])])).
cnf(c_0_348, plain, (esk14_3(X1,X2,X3)=X1|~v1_xxreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k3_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_313, c_0_312])])).
cnf(c_0_349, plain, (r1_xxreal_0(esk15_3(X1,X2,X3),X3)|~r2_hidden(X1,a_2_0_rcomp_1(X2,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3)), inference(split_conjunct,[status(thm)],[c_0_143])).
cnf(c_0_350, plain, (esk15_3(X1,X2,X3)=X2|~v1_xxreal_0(esk15_3(X1,X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(esk15_3(X1,X2,X3),X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3))), inference(csr,[status(thm)],[c_0_314, c_0_221])).
cnf(c_0_351, plain, (esk15_3(X1,X2,X3)=X1|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_232])])).
cnf(c_0_352, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X3)|~r2_hidden(X1,k2_tarski(X2,X3))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_316, c_0_221]), c_0_221])).
cnf(c_0_353, plain, (r2_hidden(X1,k2_tarski(X2,X1))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_317])])).
cnf(c_0_354, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~r1_xxreal_0(X2,k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_319])])).
cnf(c_0_355, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_320])).
cnf(c_0_356, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_321, c_0_322])])).
cnf(c_0_357, plain, (v1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_323, c_0_324])])).
cnf(c_0_358, plain, (m1_subset_1(k1_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_325])).
cnf(c_0_359, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),k1_rcomp_1(k6_numbers,k32_sin_cos))|~r1_xxreal_0(k6_numbers,k32_sin_cos)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_326, c_0_327]), c_0_328]), c_0_329])])])).
cnf(c_0_360, plain, (r1_xxreal_0(k6_numbers,k32_sin_cos)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_330, c_0_328]), c_0_331])])).
cnf(c_0_361, plain, (r1_tarski(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_relset_1(k1_numbers,k2_fdiff_9))), inference(evalgc,[status(thm)],[c_0_332])).
cnf(c_0_362, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_real_1(X2),k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_333, c_0_190])).
cnf(c_0_363, plain, (v1_xxreal_0(k10_real_1(k32_sin_cos,np__4))|~v2_membered(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_334])])).
cnf(c_0_364, plain, (~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,a_2_1_rcomp_1(X2,X3))|~r1_xxreal_0(esk16_3(X1,X2,X3),X2)), inference(evalgc,[status(thm)],[c_0_335])).
cnf(c_0_365, plain, (esk16_3(X1,X2,X3)=X1|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k2_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_336, c_0_337])])).
cnf(c_0_366, plain, (v1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~v2_membered(k2_relset_1(k1_numbers,k3_sincos10))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_338])])).
cnf(c_0_367, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__2),k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[c_0_339])).
fof(c_0_368, plain, ![X139, X140]:(~v3_membered(X139)|(~m1_subset_1(X140,k1_zfmisc_1(X139))|v3_membered(X140))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc21_membered])])])).
cnf(c_0_369, negated_conjecture, (~r2_hidden(k1_seq_1(k2_fdiff_9,esk1_0),k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1)))), inference(split_conjunct,[status(thm)],[c_0_340])).
cnf(c_0_370, plain, (k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))=k7_square_1(np__2)), inference(split_conjunct,[status(thm)],[t32_sincos10])).
cnf(c_0_371, plain, (r2_hidden(X2,k1_xxreal_1(X1,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_341])).
cnf(c_0_372, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_342])).
cnf(c_0_373, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,X3)|~r2_hidden(X2,X3)), inference(csr,[status(thm)],[c_0_343, c_0_344])).
cnf(c_0_374, plain, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_seq_1(k2_fdiff_9,X2))|r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_345, c_0_309])])).
cnf(c_0_375, plain, (r1_tarski(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_346]), c_0_244]), c_0_226])])])).
cnf(c_0_376, plain, (k1_seq_1(k2_fdiff_9,k1_real_1(k10_real_1(k32_sin_cos,np__4)))=k1_real_1(k7_square_1(np__2))), inference(split_conjunct,[status(thm)],[t32_sincos10])).
cnf(c_0_377, plain, (v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v3_membered(k2_relset_1(k1_numbers,k3_sincos10))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_293, c_0_192])])).
cnf(c_0_378, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k2_relset_1(k1_numbers,k3_sincos10))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_324])])).
cnf(c_0_379, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xreal_0(X1)|~r2_hidden(X2,k3_rcomp_1(X1,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_347, c_0_348])])).
cnf(c_0_380, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_188])).
cnf(c_0_381, plain, (r1_xxreal_0(esk15_3(X1,X2,X3),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,a_2_0_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[c_0_349])).
cnf(c_0_382, plain, (X1=X2|~v1_xxreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_350, c_0_351])])).
cnf(c_0_383, plain, (r2_hidden(X1,k1_rcomp_1(X2,X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_352, c_0_353])])).
cnf(c_0_384, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~r1_xxreal_0(X2,k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_354, c_0_355])).
cnf(c_0_385, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_357])])).
cnf(c_0_386, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_358])])).
cnf(c_0_387, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),k1_rcomp_1(k6_numbers,k32_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_359, c_0_360])])])).
cnf(c_0_388, plain, (v1_xreal_0(k10_real_1(k32_sin_cos,np__4))|~v3_membered(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_293, c_0_334])])).
cnf(c_0_389, plain, (r1_tarski(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_361]), c_0_244]), c_0_226])])])).
cnf(c_0_390, plain, (r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(k1_real_1(X2),X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_362, c_0_319])])).
cnf(c_0_391, plain, (v1_xxreal_0(k10_real_1(k32_sin_cos,np__4))|~v3_membered(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_363, c_0_239])])).
cnf(c_0_392, plain, (k1_real_1(X1)=k1_real_1(X2)|~v1_xxreal_0(k1_real_1(X2))|~v1_xxreal_0(k1_real_1(X1))|~r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_274, c_0_318])])).
fof(c_0_393, plain, ![X216, X217]:(~r2_hidden(X216,X217)|~v1_xboole_0(X217)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
cnf(c_0_394, plain, (~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(X3,X2)|~r2_hidden(X3,a_2_1_rcomp_1(X2,X1))|~r2_hidden(X3,k2_rcomp_1(X2,X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_364, c_0_365])])).
cnf(c_0_395, plain, (k2_relset_1(k1_numbers,k1_sin_cos9)=k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2))), inference(split_conjunct,[status(thm)],[t11_sin_cos9])).
cnf(c_0_396, plain, (v1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~v3_membered(k2_relset_1(k1_numbers,k3_sincos10))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_239])])).
cnf(c_0_397, plain, (m1_subset_1(k10_real_1(k32_sin_cos,np__2),k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_367])])).
cnf(c_0_398, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_368])).
cnf(c_0_399, negated_conjecture, (~r2_hidden(k1_seq_1(k2_fdiff_9,esk1_0),k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1)))), inference(evalgc,[status(thm)],[c_0_369])).
cnf(c_0_400, plain, (k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))=k7_square_1(np__2)), inference(evalgc,[status(thm)],[c_0_370])).
cnf(c_0_401, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_371, c_0_246])])).
cnf(c_0_402, plain, (k1_seq_1(k2_fdiff_9,k1_real_1(k10_real_1(k32_sin_cos,np__2)))=k1_real_1(np__1)), inference(split_conjunct,[status(thm)],[t32_sincos10])).
cnf(c_0_403, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_372])).
cnf(c_0_404, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_seq_1(k2_fdiff_9,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,X3)|~r2_hidden(X2,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_374])])).
cnf(c_0_405, plain, (m1_subset_1(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))),k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_310])])).
cnf(c_0_406, plain, (m1_subset_1(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))),k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_375])])).
cnf(c_0_407, plain, (k1_seq_1(k2_fdiff_9,k1_real_1(k10_real_1(k32_sin_cos,np__4)))=k1_real_1(k7_square_1(np__2))), inference(evalgc,[status(thm)],[c_0_376])).
cnf(c_0_408, plain, (v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(k2_relset_1(k1_numbers,k3_sincos10),k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_285])])).
cnf(c_0_409, plain, (m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_378, c_0_302])])).
cnf(c_0_410, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r2_hidden(X1,k2_relset_1(k1_numbers,k3_sincos10))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_379, c_0_116]), c_0_296])])])).
cnf(c_0_411, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(evalgc,[status(thm)],[c_0_380])).
cnf(c_0_412, plain, (r1_xxreal_0(esk15_3(X1,X2,X3),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_381, c_0_232])])).
cnf(c_0_413, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))), inference(split_conjunct,[status(thm)],[c_0_340])).
cnf(c_0_414, plain, (X1=X2|~v1_xxreal_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_382, c_0_383])])).
cnf(c_0_415, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__4))))|~m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_numbers)|~m1_subset_1(k10_real_1(k32_sin_cos,np__4),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_384, c_0_385])])).
cnf(c_0_416, plain, (m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_378, c_0_150])])).
cnf(c_0_417, plain, (m1_subset_1(k10_real_1(k32_sin_cos,np__4),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_386, c_0_387]), c_0_328]), c_0_329])])])).
cnf(c_0_418, plain, (v1_xreal_0(k10_real_1(k32_sin_cos,np__4))|~m1_subset_1(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_388, c_0_285])])).
cnf(c_0_419, plain, (m1_subset_1(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_389])])).
cnf(c_0_420, plain, (r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(k1_real_1(X2),X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_390, c_0_355])).
cnf(c_0_421, plain, (v1_xxreal_0(k10_real_1(k32_sin_cos,np__4))|~m1_subset_1(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_285])])).
cnf(c_0_422, plain, (k1_real_1(X1)=k1_real_1(X2)|~v1_xxreal_0(k1_real_1(X2))|~v1_xxreal_0(k1_real_1(X1))|~r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_392, c_0_318])])).
fof(c_0_423, plain, ![X86, X87]:(~m1_subset_1(X86,X87)|(v1_xboole_0(X87)|r2_hidden(X86,X87))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_424, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_393])).
cnf(c_0_425, plain, (~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(X3,X2)|~r2_hidden(X3,k2_rcomp_1(X2,X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_394, c_0_337])])).
cnf(c_0_426, plain, (k2_relset_1(k1_numbers,k1_sin_cos9)=k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2))), inference(evalgc,[status(thm)],[c_0_395])).
cnf(c_0_427, plain, (v1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~m1_subset_1(k2_relset_1(k1_numbers,k3_sincos10),k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_285])])).
cnf(c_0_428, plain, (v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v2_membered(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_397])])).
cnf(c_0_429, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_398])).
cnf(c_0_430, negated_conjecture, (~r2_hidden(k1_seq_1(k2_fdiff_9,esk1_0),k1_rcomp_1(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),k1_real_1(np__1)))), inference(rw,[status(thm)],[c_0_399, c_0_400])).
cnf(c_0_431, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~v1_xxreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X2,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_401, c_0_221]), c_0_221])).
cnf(c_0_432, plain, (k1_seq_1(k2_fdiff_9,k1_real_1(k10_real_1(k32_sin_cos,np__2)))=k1_real_1(np__1)), inference(evalgc,[status(thm)],[c_0_402])).
cnf(c_0_433, plain, (m1_subset_1(k1_seq_1(k2_fdiff_9,X1),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_403, c_0_269]), c_0_226]), c_0_219])])])).
cnf(c_0_434, plain, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_seq_1(k2_fdiff_9,X2))|r1_xxreal_0(X1,X2)|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_404, c_0_405]), c_0_406])])])).
cnf(c_0_435, plain, (k1_seq_1(k2_fdiff_9,k1_real_1(k10_real_1(k32_sin_cos,np__4)))=k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)))), inference(rw,[status(thm)],[c_0_407, c_0_400])).
cnf(c_0_436, plain, (v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_408, c_0_324])])])).
cnf(c_0_437, plain, (v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_293, c_0_409]), c_0_295])])])).
cnf(c_0_438, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_410, c_0_150])])).
cnf(c_0_439, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X1,X3)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_411])])).
cnf(c_0_440, plain, (r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_412, c_0_351])])).
cnf(c_0_441, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))), inference(evalgc,[status(thm)],[c_0_413])).
cnf(c_0_442, plain, (X1=X2|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)), inference(csr,[status(thm)],[c_0_414, c_0_221])).
cnf(c_0_443, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__4))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_415, c_0_416]), c_0_417])])).
cnf(c_0_444, plain, (v1_xreal_0(k10_real_1(k32_sin_cos,np__4))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_418, c_0_419])])])).
cnf(c_0_445, plain, (r1_xxreal_0(k1_real_1(k1_real_1(X1)),X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_420, c_0_318])])).
cnf(c_0_446, plain, (v1_xxreal_0(k10_real_1(k32_sin_cos,np__4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_421, c_0_419])])).
cnf(c_0_447, plain, (k1_real_1(X1)=k1_real_1(X2)|~v1_xxreal_0(k1_real_1(X1))|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(k1_real_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_422, c_0_221])])).
cnf(c_0_448, plain, (v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_293, c_0_355]), c_0_295])])])).
cnf(c_0_449, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r2_hidden(X2,k1_rcomp_1(X1,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_351])])).
cnf(c_0_450, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_423])).
cnf(c_0_451, plain, (m1_subset_1(X1,k2_relset_1(k1_numbers,k3_sincos10))|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_405])])).
cnf(c_0_452, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_424])).
cnf(c_0_453, plain, (~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r2_hidden(X1,k2_relset_1(k1_numbers,k1_sin_cos9))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_425, c_0_426])])).
cnf(c_0_454, plain, (v1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_427, c_0_324])])).
cnf(c_0_455, plain, (v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v3_membered(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_428, c_0_239])])).
cnf(c_0_456, plain, (v3_membered(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_429, c_0_419]), c_0_295])])])).
cnf(c_0_457, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)))), inference(split_conjunct,[status(thm)],[l24_sin_cos9])).
fof(c_0_458, plain, ![X384]:((r2_hidden(k1_seq_1(k16_sin_cos,X384),k1_rcomp_1(k1_real_1(np__1),np__1))|~v1_xreal_0(X384))&(r2_hidden(k1_seq_1(k19_sin_cos,X384),k1_rcomp_1(k1_real_1(np__1),np__1))|~v1_xreal_0(X384))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t27_comptrig])])])).
cnf(c_0_459, negated_conjecture, (~v1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0))|~v1_xreal_0(k1_real_1(np__1))|~v1_xreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))))|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_real_1(np__1))|~r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),k1_seq_1(k2_fdiff_9,esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_430, c_0_431])])).
cnf(c_0_460, plain, (v1_xreal_0(k1_real_1(np__1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_432])])).
cnf(c_0_461, plain, (v1_xxreal_0(k1_seq_1(k2_fdiff_9,X1))|~v2_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_433])])).
cnf(c_0_462, plain, (r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),k1_seq_1(k2_fdiff_9,X1))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1)|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_434, c_0_383]), c_0_435]), c_0_436]), c_0_437])])])).
cnf(c_0_463, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_438, c_0_437])])])).
cnf(c_0_464, plain, (r2_hidden(X1,k2_tarski(X1,X2))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_439])])).
cnf(c_0_465, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_440, c_0_441])])).
cnf(c_0_466, plain, (k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)))=k10_real_1(k32_sin_cos,np__4)|~v1_xreal_0(k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__4))))|~r1_xxreal_0(k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__4))),k10_real_1(k32_sin_cos,np__4))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_442, c_0_443]), c_0_444])])])).
cnf(c_0_467, plain, (r1_xxreal_0(k1_real_1(k1_real_1(X1)),X2)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_445, c_0_355])).
cnf(c_0_468, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k10_real_1(k32_sin_cos,np__4))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_446])])).
cnf(c_0_469, plain, (k1_real_1(X1)=k1_real_1(X2)|~v1_xxreal_0(k1_real_1(X1))|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_447, c_0_448])).
cnf(c_0_470, negated_conjecture, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),esk1_0)|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_449, c_0_441])])).
cnf(c_0_471, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_450])).
cnf(c_0_472, negated_conjecture, (m1_subset_1(esk1_0,k2_relset_1(k1_numbers,k3_sincos10))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_451, c_0_441])])).
cnf(c_0_473, plain, (~v1_xboole_0(k2_relset_1(k1_numbers,k3_sincos10))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_452, c_0_302])])).
cnf(c_0_474, plain, (~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r2_hidden(X1,k2_relset_1(k1_numbers,k1_sin_cos9))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_453, c_0_454])])])).
cnf(c_0_475, plain, (v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_455, c_0_456])])).
cnf(c_0_476, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[c_0_457])).
cnf(c_0_477, plain, (r2_hidden(k1_seq_1(k19_sin_cos,X1),k1_rcomp_1(k1_real_1(np__1),np__1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_458])).
cnf(c_0_478, plain, (k1_seq_1(k19_sin_cos,k32_sin_cos)=k1_real_1(np__1)), inference(split_conjunct,[status(thm)],[t76_sin_cos])).
cnf(c_0_479, plain, (m1_subset_1(np__1,k1_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_480, negated_conjecture, (~v1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0))|~v1_xreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))))|~r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),k1_seq_1(k2_fdiff_9,esk1_0))|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_real_1(np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_459, c_0_460])])])).
cnf(c_0_481, plain, (v1_xxreal_0(k1_seq_1(k2_fdiff_9,X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_461, c_0_239]), c_0_295])])])).
cnf(c_0_482, plain, (v1_xreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_435])])).
cnf(c_0_483, plain, (r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),k1_seq_1(k2_fdiff_9,X1))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1)|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_462, c_0_463])])).
cnf(c_0_484, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_seq_1(k2_fdiff_9,X1))|r1_xxreal_0(esk1_0,X1)|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_434, c_0_441])])).
cnf(c_0_485, plain, (r2_hidden(X1,k1_rcomp_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_352, c_0_464])])).
cnf(c_0_486, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_465, c_0_190])])).
cnf(c_0_487, plain, (r2_hidden(k1_seq_1(k16_sin_cos,X1),k1_rcomp_1(k1_real_1(np__1),np__1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_458])).
cnf(c_0_488, plain, (k1_seq_1(k16_sin_cos,k10_real_1(k32_sin_cos,np__2))=np__1), inference(split_conjunct,[status(thm)],[t76_sin_cos])).
cnf(c_0_489, plain, (v1_xreal_0(k10_real_1(k32_sin_cos,np__2))|~v3_membered(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_293, c_0_397])])).
cnf(c_0_490, plain, (k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)))=k10_real_1(k32_sin_cos,np__4)|~v1_xreal_0(k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__4))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_466, c_0_467]), c_0_468]), c_0_417])])])).
cnf(c_0_491, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(k1_real_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_469, c_0_221])])).
cnf(c_0_492, negated_conjecture, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),esk1_0)|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_470, c_0_436])])).
cnf(c_0_493, negated_conjecture, (r2_hidden(esk1_0,k2_relset_1(k1_numbers,k3_sincos10))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_471, c_0_472]), c_0_473])])).
cnf(c_0_494, plain, (~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r2_hidden(X1,k2_relset_1(k1_numbers,k1_sin_cos9))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_474, c_0_475])])])).
cnf(c_0_495, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),k2_relset_1(k1_numbers,k1_sin_cos9))), inference(rw,[status(thm)],[c_0_476, c_0_426])).
cnf(c_0_496, plain, (r2_hidden(k1_seq_1(k19_sin_cos,X1),k1_rcomp_1(k1_real_1(np__1),np__1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_477])).
cnf(c_0_497, plain, (k1_seq_1(k19_sin_cos,k32_sin_cos)=k1_real_1(np__1)), inference(evalgc,[status(thm)],[c_0_478])).
cnf(c_0_498, plain, (m1_subset_1(np__1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_479])).
cnf(c_0_499, negated_conjecture, (~r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),k1_seq_1(k2_fdiff_9,esk1_0))|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_real_1(np__1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_480, c_0_481]), c_0_482])])).
cnf(c_0_500, negated_conjecture, (r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),k1_seq_1(k2_fdiff_9,esk1_0))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_483, c_0_441])])).
cnf(c_0_501, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_real_1(np__1))|r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_484, c_0_485]), c_0_432]), c_0_436]), c_0_437])])])).
cnf(c_0_502, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_numbers)|~m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_486, c_0_190])])).
fof(c_0_503, plain, ![X468]:((r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k5_sin_cos9(X468))|(~r1_xxreal_0(k1_real_1(np__1),X468)|~r1_xxreal_0(X468,np__1))|~m1_subset_1(X468,k1_numbers))&(r1_xxreal_0(k5_sin_cos9(X468),k10_real_1(k32_sin_cos,np__4))|(~r1_xxreal_0(k1_real_1(np__1),X468)|~r1_xxreal_0(X468,np__1))|~m1_subset_1(X468,k1_numbers))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_sin_cos9])])])).
cnf(c_0_504, plain, (r2_hidden(k1_seq_1(k16_sin_cos,X1),k1_rcomp_1(k1_real_1(np__1),np__1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_487])).
cnf(c_0_505, plain, (k1_seq_1(k16_sin_cos,k10_real_1(k32_sin_cos,np__2))=np__1), inference(evalgc,[status(thm)],[c_0_488])).
cnf(c_0_506, plain, (v1_xreal_0(k10_real_1(k32_sin_cos,np__2))|~m1_subset_1(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_489, c_0_285])])).
cnf(c_0_507, plain, (k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)))=k10_real_1(k32_sin_cos,np__4)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_490, c_0_448]), c_0_416])])])).
cnf(c_0_508, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_491, c_0_448])).
cnf(c_0_509, negated_conjecture, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_492, c_0_437])])])).
cnf(c_0_510, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_378, c_0_493])])).
cnf(c_0_511, plain, (r1_xxreal_0(X1,k10_real_1(k32_sin_cos,np__4))|r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),X1)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_446])])).
cnf(c_0_512, plain, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_494, c_0_495])])).
cnf(c_0_513, plain, (r2_hidden(k1_real_1(np__1),k1_rcomp_1(k1_real_1(np__1),np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_496, c_0_497]), c_0_328])])])).
cnf(c_0_514, plain, (v1_xreal_0(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_293, c_0_498]), c_0_295])])])).
cnf(c_0_515, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1))|~m1_subset_1(k2_relset_1(k1_numbers,k3_sincos10),k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,k2_relset_1(k1_numbers,k3_sincos10))|~r2_hidden(X2,k2_relset_1(k1_numbers,k3_sincos10))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_175]), c_0_243]), c_0_218]), c_0_183]), c_0_219])])])).
cnf(c_0_516, plain, (r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)))), inference(split_conjunct,[status(thm)],[l23_sin_cos9])).
cnf(c_0_517, negated_conjecture, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),esk1_0)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_real_1(np__1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_499, c_0_500])])).
cnf(c_0_518, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_real_1(np__1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_501, c_0_463])])).
cnf(c_0_519, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_502, c_0_416])])])).
fof(c_0_520, plain, ![X440, X442, X443]:(((m1_subset_1(esk23_1(X440),k1_numbers)|~r2_hidden(X440,a_0_2_sin_cos9))&((X440=esk23_1(X440)|~r2_hidden(X440,a_0_2_sin_cos9))&((r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),esk23_1(X440))|~r2_hidden(X440,a_0_2_sin_cos9))&(r1_xxreal_0(esk23_1(X440),k10_real_1(k32_sin_cos,np__4))|~r2_hidden(X440,a_0_2_sin_cos9)))))&(~m1_subset_1(X443,k1_numbers)|(X442!=X443|(~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X443)|~r1_xxreal_0(X443,k10_real_1(k32_sin_cos,np__4))))|r2_hidden(X442,a_0_2_sin_cos9))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[fraenkel_a_0_2_sin_cos9])])])])])])).
cnf(c_0_521, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k5_sin_cos9(X1))|~r1_xxreal_0(k1_real_1(np__1),X1)|~r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_503])).
cnf(c_0_522, plain, (k5_sin_cos9(np__1)=k10_real_1(k32_sin_cos,np__4)), inference(split_conjunct,[status(thm)],[t39_sin_cos9])).
cnf(c_0_523, plain, (r1_xxreal_0(np__1,np__1)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r1_r1])).
cnf(c_0_524, plain, (r2_hidden(np__1,k1_rcomp_1(k1_real_1(np__1),np__1))|~v1_xreal_0(k10_real_1(k32_sin_cos,np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_504, c_0_505])])).
cnf(c_0_525, plain, (v1_xreal_0(k10_real_1(k32_sin_cos,np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_506, c_0_419])])])).
cnf(c_0_526, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_real_1(X1))|~r1_xxreal_0(X1,k10_real_1(k32_sin_cos,np__4))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_384, c_0_507]), c_0_416])])])).
cnf(c_0_527, negated_conjecture, (k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)))=k1_real_1(esk1_0)|~r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_508, c_0_509]), c_0_409]), c_0_510])])])).
cnf(c_0_528, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__4))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_511, c_0_454]), c_0_512])])).
cnf(c_0_529, plain, (m1_subset_1(k1_real_1(np__1),k1_numbers)|~v1_xreal_0(k1_real_1(np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_386, c_0_513]), c_0_514])])])).
cnf(c_0_530, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1))|~r2_hidden(X1,k2_relset_1(k1_numbers,k3_sincos10))|~r2_hidden(X2,k2_relset_1(k1_numbers,k3_sincos10))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_515, c_0_324])])).
cnf(c_0_531, plain, (r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[c_0_516])).
cnf(c_0_532, negated_conjecture, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),esk1_0)|r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_518])])).
cnf(c_0_533, negated_conjecture, (v1_xreal_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_293, c_0_510]), c_0_295])])])).
cnf(c_0_534, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__4)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_519, c_0_409])])).
cnf(c_0_535, plain, (r2_hidden(X2,a_0_2_sin_cos9)|~m1_subset_1(X1,k1_numbers)|X2!=X1|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1)|~r1_xxreal_0(X1,k10_real_1(k32_sin_cos,np__4))), inference(split_conjunct,[status(thm)],[c_0_520])).
cnf(c_0_536, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k5_sin_cos9(X1))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(k1_real_1(np__1),X1)), inference(evalgc,[status(thm)],[c_0_521])).
cnf(c_0_537, plain, (k5_sin_cos9(np__1)=k10_real_1(k32_sin_cos,np__4)), inference(evalgc,[status(thm)],[c_0_522])).
cnf(c_0_538, plain, (r1_xxreal_0(np__1,np__1)), inference(evalgc,[status(thm)],[c_0_523])).
cnf(c_0_539, plain, (r2_hidden(np__1,k1_rcomp_1(k1_real_1(np__1),np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_524, c_0_525])])])).
cnf(c_0_540, negated_conjecture, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_real_1(esk1_0))|~r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_526, c_0_527]), c_0_528]), c_0_409])])])).
cnf(c_0_541, plain, (m1_subset_1(k1_real_1(np__1),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_529, c_0_448]), c_0_498])])])).
cnf(c_0_542, plain, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k6_numbers)), inference(split_conjunct,[status(thm)],[c_0_258])).
cnf(c_0_543, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_seq_1(k2_fdiff_9,X2))|~r2_hidden(X1,k2_relset_1(k1_numbers,k3_sincos10))|~r2_hidden(X2,k2_relset_1(k1_numbers,k3_sincos10))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_530, c_0_374])])).
cnf(c_0_544, plain, (r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k2_relset_1(k1_numbers,k1_sin_cos9))), inference(rw,[status(thm)],[c_0_531, c_0_426])).
cnf(c_0_545, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__4))=esk1_0|r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_442, c_0_532]), c_0_436]), c_0_533]), c_0_534])])])).
cnf(c_0_546, plain, (r2_hidden(X2,a_0_2_sin_cos9)|X1!=X2|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,k10_real_1(k32_sin_cos,np__4))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1)), inference(evalgc,[status(thm)],[c_0_535])).
cnf(c_0_547, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k10_real_1(k32_sin_cos,np__4))|~r1_xxreal_0(k1_real_1(np__1),np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_536, c_0_537]), c_0_538]), c_0_498])])])).
cnf(c_0_548, plain, (r1_xxreal_0(k1_real_1(np__1),np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_449, c_0_539]), c_0_514]), c_0_460])])])).
cnf(c_0_549, negated_conjecture, (r1_xxreal_0(esk1_0,k10_real_1(k32_sin_cos,np__4))|~r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_384, c_0_540]), c_0_507]), c_0_416]), c_0_510])])])).
cnf(c_0_550, negated_conjecture, (k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__4)|r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_508, c_0_532]), c_0_507]), c_0_534]), c_0_416]), c_0_510])])])).
cnf(c_0_551, plain, (v1_xxreal_0(k1_real_1(np__1))|~v2_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_541])])).
cnf(c_0_552, plain, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k6_numbers)), inference(evalgc,[status(thm)],[c_0_542])).
cnf(c_0_553, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),esk23_1(X1))|~r2_hidden(X1,a_0_2_sin_cos9)), inference(split_conjunct,[status(thm)],[c_0_520])).
cnf(c_0_554, plain, (X1=esk23_1(X1)|~r2_hidden(X1,a_0_2_sin_cos9)), inference(split_conjunct,[status(thm)],[c_0_520])).
cnf(c_0_555, plain, (r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),k1_seq_1(k2_fdiff_9,X1))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1)|~r2_hidden(X1,k2_relset_1(k1_numbers,k3_sincos10))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_543, c_0_150]), c_0_435])])).
cnf(c_0_556, plain, (~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_real_1(k10_real_1(k32_sin_cos,np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_494, c_0_544])])).
cnf(c_0_557, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0|k1_real_1(k10_real_1(k32_sin_cos,np__4))=esk1_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_442, c_0_545]), c_0_533]), c_0_437]), c_0_509])])])).
cnf(c_0_558, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_221])])).
cnf(c_0_559, plain, (r2_hidden(X1,a_0_2_sin_cos9)|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1)|~r1_xxreal_0(X1,k10_real_1(k32_sin_cos,np__4))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_546])])).
cnf(c_0_560, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k10_real_1(k32_sin_cos,np__4))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_547, c_0_548])])])).
cnf(c_0_561, negated_conjecture, (r1_xxreal_0(esk1_0,k10_real_1(k32_sin_cos,np__4))|k1_real_1(k10_real_1(k32_sin_cos,np__4))=esk1_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_549, c_0_545])])).
cnf(c_0_562, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0|k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__4)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_442, c_0_550]), c_0_533]), c_0_437]), c_0_509])])])).
cnf(c_0_563, plain, (v1_xxreal_0(k1_real_1(np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_551, c_0_239]), c_0_295])])])).
cnf(c_0_564, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_386, c_0_383])])).
cnf(c_0_565, plain, (r1_xxreal_0(k6_numbers,k10_real_1(k32_sin_cos,np__2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_330, c_0_525]), c_0_552])])).
cnf(c_0_566, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),esk23_1(X1))|~r2_hidden(X1,a_0_2_sin_cos9)), inference(evalgc,[status(thm)],[c_0_553])).
cnf(c_0_567, plain, (X1=esk23_1(X1)|~r2_hidden(X1,a_0_2_sin_cos9)), inference(evalgc,[status(thm)],[c_0_554])).
cnf(c_0_568, plain, (r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),k1_real_1(np__1))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_555, c_0_302]), c_0_432]), c_0_556])])).
cnf(c_0_569, negated_conjecture, (k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)))=k1_seq_1(k2_fdiff_9,esk1_0)|k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_435, c_0_557])])).
cnf(c_0_570, plain, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_seq_1(k2_fdiff_9,X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_558, c_0_309])])).
cnf(c_0_571, negated_conjecture, (r2_hidden(esk1_0,a_0_2_sin_cos9)|r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r1_xxreal_0(esk1_0,k10_real_1(k32_sin_cos,np__4))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_559, c_0_532]), c_0_510])])])).
cnf(c_0_572, negated_conjecture, (r1_xxreal_0(esk1_0,k10_real_1(k32_sin_cos,np__4))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_560, c_0_561])])).
cnf(c_0_573, negated_conjecture, (k1_seq_1(k2_fdiff_9,esk1_0)=k1_real_1(np__1)|k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_432, c_0_562])])).
cnf(c_0_574, plain, (r1_xxreal_0(k1_real_1(np__1),k1_real_1(np__1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_563])])).
cnf(c_0_575, plain, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k10_real_1(k32_sin_cos,np__4))|~m1_subset_1(k10_real_1(k32_sin_cos,np__2),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_556, c_0_318]), c_0_417])])])).
cnf(c_0_576, plain, (m1_subset_1(k10_real_1(k32_sin_cos,np__2),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_564, c_0_565]), c_0_525]), c_0_329])])])).
cnf(c_0_577, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1)|~r2_hidden(X1,a_0_2_sin_cos9)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_566, c_0_567])])).
cnf(c_0_578, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_real_1(np__1))|k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_568, c_0_569])])).
cnf(c_0_579, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_real_1(np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_499, c_0_569]), c_0_570])])])).
cnf(c_0_580, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|r2_hidden(esk1_0,a_0_2_sin_cos9)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_571, c_0_572])])])).
cnf(c_0_581, negated_conjecture, (k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__4)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_499, c_0_573]), c_0_568]), c_0_574])])])).
cnf(c_0_582, plain, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k10_real_1(k32_sin_cos,np__4))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_575, c_0_576])])).
cnf(c_0_583, plain, (~r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__2)),a_0_2_sin_cos9)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_556, c_0_577])])).
cnf(c_0_584, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0), inference(csr,[status(thm)],[c_0_578, c_0_579])).
cnf(c_0_585, negated_conjecture, (r2_hidden(esk1_0,a_0_2_sin_cos9)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_384, c_0_580]), c_0_581]), c_0_510]), c_0_576])]), c_0_582])])).
cnf(c_0_586, plain, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_583, c_0_584]), c_0_585])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 587
# Proof object clause steps            : 445
# Proof object formula steps           : 142
# Proof object conjectures             : 47
# Proof object clause conjectures      : 44
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 103
# Proof object initial formulas used   : 83
# Proof object generating inferences   : 177
# Proof object simplifying inferences  : 235
# Training examples: 297 positive, 1585 negative
# Training: Positive examples begin
cnf(c_0_165, plain, (k1_xboole_0=k6_numbers)).# trainpos
cnf(c_0_208, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_292, plain, (k31_sin_cos=k32_sin_cos)).# trainpos
cnf(c_0_219, plain, (v1_funct_1(k2_fdiff_9))).# trainpos
cnf(c_0_291, plain, (v1_xreal_0(k31_sin_cos))).# trainpos
cnf(c_0_256, plain, (v1_xxreal_0(esk8_0))).# trainpos
cnf(c_0_322, plain, (v1_xxreal_0(esk24_0))).# trainpos
cnf(c_0_295, plain, (v3_membered(k1_numbers))).# trainpos
cnf(c_0_211, plain, (v1_xboole_0(esk8_0))).# trainpos
cnf(c_0_253, plain, (~v1_xboole_0(k1_numbers))).# trainpos
cnf(c_0_207, plain, (~v1_xboole_0(k4_ordinal1))).# trainpos
cnf(c_0_498, plain, (m1_subset_1(np__1,k1_numbers))).# trainpos
cnf(c_0_538, plain, (r1_xxreal_0(np__1,np__1))).# trainpos
cnf(c_0_331, plain, (~r1_xxreal_0(k32_sin_cos,k6_numbers))).# trainpos
cnf(c_0_164, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_537, plain, (k5_sin_cos9(np__1)=k10_real_1(k32_sin_cos,np__4))).# trainpos
cnf(c_0_497, plain, (k1_seq_1(k19_sin_cos,k32_sin_cos)=k1_real_1(np__1))).# trainpos
cnf(c_0_252, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_251, plain, (m2_subset_1(k6_numbers,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_221, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_239, plain, (v2_membered(X1)|~v3_membered(X1))).# trainpos
cnf(c_0_99, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1))).# trainpos
cnf(c_0_130, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1))).# trainpos
cnf(c_0_505, plain, (k1_seq_1(k16_sin_cos,k10_real_1(k32_sin_cos,np__2))=np__1)).# trainpos
cnf(c_0_552, plain, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k6_numbers))).# trainpos
cnf(c_0_190, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_567, plain, (X1=esk23_1(X1)|~r2_hidden(X1,a_0_2_sin_cos9))).# trainpos
cnf(c_0_452, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_400, plain, (k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))=k7_square_1(np__2))).# trainpos
cnf(c_0_183, plain, (m1_subset_1(k2_fdiff_9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_146, plain, (k4_xcmplx_0(X1)=k1_real_1(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_319, plain, (k1_real_1(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_285, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_149, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_131, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_327, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),k2_rcomp_1(k6_numbers,k32_sin_cos)))).# trainpos
cnf(c_0_432, plain, (k1_seq_1(k2_fdiff_9,k1_real_1(k10_real_1(k32_sin_cos,np__2)))=k1_real_1(np__1))).# trainpos
cnf(c_0_355, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_286, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_153, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_407, plain, (k1_seq_1(k2_fdiff_9,k1_real_1(k10_real_1(k32_sin_cos,np__4)))=k1_real_1(k7_square_1(np__2)))).# trainpos
cnf(c_0_116, plain, (k2_relset_1(k1_numbers,k3_sincos10)=k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))).# trainpos
cnf(c_0_321, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_224, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).# trainpos
cnf(c_0_293, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_191, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_278, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2))).# trainpos
cnf(c_0_411, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4))).# trainpos
cnf(c_0_367, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__2),k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_299, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_361, plain, (r1_tarski(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_relset_1(k1_numbers,k2_fdiff_9)))).# trainpos
cnf(c_0_184, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_471, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_429, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_426, plain, (k2_relset_1(k1_numbers,k1_sin_cos9)=k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)))).# trainpos
cnf(c_0_198, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_182, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).# trainpos
cnf(c_0_399, negated_conjecture, (~r2_hidden(k1_seq_1(k2_fdiff_9,esk1_0),k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1))))).# trainpos
cnf(c_0_197, plain, (r1_tarski(k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers),k1_relset_1(k1_numbers,k2_fdiff_9)))).# trainpos
cnf(c_0_246, plain, (k1_xxreal_1(X1,X2)=k1_rcomp_1(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_232, plain, (a_2_0_rcomp_1(X1,X2)=k1_rcomp_1(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_312, plain, (a_2_2_rcomp_1(X1,X2)=k3_rcomp_1(X1,X2)|~v1_xreal_0(X1)|~v1_xxreal_0(X2))).# trainpos
cnf(c_0_122, plain, (k2_rcomp_1(X1,X2)=k4_xxreal_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_337, plain, (a_2_1_rcomp_1(X1,X2)=k2_rcomp_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_152, plain, (r1_tarski(X1,X3)|~r1_tarski(X2,X3)|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_157, plain, (r2_hidden(X1,X3)|X3!=k2_xboole_0(X2,X4)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_504, plain, (r2_hidden(k1_seq_1(k16_sin_cos,X1),k1_rcomp_1(k1_real_1(np__1),np__1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_496, plain, (r2_hidden(k1_seq_1(k19_sin_cos,X1),k1_rcomp_1(k1_real_1(np__1),np__1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_566, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),esk23_1(X1))|~r2_hidden(X1,a_0_2_sin_cos9))).# trainpos
cnf(c_0_170, plain, (k9_subset_1(X2,X3,X1)=k3_xboole_0(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_263, plain, (r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))).# trainpos
cnf(c_0_115, plain, (r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))).# trainpos
cnf(c_0_441, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).# trainpos
cnf(c_0_134, plain, (k2_funct_1(k3_sincos10)=k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))).# trainpos
cnf(c_0_174, plain, (v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))))).# trainpos
cnf(c_0_358, plain, (m1_subset_1(k1_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_265, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_169, plain, (k9_subset_1(X2,X1,X3)=k9_subset_1(X2,X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_476, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_220, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_222, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).# trainpos
cnf(c_0_196, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).# trainpos
cnf(c_0_531, plain, (r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_223, plain, (k1_funct_1(X1,X2)=k1_seq_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).# trainpos
cnf(c_0_403, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).# trainpos
cnf(c_0_274, plain, (X1=X2|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_145, plain, (r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_212, plain, (r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2)))).# trainpos
cnf(c_0_313, plain, (X1=esk14_3(X1,X2,X3)|~v1_xreal_0(X2)|~v1_xxreal_0(X3)|~r2_hidden(X1,a_2_2_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_315, plain, (X1=esk15_3(X1,X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,a_2_0_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_336, plain, (X1=esk16_3(X1,X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,a_2_1_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_270, plain, (r1_tarski(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))),k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))).# trainpos
cnf(c_0_311, plain, (r1_xxreal_0(X1,esk14_3(X2,X1,X3))|~v1_xreal_0(X1)|~v1_xxreal_0(X3)|~r2_hidden(X2,a_2_2_rcomp_1(X1,X3)))).# trainpos
cnf(c_0_231, plain, (r1_xxreal_0(X1,esk15_3(X2,X1,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r2_hidden(X2,a_2_0_rcomp_1(X1,X3)))).# trainpos
cnf(c_0_381, plain, (r1_xxreal_0(esk15_3(X1,X2,X3),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,a_2_0_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_364, plain, (~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,a_2_1_rcomp_1(X2,X3))|~r1_xxreal_0(esk16_3(X1,X2,X3),X2))).# trainpos
cnf(c_0_98, plain, (k1_xxreal_1(X1,X2)=k2_xboole_0(k4_xxreal_1(X1,X2),k2_tarski(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_250, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_536, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k5_sin_cos9(X1))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(k1_real_1(np__1),X1))).# trainpos
cnf(c_0_371, plain, (r2_hidden(X2,k1_xxreal_1(X1,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_546, plain, (r2_hidden(X2,a_0_2_sin_cos9)|X1!=X2|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,k10_real_1(k32_sin_cos,np__4))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1))).# trainpos
cnf(c_0_217, plain, (v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X1,X2))|~v1_funct_1(X1)|~r1_tarski(X2,X3)|~v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X1,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_168, plain, (r1_xxreal_0(X4,X3)|~v1_funct_1(X1)|~m1_subset_1(X4,k1_numbers)|~m1_subset_1(X3,k1_numbers)|~v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X1,X2))|~r1_xxreal_0(k1_seq_1(X1,X3),k1_seq_1(X1,X4))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_hidden(X4,k9_subset_1(k1_numbers,X2,k1_relset_1(k1_numbers,X1)))|~r2_hidden(X3,k9_subset_1(k1_numbers,X2,k1_relset_1(k1_numbers,X1))))).# trainpos
cnf(c_0_439, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X1,X3))).# trainpos
cnf(c_0_351, plain, (esk15_3(X1,X2,X3)=X1|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_199, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_365, plain, (esk16_3(X1,X2,X3)=X1|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k2_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_294, plain, (m1_subset_1(k6_numbers,k1_numbers))).# trainpos
cnf(c_0_514, plain, (v1_xreal_0(np__1))).# trainpos
cnf(c_0_275, plain, (r1_xxreal_0(X1,esk15_3(X2,X1,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r2_hidden(X2,k1_rcomp_1(X1,X3)))).# trainpos
cnf(c_0_317, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X3,X1))).# trainpos
cnf(c_0_158, plain, (k2_xboole_0(k2_tarski(X1,X2),k2_rcomp_1(X1,X2))=k1_xxreal_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_267, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_268, plain, (v1_xreal_0(k1_seq_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).# trainpos
cnf(c_0_412, plain, (r1_xxreal_0(esk15_3(X1,X2,X3),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_401, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_577, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1)|~r2_hidden(X1,a_0_2_sin_cos9))).# trainpos
cnf(c_0_192, plain, (m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k2_relset_1(k1_numbers,k3_sincos10)))).# trainpos
cnf(c_0_226, plain, (v1_relat_1(k2_fdiff_9))).# trainpos
cnf(c_0_189, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_338, plain, (m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k2_relset_1(k1_numbers,k3_sincos10)))).# trainpos
cnf(c_0_356, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_225, plain, (v5_relat_1(k2_fdiff_9,k1_numbers))).# trainpos
cnf(c_0_559, plain, (r2_hidden(X1,a_0_2_sin_cos9)|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1)|~r1_xxreal_0(X1,k10_real_1(k32_sin_cos,np__4))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_348, plain, (esk14_3(X1,X2,X3)=X1|~v1_xxreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k3_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_171, plain, (k3_xboole_0(X2,X1)=X1|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_257, plain, (esk8_0=k6_numbers)).# trainpos
cnf(c_0_397, plain, (m1_subset_1(k10_real_1(k32_sin_cos,np__2),k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_266, plain, (v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))|~r1_tarski(X1,k2_relset_1(k1_numbers,k3_sincos10)))).# trainpos
cnf(c_0_195, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_307, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k2_fdiff_9))|~r1_tarski(X1,k2_relset_1(k1_numbers,k3_sincos10)))).# trainpos
cnf(c_0_259, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_513, plain, (r2_hidden(k1_real_1(np__1),k1_rcomp_1(k1_real_1(np__1),np__1)))).# trainpos
cnf(c_0_448, plain, (v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_244, plain, (v4_relat_1(k2_fdiff_9,k1_numbers))).# trainpos
cnf(c_0_347, plain, (r1_xxreal_0(X1,esk14_3(X2,X1,X3))|~v1_xxreal_0(X3)|~v1_xreal_0(X1)|~r2_hidden(X2,k3_rcomp_1(X1,X3)))).# trainpos
cnf(c_0_214, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(X3,X2),k1_seq_1(X3,X1))|~v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_1(X3)|~r2_hidden(X1,k9_subset_1(k1_numbers,k1_relset_1(k1_numbers,X3),X4))|~r2_hidden(X2,k9_subset_1(k1_numbers,k1_relset_1(k1_numbers,X3),X4))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_386, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_334, plain, (m1_subset_1(k10_real_1(k32_sin_cos,np__4),k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_405, plain, (m1_subset_1(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))),k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10))))).# trainpos
cnf(c_0_473, plain, (~v1_xboole_0(k2_relset_1(k1_numbers,k3_sincos10)))).# trainpos
cnf(c_0_464, plain, (r2_hidden(X1,k2_tarski(X1,X2)))).# trainpos
cnf(c_0_296, plain, (v1_xxreal_0(k6_numbers))).# trainpos
cnf(c_0_308, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_558, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_329, plain, (v1_xreal_0(k6_numbers))).# trainpos
cnf(c_0_346, plain, (r1_tarski(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))),k1_relset_1(k1_numbers,k2_fdiff_9)))).# trainpos
cnf(c_0_353, plain, (r2_hidden(X1,k2_tarski(X2,X1)))).# trainpos
cnf(c_0_279, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_529, plain, (m1_subset_1(k1_real_1(np__1),k1_numbers)|~v1_xreal_0(k1_real_1(np__1)))).# trainpos
cnf(c_0_449, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r2_hidden(X2,k1_rcomp_1(X1,X3)))).# trainpos
cnf(c_0_242, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(X2,X3))|~v4_relat_1(X3,X2)|~v1_relat_1(X3))).# trainpos
cnf(c_0_440, plain, (r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_379, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xreal_0(X1)|~r2_hidden(X2,k3_rcomp_1(X1,X3)))).# trainpos
cnf(c_0_333, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k1_real_1(X2),k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_306, plain, (v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10))))).# trainpos
cnf(c_0_269, plain, (v3_valued_0(k2_fdiff_9))).# trainpos
cnf(c_0_238, plain, (v1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v2_membered(k2_relset_1(k1_numbers,k3_sincos10)))).# trainpos
cnf(c_0_366, plain, (v1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~v2_membered(k2_relset_1(k1_numbers,k3_sincos10)))).# trainpos
cnf(c_0_277, plain, (r2_hidden(X1,k1_xxreal_1(X2,X3))|~r2_hidden(X1,k2_tarski(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3))).# trainpos
cnf(c_0_388, plain, (v1_xreal_0(k10_real_1(k32_sin_cos,np__4))|~v3_membered(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_489, plain, (v1_xreal_0(k10_real_1(k32_sin_cos,np__2))|~v3_membered(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_245, plain, (r2_hidden(X1,k1_xxreal_1(X2,X3))|~r2_hidden(X1,k2_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3))).# trainpos
cnf(c_0_363, plain, (v1_xxreal_0(k10_real_1(k32_sin_cos,np__4))|~v2_membered(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_215, plain, (k9_subset_1(X1,X2,X3)=X3|~m1_subset_1(X3,k1_zfmisc_1(X1))|~r1_tarski(X3,X2))).# trainpos
cnf(c_0_344, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k2_fdiff_9))|~m1_subset_1(X1,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10))))).# trainpos
cnf(c_0_314, plain, (esk15_3(X1,X2,X3)=X2|~v1_xxreal_0(X2)|~v1_xxreal_0(esk15_3(X1,X2,X3))|~r1_xxreal_0(esk15_3(X1,X2,X3),X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_428, plain, (v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v2_membered(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_394, plain, (~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(X3,X2)|~r2_hidden(X3,a_2_1_rcomp_1(X2,X1))|~r2_hidden(X3,k2_rcomp_1(X2,X1)))).# trainpos
cnf(c_0_377, plain, (v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v3_membered(k2_relset_1(k1_numbers,k3_sincos10)))).# trainpos
cnf(c_0_459, negated_conjecture, (~v1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0))|~v1_xreal_0(k1_real_1(np__1))|~v1_xreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))))|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_real_1(np__1))|~r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),k1_seq_1(k2_fdiff_9,esk1_0)))).# trainpos
cnf(c_0_524, plain, (r2_hidden(np__1,k1_rcomp_1(k1_real_1(np__1),np__1))|~v1_xreal_0(k10_real_1(k32_sin_cos,np__2)))).# trainpos
cnf(c_0_547, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k10_real_1(k32_sin_cos,np__4))|~r1_xxreal_0(k1_real_1(np__1),np__1))).# trainpos
cnf(c_0_451, plain, (m1_subset_1(X1,k2_relset_1(k1_numbers,k3_sincos10))|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).# trainpos
cnf(c_0_309, plain, (v1_xreal_0(k1_seq_1(k2_fdiff_9,X1)))).# trainpos
cnf(c_0_465, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_354, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~r1_xxreal_0(X2,k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_433, plain, (m1_subset_1(k1_seq_1(k2_fdiff_9,X1),k1_numbers))).# trainpos
cnf(c_0_382, plain, (X1=X2|~v1_xxreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_541, plain, (m1_subset_1(k1_real_1(np__1),k1_numbers))).# trainpos
cnf(c_0_390, plain, (r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(k1_real_1(X2),X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_472, negated_conjecture, (m1_subset_1(esk1_0,k2_relset_1(k1_numbers,k3_sincos10)))).# trainpos
cnf(c_0_506, plain, (v1_xreal_0(k10_real_1(k32_sin_cos,np__2))|~m1_subset_1(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_264, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(X3,X2),k1_seq_1(X3,X1))|~v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_1(X3)|~r2_hidden(X1,X4)|~r2_hidden(X2,X4)|~r1_tarski(X4,k1_relset_1(k1_numbers,X3)))).# trainpos
cnf(c_0_330, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_316, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3)|~r2_hidden(X1,k2_tarski(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_289, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3)|~r2_hidden(X1,k2_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_418, plain, (v1_xreal_0(k10_real_1(k32_sin_cos,np__4))|~m1_subset_1(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_410, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r2_hidden(X1,k2_relset_1(k1_numbers,k3_sincos10)))).# trainpos
cnf(c_0_425, plain, (~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(X3,X2)|~r2_hidden(X3,k2_rcomp_1(X2,X1)))).# trainpos
cnf(c_0_284, plain, (v1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v3_membered(k2_relset_1(k1_numbers,k3_sincos10)))).# trainpos
cnf(c_0_392, plain, (k1_real_1(X1)=k1_real_1(X2)|~v1_xxreal_0(k1_real_1(X2))|~v1_xxreal_0(k1_real_1(X1))|~r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_408, plain, (v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(k2_relset_1(k1_numbers,k3_sincos10),k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_396, plain, (v1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~v3_membered(k2_relset_1(k1_numbers,k3_sincos10)))).# trainpos
cnf(c_0_375, plain, (r1_tarski(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))),k1_numbers))).# trainpos
cnf(c_0_391, plain, (v1_xxreal_0(k10_real_1(k32_sin_cos,np__4))|~v3_membered(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_287, plain, (r1_tarski(k2_relset_1(k1_numbers,k3_sincos10),k1_numbers))).# trainpos
cnf(c_0_389, plain, (r1_tarski(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_numbers))).# trainpos
cnf(c_0_493, negated_conjecture, (r2_hidden(esk1_0,k2_relset_1(k1_numbers,k3_sincos10)))).# trainpos
cnf(c_0_324, plain, (m1_subset_1(k2_relset_1(k1_numbers,k3_sincos10),k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_419, plain, (m1_subset_1(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_485, plain, (r2_hidden(X1,k1_rcomp_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_460, plain, (v1_xreal_0(k1_real_1(np__1)))).# trainpos
cnf(c_0_383, plain, (r2_hidden(X1,k1_rcomp_1(X2,X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1))).# trainpos
cnf(c_0_406, plain, (m1_subset_1(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))),k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_359, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),k1_rcomp_1(k6_numbers,k32_sin_cos))|~r1_xxreal_0(k6_numbers,k32_sin_cos))).# trainpos
cnf(c_0_486, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_numbers))).# trainpos
cnf(c_0_515, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1))|~m1_subset_1(k2_relset_1(k1_numbers,k3_sincos10),k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,k2_relset_1(k1_numbers,k3_sincos10))|~r2_hidden(X2,k2_relset_1(k1_numbers,k3_sincos10)))).# trainpos
cnf(c_0_461, plain, (v1_xxreal_0(k1_seq_1(k2_fdiff_9,X1))|~v2_membered(k1_numbers))).# trainpos
cnf(c_0_453, plain, (~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r2_hidden(X1,k2_relset_1(k1_numbers,k1_sin_cos9)))).# trainpos
cnf(c_0_323, plain, (v1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(k2_relset_1(k1_numbers,k3_sincos10),k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_421, plain, (v1_xxreal_0(k10_real_1(k32_sin_cos,np__4))|~m1_subset_1(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_422, plain, (k1_real_1(X1)=k1_real_1(X2)|~v1_xxreal_0(k1_real_1(X2))|~v1_xxreal_0(k1_real_1(X1))|~r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X2,X1))).# trainpos
cnf(c_0_445, plain, (r1_xxreal_0(k1_real_1(k1_real_1(X1)),X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_345, plain, (r1_xxreal_0(X1,k1_seq_1(k2_fdiff_9,X2))|r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_551, plain, (v1_xxreal_0(k1_real_1(np__1))|~v2_membered(k1_numbers))).# trainpos
cnf(c_0_570, plain, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_seq_1(k2_fdiff_9,X1)))).# trainpos
cnf(c_0_482, plain, (v1_xreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)))))).# trainpos
cnf(c_0_360, plain, (r1_xxreal_0(k6_numbers,k32_sin_cos))).# trainpos
cnf(c_0_438, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_427, plain, (v1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~m1_subset_1(k2_relset_1(k1_numbers,k3_sincos10),k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_343, plain, (r1_xxreal_0(X1,X2)|~r1_tarski(X3,k1_relset_1(k1_numbers,k2_fdiff_9))|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,X3)|~r2_hidden(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10))))).# trainpos
cnf(c_0_481, plain, (v1_xxreal_0(k1_seq_1(k2_fdiff_9,X1)))).# trainpos
cnf(c_0_387, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),k1_rcomp_1(k6_numbers,k32_sin_cos)))).# trainpos
cnf(c_0_374, plain, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_seq_1(k2_fdiff_9,X2))|r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1)))).# trainpos
cnf(c_0_436, plain, (v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4))))).# trainpos
cnf(c_0_525, plain, (v1_xreal_0(k10_real_1(k32_sin_cos,np__2)))).# trainpos
cnf(c_0_444, plain, (v1_xreal_0(k10_real_1(k32_sin_cos,np__4)))).# trainpos
cnf(c_0_563, plain, (v1_xxreal_0(k1_real_1(np__1)))).# trainpos
cnf(c_0_502, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_numbers)|~m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_numbers))).# trainpos
cnf(c_0_474, plain, (~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r2_hidden(X1,k2_relset_1(k1_numbers,k1_sin_cos9)))).# trainpos
cnf(c_0_414, plain, (X1=X2|~v1_xxreal_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1))).# trainpos
cnf(c_0_385, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))).# trainpos
cnf(c_0_468, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k10_real_1(k32_sin_cos,np__4)))).# trainpos
cnf(c_0_456, plain, (v3_membered(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_447, plain, (k1_real_1(X1)=k1_real_1(X2)|~v1_xxreal_0(k1_real_1(X1))|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(k1_real_1(X2)))).# trainpos
cnf(c_0_378, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k2_relset_1(k1_numbers,k3_sincos10)))).# trainpos
cnf(c_0_455, plain, (v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v3_membered(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_416, plain, (m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_numbers))).# trainpos
cnf(c_0_543, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_seq_1(k2_fdiff_9,X2))|~r2_hidden(X1,k2_relset_1(k1_numbers,k3_sincos10))|~r2_hidden(X2,k2_relset_1(k1_numbers,k3_sincos10)))).# trainpos
cnf(c_0_565, plain, (r1_xxreal_0(k6_numbers,k10_real_1(k32_sin_cos,np__2)))).# trainpos
cnf(c_0_494, plain, (~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r2_hidden(X1,k2_relset_1(k1_numbers,k1_sin_cos9)))).# trainpos
cnf(c_0_417, plain, (m1_subset_1(k10_real_1(k32_sin_cos,np__4),k1_numbers))).# trainpos
cnf(c_0_510, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers))).# trainpos
cnf(c_0_409, plain, (m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_numbers))).# trainpos
cnf(c_0_574, plain, (r1_xxreal_0(k1_real_1(np__1),k1_real_1(np__1)))).# trainpos
cnf(c_0_539, plain, (r2_hidden(np__1,k1_rcomp_1(k1_real_1(np__1),np__1)))).# trainpos
cnf(c_0_404, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_seq_1(k2_fdiff_9,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,X3)|~r2_hidden(X2,X3))).# trainpos
cnf(c_0_415, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__4))))|~m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_numbers)|~m1_subset_1(k10_real_1(k32_sin_cos,np__4),k1_numbers))).# trainpos
cnf(c_0_491, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(k1_real_1(X1)))).# trainpos
cnf(c_0_470, negated_conjecture, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),esk1_0)|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_564, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1))).# trainpos
cnf(c_0_533, negated_conjecture, (v1_xreal_0(esk1_0))).# trainpos
cnf(c_0_556, plain, (~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_real_1(k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_437, plain, (v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_434, plain, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_seq_1(k2_fdiff_9,X2))|r1_xxreal_0(X1,X2)|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).# trainpos
cnf(c_0_519, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_numbers))).# trainpos
cnf(c_0_512, plain, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_576, plain, (m1_subset_1(k10_real_1(k32_sin_cos,np__2),k1_numbers))).# trainpos
cnf(c_0_548, plain, (r1_xxreal_0(k1_real_1(np__1),np__1))).# trainpos
cnf(c_0_555, plain, (r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),k1_seq_1(k2_fdiff_9,X1))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1)|~r2_hidden(X1,k2_relset_1(k1_numbers,k3_sincos10)))).# trainpos
cnf(c_0_466, plain, (k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)))=k10_real_1(k32_sin_cos,np__4)|~v1_xreal_0(k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__4))))|~r1_xxreal_0(k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__4))),k10_real_1(k32_sin_cos,np__4)))).# trainpos
cnf(c_0_480, negated_conjecture, (~v1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0))|~v1_xreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))))|~r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),k1_seq_1(k2_fdiff_9,esk1_0))|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_real_1(np__1)))).# trainpos
cnf(c_0_583, plain, (~r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__2)),a_0_2_sin_cos9))).# trainpos
cnf(c_0_484, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_seq_1(k2_fdiff_9,X1))|r1_xxreal_0(esk1_0,X1)|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).# trainpos
cnf(c_0_575, plain, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k10_real_1(k32_sin_cos,np__4))|~m1_subset_1(k10_real_1(k32_sin_cos,np__2),k1_numbers))).# trainpos
cnf(c_0_509, negated_conjecture, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),esk1_0))).# trainpos
cnf(c_0_490, plain, (k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)))=k10_real_1(k32_sin_cos,np__4)|~v1_xreal_0(k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).# trainpos
cnf(c_0_560, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k10_real_1(k32_sin_cos,np__4)))).# trainpos
cnf(c_0_463, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))).# trainpos
cnf(c_0_568, plain, (r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),k1_real_1(np__1)))).# trainpos
cnf(c_0_462, plain, (r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),k1_seq_1(k2_fdiff_9,X1))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1)|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))).# trainpos
cnf(c_0_511, plain, (r1_xxreal_0(X1,k10_real_1(k32_sin_cos,np__4))|r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),X1)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_507, plain, (k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)))=k10_real_1(k32_sin_cos,np__4))).# trainpos
cnf(c_0_528, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__4)))).# trainpos
cnf(c_0_527, negated_conjecture, (k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)))=k1_real_1(esk1_0)|~r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_500, negated_conjecture, (r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),k1_seq_1(k2_fdiff_9,esk1_0))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),esk1_0))).# trainpos
cnf(c_0_501, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_real_1(np__1))|r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))).# trainpos
cnf(c_0_517, negated_conjecture, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),esk1_0)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_real_1(np__1)))).# trainpos
cnf(c_0_526, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_real_1(X1))|~r1_xxreal_0(X1,k10_real_1(k32_sin_cos,np__4))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_532, negated_conjecture, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),esk1_0)|r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_540, negated_conjecture, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_real_1(esk1_0))|~r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_545, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__4))=esk1_0|r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_550, negated_conjecture, (k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__4)|r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_549, negated_conjecture, (r1_xxreal_0(esk1_0,k10_real_1(k32_sin_cos,np__4))|~r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).# trainpos
cnf(c_0_571, negated_conjecture, (r2_hidden(esk1_0,a_0_2_sin_cos9)|r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r1_xxreal_0(esk1_0,k10_real_1(k32_sin_cos,np__4)))).# trainpos
cnf(c_0_557, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0|k1_real_1(k10_real_1(k32_sin_cos,np__4))=esk1_0)).# trainpos
cnf(c_0_562, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0|k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__4))).# trainpos
cnf(c_0_561, negated_conjecture, (r1_xxreal_0(esk1_0,k10_real_1(k32_sin_cos,np__4))|k1_real_1(k10_real_1(k32_sin_cos,np__4))=esk1_0)).# trainpos
cnf(c_0_572, negated_conjecture, (r1_xxreal_0(esk1_0,k10_real_1(k32_sin_cos,np__4)))).# trainpos
cnf(c_0_573, negated_conjecture, (k1_seq_1(k2_fdiff_9,esk1_0)=k1_real_1(np__1)|k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__4))).# trainpos
cnf(c_0_569, negated_conjecture, (k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)))=k1_seq_1(k2_fdiff_9,esk1_0)|k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0)).# trainpos
cnf(c_0_581, negated_conjecture, (k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__4))).# trainpos
cnf(c_0_579, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_real_1(np__1)))).# trainpos
cnf(c_0_580, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|r2_hidden(esk1_0,a_0_2_sin_cos9))).# trainpos
cnf(c_0_585, negated_conjecture, (r2_hidden(esk1_0,a_0_2_sin_cos9))).# trainpos
cnf(c_0_578, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_real_1(np__1))|k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0)).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_149, plain, (k7_square_1(np__4)=np__2)).#trainneg
cnf(i_0_174, plain, (k4_xcmplx_0(np__0)=np__0)).#trainneg
cnf(i_0_325, plain, (k18_sin_cos(k32_sin_cos)=k6_numbers)).#trainneg
cnf(i_0_294, plain, (v1_funct_1(k16_sin_cos))).#trainneg
cnf(i_0_377, plain, (v1_funct_1(k1_sin_cos9))).#trainneg
cnf(i_0_406, plain, (v1_funct_1(k19_sin_cos))).#trainneg
cnf(i_0_205, plain, (v1_funct_1(k1_fdiff_9))).#trainneg
cnf(i_0_349, plain, (v1_funct_1(k2_sin_cos9))).#trainneg
cnf(i_0_346, plain, (v1_funct_1(k29_sin_cos))).#trainneg
cnf(i_0_421, plain, (v1_funct_1(k30_sin_cos))).#trainneg
cnf(i_0_401, plain, (v1_funct_1(k1_sin_cos6))).#trainneg
cnf(i_0_76, plain, (v1_xreal_0(esk8_0))).#trainneg
cnf(i_0_431, plain, (v1_xreal_0(esk18_0))).#trainneg
cnf(i_0_132, plain, (v1_xxreal_0(esk9_0))).#trainneg
cnf(i_0_433, plain, (v1_xxreal_0(esk18_0))).#trainneg
cnf(i_0_28, plain, (v2_xxreal_0(np__2))).#trainneg
cnf(i_0_41, plain, (v2_xxreal_0(np__4))).#trainneg
cnf(i_0_18, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_126, plain, (v2_xxreal_0(np__3))).#trainneg
cnf(i_0_520, plain, (v2_xxreal_0(k31_sin_cos))).#trainneg
cnf(i_0_232, plain, (v2_xxreal_0(np__5))).#trainneg
cnf(i_0_251, plain, (v2_xxreal_0(np__7))).#trainneg
cnf(i_0_311, plain, (v2_xxreal_0(np__15))).#trainneg
cnf(i_0_78, plain, (v1_xcmplx_0(esk8_0))).#trainneg
cnf(i_0_432, plain, (v1_xcmplx_0(esk18_0))).#trainneg
cnf(i_0_91, plain, (v6_membered(k4_ordinal1))).#trainneg
cnf(i_0_506, plain, (v7_ordinal1(esk19_0))).#trainneg
cnf(i_0_407, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_208, plain, (v1_xboole_0(np__0))).#trainneg
cnf(i_0_133, plain, (v1_xboole_0(esk9_0))).#trainneg
cnf(i_0_538, plain, (v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_430, plain, (v1_int_1(esk18_0))).#trainneg
cnf(i_0_191, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_561, plain, (v3_xxreal_0(esk24_0))).#trainneg
cnf(i_0_507, plain, (~v1_xboole_0(esk19_0))).#trainneg
cnf(i_0_326, plain, (k21_sin_cos(k32_sin_cos)=k1_real_1(np__1))).#trainneg
cnf(i_0_454, plain, (k4_xcmplx_0(k4_xcmplx_0(np__2))=np__2)).#trainneg
cnf(i_0_307, plain, (k4_xcmplx_0(k4_xcmplx_0(np__4))=np__4)).#trainneg
cnf(i_0_83, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_378, plain, (k4_xcmplx_0(k4_xcmplx_0(np__3))=np__3)).#trainneg
cnf(i_0_612, plain, (k4_xcmplx_0(k4_xcmplx_0(np__5))=np__5)).#trainneg
cnf(i_0_165, plain, (k1_seq_1(k16_sin_cos,k32_sin_cos)=k6_numbers)).#trainneg
cnf(i_0_215, plain, (k1_seq_1(k1_fdiff_9,k6_numbers)=np__1)).#trainneg
cnf(i_0_265, plain, (k7_xcmplx_0(np__2,np__2)=np__1)).#trainneg
cnf(i_0_341, plain, (k7_xcmplx_0(np__4,np__4)=np__1)).#trainneg
cnf(i_0_338, plain, (k7_xcmplx_0(np__1,np__1)=np__1)).#trainneg
cnf(i_0_462, plain, (k6_xcmplx_0(np__2,np__2)=np__0)).#trainneg
cnf(i_0_574, plain, (k6_xcmplx_0(np__2,np__0)=np__2)).#trainneg
cnf(i_0_502, plain, (k6_xcmplx_0(np__4,np__2)=np__2)).#trainneg
cnf(i_0_488, plain, (k6_xcmplx_0(np__4,np__4)=np__0)).#trainneg
cnf(i_0_422, plain, (k6_xcmplx_0(np__4,np__1)=np__3)).#trainneg
cnf(i_0_596, plain, (k6_xcmplx_0(np__4,np__3)=np__1)).#trainneg
cnf(i_0_379, plain, (k6_xcmplx_0(np__1,np__1)=np__0)).#trainneg
cnf(i_0_372, plain, (k6_xcmplx_0(np__3,np__1)=np__2)).#trainneg
cnf(i_0_595, plain, (k6_xcmplx_0(np__3,np__3)=np__0)).#trainneg
cnf(i_0_486, plain, (k6_xcmplx_0(np__5,np__2)=np__3)).#trainneg
cnf(i_0_567, plain, (k6_xcmplx_0(np__5,np__1)=np__4)).#trainneg
cnf(i_0_410, plain, (k2_xcmplx_0(np__2,np__2)=np__4)).#trainneg
cnf(i_0_554, plain, (k2_xcmplx_0(np__2,np__1)=np__3)).#trainneg
cnf(i_0_589, plain, (k2_xcmplx_0(np__4,np__1)=np__5)).#trainneg
cnf(i_0_555, plain, (k2_xcmplx_0(np__1,np__2)=np__3)).#trainneg
cnf(i_0_357, plain, (k2_xcmplx_0(np__1,np__1)=np__2)).#trainneg
cnf(i_0_614, plain, (k2_xcmplx_0(np__1,np__0)=np__1)).#trainneg
cnf(i_0_423, plain, (k2_xcmplx_0(np__0,np__1)=np__1)).#trainneg
cnf(i_0_427, plain, (k3_xcmplx_0(np__2,np__2)=np__4)).#trainneg
cnf(i_0_492, plain, (k3_xcmplx_0(np__2,np__1)=np__2)).#trainneg
cnf(i_0_527, plain, (k3_xcmplx_0(np__2,np__0)=np__0)).#trainneg
cnf(i_0_608, plain, (k3_xcmplx_0(np__1,np__2)=np__2)).#trainneg
cnf(i_0_426, plain, (k3_xcmplx_0(np__1,np__1)=np__1)).#trainneg
cnf(i_0_497, plain, (k3_xcmplx_0(np__0,np__2)=np__0)).#trainneg
cnf(i_0_537, plain, (k3_xcmplx_0(np__0,np__0)=np__0)).#trainneg
cnf(i_0_61, plain, (m1_subset_1(k32_sin_cos,k1_numbers))).#trainneg
cnf(i_0_25, plain, (m1_subset_1(np__2,k1_numbers))).#trainneg
cnf(i_0_26, plain, (m1_subset_1(np__2,k5_numbers))).#trainneg
cnf(i_0_38, plain, (m1_subset_1(np__4,k1_numbers))).#trainneg
cnf(i_0_39, plain, (m1_subset_1(np__4,k5_numbers))).#trainneg
cnf(i_0_16, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_335, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_98, plain, (m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_99, plain, (m1_subset_1(np__0,k5_numbers))).#trainneg
cnf(i_0_123, plain, (m1_subset_1(np__3,k1_numbers))).#trainneg
cnf(i_0_124, plain, (m1_subset_1(np__3,k5_numbers))).#trainneg
cnf(i_0_229, plain, (m1_subset_1(np__5,k1_numbers))).#trainneg
cnf(i_0_230, plain, (m1_subset_1(np__5,k5_numbers))).#trainneg
cnf(i_0_248, plain, (m1_subset_1(np__7,k1_numbers))).#trainneg
cnf(i_0_249, plain, (m1_subset_1(np__7,k5_numbers))).#trainneg
cnf(i_0_308, plain, (m1_subset_1(np__15,k1_numbers))).#trainneg
cnf(i_0_309, plain, (m1_subset_1(np__15,k5_numbers))).#trainneg
cnf(i_0_434, plain, (m1_subset_1(esk18_0,k1_numbers))).#trainneg
cnf(i_0_552, plain, (r1_xxreal_0(np__4,np__4))).#trainneg
cnf(i_0_615, plain, (r1_xxreal_0(np__1,np__2))).#trainneg
cnf(i_0_547, plain, (r1_xxreal_0(np__0,np__2))).#trainneg
cnf(i_0_602, plain, (r1_xxreal_0(np__0,np__4))).#trainneg
cnf(i_0_475, plain, (r1_xxreal_0(np__0,np__1))).#trainneg
cnf(i_0_556, plain, (r1_xxreal_0(np__0,np__0))).#trainneg
cnf(i_0_344, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_245, plain, (k2_xboole_0(X1,k1_xboole_0)=X1)).#trainneg
cnf(i_0_257, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_200, plain, (k2_xboole_0(X1,X1)=X1)).#trainneg
cnf(i_0_159, plain, (k3_xboole_0(X1,X1)=X1)).#trainneg
cnf(i_0_510, plain, (~r1_xxreal_0(np__2,np__1))).#trainneg
cnf(i_0_570, plain, (~r1_xxreal_0(np__2,np__0))).#trainneg
cnf(i_0_607, plain, (~r1_xxreal_0(np__4,np__2))).#trainneg
cnf(i_0_568, plain, (~r1_xxreal_0(np__4,np__1))).#trainneg
cnf(i_0_489, plain, (~r1_xxreal_0(np__1,np__0))).#trainneg
cnf(i_0_584, plain, (~r1_xxreal_0(np__3,np__1))).#trainneg
cnf(i_0_190, plain, (k1_comptrig(k7_complex1)=k10_real_1(k32_sin_cos,np__2))).#trainneg
cnf(i_0_212, plain, (k1_seq_1(k1_fdiff_9,k32_sin_cos)=k1_real_1(np__1))).#trainneg
cnf(i_0_286, plain, (k6_xcmplx_0(np__2,np__4)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_184, plain, (k6_xcmplx_0(np__2,np__3)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_447, plain, (k6_xcmplx_0(np__2,np__5)=k4_xcmplx_0(np__3))).#trainneg
cnf(i_0_169, plain, (k6_xcmplx_0(np__1,np__2)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_364, plain, (k6_xcmplx_0(np__1,np__4)=k4_xcmplx_0(np__3))).#trainneg
cnf(i_0_320, plain, (k6_xcmplx_0(np__1,np__3)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_370, plain, (k6_xcmplx_0(np__1,np__5)=k4_xcmplx_0(np__4))).#trainneg
cnf(i_0_314, plain, (k6_xcmplx_0(np__0,np__2)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_260, plain, (k6_xcmplx_0(np__0,np__4)=k4_xcmplx_0(np__4))).#trainneg
cnf(i_0_170, plain, (k6_xcmplx_0(np__0,np__1)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_402, plain, (k6_xcmplx_0(np__0,np__3)=k4_xcmplx_0(np__3))).#trainneg
cnf(i_0_587, plain, (k6_xcmplx_0(np__0,np__5)=k4_xcmplx_0(np__5))).#trainneg
cnf(i_0_550, plain, (k6_xcmplx_0(np__0,np__7)=k4_xcmplx_0(np__7))).#trainneg
cnf(i_0_188, plain, (k6_xcmplx_0(np__3,np__4)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_500, plain, (k6_xcmplx_0(np__3,np__5)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_565, plain, (k6_xcmplx_0(np__3,np__7)=k4_xcmplx_0(np__4))).#trainneg
cnf(i_0_551, plain, (k6_xcmplx_0(np__5,np__7)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_617, plain, (k6_xcmplx_0(np__2,k4_xcmplx_0(np__1))=np__3)).#trainneg
cnf(i_0_534, plain, (k6_xcmplx_0(np__1,k4_xcmplx_0(np__1))=np__2)).#trainneg
cnf(i_0_588, plain, (k6_xcmplx_0(np__3,k4_xcmplx_0(np__1))=np__4)).#trainneg
cnf(i_0_553, plain, (k2_xcmplx_0(np__2,k4_xcmplx_0(np__2))=np__0)).#trainneg
cnf(i_0_460, plain, (k2_xcmplx_0(np__2,k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_459, plain, (k2_xcmplx_0(np__1,k4_xcmplx_0(np__1))=np__0)).#trainneg
cnf(i_0_566, plain, (k2_xcmplx_0(k4_xcmplx_0(np__2),np__2)=np__0)).#trainneg
cnf(i_0_611, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),np__2)=np__1)).#trainneg
cnf(i_0_540, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__0))).#trainneg
cnf(i_0_509, plain, (m1_subset_1(esk20_1(X1),X1))).#trainneg
cnf(i_0_327, plain, (k18_sin_cos(k10_real_1(k32_sin_cos,np__2))=np__1)).#trainneg
cnf(i_0_321, plain, (k18_sin_cos(k8_real_1(np__2,k32_sin_cos))=k6_numbers)).#trainneg
cnf(i_0_281, plain, (k2_sin_cos4(k10_real_1(k32_sin_cos,np__2))=k6_numbers)).#trainneg
cnf(i_0_473, plain, (k2_sin_cos4(k10_real_1(k32_sin_cos,np__4))=np__1)).#trainneg
cnf(i_0_328, plain, (k21_sin_cos(k10_real_1(k32_sin_cos,np__2))=k6_numbers)).#trainneg
cnf(i_0_322, plain, (k21_sin_cos(k8_real_1(np__2,k32_sin_cos))=np__1)).#trainneg
cnf(i_0_437, plain, (~r1_xxreal_0(np__2,k7_square_1(np__2)))).#trainneg
cnf(i_0_461, plain, (~r1_xxreal_0(np__0,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_118, plain, (~r1_xxreal_0(k7_square_1(np__2),np__1))).#trainneg
cnf(i_0_526, plain, (~v1_xboole_0(k2_tarski(X1,X2)))).#trainneg
cnf(i_0_101, plain, (k2_fdiff_9=k6_rfunct_1(k1_numbers,k1_numbers,k16_sin_cos))).#trainneg
cnf(i_0_27, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_40, plain, (m2_subset_1(np__4,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_17, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_100, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_125, plain, (m2_subset_1(np__3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_231, plain, (m2_subset_1(np__5,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_250, plain, (m2_subset_1(np__7,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_310, plain, (m2_subset_1(np__15,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_293, plain, (v1_funct_2(k16_sin_cos,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_405, plain, (v1_funct_2(k19_sin_cos,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_515, plain, (v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_254, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_144, plain, (v1_xreal_0(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_140, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_438, plain, (v3_membered(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_56, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_129, plain, (v1_xcmplx_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_366, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_171, plain, (v1_int_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_303, plain, (v1_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_300, plain, (v5_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_428, plain, (v4_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_224, plain, (k10_xtuple_0(k16_sin_cos)=k1_rcomp_1(k1_real_1(np__1),np__1))).#trainneg
cnf(i_0_181, plain, (k7_xcmplx_0(np__1,k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_505, plain, (k2_xcmplx_0(np__1,k4_xcmplx_0(np__2))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_546, plain, (k2_xcmplx_0(np__0,k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_560, plain, (k3_xcmplx_0(np__2,k4_xcmplx_0(np__2))=k4_xcmplx_0(np__4))).#trainneg
cnf(i_0_548, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),np__2)=k4_xcmplx_0(np__4))).#trainneg
cnf(i_0_469, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),np__2)=k4_xcmplx_0(np__3))).#trainneg
cnf(i_0_468, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),np__1)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_572, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),np__0)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_455, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),np__3)=k4_xcmplx_0(np__4))).#trainneg
cnf(i_0_603, plain, (k2_xcmplx_0(k4_xcmplx_0(np__2),np__1)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_535, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),np__0)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_559, plain, (k6_xcmplx_0(k4_xcmplx_0(np__4),k4_xcmplx_0(np__4))=np__0)).#trainneg
cnf(i_0_482, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))=np__0)).#trainneg
cnf(i_0_112, plain, (k1_seq_1(k1_sin_cos9,np__1)=k10_real_1(k32_sin_cos,np__4))).#trainneg
cnf(i_0_244, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1))).#trainneg
cnf(i_0_51, plain, (k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__2))=np__1)).#trainneg
cnf(i_0_161, plain, (k1_seq_1(k16_sin_cos,k8_real_1(np__2,k32_sin_cos))=k6_numbers)).#trainneg
cnf(i_0_168, plain, (k1_seq_1(k19_sin_cos,k10_real_1(k32_sin_cos,np__2))=k6_numbers)).#trainneg
cnf(i_0_162, plain, (k1_seq_1(k19_sin_cos,k8_real_1(np__2,k32_sin_cos))=np__1)).#trainneg
cnf(i_0_282, plain, (k1_seq_1(k30_sin_cos,k10_real_1(k32_sin_cos,np__2))=k6_numbers)).#trainneg
cnf(i_0_474, plain, (k1_seq_1(k30_sin_cos,k10_real_1(k32_sin_cos,np__4))=np__1)).#trainneg
cnf(i_0_456, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,np__2))=np__2)).#trainneg
cnf(i_0_397, plain, (~r1_xxreal_0(k32_sin_cos,k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_536, plain, (~r1_xxreal_0(np__2,k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_539, plain, (~r1_xxreal_0(np__2,k7_xcmplx_0(np__3,np__2)))).#trainneg
cnf(i_0_491, plain, (~r1_xxreal_0(np__2,k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_531, plain, (~r1_xxreal_0(np__4,k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_613, plain, (~r1_xxreal_0(np__4,k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_425, plain, (~r1_xxreal_0(np__4,k7_xcmplx_0(np__3,np__2)))).#trainneg
cnf(i_0_579, plain, (~r1_xxreal_0(np__4,k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_439, plain, (~r1_xxreal_0(np__1,k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_569, plain, (~r1_xxreal_0(np__3,k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_533, plain, (~r1_xxreal_0(np__3,k7_xcmplx_0(np__3,np__2)))).#trainneg
cnf(i_0_493, plain, (~r1_xxreal_0(k7_xcmplx_0(np__1,np__2),np__0))).#trainneg
cnf(i_0_463, plain, (~r1_xxreal_0(k7_xcmplx_0(np__1,np__4),np__0))).#trainneg
cnf(i_0_494, plain, (~r1_xxreal_0(k7_xcmplx_0(np__3,np__2),np__1))).#trainneg
cnf(i_0_429, plain, (~r1_xxreal_0(k7_xcmplx_0(np__3,np__2),np__0))).#trainneg
cnf(i_0_458, plain, (~r1_xxreal_0(k7_xcmplx_0(np__3,np__4),np__0))).#trainneg
cnf(i_0_394, plain, (~r1_xxreal_0(k8_real_1(np__2,k32_sin_cos),k32_sin_cos))).#trainneg
cnf(i_0_391, plain, (~r1_xxreal_0(k8_real_1(np__2,k32_sin_cos),k6_numbers))).#trainneg
cnf(i_0_32, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_302, plain, (v1_xreal_0(k6_square_1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_135, plain, (v1_xreal_0(k5_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_33, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_107, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_96, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_136, plain, (v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_121, plain, (v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_95, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_60, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_156, plain, (k5_xcmplx_0(k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_436, plain, (k3_xcmplx_0(X1,k6_numbers)=k6_numbers|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_387, plain, (k7_xcmplx_0(k6_numbers,X1)=k6_numbers|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_189, plain, (k7_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_274, plain, (k6_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_259, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_312, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_445, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_93, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_225, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_383, plain, (X1=esk17_1(X1)|~r2_hidden(X1,a_0_0_sin_cos9))).#trainneg
cnf(i_0_593, plain, (X1=esk26_1(X1)|~r2_hidden(X1,a_0_1_sin_cos9))).#trainneg
cnf(i_0_154, plain, (k5_sin_cos9(k1_real_1(np__1))=k1_real_1(k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_122, plain, (k3_sin_cos6(k1_real_1(np__1))=k1_real_1(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_214, plain, (k1_seq_1(k1_fdiff_9,k10_real_1(k32_sin_cos,np__4))=k7_square_1(np__2))).#trainneg
cnf(i_0_194, plain, (k18_sin_cos(k1_real_1(k10_real_1(k32_sin_cos,np__2)))=k1_real_1(np__1))).#trainneg
cnf(i_0_222, plain, (k1_sin_cos4(k1_real_1(k10_real_1(k32_sin_cos,np__4)))=k1_real_1(np__1))).#trainneg
cnf(i_0_292, plain, (m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_376, plain, (m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_404, plain, (m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_204, plain, (m1_subset_1(k1_fdiff_9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_348, plain, (m1_subset_1(k2_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_345, plain, (m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_420, plain, (m1_subset_1(k30_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_400, plain, (m1_subset_1(k1_sin_cos6,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_324, plain, (k21_sin_cos(k7_real_1(k32_sin_cos,k10_real_1(k32_sin_cos,np__2)))=k6_numbers)).#trainneg
cnf(i_0_392, plain, (~r1_xxreal_0(k6_numbers,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_330, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_115, plain, (k20_sin_cos(X1)=k1_seq_1(k19_sin_cos,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_152, plain, (k17_sin_cos(X1)=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_89, plain, (k7_xcmplx_0(np__1,X1)=k5_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_130, plain, (k6_square_1(X1)=k7_square_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_450, plain, (k7_xcmplx_0(np__1,k4_xcmplx_0(np__2))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_481, plain, (v1_relat_1(k5_relat_1(X1,X2))|~v1_relat_1(X1))).#trainneg
cnf(i_0_207, plain, (k6_sin_cos9(k1_real_1(np__1))=k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))).#trainneg
cnf(i_0_490, plain, (k4_xcmplx_0(k7_xcmplx_0(np__4,np__3))=k7_xcmplx_0(k4_xcmplx_0(np__4),np__3))).#trainneg
cnf(i_0_483, plain, (k4_xcmplx_0(k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_361, plain, (k4_xcmplx_0(k7_xcmplx_0(np__1,np__4))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))).#trainneg
cnf(i_0_549, plain, (k4_xcmplx_0(k7_xcmplx_0(np__1,np__3))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__3))).#trainneg
cnf(i_0_573, plain, (k4_xcmplx_0(k7_xcmplx_0(np__3,np__4))=k7_xcmplx_0(k4_xcmplx_0(np__3),np__4))).#trainneg
cnf(i_0_153, plain, (k1_seq_1(k1_sin_cos9,k1_real_1(np__1))=k1_real_1(k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_499, plain, (k3_xcmplx_0(np__3,k7_xcmplx_0(np__1,np__4))=k7_xcmplx_0(np__3,np__4))).#trainneg
cnf(i_0_578, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,np__4),np__2)=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_498, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,np__4),np__3)=k7_xcmplx_0(np__3,np__4))).#trainneg
cnf(i_0_464, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_413, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))=k7_xcmplx_0(np__1,np__4))).#trainneg
cnf(i_0_479, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_193, plain, (k1_seq_1(k16_sin_cos,k1_real_1(k10_real_1(k32_sin_cos,np__2)))=k1_real_1(np__1))).#trainneg
cnf(i_0_223, plain, (k1_seq_1(k29_sin_cos,k1_real_1(k10_real_1(k32_sin_cos,np__4)))=k1_real_1(np__1))).#trainneg
cnf(i_0_323, plain, (k18_sin_cos(k7_real_1(k32_sin_cos,k10_real_1(k32_sin_cos,np__2)))=k1_real_1(np__1))).#trainneg
cnf(i_0_471, plain, (k2_sin_cos4(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))=k1_real_1(np__1))).#trainneg
cnf(i_0_164, plain, (k1_seq_1(k19_sin_cos,k7_real_1(k32_sin_cos,k10_real_1(k32_sin_cos,np__2)))=k6_numbers)).#trainneg
cnf(i_0_210, plain, (r2_hidden(k6_numbers,k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_581, plain, (~r1_xxreal_0(k7_xcmplx_0(np__3,np__2),k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_597, plain, (~r1_xxreal_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_390, plain, (~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos),k32_sin_cos))).#trainneg
cnf(i_0_388, plain, (~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos),k6_numbers))).#trainneg
cnf(i_0_511, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_453, plain, (~v1_xxreal_0(X1)|~v2_xxreal_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_451, plain, (~v1_xxreal_0(X1)|~v2_xxreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_518, plain, (~v1_xxreal_0(X1)|~v1_xboole_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_301, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_160, plain, (m1_subset_1(k7_square_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_384, plain, (m1_subset_1(esk17_1(X1),k1_numbers)|~r2_hidden(X1,a_0_0_sin_cos9))).#trainneg
cnf(i_0_545, plain, (m1_subset_1(esk23_1(X1),k1_numbers)|~r2_hidden(X1,a_0_2_sin_cos9))).#trainneg
cnf(i_0_594, plain, (m1_subset_1(esk26_1(X1),k1_numbers)|~r2_hidden(X1,a_0_1_sin_cos9))).#trainneg
cnf(i_0_206, plain, (k1_seq_1(k2_sin_cos9,k1_real_1(np__1))=k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))).#trainneg
cnf(i_0_228, plain, (k1_seq_1(k16_sin_cos,k10_real_1(k32_sin_cos,np__4))=k10_real_1(np__1,k7_square_1(np__2)))).#trainneg
cnf(i_0_227, plain, (k1_seq_1(k19_sin_cos,k10_real_1(k32_sin_cos,np__4))=k10_real_1(np__1,k7_square_1(np__2)))).#trainneg
cnf(i_0_163, plain, (k1_seq_1(k16_sin_cos,k7_real_1(k32_sin_cos,k10_real_1(k32_sin_cos,np__2)))=k1_real_1(np__1))).#trainneg
cnf(i_0_472, plain, (k1_seq_1(k30_sin_cos,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))=k1_real_1(np__1))).#trainneg
cnf(i_0_395, plain, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_610, plain, (v1_relat_1(k5_relat_1(k16_sin_cos,k1_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))))).#trainneg
cnf(i_0_609, plain, (v2_funct_1(k5_relat_1(k16_sin_cos,k1_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))))).#trainneg
cnf(i_0_23, plain, (v6_xxreal_2(X1)|v1_xxreal_0(esk4_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_22, plain, (v6_xxreal_2(X1)|v1_xxreal_0(esk5_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_457, plain, (v3_membered(k10_xtuple_0(X1))|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_284, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_448, plain, (v1_funct_1(X1)|~v4_funct_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_449, plain, (v1_relat_1(X1)|~v4_funct_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_185, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_73, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_352, plain, (v1_int_1(X2)|~v5_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_276, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_273, plain, (k10_xtuple_0(k5_relat_1(X1,X2))=k7_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_524, plain, (r2_hidden(esk21_2(X1,X2),X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_525, plain, (k2_xcmplx_0(k7_xcmplx_0(np__1,np__4),k7_xcmplx_0(np__1,np__4))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_237, plain, (k1_seq_1(k19_sin_cos,k1_real_1(k10_real_1(k32_sin_cos,np__4)))=k10_real_1(np__1,k7_square_1(np__2)))).#trainneg
cnf(i_0_213, plain, (k1_seq_1(k1_fdiff_9,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))=k1_real_1(k7_square_1(np__2)))).#trainneg
cnf(i_0_258, plain, (k9_subset_1(X2,X3,X3)=X3|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_209, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_389, plain, (~r1_xxreal_0(k8_real_1(np__2,k32_sin_cos),k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos)))).#trainneg
cnf(i_0_393, plain, (~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos),k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_187, plain, (v1_relat_1(k5_relat_1(k2_fdiff_9,k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))))).#trainneg
cnf(i_0_583, plain, (v1_relat_1(k5_relat_1(k16_sin_cos,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))))).#trainneg
cnf(i_0_186, plain, (v2_funct_1(k5_relat_1(k2_fdiff_9,k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))))).#trainneg
cnf(i_0_582, plain, (v2_funct_1(k5_relat_1(k16_sin_cos,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))))).#trainneg
cnf(i_0_21, plain, (v6_xxreal_2(X1)|r2_hidden(esk4_1(X1),X1)|~v2_membered(X1))).#trainneg
cnf(i_0_20, plain, (v6_xxreal_2(X1)|r2_hidden(esk5_1(X1),X1)|~v2_membered(X1))).#trainneg
cnf(i_0_334, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_408, plain, (v1_funct_1(k5_relat_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_94, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_142, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_306, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_337, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_503, plain, (v6_xxreal_2(k1_xxreal_1(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_355, plain, (v6_xxreal_2(k4_xxreal_1(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_158, plain, (v6_xxreal_2(k2_xxreal_1(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_246, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_313, plain, (v1_xcmplx_0(k6_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_351, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_343, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_371, plain, (v1_int_1(k6_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_501, plain, (v1_int_1(k2_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_530, plain, (v1_int_1(k3_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_476, plain, (v3_rcomp_1(k4_xxreal_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_542, plain, (r1_xxreal_0(esk23_1(X1),k10_real_1(k32_sin_cos,np__4))|~r2_hidden(X1,a_0_2_sin_cos9))).#trainneg
cnf(i_0_375, plain, (k5_relat_1(k5_relat_1(X1,X2),X2)=k5_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_381, plain, (~r2_hidden(X1,a_0_0_sin_cos9)|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),esk17_1(X1)))).#trainneg
cnf(i_0_591, plain, (~r2_hidden(X1,a_0_1_sin_cos9)|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),esk26_1(X1)))).#trainneg
cnf(i_0_411, plain, (k2_relset_1(k1_numbers,k1_sin_cos6)=k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_236, plain, (k1_seq_1(k16_sin_cos,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))=k10_real_1(np__1,k7_square_1(np__2)))).#trainneg
cnf(i_0_238, plain, (k1_seq_1(k16_sin_cos,k1_real_1(k10_real_1(k32_sin_cos,np__4)))=k1_real_1(k10_real_1(np__1,k7_square_1(np__2))))).#trainneg
cnf(i_0_241, plain, (r1_tarski(k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers),k1_relset_1(k1_numbers,k2_fdiff_9)))).#trainneg
cnf(i_0_365, plain, (r2_fdiff_1(k29_sin_cos,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_198, plain, (v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))))).#trainneg
cnf(i_0_64, plain, (k2_xxreal_1(X1,X2)=k3_rcomp_1(X1,X2)|~v1_xreal_0(X1)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_203, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_360, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_155, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_172, plain, (k2_relset_1(X2,X1)=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).#trainneg
cnf(i_0_71, plain, (X1=X4|X1=X3|X2!=k2_tarski(X3,X4)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_45, plain, (r2_hidden(X1,X3)|X3!=k2_xboole_0(X4,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_216, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_235, plain, (k1_seq_1(k19_sin_cos,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))=k1_real_1(k10_real_1(np__1,k7_square_1(np__2))))).#trainneg
cnf(i_0_382, plain, (~r2_hidden(X1,a_0_0_sin_cos9)|~r1_xxreal_0(esk17_1(X1),k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_592, plain, (~r2_hidden(X1,a_0_1_sin_cos9)|~r1_xxreal_0(esk26_1(X1),k1_real_1(k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_252, plain, (r1_tarski(k1_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__4)),k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_414, plain, (r1_tarski(k2_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_218, plain, (r1_tarski(k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers),k2_rcomp_1(k1_real_1(k32_sin_cos),k6_numbers)))).#trainneg
cnf(i_0_362, plain, (v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))))).#trainneg
cnf(i_0_563, plain, (v1_fcont_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))))).#trainneg
cnf(i_0_336, plain, (v2_funct_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))))).#trainneg
cnf(i_0_255, plain, (v1_funct_1(k4_rfunct_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_256, plain, (v1_relat_1(k4_rfunct_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_465, plain, (v3_ordinal1(k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v5_ordinal1(X1))).#trainneg
cnf(i_0_443, plain, (k2_funct_1(k2_funct_1(X1))=X1|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_funct_1(X1))).#trainneg
cnf(i_0_333, plain, (k7_xcmplx_0(X1,X2)=k3_xcmplx_0(X1,k5_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_177, plain, (k6_xcmplx_0(X1,X2)=k2_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_34, plain, (k7_xcmplx_0(X1,X2)=k10_real_1(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_63, plain, (m1_subset_1(k3_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_283, plain, (m1_subset_1(k2_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_386, plain, (r2_hidden(k1_seq_1(k6_real_3(X2),X1),k5_numbers)|~v1_xreal_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_50, plain, (m1_subset_1(k10_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_226, plain, (m1_subset_1(k8_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_247, plain, (k9_xtuple_0(k5_relat_1(X1,X2))=X2|~v1_relat_1(X1)|~r1_tarski(X2,k9_xtuple_0(X1)))).#trainneg
cnf(i_0_279, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_280, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_90, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_369, plain, (v1_valued_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_532, plain, (v2_valued_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_157, plain, (m1_subset_1(k9_subset_1(X2,X3,X1),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_277, plain, (r1_tarski(k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)),k1_relset_1(k1_numbers,k29_sin_cos)))).#trainneg
cnf(i_0_339, plain, (k9_xtuple_0(X1)=k10_xtuple_0(k2_funct_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_funct_1(X1))).#trainneg
cnf(i_0_340, plain, (k10_xtuple_0(X1)=k9_xtuple_0(k2_funct_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_funct_1(X1))).#trainneg
cnf(i_0_442, plain, (v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_350, plain, (v3_valued_0(X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_496, plain, (v1_valued_0(X2)|~v1_relat_1(X1)|~v1_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_264, plain, (v1_xxreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_valued_0(X1))).#trainneg
cnf(i_0_219, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_557, plain, (v7_ordinal1(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_329, plain, (k7_xcmplx_0(k5_xcmplx_0(X1),k5_xcmplx_0(X2))=k7_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_356, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_30, plain, (r1_tarski(k4_xxreal_1(X1,X2),k1_xxreal_1(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_523, plain, (~r2_hidden(X3,X2)|~r2_hidden(X1,X2)|~r2_hidden(X1,esk21_2(X3,X2)))).#trainneg
cnf(i_0_19, plain, (v6_xxreal_2(X1)|~v2_membered(X1)|~r1_tarski(k1_xxreal_1(esk4_1(X1),esk5_1(X1)),X1))).#trainneg
cnf(i_0_253, plain, (r1_tarski(k1_rcomp_1(k10_real_1(k32_sin_cos,np__4),k10_real_1(k32_sin_cos,np__2)),k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_278, plain, (r1_tarski(k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers),k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))).#trainneg
cnf(i_0_604, plain, (k3_relat_1(k4_sin_cos6,k5_relset_1(k1_numbers,k1_numbers,k19_sin_cos,k1_rcomp_1(k6_numbers,k32_sin_cos)))=k6_partfun1(k1_rcomp_1(k1_real_1(np__1),np__1)))).#trainneg
cnf(i_0_211, plain, (k3_sincos10=k2_partfun2(k1_numbers,k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k3_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))))).#trainneg
cnf(i_0_478, plain, (v2_funct_1(k2_partfun1(k1_numbers,k1_numbers,k29_sin_cos,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_358, plain, (v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,k16_sin_cos,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)))))).#trainneg
cnf(i_0_519, plain, (v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,k29_sin_cos,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)))))).#trainneg
cnf(i_0_417, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_495, plain, (k3_xcmplx_0(k5_xcmplx_0(X1),k5_xcmplx_0(X2))=k5_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_347, plain, (k5_relat_1(k5_relat_1(X1,X3),X2)=k5_relat_1(X1,X2)|~v1_relat_1(X1)|~r1_tarski(X2,X3))).#trainneg
cnf(i_0_305, plain, (r1_xxreal_0(k6_numbers,k1_seq_1(k16_sin_cos,X1))|~v1_xreal_0(X1)|~r2_hidden(X1,k1_rcomp_1(k6_numbers,k32_sin_cos)))).#trainneg
cnf(i_0_418, plain, (r1_xxreal_0(k6_numbers,k1_seq_1(k16_sin_cos,X1))|~v1_xreal_0(X1)|~r2_hidden(X1,k1_rcomp_1(k6_numbers,np__1)))).#trainneg
cnf(i_0_485, plain, (~v1_xreal_0(X1)|~r2_hidden(X1,k1_rcomp_1(k6_numbers,np__1))|~r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),k6_numbers))).#trainneg
cnf(i_0_367, plain, (r1_tarski(k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_150, plain, (k7_relset_1(X2,X3,X1,X4)=k7_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_353, plain, (k5_relset_1(X2,X3,X1,X4)=k5_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_513, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X2)|r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_576, plain, (v2_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_529, plain, (k3_relat_1(X1,k2_funct_1(X1))=k6_partfun1(k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_funct_1(X1))).#trainneg
cnf(i_0_528, plain, (k3_relat_1(k2_funct_1(X1),X1)=k6_partfun1(k10_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_funct_1(X1))).#trainneg
cnf(i_0_575, plain, (k18_sin_cos(k3_sin_cos6(X1))=X1|~v1_xreal_0(X1)|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(k1_real_1(np__1),X1))).#trainneg
cnf(i_0_47, plain, (r2_hidden(X1,X4)|r2_hidden(X1,X3)|X2!=k2_xboole_0(X3,X4)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_67, plain, (X3=k2_tarski(X1,X2)|esk7_3(X1,X2,X3)!=X2|~r2_hidden(esk7_3(X1,X2,X3),X3))).#trainneg
cnf(i_0_68, plain, (X3=k2_tarski(X1,X2)|esk7_3(X1,X2,X3)!=X1|~r2_hidden(esk7_3(X1,X2,X3),X3))).#trainneg
cnf(i_0_564, plain, (v6_xxreal_2(k3_xboole_0(X1,X2))|~v2_membered(X2)|~v2_membered(X1)|~v6_xxreal_2(X2)|~v6_xxreal_2(X1))).#trainneg
cnf(i_0_416, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xboole_0(k7_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_504, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xboole_0(k3_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_399, plain, (k1_sin_cos4(k5_sin_cos9(X1))=X1|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(k1_real_1(np__1),X1))).#trainneg
cnf(i_0_477, plain, (k2_sin_cos4(k6_sin_cos9(X1))=X1|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(k1_real_1(np__1),X1))).#trainneg
cnf(i_0_487, plain, (r2_hidden(k1_seq_1(k4_real_3(X2),X1),k5_numbers)|~v1_xreal_0(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_484, plain, (r1_xxreal_0(k10_real_1(np__1,np__2),k1_seq_1(k19_sin_cos,X1))|~v1_xreal_0(X1)|~r2_hidden(X1,k1_rcomp_1(k6_numbers,np__1)))).#trainneg
cnf(i_0_368, plain, (r1_xxreal_0(k1_seq_1(k16_sin_cos,X1),k6_numbers)|~v1_xreal_0(X1)|~r2_hidden(X1,k1_rcomp_1(k32_sin_cos,k8_real_1(np__2,k32_sin_cos))))).#trainneg
cnf(i_0_558, plain, (r2_hidden(k1_seq_1(k1_fdiff_9,X1),k1_rcomp_1(np__1,k7_square_1(np__2)))|~r2_hidden(X1,k1_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_117, plain, (v1_funct_1(k2_partfun1(X1,X2,X3,X4))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_571, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m2_subset_1(esk25_2(X1,X2),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_131, plain, (k2_partfun1(X2,X3,X1,X4)=k5_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_43, plain, (X3=k2_xboole_0(X1,X2)|~r2_hidden(esk6_3(X1,X2,X3),X3)|~r2_hidden(esk6_3(X1,X2,X3),X2))).#trainneg
cnf(i_0_44, plain, (X3=k2_xboole_0(X1,X2)|~r2_hidden(esk6_3(X1,X2,X3),X3)|~r2_hidden(esk6_3(X1,X2,X3),X1))).#trainneg
cnf(i_0_199, plain, (m1_subset_1(k5_relset_1(X2,X3,X1,X4),k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_13, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(X2,k1_xxreal_1(X1,X3)))).#trainneg
cnf(i_0_12, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(X1,k1_xxreal_1(X3,X2)))).#trainneg
cnf(i_0_605, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,k1_real_1(np__1))|~v1_xreal_0(X1)|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k3_sin_cos6(X1)))).#trainneg
cnf(i_0_342, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X2,X3))=k7_xcmplx_0(k3_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_419, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_580, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_217, plain, (k6_rfunct_1(X3,X1,X2)=k4_rfunct_1(X2)|~v1_funct_1(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_240, plain, (k2_partfun2(X2,X3,X1)=k2_funct_1(X1)|~v1_funct_1(X1)|~v2_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_291, plain, (m1_subset_1(esk14_3(X1,X2,X3),k1_numbers)|~v1_xreal_0(X2)|~v1_xxreal_0(X3)|~r2_hidden(X1,a_2_2_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_299, plain, (m1_subset_1(esk15_3(X1,X2,X3),k1_numbers)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,a_2_0_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_319, plain, (m1_subset_1(esk16_3(X1,X2,X3),k1_numbers)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,a_2_1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_374, plain, (v1_funct_1(k6_rfunct_1(X1,X2,X3))|~v1_funct_1(X3)|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_288, plain, (~v1_xreal_0(X3)|~v1_xxreal_0(X1)|~r2_hidden(X2,a_2_2_rcomp_1(X3,X1))|~r1_xxreal_0(X1,esk14_3(X2,X3,X1)))).#trainneg
cnf(i_0_316, plain, (~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r2_hidden(X2,a_2_1_rcomp_1(X3,X1))|~r1_xxreal_0(X1,esk16_3(X2,X3,X1)))).#trainneg
cnf(i_0_435, plain, (k1_seq_1(k29_sin_cos,X1)=k1_sin_cos4(X1)|~v1_xreal_0(X1)|~r2_hidden(X1,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_62, plain, (r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k6_numbers,X1)|~r1_xxreal_0(k6_square_1(X2),k6_square_1(X1)))).#trainneg
cnf(i_0_600, plain, (r1_xxreal_0(k5_sin_cos9(X1),k10_real_1(k32_sin_cos,np__4))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(k1_real_1(np__1),X1))).#trainneg
cnf(i_0_598, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,k1_real_1(np__1))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k5_sin_cos9(X1)))).#trainneg
cnf(i_0_196, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_606, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,k1_real_1(np__1))|~v1_xreal_0(X1)|~r1_xxreal_0(k3_sin_cos6(X1),k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_197, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_331, plain, (k3_sin_cos6(k17_sin_cos(X1))=X1|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k10_real_1(k32_sin_cos,np__2))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1))).#trainneg
cnf(i_0_480, plain, (r2_hidden(k1_seq_1(k1_sin_cos9,X1),k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k10_real_1(k32_sin_cos,np__4)))|~r2_hidden(X1,k1_rcomp_1(k1_real_1(np__1),np__1)))).#trainneg
cnf(i_0_412, plain, (r1_xxreal_0(k6_numbers,k1_seq_1(k19_sin_cos,X1))|~v1_xreal_0(X1)|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_304, plain, (~v1_xreal_0(X1)|~r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),k6_numbers)|~r2_hidden(X1,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_599, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,k1_real_1(np__1))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k5_sin_cos9(X1),k1_real_1(k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_66, plain, (X3=k2_tarski(X1,X2)|esk7_3(X1,X2,X3)=X2|esk7_3(X1,X2,X3)=X1|r2_hidden(esk7_3(X1,X2,X3),X3))).#trainneg
cnf(i_0_586, plain, (r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),k6_numbers)|~v1_xreal_0(X1)|~r2_hidden(X1,k1_rcomp_1(k10_real_1(k32_sin_cos,np__2),k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos))))).#trainneg
cnf(i_0_145, plain, (v1_fcont_1(k2_partfun1(k1_numbers,k1_numbers,X1,X2))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_116, plain, (m1_subset_1(k2_partfun1(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_616, plain, (k1_relset_1(k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k29_sin_cos,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k10_real_1(k32_sin_cos,np__4))))=k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_380, plain, (r2_hidden(X2,a_0_0_sin_cos9)|r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),X1)|r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|X1!=X2|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_590, plain, (r2_hidden(X2,a_0_1_sin_cos9)|r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),X1)|r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|X1!=X2|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_373, plain, (m1_subset_1(k6_rfunct_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X3)|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_8, plain, (m1_subset_1(esk2_2(X1,X2),k1_numbers)|v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X2,X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_7, plain, (m1_subset_1(esk3_2(X1,X2),k1_numbers)|v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X2,X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_585, plain, (k1_fdiff_1(k29_sin_cos,X1)=k10_real_1(np__1,k5_square_1(k1_seq_1(k19_sin_cos,X1)))|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_261, plain, (v2_funct_1(k2_partfun1(k1_numbers,k1_numbers,X1,X2))|~v1_funct_1(X1)|~v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X1,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_262, plain, (v2_funct_1(k2_partfun1(k1_numbers,k1_numbers,X1,X2))|~v1_funct_1(X1)|~v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X1,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_179, plain, (k4_relset_1(X2,X3,X5,X6,X1,X4)=k3_relat_1(X1,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_270, plain, (r1_fcont_1(X1,X2)|m1_rcomp_1(esk12_2(X1,X2),k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_42, plain, (X3=k2_xboole_0(X1,X2)|r2_hidden(esk6_3(X1,X2,X3),X3)|r2_hidden(esk6_3(X1,X2,X3),X2)|r2_hidden(esk6_3(X1,X2,X3),X1))).#trainneg
cnf(i_0_4, plain, (v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X2,X1))|~v1_funct_1(X2)|~r1_xxreal_0(esk3_2(X1,X2),esk2_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_522, plain, (v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k29_sin_cos,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k10_real_1(k32_sin_cos,np__4))),k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_315, plain, (r1_xxreal_0(X4,X1)|r1_xxreal_0(X1,X3)|r2_hidden(X2,a_2_1_rcomp_1(X3,X4))|X1!=X2|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_24, plain, (r1_tarski(k1_xxreal_1(X2,X3),X1)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v2_membered(X1)|~v6_xxreal_2(X1)|~r2_hidden(X3,X1)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_287, plain, (r1_xxreal_0(X4,X1)|r2_hidden(X2,a_2_2_rcomp_1(X3,X4))|X1!=X2|~v1_xreal_0(X3)|~v1_xxreal_0(X4)|~r1_xxreal_0(X3,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_295, plain, (r2_hidden(X2,a_2_0_rcomp_1(X3,X4))|X1!=X2|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X1,X4)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_269, plain, (r1_fcont_1(X1,X2)|v1_xreal_0(esk13_3(X1,X2,X3))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_rcomp_1(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_267, plain, (r1_fcont_1(X1,X2)|r2_hidden(esk13_3(X1,X2,X3),X3)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_rcomp_1(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_467, plain, (r1_xxreal_0(k18_complex1(k5_real_1(k2_xcmplx_0(X2,X3),k8_real_1(np__2,X1))),k6_xcmplx_0(X3,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_466, plain, (r2_hidden(X3,k1_rcomp_1(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k18_complex1(k5_real_1(k2_xcmplx_0(X1,X2),k8_real_1(np__2,X3))),k6_xcmplx_0(X2,X1)))).#trainneg
cnf(i_0_6, plain, (v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X2,X1))|r2_hidden(esk2_2(X1,X2),k9_subset_1(k1_numbers,X1,k1_relset_1(k1_numbers,X2)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_5, plain, (v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X2,X1))|r2_hidden(esk3_2(X1,X2),k9_subset_1(k1_numbers,X1,k1_relset_1(k1_numbers,X2)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_10, plain, (r1_xxreal_0(X3,X2)|r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k7_xcmplx_0(X1,X2),k7_xcmplx_0(X1,X3)))).#trainneg
cnf(i_0_3, plain, (v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X1,X2))|r1_xxreal_0(k1_seq_1(X1,esk2_2(X2,X1)),k1_seq_1(X1,esk3_2(X2,X1)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_268, plain, (r1_fcont_1(X1,X2)|r2_hidden(esk13_3(X1,X2,X3),k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_rcomp_1(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_272, plain, (m1_rcomp_1(esk11_3(X1,X2,X3),X2)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~r1_fcont_1(X1,X2)|~m1_rcomp_1(X3,k1_seq_1(X1,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_180, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k3_relat_1(X1,X4)|~v1_funct_1(X4)|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_266, plain, (r1_fcont_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_rcomp_1(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_hidden(k1_seq_1(X1,esk13_3(X1,X2,X3)),esk12_2(X1,X2)))).#trainneg
cnf(i_0_385, plain, (r1_xxreal_0(k7_real_1(k10_real_1(k32_sin_cos,np__2),k8_real_1(k8_real_1(np__2,k32_sin_cos),X2)),X1)|r1_xxreal_0(X1,k7_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k8_real_1(k8_real_1(np__2,k32_sin_cos),X2)))|~v1_xreal_0(X1)|~v1_int_1(X2)|~r1_xxreal_0(k20_sin_cos(X1),k6_numbers))).#trainneg
cnf(i_0_444, plain, (r1_xxreal_0(k6_numbers,k20_sin_cos(X1))|~v1_xreal_0(X1)|~v1_int_1(X2)|~r1_xxreal_0(X1,k7_real_1(k10_real_1(k32_sin_cos,np__2),k8_real_1(k8_real_1(np__2,k32_sin_cos),X2)))|~r1_xxreal_0(k7_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k8_real_1(k8_real_1(np__2,k32_sin_cos),X2)),X1))).#trainneg
cnf(i_0_271, plain, (r2_hidden(k1_seq_1(X2,X1),X4)|~v1_funct_1(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r1_fcont_1(X2,X3)|~m1_rcomp_1(X4,k1_seq_1(X2,X3))|~r2_hidden(X1,k1_relset_1(k1_numbers,X2))|~r2_hidden(X1,esk11_3(X2,X3,X4))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_148, plain, (m1_subset_1(esk10_4(X1,X2,X3,X4),k1_numbers)|~v1_funct_1(X3)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X4,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r1_tarski(k1_rcomp_1(X1,X2),k1_relset_1(k1_numbers,X3))|~v1_fcont_1(k2_partfun1(k1_numbers,k1_numbers,X3,k1_rcomp_1(X1,X2)))|~r2_hidden(X4,k4_subset_1(k1_numbers,k1_rcomp_1(k1_seq_1(X3,X1),k1_seq_1(X3,X2)),k1_rcomp_1(k1_seq_1(X3,X2),k1_seq_1(X3,X1)))))).#trainneg
cnf(i_0_146, plain, (X1=k1_seq_1(X2,esk10_4(X3,X4,X2,X1))|~v1_funct_1(X2)|~r1_xxreal_0(X3,X4)|~m1_subset_1(X4,k1_numbers)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r1_tarski(k1_rcomp_1(X3,X4),k1_relset_1(k1_numbers,X2))|~v1_fcont_1(k2_partfun1(k1_numbers,k1_numbers,X2,k1_rcomp_1(X3,X4)))|~r2_hidden(X1,k4_subset_1(k1_numbers,k1_rcomp_1(k1_seq_1(X2,X3),k1_seq_1(X2,X4)),k1_rcomp_1(k1_seq_1(X2,X4),k1_seq_1(X2,X3)))))).#trainneg
cnf(i_0_147, plain, (r2_hidden(esk10_4(X1,X2,X3,X4),k1_rcomp_1(X1,X2))|~v1_funct_1(X3)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X4,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r1_tarski(k1_rcomp_1(X1,X2),k1_relset_1(k1_numbers,X3))|~v1_fcont_1(k2_partfun1(k1_numbers,k1_numbers,X3,k1_rcomp_1(X1,X2)))|~r2_hidden(X4,k4_subset_1(k1_numbers,k1_rcomp_1(k1_seq_1(X3,X1),k1_seq_1(X3,X2)),k1_rcomp_1(k1_seq_1(X3,X2),k1_seq_1(X3,X1)))))).#trainneg
cnf(i_0_1520, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_2440, plain, (m1_subset_1(k2_funct_1(k3_sincos10),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1811, plain, (~v1_xboole_0(k2_relset_1(k1_numbers,k1_sin_cos9)))).#trainneg
cnf(i_0_2084, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k2_xboole_0(X2,X3)))).#trainneg
cnf(i_0_1512, plain, (X1=X2|X1=X3|~r2_hidden(X1,k2_tarski(X2,X3)))).#trainneg
cnf(i_0_1583, negated_conjecture, (m1_subset_1(esk1_0,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_2172, plain, (k5_relat_1(k2_fdiff_9,X1)=k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))).#trainneg
cnf(i_0_621, plain, (np__0=k6_numbers)).#trainneg
cnf(i_0_1423, plain, (r1_xxreal_0(X1,k10_real_1(k32_sin_cos,np__4))|~r2_hidden(X1,a_0_2_sin_cos9))).#trainneg
cnf(i_0_2164, plain, (k5_relat_1(k29_sin_cos,X1)=k2_partfun1(k1_numbers,k1_numbers,k29_sin_cos,X1))).#trainneg
cnf(i_0_2160, plain, (v1_funct_1(k2_funct_1(k3_sincos10)))).#trainneg
cnf(i_0_869, plain, (v1_xreal_0(np__2))).#trainneg
cnf(i_0_682, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2159, plain, (r1_tarski(X1,k2_relset_1(k1_numbers,k3_sincos10))|~r1_tarski(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_719, plain, (k4_xcmplx_0(k1_real_1(np__1))=np__1)).#trainneg
cnf(i_0_2200, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(X2,k1_rcomp_1(X1,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_777, plain, (m1_subset_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)),k1_numbers))).#trainneg
cnf(i_0_1173, plain, (v3_membered(k5_numbers))).#trainneg
cnf(i_0_2204, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(X1,k1_rcomp_1(X3,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_1950, plain, (v1_xxreal_0(k1_seq_1(X1,X2))|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_2167, plain, (k5_relat_1(k19_sin_cos,X1)=k2_partfun1(k1_numbers,k1_numbers,k19_sin_cos,X1))).#trainneg
cnf(i_0_866, plain, (v1_xreal_0(np__4))).#trainneg
cnf(i_0_2061, plain, (r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1812, plain, (m1_subset_1(k10_real_1(k32_sin_cos,np__4),k2_relset_1(k1_numbers,k1_sin_cos9)))).#trainneg
cnf(i_0_2294, plain, (k2_funct_1(k2_funct_1(k3_sincos10))=k3_sincos10|~m1_subset_1(k2_funct_1(k3_sincos10),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(k2_funct_1(k3_sincos10)))).#trainneg
cnf(i_0_2199, plain, (m1_subset_1(k5_relat_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_684, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1814, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k29_sin_cos))|~r1_tarski(X1,k2_relset_1(k1_numbers,k1_sin_cos9)))).#trainneg
cnf(i_0_1582, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_1946, plain, (m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k2_relset_1(k1_numbers,k1_sin_cos9)))).#trainneg
cnf(i_0_618, plain, (esk22_0=k6_numbers)).#trainneg
cnf(i_0_1001, plain, (v7_ordinal1(np__1))).#trainneg
cnf(i_0_295, plain, (r2_hidden(X1,a_2_0_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2169, plain, (k5_relat_1(k16_sin_cos,X1)=k2_partfun1(k1_numbers,k1_numbers,k16_sin_cos,X1))).#trainneg
cnf(i_0_1450, plain, (m1_subset_1(k2_relset_1(k1_numbers,k3_sincos10),k1_zfmisc_1(k1_relset_1(k1_numbers,k2_fdiff_9))))).#trainneg
cnf(i_0_2168, plain, (k5_relat_1(k1_sin_cos9,X1)=k2_partfun1(k1_numbers,k1_numbers,k1_sin_cos9,X1))).#trainneg
cnf(i_0_619, plain, (esk9_0=k6_numbers)).#trainneg
cnf(i_0_1521, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k6_numbers))).#trainneg
cnf(i_0_1813, plain, (m1_subset_1(k2_relset_1(k1_numbers,k1_sin_cos9),k1_zfmisc_1(k1_relset_1(k1_numbers,k29_sin_cos))))).#trainneg
cnf(i_0_780, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,a_0_1_sin_cos9))).#trainneg
cnf(i_0_637, plain, (k2_xboole_0(k6_numbers,X1)=X1)).#trainneg
cnf(i_0_2431, plain, (k2_partfun1(X1,X2,k2_partfun1(X1,X2,X3,X4),X5)=k5_relat_1(k2_partfun1(X1,X2,X3,X4),X5)|~v1_funct_1(k2_partfun1(X1,X2,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1008, plain, (v7_ordinal1(np__2))).#trainneg
cnf(i_0_1784, plain, (k9_subset_1(X1,X3,X2)=k3_xboole_0(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1103, plain, (m1_subset_1(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_zfmisc_1(k1_relset_1(k1_numbers,k2_fdiff_9))))).#trainneg
cnf(i_0_778, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,a_0_0_sin_cos9))).#trainneg
cnf(i_0_853, plain, (v1_xreal_0(np__3))).#trainneg
cnf(i_0_779, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,a_0_2_sin_cos9))).#trainneg
cnf(i_0_1729, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_1576, plain, (m1_subset_1(k2_relset_1(k1_numbers,k3_sincos10),k1_zfmisc_1(k2_rcomp_1(k1_real_1(k32_sin_cos),k6_numbers))))).#trainneg
cnf(i_0_590, plain, (r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),X1)|r2_hidden(X1,a_0_1_sin_cos9)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1570, plain, (~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r2_hidden(X1,a_0_0_sin_cos9))).#trainneg
cnf(i_0_380, plain, (r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),X1)|r2_hidden(X1,a_0_0_sin_cos9)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1553, plain, (r2_hidden(k6_numbers,k1_rcomp_1(k1_real_1(np__1),np__1)))).#trainneg
cnf(i_0_1571, plain, (~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r2_hidden(X1,a_0_1_sin_cos9))).#trainneg
cnf(i_0_631, plain, (v1_xreal_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1937, plain, (r1_tarski(k2_rcomp_1(X1,X2),k1_xxreal_1(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1149, plain, (r2_hidden(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2006, plain, (r1_tarski(X1,k2_relset_1(k1_numbers,k1_sin_cos9))|~r1_tarski(X1,k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_1933, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,X3))|~v4_relat_1(X3,X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_2439, plain, (v4_relat_1(k2_partfun1(X1,X2,X3,X4),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_2778, plain, (v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))|~r1_tarski(X1,k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_1323, plain, (v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_2335, plain, (r2_hidden(X1,X2)|X2!=k1_xxreal_1(X3,X4)|~r2_hidden(X1,k2_tarski(X3,X4))|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~r1_xxreal_0(X3,X4))).#trainneg
cnf(i_0_1152, plain, (r2_hidden(k32_sin_cos,k1_numbers))).#trainneg
cnf(i_0_1951, plain, (v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_2378, plain, (~v1_xreal_0(k10_real_1(k32_sin_cos,np__4))|~r1_xxreal_0(k1_seq_1(k16_sin_cos,k10_real_1(k32_sin_cos,np__4)),k6_numbers))).#trainneg
cnf(i_0_1144, plain, (r2_hidden(np__1,k1_numbers))).#trainneg
cnf(i_0_739, plain, (X1=k6_numbers|~r1_tarski(X1,k6_numbers))).#trainneg
cnf(i_0_729, plain, (k4_xcmplx_0(k1_real_1(np__2))=np__2)).#trainneg
cnf(i_0_1699, plain, (v1_xreal_0(k10_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_633, plain, (v1_xcmplx_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2377, plain, (~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k1_seq_1(k16_sin_cos,k10_real_1(k32_sin_cos,np__4)),k6_numbers))).#trainneg
cnf(i_0_1271, plain, (v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_1941, plain, (v6_xxreal_2(X1)|~r1_tarski(k1_rcomp_1(esk4_1(X1),esk5_1(X1)),X1)|~v2_membered(X1)|~v1_xreal_0(esk5_1(X1))|~v1_xreal_0(esk4_1(X1)))).#trainneg
cnf(i_0_2609, plain, (r1_tarski(k1_rcomp_1(X1,X2),X3)|~v6_xxreal_2(X3)|~v2_membered(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(X2,X3)|~r2_hidden(X1,X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_781, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1706, plain, (r2_hidden(k10_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1428, plain, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),X1)|~r2_hidden(X1,a_0_0_sin_cos9))).#trainneg
cnf(i_0_3676, plain, (v1_xxreal_0(np__2))).#trainneg
cnf(i_0_3547, plain, (v1_xxreal_0(np__1))).#trainneg
cnf(i_0_3520, plain, (X2=X1|~r1_tarski(X1,X2)|~r1_tarski(X2,X1))).#trainneg
cnf(i_0_3751, plain, (k9_xtuple_0(k2_fdiff_9)=k1_relset_1(k1_numbers,k2_fdiff_9))).#trainneg
cnf(i_0_3362, plain, (v1_xreal_0(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_570, plain, (~r1_xxreal_0(np__2,k6_numbers))).#trainneg
cnf(i_0_3969, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_174, plain, (k4_xcmplx_0(k6_numbers)=k6_numbers)).#trainneg
cnf(i_0_3381, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v1_xxreal_0(esk1_0)|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_475, plain, (r1_xxreal_0(k6_numbers,np__1))).#trainneg
cnf(i_0_3874, plain, (k1_relset_1(X1,k2_partfun1(X1,X2,X3,X4))=k9_xtuple_0(k2_partfun1(X1,X2,X3,X4))|~v1_relat_1(k2_partfun1(X1,X2,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_3509, plain, (k2_relset_1(k1_numbers,k2_fdiff_9)=k10_xtuple_0(k2_fdiff_9))).#trainneg
cnf(i_0_3560, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3794, plain, (r2_hidden(k6_numbers,a_0_0_sin_cos9))).#trainneg
cnf(i_0_3558, plain, (r1_xxreal_0(esk15_3(X1,X2,X3),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_556, plain, (r1_xxreal_0(k6_numbers,k6_numbers))).#trainneg
cnf(i_0_3363, plain, (v1_xxreal_0(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_3846, plain, (r1_tarski(k2_rcomp_1(X1,X2),k1_rcomp_1(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3530, plain, (v1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v2_membered(k2_relset_1(k1_numbers,k1_sin_cos9)))).#trainneg
cnf(i_0_547, plain, (r1_xxreal_0(k6_numbers,np__2))).#trainneg
cnf(i_0_3338, plain, (k1_real_1(k1_real_1(np__1))=np__1|~m1_subset_1(k1_real_1(np__1),k1_numbers))).#trainneg
cnf(i_0_3714, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k2_fdiff_9))|~r1_tarski(X1,k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_489, plain, (~r1_xxreal_0(np__1,k6_numbers))).#trainneg
cnf(i_0_3685, plain, (k9_subset_1(X3,X2,X1)=k3_xboole_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_3548, plain, (v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_3926, plain, (k1_real_1(k1_real_1(np__2))=np__2|~m1_subset_1(k1_real_1(np__2),k1_numbers))).#trainneg
cnf(i_0_3825, plain, (~r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__4)),a_0_1_sin_cos9)|~r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__4)),a_0_2_sin_cos9))).#trainneg
cnf(i_0_3678, plain, (v1_int_1(np__2))).#trainneg
cnf(i_0_3421, plain, (r1_xxreal_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(X2)|~r1_xxreal_0(X2,k1_real_1(X1))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3556, plain, (esk15_3(X1,X2,X3)=X1|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3677, plain, (v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_3508, plain, (m1_subset_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k2_fdiff_9,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_3310, plain, (k2_xboole_0(k2_rcomp_1(X1,X2),k2_tarski(X2,X1))=k1_xxreal_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_3807, plain, (esk14_3(X1,X2,X3)=X2|~v1_xxreal_0(X2)|~v1_xxreal_0(esk14_3(X1,X2,X3))|~r1_xxreal_0(esk14_3(X1,X2,X3),X2)|~v1_xxreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k3_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_3835, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers)|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_3305, plain, (r2_hidden(X1,X2)|X2!=k1_xxreal_1(X3,X4)|~r2_hidden(X1,k2_rcomp_1(X3,X4))|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~r1_xxreal_0(X3,X4))).#trainneg
cnf(i_0_3667, plain, (k9_xtuple_0(k2_partfun1(X1,X2,k2_partfun1(X1,X2,X3,X4),X5))=X5|~r1_tarski(X5,k9_xtuple_0(k2_partfun1(X1,X2,X3,X4)))|~v1_relat_1(k2_partfun1(X1,X2,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_3622, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k9_subset_1(X2,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_3524, plain, (r1_tarski(k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_relset_1(k1_numbers,k29_sin_cos)))).#trainneg
cnf(i_0_3549, plain, (v1_int_1(np__1))).#trainneg
cnf(i_0_3968, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_3686, plain, (k9_subset_1(X3,X2,X1)=X1|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_3990, plain, (~r2_hidden(k10_real_1(k32_sin_cos,np__4),a_0_0_sin_cos9)|~r2_hidden(k10_real_1(k32_sin_cos,np__2),a_0_2_sin_cos9))).#trainneg
cnf(i_0_3297, plain, (k9_xtuple_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))=X1|~r1_tarski(X1,k9_xtuple_0(k2_fdiff_9))|~v1_relat_1(k2_fdiff_9))).#trainneg
cnf(i_0_3863, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k29_sin_cos))|~r1_tarski(X1,k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_3450, plain, (v1_xxreal_0(k10_real_1(k32_sin_cos,np__4))|~v2_membered(k2_relset_1(k1_numbers,k1_sin_cos9)))).#trainneg
cnf(i_0_3789, plain, (~r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__4)),a_0_0_sin_cos9)|~r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__2)),a_0_2_sin_cos9))).#trainneg
cnf(i_0_3307, plain, (r2_hidden(X1,k2_rcomp_1(X2,X3))|r2_hidden(X1,k2_tarski(X2,X3))|X4!=k1_xxreal_1(X2,X3)|~r2_hidden(X1,X4)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3))).#trainneg
cnf(i_0_3844, plain, (m1_subset_1(k2_rcomp_1(X1,X2),k1_zfmisc_1(k1_xxreal_1(X1,X2)))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_3559, plain, (r1_xxreal_0(X1,esk15_3(X2,X1,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3973, plain, (k9_subset_1(X1,X2,X2)=X2)).#trainneg
cnf(i_0_3306, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k2_rcomp_1(X2,X3))|~r2_hidden(X1,k1_xxreal_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3))).#trainneg
cnf(i_0_1143, plain, (r2_hidden(np__1,k5_numbers))).#trainneg
cnf(i_0_1681, plain, (m1_subset_1(k2_relset_1(k1_numbers,k3_sincos10),k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(k6_numbers)|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_78, plain, (v1_xcmplx_0(k6_numbers))).#trainneg
cnf(i_0_2436, plain, (v5_relat_1(k2_partfun1(X1,X2,X3,X4),X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_2072, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1005, plain, (v7_ordinal1(np__4))).#trainneg
cnf(i_0_3924, plain, (r1_xxreal_0(k1_real_1(X1),np__2)|~v1_xreal_0(k1_real_1(np__2))|~r1_xxreal_0(k1_real_1(np__2),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1429, plain, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),X1)|~r2_hidden(X1,a_0_1_sin_cos9))).#trainneg
cnf(i_0_2741, plain, (r2_hidden(X1,k1_rcomp_1(k6_numbers,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(k6_numbers)|~r1_xxreal_0(k18_complex1(k5_real_1(k2_xcmplx_0(k6_numbers,X2),k8_real_1(np__2,X1))),X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_614, plain, (k2_xcmplx_0(np__1,k6_numbers)=np__1)).#trainneg
cnf(i_0_118, plain, (~r1_xxreal_0(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)),np__1))).#trainneg
cnf(i_0_2831, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(X3,X2),k1_seq_1(X3,X1))|~v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_1(X3)|~r2_hidden(X1,k9_subset_1(k1_numbers,k1_relset_1(k1_numbers,X3),X4))|~r2_hidden(X2,k9_subset_1(k1_numbers,k1_relset_1(k1_numbers,X3),X4))|~m1_subset_1(k1_relset_1(k1_numbers,X3),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1810, plain, (m1_subset_1(k9_subset_1(X1,X3,X2),k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2338, plain, (k2_xboole_0(k2_tarski(X2,X1),k4_xxreal_1(X1,X2))=k1_xxreal_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_3504, plain, (m1_subset_1(k2_partfun1(k1_numbers,k1_numbers,k29_sin_cos,X1),k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_768, plain, (~v1_xboole_0(np__1)|~v1_xxreal_0(np__1))).#trainneg
cnf(i_0_683, plain, (v1_xxreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_848, plain, (v1_xreal_0(np__5))).#trainneg
cnf(i_0_3529, plain, (v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v3_membered(k2_relset_1(k1_numbers,k1_sin_cos9)))).#trainneg
cnf(i_0_437, plain, (~r1_xxreal_0(np__2,k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_1565, plain, (k9_subset_1(X3,X1,X2)=X1|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_2293, plain, (k4_rfunct_1(k16_sin_cos)=k2_fdiff_9)).#trainneg
cnf(i_0_287, plain, (r1_xxreal_0(X1,X2)|r2_hidden(X2,a_2_2_rcomp_1(X3,X1))|~v1_xxreal_0(X1)|~v1_xreal_0(X3)|~r1_xxreal_0(X3,X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1151, plain, (r2_hidden(np__2,k1_numbers))).#trainneg
cnf(i_0_2607, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_xxreal_1(X3,X4))|~v6_xxreal_2(X2)|~v2_membered(X2)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~r2_hidden(X4,X2)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_3298, plain, (k5_relat_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),X2)=k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X2)|~r1_tarski(X2,X1)|~v1_relat_1(k2_fdiff_9))).#trainneg
cnf(i_0_2677, plain, (r1_xxreal_0(k18_complex1(k5_real_1(k2_xcmplx_0(k6_numbers,X1),k8_real_1(np__2,X2))),X1)|~v1_xreal_0(X1)|~v1_xreal_0(k6_numbers)|~v1_xreal_0(X2)|~r2_hidden(X2,k1_rcomp_1(k6_numbers,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2326, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k6_square_1(X1),k7_square_1(X2))|~r1_xxreal_0(k6_numbers,X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3594, plain, (v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))))).#trainneg
cnf(i_0_3369, plain, (r2_hidden(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)),k1_numbers))).#trainneg
cnf(i_0_3801, plain, (m1_subset_1(k6_numbers,k1_rcomp_1(k1_real_1(np__1),np__1)))).#trainneg
cnf(i_0_2413, plain, (v1_fcont_1(k2_funct_1(k3_sincos10))|~r2_fdiff_1(k2_fdiff_9,k2_relset_1(k1_numbers,k3_sincos10)))).#trainneg
cnf(i_0_695, plain, (k5_xcmplx_0(k7_xcmplx_0(np__1,X1))=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3261, plain, (esk13_3(X1,X2,k2_tarski(X3,X4))=X4|esk13_3(X1,X2,k2_tarski(X3,X4))=X3|r1_fcont_1(X1,X2)|~m1_rcomp_1(k2_tarski(X3,X4),X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1))).#trainneg
cnf(i_0_3296, plain, (k10_xtuple_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))=k7_relat_1(k2_fdiff_9,X1)|~v1_relat_1(k2_fdiff_9))).#trainneg
cnf(i_0_3935, plain, (k4_xcmplx_0(np__2)=k1_real_1(np__2)|~v1_xcmplx_0(k1_real_1(np__2)))).#trainneg
cnf(i_0_2528, plain, (k2_xboole_0(X6,X7)=X6|r2_hidden(esk6_3(X6,X7,X6),X7)|r2_hidden(esk6_3(X6,X7,X6),X6))).#trainneg
cnf(i_0_3796, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k10_real_1(k32_sin_cos,np__2))|r2_hidden(k10_real_1(k32_sin_cos,np__2),a_0_0_sin_cos9)|~m1_subset_1(k10_real_1(k32_sin_cos,np__2),k1_numbers))).#trainneg
cnf(i_0_846, plain, (v1_xreal_0(np__7))).#trainneg
cnf(i_0_461, plain, (~r1_xxreal_0(k6_numbers,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_3725, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~v1_xreal_0(X1)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3688, plain, (k9_subset_1(X3,X2,X1)=X2|~r1_tarski(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_843, plain, (v1_xreal_0(np__15))).#trainneg
cnf(i_0_1558, plain, (r2_hidden(k1_seq_1(k16_sin_cos,k10_real_1(k32_sin_cos,np__4)),k1_rcomp_1(k1_real_1(np__1),np__1))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_3437, plain, (r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k1_real_1(X2),X1)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3347, plain, (k4_xcmplx_0(np__1)=k1_real_1(np__1)|~v1_xcmplx_0(k1_real_1(np__1)))).#trainneg
cnf(i_0_3862, plain, (r1_tarski(X1,k2_relset_1(k1_numbers,k1_sin_cos9))|~r1_tarski(X1,X2)|~r1_tarski(X2,k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_423, plain, (k2_xcmplx_0(k6_numbers,np__1)=np__1)).#trainneg
cnf(i_0_3419, plain, (r1_xxreal_0(k1_real_1(X1),np__1)|~v1_xreal_0(k1_real_1(np__1))|~r1_xxreal_0(k1_real_1(np__1),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3729, plain, (m1_subset_1(k1_real_1(np__1),k1_rcomp_1(k1_real_1(np__1),np__1)))).#trainneg
cnf(i_0_837, plain, (v1_xreal_0(esk23_1(X1))|~r2_hidden(X1,a_0_2_sin_cos9))).#trainneg
cnf(i_0_1807, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k9_subset_1(X2,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_3591, plain, (m1_subset_1(k2_partfun1(k1_numbers,k1_numbers,k1_sin_cos9,X1),k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_3241, plain, (r2_hidden(X1,k4_xxreal_1(X2,X3))|r2_hidden(X1,k2_tarski(X2,X3))|~r2_hidden(X1,k1_xxreal_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3))).#trainneg
cnf(i_0_3396, plain, (r1_xxreal_0(k6_numbers,k1_seq_1(k16_sin_cos,X1))|~v1_xreal_0(X1)|~v1_xxreal_0(X1)|~v1_xreal_0(k6_numbers)|~r1_xxreal_0(X1,k32_sin_cos)|~r1_xxreal_0(k6_numbers,X1))).#trainneg
cnf(i_0_3728, plain, (~v1_xboole_0(k1_rcomp_1(k1_real_1(np__1),np__1)))).#trainneg
cnf(i_0_3937, plain, (r1_xxreal_0(np__2,k1_real_1(X1))|~v1_xreal_0(k1_real_1(np__2))|~r1_xxreal_0(X1,k1_real_1(np__2))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2190, plain, (k2_partfun1(X1,X2,k5_relset_1(X1,X2,X3,X4),X5)=k5_relat_1(k5_relset_1(X1,X2,X3,X4),X5)|~v1_funct_1(k5_relset_1(X1,X2,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2471, plain, (k2_partfun1(X1,k1_numbers,k6_rfunct_1(X1,X2,X3),X4)=k5_relat_1(k6_rfunct_1(X1,X2,X3),X4)|~v1_funct_1(k6_rfunct_1(X1,X2,X3))|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1716, plain, (r2_hidden(k8_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3845, plain, (r1_tarski(X1,k1_xxreal_1(X2,X3))|~r1_tarski(X1,k2_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_4537, plain, (v1_xxreal_0(np__4))).#trainneg
cnf(i_0_4550, plain, (~r2_hidden(k10_real_1(k32_sin_cos,np__4),a_0_1_sin_cos9)|~r2_hidden(k10_real_1(k32_sin_cos,np__4),a_0_2_sin_cos9))).#trainneg
cnf(i_0_4338, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3)|~r2_hidden(X2,k2_rcomp_1(X1,X3)))).#trainneg
cnf(i_0_4202, plain, (v1_xxreal_0(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_4276, plain, (r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(k6_numbers,k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4976, plain, (v1_xreal_0(X1)|~r2_hidden(X1,a_0_2_sin_cos9))).#trainneg
cnf(i_0_3996, plain, (r1_xxreal_0(np__2,np__2))).#trainneg
cnf(i_0_4067, plain, (m1_subset_1(X1,k2_tarski(X2,X1)))).#trainneg
cnf(i_0_4290, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3)|~r2_hidden(X2,k2_tarski(X1,X3)))).#trainneg
cnf(i_0_4109, plain, (k1_real_1(k6_numbers)=k6_numbers)).#trainneg
cnf(i_0_4770, plain, (k4_xcmplx_0(np__2)=k1_real_1(np__2)|~v1_xreal_0(k1_real_1(np__2)))).#trainneg
cnf(i_0_4177, plain, (m1_subset_1(k6_numbers,a_0_0_sin_cos9))).#trainneg
cnf(i_0_4063, plain, (m1_subset_1(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))),k1_zfmisc_1(k1_relset_1(k1_numbers,k2_fdiff_9))))).#trainneg
cnf(i_0_4702, plain, (k2_partfun1(k1_numbers,k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),X2)=k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X2)|~r1_tarski(X2,X1))).#trainneg
cnf(i_0_3993, plain, (r2_hidden(X1,X2)|X2!=k1_xxreal_1(X1,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3))).#trainneg
cnf(i_0_4065, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k2_fdiff_9))|~r1_tarski(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_4658, plain, (v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(k2_relset_1(k1_numbers,k1_sin_cos9),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4339, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X3,X2)|~r2_hidden(X1,k2_rcomp_1(X3,X2)))).#trainneg
cnf(i_0_4739, plain, (k5_xcmplx_0(np__1)=np__1)).#trainneg
cnf(i_0_4455, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(X3,X2),k1_seq_1(X3,X1))|~v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,k1_relset_1(k1_numbers,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_1(X3)|~r2_hidden(X1,k1_relset_1(k1_numbers,X3))|~r2_hidden(X2,k1_relset_1(k1_numbers,X3)))).#trainneg
cnf(i_0_4003, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_4135, plain, (r1_tarski(k1_relset_1(X1,X2),X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_4068, plain, (r2_hidden(X1,X2)|X2!=k1_xxreal_1(X3,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X3)|~r1_xxreal_0(X3,X1))).#trainneg
cnf(i_0_4873, plain, (~r1_xxreal_0(k6_numbers,k1_real_1(np__1)))).#trainneg
cnf(i_0_4059, plain, (X1=X2|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4041, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4612, plain, (m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))|~r1_tarski(X3,X2))).#trainneg
cnf(i_0_4769, plain, (k4_xcmplx_0(np__2)=k1_real_1(np__2)|~m1_subset_1(k1_real_1(np__2),k1_numbers))).#trainneg
cnf(i_0_4054, plain, (k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))=k2_relset_1(k1_numbers,k3_sincos10)|~r1_tarski(k2_relset_1(k1_numbers,k3_sincos10),k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_3992, plain, (m1_subset_1(X1,k2_tarski(X1,X2)))).#trainneg
cnf(i_0_4040, plain, (r1_xxreal_0(k32_sin_cos,k32_sin_cos))).#trainneg
cnf(i_0_4684, plain, (r1_xxreal_0(X1,esk14_3(X2,X1,X3))|r1_xxreal_0(X3,X2)|~v1_xxreal_0(X3)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4291, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X3,X2)|~r2_hidden(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_4443, plain, (m1_subset_1(k2_rcomp_1(X1,X2),k1_zfmisc_1(k1_rcomp_1(X1,X2)))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4930, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k1_real_1(k1_real_1(X1)),X2)|~m1_subset_1(k1_real_1(X1),k1_numbers))).#trainneg
cnf(i_0_4376, plain, (m1_subset_1(k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_zfmisc_1(k1_relset_1(k1_numbers,k29_sin_cos))))).#trainneg
cnf(i_0_4049, plain, (k1_xxreal_1(X1,X2)=k2_rcomp_1(X1,X2)|~r1_tarski(k1_xxreal_1(X1,X2),k2_rcomp_1(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_4207, plain, (r1_tarski(k2_relset_1(k1_numbers,k1_sin_cos9),k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)))|~v1_xreal_0(k10_real_1(k32_sin_cos,np__2))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_4107, plain, (r1_xxreal_0(k1_real_1(X1),k6_numbers)|~r1_xxreal_0(k6_numbers,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4051, plain, (k2_relset_1(k1_numbers,k1_sin_cos9)=k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))|~r1_tarski(k2_relset_1(k1_numbers,k1_sin_cos9),k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_4232, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k2_fdiff_9))|~m1_subset_1(X1,k1_zfmisc_1(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))))).#trainneg
cnf(i_0_4341, plain, (X1=X2|~v1_xxreal_0(X1)|~v1_xxreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)|~r2_hidden(X1,k3_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_4689, plain, (r1_tarski(k2_rcomp_1(X1,X2),X3)|~v6_xxreal_2(X3)|~v2_membered(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(X2,X3)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_4690, plain, (r1_tarski(k4_xxreal_1(X1,X2),X3)|~v6_xxreal_2(X3)|~v2_membered(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(X2,X3)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_4399, plain, (k9_xtuple_0(k2_funct_1(k3_sincos10))=k2_relset_1(k1_numbers,k3_sincos10))).#trainneg
cnf(i_0_4418, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k2_rcomp_1(X2,X3))|X4!=k1_rcomp_1(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3)|~r2_hidden(X1,X4)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4403, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k29_sin_cos))|~m1_subset_1(X1,k1_zfmisc_1(k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))))).#trainneg
cnf(i_0_4481, plain, (r2_hidden(X1,k2_rcomp_1(X2,X3))|r2_hidden(X1,k2_tarski(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X2,X1))).#trainneg
cnf(i_0_4347, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers)|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_numbers))).#trainneg
cnf(i_0_4569, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(X3,X2),k1_seq_1(X3,X1))|~v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_1(X3)|~r2_hidden(X1,X4)|~r2_hidden(X2,X4)|~r1_tarski(X4,k1_relset_1(k1_numbers,X3)))).#trainneg
cnf(i_0_4538, plain, (v1_xcmplx_0(np__4))).#trainneg
cnf(i_0_4255, plain, (r1_xxreal_0(X1,k1_real_1(np__2))|~r1_xxreal_0(np__2,k1_real_1(X1))|~m1_subset_1(k1_real_1(np__2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4375, plain, (r1_tarski(k9_subset_1(X1,X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4522, plain, (r1_xxreal_0(X1,k1_real_1(np__2))|~v1_xreal_0(k1_real_1(np__2))|~r1_xxreal_0(np__2,k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3997, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_4211, plain, (v1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v3_membered(k2_relset_1(k1_numbers,k1_sin_cos9)))).#trainneg
cnf(i_0_4349, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers)|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~m1_subset_1(k10_real_1(k32_sin_cos,np__4),k1_numbers))).#trainneg
cnf(i_0_4506, plain, (k2_relset_1(X1,k2_partfun1(X2,X1,X3,X4))=k10_xtuple_0(k2_partfun1(X2,X1,X3,X4))|~v1_relat_1(k2_partfun1(X2,X1,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_4154, plain, (r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),np__1)|~v1_xreal_0(k1_real_1(np__1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4000, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_4228, plain, (r1_tarski(k2_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_relset_1(k1_numbers,k2_fdiff_9)))).#trainneg
cnf(i_0_5044, plain, (r1_xxreal_0(np__2,k1_real_1(X1))|~r1_xxreal_0(X1,k1_real_1(np__2))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4359, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),X1)|X1!=k1_xxreal_1(k6_numbers,k32_sin_cos)|~v1_xxreal_0(k32_sin_cos)|~r1_xxreal_0(k6_numbers,k32_sin_cos))).#trainneg
cnf(i_0_4472, plain, (r2_hidden(X1,k2_rcomp_1(X2,X3))|r2_hidden(X1,k2_tarski(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3)|~r2_hidden(X1,k1_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4539, plain, (v1_int_1(np__4))).#trainneg
cnf(i_0_4402, plain, (r1_tarski(k1_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__4)),k1_relset_1(k1_numbers,k29_sin_cos)))).#trainneg
cnf(i_0_4682, plain, (esk14_3(X1,X2,X3)=X1|r1_xxreal_0(X3,X1)|~v1_xxreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4949, plain, (r1_tarski(X1,k2_relset_1(k1_numbers,k1_sin_cos9))|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))))).#trainneg
cnf(i_0_4204, plain, (k2_rcomp_1(X1,X2)=k1_rcomp_1(X1,X2)|~r1_tarski(k1_rcomp_1(X1,X2),k2_rcomp_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4231, plain, (r1_tarski(k1_rcomp_1(k10_real_1(k32_sin_cos,np__4),k10_real_1(k32_sin_cos,np__2)),k1_relset_1(k1_numbers,k2_fdiff_9)))).#trainneg
cnf(i_0_4205, plain, (r1_tarski(X1,k1_rcomp_1(X2,X3))|~r1_tarski(X1,k2_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4189, plain, (r1_xxreal_0(k6_numbers,k10_real_1(k32_sin_cos,np__4))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_4095, plain, (r1_xxreal_0(k1_real_1(np__1),k6_numbers)|~v1_xreal_0(k1_real_1(np__1)))).#trainneg
cnf(i_0_4717, plain, (v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))|~r1_tarski(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_4892, plain, (r1_xxreal_0(X1,k1_real_1(k1_real_1(X2)))|~v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4740, plain, (k5_xcmplx_0(k10_real_1(np__1,X1))=X1|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4021, plain, (r1_xxreal_0(X1,k32_sin_cos)|r1_xxreal_0(k32_sin_cos,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4022, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4044, plain, (k2_relset_1(k1_numbers,k3_sincos10)=k1_relset_1(k1_numbers,k2_fdiff_9)|~r1_tarski(k1_relset_1(k1_numbers,k2_fdiff_9),k2_relset_1(k1_numbers,k3_sincos10)))).#trainneg
cnf(i_0_4423, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k2_fdiff_9))|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10))))).#trainneg
cnf(i_0_3998, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4944, plain, (k4_xcmplx_0(np__1)=k1_real_1(np__1)|~v1_xreal_0(k1_real_1(np__1)))).#trainneg
cnf(i_0_4074, plain, (~r2_hidden(k1_real_1(X1),a_0_1_sin_cos9)|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),X1)|~m1_subset_1(k10_real_1(k32_sin_cos,np__4),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4240, plain, (k9_subset_1(X1,X2,X3)=k9_subset_1(X4,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X4)))).#trainneg
cnf(i_0_4064, plain, (k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))=k1_relset_1(k1_numbers,k2_fdiff_9)|~r1_tarski(k1_relset_1(k1_numbers,k2_fdiff_9),k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_4948, plain, (r1_tarski(X1,k2_relset_1(k1_numbers,k1_sin_cos9))|~r1_tarski(X1,k1_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_4960, plain, (r2_hidden(X1,k1_rcomp_1(k6_numbers,np__1))|~v1_xreal_0(X1)|~r1_xxreal_0(k18_complex1(k5_real_1(np__1,k8_real_1(np__2,X1))),np__1))).#trainneg
cnf(i_0_4398, plain, (k1_relset_1(k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))=X1|~v1_relat_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))|~r1_tarski(X1,k1_relset_1(k1_numbers,k2_fdiff_9)))).#trainneg
cnf(i_0_4471, plain, (esk15_3(X1,X2,X2)=X2|~v1_xxreal_0(esk15_3(X1,X2,X2))|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X2)))).#trainneg
cnf(i_0_4416, plain, (v1_xxreal_0(k10_real_1(k32_sin_cos,np__4))|~v3_membered(k2_relset_1(k1_numbers,k1_sin_cos9)))).#trainneg
cnf(i_0_4214, plain, (v1_xcmplx_0(k1_seq_1(k2_fdiff_9,X1)))).#trainneg
cnf(i_0_4108, plain, (v3_xxreal_0(k6_numbers)|~v2_xxreal_0(k6_numbers))).#trainneg
cnf(i_0_4001, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4421, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k2_fdiff_9)))|~m1_subset_1(X1,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10))))).#trainneg
cnf(i_0_5051, plain, (m1_subset_1(k2_partfun1(X1,X2,k5_relset_1(X1,X2,X3,X4),X5),k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(k5_relset_1(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(k5_relset_1(X1,X2,X3,X4)))).#trainneg
cnf(i_0_4134, plain, (r1_tarski(X1,X2)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(X2,X3))))).#trainneg
cnf(i_0_4244, plain, (k3_xboole_0(X2,X3)=k9_subset_1(X1,X3,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4941, plain, (r1_xxreal_0(np__2,X1)|~v1_xreal_0(k1_real_1(np__2))|~r1_xxreal_0(k1_real_1(X1),k1_real_1(np__2))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4278, plain, (r1_xxreal_0(X1,np__2)|~v1_xreal_0(k1_real_1(np__2))|~r1_xxreal_0(k1_real_1(np__2),k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5253, plain, (r2_hidden(X1,k1_xxreal_1(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_3338, plain, (k1_real_1(k1_real_1(np__1))=np__1)).#trainneg
cnf(i_0_5396, plain, (~r2_hidden(k10_real_1(k32_sin_cos,np__4),a_0_1_sin_cos9)|~m1_subset_1(k10_real_1(k32_sin_cos,np__4),k1_numbers))).#trainneg
cnf(i_0_5254, plain, (r2_hidden(X1,X2)|X2!=k1_rcomp_1(X1,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5347, plain, (r2_hidden(X1,k1_xxreal_1(X2,X1))|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X1))).#trainneg
cnf(i_0_6096, plain, (~v2_xxreal_0(k6_numbers))).#trainneg
cnf(i_0_5395, plain, (r1_xxreal_0(k1_real_1(k1_real_1(X1)),X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers))).#trainneg
cnf(i_0_5212, plain, (k4_xcmplx_0(np__2)=k1_real_1(np__2))).#trainneg
cnf(i_0_5450, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1))|~m1_subset_1(k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))|~r2_hidden(X2,k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_5454, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1))|~m1_subset_1(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))),k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_5348, plain, (r2_hidden(X1,X2)|X2!=k1_rcomp_1(X3,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X3)|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_5397, plain, (~r2_hidden(k10_real_1(k32_sin_cos,np__2),a_0_0_sin_cos9)|~m1_subset_1(k10_real_1(k32_sin_cos,np__2),k1_numbers))).#trainneg
cnf(i_0_5451, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1))|~m1_subset_1(k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers),k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))|~r2_hidden(X2,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))).#trainneg
cnf(i_0_6202, plain, (~v1_xreal_0(k1_real_1(np__2))|~r1_xxreal_0(k6_numbers,k1_real_1(np__2)))).#trainneg
cnf(i_0_5761, plain, (~v1_xxreal_0(k32_sin_cos)|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k6_numbers))).#trainneg
cnf(i_0_5401, plain, (r1_xxreal_0(X1,k1_real_1(k1_real_1(X1)))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5115, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~m1_subset_1(k10_real_1(k32_sin_cos,np__4),k1_numbers))).#trainneg
cnf(i_0_6093, plain, (v1_xcmplx_0(k1_real_1(np__1)))).#trainneg
cnf(i_0_5848, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k29_sin_cos))|~r1_tarski(X1,k1_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_5709, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1))|~m1_subset_1(k1_relset_1(k1_numbers,k2_fdiff_9),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))|~r2_hidden(X2,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))).#trainneg
cnf(i_0_5417, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r1_tarski(X1,k19_sin_cos))).#trainneg
cnf(i_0_5433, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10)))|~r1_tarski(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_5432, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r1_tarski(X1,k2_fdiff_9))).#trainneg
cnf(i_0_6088, plain, (X1=X2|~v1_xxreal_0(X1)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X2)))).#trainneg
cnf(i_0_5698, plain, (r2_hidden(X1,X2)|r2_hidden(X1,k2_rcomp_1(X3,X4))|X2!=k1_xxreal_1(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X4)|~r1_xxreal_0(X3,X1))).#trainneg
cnf(i_0_5828, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),k1_xxreal_1(k6_numbers,k32_sin_cos))|~v1_xxreal_0(k32_sin_cos)|~r1_xxreal_0(k6_numbers,k32_sin_cos))).#trainneg
cnf(i_0_5610, plain, (~r1_xxreal_0(k6_numbers,k1_real_1(np__2))|~m1_subset_1(k1_real_1(np__2),k1_numbers))).#trainneg
cnf(i_0_5502, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X3,X2)|~r2_hidden(X1,k2_tarski(X2,X3)))).#trainneg
cnf(i_0_5399, plain, (~r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__2)),a_0_0_sin_cos9)|~m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_numbers))).#trainneg
cnf(i_0_5648, plain, (r1_tarski(k2_relset_1(k1_numbers,k1_sin_cos9),X1)|~v6_xxreal_2(X1)|~v2_membered(X1)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r2_hidden(k10_real_1(k32_sin_cos,np__2),X1)|~r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1))).#trainneg
cnf(i_0_5434, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k2_fdiff_9)))|~r1_tarski(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_5883, plain, (r1_xxreal_0(k6_numbers,k10_real_1(k32_sin_cos,np__4))|~v1_xreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_5799, plain, (r1_xxreal_0(np__1,k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_5420, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k2_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_5541, plain, (r1_xxreal_0(X1,esk15_3(X2,k1_real_1(k1_real_1(X1)),X3))|~v1_xreal_0(esk15_3(X2,k1_real_1(k1_real_1(X1)),X3))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X3)|~v1_xreal_0(k1_real_1(k1_real_1(X1)))|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k1_real_1(X1)),X3)))).#trainneg
cnf(i_0_5435, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_5412, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k2_fdiff_9)))|~r1_tarski(X1,k2_relset_1(k1_numbers,k3_sincos10)))).#trainneg
cnf(i_0_5416, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k2_fdiff_9)))|~r1_tarski(X1,k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_5647, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k2_rcomp_1(X3,X4))|~v6_xxreal_2(X2)|~v2_membered(X2)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~r2_hidden(X4,X2)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_5290, plain, (X1=X2|~v1_xxreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5769, plain, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)),np__2))).#trainneg
cnf(i_0_5291, plain, (X1=X2|~v1_xxreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X2,X1))).#trainneg
cnf(i_0_5414, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k29_sin_cos)))|~r1_tarski(X1,k2_relset_1(k1_numbers,k1_sin_cos9)))).#trainneg
cnf(i_0_5552, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X3)|~r2_hidden(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_5400, plain, (~r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__4)),a_0_1_sin_cos9)|~m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_numbers))).#trainneg
cnf(i_0_5829, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),X1)|X1!=k1_rcomp_1(k6_numbers,k32_sin_cos)|~v1_xxreal_0(k32_sin_cos)|~r1_xxreal_0(k6_numbers,k32_sin_cos))).#trainneg
cnf(i_0_5681, plain, (k2_relset_1(k1_numbers,k3_sincos10)=k10_xtuple_0(k3_sincos10)|~v2_funct_1(k3_sincos10)|~v1_relat_1(k3_sincos10)|~v1_funct_1(k3_sincos10))).#trainneg
cnf(i_0_5271, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k32_sin_cos)|~v1_xxreal_0(k32_sin_cos)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__4))|~r1_xxreal_0(k6_numbers,k32_sin_cos))).#trainneg
cnf(i_0_5596, plain, (k1_xxreal_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2))=k2_relset_1(k1_numbers,k1_sin_cos9)|~r1_tarski(k1_xxreal_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)),k2_relset_1(k1_numbers,k1_sin_cos9))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_5192, plain, (r2_hidden(k1_seq_1(k2_fdiff_9,X1),k1_numbers))).#trainneg
cnf(i_0_5424, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k5_numbers))).#trainneg
cnf(i_0_5418, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r1_tarski(X1,k1_sin_cos9))).#trainneg
cnf(i_0_5629, plain, (k10_real_1(k32_sin_cos,np__4)=k6_numbers|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__4))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k6_numbers))).#trainneg
cnf(i_0_5770, plain, (r1_xxreal_0(np__2,np__4))).#trainneg
cnf(i_0_5980, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_5353, plain, (r1_xxreal_0(k1_real_1(esk15_3(X1,k1_real_1(X2),X3)),X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(esk15_3(X1,k1_real_1(X2),X3),k1_numbers)|~v1_xreal_0(X3)|~v1_xreal_0(k1_real_1(X2))|~r2_hidden(X1,k1_rcomp_1(k1_real_1(X2),X3)))).#trainneg
cnf(i_0_5807, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k2_fdiff_9))|~r1_tarski(X1,k2_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_5603, plain, (r1_tarski(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)))|~r1_tarski(X1,k2_relset_1(k1_numbers,k1_sin_cos9))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~v1_xreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_6070, plain, (r1_tarski(X1,k2_relset_1(k1_numbers,k1_sin_cos9))|~r1_tarski(X1,X2)|~r1_tarski(X2,k1_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_5583, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k29_sin_cos)))|~r1_tarski(X1,k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_5199, plain, (r2_hidden(k6_numbers,k1_numbers))).#trainneg
cnf(i_0_5231, plain, (~v1_xboole_0(k1_relset_1(k1_numbers,k2_fdiff_9))|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_5605, plain, (k1_real_1(X1)=k6_numbers|~v1_xxreal_0(k1_real_1(X1))|~r1_xxreal_0(k6_numbers,k1_real_1(X1))|~r1_xxreal_0(k6_numbers,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5618, plain, (X1=k1_real_1(k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(X1)|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r2_hidden(X1,k2_relset_1(k1_numbers,k3_sincos10)))).#trainneg
cnf(i_0_5237, plain, (m1_subset_1(X1,k1_relset_1(k1_numbers,k2_fdiff_9))|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_6019, plain, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(X1))|~m1_subset_1(k10_real_1(k32_sin_cos,np__4),k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~r2_hidden(X1,a_0_1_sin_cos9)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5841, plain, (r2_hidden(X1,k2_tarski(X3,X2))|r2_hidden(X1,k2_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X3)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_6089, plain, (esk15_3(X1,X2,X2)=X2|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X2))|~v1_xreal_0(esk15_3(X1,X2,X2)))).#trainneg
cnf(i_0_5867, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k2_fdiff_9))|~r1_tarski(X1,k1_rcomp_1(k10_real_1(k32_sin_cos,np__4),k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_5688, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k29_sin_cos))|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))))).#trainneg
cnf(i_0_5838, plain, (r2_hidden(X1,X2)|r2_hidden(X1,k2_rcomp_1(X3,X4))|X2!=k1_xxreal_1(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r2_hidden(X1,k1_rcomp_1(X3,X4)))).#trainneg
cnf(i_0_5240, plain, (m1_subset_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))|~r1_tarski(X2,X1))).#trainneg
cnf(i_0_5127, plain, (r1_xxreal_0(k6_numbers,k10_real_1(k32_sin_cos,np__4))|~v1_xxreal_0(k32_sin_cos)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__4))|~r1_xxreal_0(k6_numbers,k32_sin_cos))).#trainneg
cnf(i_0_5315, plain, (r2_hidden(k1_real_1(np__1),k1_numbers))).#trainneg
cnf(i_0_5845, plain, (m1_subset_1(k1_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__4)),k1_zfmisc_1(k1_relset_1(k1_numbers,k29_sin_cos))))).#trainneg
cnf(i_0_5533, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_rcomp_1(X2,X3)))|~r1_tarski(X1,k2_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_5881, plain, (r1_tarski(X1,k1_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_rcomp_1(X2,X3))))).#trainneg
cnf(i_0_5286, plain, (k1_seq_1(k19_sin_cos,X1)=k1_real_1(np__1)|~v1_xxreal_0(k1_seq_1(k19_sin_cos,X1))|~v1_xreal_0(k1_real_1(np__1))|~r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),k1_real_1(np__1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6201, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k2_relset_1(k1_numbers,k3_sincos10)))).#trainneg
cnf(i_0_5431, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r1_tarski(X1,k2_partfun1(X2,X3,X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_4885, plain, (r1_xxreal_0(k1_real_1(np__2),k1_real_1(X1))|~v1_xreal_0(k1_real_1(np__2))|~r1_xxreal_0(X1,np__2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5252, plain, (k2_partfun1(k1_numbers,k1_numbers,k2_funct_1(k3_sincos10),X1)=k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1)|~r1_tarski(X1,k2_relset_1(k1_numbers,k3_sincos10)))).#trainneg
cnf(i_0_5410, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r1_tarski(X1,k5_relset_1(X2,X3,X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_6383, plain, (v1_xreal_0(k1_real_1(np__2)))).#trainneg
cnf(i_0_7058, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))).#trainneg
cnf(i_0_6421, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_seq_1(k2_fdiff_9,esk1_0))|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_6434, plain, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k6_numbers))).#trainneg
cnf(i_0_6682, plain, (m1_subset_1(k2_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6272, plain, (r1_xxreal_0(k1_real_1(np__1),k6_numbers))).#trainneg
cnf(i_0_7181, plain, (r1_xxreal_0(k6_numbers,k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_6548, plain, (v1_relat_1(X1)|~r1_tarski(X1,k2_fdiff_9))).#trainneg
cnf(i_0_6460, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(k10_real_1(k32_sin_cos,np__4),k1_numbers)|~m1_subset_1(k10_real_1(k32_sin_cos,np__2),k1_numbers))).#trainneg
cnf(i_0_7056, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_6429, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_numbers)|~m1_subset_1(k10_real_1(k32_sin_cos,np__2),k1_numbers))).#trainneg
cnf(i_0_5828, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),k1_xxreal_1(k6_numbers,k32_sin_cos))|~v1_xxreal_0(k32_sin_cos))).#trainneg
cnf(i_0_6583, plain, (r2_hidden(X1,X2)|X2!=k1_rcomp_1(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X1,X4)|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_6430, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1))|~r2_hidden(X1,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))|~r2_hidden(X2,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))|~v1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_7157, plain, (X1=X2|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X2)))).#trainneg
cnf(i_0_6292, plain, (r1_xxreal_0(k6_numbers,k1_seq_1(k16_sin_cos,k6_numbers)))).#trainneg
cnf(i_0_6240, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k2_relset_1(k1_numbers,k3_sincos10)))).#trainneg
cnf(i_0_6409, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),X1)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)))|~r2_hidden(X1,k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_6458, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(k10_real_1(k32_sin_cos,np__4),k1_numbers)|~m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_numbers))).#trainneg
cnf(i_0_7173, plain, (r2_hidden(X1,k2_rcomp_1(X2,X3))|r2_hidden(X1,k1_xxreal_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X3)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_6692, plain, (m1_subset_1(k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers),k1_zfmisc_1(k1_relset_1(k1_numbers,k2_fdiff_9))))).#trainneg
cnf(i_0_6384, plain, (v1_xcmplx_0(k1_real_1(np__2)))).#trainneg
cnf(i_0_6258, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_6342, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_4944, plain, (k4_xcmplx_0(np__1)=k1_real_1(np__1))).#trainneg
cnf(i_0_7175, plain, (m1_subset_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r1_tarski(X1,k2_relset_1(k1_numbers,k3_sincos10)))).#trainneg
cnf(i_0_5271, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k32_sin_cos)|~v1_xxreal_0(k32_sin_cos))).#trainneg
cnf(i_0_6481, plain, (r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r1_xxreal_0(k1_real_1(np__1),k1_seq_1(k2_fdiff_9,X1))|~r2_hidden(X1,k2_relset_1(k1_numbers,k3_sincos10)))).#trainneg
cnf(i_0_6336, plain, (v3_membered(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_5829, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),X1)|X1!=k1_rcomp_1(k6_numbers,k32_sin_cos)|~v1_xxreal_0(k32_sin_cos))).#trainneg
cnf(i_0_7119, negated_conjecture, (m1_subset_1(esk1_0,k1_relset_1(k1_numbers,k2_fdiff_9)))).#trainneg
cnf(i_0_6484, plain, (r1_xxreal_0(X1,X1)|~r2_hidden(X1,k2_relset_1(k1_numbers,k3_sincos10)))).#trainneg
cnf(i_0_6562, plain, (X1=X2|~v1_xxreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1))).#trainneg
cnf(i_0_7060, plain, (r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_6234, plain, (v3_membered(k2_relset_1(k1_numbers,k3_sincos10)))).#trainneg
cnf(i_0_6542, plain, (k5_relat_1(X1,X2)=k2_partfun1(k1_numbers,k1_numbers,X1,X2)|~v1_funct_1(X1)|~r1_tarski(X1,k2_fdiff_9))).#trainneg
cnf(i_0_6482, plain, (r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),k1_seq_1(k2_fdiff_9,X1))|~r2_hidden(X1,k2_relset_1(k1_numbers,k3_sincos10)))).#trainneg
cnf(i_0_7061, plain, (r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6580, plain, (r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_6511, plain, (v1_relat_1(X1)|~r1_tarski(X1,k19_sin_cos))).#trainneg
cnf(i_0_5648, plain, (r1_tarski(k2_relset_1(k1_numbers,k1_sin_cos9),X1)|~v6_xxreal_2(X1)|~v2_membered(X1)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~r2_hidden(k10_real_1(k32_sin_cos,np__2),X1))).#trainneg
cnf(i_0_6294, plain, (r1_xxreal_0(k10_real_1(np__1,np__2),k1_seq_1(k19_sin_cos,k6_numbers)))).#trainneg
cnf(i_0_4359, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),X1)|X1!=k1_xxreal_1(k6_numbers,k32_sin_cos)|~v1_xxreal_0(k32_sin_cos))).#trainneg
cnf(i_0_7050, plain, (k1_real_1(k10_real_1(k32_sin_cos,np__4))=k1_real_1(k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_6581, plain, (r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5596, plain, (k1_xxreal_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2))=k2_relset_1(k1_numbers,k1_sin_cos9)|~r1_tarski(k1_xxreal_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)),k2_relset_1(k1_numbers,k1_sin_cos9))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_7084, negated_conjecture, (~v1_xboole_0(k1_relset_1(k1_numbers,k2_fdiff_9)))).#trainneg
cnf(i_0_7277, plain, (k2_partfun1(k1_numbers,k1_numbers,k2_funct_1(k3_sincos10),k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))=k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))).#trainneg
cnf(i_0_286, plain, (k6_xcmplx_0(np__2,np__4)=k1_real_1(np__2))).#trainneg
cnf(i_0_6615, plain, (r1_xxreal_0(X1,X2)|r2_hidden(X1,k2_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X3)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_6587, plain, (r1_xxreal_0(np__2,k10_real_1(k32_sin_cos,np__4))|r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),np__2))).#trainneg
cnf(i_0_6878, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|r2_hidden(X1,k2_rcomp_1(X3,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X3,X2)|~r2_hidden(X1,k1_rcomp_1(X3,X2)))).#trainneg
cnf(i_0_7283, plain, (k2_partfun1(k1_numbers,k1_numbers,k2_funct_1(k3_sincos10),k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))=k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_6749, plain, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k6_numbers)|r1_xxreal_0(k6_numbers,k1_seq_1(k2_fdiff_9,X1)))).#trainneg
cnf(i_0_6428, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_numbers)|~m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k5_numbers))).#trainneg
cnf(i_0_6894, plain, (r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_xtuple_0(k3_sincos10))|~v2_funct_1(k3_sincos10)|~v1_relat_1(k3_sincos10)|~v1_funct_1(k3_sincos10))).#trainneg
cnf(i_0_6315, plain, (r1_xxreal_0(k10_real_1(np__1,np__2),k1_seq_1(k19_sin_cos,np__1)))).#trainneg
cnf(i_0_6459, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(k10_real_1(k32_sin_cos,np__4),k1_numbers)|~m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k5_numbers))).#trainneg
cnf(i_0_6386, plain, (v1_int_1(k1_real_1(np__2)))).#trainneg
cnf(i_0_6313, plain, (r1_xxreal_0(k6_numbers,k1_seq_1(k16_sin_cos,np__1)))).#trainneg
cnf(i_0_6985, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7059, plain, (r1_xxreal_0(np__2,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),np__2))).#trainneg
cnf(i_0_6895, plain, (r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k10_xtuple_0(k3_sincos10))|~v2_funct_1(k3_sincos10)|~v1_relat_1(k3_sincos10)|~v1_funct_1(k3_sincos10))).#trainneg
cnf(i_0_6579, plain, (r1_xxreal_0(np__2,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),np__2))).#trainneg
cnf(i_0_6914, negated_conjecture, (r2_hidden(esk1_0,k10_xtuple_0(k3_sincos10))|~v2_funct_1(k3_sincos10)|~v1_relat_1(k3_sincos10)|~v1_funct_1(k3_sincos10))).#trainneg
cnf(i_0_6483, plain, (r1_xxreal_0(X1,k10_real_1(k32_sin_cos,np__2))|~r1_xxreal_0(np__1,k1_seq_1(k2_fdiff_9,X1))|~r2_hidden(X1,k2_relset_1(k1_numbers,k3_sincos10))|~r2_hidden(k10_real_1(k32_sin_cos,np__2),k2_relset_1(k1_numbers,k3_sincos10)))).#trainneg
cnf(i_0_6463, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k29_sin_cos)))|~r1_tarski(X1,k1_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_7284, plain, (k2_partfun1(k1_numbers,k1_numbers,k2_funct_1(k3_sincos10),X1)=k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10))))).#trainneg
cnf(i_0_6260, plain, (m1_subset_1(X1,k1_xxreal_1(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_7179, plain, (m1_subset_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r1_tarski(X1,X2)|~v1_funct_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X2)))).#trainneg
cnf(i_0_6896, plain, (r1_tarski(k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers),k10_xtuple_0(k3_sincos10))|~v2_funct_1(k3_sincos10)|~v1_relat_1(k3_sincos10)|~v1_funct_1(k3_sincos10))).#trainneg
cnf(i_0_6423, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1))|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))|~v1_xxreal_0(X1)|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1))).#trainneg
cnf(i_0_4620, plain, (r2_hidden(X1,k1_xxreal_1(X3,X2))|~r2_hidden(X1,k4_xxreal_1(X3,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~r1_xxreal_0(X3,X2))).#trainneg
cnf(i_0_6422, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1))|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5470, plain, (r1_xxreal_0(k1_real_1(esk14_3(X1,k1_real_1(X2),X3)),X2)|r1_xxreal_0(X3,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(esk14_3(X1,k1_real_1(X2),X3),k1_numbers)|~v1_xxreal_0(X3)|~v1_xreal_0(k1_real_1(X2))|~r1_xxreal_0(k1_real_1(X2),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5872, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_5255, plain, (r1_tarski(k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))),k1_relset_1(k1_numbers,k2_fdiff_9))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_6756, plain, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),np__2)|r1_xxreal_0(np__2,k1_seq_1(k2_fdiff_9,X1)))).#trainneg
cnf(i_0_6219, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_5135, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k6_numbers)|r2_hidden(k6_numbers,a_0_1_sin_cos9)|~m1_subset_1(k10_real_1(k32_sin_cos,np__4),k1_numbers))).#trainneg
cnf(i_0_5419, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_xxreal_1(X2,X3)))|~r1_tarski(X1,k2_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_5407, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r1_tarski(X1,k29_sin_cos))).#trainneg
cnf(i_0_5644, plain, (m1_subset_1(k2_rcomp_1(X1,X2),k1_zfmisc_1(X3))|~v6_xxreal_2(X3)|~v2_membered(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(X2,X3)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_4327, plain, (r2_hidden(X1,k1_xxreal_1(X2,X3))|~r2_hidden(X1,k2_tarski(X3,X2))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3))).#trainneg
cnf(i_0_4397, plain, (k9_xtuple_0(k2_partfun1(k1_numbers,k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),X2))=X2|~r1_tarski(X2,X1)|~v1_relat_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))|~r1_tarski(X1,k1_relset_1(k1_numbers,k2_fdiff_9)))).#trainneg
cnf(i_0_5156, plain, (r1_xxreal_0(X1,k1_real_1(k1_real_1(X2)))|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6775, plain, (r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_6390, plain, (r1_xxreal_0(np__2,X1)|~r1_xxreal_0(k1_real_1(X1),k1_real_1(np__2))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5474, plain, (esk14_3(X1,X2,X3)=X2|r1_xxreal_0(X3,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(esk14_3(X1,X2,X3))|~r1_xxreal_0(esk14_3(X1,X2,X3),X2)|~v1_xxreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4151, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_7385, negated_conjecture, (r1_xxreal_0(esk1_0,esk1_0))).#trainneg
cnf(i_0_8251, plain, (r2_hidden(k6_numbers,a_0_1_sin_cos9))).#trainneg
cnf(i_0_7499, plain, (m1_subset_1(k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6202, plain, (~r1_xxreal_0(k6_numbers,k1_real_1(np__2)))).#trainneg
cnf(i_0_7302, plain, (v1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_7358, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_7648, plain, (m1_subset_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_7308, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k32_sin_cos))).#trainneg
cnf(i_0_7304, plain, (~v1_xboole_0(k1_rcomp_1(k6_numbers,k32_sin_cos)))).#trainneg
cnf(i_0_8144, plain, (m1_subset_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_8007, plain, (m1_subset_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_8116, plain, (m1_subset_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10))))).#trainneg
cnf(i_0_4207, plain, (r1_tarski(k2_relset_1(k1_numbers,k1_sin_cos9),k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_7796, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),k1_seq_1(k2_fdiff_9,esk1_0)))).#trainneg
cnf(i_0_7519, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_7417, plain, (m1_subset_1(k2_relset_1(k1_numbers,k1_sin_cos9),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_7305, plain, (m1_subset_1(k10_real_1(k32_sin_cos,np__4),k1_rcomp_1(k6_numbers,k32_sin_cos)))).#trainneg
cnf(i_0_7316, plain, (r1_xxreal_0(k1_real_1(np__2),k1_real_1(np__2)))).#trainneg
cnf(i_0_7551, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k2_fdiff_9))|~r1_tarski(X1,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))).#trainneg
cnf(i_0_5648, plain, (r1_tarski(k2_relset_1(k1_numbers,k1_sin_cos9),X1)|~v6_xxreal_2(X1)|~v2_membered(X1)|~r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~r2_hidden(k10_real_1(k32_sin_cos,np__2),X1))).#trainneg
cnf(i_0_7824, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_7688, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k2_fdiff_9))|r2_hidden(esk1_0,k1_relset_1(k1_numbers,k2_fdiff_9)))).#trainneg
cnf(i_0_7953, plain, (v1_funct_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))))).#trainneg
cnf(i_0_7615, plain, (m1_subset_1(k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_7357, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k32_sin_cos))).#trainneg
cnf(i_0_8006, plain, (v1_funct_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))))).#trainneg
cnf(i_0_7670, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r1_xxreal_0(k1_real_1(np__1),k1_seq_1(k2_fdiff_9,esk1_0)))).#trainneg
cnf(i_0_7862, plain, (v1_relat_1(k19_sin_cos))).#trainneg
cnf(i_0_7781, plain, (k9_xtuple_0(k2_partfun1(k1_numbers,k1_numbers,X1,X2))=X2|~r1_tarski(X2,k9_xtuple_0(X1))|~v1_relat_1(X1)|~r1_tarski(X1,k2_fdiff_9)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7319, plain, (r1_xxreal_0(X1,k1_real_1(np__2))|r1_xxreal_0(k1_real_1(np__2),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7299, plain, (r1_xxreal_0(X1,k1_seq_1(k2_fdiff_9,X2))|r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_7784, plain, (m1_subset_1(k2_partfun1(k1_numbers,k1_numbers,X1,X2),k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~r1_tarski(X1,k2_fdiff_9)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7915, plain, (k1_real_1(k10_real_1(k32_sin_cos,np__4))=k1_real_1(k10_real_1(k32_sin_cos,np__2))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__2)),a_0_2_sin_cos9))).#trainneg
cnf(i_0_7419, plain, (~v2_xxreal_0(k1_real_1(np__1)))).#trainneg
cnf(i_0_7964, plain, (v2_funct_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))))).#trainneg
cnf(i_0_7454, plain, (m1_subset_1(k10_real_1(k32_sin_cos,np__4),k1_xxreal_1(k6_numbers,k32_sin_cos))|~v1_xxreal_0(k32_sin_cos))).#trainneg
cnf(i_0_7774, plain, (k2_partfun1(k1_numbers,k1_numbers,k5_relat_1(X1,X2),X3)=k5_relat_1(X1,X3)|~r1_tarski(X3,X2)|~v1_relat_1(X1)|~r1_tarski(k5_relat_1(X1,X2),k2_fdiff_9)|~v1_funct_1(k5_relat_1(X1,X2)))).#trainneg
cnf(i_0_7350, plain, (r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_460, plain, (k2_xcmplx_0(np__2,k1_real_1(np__1))=np__1)).#trainneg
cnf(i_0_7606, plain, (esk14_3(X2,X1,X3)=X1|~v1_xreal_0(X1)|~v1_xreal_0(esk14_3(X2,X1,X3))|~r1_xxreal_0(esk14_3(X2,X1,X3),X1)|~v1_xxreal_0(X3)|~r2_hidden(X2,k3_rcomp_1(X1,X3)))).#trainneg
cnf(i_0_7604, plain, (esk15_3(X2,X1,X3)=X1|~v1_xreal_0(X1)|~v1_xreal_0(esk15_3(X2,X1,X3))|~r1_xxreal_0(esk15_3(X2,X1,X3),X1)|~v1_xreal_0(X3)|~r2_hidden(X2,k1_rcomp_1(X1,X3)))).#trainneg
cnf(i_0_7431, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_fdiff_9)))).#trainneg
cnf(i_0_8115, plain, (v1_funct_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10))))).#trainneg
cnf(i_0_7782, plain, (k5_relat_1(k2_partfun1(k1_numbers,k1_numbers,X1,X2),X3)=k5_relat_1(X1,X3)|~r1_tarski(X3,X2)|~v1_relat_1(X1)|~r1_tarski(X1,k2_fdiff_9)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8017, plain, (v2_funct_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))))).#trainneg
cnf(i_0_7780, plain, (k10_xtuple_0(k2_partfun1(k1_numbers,k1_numbers,X1,X2))=k7_relat_1(X1,X2)|~v1_relat_1(X1)|~r1_tarski(X1,k2_fdiff_9)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7363, plain, (~v2_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_7429, plain, (v1_relat_1(k1_relset_1(k2_fdiff_9,X1))|~v4_relat_1(X1,k2_fdiff_9)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7682, negated_conjecture, (v1_xxreal_0(esk1_0)|~v2_membered(k1_relset_1(k1_numbers,k2_fdiff_9)))).#trainneg
cnf(i_0_5596, plain, (k1_xxreal_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2))=k2_relset_1(k1_numbers,k1_sin_cos9)|~r1_tarski(k1_xxreal_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)),k2_relset_1(k1_numbers,k1_sin_cos9)))).#trainneg
cnf(i_0_7415, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k2_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8121, plain, (k9_xtuple_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))=k1_relset_1(k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))|~v1_relat_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10))))).#trainneg
cnf(i_0_7361, plain, (r1_xxreal_0(X1,k10_real_1(k32_sin_cos,np__2))|r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5603, plain, (r1_tarski(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)))|~r1_tarski(X1,k2_relset_1(k1_numbers,k1_sin_cos9))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_7958, plain, (v4_relat_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)),k1_numbers))).#trainneg
cnf(i_0_7969, plain, (r1_xxreal_0(X1,X2)|~r1_tarski(k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers),k1_relset_1(k1_numbers,k2_funct_1(k3_sincos10)))|~r1_xxreal_0(k1_seq_1(k2_funct_1(k3_sincos10),X2),k1_seq_1(k2_funct_1(k3_sincos10),X1))|~m1_subset_1(k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers),k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))|~r2_hidden(X2,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))).#trainneg
cnf(i_0_611, plain, (k2_xcmplx_0(k1_real_1(np__1),np__2)=np__1)).#trainneg
cnf(i_0_7394, plain, (r1_xxreal_0(X1,k1_real_1(np__1))|r1_xxreal_0(k1_real_1(np__1),X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_7317, plain, (r1_xxreal_0(X1,k1_real_1(np__2))|r1_xxreal_0(k1_real_1(np__2),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8122, plain, (k2_relset_1(k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))=k7_relat_1(k2_fdiff_9,X1)|~v1_relat_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10))))).#trainneg
cnf(i_0_7395, plain, (r1_xxreal_0(X1,k1_real_1(np__1))|r1_xxreal_0(k1_real_1(np__1),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8216, plain, (r1_xxreal_0(X1,k10_real_1(k32_sin_cos,np__2))|r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_7347, plain, (r1_xxreal_0(k6_numbers,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k6_numbers))).#trainneg
cnf(i_0_7351, plain, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_real_1(k10_real_1(k32_sin_cos,np__4)))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_seq_1(k2_fdiff_9,X1)))).#trainneg
cnf(i_0_7916, plain, (k1_real_1(k10_real_1(k32_sin_cos,np__4))=k1_real_1(k10_real_1(k32_sin_cos,np__2))|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k10_real_1(k32_sin_cos,np__4))|~m1_subset_1(k10_real_1(k32_sin_cos,np__2),k1_numbers)|~m1_subset_1(k10_real_1(k32_sin_cos,np__4),k1_numbers))).#trainneg
cnf(i_0_7959, plain, (k9_xtuple_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))=k1_relset_1(k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))|~v1_relat_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))))).#trainneg
cnf(i_0_4522, plain, (r1_xxreal_0(X1,k1_real_1(np__2))|~r1_xxreal_0(np__2,k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7783, plain, (k5_relat_1(k2_partfun1(k1_numbers,k1_numbers,X1,X2),X2)=k2_partfun1(k1_numbers,k1_numbers,X1,X2)|~v1_relat_1(X1)|~r1_tarski(X1,k2_fdiff_9)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7506, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10))))).#trainneg
cnf(i_0_8100, plain, (r2_hidden(np__2,a_0_2_sin_cos9)|r1_xxreal_0(np__2,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(np__2,k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_7354, plain, (r1_xxreal_0(np__2,k10_real_1(k32_sin_cos,np__2))|r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),np__2))).#trainneg
cnf(i_0_7677, plain, (r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__4)),a_0_2_sin_cos9)|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k10_real_1(k32_sin_cos,np__4))|~m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_numbers))).#trainneg
cnf(i_0_7828, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7961, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r1_tarski(X1,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))))).#trainneg
cnf(i_0_8004, plain, (r2_hidden(X3,k1_rcomp_1(X1,X2))|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)|~r2_hidden(X3,k1_rcomp_1(X2,X1)))).#trainneg
cnf(i_0_8102, negated_conjecture, (m1_subset_1(esk1_0,k10_xtuple_0(k3_sincos10))|~v2_funct_1(k3_sincos10)|~v1_relat_1(k3_sincos10)|~v1_funct_1(k3_sincos10))).#trainneg
cnf(i_0_8011, plain, (v4_relat_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))),k1_numbers))).#trainneg
cnf(i_0_8036, plain, (~v1_xboole_0(k10_xtuple_0(k3_sincos10))|~v2_funct_1(k3_sincos10)|~v1_relat_1(k3_sincos10)|~v1_funct_1(k3_sincos10))).#trainneg
cnf(i_0_8332, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3)|~r2_hidden(X2,k2_tarski(X3,X1)))).#trainneg
cnf(i_0_5543, plain, (r1_xxreal_0(X1,esk14_3(X2,k1_real_1(k1_real_1(X1)),X3))|~v1_xreal_0(esk14_3(X2,k1_real_1(k1_real_1(X1)),X3))|~m1_subset_1(X1,k1_numbers)|~v1_xxreal_0(X3)|~v1_xreal_0(k1_real_1(k1_real_1(X1)))|~r2_hidden(X2,k3_rcomp_1(k1_real_1(k1_real_1(X1)),X3)))).#trainneg
cnf(i_0_6897, plain, (r1_tarski(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))),k10_xtuple_0(k3_sincos10))|~v2_funct_1(k3_sincos10)|~v1_relat_1(k3_sincos10)|~v1_funct_1(k3_sincos10))).#trainneg
cnf(i_0_7051, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__2))))|~m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_numbers)|~m1_subset_1(k10_real_1(k32_sin_cos,np__4),k1_numbers)|~v1_xreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_7552, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k2_fdiff_9)))|~r1_tarski(X1,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))).#trainneg
cnf(i_0_5804, plain, (m1_subset_1(k2_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_zfmisc_1(k1_relset_1(k1_numbers,k2_fdiff_9))))).#trainneg
cnf(i_0_7957, plain, (v5_relat_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)),k1_numbers))).#trainneg
cnf(i_0_6758, plain, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),X2)|r1_xxreal_0(X2,k1_seq_1(k2_fdiff_9,X1))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_7956, plain, (k5_relat_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)),X1)=k2_partfun1(k1_numbers,k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)),X1))).#trainneg
cnf(i_0_6879, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|r2_hidden(X1,k2_rcomp_1(X3,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X3,X1))).#trainneg
cnf(i_0_8449, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_seq_1(k2_fdiff_9,X1))|r1_xxreal_0(esk1_0,X1)|~r2_hidden(X1,k2_relset_1(k1_numbers,k3_sincos10)))).#trainneg
cnf(i_0_8564, plain, (m1_subset_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_9533, plain, (m1_subset_1(k1_real_1(np__2),k1_numbers))).#trainneg
cnf(i_0_8643, plain, (~r2_hidden(k10_real_1(k32_sin_cos,np__2),a_0_1_sin_cos9))).#trainneg
cnf(i_0_8408, plain, (m1_subset_1(k6_numbers,a_0_1_sin_cos9))).#trainneg
cnf(i_0_8765, plain, (~r2_hidden(k32_sin_cos,a_0_0_sin_cos9))).#trainneg
cnf(i_0_8513, plain, (~r2_hidden(k32_sin_cos,a_0_1_sin_cos9))).#trainneg
cnf(i_0_8450, plain, (r1_xxreal_0(k1_real_1(np__1),k1_seq_1(k2_fdiff_9,X1))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~r2_hidden(X1,k2_relset_1(k1_numbers,k3_sincos10)))).#trainneg
cnf(i_0_8797, plain, (r1_xxreal_0(k1_real_1(np__2),k6_numbers))).#trainneg
cnf(i_0_8937, plain, (X1=X2|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~r1_xxreal_0(X1,X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_9233, plain, (r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__4)),a_0_2_sin_cos9)|~r1_xxreal_0(k1_real_1(np__1),np__1))).#trainneg
cnf(i_0_9487, plain, (r2_hidden(X3,k1_rcomp_1(X1,X2))|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X3))).#trainneg
cnf(i_0_9209, plain, (~r2_hidden(k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__4))),a_0_1_sin_cos9))).#trainneg
cnf(i_0_9474, plain, (k2_partfun1(k1_numbers,k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)),k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))=k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))).#trainneg
cnf(i_0_9006, plain, (k1_relset_1(k1_numbers,k2_funct_1(k3_sincos10))=k2_relset_1(k1_numbers,k3_sincos10)|~v1_relat_1(k2_funct_1(k3_sincos10)))).#trainneg
cnf(i_0_8928, plain, (X1=X2|~v1_xxreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)|~r2_hidden(X1,k3_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_8966, plain, (k5_relat_1(k2_funct_1(k3_sincos10),X1)=k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1)|~r1_tarski(X1,k2_relset_1(k1_numbers,k3_sincos10)))).#trainneg
cnf(i_0_9286, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6019, plain, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(X1))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~r2_hidden(X1,a_0_1_sin_cos9))).#trainneg
cnf(i_0_8672, plain, (r1_xxreal_0(np__2,k1_real_1(k1_real_1(np__2)))|~m1_subset_1(k1_real_1(np__2),k1_numbers))).#trainneg
cnf(i_0_9232, plain, (r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__4)),a_0_2_sin_cos9)|~r2_hidden(k10_real_1(k32_sin_cos,np__4),a_0_2_sin_cos9))).#trainneg
cnf(i_0_9529, plain, (m1_subset_1(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),k1_numbers))).#trainneg
cnf(i_0_8506, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),k1_numbers))).#trainneg
cnf(i_0_9179, plain, (k1_relset_1(k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))=k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)|~v1_relat_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))))).#trainneg
cnf(i_0_8407, plain, (~v1_xboole_0(a_0_1_sin_cos9))).#trainneg
cnf(i_0_9381, plain, (k1_real_1(k10_real_1(k32_sin_cos,np__4))=k1_real_1(k10_real_1(k32_sin_cos,np__2))|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k10_real_1(k32_sin_cos,np__4))|~m1_subset_1(k10_real_1(k32_sin_cos,np__2),k1_numbers))).#trainneg
cnf(i_0_8509, plain, (k10_real_1(k32_sin_cos,np__4)=k32_sin_cos|~r1_xxreal_0(k32_sin_cos,k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_8964, plain, (k2_partfun1(k1_numbers,k1_numbers,k2_partfun1(k1_numbers,k1_numbers,X1,X2),X3)=k5_relat_1(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~r1_tarski(X1,k2_fdiff_9)|~r1_tarski(X3,X2))).#trainneg
cnf(i_0_9197, plain, (k2_partfun1(k1_numbers,k1_numbers,k2_partfun1(k1_numbers,k1_numbers,X1,X2),X2)=k2_partfun1(k1_numbers,k1_numbers,X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~r1_tarski(X1,k2_fdiff_9))).#trainneg
cnf(i_0_8869, plain, (m1_subset_1(np__1,k1_rcomp_1(k1_real_1(np__1),np__1)))).#trainneg
cnf(i_0_8989, plain, (k1_xxreal_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2))=k2_relset_1(k1_numbers,k1_sin_cos9)|~r1_tarski(k1_xxreal_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)),k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_8427, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))).#trainneg
cnf(i_0_8870, plain, (k1_real_1(np__1)=np__1|~r1_xxreal_0(np__1,k1_real_1(np__1)))).#trainneg
cnf(i_0_8916, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_real_1(X1))|r1_xxreal_0(k1_real_1(X1),k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8888, plain, (k2_partfun1(k1_numbers,k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k19_sin_cos,X1),X2)=k2_partfun1(k1_numbers,k1_numbers,k19_sin_cos,X2)|~r1_tarski(k2_partfun1(k1_numbers,k1_numbers,k19_sin_cos,X1),k2_fdiff_9)|~r1_tarski(X2,X1)|~v1_funct_1(k2_partfun1(k1_numbers,k1_numbers,k19_sin_cos,X1)))).#trainneg
cnf(i_0_8827, plain, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k10_real_1(k32_sin_cos,np__2))|r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k1_seq_1(k2_fdiff_9,X1)))).#trainneg
cnf(i_0_8859, plain, (m1_subset_1(k2_funct_1(k3_sincos10),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_fdiff_9,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_8976, plain, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_seq_1(k2_fdiff_9,X2))|r1_xxreal_0(X1,X2)|~m1_subset_1(k1_relset_1(k2_relset_1(k1_numbers,k3_sincos10),X3),k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,k1_relset_1(k2_relset_1(k1_numbers,k3_sincos10),X3))|~r2_hidden(X2,k1_relset_1(k2_relset_1(k1_numbers,k3_sincos10),X3))|~v4_relat_1(X3,k2_relset_1(k1_numbers,k3_sincos10))|~v1_relat_1(X3))).#trainneg
cnf(i_0_8824, plain, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_real_1(k10_real_1(k32_sin_cos,np__2)))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_seq_1(k2_fdiff_9,X1)))).#trainneg
cnf(i_0_9183, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(np__2))|~r1_xxreal_0(np__2,X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8639, plain, (k10_real_1(k32_sin_cos,np__4)=k10_real_1(k32_sin_cos,np__2)|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_8941, plain, (esk15_3(X1,k1_real_1(k1_real_1(X2)),X3)=k1_real_1(k1_real_1(X2))|~v1_xreal_0(esk15_3(X1,k1_real_1(k1_real_1(X2)),X3))|~v1_xreal_0(k1_real_1(k1_real_1(X2)))|~v1_xreal_0(X3)|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k1_real_1(X2)),X3))|~r1_xxreal_0(esk15_3(X1,k1_real_1(k1_real_1(X2)),X3),X2)|~m1_subset_1(esk15_3(X1,k1_real_1(k1_real_1(X2)),X3),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_8573, plain, (r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_numbers))).#trainneg
cnf(i_0_9384, plain, (k1_real_1(k10_real_1(k32_sin_cos,np__4))=k1_real_1(k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_9224, plain, (r2_hidden(np__2,a_0_2_sin_cos9)|~r2_hidden(np__2,a_0_1_sin_cos9)|~r1_xxreal_0(np__2,k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_9546, plain, (m1_subset_1(k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__4))),k1_numbers)|~v1_xreal_0(k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_9253, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(k10_real_1(k32_sin_cos,np__2)))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8674, plain, (r1_xxreal_0(k1_real_1(k1_real_1(np__2)),np__2)|~m1_subset_1(k1_real_1(np__2),k1_numbers))).#trainneg
cnf(i_0_4264, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X1,k1_real_1(k1_real_1(X2)))|~m1_subset_1(k1_real_1(X2),k1_numbers))).#trainneg
cnf(i_0_8404, plain, (r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_numbers))).#trainneg
cnf(i_0_8711, plain, (m1_subset_1(k2_relset_1(k1_numbers,k1_sin_cos9),k1_zfmisc_1(k1_numbers))|~v6_xxreal_2(k5_numbers)|~v2_membered(k5_numbers)|~r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k5_numbers)|~r2_hidden(k10_real_1(k32_sin_cos,np__2),k5_numbers))).#trainneg
cnf(i_0_5404, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r1_tarski(X1,k1_sin_cos6))).#trainneg
cnf(i_0_4277, plain, (r1_xxreal_0(X1,np__1)|~v1_xreal_0(k1_real_1(np__1))|~r1_xxreal_0(k1_real_1(np__1),k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7553, plain, (r1_tarski(k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers),k1_numbers))).#trainneg
cnf(i_0_4037, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_169, plain, (k6_xcmplx_0(np__1,np__2)=k1_real_1(np__1))).#trainneg
cnf(i_0_4616, plain, (r2_hidden(X1,X2)|X2!=k1_xxreal_1(X4,X3)|~r2_hidden(X1,k4_xxreal_1(X4,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X4)|~r1_xxreal_0(X4,X3))).#trainneg
cnf(i_0_7555, plain, (m1_subset_1(k2_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5423, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k3_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8357, plain, (X1=k1_real_1(k1_real_1(X2))|~v1_xxreal_0(X1)|~v1_xreal_0(k1_real_1(k1_real_1(X2)))|~r1_xxreal_0(k1_real_1(k1_real_1(X2)),X1)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_8012, plain, (k9_xtuple_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))=k1_relset_1(k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))|~v1_relat_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))))).#trainneg
cnf(i_0_4454, plain, (v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X1,k1_relset_1(k1_numbers,X1)))|r2_hidden(esk3_2(k1_relset_1(k1_numbers,X1),X1),k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1))).#trainneg
cnf(i_0_8799, plain, (r1_xxreal_0(k1_real_1(np__2),k1_seq_1(k2_fdiff_9,X1))|r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_real_1(np__2)))).#trainneg
cnf(i_0_5343, plain, (k1_relset_1(X1,X2)=X1|~r1_tarski(X1,k1_relset_1(X1,X2))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_8805, plain, (r1_xxreal_0(k1_real_1(np__2),k1_real_1(X1))|r1_xxreal_0(k1_real_1(X1),k1_real_1(np__2))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8010, plain, (v5_relat_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))),k1_numbers))).#trainneg
cnf(i_0_6176, plain, (r1_tarski(k3_xboole_0(X3,X2),X1)|~m1_subset_1(X3,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_8126, plain, (v2_funct_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))|~v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10))))).#trainneg
cnf(i_0_8826, plain, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k10_real_1(k32_sin_cos,np__4))|r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_seq_1(k2_fdiff_9,X1)))).#trainneg
cnf(i_0_6340, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_4715, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k7_square_1(X1),k7_square_1(X2))|~r1_xxreal_0(k6_numbers,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4325, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k2_rcomp_1(X3,X2))|~r2_hidden(X1,k1_xxreal_1(X3,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~r1_xxreal_0(X3,X2))).#trainneg
cnf(i_0_5976, plain, (r1_xxreal_0(X1,k1_real_1(X2))|r1_xxreal_0(k1_real_1(X2),X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6466, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k29_sin_cos))|~r1_tarski(X1,X2)|~r1_tarski(X2,k1_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_8337, plain, (k1_relset_1(k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),X2))=X2|~v1_relat_1(k2_partfun1(k1_numbers,k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),X2))|~m1_subset_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))|~r1_tarski(X1,k1_relset_1(k1_numbers,k2_fdiff_9))|~r1_tarski(X2,X1)|~v1_relat_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1)))).#trainneg
cnf(i_0_4176, plain, (~v1_xboole_0(a_0_0_sin_cos9))).#trainneg
cnf(i_0_8123, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r1_tarski(X1,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10))))).#trainneg
cnf(i_0_9642, plain, (v1_relat_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1)))).#trainneg
cnf(i_0_3926, plain, (k1_real_1(k1_real_1(np__2))=np__2)).#trainneg
cnf(i_0_7417, plain, (m1_subset_1(k2_relset_1(k1_numbers,k1_sin_cos9),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_9233, plain, (r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__4)),a_0_2_sin_cos9))).#trainneg
cnf(i_0_4207, plain, (r1_tarski(k2_relset_1(k1_numbers,k1_sin_cos9),k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_10454, plain, (m1_subset_1(k1_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__4)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4179, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~r2_hidden(X1,k2_relset_1(k1_numbers,k3_sincos10)))).#trainneg
cnf(i_0_9683, negated_conjecture, (~r2_hidden(esk1_0,a_0_1_sin_cos9))).#trainneg
cnf(i_0_7363, plain, (~v2_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_9644, plain, (v4_relat_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),k1_numbers))).#trainneg
cnf(i_0_9702, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__2),k1_numbers))).#trainneg
cnf(i_0_9682, negated_conjecture, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(esk1_0)))).#trainneg
cnf(i_0_9726, plain, (~v2_xxreal_0(k1_real_1(np__2)))).#trainneg
cnf(i_0_5603, plain, (r1_tarski(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)))|~r1_tarski(X1,k2_relset_1(k1_numbers,k1_sin_cos9)))).#trainneg
cnf(i_0_10094, plain, (v4_relat_1(k2_funct_1(k3_sincos10),X1)|~m1_subset_1(k2_fdiff_9,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10646, plain, (r1_xxreal_0(X1,np__2)|~r1_xxreal_0(k7_square_1(X1),k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9634, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r1_tarski(X1,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X2)))).#trainneg
cnf(i_0_9663, plain, (v1_xxreal_0(k1_real_1(np__2))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_9684, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k10_real_1(k32_sin_cos,np__4))|r2_hidden(k10_real_1(k32_sin_cos,np__4),a_0_0_sin_cos9))).#trainneg
cnf(i_0_9942, plain, (k1_real_1(k10_real_1(k32_sin_cos,np__4))=k1_real_1(k10_real_1(k32_sin_cos,np__2))|~r2_hidden(k10_real_1(k32_sin_cos,np__2),a_0_2_sin_cos9))).#trainneg
cnf(i_0_9776, plain, (X1=X2|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X3,X1))).#trainneg
cnf(i_0_9585, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|r1_xxreal_0(esk1_0,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_10002, plain, (k2_partfun1(k1_numbers,k1_numbers,k2_funct_1(k3_sincos10),k2_relset_1(k1_numbers,k3_sincos10))=k2_funct_1(k3_sincos10))).#trainneg
cnf(i_0_10137, plain, (esk15_3(X1,X2,X3)=X2|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~r1_xxreal_0(esk15_3(X1,X2,X3),X2)|~m1_subset_1(esk15_3(X1,X2,X3),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_10132, plain, (k10_real_1(k32_sin_cos,np__4)=k10_real_1(k32_sin_cos,np__2)|~r2_hidden(k10_real_1(k32_sin_cos,np__2),a_0_2_sin_cos9))).#trainneg
cnf(i_0_9589, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_seq_1(k2_fdiff_9,X1))|r1_xxreal_0(esk1_0,X1)|~r2_hidden(X1,k10_xtuple_0(k3_sincos10))|~v2_funct_1(k3_sincos10)|~v1_relat_1(k3_sincos10)|~v1_funct_1(k3_sincos10))).#trainneg
cnf(i_0_9669, plain, (r2_hidden(k1_real_1(np__2),k1_numbers))).#trainneg
cnf(i_0_10260, plain, (m1_subset_1(k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__4))),k1_numbers))).#trainneg
cnf(i_0_5255, plain, (r1_tarski(k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))),k1_relset_1(k1_numbers,k2_fdiff_9)))).#trainneg
cnf(i_0_9573, negated_conjecture, (r1_xxreal_0(np__2,esk1_0)|r1_xxreal_0(esk1_0,np__2))).#trainneg
cnf(i_0_5618, plain, (X1=k1_real_1(k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r2_hidden(X1,k2_relset_1(k1_numbers,k3_sincos10)))).#trainneg
cnf(i_0_10480, plain, (X1=k1_real_1(k1_real_1(X2))|~v1_xxreal_0(X1)|~v1_xreal_0(k1_real_1(k1_real_1(X2)))|~r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X2,X1))).#trainneg
cnf(i_0_6422, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_10199, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k32_sin_cos)))|~r1_xxreal_0(k32_sin_cos,k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_9636, plain, (k5_relat_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),X2)=k2_partfun1(k1_numbers,k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),X2)|~v1_funct_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1)))).#trainneg
cnf(i_0_9892, plain, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),np__1)|~r2_hidden(k1_real_1(np__1),a_0_1_sin_cos9))).#trainneg
cnf(i_0_7051, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)))))).#trainneg
cnf(i_0_9676, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__4))=esk1_0|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),esk1_0))).#trainneg
cnf(i_0_9839, plain, (X1=k1_real_1(k10_real_1(k32_sin_cos,np__2))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r2_hidden(X1,k2_relset_1(k1_numbers,k3_sincos10)))).#trainneg
cnf(i_0_10157, plain, (~r1_xxreal_0(k6_numbers,k1_real_1(k32_sin_cos))|~r1_xxreal_0(k32_sin_cos,k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_9854, plain, (k5_relat_1(k2_funct_1(k3_sincos10),k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))=k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))).#trainneg
cnf(i_0_9775, plain, (X1=X2|~v1_xxreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X3,X1))).#trainneg
cnf(i_0_9657, plain, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_seq_1(k2_fdiff_9,X2))|r1_xxreal_0(X1,X2)|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9779, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X3,X2))).#trainneg
cnf(i_0_10455, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k3_rcomp_1(X2,X3))))).#trainneg
cnf(i_0_6423, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1))|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_10241, plain, (~r1_xxreal_0(k1_real_1(k32_sin_cos),k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r1_xxreal_0(k32_sin_cos,k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_9583, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|r1_xxreal_0(esk1_0,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10648, plain, (r1_xxreal_0(np__2,X1)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)),k7_square_1(X1))|~r1_xxreal_0(k6_numbers,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10389, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))).#trainneg
cnf(i_0_9860, plain, (k5_relat_1(k2_funct_1(k3_sincos10),k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))=k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_10195, plain, (r2_hidden(k1_real_1(k32_sin_cos),k2_relset_1(k1_numbers,k3_sincos10))|~r1_xxreal_0(k32_sin_cos,k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_9580, negated_conjecture, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),esk1_0)|r1_xxreal_0(esk1_0,k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_9658, plain, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_seq_1(k2_fdiff_9,X2))|r1_xxreal_0(X1,X2)|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1))).#trainneg
cnf(i_0_10020, plain, (k1_xxreal_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2))=k2_relset_1(k1_numbers,k1_sin_cos9)|~v6_xxreal_2(k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))|~v2_membered(k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))|~r2_hidden(k10_real_1(k32_sin_cos,np__2),k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))|~r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_10066, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),np__1)|r1_xxreal_0(np__1,k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_9740, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0|~r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_10198, plain, (r2_hidden(k1_real_1(k32_sin_cos),k2_relset_1(k1_numbers,k1_sin_cos9))|~r1_xxreal_0(k32_sin_cos,k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_10153, plain, (v1_xreal_0(k1_real_1(k32_sin_cos))|~r1_xxreal_0(k32_sin_cos,k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_7615, plain, (m1_subset_1(k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_10219, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1))|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k32_sin_cos)))|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k32_sin_cos)))|~r1_xxreal_0(k32_sin_cos,k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_9641, plain, (v5_relat_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),k1_numbers))).#trainneg
cnf(i_0_9707, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k10_real_1(k32_sin_cos,np__2))|r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k1_real_1(k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_9660, plain, (r1_xxreal_0(k1_real_1(np__1),k1_seq_1(k2_fdiff_9,X1))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_4097, plain, (r1_xxreal_0(k1_real_1(np__1),k1_seq_1(k19_sin_cos,X1))|~v1_xreal_0(k1_real_1(np__1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4408, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(X3,X2),k1_seq_1(X3,X1))|~v6_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_1(X3)|~r2_hidden(X1,k1_relset_1(k1_numbers,X3))|~r2_hidden(X2,k1_relset_1(k1_numbers,X3))|~r1_tarski(k1_relset_1(k1_numbers,X3),X4)|~m1_subset_1(k1_relset_1(k1_numbers,X3),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4074, plain, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),X1)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(k1_real_1(X1),a_0_1_sin_cos9))).#trainneg
cnf(i_0_9312, plain, (r1_xxreal_0(X1,X2)|r2_hidden(X2,k2_rcomp_1(X3,X1))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~r1_xxreal_0(X3,X1)|~r2_hidden(X2,k1_rcomp_1(X3,X1)))).#trainneg
cnf(i_0_6192, plain, (k9_subset_1(X4,X2,X1)=k9_subset_1(X3,X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X4)))).#trainneg
cnf(i_0_9199, plain, (k5_relat_1(k2_funct_1(k3_sincos10),k2_relset_1(k1_numbers,k3_sincos10))=k2_funct_1(k3_sincos10))).#trainneg
cnf(i_0_5388, plain, (X1=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_10626, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),np__1)|r1_xxreal_0(np__1,k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_7158, plain, (esk15_3(X1,X2,X2)=X2|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X2))|~m1_subset_1(esk15_3(X1,X2,X2),k1_numbers))).#trainneg
cnf(i_0_8356, plain, (X1=k1_real_1(k1_real_1(X2))|~v1_xreal_0(X1)|~v1_xreal_0(k1_real_1(k1_real_1(X2)))|~r1_xxreal_0(k1_real_1(k1_real_1(X2)),X1)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4768, plain, (k7_relat_1(k2_fdiff_9,k2_relset_1(k1_numbers,k3_sincos10))=k10_xtuple_0(k2_funct_1(k3_sincos10)))).#trainneg
cnf(i_0_4685, plain, (r1_xxreal_0(X1,X2)|r2_hidden(X2,k3_rcomp_1(X3,X1))|~v1_xxreal_0(X1)|~v1_xreal_0(X3)|~r1_xxreal_0(X3,X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4154, plain, (r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5864, plain, (m1_subset_1(k1_rcomp_1(k10_real_1(k32_sin_cos,np__4),k10_real_1(k32_sin_cos,np__2)),k1_zfmisc_1(k1_relset_1(k1_numbers,k2_fdiff_9))))).#trainneg
cnf(i_0_11473, plain, (~r1_xxreal_0(k32_sin_cos,k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_10931, plain, (v1_xxreal_0(k1_real_1(np__2)))).#trainneg
cnf(i_0_10833, plain, (~r2_hidden(k10_real_1(k32_sin_cos,np__2),a_0_2_sin_cos9))).#trainneg
cnf(i_0_10945, plain, (m1_subset_1(k10_real_1(k32_sin_cos,np__4),a_0_0_sin_cos9))).#trainneg
cnf(i_0_10840, negated_conjecture, (r1_xxreal_0(k1_real_1(esk1_0),k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_10824, plain, (m1_subset_1(k1_real_1(k10_real_1(k32_sin_cos,np__4)),a_0_2_sin_cos9))).#trainneg
cnf(i_0_10884, plain, (r2_hidden(k10_real_1(k32_sin_cos,np__4),a_0_2_sin_cos9))).#trainneg
cnf(i_0_9179, plain, (k1_relset_1(k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))=k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))).#trainneg
cnf(i_0_11650, plain, (k2_relset_1(k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))=k7_relat_1(k2_fdiff_9,X1))).#trainneg
cnf(i_0_10895, negated_conjecture, (m1_subset_1(k1_real_1(esk1_0),k1_numbers)|~v1_xreal_0(k1_real_1(esk1_0)))).#trainneg
cnf(i_0_11201, plain, (k2_partfun1(k1_numbers,k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),X1)=k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1)|~v1_funct_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1)))).#trainneg
cnf(i_0_10738, plain, (v1_relat_1(k2_funct_1(k3_sincos10)))).#trainneg
cnf(i_0_11564, negated_conjecture, (r1_xxreal_0(esk1_0,k10_real_1(k32_sin_cos,np__2))|~r2_hidden(esk1_0,a_0_0_sin_cos9))).#trainneg
cnf(i_0_10822, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k2_relset_1(k1_numbers,k1_sin_cos9)))).#trainneg
cnf(i_0_4398, plain, (k1_relset_1(k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))=X1|~r1_tarski(X1,k1_relset_1(k1_numbers,k2_fdiff_9)))).#trainneg
cnf(i_0_10799, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_seq_1(k2_fdiff_9,X1))|r1_xxreal_0(esk1_0,X1)|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10802, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_seq_1(k2_fdiff_9,X1))|r1_xxreal_0(esk1_0,X1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1))).#trainneg
cnf(i_0_11064, plain, (k1_relset_1(k1_numbers,k2_funct_1(k3_sincos10))=k2_relset_1(k1_numbers,k3_sincos10)|~v1_relat_1(k2_funct_1(k3_sincos10)))).#trainneg
cnf(i_0_11051, negated_conjecture, (r1_xxreal_0(esk1_0,k10_real_1(k32_sin_cos,np__4))|r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),esk1_0))).#trainneg
cnf(i_0_10902, plain, (r1_tarski(k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_10823, plain, (~v1_xboole_0(a_0_2_sin_cos9))).#trainneg
cnf(i_0_11902, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),X1)|r1_xxreal_0(X1,k10_real_1(k32_sin_cos,np__4))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_11177, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_seq_1(k2_fdiff_9,X1))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11452, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_seq_1(k2_fdiff_9,X1))|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1))).#trainneg
cnf(i_0_4397, plain, (k9_xtuple_0(k2_partfun1(k1_numbers,k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),X2))=X2|~r1_tarski(X1,k1_relset_1(k1_numbers,k2_fdiff_9))|~r1_tarski(X2,X1))).#trainneg
cnf(i_0_11217, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__4))=esk1_0|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_11896, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))))|r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_10828, plain, (m1_subset_1(k2_relset_1(k1_numbers,k1_sin_cos9),k1_zfmisc_1(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)))))).#trainneg
cnf(i_0_10890, negated_conjecture, (k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__4)|~v1_xreal_0(k1_real_1(esk1_0))|~r1_xxreal_0(k1_real_1(esk1_0),k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_11055, negated_conjecture, (r1_xxreal_0(esk1_0,k32_sin_cos)|r1_xxreal_0(k32_sin_cos,esk1_0))).#trainneg
cnf(i_0_10886, plain, (k9_xtuple_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))=k1_relset_1(k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1)))).#trainneg
cnf(i_0_11210, plain, (m1_subset_1(k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__2))),k1_numbers)|~v1_xreal_0(k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)))))).#trainneg
cnf(i_0_11884, plain, (X1=k1_real_1(k1_real_1(X2))|~v1_xreal_0(k1_real_1(k1_real_1(X2)))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11183, plain, (r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r1_xxreal_0(k1_real_1(np__1),k1_seq_1(k2_fdiff_9,X1))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11178, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X2)|~r1_xxreal_0(X2,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_11458, plain, (r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~v1_xxreal_0(X1)|~r1_xxreal_0(k1_real_1(np__1),k1_seq_1(k2_fdiff_9,X1))|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1))).#trainneg
cnf(i_0_11083, plain, (X1=X2|~v1_xreal_0(X3)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k1_rcomp_1(X2,X3))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_11859, plain, (X1=X2|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_rcomp_1(X2,X2)))).#trainneg
cnf(i_0_10877, plain, (k1_real_1(k10_real_1(k32_sin_cos,np__4))=k10_real_1(k32_sin_cos,np__4)|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_10897, negated_conjecture, (~r2_hidden(k1_real_1(esk1_0),a_0_1_sin_cos9))).#trainneg
cnf(i_0_11468, plain, (r1_xxreal_0(np__1,k1_real_1(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))))))).#trainneg
cnf(i_0_11592, plain, (r1_xxreal_0(np__1,k10_real_1(k32_sin_cos,np__2))|~r2_hidden(np__1,a_0_0_sin_cos9))).#trainneg
cnf(i_0_10911, plain, (r1_tarski(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)))|~r1_tarski(X1,k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_11456, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1))|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X2))).#trainneg
cnf(i_0_11731, plain, (r1_xxreal_0(X1,X3)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)|~r2_hidden(X3,k1_rcomp_1(X2,X1)))).#trainneg
cnf(i_0_11631, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_10894, negated_conjecture, (k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__4)|~v1_xxreal_0(k1_real_1(esk1_0))|~r1_xxreal_0(k1_real_1(esk1_0),k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_10939, plain, (~r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__2)),a_0_1_sin_cos9))).#trainneg
cnf(i_0_11453, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1))|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X2)|~r1_xxreal_0(X2,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_11062, plain, (v5_relat_1(k2_funct_1(k3_sincos10),k1_numbers))).#trainneg
cnf(i_0_11132, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k2_fdiff_9))|~r1_tarski(X1,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_11046, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))))|r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),esk1_0))).#trainneg
cnf(i_0_11376, negated_conjecture, (r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,esk1_0)|~r1_xxreal_0(esk1_0,X1)|~r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_11696, plain, (r1_xxreal_0(k1_real_1(np__1),k1_seq_1(k16_sin_cos,k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_11043, negated_conjecture, (r1_xxreal_0(esk1_0,k1_seq_1(k2_fdiff_9,X1))|r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),esk1_0))).#trainneg
cnf(i_0_10920, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r1_tarski(X1,k2_funct_1(k3_sincos10)))).#trainneg
cnf(i_0_10773, plain, (r1_xxreal_0(X1,k1_real_1(np__2))|~v1_xreal_0(X1)|~r1_xxreal_0(np__2,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_11049, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(np__2))|r1_xxreal_0(k1_real_1(np__2),esk1_0))).#trainneg
cnf(i_0_11042, negated_conjecture, (r1_xxreal_0(esk1_0,np__1)|r1_xxreal_0(np__1,esk1_0))).#trainneg
cnf(i_0_11331, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_seq_1(k2_fdiff_9,esk1_0))|r1_xxreal_0(X1,esk1_0)|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11459, plain, (k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)))=k1_real_1(np__1)|~r1_xxreal_0(k1_real_1(np__1),k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_10887, plain, (v4_relat_1(k2_funct_1(k3_sincos10),k1_numbers))).#trainneg
cnf(i_0_10761, plain, (k1_real_1(X1)=np__2|~v1_xxreal_0(k1_real_1(X1))|~r1_xxreal_0(np__2,k1_real_1(X1))|~r1_xxreal_0(k1_real_1(np__2),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11576, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_seq_1(k2_fdiff_9,esk1_0))|r1_xxreal_0(X1,esk1_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1))).#trainneg
cnf(i_0_11332, plain, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_seq_1(k2_fdiff_9,X2))|r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X2)|~r1_xxreal_0(X2,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_11651, plain, (v3_valued_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1)))).#trainneg
cnf(i_0_11709, plain, (r1_xxreal_0(X1,X2)|~r1_tarski(k1_relset_1(k1_numbers,k2_fdiff_9),k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1))|~m1_subset_1(k1_relset_1(k1_numbers,k2_fdiff_9),k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,k1_relset_1(k1_numbers,k2_fdiff_9))|~r2_hidden(X2,k1_relset_1(k1_numbers,k2_fdiff_9)))).#trainneg
cnf(i_0_11181, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X2),k1_seq_1(k2_fdiff_9,X1))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(X1,k1_numbers)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X2))).#trainneg
cnf(i_0_11404, plain, (r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X2,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r2_hidden(X2,a_0_2_sin_cos9))).#trainneg
cnf(i_0_11129, plain, (m1_subset_1(k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))),k1_zfmisc_1(k1_relset_1(k1_numbers,k2_fdiff_9))))).#trainneg
cnf(i_0_10821, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k2_relset_1(k1_numbers,k1_sin_cos9)))).#trainneg
cnf(i_0_11516, plain, (r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),k1_seq_1(k2_fdiff_9,X1))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1)|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11580, plain, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_seq_1(k2_fdiff_9,X2))|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X2))).#trainneg
cnf(i_0_11860, plain, (X1=k6_numbers|~v1_xreal_0(X1)|~r1_xxreal_0(k6_numbers,X1)|~r1_xxreal_0(X1,k6_numbers))).#trainneg
cnf(i_0_11519, plain, (r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))),k1_seq_1(k2_fdiff_9,X1))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1))).#trainneg
cnf(i_0_11927, plain, (r1_xxreal_0(k1_seq_1(k16_sin_cos,k10_real_1(k32_sin_cos,np__4)),np__1))).#trainneg
cnf(i_0_11866, plain, (X1=np__2|~v1_xreal_0(X1)|~r1_xxreal_0(np__2,X1)|~r1_xxreal_0(X1,np__2))).#trainneg
cnf(i_0_11716, plain, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(np__1))|~r2_hidden(np__1,a_0_1_sin_cos9))).#trainneg
cnf(i_0_11432, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k10_real_1(k32_sin_cos,np__4))|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),X1)|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),X2)|~r2_hidden(X2,a_0_2_sin_cos9))).#trainneg
cnf(i_0_11717, plain, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(np__2))|~r2_hidden(np__2,a_0_1_sin_cos9))).#trainneg
cnf(i_0_11061, plain, (k5_relat_1(k2_funct_1(k3_sincos10),X1)=k2_partfun1(k1_numbers,k1_numbers,k2_funct_1(k3_sincos10),X1))).#trainneg
cnf(i_0_11487, negated_conjecture, (r1_xxreal_0(esk1_0,k7_square_1(X1))|r1_xxreal_0(k7_square_1(X1),esk1_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10892, negated_conjecture, (k1_real_1(k1_real_1(esk1_0))=k1_real_1(k10_real_1(k32_sin_cos,np__4))|~r1_xxreal_0(k1_real_1(esk1_0),k10_real_1(k32_sin_cos,np__4))|~m1_subset_1(k1_real_1(esk1_0),k1_numbers))).#trainneg
cnf(i_0_11045, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(np__1))|r1_xxreal_0(k1_real_1(np__1),esk1_0))).#trainneg
cnf(i_0_11865, plain, (X1=np__1|~v1_xreal_0(X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_11195, plain, (k9_xtuple_0(k2_partfun1(k1_numbers,k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),X2))=X2|~r1_tarski(X2,k9_xtuple_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1)))|~v1_funct_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1)))).#trainneg
cnf(i_0_10762, plain, (r1_xxreal_0(np__2,k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k1_real_1(np__2)))).#trainneg
cnf(i_0_11577, plain, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_seq_1(k2_fdiff_9,X2))|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X2)|~r1_xxreal_0(X2,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_12059, plain, (m1_subset_1(k10_real_1(k32_sin_cos,np__4),a_0_2_sin_cos9))).#trainneg
cnf(i_0_12037, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_12080, negated_conjecture, (m1_subset_1(k1_real_1(esk1_0),k1_numbers))).#trainneg
cnf(i_0_12454, plain, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k1_real_1(k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_11951, plain, (~r2_hidden(k32_sin_cos,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12077, plain, (k2_relset_1(k1_numbers,k2_funct_1(k3_sincos10))=k10_xtuple_0(k2_funct_1(k3_sincos10)))).#trainneg
cnf(i_0_11957, plain, (~r1_xxreal_0(k6_numbers,k1_real_1(k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_12923, plain, (r1_tarski(k2_relset_1(k1_numbers,k1_sin_cos9),k1_numbers))).#trainneg
cnf(i_0_12321, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),esk1_0))|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_11958, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k6_numbers))).#trainneg
cnf(i_0_12118, plain, (m1_subset_1(k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_12266, negated_conjecture, (k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__4)|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12320, negated_conjecture, (r2_hidden(esk1_0,k2_relset_1(k1_numbers,k1_sin_cos9))|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12362, negated_conjecture, (~r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12127, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_11984, plain, (r1_xxreal_0(k1_real_1(X1),k10_real_1(k32_sin_cos,np__4))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11986, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12143, plain, (k1_relset_1(k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))=X1|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k2_fdiff_9))))).#trainneg
cnf(i_0_12707, negated_conjecture, (r1_xxreal_0(k1_real_1(esk1_0),k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)))|r1_xxreal_0(esk1_0,k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_12399, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)))|r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_12449, plain, (X1=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(k1_real_1(X2)),k1_numbers))).#trainneg
cnf(i_0_12141, plain, (k1_relset_1(k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))=k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_11985, plain, (r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13100, negated_conjecture, (r1_xxreal_0(k1_real_1(esk1_0),np__1)|r1_xxreal_0(esk1_0,k1_real_1(np__1)))).#trainneg
cnf(i_0_7969, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_funct_1(k3_sincos10),X2),k1_seq_1(k2_funct_1(k3_sincos10),X1))|~r2_hidden(X1,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers))|~r2_hidden(X2,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))).#trainneg
cnf(i_0_12685, plain, (v3_valued_0(k2_funct_1(k3_sincos10)))).#trainneg
cnf(i_0_12578, negated_conjecture, (r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,esk1_0)),k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),esk1_0))).#trainneg
cnf(i_0_12606, plain, (r1_tarski(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)))|~r1_tarski(X1,X2)|~r1_tarski(X2,k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_12342, negated_conjecture, (v1_xxreal_0(esk1_0)|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12224, plain, (r1_xxreal_0(k1_real_1(np__1),k10_real_1(k32_sin_cos,np__4))|r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(np__1)))).#trainneg
cnf(i_0_12587, plain, (m1_subset_1(k1_real_1(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)))),k1_numbers)|~v1_xreal_0(k1_real_1(k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))))))).#trainneg
cnf(i_0_12146, plain, (k1_relset_1(k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))=X1|~m1_subset_1(X1,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10))))).#trainneg
cnf(i_0_12330, negated_conjecture, (m1_subset_1(esk1_0,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),esk1_0))|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12419, negated_conjecture, (k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__4)|~v1_xreal_0(k1_real_1(esk1_0))|~r2_hidden(k1_real_1(esk1_0),a_0_2_sin_cos9))).#trainneg
cnf(i_0_11971, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(X1))|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12230, plain, (r1_xxreal_0(k1_real_1(X1),k10_real_1(k32_sin_cos,np__4))|r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12376, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_seq_1(k2_fdiff_9,X1))|r1_xxreal_0(esk1_0,X1)|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),esk1_0))|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12437, plain, (m1_subset_1(k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__2))),k1_numbers))).#trainneg
cnf(i_0_12909, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k2_fdiff_9)))|~r1_tarski(X1,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_12009, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_real_1(np__2))|r1_xxreal_0(k1_real_1(np__2),k1_real_1(k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_12273, negated_conjecture, (~r1_xxreal_0(k6_numbers,esk1_0)|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12509, negated_conjecture, (k1_real_1(k1_real_1(esk1_0))=k1_real_1(k10_real_1(k32_sin_cos,np__2))|~r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_12309, negated_conjecture, (k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)))=k1_seq_1(k2_fdiff_9,esk1_0)|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12239, negated_conjecture, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),esk1_0)|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_12277, negated_conjecture, (r1_xxreal_0(k1_real_1(esk1_0),k10_real_1(k32_sin_cos,np__4))|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12015, plain, (r1_xxreal_0(X1,k1_real_1(np__2))|r1_xxreal_0(k1_real_1(np__2),X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_12453, plain, (~r1_xxreal_0(k1_real_1(np__1),k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_13050, plain, (X1=np__2|~r1_xxreal_0(np__2,X1)|~r1_xxreal_0(X1,np__2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12259, plain, (k1_relset_1(k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),X2))=X2|~v1_relat_1(k2_partfun1(k1_numbers,k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),X2))|~v1_funct_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))|~r1_tarski(X1,k1_relset_1(k1_numbers,k2_fdiff_9))|~r1_tarski(X2,X1))).#trainneg
cnf(i_0_12236, plain, (r1_xxreal_0(X1,k10_real_1(k32_sin_cos,np__4))|r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12554, negated_conjecture, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k1_real_1(esk1_0))|~r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_12301, negated_conjecture, (r1_xxreal_0(k1_real_1(X1),k10_real_1(k32_sin_cos,np__4))|~r1_xxreal_0(esk1_0,X1)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12550, negated_conjecture, (k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__2)|~r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_12511, negated_conjecture, (v1_xreal_0(k1_real_1(esk1_0))|~r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_12311, negated_conjecture, (r1_xxreal_0(esk1_0,esk23_1(X1))|~r2_hidden(X1,a_0_2_sin_cos9)|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12418, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2))))|~r1_tarski(X1,k2_relset_1(k1_numbers,k1_sin_cos9)))).#trainneg
cnf(i_0_11982, plain, (X1=k1_real_1(k10_real_1(k32_sin_cos,np__4))|~v1_xreal_0(X1)|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_12910, plain, (r1_tarski(k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4))),k1_numbers))).#trainneg
cnf(i_0_12613, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~v2_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0)))).#trainneg
cnf(i_0_12323, negated_conjecture, (r1_tarski(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),esk1_0),k2_relset_1(k1_numbers,k3_sincos10))|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12271, negated_conjecture, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(X1))|~r1_xxreal_0(X1,esk1_0)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_13070, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),X1)|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),X2)|~r2_hidden(X2,a_0_2_sin_cos9)|~r2_hidden(X1,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12272, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12201, plain, (k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2))=k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))|~r1_tarski(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)),k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_12126, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_11991, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k10_real_1(k32_sin_cos,np__4))|~v1_xreal_0(X1)|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1))).#trainneg
cnf(i_0_11973, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(np__2))|~r1_xxreal_0(np__2,k1_real_1(k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_12343, negated_conjecture, (m1_subset_1(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),esk1_0),k1_zfmisc_1(k1_numbers))|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12199, plain, (m1_subset_1(k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_zfmisc_1(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)))))).#trainneg
cnf(i_0_12874, negated_conjecture, (r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,esk1_0)|~r1_xxreal_0(esk1_0,X1)|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_13163, plain, (~r1_xxreal_0(k1_real_1(np__1),k1_real_1(np__2)))).#trainneg
cnf(i_0_12617, negated_conjecture, (r1_xxreal_0(np__1,k1_real_1(k1_seq_1(k2_fdiff_9,esk1_0)))|r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_11983, plain, (X1=k1_real_1(k10_real_1(k32_sin_cos,np__4))|~v1_xxreal_0(X1)|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12763, negated_conjecture, (r1_xxreal_0(k1_real_1(esk1_0),np__2)|r1_xxreal_0(esk1_0,k1_real_1(np__2)))).#trainneg
cnf(i_0_12010, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(np__2))|r1_xxreal_0(k1_real_1(np__2),k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_12603, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2))))|~r1_tarski(X1,k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_12388, negated_conjecture, (r1_tarski(X1,k2_relset_1(k1_numbers,k3_sincos10))|~r1_tarski(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),esk1_0))|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12513, negated_conjecture, (r1_xxreal_0(k1_real_1(esk1_0),k1_real_1(esk1_0))|~r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_12300, negated_conjecture, (r1_xxreal_0(X1,k10_real_1(k32_sin_cos,np__4))|~r1_xxreal_0(esk1_0,k1_real_1(X1))|~m1_subset_1(X1,k1_numbers)|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12925, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_12368, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_seq_1(k2_fdiff_9,X2))|r1_xxreal_0(X1,X2)|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),esk1_0))|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),esk1_0))|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12011, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(np__2))|r1_xxreal_0(k1_real_1(np__2),k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_12182, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_funct_1(k3_sincos10),X2),k1_seq_1(k2_funct_1(k3_sincos10),X1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k2_relset_1(k1_numbers,k3_sincos10))|~r2_hidden(X2,k2_relset_1(k1_numbers,k3_sincos10)))).#trainneg
cnf(i_0_11582, plain, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_real_1(np__1))|r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1))).#trainneg
cnf(i_0_12012, plain, (r1_xxreal_0(k1_real_1(np__1),k1_real_1(np__2))|r1_xxreal_0(k1_real_1(np__2),k1_real_1(np__1)))).#trainneg
cnf(i_0_11857, plain, (r1_xxreal_0(np__1,k10_real_1(k32_sin_cos,np__4))|~r2_hidden(np__1,a_0_1_sin_cos9))).#trainneg
cnf(i_0_13078, plain, (m1_subset_1(k2_partfun1(k1_numbers,k1_numbers,k2_funct_1(k3_sincos10),X1),k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k2_funct_1(k3_sincos10),k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_12001, plain, (X1=k10_real_1(k32_sin_cos,np__4)|~v1_xreal_0(X1)|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),X1)|~r1_xxreal_0(X1,k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_12440, plain, (X1=k1_real_1(k10_real_1(k32_sin_cos,np__4))|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10910, plain, (r1_tarski(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)))|~r1_tarski(X1,k1_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_12254, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|r1_xxreal_0(esk1_0,X1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1))).#trainneg
cnf(i_0_13213, negated_conjecture, (v1_xreal_0(k1_real_1(esk1_0)))).#trainneg
cnf(i_0_13598, plain, (v1_xreal_0(k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)))))).#trainneg
cnf(i_0_13220, negated_conjecture, (r2_hidden(k1_real_1(esk1_0),k1_numbers))).#trainneg
cnf(i_0_13570, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1)|r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13214, negated_conjecture, (v1_xxreal_0(k1_real_1(esk1_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_13231, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(esk1_0))|r1_xxreal_0(k1_real_1(esk1_0),esk1_0))).#trainneg
cnf(i_0_14371, plain, (m1_subset_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k2_funct_1(k3_sincos10),k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10))))).#trainneg
cnf(i_0_13226, negated_conjecture, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_real_1(esk1_0))|r1_xxreal_0(k1_real_1(esk1_0),k1_real_1(k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_13599, plain, (v1_xxreal_0(k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__2))))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_13503, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),np__1)|r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(np__1)))).#trainneg
cnf(i_0_13227, negated_conjecture, (r1_xxreal_0(k1_real_1(esk1_0),k1_real_1(k10_real_1(k32_sin_cos,np__2)))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(esk1_0)))).#trainneg
cnf(i_0_13223, negated_conjecture, (r1_xxreal_0(k1_real_1(esk1_0),k1_seq_1(k2_fdiff_9,X1))|r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_real_1(esk1_0)))).#trainneg
cnf(i_0_13641, plain, (r1_xxreal_0(np__2,k10_real_1(k32_sin_cos,np__4))|r1_xxreal_0(k1_real_1(np__2),k1_real_1(k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_13361, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k2_fdiff_9))))).#trainneg
cnf(i_0_14209, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),np__2)|r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(np__2)))).#trainneg
cnf(i_0_13222, negated_conjecture, (r1_xxreal_0(X1,k1_real_1(esk1_0))|r1_xxreal_0(k1_real_1(esk1_0),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14239, negated_conjecture, (r1_xxreal_0(k1_real_1(k1_real_1(esk1_0)),esk1_0)|~r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_13412, plain, (X1=X2|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13605, plain, (r2_hidden(k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__2))),k1_numbers))).#trainneg
cnf(i_0_13221, negated_conjecture, (r1_xxreal_0(k1_real_1(X1),k1_real_1(esk1_0))|r1_xxreal_0(k1_real_1(esk1_0),k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14238, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k1_real_1(esk1_0)))|~r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_13224, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_real_1(esk1_0))|r1_xxreal_0(k1_real_1(esk1_0),k6_numbers))).#trainneg
cnf(i_0_13704, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_real_1(np__1))|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_13554, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(X1))|r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),X1)|r2_hidden(X1,a_0_1_sin_cos9)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14369, plain, (m1_subset_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_funct_1(k3_sincos10),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_13225, negated_conjecture, (r1_xxreal_0(k1_real_1(np__2),k1_real_1(esk1_0))|r1_xxreal_0(k1_real_1(esk1_0),k1_real_1(np__2)))).#trainneg
cnf(i_0_13257, plain, (~v2_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_13913, negated_conjecture, (r1_tarski(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),esk1_0),k1_relset_1(k1_numbers,k2_fdiff_9))|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_14218, plain, (m1_subset_1(k1_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__4)),k1_zfmisc_1(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)))))).#trainneg
cnf(i_0_13885, plain, (k1_real_1(k10_real_1(k32_sin_cos,np__4))=k1_real_1(k10_real_1(k32_sin_cos,np__2))|r2_hidden(k1_real_1(k10_real_1(k32_sin_cos,np__4)),a_0_0_sin_cos9))).#trainneg
cnf(i_0_13416, plain, (k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))=k1_numbers|~r1_tarski(k1_numbers,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))))).#trainneg
cnf(i_0_11976, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_13366, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,k2_fdiff_9))))).#trainneg
cnf(i_0_12669, negated_conjecture, (k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__4)|~r1_xxreal_0(k1_real_1(esk1_0),k10_real_1(k32_sin_cos,np__4))|~v1_xreal_0(k1_real_1(esk1_0)))).#trainneg
cnf(i_0_12067, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k2_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k6_numbers)))).#trainneg
cnf(i_0_11862, plain, (X1=k1_real_1(np__2)|~v1_xreal_0(X1)|~r1_xxreal_0(k1_real_1(np__2),X1)|~r1_xxreal_0(X1,k1_real_1(np__2)))).#trainneg
cnf(i_0_12451, plain, (X1=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(X2),k1_numbers))).#trainneg
cnf(i_0_13132, plain, (X1=np__1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13616, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__2))))|r1_xxreal_0(k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__2))),esk1_0))).#trainneg
cnf(i_0_13229, negated_conjecture, (r1_xxreal_0(k1_real_1(k1_real_1(esk1_0)),k1_real_1(k10_real_1(k32_sin_cos,np__4)))|r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_real_1(k1_real_1(esk1_0))))).#trainneg
cnf(i_0_12013, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k1_real_1(np__2))|r1_xxreal_0(k1_real_1(np__2),k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_12994, plain, (X1=k6_numbers|~r1_xxreal_0(k6_numbers,X1)|~r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13362, plain, (r1_tarski(X1,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k2_fdiff_9))))).#trainneg
cnf(i_0_12916, plain, (r1_tarski(k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)),k1_numbers))).#trainneg
cnf(i_0_13611, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__2))))|r1_xxreal_0(k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__2))),k1_real_1(k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_10769, plain, (r1_xxreal_0(k1_real_1(X1),np__2)|~r1_xxreal_0(k1_real_1(np__2),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13479, plain, (r1_tarski(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k10_real_1(k32_sin_cos,np__2)))|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k3_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2)))))).#trainneg
cnf(i_0_14294, plain, (r1_xxreal_0(np__2,k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__2))))|r1_xxreal_0(k1_real_1(np__2),k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_11024, plain, (X1=k10_real_1(k32_sin_cos,np__4)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),X2)|~r1_xxreal_0(X2,k10_real_1(k32_sin_cos,np__4))|~r1_xxreal_0(X2,X1)|~r2_hidden(X1,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12389, negated_conjecture, (m1_subset_1(esk1_0,k2_relset_1(k1_numbers,k1_sin_cos9))|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12801, plain, (k1_real_1(X1)=np__2|~r1_xxreal_0(np__2,k1_real_1(X1))|~r1_xxreal_0(k1_real_1(np__2),X1)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(k1_real_1(X1)))).#trainneg
cnf(i_0_12367, negated_conjecture, (m1_subset_1(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),esk1_0),k1_zfmisc_1(k1_relset_1(k1_numbers,k2_fdiff_9)))|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12387, negated_conjecture, (m1_subset_1(k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),esk1_0),k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10)))|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_12038, negated_conjecture, (k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__2)|~v1_xreal_0(k1_real_1(esk1_0))|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k1_real_1(esk1_0)))).#trainneg
cnf(i_0_10775, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__2)|~v1_xreal_0(X1)|~r1_xxreal_0(k1_real_1(np__2),X1))).#trainneg
cnf(i_0_13328, plain, (r1_xxreal_0(k1_real_1(np__2),k10_real_1(k32_sin_cos,np__4))|r1_xxreal_0(np__2,k1_real_1(k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_12332, negated_conjecture, (r1_tarski(X1,k1_relset_1(k1_numbers,k2_fdiff_9))|~r1_tarski(X1,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),esk1_0))|~r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_11975, plain, (k1_real_1(X1)=k10_real_1(k32_sin_cos,np__4)|~v1_xxreal_0(k1_real_1(X1))|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(X1))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11337, plain, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),k1_real_1(np__1))|r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13037, plain, (k1_seq_1(k2_fdiff_9,X1)=np__2|~r1_xxreal_0(np__2,k1_seq_1(k2_fdiff_9,X1))|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),np__2))).#trainneg
cnf(i_0_7746, plain, (esk14_3(X2,X1,X3)=X1|r1_xxreal_0(X3,X2)|~v1_xxreal_0(X1)|~v1_xreal_0(esk14_3(X2,X1,X3))|~r1_xxreal_0(esk14_3(X2,X1,X3),X1)|~v1_xxreal_0(X3)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_7613, plain, (X1=k10_real_1(k32_sin_cos,np__4)|~v1_xreal_0(X1)|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),X1)|~r2_hidden(X1,a_0_2_sin_cos9))).#trainneg
cnf(i_0_8117, plain, (k9_xtuple_0(k2_partfun1(k1_numbers,k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1),X2))=X2|~r1_tarski(X2,k9_xtuple_0(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1)))|~v1_relat_1(k2_partfun1(k1_numbers,k1_numbers,k2_fdiff_9,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_relset_1(k1_numbers,k3_sincos10))))).#trainneg
cnf(i_0_15165, negated_conjecture, (k1_real_1(k1_real_1(esk1_0))=esk1_0)).#trainneg
cnf(i_0_15461, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)))))).#trainneg
cnf(i_0_14614, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__2))))|r1_xxreal_0(k1_real_1(esk1_0),k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_15260, plain, (k1_real_1(k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__2))))=k1_real_1(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_15307, plain, (k1_real_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)))=k10_real_1(k32_sin_cos,np__2))).#trainneg
cnf(i_0_14663, plain, (m1_subset_1(k1_relset_1(k1_numbers,k2_fdiff_9),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_14537, negated_conjecture, (v1_xxreal_0(k1_real_1(esk1_0)))).#trainneg
cnf(i_0_14471, negated_conjecture, (r1_xxreal_0(k1_real_1(esk1_0),k1_real_1(esk1_0)))).#trainneg
cnf(i_0_14454, negated_conjecture, (r1_xxreal_0(k1_real_1(esk1_0),k10_real_1(k32_sin_cos,np__4))|r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_14564, negated_conjecture, (r1_xxreal_0(esk1_0,k10_real_1(k32_sin_cos,np__4))|r1_xxreal_0(k1_real_1(esk1_0),k1_real_1(k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_14617, negated_conjecture, (r1_xxreal_0(esk1_0,k10_real_1(k32_sin_cos,np__2))|r1_xxreal_0(k1_real_1(esk1_0),k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_15039, negated_conjecture, (r1_xxreal_0(np__2,esk1_0)|r1_xxreal_0(k1_real_1(np__2),k1_real_1(esk1_0)))).#trainneg
cnf(i_0_15316, plain, (k1_real_1(k1_real_1(k32_sin_cos))=k32_sin_cos)).#trainneg
cnf(i_0_15222, plain, (k1_real_1(k1_real_1(np__4))=np__4)).#trainneg
cnf(i_0_14627, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k1_seq_1(k2_fdiff_9,X1)))|r1_xxreal_0(k1_real_1(esk1_0),k1_seq_1(k2_fdiff_9,X1)))).#trainneg
cnf(i_0_14479, negated_conjecture, (r1_xxreal_0(X1,k1_real_1(esk1_0))|r1_xxreal_0(k1_real_1(esk1_0),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_14462, negated_conjecture, (r1_xxreal_0(np__2,k1_real_1(esk1_0))|r1_xxreal_0(k1_real_1(esk1_0),np__2))).#trainneg
cnf(i_0_14465, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_real_1(esk1_0))|r1_xxreal_0(k1_real_1(esk1_0),k32_sin_cos))).#trainneg
cnf(i_0_15404, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),np__2)|r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k1_real_1(np__2)))).#trainneg
cnf(i_0_15118, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k2_fdiff_9)))).#trainneg
cnf(i_0_14453, negated_conjecture, (r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,esk1_0)|~r1_xxreal_0(esk1_0,X1))).#trainneg
cnf(i_0_15204, plain, (k1_real_1(k1_real_1(k1_seq_1(k2_fdiff_9,X1)))=k1_seq_1(k2_fdiff_9,X1))).#trainneg
cnf(i_0_14927, negated_conjecture, (r1_xxreal_0(k6_numbers,esk1_0)|r1_xxreal_0(k6_numbers,k1_real_1(esk1_0)))).#trainneg
cnf(i_0_14705, negated_conjecture, (r1_xxreal_0(k1_real_1(esk1_0),k10_real_1(X1,X2))|r1_xxreal_0(k10_real_1(X1,X2),k1_real_1(esk1_0))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14976, plain, (r1_xxreal_0(np__2,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),np__2)|r2_hidden(np__2,a_0_1_sin_cos9))).#trainneg
cnf(i_0_14940, negated_conjecture, (r2_hidden(k1_real_1(esk1_0),a_0_2_sin_cos9)|~r1_xxreal_0(k1_real_1(esk1_0),k10_real_1(k32_sin_cos,np__4))|~r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_14864, plain, (X1=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(k1_real_1(k1_real_1(X2)),X1)|~m1_subset_1(k1_real_1(k1_real_1(X2)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_14566, negated_conjecture, (r2_hidden(k1_real_1(esk1_0),a_0_2_sin_cos9)|r1_xxreal_0(k1_real_1(esk1_0),k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k1_real_1(esk1_0),k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_14692, negated_conjecture, (r1_xxreal_0(k1_real_1(esk1_0),k7_square_1(X1))|r1_xxreal_0(k7_square_1(X1),k1_real_1(esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14992, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),np__2)|r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(np__2))|r2_hidden(k1_real_1(np__2),a_0_1_sin_cos9))).#trainneg
cnf(i_0_15108, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k4_rcomp_1(k6_numbers,k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_14897, negated_conjecture, (r1_xxreal_0(k1_real_1(esk1_0),k1_real_1(k1_real_1(X1)))|r1_xxreal_0(k1_real_1(k1_real_1(X1)),k1_real_1(esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14569, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__4))=k1_real_1(esk1_0)|r1_xxreal_0(k1_real_1(esk1_0),k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~v1_xxreal_0(k1_real_1(esk1_0))|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),esk1_0))).#trainneg
cnf(i_0_14928, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__4))=k1_real_1(esk1_0)|~r1_xxreal_0(k1_real_1(esk1_0),k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_14943, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__4))=k1_real_1(esk1_0)|~v1_xxreal_0(k1_real_1(esk1_0))|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),esk1_0)|~r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_15289, plain, (k1_real_1(k1_real_1(X1))=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(X2,k1_real_1(k1_real_1(X1)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_real_1(k1_real_1(X1)),k1_numbers)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14888, negated_conjecture, (r1_xxreal_0(k1_real_1(esk1_0),k1_real_1(k1_seq_1(k2_fdiff_9,X1)))|r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,X1)),k1_real_1(esk1_0)))).#trainneg
cnf(i_0_15663, plain, (k1_real_1(k10_real_1(k32_sin_cos,np__4))=np__2|r1_xxreal_0(np__2,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(np__2)))).#trainneg
cnf(i_0_15661, plain, (k1_real_1(np__2)=k10_real_1(k32_sin_cos,np__4)|r1_xxreal_0(np__2,k1_real_1(k10_real_1(k32_sin_cos,np__4)))|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(np__2)))).#trainneg
cnf(i_0_15040, negated_conjecture, (k1_real_1(esk1_0)=k1_real_1(np__2)|r1_xxreal_0(k1_real_1(np__2),k1_real_1(esk1_0))|~v1_xxreal_0(k1_real_1(esk1_0))|~r1_xxreal_0(esk1_0,np__2))).#trainneg
cnf(i_0_15101, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(np__1))|~r1_xxreal_0(np__1,k1_real_1(k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_14860, plain, (esk14_3(X2,X1,X3)=X1|~r1_xxreal_0(esk14_3(X2,X1,X3),X1)|~m1_subset_1(esk14_3(X2,X1,X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xxreal_0(X3)|~v1_xreal_0(X1)|~r2_hidden(X2,k3_rcomp_1(X1,X3)))).#trainneg
cnf(i_0_14616, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__2))=k1_real_1(esk1_0)|r1_xxreal_0(k1_real_1(esk1_0),k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~v1_xxreal_0(k1_real_1(esk1_0))|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),esk1_0))).#trainneg
cnf(i_0_15103, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),np__2)|~r1_xxreal_0(k1_real_1(np__2),k1_real_1(k10_real_1(k32_sin_cos,np__4))))).#trainneg
cnf(i_0_14989, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),np__1)|r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(np__1))|r2_hidden(k1_real_1(np__1),a_0_1_sin_cos9))).#trainneg
cnf(i_0_14514, plain, (r1_xxreal_0(X1,k10_real_1(k32_sin_cos,np__4))|r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(k1_real_1(X1)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers))).#trainneg
cnf(i_0_14532, plain, (r2_hidden(X1,a_0_2_sin_cos9)|r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(X1))|~r1_xxreal_0(X1,k10_real_1(k32_sin_cos,np__4))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14889, negated_conjecture, (r1_xxreal_0(k1_real_1(esk1_0),k1_real_1(np__4))|r1_xxreal_0(k1_real_1(np__4),k1_real_1(esk1_0)))).#trainneg
cnf(i_0_14688, negated_conjecture, (r1_xxreal_0(k1_real_1(esk1_0),k8_real_1(X1,X2))|r1_xxreal_0(k8_real_1(X1,X2),k1_real_1(esk1_0))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14685, negated_conjecture, (r1_xxreal_0(k1_real_1(esk1_0),np__15)|r1_xxreal_0(np__15,k1_real_1(esk1_0)))).#trainneg
cnf(i_0_14463, negated_conjecture, (r1_xxreal_0(np__1,k1_real_1(esk1_0))|r1_xxreal_0(k1_real_1(esk1_0),np__1))).#trainneg
cnf(i_0_14816, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(X2),k1_real_1(X1))|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14470, negated_conjecture, (r1_xxreal_0(k1_real_1(np__1),k1_real_1(esk1_0))|r1_xxreal_0(k1_real_1(esk1_0),k1_real_1(np__1)))).#trainneg
cnf(i_0_16101, negated_conjecture, (r1_xxreal_0(esk1_0,k10_real_1(k32_sin_cos,np__4))|k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__4))).#trainneg
cnf(i_0_16048, negated_conjecture, (k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__2)|k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__4))).#trainneg
cnf(i_0_15866, negated_conjecture, (k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__2)|k1_real_1(k10_real_1(k32_sin_cos,np__4))=esk1_0)).#trainneg
cnf(i_0_16962, plain, (r2_hidden(k6_numbers,a_0_2_sin_cos9))).#trainneg
cnf(i_0_15860, negated_conjecture, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k1_real_1(esk1_0))|k1_real_1(k10_real_1(k32_sin_cos,np__4))=esk1_0)).#trainneg
cnf(i_0_16057, negated_conjecture, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k1_real_1(esk1_0))|k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__4))).#trainneg
cnf(i_0_16023, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_relset_1(k1_numbers,k2_fdiff_9)))).#trainneg
cnf(i_0_17027, negated_conjecture, (r1_xxreal_0(k1_real_1(np__1),esk1_0)|r1_xxreal_0(np__1,k1_real_1(esk1_0)))).#trainneg
cnf(i_0_15785, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(X1))|~r1_xxreal_0(X1,k1_real_1(esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15963, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k1_real_1(X1))|~r1_xxreal_0(X1,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_17015, negated_conjecture, (r1_xxreal_0(k1_real_1(np__15),esk1_0)|r1_xxreal_0(np__15,k1_real_1(esk1_0)))).#trainneg
cnf(i_0_15781, negated_conjecture, (r1_xxreal_0(np__2,k1_real_1(esk1_0))|~r1_xxreal_0(esk1_0,k1_real_1(np__2)))).#trainneg
cnf(i_0_16777, negated_conjecture, (r1_xxreal_0(esk1_0,k1_seq_1(k2_fdiff_9,X1))|r1_xxreal_0(k1_real_1(esk1_0),k1_real_1(k1_seq_1(k2_fdiff_9,X1))))).#trainneg
cnf(i_0_16401, plain, (r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,X1)),k1_real_1(X2))|~r1_xxreal_0(X2,k1_seq_1(k2_fdiff_9,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_real_1(k1_seq_1(k2_fdiff_9,X1)),k1_numbers))).#trainneg
cnf(i_0_15770, negated_conjecture, (r1_xxreal_0(k1_real_1(esk1_0),k1_real_1(X1))|~r1_xxreal_0(X1,esk1_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15823, negated_conjecture, (X1=esk1_0|~v1_xreal_0(X1)|~r1_xxreal_0(esk1_0,X1)|~r1_xxreal_0(X1,esk1_0))).#trainneg
cnf(i_0_15809, negated_conjecture, (r1_xxreal_0(k1_real_1(X1),k1_real_1(esk1_0))|~r1_xxreal_0(esk1_0,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_17057, plain, (k1_real_1(k1_real_1(X1))=k1_real_1(X2)|~r1_xxreal_0(X2,k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~r1_xxreal_0(k1_real_1(X2),X1)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_17310, negated_conjecture, (k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__4)|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),esk1_0))).#trainneg
cnf(i_0_17931, plain, (m1_subset_1(k6_numbers,a_0_2_sin_cos9))).#trainneg
cnf(i_0_17813, negated_conjecture, (k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)))=k1_seq_1(k2_fdiff_9,esk1_0)|k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__2))).#trainneg
cnf(i_0_18283, plain, (k1_real_1(k1_real_1(X1))=X2|~r1_xxreal_0(k1_real_1(X2),k1_real_1(X1))|~r1_xxreal_0(X2,X1)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_17890, negated_conjecture, (r2_hidden(esk1_0,a_0_2_sin_cos9)|k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__2))).#trainneg
cnf(i_0_17825, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),esk1_0))|k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__2))).#trainneg
cnf(i_0_17175, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(k1_real_1(k10_real_1(k32_sin_cos,np__2)),esk1_0))|k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0)).#trainneg
cnf(i_0_17895, negated_conjecture, (m1_subset_1(esk1_0,a_0_2_sin_cos9)|k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__2))).#trainneg
cnf(i_0_17824, negated_conjecture, (r2_hidden(esk1_0,k2_relset_1(k1_numbers,k1_sin_cos9))|k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__2))).#trainneg
cnf(i_0_17241, negated_conjecture, (k10_real_1(k32_sin_cos,np__4)=esk1_0|k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),esk1_0))).#trainneg
cnf(i_0_17245, negated_conjecture, (m1_subset_1(esk1_0,a_0_2_sin_cos9)|k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0)).#trainneg
cnf(i_0_17135, negated_conjecture, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),esk1_0)|k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0)).#trainneg
cnf(i_0_17240, negated_conjecture, (r2_hidden(esk1_0,a_0_2_sin_cos9)|k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0)).#trainneg
cnf(i_0_17174, negated_conjecture, (r2_hidden(esk1_0,k2_relset_1(k1_numbers,k1_sin_cos9))|k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0)).#trainneg
cnf(i_0_17550, negated_conjecture, (r1_xxreal_0(esk1_0,k10_real_1(k32_sin_cos,np__2))|k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__2))).#trainneg
cnf(i_0_17504, negated_conjecture, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__4)),esk1_0)|k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__2))).#trainneg
cnf(i_0_17399, negated_conjecture, (r1_xxreal_0(esk1_0,k10_real_1(k32_sin_cos,np__2))|k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__4))).#trainneg
cnf(i_0_17260, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),esk1_0))).#trainneg
cnf(i_0_19624, negated_conjecture, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),esk1_0))).#trainneg
cnf(i_0_18398, negated_conjecture, (k10_real_1(k32_sin_cos,np__4)=esk1_0|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),esk1_0))).#trainneg
cnf(i_0_19340, negated_conjecture, (r2_hidden(esk1_0,a_0_2_sin_cos9)|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),esk1_0))).#trainneg
cnf(i_0_19307, negated_conjecture, (k1_real_1(esk1_0)=k10_real_1(k32_sin_cos,np__2)|r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_19406, negated_conjecture, (r2_hidden(esk1_0,a_0_2_sin_cos9)|~r2_hidden(esk1_0,a_0_0_sin_cos9))).#trainneg
cnf(i_0_19069, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0|~r2_hidden(k1_seq_1(k2_fdiff_9,esk1_0),k1_rcomp_1(k1_seq_1(k2_fdiff_9,esk1_0),k1_real_1(np__1))))).#trainneg
cnf(i_0_17255, negated_conjecture, (r1_xxreal_0(esk1_0,k10_real_1(k32_sin_cos,np__2))|k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0)).#trainneg
cnf(i_0_19005, negated_conjecture, (k1_real_1(k1_seq_1(k2_fdiff_9,esk1_0))=k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4))|k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0)).#trainneg
cnf(i_0_19432, negated_conjecture, (r1_xxreal_0(esk1_0,k10_real_1(k32_sin_cos,np__2))|r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_19378, negated_conjecture, (k1_seq_1(k2_fdiff_9,esk1_0)=k1_real_1(np__1)|r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_17039, plain, (X1=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(X2),X1)|~r1_xxreal_0(X2,k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_19029, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,esk1_0),k1_real_1(np__2)))).#trainneg
cnf(i_0_14294, plain, (r1_xxreal_0(k1_real_1(np__2),k1_real_1(k10_real_1(k32_sin_cos,np__2)))|r1_xxreal_0(np__2,k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_17119, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0|~r1_xxreal_0(k6_numbers,esk1_0))).#trainneg
cnf(i_0_17271, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0|r2_hidden(esk1_0,a_0_0_sin_cos9))).#trainneg
cnf(i_0_19381, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__4))))|r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_15861, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__4))=esk1_0|~r2_hidden(esk1_0,a_0_0_sin_cos9))).#trainneg
cnf(i_0_18222, negated_conjecture, (X1=esk1_0|~r1_xxreal_0(esk1_0,X1)|~r1_xxreal_0(X1,esk1_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_19769, negated_conjecture, (~r1_xxreal_0(k6_numbers,esk1_0))).#trainneg
cnf(i_0_20528, negated_conjecture, (r2_hidden(esk1_0,a_0_0_sin_cos9))).#trainneg
cnf(i_0_15860, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__4))=esk1_0)).#trainneg
cnf(i_0_9580, negated_conjecture, (r1_xxreal_0(esk1_0,k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_12554, negated_conjecture, (~r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_19770, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers))).#trainneg
cnf(i_0_15866, negated_conjecture, (k1_real_1(k10_real_1(k32_sin_cos,np__4))=esk1_0|k10_real_1(k32_sin_cos,np__4)=k10_real_1(k32_sin_cos,np__2))).#trainneg
cnf(i_0_17890, negated_conjecture, (k10_real_1(k32_sin_cos,np__4)=k10_real_1(k32_sin_cos,np__2)|r2_hidden(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_17550, negated_conjecture, (k10_real_1(k32_sin_cos,np__4)=k10_real_1(k32_sin_cos,np__2)|r1_xxreal_0(esk1_0,k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_20282, plain, (k1_real_1(esk15_3(X2,k1_real_1(X1),X3))=X1|~r1_xxreal_0(esk15_3(X2,k1_real_1(X1),X3),k1_real_1(X1))|~m1_subset_1(esk15_3(X2,k1_real_1(X1),X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X3)|~r2_hidden(X2,k1_rcomp_1(k1_real_1(X1),X3)))).#trainneg
cnf(i_0_17813, negated_conjecture, (k1_real_1(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)))=k1_seq_1(k2_fdiff_9,esk1_0)|k10_real_1(k32_sin_cos,np__4)=k10_real_1(k32_sin_cos,np__2))).#trainneg
cnf(i_0_19784, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|~r1_xxreal_0(k1_real_1(X1),k10_real_1(k32_sin_cos,np__4))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14627, negated_conjecture, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_seq_1(k2_fdiff_9,X1))|r1_xxreal_0(esk1_0,k1_real_1(k1_seq_1(k2_fdiff_9,X1))))).#trainneg
cnf(i_0_16764, plain, (k1_real_1(k1_real_1(X1))=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(X1,X2)|~m1_subset_1(k1_real_1(k1_real_1(X1)),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,X1))).#trainneg
cnf(i_0_18710, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k10_real_1(k32_sin_cos,np__2)))|m1_subset_1(esk1_0,a_0_2_sin_cos9))).#trainneg
cnf(i_0_15770, negated_conjecture, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),k1_real_1(X1))|~r1_xxreal_0(X1,esk1_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_18207, negated_conjecture, (k1_seq_1(k2_fdiff_9,X1)=esk1_0|~r1_xxreal_0(esk1_0,k1_seq_1(k2_fdiff_9,X1))|~r1_xxreal_0(k1_seq_1(k2_fdiff_9,X1),esk1_0))).#trainneg
cnf(i_0_18395, plain, (k1_real_1(k1_real_1(X1))=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(k1_real_1(X2),k1_real_1(X1))|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_18058, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k1_real_1(X1))|r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),X1)|r2_hidden(X1,a_0_0_sin_cos9)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_19780, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(np__2))|~r1_xxreal_0(np__2,k10_real_1(k32_sin_cos,np__4)))).#trainneg
cnf(i_0_20210, plain, (X1=X2|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(k1_real_1(X2),k1_real_1(X1))|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_15948, plain, (r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),k1_real_1(X1))|~r1_xxreal_0(X1,k10_real_1(k32_sin_cos,np__2))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15965, plain, (r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k1_real_1(np__2))|~r1_xxreal_0(np__2,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_16438, plain, (r1_xxreal_0(k1_real_1(X1),k1_seq_1(k2_fdiff_9,X2))|~r1_xxreal_0(k1_real_1(k1_seq_1(k2_fdiff_9,X2)),X1)|~m1_subset_1(k1_real_1(k1_seq_1(k2_fdiff_9,X2)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15987, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(k10_real_1(k32_sin_cos,np__2)))|~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13100, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(np__1))|r1_xxreal_0(k10_real_1(k32_sin_cos,np__4),np__1))).#trainneg
cnf(i_0_15985, plain, (r1_xxreal_0(k1_real_1(X1),k10_real_1(k32_sin_cos,np__2))|~r1_xxreal_0(k1_real_1(k10_real_1(k32_sin_cos,np__2)),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_18586, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(k1_real_1(X2),X1)|~r1_xxreal_0(X2,k1_real_1(X1))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_19018, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_fdiff_9,k10_real_1(k32_sin_cos,np__4)),k1_real_1(X1))|k1_real_1(k10_real_1(k32_sin_cos,np__2))=esk1_0|~r1_xxreal_0(X1,k1_seq_1(k2_fdiff_9,esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 437
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 617
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 612
# Processed clauses                    : 3092
# ...of these trivial                  : 151
# ...subsumed                          : 1059
# ...remaining for further processing  : 1882
# Other redundant clauses eliminated   : 17
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 45
# Backward-rewritten                   : 628
# Generated clauses                    : 21182
# ...of the previous two non-trivial   : 18616
# Contextual simplify-reflections      : 133
# Paramodulations                      : 21131
# Factorizations                       : 11
# Equation resolutions                 : 36
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1197
#    Positive orientable unit clauses  : 386
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 64
#    Non-unit-clauses                  : 744
# Current number of unprocessed clauses: 14711
# ...number of literals in the above   : 63420
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2560
# Proof object given clauses           : 297
# Proof search given clauses           : 1882
# Clause-clause subsumption calls (NU) : 252157
# Rec. Clause-clause subsumption calls : 88683
# Non-unit clause-clause subsumptions  : 943
# Unit Clause-clause subsumption calls : 14634
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 335
# BW rewrite match successes           : 98
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 472707

# -------------------------------------------------
# User time                : 24.423 s
# System time              : 0.708 s
# Total time               : 25.131 s
# Maximum resident set size: 360792 pages
24.42user 0.81system 0:25.48elapsed 99%CPU (0avgtext+0avgdata 360792maxresident)k
0inputs+784outputs (0major+207151minor)pagefaults 0swaps
