2020-08-31 10:41:28.484378: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_67
2020-08-31 10:41:28.558839: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-08-31 10:41:28.668025: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-08-31 10:41:28.993715: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-08-31 10:41:30.416334: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 1931957 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_67' loaded (query=256; context=1024; weigths=real; len_mult=0.000000).
# Preprocessing time       : 2.523 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', t6_boole)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', redefinition_k6_numbers)).
fof(t9_nat_4, axiom, ![X1]:(v7_ordinal1(X1)=>![X2]:(v7_ordinal1(X2)=>![X3]:(v7_ordinal1(X3)=>![X4]:(v7_ordinal1(X4)=>~((((X1=k2_nat_1(k24_binop_2(X2,X3),X4)&~(r1_xxreal_0(X2,X4)))&~(r1_xxreal_0(X4,k6_numbers)))&r1_nat_d(X2,X1))))))), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', t9_nat_4)).
fof(spc0_boole, axiom, v1_xboole_0(np__0), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc0_boole)).
fof(commutativity_k2_nat_1, axiom, ![X1, X2]:((m1_subset_1(X1,k5_numbers)&v7_ordinal1(X2))=>k2_nat_1(X1,X2)=k2_nat_1(X2,X1)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', commutativity_k2_nat_1)).
fof(cc8_ordinal1, axiom, ![X1]:(m1_subset_1(X1,k4_ordinal1)=>v7_ordinal1(X1)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', cc8_ordinal1)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', redefinition_k5_numbers)).
fof(redefinition_k2_nat_1, axiom, ![X1, X2]:((m1_subset_1(X1,k5_numbers)&v7_ordinal1(X2))=>k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', redefinition_k2_nat_1)).
fof(redefinition_k24_binop_2, axiom, ![X1, X2]:((v7_ordinal1(X1)&v7_ordinal1(X2))=>k24_binop_2(X1,X2)=k3_xcmplx_0(X1,X2)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', redefinition_k24_binop_2)).
fof(commutativity_k24_binop_2, axiom, ![X1, X2]:((v7_ordinal1(X1)&v7_ordinal1(X2))=>k24_binop_2(X1,X2)=k24_binop_2(X2,X1)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', commutativity_k24_binop_2)).
fof(fc2_nat_1, axiom, ![X1, X2]:((v7_ordinal1(X1)&v7_ordinal1(X2))=>v7_ordinal1(k3_xcmplx_0(X1,X2))), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', fc2_nat_1)).
fof(cc17_membered, axiom, ![X1]:(v6_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v7_ordinal1(X2))), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', cc17_membered)).
fof(fc6_membered, axiom, v6_membered(k4_ordinal1), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', fc6_membered)).
fof(l36_nat_4, axiom, ![X1]:(m2_subset_1(X1,k1_numbers,k5_numbers)=>(~(r1_xxreal_0(np__841,X1))=>![X2]:(m2_subset_1(X2,k1_numbers,k5_numbers)=>~((((((((((((~(r1_xxreal_0(X2,np__1))&r1_xxreal_0(k4_nat_1(X2,X2),X1))&v1_int_2(X2))&X2!=np__2)&X2!=np__3)&X2!=np__5)&X2!=np__7)&X2!=np__11)&X2!=np__13)&X2!=np__17)&X2!=np__19)&X2!=np__23))))), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', l36_nat_4)).
fof(t14_nat_4, axiom, ![X1]:(v7_ordinal1(X1)=>(v1_int_2(X1)<=>(~(r1_xxreal_0(X1,np__1))&![X2]:(m2_subset_1(X2,k1_numbers,k5_numbers)=>~((((~(r1_xxreal_0(X2,np__1))&r1_xxreal_0(k4_nat_1(X2,X2),X1))&v1_int_2(X2))&r1_nat_d(X2,X1))))))), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', t14_nat_4)).
fof(spc8_numerals, axiom, (((v2_xxreal_0(np__8)&m2_subset_1(np__8,k1_numbers,k5_numbers))&m1_subset_1(np__8,k5_numbers))&m1_subset_1(np__8,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc8_numerals)).
fof(spc19_numerals, axiom, (((v2_xxreal_0(np__19)&m2_subset_1(np__19,k1_numbers,k5_numbers))&m1_subset_1(np__19,k5_numbers))&m1_subset_1(np__19,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc19_numerals)).
fof(rqRealMult__k3_xcmplx_0__r19_r8_r152, axiom, k3_xcmplx_0(np__19,np__8)=np__152, file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqRealMult__k3_xcmplx_0__r19_r8_r152)).
fof(rqLessOrEqual__r1_xxreal_0__r19_r11, axiom, ~(r1_xxreal_0(np__19,np__11)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqLessOrEqual__r1_xxreal_0__r19_r11)).
fof(rqLessOrEqual__r1_xxreal_0__r11_r0, axiom, ~(r1_xxreal_0(np__11,np__0)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqLessOrEqual__r1_xxreal_0__r11_r0)).
fof(spc7_numerals, axiom, (((v2_xxreal_0(np__7)&m2_subset_1(np__7,k1_numbers,k5_numbers))&m1_subset_1(np__7,k5_numbers))&m1_subset_1(np__7,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc7_numerals)).
fof(spc23_numerals, axiom, (((v2_xxreal_0(np__23)&m2_subset_1(np__23,k1_numbers,k5_numbers))&m1_subset_1(np__23,k5_numbers))&m1_subset_1(np__23,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc23_numerals)).
fof(rqRealAdd__k2_xcmplx_0__r11_r152_r163, axiom, k2_xcmplx_0(np__11,np__152)=np__163, file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqRealAdd__k2_xcmplx_0__r11_r152_r163)).
fof(spc11_numerals, axiom, (((v2_xxreal_0(np__11)&m2_subset_1(np__11,k1_numbers,k5_numbers))&m1_subset_1(np__11,k5_numbers))&m1_subset_1(np__11,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc11_numerals)).
fof(rqRealMult__k3_xcmplx_0__r7_r23_r161, axiom, k3_xcmplx_0(np__7,np__23)=np__161, file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqRealMult__k3_xcmplx_0__r7_r23_r161)).
fof(rqLessOrEqual__r1_xxreal_0__r23_r2, axiom, ~(r1_xxreal_0(np__23,np__2)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqLessOrEqual__r1_xxreal_0__r23_r2)).
fof(rqLessOrEqual__r1_xxreal_0__r2_r0, axiom, ~(r1_xxreal_0(np__2,np__0)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqLessOrEqual__r1_xxreal_0__r2_r0)).
fof(rqRealAdd__k2_xcmplx_0__r2_r161_r163, axiom, k2_xcmplx_0(np__2,np__161)=np__163, file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqRealAdd__k2_xcmplx_0__r2_r161_r163)).
fof(spc2_numerals, axiom, (((v2_xxreal_0(np__2)&m2_subset_1(np__2,k1_numbers,k5_numbers))&m1_subset_1(np__2,k5_numbers))&m1_subset_1(np__2,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc2_numerals)).
fof(spc9_numerals, axiom, (((v2_xxreal_0(np__9)&m2_subset_1(np__9,k1_numbers,k5_numbers))&m1_subset_1(np__9,k5_numbers))&m1_subset_1(np__9,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc9_numerals)).
fof(spc17_numerals, axiom, (((v2_xxreal_0(np__17)&m2_subset_1(np__17,k1_numbers,k5_numbers))&m1_subset_1(np__17,k5_numbers))&m1_subset_1(np__17,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc17_numerals)).
fof(rqRealMult__k3_xcmplx_0__r17_r9_r153, axiom, k3_xcmplx_0(np__17,np__9)=np__153, file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqRealMult__k3_xcmplx_0__r17_r9_r153)).
fof(spc153_numerals, axiom, (((v2_xxreal_0(np__153)&m2_subset_1(np__153,k1_numbers,k5_numbers))&m1_subset_1(np__153,k5_numbers))&m1_subset_1(np__153,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc153_numerals)).
fof(spc10_numerals, axiom, (((v2_xxreal_0(np__10)&m2_subset_1(np__10,k1_numbers,k5_numbers))&m1_subset_1(np__10,k5_numbers))&m1_subset_1(np__10,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc10_numerals)).
fof(rqLessOrEqual__r1_xxreal_0__r17_r10, axiom, ~(r1_xxreal_0(np__17,np__10)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqLessOrEqual__r1_xxreal_0__r17_r10)).
fof(rqLessOrEqual__r1_xxreal_0__r10_r0, axiom, ~(r1_xxreal_0(np__10,np__0)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqLessOrEqual__r1_xxreal_0__r10_r0)).
fof(spc13_numerals, axiom, (((v2_xxreal_0(np__13)&m2_subset_1(np__13,k1_numbers,k5_numbers))&m1_subset_1(np__13,k5_numbers))&m1_subset_1(np__13,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc13_numerals)).
fof(rqRealAdd__k2_xcmplx_0__r153_r10_r163, axiom, k2_xcmplx_0(np__153,np__10)=np__163, file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqRealAdd__k2_xcmplx_0__r153_r10_r163)).
fof(rqRealMult__k3_xcmplx_0__r13_r12_r156, axiom, k3_xcmplx_0(np__13,np__12)=np__156, file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqRealMult__k3_xcmplx_0__r13_r12_r156)).
fof(spc156_numerals, axiom, (((v2_xxreal_0(np__156)&m2_subset_1(np__156,k1_numbers,k5_numbers))&m1_subset_1(np__156,k5_numbers))&m1_subset_1(np__156,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc156_numerals)).
fof(spc12_numerals, axiom, (((v2_xxreal_0(np__12)&m2_subset_1(np__12,k1_numbers,k5_numbers))&m1_subset_1(np__12,k5_numbers))&m1_subset_1(np__12,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc12_numerals)).
fof(rqLessOrEqual__r1_xxreal_0__r13_r7, axiom, ~(r1_xxreal_0(np__13,np__7)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqLessOrEqual__r1_xxreal_0__r13_r7)).
fof(rqLessOrEqual__r1_xxreal_0__r7_r0, axiom, ~(r1_xxreal_0(np__7,np__0)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqLessOrEqual__r1_xxreal_0__r7_r0)).
fof(rqRealAdd__k2_xcmplx_0__r156_r7_r163, axiom, k2_xcmplx_0(np__156,np__7)=np__163, file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqRealAdd__k2_xcmplx_0__r156_r7_r163)).
fof(spc32_numerals, axiom, (((v2_xxreal_0(np__32)&m2_subset_1(np__32,k1_numbers,k5_numbers))&m1_subset_1(np__32,k5_numbers))&m1_subset_1(np__32,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc32_numerals)).
fof(spc5_numerals, axiom, (((v2_xxreal_0(np__5)&m2_subset_1(np__5,k1_numbers,k5_numbers))&m1_subset_1(np__5,k5_numbers))&m1_subset_1(np__5,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc5_numerals)).
fof(rqRealMult__k3_xcmplx_0__r32_r5_r160, axiom, k3_xcmplx_0(np__32,np__5)=np__160, file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqRealMult__k3_xcmplx_0__r32_r5_r160)).
fof(spc160_numerals, axiom, (((v2_xxreal_0(np__160)&m2_subset_1(np__160,k1_numbers,k5_numbers))&m1_subset_1(np__160,k5_numbers))&m1_subset_1(np__160,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc160_numerals)).
fof(spc3_numerals, axiom, (((v2_xxreal_0(np__3)&m2_subset_1(np__3,k1_numbers,k5_numbers))&m1_subset_1(np__3,k5_numbers))&m1_subset_1(np__3,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc3_numerals)).
fof(rqLessOrEqual__r1_xxreal_0__r5_r3, axiom, ~(r1_xxreal_0(np__5,np__3)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqLessOrEqual__r1_xxreal_0__r5_r3)).
fof(rqLessOrEqual__r1_xxreal_0__r3_r0, axiom, ~(r1_xxreal_0(np__3,np__0)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqLessOrEqual__r1_xxreal_0__r3_r0)).
fof(rqRealAdd__k2_xcmplx_0__r160_r3_r163, axiom, k2_xcmplx_0(np__160,np__3)=np__163, file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqRealAdd__k2_xcmplx_0__r160_r3_r163)).
fof(spc161_numerals, axiom, (((v2_xxreal_0(np__161)&m2_subset_1(np__161,k1_numbers,k5_numbers))&m1_subset_1(np__161,k5_numbers))&m1_subset_1(np__161,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc161_numerals)).
fof(rqLessOrEqual__r1_xxreal_0__r7_r2, axiom, ~(r1_xxreal_0(np__7,np__2)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqLessOrEqual__r1_xxreal_0__r7_r2)).
fof(rqRealAdd__k2_xcmplx_0__r161_r2_r163, axiom, k2_xcmplx_0(np__161,np__2)=np__163, file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqRealAdd__k2_xcmplx_0__r161_r2_r163)).
fof(rqRealMult__k3_xcmplx_0__r11_r14_r154, axiom, k3_xcmplx_0(np__11,np__14)=np__154, file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqRealMult__k3_xcmplx_0__r11_r14_r154)).
fof(spc154_numerals, axiom, (((v2_xxreal_0(np__154)&m2_subset_1(np__154,k1_numbers,k5_numbers))&m1_subset_1(np__154,k5_numbers))&m1_subset_1(np__154,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc154_numerals)).
fof(spc14_numerals, axiom, (((v2_xxreal_0(np__14)&m2_subset_1(np__14,k1_numbers,k5_numbers))&m1_subset_1(np__14,k5_numbers))&m1_subset_1(np__14,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc14_numerals)).
fof(rqLessOrEqual__r1_xxreal_0__r11_r9, axiom, ~(r1_xxreal_0(np__11,np__9)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqLessOrEqual__r1_xxreal_0__r11_r9)).
fof(rqLessOrEqual__r1_xxreal_0__r9_r0, axiom, ~(r1_xxreal_0(np__9,np__0)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqLessOrEqual__r1_xxreal_0__r9_r0)).
fof(rqRealAdd__k2_xcmplx_0__r154_r9_r163, axiom, k2_xcmplx_0(np__154,np__9)=np__163, file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqRealAdd__k2_xcmplx_0__r154_r9_r163)).
fof(spc54_numerals, axiom, (((v2_xxreal_0(np__54)&m2_subset_1(np__54,k1_numbers,k5_numbers))&m1_subset_1(np__54,k5_numbers))&m1_subset_1(np__54,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc54_numerals)).
fof(rqRealMult__k3_xcmplx_0__r54_r3_r162, axiom, k3_xcmplx_0(np__54,np__3)=np__162, file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqRealMult__k3_xcmplx_0__r54_r3_r162)).
fof(rqLessOrEqual__r1_xxreal_0__r3_r1, axiom, ~(r1_xxreal_0(np__3,np__1)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqLessOrEqual__r1_xxreal_0__r3_r1)).
fof(rqLessOrEqual__r1_xxreal_0__r1_r0, axiom, ~(r1_xxreal_0(np__1,np__0)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqLessOrEqual__r1_xxreal_0__r1_r0)).
fof(rqRealAdd__k2_xcmplx_0__r1_r162_r163, axiom, k2_xcmplx_0(np__1,np__162)=np__163, file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqRealAdd__k2_xcmplx_0__r1_r162_r163)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc1_numerals)).
fof(spc81_numerals, axiom, (((v2_xxreal_0(np__81)&m2_subset_1(np__81,k1_numbers,k5_numbers))&m1_subset_1(np__81,k5_numbers))&m1_subset_1(np__81,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc81_numerals)).
fof(rqRealMult__k3_xcmplx_0__r81_r2_r162, axiom, k3_xcmplx_0(np__81,np__2)=np__162, file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqRealMult__k3_xcmplx_0__r81_r2_r162)).
fof(rqLessOrEqual__r1_xxreal_0__r2_r1, axiom, ~(r1_xxreal_0(np__2,np__1)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqLessOrEqual__r1_xxreal_0__r2_r1)).
fof(t35_nat_4, conjecture, v1_int_2(np__163), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', t35_nat_4)).
fof(spc163_numerals, axiom, (((v2_xxreal_0(np__163)&m2_subset_1(np__163,k1_numbers,k5_numbers))&m1_subset_1(np__163,k5_numbers))&m1_subset_1(np__163,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', spc163_numerals)).
fof(rqLessOrEqual__r1_xxreal_0__r841_r163, axiom, ~(r1_xxreal_0(np__841,np__163)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqLessOrEqual__r1_xxreal_0__r841_r163)).
fof(rqLessOrEqual__r1_xxreal_0__r163_r1, axiom, ~(r1_xxreal_0(np__163,np__1)), file('/scratch/mptp/hardpred/preds__0.005//t35_nat_4', rqLessOrEqual__r1_xxreal_0__r163_r1)).
fof(c_0_74, plain, ![X31]:(~v1_xboole_0(X31)|X31=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
cnf(c_0_75, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_76, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
fof(c_0_77, plain, ![X10, X11, X12, X13]:(~v7_ordinal1(X10)|(~v7_ordinal1(X11)|(~v7_ordinal1(X12)|(~v7_ordinal1(X13)|(X10!=k2_nat_1(k24_binop_2(X11,X12),X13)|r1_xxreal_0(X11,X13)|r1_xxreal_0(X13,k6_numbers)|~r1_nat_d(X11,X10)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_nat_4])])])])).
cnf(c_0_78, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_75])).
cnf(c_0_79, plain, (k1_xboole_0=k6_numbers), inference(evalgc,[status(thm)],[c_0_76])).
cnf(c_0_80, plain, (v1_xboole_0(np__0)), inference(split_conjunct,[status(thm)],[spc0_boole])).
cnf(c_0_81, plain, (r1_xxreal_0(X2,X4)|r1_xxreal_0(X4,k6_numbers)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|X1!=k2_nat_1(k24_binop_2(X2,X3),X4)|~r1_nat_d(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_82, plain, (X1=k6_numbers|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_83, plain, (v1_xboole_0(np__0)), inference(evalgc,[status(thm)],[c_0_80])).
fof(c_0_84, plain, ![X22, X23]:(~m1_subset_1(X22,k5_numbers)|~v7_ordinal1(X23)|k2_nat_1(X22,X23)=k2_nat_1(X23,X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k2_nat_1])])).
fof(c_0_85, plain, ![X36]:(~m1_subset_1(X36,k4_ordinal1)|v7_ordinal1(X36)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_ordinal1])])).
cnf(c_0_86, plain, (r1_xxreal_0(X2,X4)|r1_xxreal_0(X4,k6_numbers)|X1!=k2_nat_1(k24_binop_2(X2,X3),X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_nat_d(X2,X1)), inference(evalgc,[status(thm)],[c_0_81])).
cnf(c_0_87, plain, (k6_numbers=np__0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83])])).
cnf(c_0_88, plain, (k2_nat_1(X1,X2)=k2_nat_1(X2,X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_89, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_90, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
cnf(c_0_91, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(k24_binop_2(X2,X4),X1)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_86, c_0_87])])).
cnf(c_0_92, plain, (k2_nat_1(X1,X2)=k2_nat_1(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_88])).
cnf(c_0_93, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_94, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_90])).
fof(c_0_95, plain, ![X16, X17]:(~m1_subset_1(X16,k5_numbers)|~v7_ordinal1(X17)|k2_nat_1(X16,X17)=k2_xcmplx_0(X16,X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_nat_1])])).
cnf(c_0_96, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(X1,k24_binop_2(X2,X4))|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(k24_binop_2(X2,X4))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92])])).
cnf(c_0_97, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)), inference(rw,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_98, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_95])).
fof(c_0_99, plain, ![X14, X15]:(~v7_ordinal1(X14)|~v7_ordinal1(X15)|k24_binop_2(X14,X15)=k3_xcmplx_0(X14,X15)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k24_binop_2])])).
fof(c_0_100, plain, ![X20, X21]:(~v7_ordinal1(X20)|~v7_ordinal1(X21)|k24_binop_2(X20,X21)=k24_binop_2(X21,X20)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k24_binop_2])])).
cnf(c_0_101, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(X1,k24_binop_2(X2,X4))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(k24_binop_2(X2,X4))|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3)), inference(csr,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_102, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_98])).
cnf(c_0_103, plain, (k24_binop_2(X1,X2)=k3_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_99])).
fof(c_0_104, plain, ![X37, X38]:(~v7_ordinal1(X37)|~v7_ordinal1(X38)|v7_ordinal1(k3_xcmplx_0(X37,X38))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc2_nat_1])])).
fof(c_0_105, plain, ![X26, X27]:(~v6_membered(X26)|(~m1_subset_1(X27,X26)|v7_ordinal1(X27))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc17_membered])])])).
cnf(c_0_106, plain, (v6_membered(k4_ordinal1)), inference(split_conjunct,[status(thm)],[fc6_membered])).
fof(c_0_107, plain, ![X8, X9]:(~m2_subset_1(X8,k1_numbers,k5_numbers)|(r1_xxreal_0(np__841,X8)|(~m2_subset_1(X9,k1_numbers,k5_numbers)|(r1_xxreal_0(X9,np__1)|~r1_xxreal_0(k4_nat_1(X9,X9),X8)|~v1_int_2(X9)|X9=np__2|X9=np__3|X9=np__5|X9=np__7|X9=np__11|X9=np__13|X9=np__17|X9=np__19|X9=np__23)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l36_nat_4])])])])).
fof(c_0_108, plain, ![X5, X6]:(((~r1_xxreal_0(X5,np__1)|~v1_int_2(X5)|~v7_ordinal1(X5))&(~m2_subset_1(X6,k1_numbers,k5_numbers)|(r1_xxreal_0(X6,np__1)|~r1_xxreal_0(k4_nat_1(X6,X6),X5)|~v1_int_2(X6)|~r1_nat_d(X6,X5))|~v1_int_2(X5)|~v7_ordinal1(X5)))&((m2_subset_1(esk1_1(X5),k1_numbers,k5_numbers)|r1_xxreal_0(X5,np__1)|v1_int_2(X5)|~v7_ordinal1(X5))&((((~r1_xxreal_0(esk1_1(X5),np__1)|r1_xxreal_0(X5,np__1)|v1_int_2(X5)|~v7_ordinal1(X5))&(r1_xxreal_0(k4_nat_1(esk1_1(X5),esk1_1(X5)),X5)|r1_xxreal_0(X5,np__1)|v1_int_2(X5)|~v7_ordinal1(X5)))&(v1_int_2(esk1_1(X5))|r1_xxreal_0(X5,np__1)|v1_int_2(X5)|~v7_ordinal1(X5)))&(r1_nat_d(esk1_1(X5),X5)|r1_xxreal_0(X5,np__1)|v1_int_2(X5)|~v7_ordinal1(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t14_nat_4])])])])])])).
cnf(c_0_109, plain, (k24_binop_2(X1,X2)=k24_binop_2(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_110, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X1,k24_binop_2(X2,X4))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(k24_binop_2(X2,X4))|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102])])).
cnf(c_0_111, plain, (k24_binop_2(X1,X2)=k3_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_103])).
cnf(c_0_112, plain, (v7_ordinal1(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_113, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_114, plain, (m1_subset_1(np__8,k5_numbers)), inference(split_conjunct,[status(thm)],[spc8_numerals])).
cnf(c_0_115, plain, (v6_membered(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_106])).
cnf(c_0_116, plain, (m1_subset_1(np__19,k5_numbers)), inference(split_conjunct,[status(thm)],[spc19_numerals])).
cnf(c_0_117, plain, (r1_xxreal_0(np__841,X1)|r1_xxreal_0(X2,np__1)|X2=np__2|X2=np__3|X2=np__5|X2=np__7|X2=np__11|X2=np__13|X2=np__17|X2=np__19|X2=np__23|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(k4_nat_1(X2,X2),X1)|~v1_int_2(X2)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_118, plain, (r1_xxreal_0(k4_nat_1(esk1_1(X1),esk1_1(X1)),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_108])).
cnf(c_0_119, plain, (k24_binop_2(X1,X2)=k24_binop_2(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_109])).
cnf(c_0_120, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X1,k3_xcmplx_0(X2,X4))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(k3_xcmplx_0(X2,X4))|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111])])).
cnf(c_0_121, plain, (v7_ordinal1(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_112])).
cnf(c_0_122, plain, (k3_xcmplx_0(np__19,np__8)=np__152), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__r19_r8_r152])).
cnf(c_0_123, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_113])).
cnf(c_0_124, plain, (m1_subset_1(np__8,k5_numbers)), inference(evalgc,[status(thm)],[c_0_114])).
cnf(c_0_125, plain, (v6_membered(k5_numbers)), inference(rw,[status(thm)],[c_0_115, c_0_94])).
cnf(c_0_126, plain, (m1_subset_1(np__19,k5_numbers)), inference(evalgc,[status(thm)],[c_0_116])).
fof(c_0_127, plain, ~r1_xxreal_0(np__19,np__11), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r19_r11])).
fof(c_0_128, plain, ~r1_xxreal_0(np__11,np__0), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r11_r0])).
cnf(c_0_129, plain, (X2=np__2|X2=np__3|X2=np__5|X2=np__7|X2=np__11|X2=np__13|X2=np__17|X2=np__19|X2=np__23|r1_xxreal_0(X2,np__1)|r1_xxreal_0(np__841,X1)|~v1_int_2(X2)|~r1_xxreal_0(k4_nat_1(X2,X2),X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_117])).
cnf(c_0_130, plain, (v1_int_2(X1)|r1_xxreal_0(X1,np__1)|r1_xxreal_0(k4_nat_1(esk1_1(X1),esk1_1(X1)),X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_118])).
cnf(c_0_131, plain, (v1_int_2(esk1_1(X1))|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_108])).
cnf(c_0_132, plain, (m2_subset_1(esk1_1(X1),k1_numbers,k5_numbers)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_108])).
cnf(c_0_133, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(X1)|~r1_xxreal_0(esk1_1(X1),np__1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_108])).
cnf(c_0_134, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X1,k24_binop_2(X4,X2))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(k24_binop_2(X4,X2))|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_119])])).
cnf(c_0_135, plain, (m1_subset_1(np__7,k5_numbers)), inference(split_conjunct,[status(thm)],[spc7_numerals])).
cnf(c_0_136, plain, (m1_subset_1(np__23,k5_numbers)), inference(split_conjunct,[status(thm)],[spc23_numerals])).
cnf(c_0_137, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X1,k3_xcmplx_0(X2,X4))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3)), inference(csr,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_138, plain, (k3_xcmplx_0(np__19,np__8)=np__152), inference(evalgc,[status(thm)],[c_0_122])).
cnf(c_0_139, plain, (v7_ordinal1(np__8)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_124]), c_0_125])])])).
cnf(c_0_140, plain, (v7_ordinal1(np__19)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_126]), c_0_125])])])).
cnf(c_0_141, plain, (k2_xcmplx_0(np__11,np__152)=np__163), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r11_r152_r163])).
cnf(c_0_142, plain, (m1_subset_1(np__11,k5_numbers)), inference(split_conjunct,[status(thm)],[spc11_numerals])).
cnf(c_0_143, plain, (~r1_xxreal_0(np__19,np__11)), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_144, plain, (~r1_xxreal_0(np__11,np__0)), inference(split_conjunct,[status(thm)],[c_0_128])).
cnf(c_0_145, plain, (r1_nat_d(esk1_1(X1),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_108])).
cnf(c_0_146, plain, (esk1_1(X1)=np__23|esk1_1(X1)=np__19|esk1_1(X1)=np__17|esk1_1(X1)=np__13|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|r1_xxreal_0(esk1_1(X1),np__1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|~m2_subset_1(esk1_1(X1),k1_numbers,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_int_2(esk1_1(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130])])).
cnf(c_0_147, plain, (v1_int_2(X1)|v1_int_2(esk1_1(X1))|r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_131])).
cnf(c_0_148, plain, (v1_int_2(X1)|r1_xxreal_0(X1,np__1)|m2_subset_1(esk1_1(X1),k1_numbers,k5_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_132])).
cnf(c_0_149, plain, (v1_int_2(X1)|r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1)|~r1_xxreal_0(esk1_1(X1),np__1)), inference(evalgc,[status(thm)],[c_0_133])).
cnf(c_0_150, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X1,k3_xcmplx_0(X4,X2))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(k3_xcmplx_0(X4,X2))|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_111])])).
cnf(c_0_151, plain, (k3_xcmplx_0(np__7,np__23)=np__161), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__r7_r23_r161])).
cnf(c_0_152, plain, (m1_subset_1(np__7,k5_numbers)), inference(evalgc,[status(thm)],[c_0_135])).
cnf(c_0_153, plain, (m1_subset_1(np__23,k5_numbers)), inference(evalgc,[status(thm)],[c_0_136])).
fof(c_0_154, plain, ~r1_xxreal_0(np__23,np__2), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r23_r2])).
fof(c_0_155, plain, ~r1_xxreal_0(np__2,np__0), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r2_r0])).
cnf(c_0_156, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__19,X1)|X2!=k2_xcmplx_0(X1,np__152)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__19,X2)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_138]), c_0_139]), c_0_140])])])).
cnf(c_0_157, plain, (k2_xcmplx_0(np__11,np__152)=np__163), inference(evalgc,[status(thm)],[c_0_141])).
cnf(c_0_158, plain, (m1_subset_1(np__11,k5_numbers)), inference(evalgc,[status(thm)],[c_0_142])).
cnf(c_0_159, plain, (~r1_xxreal_0(np__19,np__11)), inference(evalgc,[status(thm)],[c_0_143])).
cnf(c_0_160, plain, (~r1_xxreal_0(np__11,np__0)), inference(evalgc,[status(thm)],[c_0_144])).
cnf(c_0_161, plain, (v1_int_2(X1)|r1_xxreal_0(X1,np__1)|r1_nat_d(esk1_1(X1),X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_145])).
cnf(c_0_162, plain, (esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|esk1_1(X1)=np__17|esk1_1(X1)=np__19|esk1_1(X1)=np__23|r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__841,X1)|v1_int_2(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_146, c_0_147]), c_0_148]), c_0_149])).
cnf(c_0_163, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X1,k3_xcmplx_0(X4,X2))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3)), inference(csr,[status(thm)],[c_0_150, c_0_121])).
cnf(c_0_164, plain, (k3_xcmplx_0(np__7,np__23)=np__161), inference(evalgc,[status(thm)],[c_0_151])).
cnf(c_0_165, plain, (v7_ordinal1(np__7)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_152]), c_0_125])])])).
cnf(c_0_166, plain, (v7_ordinal1(np__23)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_153]), c_0_125])])])).
cnf(c_0_167, plain, (k2_xcmplx_0(np__2,np__161)=np__163), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r2_r161_r163])).
cnf(c_0_168, plain, (m1_subset_1(np__2,k5_numbers)), inference(split_conjunct,[status(thm)],[spc2_numerals])).
cnf(c_0_169, plain, (~r1_xxreal_0(np__23,np__2)), inference(split_conjunct,[status(thm)],[c_0_154])).
cnf(c_0_170, plain, (~r1_xxreal_0(np__2,np__0)), inference(split_conjunct,[status(thm)],[c_0_155])).
cnf(c_0_171, plain, (X1!=np__163|~r1_nat_d(np__19,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_157]), c_0_158])]), c_0_159]), c_0_160])])).
cnf(c_0_172, plain, (r1_nat_d(np__19,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|esk1_1(X1)=np__17|esk1_1(X1)=np__23|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_162])])).
cnf(c_0_173, plain, (m1_subset_1(np__9,k5_numbers)), inference(split_conjunct,[status(thm)],[spc9_numerals])).
cnf(c_0_174, plain, (m1_subset_1(np__17,k5_numbers)), inference(split_conjunct,[status(thm)],[spc17_numerals])).
cnf(c_0_175, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__23,X1)|X2!=k2_xcmplx_0(X1,np__161)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__23,X2)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_164]), c_0_165]), c_0_166])])])).
cnf(c_0_176, plain, (k2_xcmplx_0(np__2,np__161)=np__163), inference(evalgc,[status(thm)],[c_0_167])).
cnf(c_0_177, plain, (m1_subset_1(np__2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_168])).
cnf(c_0_178, plain, (~r1_xxreal_0(np__23,np__2)), inference(evalgc,[status(thm)],[c_0_169])).
cnf(c_0_179, plain, (~r1_xxreal_0(np__2,np__0)), inference(evalgc,[status(thm)],[c_0_170])).
cnf(c_0_180, plain, (esk1_1(X1)=np__23|esk1_1(X1)=np__17|esk1_1(X1)=np__13|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__163|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_172])])).
cnf(c_0_181, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(k24_binop_2(X2,X4),X1)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(k24_binop_2(X2,X4),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_102])])).
cnf(c_0_182, plain, (k3_xcmplx_0(np__17,np__9)=np__153), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__r17_r9_r153])).
cnf(c_0_183, plain, (m1_subset_1(np__153,k5_numbers)), inference(split_conjunct,[status(thm)],[spc153_numerals])).
cnf(c_0_184, plain, (m1_subset_1(np__9,k5_numbers)), inference(evalgc,[status(thm)],[c_0_173])).
cnf(c_0_185, plain, (m1_subset_1(np__17,k5_numbers)), inference(evalgc,[status(thm)],[c_0_174])).
cnf(c_0_186, plain, (m1_subset_1(np__10,k5_numbers)), inference(split_conjunct,[status(thm)],[spc10_numerals])).
fof(c_0_187, plain, ~r1_xxreal_0(np__17,np__10), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r17_r10])).
fof(c_0_188, plain, ~r1_xxreal_0(np__10,np__0), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r10_r0])).
cnf(c_0_189, plain, (m1_subset_1(np__13,k5_numbers)), inference(split_conjunct,[status(thm)],[spc13_numerals])).
cnf(c_0_190, plain, (X1!=np__163|~r1_nat_d(np__23,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_176]), c_0_177])]), c_0_178]), c_0_179])])).
cnf(c_0_191, plain, (r1_nat_d(np__23,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|esk1_1(X1)=np__17|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_180])])).
cnf(c_0_192, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(k3_xcmplx_0(X2,X4),X1)|~m1_subset_1(k3_xcmplx_0(X2,X4),k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_111])])).
cnf(c_0_193, plain, (k3_xcmplx_0(np__17,np__9)=np__153), inference(evalgc,[status(thm)],[c_0_182])).
cnf(c_0_194, plain, (m1_subset_1(np__153,k5_numbers)), inference(evalgc,[status(thm)],[c_0_183])).
cnf(c_0_195, plain, (v7_ordinal1(np__9)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_184]), c_0_125])])])).
cnf(c_0_196, plain, (v7_ordinal1(np__17)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_185]), c_0_125])])])).
cnf(c_0_197, plain, (k2_xcmplx_0(np__153,np__10)=np__163), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r153_r10_r163])).
cnf(c_0_198, plain, (m1_subset_1(np__10,k5_numbers)), inference(evalgc,[status(thm)],[c_0_186])).
cnf(c_0_199, plain, (~r1_xxreal_0(np__17,np__10)), inference(split_conjunct,[status(thm)],[c_0_187])).
cnf(c_0_200, plain, (~r1_xxreal_0(np__10,np__0)), inference(split_conjunct,[status(thm)],[c_0_188])).
cnf(c_0_201, plain, (k3_xcmplx_0(np__13,np__12)=np__156), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__r13_r12_r156])).
cnf(c_0_202, plain, (m1_subset_1(np__156,k5_numbers)), inference(split_conjunct,[status(thm)],[spc156_numerals])).
cnf(c_0_203, plain, (m1_subset_1(np__13,k5_numbers)), inference(evalgc,[status(thm)],[c_0_189])).
cnf(c_0_204, plain, (m1_subset_1(np__12,k5_numbers)), inference(split_conjunct,[status(thm)],[spc12_numerals])).
cnf(c_0_205, plain, (esk1_1(X1)=np__17|esk1_1(X1)=np__13|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__163|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_190, c_0_191])])).
cnf(c_0_206, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__17,X1)|X2!=k2_xcmplx_0(np__153,X1)|~r1_nat_d(np__17,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_193]), c_0_194]), c_0_195]), c_0_196])])])).
cnf(c_0_207, plain, (k2_xcmplx_0(np__153,np__10)=np__163), inference(evalgc,[status(thm)],[c_0_197])).
cnf(c_0_208, plain, (v7_ordinal1(np__10)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_198]), c_0_125])])])).
cnf(c_0_209, plain, (~r1_xxreal_0(np__17,np__10)), inference(evalgc,[status(thm)],[c_0_199])).
cnf(c_0_210, plain, (~r1_xxreal_0(np__10,np__0)), inference(evalgc,[status(thm)],[c_0_200])).
cnf(c_0_211, plain, (k3_xcmplx_0(np__13,np__12)=np__156), inference(evalgc,[status(thm)],[c_0_201])).
cnf(c_0_212, plain, (m1_subset_1(np__156,k5_numbers)), inference(evalgc,[status(thm)],[c_0_202])).
cnf(c_0_213, plain, (v7_ordinal1(np__13)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_203]), c_0_125])])])).
cnf(c_0_214, plain, (m1_subset_1(np__12,k5_numbers)), inference(evalgc,[status(thm)],[c_0_204])).
fof(c_0_215, plain, ~r1_xxreal_0(np__13,np__7), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r13_r7])).
fof(c_0_216, plain, ~r1_xxreal_0(np__7,np__0), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r7_r0])).
cnf(c_0_217, plain, (r1_nat_d(np__17,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_205])])).
cnf(c_0_218, plain, (X1!=np__163|~r1_nat_d(np__17,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_207]), c_0_208])]), c_0_209]), c_0_210])])).
cnf(c_0_219, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__13,X1)|X2!=k2_xcmplx_0(np__156,X1)|~r1_nat_d(np__13,X2)|~v7_ordinal1(X1)|~v7_ordinal1(np__12)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_211]), c_0_212]), c_0_213])])])).
cnf(c_0_220, plain, (v7_ordinal1(np__12)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_214]), c_0_125])])])).
cnf(c_0_221, plain, (k2_xcmplx_0(np__156,np__7)=np__163), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r156_r7_r163])).
cnf(c_0_222, plain, (~r1_xxreal_0(np__13,np__7)), inference(split_conjunct,[status(thm)],[c_0_215])).
cnf(c_0_223, plain, (~r1_xxreal_0(np__7,np__0)), inference(split_conjunct,[status(thm)],[c_0_216])).
cnf(c_0_224, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(k24_binop_2(X4,X2),X1)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_119])])).
cnf(c_0_225, plain, (m1_subset_1(np__32,k5_numbers)), inference(split_conjunct,[status(thm)],[spc32_numerals])).
cnf(c_0_226, plain, (m1_subset_1(np__5,k5_numbers)), inference(split_conjunct,[status(thm)],[spc5_numerals])).
cnf(c_0_227, plain, (esk1_1(X1)=np__13|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_217, c_0_218])).
cnf(c_0_228, plain, (r1_xxreal_0(np__13,X1)|r1_xxreal_0(X1,np__0)|X2!=k2_xcmplx_0(np__156,X1)|~r1_nat_d(np__13,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_219, c_0_220])])).
cnf(c_0_229, plain, (k2_xcmplx_0(np__156,np__7)=np__163), inference(evalgc,[status(thm)],[c_0_221])).
cnf(c_0_230, plain, (~r1_xxreal_0(np__13,np__7)), inference(evalgc,[status(thm)],[c_0_222])).
cnf(c_0_231, plain, (~r1_xxreal_0(np__7,np__0)), inference(evalgc,[status(thm)],[c_0_223])).
cnf(c_0_232, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(k24_binop_2(X4,X2),X1)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(k24_binop_2(X4,X2),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_224, c_0_102])])).
cnf(c_0_233, plain, (k3_xcmplx_0(np__32,np__5)=np__160), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__r32_r5_r160])).
cnf(c_0_234, plain, (m1_subset_1(np__160,k5_numbers)), inference(split_conjunct,[status(thm)],[spc160_numerals])).
cnf(c_0_235, plain, (m1_subset_1(np__32,k5_numbers)), inference(evalgc,[status(thm)],[c_0_225])).
cnf(c_0_236, plain, (m1_subset_1(np__5,k5_numbers)), inference(evalgc,[status(thm)],[c_0_226])).
cnf(c_0_237, plain, (m1_subset_1(np__3,k5_numbers)), inference(split_conjunct,[status(thm)],[spc3_numerals])).
fof(c_0_238, plain, ~r1_xxreal_0(np__5,np__3), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r5_r3])).
fof(c_0_239, plain, ~r1_xxreal_0(np__3,np__0), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r3_r0])).
cnf(c_0_240, plain, (r1_nat_d(np__13,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_227])])).
cnf(c_0_241, plain, (X1!=np__163|~r1_nat_d(np__13,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_229]), c_0_165])]), c_0_230]), c_0_231])])).
cnf(c_0_242, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(k3_xcmplx_0(X4,X2),X1)|~m1_subset_1(k3_xcmplx_0(X4,X2),k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_111])])).
cnf(c_0_243, plain, (k3_xcmplx_0(np__32,np__5)=np__160), inference(evalgc,[status(thm)],[c_0_233])).
cnf(c_0_244, plain, (m1_subset_1(np__160,k5_numbers)), inference(evalgc,[status(thm)],[c_0_234])).
cnf(c_0_245, plain, (v7_ordinal1(np__32)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_235]), c_0_125])])])).
cnf(c_0_246, plain, (v7_ordinal1(np__5)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_236]), c_0_125])])])).
cnf(c_0_247, plain, (k2_xcmplx_0(np__160,np__3)=np__163), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r160_r3_r163])).
cnf(c_0_248, plain, (m1_subset_1(np__3,k5_numbers)), inference(evalgc,[status(thm)],[c_0_237])).
cnf(c_0_249, plain, (~r1_xxreal_0(np__5,np__3)), inference(split_conjunct,[status(thm)],[c_0_238])).
cnf(c_0_250, plain, (~r1_xxreal_0(np__3,np__0)), inference(split_conjunct,[status(thm)],[c_0_239])).
cnf(c_0_251, plain, (esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_240, c_0_241])).
cnf(c_0_252, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__5,X1)|X2!=k2_xcmplx_0(np__160,X1)|~r1_nat_d(np__5,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_243]), c_0_244]), c_0_245]), c_0_246])])])).
cnf(c_0_253, plain, (k2_xcmplx_0(np__160,np__3)=np__163), inference(evalgc,[status(thm)],[c_0_247])).
cnf(c_0_254, plain, (v7_ordinal1(np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_248]), c_0_125])])])).
cnf(c_0_255, plain, (~r1_xxreal_0(np__5,np__3)), inference(evalgc,[status(thm)],[c_0_249])).
cnf(c_0_256, plain, (~r1_xxreal_0(np__3,np__0)), inference(evalgc,[status(thm)],[c_0_250])).
cnf(c_0_257, plain, (m1_subset_1(np__161,k5_numbers)), inference(split_conjunct,[status(thm)],[spc161_numerals])).
fof(c_0_258, plain, ~r1_xxreal_0(np__7,np__2), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r7_r2])).
cnf(c_0_259, plain, (r1_nat_d(np__5,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__7|esk1_1(X1)=np__11|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_251])])).
cnf(c_0_260, plain, (X1!=np__163|~r1_nat_d(np__5,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_253]), c_0_254])]), c_0_255]), c_0_256])])).
cnf(c_0_261, plain, (m1_subset_1(np__161,k5_numbers)), inference(evalgc,[status(thm)],[c_0_257])).
cnf(c_0_262, plain, (k2_xcmplx_0(np__161,np__2)=np__163), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r161_r2_r163])).
cnf(c_0_263, plain, (~r1_xxreal_0(np__7,np__2)), inference(split_conjunct,[status(thm)],[c_0_258])).
cnf(c_0_264, plain, (k3_xcmplx_0(np__11,np__14)=np__154), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__r11_r14_r154])).
cnf(c_0_265, plain, (m1_subset_1(np__154,k5_numbers)), inference(split_conjunct,[status(thm)],[spc154_numerals])).
cnf(c_0_266, plain, (m1_subset_1(np__14,k5_numbers)), inference(split_conjunct,[status(thm)],[spc14_numerals])).
cnf(c_0_267, plain, (esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_259, c_0_260])).
cnf(c_0_268, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__7,X1)|X2!=k2_xcmplx_0(np__161,X1)|~r1_nat_d(np__7,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_164]), c_0_261]), c_0_166]), c_0_165])])])).
cnf(c_0_269, plain, (k2_xcmplx_0(np__161,np__2)=np__163), inference(evalgc,[status(thm)],[c_0_262])).
cnf(c_0_270, plain, (v7_ordinal1(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_177]), c_0_125])])])).
cnf(c_0_271, plain, (~r1_xxreal_0(np__7,np__2)), inference(evalgc,[status(thm)],[c_0_263])).
cnf(c_0_272, plain, (k3_xcmplx_0(np__11,np__14)=np__154), inference(evalgc,[status(thm)],[c_0_264])).
cnf(c_0_273, plain, (m1_subset_1(np__154,k5_numbers)), inference(evalgc,[status(thm)],[c_0_265])).
cnf(c_0_274, plain, (v7_ordinal1(np__11)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_158]), c_0_125])])])).
cnf(c_0_275, plain, (m1_subset_1(np__14,k5_numbers)), inference(evalgc,[status(thm)],[c_0_266])).
fof(c_0_276, plain, ~r1_xxreal_0(np__11,np__9), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r11_r9])).
fof(c_0_277, plain, ~r1_xxreal_0(np__9,np__0), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r9_r0])).
cnf(c_0_278, plain, (r1_nat_d(np__7,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_267])])).
cnf(c_0_279, plain, (X1!=np__163|~r1_nat_d(np__7,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_269]), c_0_270])]), c_0_271]), c_0_179])])).
cnf(c_0_280, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__11,X1)|X2!=k2_xcmplx_0(np__154,X1)|~r1_nat_d(np__11,X2)|~v7_ordinal1(X1)|~v7_ordinal1(np__14)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_272]), c_0_273]), c_0_274])])])).
cnf(c_0_281, plain, (v7_ordinal1(np__14)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_275]), c_0_125])])])).
cnf(c_0_282, plain, (k2_xcmplx_0(np__154,np__9)=np__163), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r154_r9_r163])).
cnf(c_0_283, plain, (~r1_xxreal_0(np__11,np__9)), inference(split_conjunct,[status(thm)],[c_0_276])).
cnf(c_0_284, plain, (~r1_xxreal_0(np__9,np__0)), inference(split_conjunct,[status(thm)],[c_0_277])).
cnf(c_0_285, plain, (m1_subset_1(np__54,k5_numbers)), inference(split_conjunct,[status(thm)],[spc54_numerals])).
cnf(c_0_286, plain, (esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__11|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_278, c_0_279])).
cnf(c_0_287, plain, (r1_xxreal_0(np__11,X1)|r1_xxreal_0(X1,np__0)|X2!=k2_xcmplx_0(np__154,X1)|~r1_nat_d(np__11,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_280, c_0_281])])).
cnf(c_0_288, plain, (k2_xcmplx_0(np__154,np__9)=np__163), inference(evalgc,[status(thm)],[c_0_282])).
cnf(c_0_289, plain, (~r1_xxreal_0(np__11,np__9)), inference(evalgc,[status(thm)],[c_0_283])).
cnf(c_0_290, plain, (~r1_xxreal_0(np__9,np__0)), inference(evalgc,[status(thm)],[c_0_284])).
cnf(c_0_291, plain, (k3_xcmplx_0(np__54,np__3)=np__162), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__r54_r3_r162])).
cnf(c_0_292, plain, (m1_subset_1(np__54,k5_numbers)), inference(evalgc,[status(thm)],[c_0_285])).
fof(c_0_293, plain, ~r1_xxreal_0(np__3,np__1), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r3_r1])).
fof(c_0_294, plain, ~r1_xxreal_0(np__1,np__0), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r1_r0])).
cnf(c_0_295, plain, (r1_nat_d(np__11,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_286])])).
cnf(c_0_296, plain, (X1!=np__163|~r1_nat_d(np__11,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_287, c_0_288]), c_0_195])]), c_0_289]), c_0_290])])).
cnf(c_0_297, plain, (k3_xcmplx_0(np__54,np__3)=np__162), inference(evalgc,[status(thm)],[c_0_291])).
cnf(c_0_298, plain, (v7_ordinal1(np__54)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_292]), c_0_125])])])).
cnf(c_0_299, plain, (k2_xcmplx_0(np__1,np__162)=np__163), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r162_r163])).
cnf(c_0_300, plain, (m1_subset_1(np__1,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_301, plain, (~r1_xxreal_0(np__3,np__1)), inference(split_conjunct,[status(thm)],[c_0_293])).
cnf(c_0_302, plain, (~r1_xxreal_0(np__1,np__0)), inference(split_conjunct,[status(thm)],[c_0_294])).
cnf(c_0_303, plain, (m1_subset_1(np__81,k5_numbers)), inference(split_conjunct,[status(thm)],[spc81_numerals])).
cnf(c_0_304, plain, (esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_295, c_0_296])).
cnf(c_0_305, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__3,X1)|X2!=k2_xcmplx_0(X1,np__162)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__3,X2)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_297]), c_0_298]), c_0_254])])])).
cnf(c_0_306, plain, (k2_xcmplx_0(np__1,np__162)=np__163), inference(evalgc,[status(thm)],[c_0_299])).
cnf(c_0_307, plain, (m1_subset_1(np__1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_300])).
cnf(c_0_308, plain, (~r1_xxreal_0(np__3,np__1)), inference(evalgc,[status(thm)],[c_0_301])).
cnf(c_0_309, plain, (~r1_xxreal_0(np__1,np__0)), inference(evalgc,[status(thm)],[c_0_302])).
cnf(c_0_310, plain, (k3_xcmplx_0(np__81,np__2)=np__162), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__r81_r2_r162])).
cnf(c_0_311, plain, (m1_subset_1(np__81,k5_numbers)), inference(evalgc,[status(thm)],[c_0_303])).
fof(c_0_312, plain, ~r1_xxreal_0(np__2,np__1), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r2_r1])).
fof(c_0_313, negated_conjecture, ~(v1_int_2(np__163)), inference(assume_negation,[status(cth)],[t35_nat_4])).
cnf(c_0_314, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_304])])).
cnf(c_0_315, plain, (X1!=np__163|~r1_nat_d(np__3,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_306]), c_0_307])]), c_0_308]), c_0_309])])).
cnf(c_0_316, plain, (k3_xcmplx_0(np__81,np__2)=np__162), inference(evalgc,[status(thm)],[c_0_310])).
cnf(c_0_317, plain, (v7_ordinal1(np__81)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_311]), c_0_125])])])).
cnf(c_0_318, plain, (~r1_xxreal_0(np__2,np__1)), inference(split_conjunct,[status(thm)],[c_0_312])).
fof(c_0_319, negated_conjecture, ~v1_int_2(np__163), inference(fof_simplification,[status(thm)],[c_0_313])).
cnf(c_0_320, plain, (esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_314, c_0_315])).
cnf(c_0_321, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__2,X1)|X2!=k2_xcmplx_0(X1,np__162)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__2,X2)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_316]), c_0_317]), c_0_270])])])).
cnf(c_0_322, plain, (~r1_xxreal_0(np__2,np__1)), inference(evalgc,[status(thm)],[c_0_318])).
cnf(c_0_323, plain, (m1_subset_1(np__163,k5_numbers)), inference(split_conjunct,[status(thm)],[spc163_numerals])).
fof(c_0_324, plain, ~r1_xxreal_0(np__841,np__163), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r841_r163])).
fof(c_0_325, plain, ~r1_xxreal_0(np__163,np__1), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r163_r1])).
cnf(c_0_326, negated_conjecture, (~v1_int_2(np__163)), inference(split_conjunct,[status(thm)],[c_0_319])).
cnf(c_0_327, plain, (r1_nat_d(np__2,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_320])])).
cnf(c_0_328, plain, (X1!=np__163|~r1_nat_d(np__2,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_321, c_0_306]), c_0_307])]), c_0_322]), c_0_309])])).
cnf(c_0_329, plain, (m2_subset_1(np__163,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[spc163_numerals])).
cnf(c_0_330, plain, (m1_subset_1(np__163,k5_numbers)), inference(evalgc,[status(thm)],[c_0_323])).
cnf(c_0_331, plain, (~r1_xxreal_0(np__841,np__163)), inference(split_conjunct,[status(thm)],[c_0_324])).
cnf(c_0_332, plain, (~r1_xxreal_0(np__163,np__1)), inference(split_conjunct,[status(thm)],[c_0_325])).
cnf(c_0_333, negated_conjecture, (~v1_int_2(np__163)), inference(evalgc,[status(thm)],[c_0_326])).
cnf(c_0_334, plain, (r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_327, c_0_328])).
cnf(c_0_335, plain, (m2_subset_1(np__163,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_329])).
cnf(c_0_336, plain, (v7_ordinal1(np__163)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_330]), c_0_125])])])).
cnf(c_0_337, plain, (~r1_xxreal_0(np__841,np__163)), inference(evalgc,[status(thm)],[c_0_331])).
cnf(c_0_338, plain, (~r1_xxreal_0(np__163,np__1)), inference(evalgc,[status(thm)],[c_0_332])).
cnf(c_0_339, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_333, c_0_334]), c_0_335]), c_0_336])]), c_0_337]), c_0_338]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 340
# Proof object clause steps            : 235
# Proof object formula steps           : 105
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 79
# Proof object initial formulas used   : 74
# Proof object generating inferences   : 60
# Proof object simplifying inferences  : 128
# Training examples: 139 positive, 1950 negative
# Training: Positive examples begin
cnf(c_0_79, plain, (k1_xboole_0=k6_numbers)).# trainpos
cnf(c_0_94, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_83, plain, (v1_xboole_0(np__0))).# trainpos
cnf(c_0_333, negated_conjecture, (~v1_int_2(np__163))).# trainpos
cnf(c_0_335, plain, (m2_subset_1(np__163,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_115, plain, (v6_membered(k4_ordinal1))).# trainpos
cnf(c_0_338, plain, (~r1_xxreal_0(np__163,np__1))).# trainpos
cnf(c_0_78, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_164, plain, (k3_xcmplx_0(np__7,np__23)=np__161)).# trainpos
cnf(c_0_272, plain, (k3_xcmplx_0(np__11,np__14)=np__154)).# trainpos
cnf(c_0_211, plain, (k3_xcmplx_0(np__13,np__12)=np__156)).# trainpos
cnf(c_0_193, plain, (k3_xcmplx_0(np__17,np__9)=np__153)).# trainpos
cnf(c_0_138, plain, (k3_xcmplx_0(np__19,np__8)=np__152)).# trainpos
cnf(c_0_243, plain, (k3_xcmplx_0(np__32,np__5)=np__160)).# trainpos
cnf(c_0_297, plain, (k3_xcmplx_0(np__54,np__3)=np__162)).# trainpos
cnf(c_0_316, plain, (k3_xcmplx_0(np__81,np__2)=np__162)).# trainpos
cnf(c_0_306, plain, (k2_xcmplx_0(np__1,np__162)=np__163)).# trainpos
cnf(c_0_176, plain, (k2_xcmplx_0(np__2,np__161)=np__163)).# trainpos
cnf(c_0_309, plain, (~r1_xxreal_0(np__1,np__0))).# trainpos
cnf(c_0_337, plain, (~r1_xxreal_0(np__841,np__163))).# trainpos
cnf(c_0_157, plain, (k2_xcmplx_0(np__11,np__152)=np__163)).# trainpos
cnf(c_0_207, plain, (k2_xcmplx_0(np__153,np__10)=np__163)).# trainpos
cnf(c_0_253, plain, (k2_xcmplx_0(np__160,np__3)=np__163)).# trainpos
cnf(c_0_288, plain, (k2_xcmplx_0(np__154,np__9)=np__163)).# trainpos
cnf(c_0_229, plain, (k2_xcmplx_0(np__156,np__7)=np__163)).# trainpos
cnf(c_0_322, plain, (~r1_xxreal_0(np__2,np__1))).# trainpos
cnf(c_0_269, plain, (k2_xcmplx_0(np__161,np__2)=np__163)).# trainpos
cnf(c_0_179, plain, (~r1_xxreal_0(np__2,np__0))).# trainpos
cnf(c_0_308, plain, (~r1_xxreal_0(np__3,np__1))).# trainpos
cnf(c_0_256, plain, (~r1_xxreal_0(np__3,np__0))).# trainpos
cnf(c_0_255, plain, (~r1_xxreal_0(np__5,np__3))).# trainpos
cnf(c_0_271, plain, (~r1_xxreal_0(np__7,np__2))).# trainpos
cnf(c_0_231, plain, (~r1_xxreal_0(np__7,np__0))).# trainpos
cnf(c_0_160, plain, (~r1_xxreal_0(np__11,np__0))).# trainpos
cnf(c_0_289, plain, (~r1_xxreal_0(np__11,np__9))).# trainpos
cnf(c_0_230, plain, (~r1_xxreal_0(np__13,np__7))).# trainpos
cnf(c_0_330, plain, (m1_subset_1(np__163,k5_numbers))).# trainpos
cnf(c_0_307, plain, (m1_subset_1(np__1,k5_numbers))).# trainpos
cnf(c_0_177, plain, (m1_subset_1(np__2,k5_numbers))).# trainpos
cnf(c_0_248, plain, (m1_subset_1(np__3,k5_numbers))).# trainpos
cnf(c_0_236, plain, (m1_subset_1(np__5,k5_numbers))).# trainpos
cnf(c_0_152, plain, (m1_subset_1(np__7,k5_numbers))).# trainpos
cnf(c_0_158, plain, (m1_subset_1(np__11,k5_numbers))).# trainpos
cnf(c_0_203, plain, (m1_subset_1(np__13,k5_numbers))).# trainpos
cnf(c_0_185, plain, (m1_subset_1(np__17,k5_numbers))).# trainpos
cnf(c_0_126, plain, (m1_subset_1(np__19,k5_numbers))).# trainpos
cnf(c_0_153, plain, (m1_subset_1(np__23,k5_numbers))).# trainpos
cnf(c_0_194, plain, (m1_subset_1(np__153,k5_numbers))).# trainpos
cnf(c_0_209, plain, (~r1_xxreal_0(np__17,np__10))).# trainpos
cnf(c_0_235, plain, (m1_subset_1(np__32,k5_numbers))).# trainpos
cnf(c_0_244, plain, (m1_subset_1(np__160,k5_numbers))).# trainpos
cnf(c_0_292, plain, (m1_subset_1(np__54,k5_numbers))).# trainpos
cnf(c_0_311, plain, (m1_subset_1(np__81,k5_numbers))).# trainpos
cnf(c_0_273, plain, (m1_subset_1(np__154,k5_numbers))).# trainpos
cnf(c_0_212, plain, (m1_subset_1(np__156,k5_numbers))).# trainpos
cnf(c_0_159, plain, (~r1_xxreal_0(np__19,np__11))).# trainpos
cnf(c_0_123, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_184, plain, (m1_subset_1(np__9,k5_numbers))).# trainpos
cnf(c_0_124, plain, (m1_subset_1(np__8,k5_numbers))).# trainpos
cnf(c_0_198, plain, (m1_subset_1(np__10,k5_numbers))).# trainpos
cnf(c_0_261, plain, (m1_subset_1(np__161,k5_numbers))).# trainpos
cnf(c_0_214, plain, (m1_subset_1(np__12,k5_numbers))).# trainpos
cnf(c_0_275, plain, (m1_subset_1(np__14,k5_numbers))).# trainpos
cnf(c_0_178, plain, (~r1_xxreal_0(np__23,np__2))).# trainpos
cnf(c_0_290, plain, (~r1_xxreal_0(np__9,np__0))).# trainpos
cnf(c_0_210, plain, (~r1_xxreal_0(np__10,np__0))).# trainpos
cnf(c_0_93, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).# trainpos
cnf(c_0_147, plain, (v1_int_2(X1)|v1_int_2(esk1_1(X1))|r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_149, plain, (v1_int_2(X1)|r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1)|~r1_xxreal_0(esk1_1(X1),np__1))).# trainpos
cnf(c_0_121, plain, (v7_ordinal1(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_119, plain, (k24_binop_2(X1,X2)=k24_binop_2(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_111, plain, (k24_binop_2(X1,X2)=k3_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_92, plain, (k2_nat_1(X1,X2)=k2_nat_1(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_102, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_148, plain, (v1_int_2(X1)|r1_xxreal_0(X1,np__1)|m2_subset_1(esk1_1(X1),k1_numbers,k5_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_86, plain, (r1_xxreal_0(X2,X4)|r1_xxreal_0(X4,k6_numbers)|X1!=k2_nat_1(k24_binop_2(X2,X3),X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_nat_d(X2,X1))).# trainpos
cnf(c_0_161, plain, (v1_int_2(X1)|r1_xxreal_0(X1,np__1)|r1_nat_d(esk1_1(X1),X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_130, plain, (v1_int_2(X1)|r1_xxreal_0(X1,np__1)|r1_xxreal_0(k4_nat_1(esk1_1(X1),esk1_1(X1)),X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_129, plain, (X2=np__2|X2=np__3|X2=np__5|X2=np__7|X2=np__11|X2=np__13|X2=np__17|X2=np__19|X2=np__23|r1_xxreal_0(X2,np__1)|r1_xxreal_0(np__841,X1)|~v1_int_2(X2)|~r1_xxreal_0(k4_nat_1(X2,X2),X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_336, plain, (v7_ordinal1(np__163))).# trainpos
cnf(c_0_87, plain, (k6_numbers=np__0)).# trainpos
cnf(c_0_317, plain, (v7_ordinal1(np__81))).# trainpos
cnf(c_0_298, plain, (v7_ordinal1(np__54))).# trainpos
cnf(c_0_245, plain, (v7_ordinal1(np__32))).# trainpos
cnf(c_0_146, plain, (esk1_1(X1)=np__23|esk1_1(X1)=np__19|esk1_1(X1)=np__17|esk1_1(X1)=np__13|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|r1_xxreal_0(esk1_1(X1),np__1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|~m2_subset_1(esk1_1(X1),k1_numbers,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_int_2(esk1_1(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_91, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(k24_binop_2(X2,X4),X1)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_166, plain, (v7_ordinal1(np__23))).# trainpos
cnf(c_0_140, plain, (v7_ordinal1(np__19))).# trainpos
cnf(c_0_196, plain, (v7_ordinal1(np__17))).# trainpos
cnf(c_0_213, plain, (v7_ordinal1(np__13))).# trainpos
cnf(c_0_172, plain, (r1_nat_d(np__19,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|esk1_1(X1)=np__17|esk1_1(X1)=np__23|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_96, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(X1,k24_binop_2(X2,X4))|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(k24_binop_2(X2,X4)))).# trainpos
cnf(c_0_224, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(k24_binop_2(X4,X2),X1)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_181, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(k24_binop_2(X2,X4),X1)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(k24_binop_2(X2,X4),k5_numbers))).# trainpos
cnf(c_0_274, plain, (v7_ordinal1(np__11))).# trainpos
cnf(c_0_165, plain, (v7_ordinal1(np__7))).# trainpos
cnf(c_0_246, plain, (v7_ordinal1(np__5))).# trainpos
cnf(c_0_254, plain, (v7_ordinal1(np__3))).# trainpos
cnf(c_0_270, plain, (v7_ordinal1(np__2))).# trainpos
cnf(c_0_232, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(k24_binop_2(X4,X2),X1)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(k24_binop_2(X4,X2),k5_numbers))).# trainpos
cnf(c_0_195, plain, (v7_ordinal1(np__9))).# trainpos
cnf(c_0_139, plain, (v7_ordinal1(np__8))).# trainpos
cnf(c_0_192, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(k3_xcmplx_0(X2,X4),X1)|~m1_subset_1(k3_xcmplx_0(X2,X4),k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_208, plain, (v7_ordinal1(np__10))).# trainpos
cnf(c_0_220, plain, (v7_ordinal1(np__12))).# trainpos
cnf(c_0_281, plain, (v7_ordinal1(np__14))).# trainpos
cnf(c_0_280, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__11,X1)|X2!=k2_xcmplx_0(np__154,X1)|~r1_nat_d(np__11,X2)|~v7_ordinal1(X1)|~v7_ordinal1(np__14)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_242, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(k3_xcmplx_0(X4,X2),X1)|~m1_subset_1(k3_xcmplx_0(X4,X2),k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_296, plain, (X1!=np__163|~r1_nat_d(np__11,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_110, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X1,k24_binop_2(X2,X4))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(k24_binop_2(X2,X4))|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_268, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__7,X1)|X2!=k2_xcmplx_0(np__161,X1)|~r1_nat_d(np__7,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_252, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__5,X1)|X2!=k2_xcmplx_0(np__160,X1)|~r1_nat_d(np__5,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_206, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__17,X1)|X2!=k2_xcmplx_0(np__153,X1)|~r1_nat_d(np__17,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_219, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__13,X1)|X2!=k2_xcmplx_0(np__156,X1)|~r1_nat_d(np__13,X2)|~v7_ordinal1(X1)|~v7_ordinal1(np__12)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_218, plain, (X1!=np__163|~r1_nat_d(np__17,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_241, plain, (X1!=np__163|~r1_nat_d(np__13,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_260, plain, (X1!=np__163|~r1_nat_d(np__5,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_120, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X1,k3_xcmplx_0(X2,X4))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(k3_xcmplx_0(X2,X4))|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_279, plain, (X1!=np__163|~r1_nat_d(np__7,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_134, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X1,k24_binop_2(X4,X2))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(k24_binop_2(X4,X2))|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_156, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__19,X1)|X2!=k2_xcmplx_0(X1,np__152)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__19,X2)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_171, plain, (X1!=np__163|~r1_nat_d(np__19,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_150, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X1,k3_xcmplx_0(X4,X2))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(k3_xcmplx_0(X4,X2))|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_180, plain, (esk1_1(X1)=np__23|esk1_1(X1)=np__17|esk1_1(X1)=np__13|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__163|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_191, plain, (r1_nat_d(np__23,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|esk1_1(X1)=np__17|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_305, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__3,X1)|X2!=k2_xcmplx_0(X1,np__162)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__3,X2)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_315, plain, (X1!=np__163|~r1_nat_d(np__3,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_321, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__2,X1)|X2!=k2_xcmplx_0(X1,np__162)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__2,X2)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_328, plain, (X1!=np__163|~r1_nat_d(np__2,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_175, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__23,X1)|X2!=k2_xcmplx_0(X1,np__161)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__23,X2)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_190, plain, (X1!=np__163|~r1_nat_d(np__23,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_205, plain, (esk1_1(X1)=np__17|esk1_1(X1)=np__13|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__163|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_217, plain, (r1_nat_d(np__17,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_240, plain, (r1_nat_d(np__13,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_259, plain, (r1_nat_d(np__5,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__7|esk1_1(X1)=np__11|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_278, plain, (r1_nat_d(np__7,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_295, plain, (r1_nat_d(np__11,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_314, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_327, plain, (r1_nat_d(np__2,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_621, plain, (v1_int_2(esk7_0))).#trainneg
cnf(i_0_627, plain, (v7_ordinal1(esk7_0))).#trainneg
cnf(i_0_689, plain, (v7_ordinal1(esk9_0))).#trainneg
cnf(i_0_724, plain, (v7_ordinal1(esk11_0))).#trainneg
cnf(i_0_212, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_459, plain, (v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_518, plain, (v1_xboole_0(esk4_0))).#trainneg
cnf(i_0_532, plain, (v1_xboole_0(esk5_0))).#trainneg
cnf(i_0_67, plain, (v2_xxreal_0(np__163))).#trainneg
cnf(i_0_332, plain, (~v1_xboole_0(np__163))).#trainneg
cnf(i_0_118, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_476, plain, (~v1_int_2(X1)|~v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_363, plain, (v2_xxreal_0(np__841))).#trainneg
cnf(i_0_381, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_165, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_353, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_742, plain, (v6_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_17, plain, (v2_xxreal_0(np__2))).#trainneg
cnf(i_0_59, plain, (v2_xxreal_0(np__3))).#trainneg
cnf(i_0_55, plain, (v2_xxreal_0(np__5))).#trainneg
cnf(i_0_21, plain, (v2_xxreal_0(np__7))).#trainneg
cnf(i_0_40, plain, (v2_xxreal_0(np__11))).#trainneg
cnf(i_0_114, plain, (v2_xxreal_0(np__13))).#trainneg
cnf(i_0_51, plain, (v2_xxreal_0(np__17))).#trainneg
cnf(i_0_44, plain, (v2_xxreal_0(np__19))).#trainneg
cnf(i_0_34, plain, (v2_xxreal_0(np__23))).#trainneg
cnf(i_0_26, plain, (v2_xxreal_0(np__153))).#trainneg
cnf(i_0_30, plain, (v2_xxreal_0(np__152))).#trainneg
cnf(i_0_98, plain, (v2_xxreal_0(np__32))).#trainneg
cnf(i_0_89, plain, (v2_xxreal_0(np__160))).#trainneg
cnf(i_0_93, plain, (v2_xxreal_0(np__54))).#trainneg
cnf(i_0_467, plain, (~v1_xboole_0(np__841))).#trainneg
cnf(i_0_372, plain, (~v1_xboole_0(np__2))).#trainneg
cnf(i_0_71, plain, (v2_xxreal_0(np__162))).#trainneg
cnf(i_0_8, plain, (~v1_int_2(X1)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_63, plain, (v2_xxreal_0(np__81))).#trainneg
cnf(i_0_117, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_435, plain, (~v1_xboole_0(np__3))).#trainneg
cnf(i_0_489, plain, (~v1_xboole_0(np__5))).#trainneg
cnf(i_0_185, plain, (v1_xcmplx_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_589, plain, (v3_ordinal1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_75, plain, (v2_xxreal_0(np__154))).#trainneg
cnf(i_0_79, plain, (v2_xxreal_0(np__156))).#trainneg
cnf(i_0_85, plain, (v2_xxreal_0(np__9))).#trainneg
cnf(i_0_154, plain, (v2_xxreal_0(np__8))).#trainneg
cnf(i_0_106, plain, (v2_xxreal_0(np__10))).#trainneg
cnf(i_0_110, plain, (v2_xxreal_0(np__161))).#trainneg
cnf(i_0_127, plain, (v2_xxreal_0(np__12))).#trainneg
cnf(i_0_131, plain, (v2_xxreal_0(np__14))).#trainneg
cnf(i_0_458, plain, (v1_xcmplx_0(esk3_0))).#trainneg
cnf(i_0_545, plain, (v1_xcmplx_0(esk6_0))).#trainneg
cnf(i_0_626, plain, (v1_xcmplx_0(esk7_0))).#trainneg
cnf(i_0_148, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_628, plain, (v3_ordinal1(esk7_0))).#trainneg
cnf(i_0_424, plain, (~v1_xboole_0(np__7))).#trainneg
cnf(i_0_418, plain, (~v1_xboole_0(np__11))).#trainneg
cnf(i_0_172, plain, (v3_membered(k1_numbers))).#trainneg
cnf(i_0_539, plain, (~v1_xboole_0(np__13))).#trainneg
cnf(i_0_456, plain, (v1_xreal_0(esk3_0))).#trainneg
cnf(i_0_362, plain, (m2_subset_1(np__841,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_355, plain, (~v1_xboole_0(np__17))).#trainneg
cnf(i_0_358, plain, (~v1_xboole_0(np__19))).#trainneg
cnf(i_0_615, plain, (v3_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_544, plain, (v1_xreal_0(esk6_0))).#trainneg
cnf(i_0_625, plain, (v1_xreal_0(esk7_0))).#trainneg
cnf(i_0_543, plain, (v1_int_1(esk6_0))).#trainneg
cnf(i_0_622, plain, (v1_int_1(esk7_0))).#trainneg
cnf(i_0_620, plain, (v1_finset_1(esk7_0))).#trainneg
cnf(i_0_750, plain, (v1_finset_1(esk13_0))).#trainneg
cnf(i_0_758, plain, (v1_finset_1(esk14_0))).#trainneg
cnf(i_0_457, plain, (v1_xxreal_0(esk3_0))).#trainneg
cnf(i_0_531, plain, (v1_xxreal_0(esk5_0))).#trainneg
cnf(i_0_546, plain, (v1_xxreal_0(esk6_0))).#trainneg
cnf(i_0_624, plain, (v1_xxreal_0(esk7_0))).#trainneg
cnf(i_0_755, plain, (v1_relat_1(esk13_0))).#trainneg
cnf(i_0_763, plain, (v1_relat_1(esk14_0))).#trainneg
cnf(i_0_534, plain, (v5_membered(k4_numbers))).#trainneg
cnf(i_0_331, plain, (~v1_xboole_0(np__23))).#trainneg
cnf(i_0_149, plain, (~v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_524, plain, (v1_card_1(k4_ordinal1))).#trainneg
cnf(i_0_383, plain, (~v1_xboole_0(np__153))).#trainneg
cnf(i_0_619, plain, (v1_card_1(esk7_0))).#trainneg
cnf(i_0_16, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_336, plain, (~v1_xboole_0(np__152))).#trainneg
cnf(i_0_330, plain, (~v1_xboole_0(np__32))).#trainneg
cnf(i_0_213, plain, (v1_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_227, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_753, plain, (v1_funct_1(esk13_0))).#trainneg
cnf(i_0_760, plain, (v1_funct_1(esk14_0))).#trainneg
cnf(i_0_576, plain, (v4_membered(k3_numbers))).#trainneg
cnf(i_0_751, plain, (v3_valued_0(esk13_0))).#trainneg
cnf(i_0_630, plain, (v1_ordinal1(esk7_0))).#trainneg
cnf(i_0_629, plain, (v2_ordinal1(esk7_0))).#trainneg
cnf(i_0_641, plain, (v2_card_1(k4_ordinal1))).#trainneg
cnf(i_0_749, plain, (v1_finseq_1(esk13_0))).#trainneg
cnf(i_0_757, plain, (v1_finseq_1(esk14_0))).#trainneg
cnf(i_0_752, plain, (v1_valued_0(esk13_0))).#trainneg
cnf(i_0_762, plain, (v2_relat_1(esk14_0))).#trainneg
cnf(i_0_748, plain, (v2_finseq_1(esk13_0))).#trainneg
cnf(i_0_756, plain, (v2_finseq_1(esk14_0))).#trainneg
cnf(i_0_321, plain, (~v1_xboole_0(np__160))).#trainneg
cnf(i_0_329, plain, (~v1_xboole_0(np__54))).#trainneg
cnf(i_0_351, plain, (~v1_xboole_0(np__162))).#trainneg
cnf(i_0_307, plain, (~v1_xboole_0(np__81))).#trainneg
cnf(i_0_335, plain, (~v1_xboole_0(np__154))).#trainneg
cnf(i_0_334, plain, (~v1_xboole_0(np__156))).#trainneg
cnf(i_0_58, plain, (m2_subset_1(np__3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_516, plain, (~v1_xboole_0(np__9))).#trainneg
cnf(i_0_537, plain, (~v1_xboole_0(np__8))).#trainneg
cnf(i_0_226, plain, (v1_int_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_573, plain, (v1_finset_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_487, plain, (~v1_xboole_0(np__10))).#trainneg
cnf(i_0_333, plain, (~v1_xboole_0(np__161))).#trainneg
cnf(i_0_530, plain, (~v1_xboole_0(np__12))).#trainneg
cnf(i_0_416, plain, (~v1_xboole_0(np__14))).#trainneg
cnf(i_0_719, plain, (~v1_xboole_0(k4_numbers))).#trainneg
cnf(i_0_690, plain, (~v1_xboole_0(esk9_0))).#trainneg
cnf(i_0_759, plain, (~v1_xboole_0(esk14_0))).#trainneg
cnf(i_0_384, plain, (~v1_finset_1(k4_ordinal1))).#trainneg
cnf(i_0_623, plain, (~v3_xxreal_0(esk7_0))).#trainneg
cnf(i_0_685, plain, (~v3_xxreal_2(k1_numbers))).#trainneg
cnf(i_0_684, plain, (~v4_xxreal_2(k1_numbers))).#trainneg
cnf(i_0_639, plain, (v1_relat_1(esk8_1(X1)))).#trainneg
cnf(i_0_703, plain, (v1_relat_1(esk10_1(X1)))).#trainneg
cnf(i_0_637, plain, (v1_funct_1(esk8_1(X1)))).#trainneg
cnf(i_0_239, plain, (~r1_xxreal_0(np__163,np__2))).#trainneg
cnf(i_0_593, plain, (r1_xxreal_0(np__163,np__163))).#trainneg
cnf(i_0_267, plain, (~r1_xxreal_0(np__163,np__3))).#trainneg
cnf(i_0_54, plain, (m2_subset_1(np__5,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_438, plain, (~r1_xxreal_0(np__163,np__5))).#trainneg
cnf(i_0_379, plain, (~r1_xxreal_0(np__163,np__7))).#trainneg
cnf(i_0_408, plain, (v1_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_311, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_700, plain, (v1_funct_1(esk10_1(X1)))).#trainneg
cnf(i_0_636, plain, (v5_ordinal1(esk8_1(X1)))).#trainneg
cnf(i_0_702, plain, (v2_relat_1(esk10_1(X1)))).#trainneg
cnf(i_0_447, plain, (k3_xcmplx_0(np__1,np__1)=np__1)).#trainneg
cnf(i_0_553, plain, (k3_xcmplx_0(np__1,np__2)=np__2)).#trainneg
cnf(i_0_500, plain, (k3_xcmplx_0(np__1,np__3)=np__3)).#trainneg
cnf(i_0_540, plain, (k3_xcmplx_0(np__1,np__5)=np__5)).#trainneg
cnf(i_0_513, plain, (k3_xcmplx_0(np__1,np__7)=np__7)).#trainneg
cnf(i_0_586, plain, (k3_xcmplx_0(np__1,np__11)=np__11)).#trainneg
cnf(i_0_511, plain, (k3_xcmplx_0(np__1,np__13)=np__13)).#trainneg
cnf(i_0_490, plain, (k3_xcmplx_0(np__1,np__17)=np__17)).#trainneg
cnf(i_0_430, plain, (k3_xcmplx_0(np__1,np__19)=np__19)).#trainneg
cnf(i_0_406, plain, (k3_xcmplx_0(np__1,np__23)=np__23)).#trainneg
cnf(i_0_508, plain, (k3_xcmplx_0(np__1,np__0)=np__0)).#trainneg
cnf(i_0_341, plain, (~r1_xxreal_0(np__163,np__11))).#trainneg
cnf(i_0_478, plain, (~r1_xxreal_0(np__163,np__13))).#trainneg
cnf(i_0_552, plain, (r1_xxreal_0(np__1,np__1))).#trainneg
cnf(i_0_305, plain, (~r1_xxreal_0(np__163,np__17))).#trainneg
cnf(i_0_472, plain, (k3_xcmplx_0(np__1,np__32)=np__32)).#trainneg
cnf(i_0_20, plain, (m2_subset_1(np__7,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_293, plain, (~r1_xxreal_0(np__163,np__19))).#trainneg
cnf(i_0_276, plain, (~r1_xxreal_0(np__163,np__23))).#trainneg
cnf(i_0_387, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_443, plain, (v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_428, plain, (k3_xcmplx_0(np__1,np__9)=np__9)).#trainneg
cnf(i_0_562, plain, (k3_xcmplx_0(np__1,np__8)=np__8)).#trainneg
cnf(i_0_568, plain, (k3_xcmplx_0(np__1,np__14)=np__14)).#trainneg
cnf(i_0_242, plain, (k3_xcmplx_0(np__2,np__1)=np__2)).#trainneg
cnf(i_0_133, plain, (k3_xcmplx_0(np__2,np__5)=np__10)).#trainneg
cnf(i_0_230, plain, (k3_xcmplx_0(np__2,np__7)=np__14)).#trainneg
cnf(i_0_385, plain, (k3_xcmplx_0(np__2,np__0)=np__0)).#trainneg
cnf(i_0_100, plain, (k3_xcmplx_0(np__2,np__81)=np__162)).#trainneg
cnf(i_0_491, plain, (k3_xcmplx_0(np__3,np__1)=np__3)).#trainneg
cnf(i_0_161, plain, (k3_xcmplx_0(np__3,np__3)=np__9)).#trainneg
cnf(i_0_122, plain, (k3_xcmplx_0(np__3,np__54)=np__162)).#trainneg
cnf(i_0_720, plain, (k3_xcmplx_0(np__5,np__1)=np__5)).#trainneg
cnf(i_0_143, plain, (k3_xcmplx_0(np__5,np__2)=np__10)).#trainneg
cnf(i_0_35, plain, (k3_xcmplx_0(np__5,np__32)=np__160)).#trainneg
cnf(i_0_255, plain, (~r1_xxreal_0(np__163,np__0))).#trainneg
cnf(i_0_327, plain, (~r1_xxreal_0(np__163,np__153))).#trainneg
cnf(i_0_549, plain, (r1_xxreal_0(np__1,np__2))).#trainneg
cnf(i_0_282, plain, (~r1_xxreal_0(np__163,np__152))).#trainneg
cnf(i_0_229, plain, (k3_xcmplx_0(np__7,np__2)=np__14)).#trainneg
cnf(i_0_39, plain, (m2_subset_1(np__11,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_231, plain, (~r1_xxreal_0(np__163,np__32))).#trainneg
cnf(i_0_279, plain, (~r1_xxreal_0(np__163,np__160))).#trainneg
cnf(i_0_415, plain, (v5_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_419, plain, (v2_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_587, plain, (k3_xcmplx_0(np__11,np__1)=np__11)).#trainneg
cnf(i_0_512, plain, (k3_xcmplx_0(np__13,np__1)=np__13)).#trainneg
cnf(i_0_376, plain, (k3_xcmplx_0(np__19,np__1)=np__19)).#trainneg
cnf(i_0_392, plain, (k3_xcmplx_0(np__23,np__1)=np__23)).#trainneg
cnf(i_0_135, plain, (k3_xcmplx_0(np__23,np__7)=np__161)).#trainneg
cnf(i_0_509, plain, (k3_xcmplx_0(np__0,np__1)=np__0)).#trainneg
cnf(i_0_541, plain, (k3_xcmplx_0(np__0,np__2)=np__0)).#trainneg
cnf(i_0_462, plain, (k3_xcmplx_0(np__0,np__0)=np__0)).#trainneg
cnf(i_0_397, plain, (k3_xcmplx_0(np__0,np__32)=np__0)).#trainneg
cnf(i_0_275, plain, (~r1_xxreal_0(np__163,np__54))).#trainneg
cnf(i_0_277, plain, (~r1_xxreal_0(np__163,np__162))).#trainneg
cnf(i_0_678, plain, (r1_xxreal_0(np__1,np__3))).#trainneg
cnf(i_0_274, plain, (~r1_xxreal_0(np__163,np__81))).#trainneg
cnf(i_0_113, plain, (m2_subset_1(np__13,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_281, plain, (~r1_xxreal_0(np__163,np__154))).#trainneg
cnf(i_0_280, plain, (~r1_xxreal_0(np__163,np__156))).#trainneg
cnf(i_0_559, plain, (v1_card_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_495, plain, (v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_396, plain, (k3_xcmplx_0(np__32,np__0)=np__0)).#trainneg
cnf(i_0_120, plain, (k3_xcmplx_0(np__9,np__17)=np__153)).#trainneg
cnf(i_0_123, plain, (k3_xcmplx_0(np__9,np__9)=np__81)).#trainneg
cnf(i_0_561, plain, (k3_xcmplx_0(np__8,np__1)=np__8)).#trainneg
cnf(i_0_94, plain, (k3_xcmplx_0(np__8,np__19)=np__152)).#trainneg
cnf(i_0_461, plain, (k3_xcmplx_0(np__10,np__1)=np__10)).#trainneg
cnf(i_0_179, plain, (k3_xcmplx_0(np__12,np__13)=np__156)).#trainneg
cnf(i_0_162, plain, (k3_xcmplx_0(np__14,np__11)=np__154)).#trainneg
cnf(i_0_214, plain, (k2_xcmplx_0(np__1,np__1)=np__2)).#trainneg
cnf(i_0_297, plain, (k2_xcmplx_0(np__1,np__2)=np__3)).#trainneg
cnf(i_0_202, plain, (k2_xcmplx_0(np__1,np__7)=np__8)).#trainneg
cnf(i_0_337, plain, (k2_xcmplx_0(np__1,np__11)=np__12)).#trainneg
cnf(i_0_469, plain, (~r1_xxreal_0(np__163,np__9))).#trainneg
cnf(i_0_291, plain, (~r1_xxreal_0(np__163,np__8))).#trainneg
cnf(i_0_768, plain, (r1_xxreal_0(np__1,np__5))).#trainneg
cnf(i_0_427, plain, (~r1_xxreal_0(np__163,np__10))).#trainneg
cnf(i_0_272, plain, (k2_xcmplx_0(np__1,np__13)=np__14)).#trainneg
cnf(i_0_50, plain, (m2_subset_1(np__17,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_278, plain, (~r1_xxreal_0(np__163,np__161))).#trainneg
cnf(i_0_370, plain, (~r1_xxreal_0(np__163,np__12))).#trainneg
cnf(i_0_617, plain, (v4_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_367, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_536, plain, (k2_xcmplx_0(np__1,np__0)=np__1)).#trainneg
cnf(i_0_173, plain, (k2_xcmplx_0(np__1,np__9)=np__10)).#trainneg
cnf(i_0_182, plain, (k2_xcmplx_0(np__1,np__8)=np__9)).#trainneg
cnf(i_0_164, plain, (k2_xcmplx_0(np__1,np__10)=np__11)).#trainneg
cnf(i_0_220, plain, (k2_xcmplx_0(np__1,np__12)=np__13)).#trainneg
cnf(i_0_215, plain, (k2_xcmplx_0(np__2,np__1)=np__3)).#trainneg
cnf(i_0_298, plain, (k2_xcmplx_0(np__2,np__3)=np__5)).#trainneg
cnf(i_0_697, plain, (k2_xcmplx_0(np__2,np__0)=np__2)).#trainneg
cnf(i_0_150, plain, (k2_xcmplx_0(np__2,np__9)=np__11)).#trainneg
cnf(i_0_158, plain, (k2_xcmplx_0(np__2,np__8)=np__10)).#trainneg
cnf(i_0_324, plain, (k2_xcmplx_0(np__2,np__12)=np__14)).#trainneg
cnf(i_0_273, plain, (k2_xcmplx_0(np__3,np__2)=np__5)).#trainneg
cnf(i_0_288, plain, (~r1_xxreal_0(np__163,np__14))).#trainneg
cnf(i_0_660, plain, (r1_xxreal_0(np__1,np__7))).#trainneg
cnf(i_0_302, plain, (k2_xcmplx_0(np__3,np__5)=np__8)).#trainneg
cnf(i_0_43, plain, (m2_subset_1(np__19,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_244, plain, (~r1_xxreal_0(np__841,np__1))).#trainneg
cnf(i_0_368, plain, (~r1_xxreal_0(np__841,np__2))).#trainneg
cnf(i_0_217, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_463, plain, (v3_ordinal1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_585, plain, (k2_xcmplx_0(np__3,np__0)=np__3)).#trainneg
cnf(i_0_187, plain, (k2_xcmplx_0(np__3,np__160)=np__163)).#trainneg
cnf(i_0_186, plain, (k2_xcmplx_0(np__3,np__8)=np__11)).#trainneg
cnf(i_0_208, plain, (k2_xcmplx_0(np__5,np__3)=np__8)).#trainneg
cnf(i_0_184, plain, (k2_xcmplx_0(np__5,np__8)=np__13)).#trainneg
cnf(i_0_210, plain, (k2_xcmplx_0(np__7,np__1)=np__8)).#trainneg
cnf(i_0_514, plain, (k2_xcmplx_0(np__7,np__0)=np__7)).#trainneg
cnf(i_0_166, plain, (k2_xcmplx_0(np__7,np__156)=np__163)).#trainneg
cnf(i_0_347, plain, (k2_xcmplx_0(np__11,np__1)=np__12)).#trainneg
cnf(i_0_138, plain, (k2_xcmplx_0(np__11,np__8)=np__19)).#trainneg
cnf(i_0_314, plain, (k2_xcmplx_0(np__13,np__1)=np__14)).#trainneg
cnf(i_0_386, plain, (k2_xcmplx_0(np__0,np__1)=np__1)).#trainneg
cnf(i_0_507, plain, (k2_xcmplx_0(np__0,np__2)=np__2)).#trainneg
cnf(i_0_390, plain, (~r1_xxreal_0(np__841,np__3))).#trainneg
cnf(i_0_535, plain, (~r1_xxreal_0(np__841,np__5))).#trainneg
cnf(i_0_747, plain, (r1_xxreal_0(np__1,np__11))).#trainneg
cnf(i_0_465, plain, (~r1_xxreal_0(np__841,np__7))).#trainneg
cnf(i_0_578, plain, (k2_xcmplx_0(np__0,np__3)=np__3)).#trainneg
cnf(i_0_33, plain, (m2_subset_1(np__23,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_448, plain, (~r1_xxreal_0(np__841,np__11))).#trainneg
cnf(i_0_581, plain, (~r1_xxreal_0(np__841,np__13))).#trainneg
cnf(i_0_555, plain, (v3_membered(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_403, plain, (v1_xreal_0(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_515, plain, (k2_xcmplx_0(np__0,np__7)=np__7)).#trainneg
cnf(i_0_431, plain, (k2_xcmplx_0(np__0,np__0)=np__0)).#trainneg
cnf(i_0_580, plain, (k2_xcmplx_0(np__0,np__8)=np__8)).#trainneg
cnf(i_0_140, plain, (k2_xcmplx_0(np__152,np__11)=np__163)).#trainneg
cnf(i_0_174, plain, (k2_xcmplx_0(np__162,np__1)=np__163)).#trainneg
cnf(i_0_204, plain, (k2_xcmplx_0(np__9,np__1)=np__10)).#trainneg
cnf(i_0_102, plain, (k2_xcmplx_0(np__9,np__2)=np__11)).#trainneg
cnf(i_0_144, plain, (k2_xcmplx_0(np__9,np__154)=np__163)).#trainneg
cnf(i_0_181, plain, (k2_xcmplx_0(np__8,np__1)=np__9)).#trainneg
cnf(i_0_157, plain, (k2_xcmplx_0(np__8,np__2)=np__10)).#trainneg
cnf(i_0_410, plain, (~r1_xxreal_0(np__841,np__17))).#trainneg
cnf(i_0_402, plain, (~r1_xxreal_0(np__841,np__19))).#trainneg
cnf(i_0_729, plain, (r1_xxreal_0(np__1,np__13))).#trainneg
cnf(i_0_393, plain, (~r1_xxreal_0(np__841,np__23))).#trainneg
cnf(i_0_201, plain, (k2_xcmplx_0(np__8,np__3)=np__11)).#trainneg
cnf(i_0_650, plain, (m2_subset_1(k6_numbers,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_388, plain, (~r1_xxreal_0(np__841,np__0))).#trainneg
cnf(i_0_715, plain, (v1_finseq_1(X1)|~v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_183, plain, (k2_xcmplx_0(np__8,np__5)=np__13)).#trainneg
cnf(i_0_137, plain, (k2_xcmplx_0(np__8,np__11)=np__19)).#trainneg
cnf(i_0_579, plain, (k2_xcmplx_0(np__8,np__0)=np__8)).#trainneg
cnf(i_0_170, plain, (k2_xcmplx_0(np__10,np__1)=np__11)).#trainneg
cnf(i_0_147, plain, (k2_xcmplx_0(np__10,np__153)=np__163)).#trainneg
cnf(i_0_224, plain, (k2_xcmplx_0(np__12,np__1)=np__13)).#trainneg
cnf(i_0_325, plain, (k2_xcmplx_0(np__12,np__2)=np__14)).#trainneg
cnf(i_0_646, plain, (r1_xxreal_0(np__1,np__17))).#trainneg
cnf(i_0_643, plain, (r1_xxreal_0(np__1,np__23))).#trainneg
cnf(i_0_710, plain, (r1_xxreal_0(np__1,np__9))).#trainneg
cnf(i_0_679, plain, (r1_xxreal_0(np__1,np__8))).#trainneg
cnf(i_0_666, plain, (r1_xxreal_0(np__1,np__10))).#trainneg
cnf(i_0_584, plain, (r1_xxreal_0(np__841,np__841))).#trainneg
cnf(i_0_592, plain, (r1_xxreal_0(np__2,np__163))).#trainneg
cnf(i_0_219, plain, (~r1_xxreal_0(np__3,np__2))).#trainneg
cnf(i_0_492, plain, (r1_xxreal_0(np__2,np__2))).#trainneg
cnf(i_0_234, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_405, plain, (~r1_xxreal_0(np__5,np__1))).#trainneg
cnf(i_0_681, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_698, plain, (v5_xxreal_2(X1)|~v3_membered(X1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_572, plain, (r1_xxreal_0(np__2,np__3))).#trainneg
cnf(i_0_618, plain, (r1_xxreal_0(np__2,np__5))).#trainneg
cnf(i_0_601, plain, (r1_xxreal_0(np__2,np__7))).#trainneg
cnf(i_0_722, plain, (r1_xxreal_0(np__2,np__11))).#trainneg
cnf(i_0_661, plain, (r1_xxreal_0(np__2,np__13))).#trainneg
cnf(i_0_596, plain, (r1_xxreal_0(np__2,np__17))).#trainneg
cnf(i_0_708, plain, (r1_xxreal_0(np__2,np__23))).#trainneg
cnf(i_0_665, plain, (r1_xxreal_0(np__2,np__81))).#trainneg
cnf(i_0_674, plain, (r1_xxreal_0(np__2,np__9))).#trainneg
cnf(i_0_714, plain, (r1_xxreal_0(np__2,np__8))).#trainneg
cnf(i_0_604, plain, (r1_xxreal_0(np__2,np__10))).#trainneg
cnf(i_0_657, plain, (r1_xxreal_0(np__3,np__3))).#trainneg
cnf(i_0_688, plain, (r1_xxreal_0(np__3,np__5))).#trainneg
cnf(i_0_651, plain, (r1_xxreal_0(np__3,np__7))).#trainneg
cnf(i_0_245, plain, (~r1_xxreal_0(np__5,np__2))).#trainneg
cnf(i_0_744, plain, (r1_xxreal_0(np__3,np__11))).#trainneg
cnf(i_0_422, plain, (~r1_xxreal_0(np__5,np__0))).#trainneg
cnf(i_0_725, plain, (r1_xxreal_0(np__3,np__13))).#trainneg
cnf(i_0_25, plain, (m2_subset_1(np__153,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_228, plain, (~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_769, plain, (~v7_ordinal1(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_312, plain, (v7_ordinal1(X1)|~v3_ordinal1(X1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_644, plain, (r1_xxreal_0(np__3,np__17))).#trainneg
cnf(i_0_664, plain, (r1_xxreal_0(np__3,np__23))).#trainneg
cnf(i_0_704, plain, (r1_xxreal_0(np__3,np__9))).#trainneg
cnf(i_0_668, plain, (r1_xxreal_0(np__3,np__8))).#trainneg
cnf(i_0_659, plain, (r1_xxreal_0(np__3,np__10))).#trainneg
cnf(i_0_723, plain, (r1_xxreal_0(np__3,np__14))).#trainneg
cnf(i_0_669, plain, (r1_xxreal_0(np__5,np__5))).#trainneg
cnf(i_0_728, plain, (r1_xxreal_0(np__5,np__7))).#trainneg
cnf(i_0_745, plain, (r1_xxreal_0(np__5,np__11))).#trainneg
cnf(i_0_727, plain, (r1_xxreal_0(np__5,np__13))).#trainneg
cnf(i_0_634, plain, (r1_xxreal_0(np__5,np__17))).#trainneg
cnf(i_0_594, plain, (r1_xxreal_0(np__5,np__23))).#trainneg
cnf(i_0_707, plain, (r1_xxreal_0(np__5,np__9))).#trainneg
cnf(i_0_670, plain, (r1_xxreal_0(np__5,np__8))).#trainneg
cnf(i_0_375, plain, (~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_556, plain, (~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_663, plain, (r1_xxreal_0(np__5,np__10))).#trainneg
cnf(i_0_693, plain, (r1_xxreal_0(np__5,np__14))).#trainneg
cnf(i_0_29, plain, (m2_subset_1(np__152,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_218, plain, (~r1_xxreal_0(np__11,np__1))).#trainneg
cnf(i_0_319, plain, (~r1_xxreal_0(np__11,np__2))).#trainneg
cnf(i_0_398, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_529, plain, (~v1_xboole_0(X1)|~v2_xxreal_0(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_647, plain, (r1_xxreal_0(np__7,np__7))).#trainneg
cnf(i_0_730, plain, (r1_xxreal_0(np__7,np__11))).#trainneg
cnf(i_0_694, plain, (r1_xxreal_0(np__7,np__13))).#trainneg
cnf(i_0_600, plain, (r1_xxreal_0(np__7,np__17))).#trainneg
cnf(i_0_652, plain, (r1_xxreal_0(np__7,np__19))).#trainneg
cnf(i_0_582, plain, (r1_xxreal_0(np__7,np__23))).#trainneg
cnf(i_0_672, plain, (r1_xxreal_0(np__7,np__9))).#trainneg
cnf(i_0_640, plain, (r1_xxreal_0(np__7,np__8))).#trainneg
cnf(i_0_648, plain, (r1_xxreal_0(np__7,np__10))).#trainneg
cnf(i_0_675, plain, (r1_xxreal_0(np__7,np__14))).#trainneg
cnf(i_0_726, plain, (r1_xxreal_0(np__11,np__11))).#trainneg
cnf(i_0_767, plain, (r1_xxreal_0(np__11,np__13))).#trainneg
cnf(i_0_721, plain, (r1_xxreal_0(np__11,np__17))).#trainneg
cnf(i_0_712, plain, (r1_xxreal_0(np__11,np__23))).#trainneg
cnf(i_0_346, plain, (~r1_xxreal_0(np__11,np__3))).#trainneg
cnf(i_0_497, plain, (~r1_xxreal_0(np__11,np__5))).#trainneg
cnf(i_0_736, plain, (r1_xxreal_0(np__13,np__13))).#trainneg
cnf(i_0_409, plain, (~r1_xxreal_0(np__11,np__7))).#trainneg
cnf(i_0_692, plain, (r1_xxreal_0(np__13,np__17))).#trainneg
cnf(i_0_97, plain, (m2_subset_1(np__32,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_527, plain, (~v2_xxreal_0(X1)|~v1_xxreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_320, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_683, plain, (r1_xxreal_0(np__13,np__23))).#trainneg
cnf(i_0_746, plain, (r1_xxreal_0(np__13,np__14))).#trainneg
cnf(i_0_597, plain, (r1_xxreal_0(np__17,np__17))).#trainneg
cnf(i_0_588, plain, (r1_xxreal_0(np__17,np__23))).#trainneg
cnf(i_0_764, plain, (r1_xxreal_0(np__19,np__19))).#trainneg
cnf(i_0_616, plain, (r1_xxreal_0(np__23,np__23))).#trainneg
cnf(i_0_523, plain, (r1_xxreal_0(np__0,np__1))).#trainneg
cnf(i_0_525, plain, (r1_xxreal_0(np__0,np__841))).#trainneg
cnf(i_0_526, plain, (r1_xxreal_0(np__0,np__2))).#trainneg
cnf(i_0_533, plain, (r1_xxreal_0(np__0,np__3))).#trainneg
cnf(i_0_554, plain, (r1_xxreal_0(np__0,np__5))).#trainneg
cnf(i_0_571, plain, (r1_xxreal_0(np__0,np__7))).#trainneg
cnf(i_0_583, plain, (r1_xxreal_0(np__0,np__11))).#trainneg
cnf(i_0_570, plain, (r1_xxreal_0(np__0,np__13))).#trainneg
cnf(i_0_313, plain, (~r1_xxreal_0(np__11,np__8))).#trainneg
cnf(i_0_425, plain, (~r1_xxreal_0(np__11,np__10))).#trainneg
cnf(i_0_538, plain, (r1_xxreal_0(np__0,np__17))).#trainneg
cnf(i_0_350, plain, (~r1_xxreal_0(np__13,np__1))).#trainneg
cnf(i_0_522, plain, (r1_xxreal_0(np__0,np__0))).#trainneg
cnf(i_0_88, plain, (m2_subset_1(np__160,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_471, plain, (~r1_xxreal_0(np__13,np__2))).#trainneg
cnf(i_0_501, plain, (~r1_xxreal_0(np__13,np__3))).#trainneg
cnf(i_0_391, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_252, plain, (r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_nat_d(X1,X2))).#trainneg
cnf(i_0_558, plain, (r1_xxreal_0(np__0,np__81))).#trainneg
cnf(i_0_505, plain, (r1_xxreal_0(np__0,np__9))).#trainneg
cnf(i_0_550, plain, (r1_xxreal_0(np__0,np__8))).#trainneg
cnf(i_0_548, plain, (r1_xxreal_0(np__0,np__10))).#trainneg
cnf(i_0_645, plain, (r1_xxreal_0(np__0,np__12))).#trainneg
cnf(i_0_677, plain, (r1_xxreal_0(np__0,np__14))).#trainneg
cnf(i_0_649, plain, (r1_xxreal_0(np__81,np__81))).#trainneg
cnf(i_0_743, plain, (r1_xxreal_0(np__9,np__11))).#trainneg
cnf(i_0_741, plain, (r1_xxreal_0(np__9,np__13))).#trainneg
cnf(i_0_673, plain, (r1_xxreal_0(np__9,np__17))).#trainneg
cnf(i_0_662, plain, (r1_xxreal_0(np__9,np__23))).#trainneg
cnf(i_0_676, plain, (r1_xxreal_0(np__9,np__9))).#trainneg
cnf(i_0_687, plain, (r1_xxreal_0(np__9,np__10))).#trainneg
cnf(i_0_740, plain, (r1_xxreal_0(np__9,np__14))).#trainneg
cnf(i_0_667, plain, (~r1_xxreal_0(np__13,np__5))).#trainneg
cnf(i_0_737, plain, (r1_xxreal_0(np__8,np__11))).#trainneg
cnf(i_0_235, plain, (~r1_xxreal_0(np__13,np__11))).#trainneg
cnf(i_0_709, plain, (r1_xxreal_0(np__8,np__13))).#trainneg
cnf(i_0_92, plain, (m2_subset_1(np__54,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_498, plain, (~r1_xxreal_0(np__13,np__0))).#trainneg
cnf(i_0_598, plain, (~r1_xxreal_0(np__13,np__9))).#trainneg
cnf(i_0_223, plain, (r1_nat_d(X1,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_401, plain, (r1_int_1(X1,X1)|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_633, plain, (r1_xxreal_0(np__8,np__17))).#trainneg
cnf(i_0_602, plain, (r1_xxreal_0(np__8,np__23))).#trainneg
cnf(i_0_713, plain, (r1_xxreal_0(np__8,np__9))).#trainneg
cnf(i_0_695, plain, (r1_xxreal_0(np__8,np__8))).#trainneg
cnf(i_0_642, plain, (r1_xxreal_0(np__8,np__10))).#trainneg
cnf(i_0_705, plain, (r1_xxreal_0(np__8,np__14))).#trainneg
cnf(i_0_731, plain, (r1_xxreal_0(np__10,np__11))).#trainneg
cnf(i_0_696, plain, (r1_xxreal_0(np__10,np__13))).#trainneg
cnf(i_0_603, plain, (r1_xxreal_0(np__10,np__17))).#trainneg
cnf(i_0_599, plain, (r1_xxreal_0(np__10,np__23))).#trainneg
cnf(i_0_605, plain, (r1_xxreal_0(np__10,np__10))).#trainneg
cnf(i_0_671, plain, (r1_xxreal_0(np__10,np__14))).#trainneg
cnf(i_0_658, plain, (r1_xxreal_0(np__12,np__17))).#trainneg
cnf(i_0_766, plain, (r1_xxreal_0(np__12,np__23))).#trainneg
cnf(i_0_519, plain, (~r1_xxreal_0(np__13,np__8))).#trainneg
cnf(i_0_569, plain, (~r1_xxreal_0(np__13,np__10))).#trainneg
cnf(i_0_765, plain, (r1_xxreal_0(np__12,np__12))).#trainneg
cnf(i_0_557, plain, (~r1_xxreal_0(np__13,np__12))).#trainneg
cnf(i_0_691, plain, (r1_xxreal_0(np__14,np__17))).#trainneg
cnf(i_0_70, plain, (m2_subset_1(np__162,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_206, plain, (~r1_xxreal_0(np__17,np__1))).#trainneg
cnf(i_0_250, plain, (~r1_xxreal_0(np__17,np__2))).#trainneg
cnf(i_0_485, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_504, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).#trainneg
cnf(i_0_706, plain, (r1_xxreal_0(np__14,np__23))).#trainneg
cnf(i_0_711, plain, (r1_xxreal_0(np__14,np__14))).#trainneg
cnf(i_0_64, plain, (m1_subset_1(np__163,k1_numbers))).#trainneg
cnf(i_0_115, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_360, plain, (m1_subset_1(np__841,k1_numbers))).#trainneg
cnf(i_0_361, plain, (m1_subset_1(np__841,k5_numbers))).#trainneg
cnf(i_0_14, plain, (m1_subset_1(np__2,k1_numbers))).#trainneg
cnf(i_0_56, plain, (m1_subset_1(np__3,k1_numbers))).#trainneg
cnf(i_0_52, plain, (m1_subset_1(np__5,k1_numbers))).#trainneg
cnf(i_0_248, plain, (~r1_xxreal_0(np__17,np__3))).#trainneg
cnf(i_0_451, plain, (~r1_xxreal_0(np__17,np__5))).#trainneg
cnf(i_0_18, plain, (m1_subset_1(np__7,k1_numbers))).#trainneg
cnf(i_0_395, plain, (~r1_xxreal_0(np__17,np__7))).#trainneg
cnf(i_0_62, plain, (m2_subset_1(np__81,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_356, plain, (~r1_xxreal_0(np__17,np__11))).#trainneg
cnf(i_0_499, plain, (~r1_xxreal_0(np__17,np__13))).#trainneg
cnf(i_0_251, plain, (r1_nat_d(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_int_1(X1,X2))).#trainneg
cnf(i_0_310, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_37, plain, (m1_subset_1(np__11,k1_numbers))).#trainneg
cnf(i_0_111, plain, (m1_subset_1(np__13,k1_numbers))).#trainneg
cnf(i_0_48, plain, (m1_subset_1(np__17,k1_numbers))).#trainneg
cnf(i_0_41, plain, (m1_subset_1(np__19,k1_numbers))).#trainneg
cnf(i_0_31, plain, (m1_subset_1(np__23,k1_numbers))).#trainneg
cnf(i_0_328, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_232, plain, (m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_233, plain, (m1_subset_1(np__0,k5_numbers))).#trainneg
cnf(i_0_23, plain, (m1_subset_1(np__153,k1_numbers))).#trainneg
cnf(i_0_295, plain, (~r1_xxreal_0(np__17,np__0))).#trainneg
cnf(i_0_494, plain, (~r1_xxreal_0(np__17,np__9))).#trainneg
cnf(i_0_322, plain, (~r1_xxreal_0(np__17,np__8))).#trainneg
cnf(i_0_27, plain, (m1_subset_1(np__152,k1_numbers))).#trainneg
cnf(i_0_74, plain, (m2_subset_1(np__154,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_189, plain, (~r1_xxreal_0(np__19,np__1))).#trainneg
cnf(i_0_434, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_631, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X2)|r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_28, plain, (m1_subset_1(np__152,k5_numbers))).#trainneg
cnf(i_0_95, plain, (m1_subset_1(np__32,k1_numbers))).#trainneg
cnf(i_0_86, plain, (m1_subset_1(np__160,k1_numbers))).#trainneg
cnf(i_0_90, plain, (m1_subset_1(np__54,k1_numbers))).#trainneg
cnf(i_0_68, plain, (m1_subset_1(np__162,k1_numbers))).#trainneg
cnf(i_0_69, plain, (m1_subset_1(np__162,k5_numbers))).#trainneg
cnf(i_0_60, plain, (m1_subset_1(np__81,k1_numbers))).#trainneg
cnf(i_0_72, plain, (m1_subset_1(np__154,k1_numbers))).#trainneg
cnf(i_0_76, plain, (m1_subset_1(np__156,k1_numbers))).#trainneg
cnf(i_0_247, plain, (~r1_xxreal_0(np__19,np__2))).#trainneg
cnf(i_0_299, plain, (~r1_xxreal_0(np__19,np__3))).#trainneg
cnf(i_0_453, plain, (~r1_xxreal_0(np__19,np__5))).#trainneg
cnf(i_0_82, plain, (m1_subset_1(np__9,k1_numbers))).#trainneg
cnf(i_0_78, plain, (m2_subset_1(np__156,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_389, plain, (~r1_xxreal_0(np__19,np__7))).#trainneg
cnf(i_0_151, plain, (m1_subset_1(np__8,k1_numbers))).#trainneg
cnf(i_0_103, plain, (m1_subset_1(np__10,k1_numbers))).#trainneg
cnf(i_0_107, plain, (m1_subset_1(np__161,k1_numbers))).#trainneg
cnf(i_0_124, plain, (m1_subset_1(np__12,k1_numbers))).#trainneg
cnf(i_0_128, plain, (m1_subset_1(np__14,k1_numbers))).#trainneg
cnf(i_0_547, plain, (m1_subset_1(esk6_0,k1_numbers))).#trainneg
cnf(i_0_754, plain, (v4_relat_1(esk13_0,k5_numbers))).#trainneg
cnf(i_0_761, plain, (v4_relat_1(esk14_0,k5_numbers))).#trainneg
cnf(i_0_493, plain, (~r1_xxreal_0(np__19,np__13))).#trainneg
cnf(i_0_318, plain, (~r1_xxreal_0(np__19,np__17))).#trainneg
cnf(i_0_200, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_290, plain, (~r1_xxreal_0(np__19,np__0))).#trainneg
cnf(i_0_732, plain, (v1_xboole_0(X1)|v1_finset_1(esk12_1(X1)))).#trainneg
cnf(i_0_84, plain, (m2_subset_1(np__9,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_486, plain, (~r1_xxreal_0(np__19,np__9))).#trainneg
cnf(i_0_317, plain, (~r1_xxreal_0(np__19,np__8))).#trainneg
cnf(i_0_211, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_338, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_735, plain, (v1_finset_1(k1_zfmisc_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_193, plain, (~r1_xxreal_0(np__23,np__1))).#trainneg
cnf(i_0_266, plain, (~r1_xxreal_0(np__23,np__3))).#trainneg
cnf(i_0_437, plain, (~r1_xxreal_0(np__23,np__5))).#trainneg
cnf(i_0_378, plain, (~r1_xxreal_0(np__23,np__7))).#trainneg
cnf(i_0_340, plain, (~r1_xxreal_0(np__23,np__11))).#trainneg
cnf(i_0_477, plain, (~r1_xxreal_0(np__23,np__13))).#trainneg
cnf(i_0_304, plain, (~r1_xxreal_0(np__23,np__17))).#trainneg
cnf(i_0_292, plain, (~r1_xxreal_0(np__23,np__19))).#trainneg
cnf(i_0_254, plain, (~r1_xxreal_0(np__23,np__0))).#trainneg
cnf(i_0_426, plain, (~r1_xxreal_0(np__23,np__10))).#trainneg
cnf(i_0_287, plain, (~r1_xxreal_0(np__23,np__14))).#trainneg
cnf(i_0_354, plain, (~r1_xxreal_0(np__153,np__17))).#trainneg
cnf(i_0_323, plain, (~r1_xxreal_0(np__153,np__0))).#trainneg
cnf(i_0_521, plain, (~r1_xxreal_0(np__153,np__9))).#trainneg
cnf(i_0_153, plain, (m2_subset_1(np__8,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_480, plain, (~r1_xxreal_0(np__153,np__10))).#trainneg
cnf(i_0_196, plain, (~r1_xxreal_0(np__152,np__1))).#trainneg
cnf(i_0_105, plain, (m2_subset_1(np__10,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_343, plain, (~r1_xxreal_0(np__152,np__11))).#trainneg
cnf(i_0_349, plain, (~r1_xxreal_0(np__152,np__19))).#trainneg
cnf(i_0_222, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_595, plain, (v1_int_1(X2)|~v5_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_261, plain, (~r1_xxreal_0(np__152,np__0))).#trainneg
cnf(i_0_301, plain, (~r1_xxreal_0(np__152,np__8))).#trainneg
cnf(i_0_192, plain, (~r1_xxreal_0(np__32,np__1))).#trainneg
cnf(i_0_237, plain, (~r1_xxreal_0(np__32,np__2))).#trainneg
cnf(i_0_265, plain, (~r1_xxreal_0(np__32,np__3))).#trainneg
cnf(i_0_436, plain, (~r1_xxreal_0(np__32,np__5))).#trainneg
cnf(i_0_369, plain, (~r1_xxreal_0(np__32,np__0))).#trainneg
cnf(i_0_269, plain, (~r1_xxreal_0(np__160,np__3))).#trainneg
cnf(i_0_439, plain, (~r1_xxreal_0(np__160,np__5))).#trainneg
cnf(i_0_258, plain, (~r1_xxreal_0(np__160,np__0))).#trainneg
cnf(i_0_286, plain, (~r1_xxreal_0(np__160,np__32))).#trainneg
cnf(i_0_191, plain, (~r1_xxreal_0(np__54,np__1))).#trainneg
cnf(i_0_264, plain, (~r1_xxreal_0(np__54,np__3))).#trainneg
cnf(i_0_253, plain, (~r1_xxreal_0(np__54,np__0))).#trainneg
cnf(i_0_195, plain, (~r1_xxreal_0(np__162,np__1))).#trainneg
cnf(i_0_240, plain, (~r1_xxreal_0(np__162,np__2))).#trainneg
cnf(i_0_109, plain, (m2_subset_1(np__161,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_268, plain, (~r1_xxreal_0(np__162,np__3))).#trainneg
cnf(i_0_256, plain, (~r1_xxreal_0(np__162,np__0))).#trainneg
cnf(i_0_126, plain, (m2_subset_1(np__12,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_284, plain, (~r1_xxreal_0(np__162,np__54))).#trainneg
cnf(i_0_283, plain, (~r1_xxreal_0(np__162,np__81))).#trainneg
cnf(i_0_466, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_445, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_190, plain, (~r1_xxreal_0(np__81,np__1))).#trainneg
cnf(i_0_236, plain, (~r1_xxreal_0(np__81,np__2))).#trainneg
cnf(i_0_263, plain, (~r1_xxreal_0(np__81,np__3))).#trainneg
cnf(i_0_246, plain, (~r1_xxreal_0(np__81,np__0))).#trainneg
cnf(i_0_468, plain, (~r1_xxreal_0(np__81,np__9))).#trainneg
cnf(i_0_342, plain, (~r1_xxreal_0(np__154,np__11))).#trainneg
cnf(i_0_260, plain, (~r1_xxreal_0(np__154,np__0))).#trainneg
cnf(i_0_470, plain, (~r1_xxreal_0(np__154,np__9))).#trainneg
cnf(i_0_289, plain, (~r1_xxreal_0(np__154,np__14))).#trainneg
cnf(i_0_380, plain, (~r1_xxreal_0(np__156,np__7))).#trainneg
cnf(i_0_479, plain, (~r1_xxreal_0(np__156,np__13))).#trainneg
cnf(i_0_259, plain, (~r1_xxreal_0(np__156,np__0))).#trainneg
cnf(i_0_371, plain, (~r1_xxreal_0(np__156,np__12))).#trainneg
cnf(i_0_344, plain, (~r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_444, plain, (~r1_xxreal_0(np__9,np__2))).#trainneg
cnf(i_0_464, plain, (~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_130, plain, (m2_subset_1(np__14,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_551, plain, (~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_565, plain, (~r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_482, plain, (~r1_xxreal_0(np__9,np__8))).#trainneg
cnf(i_0_205, plain, (~r1_xxreal_0(np__8,np__1))).#trainneg
cnf(i_0_739, plain, (v1_xboole_0(X2)|v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_738, plain, (v1_xboole_0(X1)|v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_249, plain, (~r1_xxreal_0(np__8,np__2))).#trainneg
cnf(i_0_296, plain, (~r1_xxreal_0(np__8,np__3))).#trainneg
cnf(i_0_450, plain, (~r1_xxreal_0(np__8,np__5))).#trainneg
cnf(i_0_377, plain, (~r1_xxreal_0(np__8,np__7))).#trainneg
cnf(i_0_315, plain, (~r1_xxreal_0(np__8,np__0))).#trainneg
cnf(i_0_294, plain, (~r1_xxreal_0(np__10,np__1))).#trainneg
cnf(i_0_400, plain, (~r1_xxreal_0(np__10,np__2))).#trainneg
cnf(i_0_421, plain, (~r1_xxreal_0(np__10,np__3))).#trainneg
cnf(i_0_635, plain, (~r1_xxreal_0(np__10,np__5))).#trainneg
cnf(i_0_506, plain, (~r1_xxreal_0(np__10,np__7))).#trainneg
cnf(i_0_560, plain, (~r1_xxreal_0(np__10,np__9))).#trainneg
cnf(i_0_429, plain, (~r1_xxreal_0(np__10,np__8))).#trainneg
cnf(i_0_241, plain, (~r1_xxreal_0(np__161,np__2))).#trainneg
cnf(i_0_411, plain, (~r1_xxreal_0(np__161,np__7))).#trainneg
cnf(i_0_285, plain, (~r1_xxreal_0(np__161,np__23))).#trainneg
cnf(i_0_257, plain, (~r1_xxreal_0(np__161,np__0))).#trainneg
cnf(i_0_225, plain, (~r1_xxreal_0(np__12,np__1))).#trainneg
cnf(i_0_339, plain, (~r1_xxreal_0(np__12,np__2))).#trainneg
cnf(i_0_366, plain, (~r1_xxreal_0(np__12,np__3))).#trainneg
cnf(i_0_520, plain, (~r1_xxreal_0(np__12,np__5))).#trainneg
cnf(i_0_454, plain, (~r1_xxreal_0(np__12,np__7))).#trainneg
cnf(i_0_574, plain, (v2_xxreal_0(X2)|~v2_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_412, plain, (~r1_xxreal_0(np__12,np__11))).#trainneg
cnf(i_0_359, plain, (~r1_xxreal_0(np__12,np__0))).#trainneg
cnf(i_0_446, plain, (~r1_xxreal_0(np__12,np__9))).#trainneg
cnf(i_0_326, plain, (~r1_xxreal_0(np__12,np__8))).#trainneg
cnf(i_0_517, plain, (~r1_xxreal_0(np__12,np__10))).#trainneg
cnf(i_0_199, plain, (~r1_xxreal_0(np__14,np__1))).#trainneg
cnf(i_0_243, plain, (~r1_xxreal_0(np__14,np__2))).#trainneg
cnf(i_0_271, plain, (~r1_xxreal_0(np__14,np__3))).#trainneg
cnf(i_0_442, plain, (~r1_xxreal_0(np__14,np__5))).#trainneg
cnf(i_0_382, plain, (~r1_xxreal_0(np__14,np__7))).#trainneg
cnf(i_0_348, plain, (~r1_xxreal_0(np__14,np__11))).#trainneg
cnf(i_0_481, plain, (~r1_xxreal_0(np__14,np__13))).#trainneg
cnf(i_0_270, plain, (~r1_xxreal_0(np__14,np__0))).#trainneg
cnf(i_0_473, plain, (~r1_xxreal_0(np__14,np__9))).#trainneg
cnf(i_0_306, plain, (~r1_xxreal_0(np__14,np__8))).#trainneg
cnf(i_0_432, plain, (~r1_xxreal_0(np__14,np__10))).#trainneg
cnf(i_0_373, plain, (~r1_xxreal_0(np__14,np__12))).#trainneg
cnf(i_0_303, plain, (k3_xcmplx_0(X1,k6_numbers)=k6_numbers|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_413, plain, (m1_subset_1(esk2_1(X1),X1))).#trainneg
cnf(i_0_542, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_733, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk12_1(X1)))).#trainneg
cnf(i_0_364, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_209, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_701, plain, (v4_relat_1(esk10_1(X1),X1))).#trainneg
cnf(i_0_638, plain, (v5_relat_1(esk8_1(X1),X1))).#trainneg
cnf(i_0_262, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_142, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_308, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_496, plain, (v1_int_1(X1)|~m1_subset_1(X1,k4_numbers))).#trainneg
cnf(i_0_309, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_441, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_734, plain, (v1_xboole_0(X1)|m1_subset_1(esk12_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_365, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_563, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainneg
cnf(i_0_399, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_686, plain, (v1_finset_1(X2)|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_407, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_771, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_417, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_7, plain, (r1_xxreal_0(X1,np__1)|~v1_int_2(X2)|~v1_int_2(X1)|~v7_ordinal1(X2)|~r1_nat_d(X1,X2)|~r1_xxreal_0(k4_nat_1(X1,X1),X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_655, plain, (v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_159, plain, (v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_452, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_455, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_577, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_440, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_488, plain, (v1_int_1(k3_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_357, plain, (v1_int_1(k2_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_163, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_203, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_101, plain, (k4_nat_1(X1,X2)=k4_nat_1(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_47, plain, (k4_nat_1(X1,X2)=k3_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_139, plain, (m2_subset_1(k24_binop_2(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_207, plain, (m2_subset_1(k4_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_680, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xboole_0(k3_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_345, plain, (m2_subset_1(k2_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_198, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_404, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_197, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_414, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_566, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_423, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,X3),k3_xcmplx_0(X2,X3))=k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_826, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1556, plain, (k3_xcmplx_0(np__1,X1)=X1|~v7_ordinal1(X1))).#trainneg
cnf(i_0_821, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_822, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_817, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_804, plain, (r1_nat_d(X1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_787, plain, (v3_membered(k3_numbers))).#trainneg
cnf(i_0_828, plain, (v7_ordinal1(np__156))).#trainneg
cnf(i_0_775, plain, (esk5_0=k6_numbers)).#trainneg
cnf(i_0_772, plain, (~v1_xboole_0(esk7_0))).#trainneg
cnf(i_0_776, plain, (esk4_0=k6_numbers)).#trainneg
cnf(i_0_777, plain, (esk3_0=k6_numbers)).#trainneg
cnf(i_0_774, plain, (~v7_ordinal1(k5_numbers))).#trainneg
cnf(i_0_1004, plain, (v3_ordinal1(np__163))).#trainneg
cnf(i_0_1187, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_834, plain, (v7_ordinal1(np__162))).#trainneg
cnf(i_0_801, plain, (~v2_xxreal_0(X1)|~v1_xboole_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1064, plain, (v1_xreal_0(np__163))).#trainneg
cnf(i_0_1587, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(X1)|~v1_xboole_0(esk1_1(X1))|~v7_ordinal1(esk1_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_810, plain, (r1_int_1(X1,X1)|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1574, plain, (v3_membered(k5_numbers))).#trainneg
cnf(i_0_838, plain, (v7_ordinal1(np__160))).#trainneg
cnf(i_0_1818, plain, (k24_binop_2(X1,X2)=k3_xcmplx_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_879, plain, (v7_ordinal1(np__161))).#trainneg
cnf(i_0_830, plain, (v7_ordinal1(np__154))).#trainneg
cnf(i_0_1577, plain, (~v1_finset_1(k1_numbers))).#trainneg
cnf(i_0_773, plain, (~r1_xxreal_0(esk7_0,np__1))).#trainneg
cnf(i_0_788, plain, (v3_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_808, plain, (r1_int_1(X1,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_1030, plain, (v1_xreal_0(np__160))).#trainneg
cnf(i_0_1020, plain, (v1_xreal_0(np__156))).#trainneg
cnf(i_0_1012, plain, (v1_xreal_0(np__161))).#trainneg
cnf(i_0_1026, plain, (v1_xreal_0(np__162))).#trainneg
cnf(i_0_1022, plain, (v1_xreal_0(np__154))).#trainneg
cnf(i_0_800, plain, (~v2_xxreal_0(X1)|~v1_xboole_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2149, plain, (r1_xxreal_0(k3_xcmplx_0(esk1_1(X1),esk1_1(X1)),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|~v7_ordinal1(X1)|~m1_subset_1(esk1_1(X1),k5_numbers)|~v7_ordinal1(esk1_1(X1)))).#trainneg
cnf(i_0_842, plain, (v7_ordinal1(np__152))).#trainneg
cnf(i_0_844, plain, (v7_ordinal1(np__153))).#trainneg
cnf(i_0_777, plain, (esk3_0=np__0)).#trainneg
cnf(i_0_155, plain, (X1=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_776, plain, (esk4_0=np__0)).#trainneg
cnf(i_0_2357, plain, (~v7_ordinal1(k1_numbers))).#trainneg
cnf(i_0_2327, plain, (v1_xcmplx_0(np__163))).#trainneg
cnf(i_0_2361, plain, (v3_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_823, plain, (v3_xxreal_0(X1)|v2_xxreal_0(esk7_0)|r1_xxreal_0(esk7_0,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_824, plain, (v3_xxreal_0(X1)|v2_xxreal_0(esk6_0)|r1_xxreal_0(esk6_0,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2314, plain, (k3_xcmplx_0(np__1,np__163)=np__163)).#trainneg
cnf(i_0_456, plain, (v1_xreal_0(k6_numbers))).#trainneg
cnf(i_0_2158, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2345, plain, (k3_xcmplx_0(X2,X1)=k3_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_303, plain, (k3_xcmplx_0(X1,np__0)=np__0|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_458, plain, (v1_xcmplx_0(k6_numbers))).#trainneg
cnf(i_0_531, plain, (v1_xxreal_0(k6_numbers))).#trainneg
cnf(i_0_2159, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_2144, plain, (v3_xxreal_0(k4_nat_1(esk1_1(X1),esk1_1(X1)))|v2_xxreal_0(X1)|v1_xboole_0(k4_nat_1(esk1_1(X1),esk1_1(X1)))|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|~v1_xreal_0(X1)|~v1_xreal_0(k4_nat_1(esk1_1(X1),esk1_1(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2143, plain, (v3_xxreal_0(k4_nat_1(esk1_1(X1),esk1_1(X1)))|v2_xxreal_0(X1)|v1_xboole_0(X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|~v1_xreal_0(X1)|~v1_xreal_0(k4_nat_1(esk1_1(X1),esk1_1(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_846, plain, (v7_ordinal1(np__0))).#trainneg
cnf(i_0_775, plain, (esk5_0=np__0)).#trainneg
cnf(i_0_262, plain, (k2_xcmplx_0(X1,np__0)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_827, plain, (v7_ordinal1(np__9)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_11, plain, (k1_xboole_0=np__0)).#trainneg
cnf(i_0_2309, plain, (r1_xxreal_0(X1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2310, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2318, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2291, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2145, plain, (v2_xxreal_0(X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|~v1_xreal_0(X1)|~v1_xreal_0(k4_nat_1(esk1_1(X1),esk1_1(X1)))|~v2_xxreal_0(k4_nat_1(esk1_1(X1),esk1_1(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2284, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1704, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2375, plain, (v1_xcmplx_0(np__162))).#trainneg
cnf(i_0_2379, plain, (v1_xcmplx_0(np__154))).#trainneg
cnf(i_0_2367, plain, (v1_xcmplx_0(np__156))).#trainneg
cnf(i_0_2371, plain, (v1_xcmplx_0(np__161))).#trainneg
cnf(i_0_1014, plain, (v1_xreal_0(np__10))).#trainneg
cnf(i_0_1036, plain, (v1_xreal_0(np__153))).#trainneg
cnf(i_0_2305, plain, (r1_xxreal_0(esk7_0,esk7_0))).#trainneg
cnf(i_0_2438, plain, (v3_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2558, plain, (r1_xxreal_0(X1,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2425, plain, (r1_xxreal_0(k4_nat_1(np__19,np__19),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|esk1_1(X1)=np__17|esk1_1(X1)=np__23|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2427, plain, (r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|esk1_1(X1)=np__17|esk1_1(X1)=np__23|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2421, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__19)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|esk1_1(X1)=np__17|esk1_1(X1)=np__23|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2459, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(k3_xcmplx_0(X4,X2),X1)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2487, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2306, plain, (r1_xxreal_0(esk6_0,esk6_0))).#trainneg
cnf(i_0_2511, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~m1_subset_1(X3,k1_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2461, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(X1,k24_binop_2(X2,X4))|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(k24_binop_2(X2,X4),k5_numbers))).#trainneg
cnf(i_0_2455, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~r1_nat_d(X2,k2_nat_1(k24_binop_2(X2,X3),X1))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(k2_nat_1(k24_binop_2(X2,X3),X1)))).#trainneg
cnf(i_0_2458, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(k3_xcmplx_0(X2,X4),X1)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2642, plain, (k3_xcmplx_0(np__1,np__162)=np__162)).#trainneg
cnf(i_0_2485, plain, (k3_xcmplx_0(np__163,X1)=k3_xcmplx_0(X1,np__163)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2717, plain, (r1_nat_d(np__23,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__17|esk1_1(X1)=np__13|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_867, plain, (v7_ordinal1(np__841))).#trainneg
cnf(i_0_2540, plain, (r1_xxreal_0(esk11_0,esk11_0))).#trainneg
cnf(i_0_2590, plain, (r1_xxreal_0(X1,np__163)|r1_xxreal_0(np__163,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2329, plain, (v3_xxreal_0(np__163)|v2_xxreal_0(X1)|r1_xxreal_0(X1,np__163)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2734, plain, (r1_nat_d(np__23,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__17|esk1_1(X1)=np__13|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|v1_int_2(np__19)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2722, plain, (r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__17|esk1_1(X1)=np__13|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2720, plain, (r1_xxreal_0(k4_nat_1(np__23,np__23),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__17|esk1_1(X1)=np__13|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2733, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__23)|v1_int_2(X1)|esk1_1(X1)=np__17|esk1_1(X1)=np__13|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|v1_int_2(np__19)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2889, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~r1_nat_d(X2,k2_nat_1(k3_xcmplx_0(X3,X2),X1))|~v7_ordinal1(k2_nat_1(k3_xcmplx_0(X3,X2),X1))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_869, plain, (v7_ordinal1(np__1))).#trainneg
cnf(i_0_2541, plain, (r1_xxreal_0(esk9_0,esk9_0))).#trainneg
cnf(i_0_1188, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2319, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2886, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~r1_nat_d(X2,k2_nat_1(k24_binop_2(X3,X2),X1))|~v7_ordinal1(k2_nat_1(k24_binop_2(X3,X2),X1))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3178, plain, (k3_xcmplx_0(np__163,np__1)=np__163)).#trainneg
cnf(i_0_3077, plain, (r1_nat_d(np__17,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3038, plain, (r1_xxreal_0(np__163,X1)|r1_xxreal_0(X1,np__163)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3080, plain, (r1_xxreal_0(k4_nat_1(np__17,np__17),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3082, plain, (r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3099, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__17)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|r1_xxreal_0(np__841,X1)|v1_int_2(np__19)|v1_int_2(np__23)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3100, plain, (r1_nat_d(np__17,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|r1_xxreal_0(np__841,X1)|v1_int_2(np__19)|v1_int_2(np__23)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3103, plain, (r1_xxreal_0(k4_nat_1(np__17,np__17),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|r1_xxreal_0(np__841,X1)|v1_int_2(np__19)|v1_int_2(np__23)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3209, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_3076, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__17)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_948, plain, (v3_ordinal1(np__14))).#trainneg
cnf(i_0_950, plain, (v3_ordinal1(np__12))).#trainneg
cnf(i_0_952, plain, (v3_ordinal1(np__161))).#trainneg
cnf(i_0_2543, plain, (r1_xxreal_0(np__161,np__161))).#trainneg
cnf(i_0_3067, plain, (v3_xxreal_0(np__163)|v2_xxreal_0(X1)|r1_xxreal_0(X1,np__163)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2591, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3247, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__13,X1)|X2!=k2_xcmplx_0(np__13,X1)|~r1_nat_d(np__13,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3224, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3253, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__5,X1)|X2!=k2_xcmplx_0(np__10,X1)|~r1_nat_d(np__5,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3249, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__11,X1)|X2!=k2_xcmplx_0(np__11,X1)|~r1_nat_d(np__11,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3285, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~r1_nat_d(X2,k2_nat_1(k3_xcmplx_0(X2,X3),X1))|~v7_ordinal1(k2_nat_1(k3_xcmplx_0(X2,X3),X1))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3347, plain, (r1_nat_d(np__13,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_954, plain, (v3_ordinal1(np__10))).#trainneg
cnf(i_0_956, plain, (v3_ordinal1(np__8))).#trainneg
cnf(i_0_958, plain, (v3_ordinal1(np__9))).#trainneg
cnf(i_0_2544, plain, (r1_xxreal_0(np__156,np__156))).#trainneg
cnf(i_0_3393, plain, (r1_xxreal_0(X1,np__163)|r1_xxreal_0(np__163,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3350, plain, (r1_xxreal_0(k4_nat_1(np__13,np__13),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3352, plain, (r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3455, plain, (X1!=np__19|~r1_nat_d(np__11,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3311, plain, (k3_xcmplx_0(np__1,np__10)=np__10)).#trainneg
cnf(i_0_3362, plain, (r1_nat_d(np__13,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|v1_int_2(np__23)|v1_int_2(np__19)|v1_int_2(np__17)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3361, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__13)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|v1_int_2(np__23)|v1_int_2(np__19)|v1_int_2(np__17)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3346, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__13)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_960, plain, (v3_ordinal1(np__156))).#trainneg
cnf(i_0_962, plain, (v3_ordinal1(np__154))).#trainneg
cnf(i_0_964, plain, (v3_ordinal1(np__81))).#trainneg
cnf(i_0_2545, plain, (r1_xxreal_0(np__154,np__154))).#trainneg
cnf(i_0_3040, plain, (v3_xxreal_0(np__163)|v2_xxreal_0(esk11_0)|r1_xxreal_0(esk11_0,np__163))).#trainneg
cnf(i_0_3300, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__163,X1)|X2!=k2_nat_1(np__163,X1)|~r1_nat_d(np__163,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3254, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__5,X1)|X2!=k2_xcmplx_0(np__5,X1)|~r1_nat_d(np__5,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3452, plain, (r1_xxreal_0(np__5,np__153)|X1!=np__163|~r1_nat_d(np__5,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3251, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__7,X1)|X2!=k2_xcmplx_0(np__14,X1)|~r1_nat_d(np__7,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2876, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~m1_subset_1(X3,k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3242, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__23,X1)|X2!=k2_xcmplx_0(np__23,X1)|~r1_nat_d(np__23,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_966, plain, (v3_ordinal1(np__162))).#trainneg
cnf(i_0_968, plain, (v3_ordinal1(np__54))).#trainneg
cnf(i_0_970, plain, (v3_ordinal1(np__160))).#trainneg
cnf(i_0_2547, plain, (r1_xxreal_0(np__162,np__162))).#trainneg
cnf(i_0_3041, plain, (v3_xxreal_0(np__163)|v2_xxreal_0(esk9_0)|r1_xxreal_0(esk9_0,np__163))).#trainneg
cnf(i_0_3042, plain, (v3_xxreal_0(np__163)|v2_xxreal_0(esk7_0)|r1_xxreal_0(esk7_0,np__163))).#trainneg
cnf(i_0_884, plain, (v7_ordinal1(esk6_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_3257, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__3,X1)|X2!=k2_xcmplx_0(np__3,X1)|~r1_nat_d(np__3,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3261, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__2,X1)|X2!=k2_xcmplx_0(np__10,X1)|~r1_nat_d(np__2,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3453, plain, (X1!=np__11|~r1_nat_d(np__5,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3570, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__17,X1)|X2!=k2_xcmplx_0(np__17,X1)|~r1_nat_d(np__17,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3262, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__2,X1)|X2!=k2_xcmplx_0(np__2,X1)|~r1_nat_d(np__2,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3573, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__7,X1)|X2!=k2_xcmplx_0(np__7,X1)|~r1_nat_d(np__7,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3402, plain, (r1_nat_d(np__13,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(np__841,X1)|v1_int_2(np__17)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_972, plain, (v3_ordinal1(np__32))).#trainneg
cnf(i_0_974, plain, (v3_ordinal1(np__152))).#trainneg
cnf(i_0_976, plain, (v3_ordinal1(np__153))).#trainneg
cnf(i_0_1562, plain, (~m1_subset_1(k5_numbers,k5_numbers))).#trainneg
cnf(i_0_887, plain, (v7_ordinal1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2548, plain, (r1_xxreal_0(np__54,np__54))).#trainneg
cnf(i_0_2359, plain, (~m1_subset_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3055, plain, (v3_xxreal_0(np__163)|v2_xxreal_0(np__0)|r1_xxreal_0(np__0,np__163))).#trainneg
cnf(i_0_1561, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k4_numbers))).#trainneg
cnf(i_0_889, plain, (v1_xcmplx_0(np__14)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_890, plain, (v1_xcmplx_0(np__14)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_2737, plain, (r1_xxreal_0(k4_nat_1(np__23,np__23),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__17|esk1_1(X1)=np__13|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|v1_int_2(np__19)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3244, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__19,X1)|X2!=k2_xcmplx_0(np__19,X1)|~r1_nat_d(np__19,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3299, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__163,X1)|X2!=k2_xcmplx_0(np__163,X1)|~r1_nat_d(np__163,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2862, plain, (k3_xcmplx_0(np__5,X1)=k3_xcmplx_0(X1,np__5)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_978, plain, (v3_ordinal1(np__0))).#trainneg
cnf(i_0_981, plain, (v3_ordinal1(np__23))).#trainneg
cnf(i_0_983, plain, (v3_ordinal1(np__19))).#trainneg
cnf(i_0_891, plain, (v1_xcmplx_0(np__12)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_2550, plain, (r1_xxreal_0(np__160,np__160))).#trainneg
cnf(i_0_2454, plain, (v3_xxreal_0(X1)|v2_xxreal_0(np__0)|r1_xxreal_0(np__0,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2537, plain, (k2_xcmplx_0(X1,np__0)=X1|~v7_ordinal1(X1))).#trainneg
cnf(i_0_892, plain, (v1_xcmplx_0(np__12)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_3620, plain, (r1_nat_d(np__5,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__7|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3625, plain, (r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__7|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3623, plain, (r1_xxreal_0(k4_nat_1(np__5,np__5),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__7|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_985, plain, (v3_ordinal1(np__17))).#trainneg
cnf(i_0_987, plain, (v3_ordinal1(np__13))).#trainneg
cnf(i_0_989, plain, (v3_ordinal1(np__11))).#trainneg
cnf(i_0_3586, plain, (np__163!=np__11)).#trainneg
cnf(i_0_2551, plain, (r1_xxreal_0(np__32,np__32))).#trainneg
cnf(i_0_3629, plain, (np__19!=np__11)).#trainneg
cnf(i_0_3791, plain, (np__5!=np__11)).#trainneg
cnf(i_0_3819, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k4_numbers))).#trainneg
cnf(i_0_895, plain, (v1_xcmplx_0(np__10)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_896, plain, (v1_xcmplx_0(np__10)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_3651, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__5)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__7|esk1_1(X1)=np__11|r1_xxreal_0(np__841,X1)|v1_int_2(np__17)|v1_int_2(np__19)|v1_int_2(np__23)|v1_int_2(np__13)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3787, plain, (X1!=np__5|~r1_nat_d(np__3,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3652, plain, (r1_nat_d(np__5,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__7|esk1_1(X1)=np__11|r1_xxreal_0(np__841,X1)|v1_int_2(np__17)|v1_int_2(np__19)|v1_int_2(np__23)|v1_int_2(np__13)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3800, plain, (X1!=np__3|~r1_nat_d(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3790, plain, (X1!=np__11|~r1_nat_d(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3619, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__5)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__7|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_991, plain, (v3_ordinal1(np__7))).#trainneg
cnf(i_0_993, plain, (v3_ordinal1(np__5))).#trainneg
cnf(i_0_995, plain, (v3_ordinal1(np__3))).#trainneg
cnf(i_0_3812, plain, (~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_897, plain, (v1_xcmplx_0(np__8)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_2552, plain, (r1_xxreal_0(np__152,np__152))).#trainneg
cnf(i_0_1183, plain, (r1_xxreal_0(X1,esk7_0)|r1_xxreal_0(esk7_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1184, plain, (r1_xxreal_0(X1,esk6_0)|r1_xxreal_0(esk6_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_898, plain, (v1_xcmplx_0(np__8)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_899, plain, (v1_xcmplx_0(np__9)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_3666, plain, (r1_nat_d(np__5,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__7|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|v1_int_2(np__13)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2856, plain, (k3_xcmplx_0(np__7,X1)=k3_xcmplx_0(X1,np__7)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2468, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2853, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__7,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2847, plain, (r1_xxreal_0(X1,np__11)|r1_xxreal_0(np__11,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3513, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__163,X1)|~r1_nat_d(np__163,k2_nat_1(np__163,X1))|~v7_ordinal1(k2_nat_1(np__163,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3105, plain, (r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|r1_xxreal_0(np__841,X1)|v1_int_2(np__19)|v1_int_2(np__23)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3930, plain, (esk1_1(X1)=np__13|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__163|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3953, plain, (esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__163|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_997, plain, (v3_ordinal1(np__2))).#trainneg
cnf(i_0_999, plain, (v3_ordinal1(np__841))).#trainneg
cnf(i_0_1002, plain, (v3_ordinal1(np__1))).#trainneg
cnf(i_0_3952, plain, (np__163!=np__13)).#trainneg
cnf(i_0_900, plain, (v1_xcmplx_0(np__9)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_2553, plain, (r1_xxreal_0(np__153,np__153))).#trainneg
cnf(i_0_3929, plain, (np__17!=np__163)).#trainneg
cnf(i_0_3956, plain, (np__163!=np__5)).#trainneg
cnf(i_0_2316, plain, (r1_xxreal_0(X1,esk7_0)|r1_xxreal_0(esk7_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2317, plain, (r1_xxreal_0(X1,esk6_0)|r1_xxreal_0(esk6_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3957, plain, (esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__163|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3976, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__23,X1)|X2!=k2_xcmplx_0(X1,np__23)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__23,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3978, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__19,X1)|X2!=k2_xcmplx_0(X1,np__19)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__19,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3936, plain, (r1_nat_d(np__7,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3928, plain, (esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3954, plain, (esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(np__17)|v1_int_2(np__19)|v1_int_2(np__23)|v1_int_2(X1)|X1!=np__163|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4133, plain, (k3_xcmplx_0(np__7,np__1)=np__7)).#trainneg
cnf(i_0_1008, plain, (v1_xreal_0(np__14))).#trainneg
cnf(i_0_1010, plain, (v1_xreal_0(np__12))).#trainneg
cnf(i_0_1016, plain, (v1_xreal_0(np__8))).#trainneg
cnf(i_0_4018, plain, (np__163!=np__7)).#trainneg
cnf(i_0_4051, plain, (np__5!=np__3)).#trainneg
cnf(i_0_4058, plain, (np__3!=np__2)).#trainneg
cnf(i_0_2503, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_2504, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_905, plain, (v1_xcmplx_0(np__81)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_3931, plain, (esk1_1(X1)=np__13|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(np__23)|v1_int_2(np__19)|v1_int_2(X1)|X1!=np__163|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3988, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__5,X1)|X2!=k2_xcmplx_0(X1,np__5)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__5,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3939, plain, (r1_xxreal_0(k4_nat_1(np__7,np__7),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4057, plain, (esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(np__13)|v1_int_2(np__23)|v1_int_2(np__19)|v1_int_2(np__17)|v1_int_2(X1)|X1!=np__163|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3941, plain, (r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4250, plain, (r1_nat_d(np__13,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4363, plain, (r1_nat_d(np__7,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4383, plain, (r1_nat_d(np__7,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1018, plain, (v1_xreal_0(np__9))).#trainneg
cnf(i_0_1024, plain, (v1_xreal_0(np__81))).#trainneg
cnf(i_0_1028, plain, (v1_xreal_0(np__54))).#trainneg
cnf(i_0_4059, plain, (np__11!=np__2)).#trainneg
cnf(i_0_906, plain, (v1_xcmplx_0(np__81)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_2573, plain, (r1_xxreal_0(X1,esk11_0)|r1_xxreal_0(esk11_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2574, plain, (r1_xxreal_0(X1,esk9_0)|r1_xxreal_0(esk9_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4541, plain, (r1_nat_d(np__11,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4422, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__7,X1)|X2!=k2_xcmplx_0(X1,np__7)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__7,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4423, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__7,X1)|X2!=k2_nat_1(np__7,X1)|~r1_nat_d(np__7,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4042, plain, (r1_nat_d(np__7,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|v1_int_2(np__13)|v1_int_2(np__23)|v1_int_2(np__19)|v1_int_2(np__17)|v1_int_2(np__5)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4544, plain, (r1_xxreal_0(k4_nat_1(np__11,np__11),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3979, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__17,X1)|X2!=k2_xcmplx_0(X1,np__153)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__17,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1032, plain, (v1_xreal_0(np__32))).#trainneg
cnf(i_0_1034, plain, (v1_xreal_0(np__152))).#trainneg
cnf(i_0_3420, plain, (X1!=np__14|~r1_nat_d(np__13,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_909, plain, (v1_xcmplx_0(np__54)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_2576, plain, (r1_xxreal_0(X1,np__161)|r1_xxreal_0(np__161,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_910, plain, (v1_xcmplx_0(np__54)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_911, plain, (v1_xcmplx_0(np__160)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_4016, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__163,X1)|X2!=k2_xcmplx_0(X1,np__163)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__163,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4494, plain, (r1_nat_d(np__13,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|r1_xxreal_0(np__841,X1)|v1_int_2(np__19)|v1_int_2(np__23)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4386, plain, (r1_xxreal_0(k4_nat_1(np__7,np__7),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4546, plain, (r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4388, plain, (r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4555, plain, (r1_nat_d(np__5,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4574, plain, (r1_nat_d(np__11,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4809, plain, (r1_nat_d(np__11,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1041, plain, (v1_xreal_0(np__23))).#trainneg
cnf(i_0_1043, plain, (v1_xreal_0(np__19))).#trainneg
cnf(i_0_1045, plain, (v1_xreal_0(np__17))).#trainneg
cnf(i_0_4583, plain, (np__19!=np__163)).#trainneg
cnf(i_0_912, plain, (v1_xcmplx_0(np__160)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_4724, plain, (np__14!=np__13)).#trainneg
cnf(i_0_2577, plain, (r1_xxreal_0(X1,np__156)|r1_xxreal_0(np__156,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2578, plain, (r1_xxreal_0(X1,np__154)|r1_xxreal_0(np__154,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_913, plain, (v1_xcmplx_0(np__32)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_914, plain, (v1_xcmplx_0(np__32)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_4694, plain, (r1_xxreal_0(np__7,X1)|r1_xxreal_0(X1,np__0)|X2!=k2_nat_1(X1,np__7)|~r1_nat_d(np__7,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4772, plain, (r1_nat_d(np__5,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|v1_int_2(np__23)|v1_int_2(np__19)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1047, plain, (v1_xreal_0(np__13))).#trainneg
cnf(i_0_1049, plain, (v1_xreal_0(np__11))).#trainneg
cnf(i_0_1051, plain, (v1_xreal_0(np__7))).#trainneg
cnf(i_0_3457, plain, (X1!=np__12|~r1_nat_d(np__11,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_915, plain, (v1_xcmplx_0(np__152)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_2579, plain, (r1_xxreal_0(X1,np__81)|r1_xxreal_0(np__81,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2580, plain, (r1_xxreal_0(X1,np__162)|r1_xxreal_0(np__162,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_916, plain, (v1_xcmplx_0(np__152)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_917, plain, (v1_xcmplx_0(np__153)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_4913, plain, (r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|esk1_1(X1)=np__17|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4921, plain, (r1_nat_d(np__7,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4891, plain, (r1_xxreal_0(k4_nat_1(np__23,np__23),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|esk1_1(X1)=np__17|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4936, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5053, plain, (r1_nat_d(np__7,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__11|r1_xxreal_0(np__841,X1)|v1_int_2(np__19)|v1_int_2(np__23)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1053, plain, (v1_xreal_0(np__5))).#trainneg
cnf(i_0_1055, plain, (v1_xreal_0(np__3))).#trainneg
cnf(i_0_1057, plain, (v1_xreal_0(np__2))).#trainneg
cnf(i_0_5087, plain, (np__12!=np__11)).#trainneg
cnf(i_0_918, plain, (v1_xcmplx_0(np__153)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_2581, plain, (r1_xxreal_0(X1,np__54)|r1_xxreal_0(np__54,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2583, plain, (r1_xxreal_0(X1,np__160)|r1_xxreal_0(np__160,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4887, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__23)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|esk1_1(X1)=np__17|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4949, plain, (esk1_1(X1)=np__2|esk1_1(X1)=np__3|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__19|~v7_ordinal1(X1)|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5089, plain, (esk1_1(X1)=np__2|esk1_1(X1)=np__3|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__12|~v7_ordinal1(X1)|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5088, plain, (esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__12|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4939, plain, (r1_xxreal_0(k4_nat_1(np__3,np__3),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1059, plain, (v1_xreal_0(np__841))).#trainneg
cnf(i_0_1062, plain, (v1_xreal_0(np__1))).#trainneg
cnf(i_0_2360, plain, (v3_membered(k4_numbers))).#trainneg
cnf(i_0_3689, plain, (X1!=np__8|~r1_nat_d(np__5,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_922, plain, (v1_xcmplx_0(np__23)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_2584, plain, (r1_xxreal_0(X1,np__32)|r1_xxreal_0(np__32,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2585, plain, (r1_xxreal_0(X1,np__152)|r1_xxreal_0(np__152,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_923, plain, (v1_xcmplx_0(np__23)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_924, plain, (v1_xcmplx_0(np__19)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_5174, plain, (r1_nat_d(np__17,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__13|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5203, plain, (r1_nat_d(np__11,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5367, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5225, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__5|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2363, plain, (v1_xcmplx_0(np__160))).#trainneg
cnf(i_0_2675, plain, (v1_xcmplx_0(np__10))).#trainneg
cnf(i_0_5419, plain, (np__5!=np__8)).#trainneg
cnf(i_0_925, plain, (v1_xcmplx_0(np__19)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_2586, plain, (r1_xxreal_0(X1,np__153)|r1_xxreal_0(np__153,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_926, plain, (v1_xcmplx_0(np__17)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_927, plain, (v1_xcmplx_0(np__17)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_5389, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__12|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5378, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5339, plain, (r1_nat_d(np__17,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__13|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|v1_int_2(np__23)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5229, plain, (r1_nat_d(np__11,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|v1_int_2(np__23)|v1_int_2(np__19)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5420, plain, (esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__8|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4870, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__11,X1)|X2!=k2_xcmplx_0(X1,np__11)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__11,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4797, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4725, plain, (esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__14|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5502, plain, (r1_nat_d(np__13,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2679, plain, (v1_xcmplx_0(np__153))).#trainneg
cnf(i_0_3826, plain, (v1_xcmplx_0(np__14))).#trainneg
cnf(i_0_2698, plain, (v2_xxreal_0(esk7_0))).#trainneg
cnf(i_0_3804, plain, (X1!=np__8|~r1_nat_d(np__7,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_928, plain, (v1_xcmplx_0(np__13)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_3011, plain, (r1_xxreal_0(np__163,esk11_0)|r1_xxreal_0(esk11_0,np__163))).#trainneg
cnf(i_0_2589, plain, (r1_xxreal_0(X1,np__23)|r1_xxreal_0(np__23,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2684, plain, (r1_xxreal_0(X1,np__19)|r1_xxreal_0(np__19,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_929, plain, (v1_xcmplx_0(np__13)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_930, plain, (v1_xcmplx_0(np__11)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_5704, plain, (r1_nat_d(np__5,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5529, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5672, plain, (r1_nat_d(np__7,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5484, plain, (r1_xxreal_0(np__11,np__152))).#trainneg
cnf(i_0_5723, plain, (r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5547, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5707, plain, (r1_xxreal_0(k4_nat_1(np__5,np__5),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5392, plain, (r1_xxreal_0(k4_nat_1(np__3,np__3),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__12|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3880, plain, (v1_xcmplx_0(np__12))).#trainneg
cnf(i_0_4081, plain, (v1_xcmplx_0(np__8))).#trainneg
cnf(i_0_931, plain, (v1_xcmplx_0(np__11)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_3012, plain, (r1_xxreal_0(np__163,esk9_0)|r1_xxreal_0(esk9_0,np__163))).#trainneg
cnf(i_0_2689, plain, (r1_xxreal_0(X1,np__17)|r1_xxreal_0(np__17,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2694, plain, (r1_xxreal_0(X1,np__13)|r1_xxreal_0(np__13,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_932, plain, (v1_xcmplx_0(np__7)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_933, plain, (v1_xcmplx_0(np__7)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_5383, plain, (r1_xxreal_0(np__9,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5732, plain, (r1_nat_d(np__5,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5871, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__5|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5898, plain, (r1_nat_d(np__7,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5864, plain, (esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__7|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5923, plain, (r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__5|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5875, plain, (r1_nat_d(np__11,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4096, plain, (v1_xcmplx_0(np__9))).#trainneg
cnf(i_0_5780, plain, (np__7!=np__8)).#trainneg
cnf(i_0_934, plain, (v1_xcmplx_0(np__5)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_3013, plain, (r1_xxreal_0(np__163,esk7_0)|r1_xxreal_0(esk7_0,np__163))).#trainneg
cnf(i_0_2859, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__5,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2989, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__3,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_935, plain, (v1_xcmplx_0(np__5)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_936, plain, (v1_xcmplx_0(np__3)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_5866, plain, (esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__7|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__8|~v7_ordinal1(X1)|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5913, plain, (r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5887, plain, (r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5901, plain, (r1_xxreal_0(k4_nat_1(np__7,np__7),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5878, plain, (r1_xxreal_0(k4_nat_1(np__11,np__11),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4487, plain, (v1_xcmplx_0(np__81))).#trainneg
cnf(i_0_937, plain, (v1_xcmplx_0(np__3)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_4604, plain, (v1_xcmplx_0(np__54))).#trainneg
cnf(i_0_2996, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3003, plain, (r1_xxreal_0(X1,np__841)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_938, plain, (v1_xcmplx_0(np__2)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_939, plain, (v1_xcmplx_0(np__2)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_6052, plain, (r1_nat_d(np__7,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6214, plain, (r1_nat_d(np__7,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6092, plain, (r1_nat_d(np__7,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6119, plain, (esk1_1(X1)=np__2|esk1_1(X1)=np__3|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__19|~v7_ordinal1(X1)|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6120, plain, (esk1_1(X1)=np__2|esk1_1(X1)=np__3|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__12|~v7_ordinal1(X1)|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4710, plain, (v1_xcmplx_0(np__32))).#trainneg
cnf(i_0_4717, plain, (v1_xcmplx_0(np__152))).#trainneg
cnf(i_0_4951, plain, (v1_xcmplx_0(np__23))).#trainneg
cnf(i_0_940, plain, (v1_xcmplx_0(np__841)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_3594, plain, (r1_xxreal_0(np__163,esk6_0)|r1_xxreal_0(esk6_0,np__163))).#trainneg
cnf(i_0_4958, plain, (v1_xcmplx_0(np__19))).#trainneg
cnf(i_0_4965, plain, (v1_xcmplx_0(np__17))).#trainneg
cnf(i_0_3177, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3186, plain, (r1_xxreal_0(X1,np__9)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_941, plain, (v1_xcmplx_0(np__841)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_942, plain, (v1_xcmplx_0(k5_numbers)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6078, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6076, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6095, plain, (r1_xxreal_0(k4_nat_1(np__7,np__7),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6237, plain, (r1_nat_d(np__11,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6072, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__2)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6108, plain, (r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6073, plain, (r1_nat_d(np__2,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5066, plain, (v1_xcmplx_0(np__13))).#trainneg
cnf(i_0_5073, plain, (v1_xcmplx_0(np__11))).#trainneg
cnf(i_0_5080, plain, (v1_xcmplx_0(np__7))).#trainneg
cnf(i_0_943, plain, (v1_xcmplx_0(np__1)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_4086, plain, (r1_xxreal_0(esk7_0,esk6_0)|r1_xxreal_0(esk6_0,esk7_0))).#trainneg
cnf(i_0_5240, plain, (v1_xcmplx_0(np__5))).#trainneg
cnf(i_0_5247, plain, (v1_xcmplx_0(np__3))).#trainneg
cnf(i_0_3195, plain, (r1_xxreal_0(X1,np__8)|r1_xxreal_0(np__8,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3205, plain, (r1_xxreal_0(X1,esk7_0)|r1_xxreal_0(esk7_0,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_944, plain, (v1_xcmplx_0(np__1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_6381, plain, (r1_nat_d(np__11,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6647, plain, (r1_nat_d(np__11,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6672, plain, (r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__3|~v7_ordinal1(X1)|X1!=np__5|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6402, plain, (r1_nat_d(np__11,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6671, plain, (r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|X1!=np__5|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6609, plain, (esk1_1(X1)=np__2|esk1_1(X1)=np__3|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__19|~v7_ordinal1(X1)|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5254, plain, (v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_5405, plain, (v1_xcmplx_0(np__841))).#trainneg
cnf(i_0_5412, plain, (v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_2493, plain, (k3_xcmplx_0(X1,np__0)=np__0|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4087, plain, (r1_xxreal_0(esk7_0,np__0)|r1_xxreal_0(np__0,esk7_0))).#trainneg
cnf(i_0_947, plain, (v3_ordinal1(esk6_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_3206, plain, (r1_xxreal_0(X1,esk6_0)|r1_xxreal_0(esk6_0,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3207, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6810, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6843, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6825, plain, (esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__19|~v7_ordinal1(X1)|X1!=np__11|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6813, plain, (r1_xxreal_0(k4_nat_1(np__3,np__3),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6846, plain, (r1_xxreal_0(k4_nat_1(np__3,np__3),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6826, plain, (esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__12|~v7_ordinal1(X1)|X1!=np__11|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6835, plain, (esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__19|~v7_ordinal1(X1)|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1571, plain, (v3_valued_0(esk8_1(k1_numbers)))).#trainneg
cnf(i_0_2472, plain, (k3_xcmplx_0(esk7_0,X1)=k3_xcmplx_0(X1,esk7_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6836, plain, (esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__12|~v7_ordinal1(X1)|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6850, plain, (r1_xxreal_0(np__9,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6809, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__3)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6842, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__3)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6610, plain, (esk1_1(X1)=np__2|esk1_1(X1)=np__3|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__12|~v7_ordinal1(X1)|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2536, plain, (k2_xcmplx_0(np__163,np__0)=np__163)).#trainneg
cnf(i_0_2473, plain, (k3_xcmplx_0(np__161,X1)=k3_xcmplx_0(X1,np__161)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4092, plain, (r1_xxreal_0(esk6_0,np__0)|r1_xxreal_0(np__0,esk6_0))).#trainneg
cnf(i_0_3213, plain, (r1_xxreal_0(np__10,X1)|r1_xxreal_0(X1,np__10)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6416, plain, (r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6663, plain, (r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6817, plain, (r1_xxreal_0(np__9,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5394, plain, (r1_xxreal_0(np__9,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__12|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6405, plain, (r1_xxreal_0(k4_nat_1(np__11,np__11),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2474, plain, (k3_xcmplx_0(np__156,X1)=k3_xcmplx_0(X1,np__156)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4150, plain, (r1_xxreal_0(np__7,esk11_0)|r1_xxreal_0(esk11_0,np__7))).#trainneg
cnf(i_0_7003, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__5|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7008, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__5|~v7_ordinal1(X1)|X1!=np__19|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6444, plain, (r1_xxreal_0(k4_nat_1(np__3,np__3),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6431, plain, (r1_xxreal_0(k4_nat_1(np__3,np__3),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6650, plain, (r1_xxreal_0(k4_nat_1(np__11,np__11),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1544, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2475, plain, (k3_xcmplx_0(np__154,X1)=k3_xcmplx_0(X1,np__154)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4151, plain, (r1_xxreal_0(np__7,esk9_0)|r1_xxreal_0(esk9_0,np__7))).#trainneg
cnf(i_0_7101, plain, (r1_xxreal_0(k4_nat_1(np__3,np__3),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4868, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__17,X1)|X2!=k2_xcmplx_0(X1,np__17)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__17,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5381, plain, (r1_xxreal_0(k4_nat_1(np__3,np__3),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7054, plain, (r1_xxreal_0(k4_nat_1(np__3,np__3),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5370, plain, (r1_xxreal_0(k4_nat_1(np__3,np__3),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4686, plain, (esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__19|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2476, plain, (k3_xcmplx_0(np__81,X1)=k3_xcmplx_0(X1,np__81)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4152, plain, (r1_xxreal_0(np__7,esk7_0)|r1_xxreal_0(esk7_0,np__7))).#trainneg
cnf(i_0_7041, plain, (r1_xxreal_0(k4_nat_1(np__3,np__3),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__11|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5631, plain, (r1_nat_d(np__13,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|r1_xxreal_0(np__841,X1)|v1_int_2(np__23)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5762, plain, (k3_xcmplx_0(np__1,np__153)=np__153)).#trainneg
cnf(i_0_7344, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__19|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7334, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7342, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__19|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2477, plain, (k3_xcmplx_0(np__162,X1)=k3_xcmplx_0(X1,np__162)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4160, plain, (r1_xxreal_0(np__7,np__154)|r1_xxreal_0(np__154,np__7))).#trainneg
cnf(i_0_7379, plain, (v1_xboole_0(X1)|m2_subset_1(X2,X1,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_7332, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7354, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__9,X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__12|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7329, plain, (r1_nat_d(np__2,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2478, plain, (k3_xcmplx_0(np__54,X1)=k3_xcmplx_0(X1,np__54)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4161, plain, (r1_xxreal_0(np__7,np__81)|r1_xxreal_0(np__81,np__7))).#trainneg
cnf(i_0_7381, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~m2_subset_1(X1,X2,X2))).#trainneg
cnf(i_0_7328, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__2)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7574, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__153,X1)|X2!=k2_xcmplx_0(X1,np__153)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__153,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_7427, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__9,X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__12|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2480, plain, (k3_xcmplx_0(np__160,X1)=k3_xcmplx_0(X1,np__160)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4162, plain, (r1_xxreal_0(np__7,np__162)|r1_xxreal_0(np__162,np__7))).#trainneg
cnf(i_0_3317, plain, (r1_xxreal_0(X1,np__12)|r1_xxreal_0(np__12,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1060, plain, (v1_xreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1066, plain, (v1_int_1(np__14)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_7361, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__9,X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__19|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7455, plain, (r1_xxreal_0(k4_nat_1(np__3,np__3),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__19|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7338, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__2)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__19|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7339, plain, (r1_nat_d(np__2,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__19|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7727, plain, (r1_nat_d(np__2,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7653, plain, (r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__3|~v7_ordinal1(X1)|X1!=np__5|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7810, plain, (v1_int_1(np__14))).#trainneg
cnf(i_0_1941, plain, (v1_int_1(np__81)|~v1_int_1(np__9))).#trainneg
cnf(i_0_7725, plain, (X1!=np__163|~r1_nat_d(np__153,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4163, plain, (r1_xxreal_0(np__7,np__54)|r1_xxreal_0(np__54,np__7))).#trainneg
cnf(i_0_1068, plain, (v1_int_1(np__12)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1069, plain, (v1_int_1(np__12)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_3326, plain, (r1_xxreal_0(X1,np__14)|r1_xxreal_0(np__14,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1070, plain, (v1_int_1(np__161)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1071, plain, (v1_int_1(np__161)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_7652, plain, (r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|X1!=np__5|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7865, plain, (r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__3|~v7_ordinal1(X1)|X1!=np__5|X1!=np__19|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7812, plain, (r1_nat_d(np__2,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7732, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7864, plain, (r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|X1!=np__5|X1!=np__19|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7644, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1072, plain, (v1_int_1(np__10)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1966, plain, (v1_int_1(np__9)|~v1_int_1(np__3))).#trainneg
cnf(i_0_2047, plain, (v1_int_1(np__2)|~v1_int_1(np__1))).#trainneg
cnf(i_0_2481, plain, (k3_xcmplx_0(np__32,X1)=k3_xcmplx_0(X1,np__32)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1073, plain, (v1_int_1(np__10)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_4165, plain, (r1_xxreal_0(np__7,np__160)|r1_xxreal_0(np__160,np__7))).#trainneg
cnf(i_0_1074, plain, (v1_int_1(np__8)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1075, plain, (v1_int_1(np__8)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_3383, plain, (r1_xxreal_0(X1,np__161)|r1_xxreal_0(np__161,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1076, plain, (v1_int_1(np__9)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1077, plain, (v1_int_1(np__9)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_7817, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7827, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__9,X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__19|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7730, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7639, plain, (r1_nat_d(np__2,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7815, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7642, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7442, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__9,X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__19|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7624, plain, (k3_xcmplx_0(np__162,np__1)=np__162)).#trainneg
cnf(i_0_1078, plain, (v1_int_1(np__156)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_2442, plain, (v2_xxreal_0(esk6_0)|r1_xxreal_0(esk6_0,esk7_0))).#trainneg
cnf(i_0_2482, plain, (k3_xcmplx_0(np__152,X1)=k3_xcmplx_0(X1,np__152)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1079, plain, (v1_int_1(np__156)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_4166, plain, (r1_xxreal_0(np__7,np__32)|r1_xxreal_0(np__32,np__7))).#trainneg
cnf(i_0_1080, plain, (v1_int_1(np__154)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1081, plain, (v1_int_1(np__154)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_3384, plain, (r1_xxreal_0(X1,np__10)|r1_xxreal_0(np__10,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3385, plain, (r1_xxreal_0(X1,np__156)|r1_xxreal_0(np__156,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1082, plain, (v1_int_1(np__81)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1083, plain, (v1_int_1(np__81)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_8062, plain, (r1_nat_d(np__2,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7512, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__9,X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__12|X1!=np__11|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8065, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7881, plain, (v1_int_1(np__12))).#trainneg
cnf(i_0_7918, plain, (v1_int_1(np__161))).#trainneg
cnf(i_0_7941, plain, (v1_int_1(np__10))).#trainneg
cnf(i_0_7877, plain, (np__153!=np__163)).#trainneg
cnf(i_0_1084, plain, (v1_int_1(np__162)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_4167, plain, (r1_xxreal_0(np__7,np__152)|r1_xxreal_0(np__152,np__7))).#trainneg
cnf(i_0_7982, plain, (v1_int_1(np__8))).#trainneg
cnf(i_0_8019, plain, (v1_int_1(np__9))).#trainneg
cnf(i_0_3386, plain, (r1_xxreal_0(X1,np__154)|r1_xxreal_0(np__154,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3387, plain, (r1_xxreal_0(X1,np__162)|r1_xxreal_0(np__162,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1085, plain, (v1_int_1(np__162)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1086, plain, (v1_int_1(np__54)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_8025, plain, (r1_nat_d(np__2,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__19|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4815, plain, (r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5624, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__5|~v7_ordinal1(X1)|X1!=np__12|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7144, plain, (r1_xxreal_0(k4_nat_1(np__3,np__3),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8087, plain, (v1_int_1(np__156))).#trainneg
cnf(i_0_8127, plain, (v1_int_1(np__154))).#trainneg
cnf(i_0_3417, plain, (v3_xxreal_0(np__163)|r1_xxreal_0(np__0,np__163))).#trainneg
cnf(i_0_2483, plain, (k3_xcmplx_0(np__153,X1)=k3_xcmplx_0(X1,np__153)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1087, plain, (v1_int_1(np__54)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_4168, plain, (r1_xxreal_0(np__7,np__153)|r1_xxreal_0(np__153,np__7))).#trainneg
cnf(i_0_1088, plain, (v1_int_1(np__160)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1089, plain, (v1_int_1(np__160)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_3389, plain, (r1_xxreal_0(X1,np__160)|r1_xxreal_0(np__160,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3390, plain, (r1_xxreal_0(X1,np__153)|r1_xxreal_0(np__153,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1090, plain, (v1_int_1(np__32)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1091, plain, (v1_int_1(np__32)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_7141, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4874, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__2,X1)|X2!=k2_xcmplx_0(X1,np__2)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__2,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_8473, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__5|~v7_ordinal1(X1)|X1!=np__12|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8337, plain, (r1_nat_d(np__2,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__12|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8313, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8352, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__9,X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__12|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8199, plain, (v1_int_1(np__81))).#trainneg
cnf(i_0_8227, plain, (v1_int_1(np__162))).#trainneg
cnf(i_0_1092, plain, (v1_int_1(np__152)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_4185, plain, (r1_xxreal_0(np__11,esk11_0)|r1_xxreal_0(esk11_0,np__11))).#trainneg
cnf(i_0_8304, plain, (v1_int_1(np__54))).#trainneg
cnf(i_0_8395, plain, (v1_int_1(np__160))).#trainneg
cnf(i_0_1093, plain, (v1_int_1(np__152)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1094, plain, (v1_int_1(np__153)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_8208, plain, (r1_nat_d(np__2,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__11|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8310, plain, (r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__9,X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__3|~v7_ordinal1(X1)|X1!=np__19|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8309, plain, (r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__9,X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|X1!=np__19|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8316, plain, (r1_xxreal_0(k4_nat_1(np__3,np__3),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8476, plain, (r1_xxreal_0(np__2,np__161)|X1!=np__163|~r1_nat_d(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8467, plain, (v1_int_1(np__32))).#trainneg
cnf(i_0_1535, plain, (v7_ordinal1(esk2_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_3681, plain, (v2_xxreal_0(esk11_0)|r1_xxreal_0(esk11_0,np__163))).#trainneg
cnf(i_0_2528, plain, (k3_xcmplx_0(np__0,X1)=k3_xcmplx_0(X1,np__0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1095, plain, (v1_int_1(np__153)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_4186, plain, (r1_xxreal_0(np__11,esk9_0)|r1_xxreal_0(esk9_0,np__11))).#trainneg
cnf(i_0_1096, plain, (v1_int_1(np__0)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1097, plain, (v1_int_1(np__0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1099, plain, (v1_int_1(np__23)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1100, plain, (v1_int_1(np__23)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_8312, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__3)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8332, plain, (r1_xxreal_0(np__9,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8475, plain, (r1_xxreal_0(np__2,np__12)|X1!=np__14|~r1_nat_d(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8223, plain, (r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(np__9,X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__19|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1101, plain, (v1_int_1(np__19)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1536, plain, (v1_xcmplx_0(esk2_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_3777, plain, (v2_xxreal_0(esk9_0)|r1_xxreal_0(esk9_0,np__163))).#trainneg
cnf(i_0_2530, plain, (k3_xcmplx_0(np__23,X1)=k3_xcmplx_0(X1,np__23)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1102, plain, (v1_int_1(np__19)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_4187, plain, (r1_xxreal_0(np__11,esk7_0)|r1_xxreal_0(esk7_0,np__11))).#trainneg
cnf(i_0_1103, plain, (v1_int_1(np__17)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1104, plain, (v1_int_1(np__17)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1105, plain, (v1_int_1(np__13)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1106, plain, (v1_int_1(np__13)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_8206, plain, (r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__9,X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__3|~v7_ordinal1(X1)|X1!=np__19|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7931, plain, (r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__9,X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__3|~v7_ordinal1(X1)|X1!=np__19|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7265, plain, (r1_nat_d(np__2,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6441, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4873, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__3,X1)|X2!=k2_xcmplx_0(X1,np__3)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__3,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5366, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__3)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1107, plain, (v1_int_1(np__11)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1537, plain, (v3_ordinal1(esk2_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_3884, plain, (v2_xxreal_0(np__0)|r1_xxreal_0(np__0,esk7_0))).#trainneg
cnf(i_0_2687, plain, (k3_xcmplx_0(np__19,X1)=k3_xcmplx_0(X1,np__19)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1108, plain, (v1_int_1(np__11)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_4190, plain, (r1_xxreal_0(np__11,np__161)|r1_xxreal_0(np__161,np__11))).#trainneg
cnf(i_0_8512, plain, (v1_int_1(np__152))).#trainneg
cnf(i_0_8519, plain, (v1_int_1(np__153))).#trainneg
cnf(i_0_1109, plain, (v1_int_1(np__7)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1110, plain, (v1_int_1(np__7)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_8750, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__5|~v7_ordinal1(X1)|X1!=np__12|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8488, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__12|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6428, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8205, plain, (r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__9,X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|X1!=np__19|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4812, plain, (r1_xxreal_0(k4_nat_1(np__11,np__11),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7014, plain, (r1_xxreal_0(k4_nat_1(np__3,np__3),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__11|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7930, plain, (r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__9,X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|X1!=np__19|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8045, plain, (r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__9,X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|X1!=np__12|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8640, plain, (v1_int_1(np__0))).#trainneg
cnf(i_0_8642, plain, (v1_int_1(np__23))).#trainneg
cnf(i_0_8675, plain, (v1_int_1(np__19))).#trainneg
cnf(i_0_2692, plain, (k3_xcmplx_0(np__17,X1)=k3_xcmplx_0(X1,np__17)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1111, plain, (v1_int_1(np__5)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_4194, plain, (r1_xxreal_0(np__11,np__156)|r1_xxreal_0(np__156,np__11))).#trainneg
cnf(i_0_8728, plain, (v1_int_1(np__17))).#trainneg
cnf(i_0_8730, plain, (v1_int_1(np__13))).#trainneg
cnf(i_0_1112, plain, (v1_int_1(np__5)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1113, plain, (v1_int_1(np__3)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_4869, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__13,X1)|X2!=k2_xcmplx_0(X1,np__13)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__13,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5372, plain, (r1_xxreal_0(np__9,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6256, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7038, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__11|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8046, plain, (r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__9,X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__3|~v7_ordinal1(X1)|X1!=np__12|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7011, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__11|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5945, plain, (k3_xcmplx_0(np__1,np__12)=np__12)).#trainneg
cnf(i_0_7051, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5619, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__5|~v7_ordinal1(X1)|X1!=np__12|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8767, plain, (v1_int_1(np__11))).#trainneg
cnf(i_0_1538, plain, (v1_xreal_0(esk2_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_2697, plain, (k3_xcmplx_0(np__13,X1)=k3_xcmplx_0(X1,np__13)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1114, plain, (v1_int_1(np__3)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_4196, plain, (r1_xxreal_0(np__11,np__81)|r1_xxreal_0(np__81,np__11))).#trainneg
cnf(i_0_1115, plain, (v1_int_1(np__2)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1116, plain, (v1_int_1(np__2)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1117, plain, (v1_int_1(np__841)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1118, plain, (v1_int_1(np__841)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_5246, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__5,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8499, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__5|~v7_ordinal1(X1)|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7110, plain, (r1_nat_d(np__2,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7253, plain, (r1_xxreal_0(k4_nat_1(np__3,np__3),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5388, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__3)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__12|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7098, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5783, plain, (esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__11|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(np__5)|v1_int_2(np__17)|v1_int_2(np__19)|v1_int_2(np__23)|v1_int_2(np__13)|v1_int_2(X1)|X1!=np__8|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8811, plain, (v1_int_1(np__7))).#trainneg
cnf(i_0_8892, plain, (v1_int_1(np__5))).#trainneg
cnf(i_0_8899, plain, (v1_int_1(np__3))).#trainneg
cnf(i_0_2850, plain, (k3_xcmplx_0(np__11,X1)=k3_xcmplx_0(X1,np__11)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1119, plain, (v1_int_1(k5_numbers)|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4197, plain, (r1_xxreal_0(np__11,np__162)|r1_xxreal_0(np__162,np__11))).#trainneg
cnf(i_0_9018, plain, (v1_int_1(np__2))).#trainneg
cnf(i_0_9020, plain, (v1_int_1(np__841))).#trainneg
cnf(i_0_1120, plain, (v1_int_1(np__1)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1121, plain, (v1_int_1(np__1)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_8888, plain, (k3_xcmplx_0(np__17,np__1)=np__17)).#trainneg
cnf(i_0_8830, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__5|~v7_ordinal1(X1)|X1!=np__19|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8916, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__5|~v7_ordinal1(X1)|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5377, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__3)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8950, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__5|~v7_ordinal1(X1)|X1!=np__19|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7237, plain, (r1_xxreal_0(k4_nat_1(np__3,np__3),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7452, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__19|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1122, plain, (v1_int_1(np__163)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1539, plain, (v1_int_1(esk2_1(X1))|~v5_membered(X1))).#trainneg
cnf(i_0_2993, plain, (k3_xcmplx_0(np__3,X1)=k3_xcmplx_0(X1,np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1123, plain, (v1_int_1(np__163)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_4198, plain, (r1_xxreal_0(np__11,np__54)|r1_xxreal_0(np__54,np__11))).#trainneg
cnf(i_0_1125, plain, (v1_xxreal_0(np__14)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1126, plain, (v1_xxreal_0(np__14)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1127, plain, (v1_xxreal_0(np__12)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1128, plain, (v1_xxreal_0(np__12)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_5542, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__5|~v7_ordinal1(X1)|X1!=np__19|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9099, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__5|~v7_ordinal1(X1)|X1!=np__12|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7113, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4820, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__11,X1)|X2!=k2_xcmplx_0(X1,np__154)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__11,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4916, plain, (v3_xxreal_0(k4_nat_1(np__23,np__23))|v2_xxreal_0(X1)|v1_xboole_0(X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|esk1_1(X1)=np__17|r1_xxreal_0(np__841,X1)|~v1_xreal_0(k4_nat_1(np__23,np__23))|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1129, plain, (v1_xxreal_0(np__161)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1540, plain, (v1_xxreal_0(esk2_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_2863, plain, (v3_xxreal_0(esk6_0)|v2_xxreal_0(esk6_0)|v1_xboole_0(esk6_0))).#trainneg
cnf(i_0_3000, plain, (k3_xcmplx_0(np__2,X1)=k3_xcmplx_0(X1,np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1130, plain, (v1_xxreal_0(np__161)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_4200, plain, (r1_xxreal_0(np__11,np__160)|r1_xxreal_0(np__160,np__11))).#trainneg
cnf(i_0_1131, plain, (v1_xxreal_0(np__10)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1132, plain, (v1_xxreal_0(np__10)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1133, plain, (v1_xxreal_0(np__8)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1134, plain, (v1_xxreal_0(np__8)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_8387, plain, (k3_xcmplx_0(np__153,np__1)=np__153)).#trainneg
cnf(i_0_4821, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__13,X1)|X2!=k2_xcmplx_0(X1,np__156)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__13,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5253, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__3,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_9155, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__163|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9344, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__12|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9266, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__19|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9214, plain, (v1_int_1(np__1))).#trainneg
cnf(i_0_9285, plain, (v1_int_1(np__163))).#trainneg
cnf(i_0_9326, plain, (v1_xxreal_0(np__14))).#trainneg
cnf(i_0_9154, plain, (np__163!=np__3)).#trainneg
cnf(i_0_1135, plain, (v1_xxreal_0(np__9)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_4201, plain, (r1_xxreal_0(np__11,np__32)|r1_xxreal_0(np__32,np__11))).#trainneg
cnf(i_0_9328, plain, (v1_xxreal_0(np__12))).#trainneg
cnf(i_0_9365, plain, (v1_xxreal_0(np__161))).#trainneg
cnf(i_0_1136, plain, (v1_xxreal_0(np__9)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1137, plain, (v1_xxreal_0(np__156)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_9346, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__12|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9229, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__17,X1)|X2!=k2_nat_1(np__17,X1)|~r1_nat_d(np__17,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_9236, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__19|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9414, plain, (v1_xxreal_0(np__10))).#trainneg
cnf(i_0_9416, plain, (v1_xxreal_0(np__8))).#trainneg
cnf(i_0_3007, plain, (k3_xcmplx_0(np__841,X1)=k3_xcmplx_0(X1,np__841)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1138, plain, (v1_xxreal_0(np__156)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_4203, plain, (r1_xxreal_0(np__11,np__153)|r1_xxreal_0(np__153,np__11))).#trainneg
cnf(i_0_1139, plain, (v1_xxreal_0(np__154)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1140, plain, (v1_xxreal_0(np__154)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1141, plain, (v1_xxreal_0(np__81)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1142, plain, (v1_xxreal_0(np__81)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_9246, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9248, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9268, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__19|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1143, plain, (v1_xxreal_0(np__162)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1546, plain, (v1_xboole_0(esk10_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3014, plain, (r1_xxreal_0(np__161,np__163))).#trainneg
cnf(i_0_3183, plain, (k3_xcmplx_0(np__1,X1)=k3_xcmplx_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1144, plain, (v1_xxreal_0(np__162)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_4289, plain, (r1_xxreal_0(esk7_0,esk11_0)|r1_xxreal_0(esk11_0,esk7_0))).#trainneg
cnf(i_0_1145, plain, (v1_xxreal_0(np__54)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1146, plain, (v1_xxreal_0(np__54)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1147, plain, (v1_xxreal_0(np__160)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1148, plain, (v1_xxreal_0(np__160)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_9284, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__5|~v7_ordinal1(X1)|X1!=np__19|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9409, plain, (k3_xcmplx_0(np__2,X1)=k3_xcmplx_0(X1,np__2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_7278, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__12|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6250, plain, (r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4836, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__7,X1)|X2!=k2_xcmplx_0(X1,np__161)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__7,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6240, plain, (r1_xxreal_0(k4_nat_1(np__11,np__11),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7275, plain, (r1_nat_d(np__2,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__12|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1149, plain, (v1_xxreal_0(np__32)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1547, plain, (v1_xboole_0(esk8_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3015, plain, (r1_xxreal_0(np__156,np__163))).#trainneg
cnf(i_0_3192, plain, (k3_xcmplx_0(np__9,X1)=k3_xcmplx_0(X1,np__9)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1150, plain, (v1_xxreal_0(np__32)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_4290, plain, (r1_xxreal_0(esk7_0,esk9_0)|r1_xxreal_0(esk9_0,esk7_0))).#trainneg
cnf(i_0_9505, plain, (v1_xxreal_0(np__9))).#trainneg
cnf(i_0_9524, plain, (v1_xxreal_0(np__156))).#trainneg
cnf(i_0_9615, plain, (X1=esk10_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_9781, plain, (X1=esk8_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1151, plain, (v1_xxreal_0(np__152)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1152, plain, (v1_xxreal_0(np__152)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_9530, plain, (r1_xxreal_0(np__17,X1)|r1_xxreal_0(X1,np__0)|X2!=k2_nat_1(X1,np__17)|~r1_nat_d(np__17,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_9623, plain, (k3_xcmplx_0(np__14,np__1)=np__14)).#trainneg
cnf(i_0_9624, plain, (k3_xcmplx_0(np__12,np__1)=np__12)).#trainneg
cnf(i_0_9473, plain, (r1_nat_d(np__2,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9721, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9689, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(k3_xcmplx_0(np__2,X2),X1)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_9688, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X1,k3_xcmplx_0(np__2,X2))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_9592, plain, (v1_xxreal_0(np__154))).#trainneg
cnf(i_0_9594, plain, (v1_xxreal_0(np__81))).#trainneg
cnf(i_0_9609, plain, (v1_xxreal_0(np__162))).#trainneg
cnf(i_0_3201, plain, (k3_xcmplx_0(np__8,X1)=k3_xcmplx_0(X1,np__8)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1153, plain, (v1_xxreal_0(np__153)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_4292, plain, (r1_xxreal_0(esk7_0,np__14)|r1_xxreal_0(np__14,esk7_0))).#trainneg
cnf(i_0_9658, plain, (v1_xxreal_0(np__54))).#trainneg
cnf(i_0_9660, plain, (v1_xxreal_0(np__160))).#trainneg
cnf(i_0_1154, plain, (v1_xxreal_0(np__153)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_9628, plain, (k3_xcmplx_0(np__9,np__1)=np__9)).#trainneg
cnf(i_0_8030, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__19|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9362, plain, (esk1_1(X1)=np__17|esk1_1(X1)=np__13|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|v3_xxreal_0(k4_nat_1(np__23,np__23))|v2_xxreal_0(X1)|v1_xboole_0(X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~v7_ordinal1(k4_nat_1(np__23,np__23)))).#trainneg
cnf(i_0_5072, plain, (r1_xxreal_0(X1,np__13)|r1_xxreal_0(np__13,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8028, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__19|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9775, plain, (v1_xxreal_0(np__32))).#trainneg
cnf(i_0_8962, plain, (v1_xcmplx_0(esk2_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_3016, plain, (r1_xxreal_0(np__154,np__163))).#trainneg
cnf(i_0_3314, plain, (k3_xcmplx_0(np__10,X1)=k3_xcmplx_0(X1,np__10)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4293, plain, (r1_xxreal_0(esk7_0,np__12)|r1_xxreal_0(np__12,esk7_0))).#trainneg
cnf(i_0_1158, plain, (v1_xxreal_0(np__23)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1159, plain, (v1_xxreal_0(np__23)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1160, plain, (v1_xxreal_0(np__19)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1161, plain, (v1_xxreal_0(np__19)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_6091, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__7)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_10126, plain, (v1_xxreal_0(np__19))).#trainneg
cnf(i_0_10020, plain, (esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|esk1_1(X1)=np__17|v2_xxreal_0(X1)|v1_xboole_0(X1)|r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__841,X1)|v1_int_2(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_1(np__23,np__23),k5_numbers))).#trainneg
cnf(i_0_9917, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__7,X1)|X2!=k2_xcmplx_0(X1,np__14)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__7,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_9341, plain, (r1_nat_d(np__2,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__12|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9334, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__19|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9844, plain, (v1_xxreal_0(np__152))).#trainneg
cnf(i_0_9978, plain, (v1_xxreal_0(np__153))).#trainneg
cnf(i_0_10124, plain, (v1_xxreal_0(np__23))).#trainneg
cnf(i_0_9846, plain, (np__163!=np__2)).#trainneg
cnf(i_0_1162, plain, (v1_xxreal_0(np__17)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_4294, plain, (r1_xxreal_0(esk7_0,np__161)|r1_xxreal_0(np__161,esk7_0))).#trainneg
cnf(i_0_1163, plain, (v1_xxreal_0(np__17)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1164, plain, (v1_xxreal_0(np__13)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1165, plain, (v1_xxreal_0(np__13)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1166, plain, (v1_xxreal_0(np__11)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_5897, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__7)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7870, plain, (r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__9,X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|X1!=np__12|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4800, plain, (r1_xxreal_0(k4_nat_1(np__3,np__3),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6471, plain, (k3_xcmplx_0(np__1,np__152)=np__152)).#trainneg
cnf(i_0_5698, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__5|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9538, plain, (r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__5|X1!=np__19|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7871, plain, (r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__9,X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__3|~v7_ordinal1(X1)|X1!=np__12|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1167, plain, (v1_xxreal_0(np__11)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_9286, plain, (v1_xreal_0(esk2_1(X1))|~v5_membered(X1))).#trainneg
cnf(i_0_3017, plain, (r1_xxreal_0(np__81,np__163))).#trainneg
cnf(i_0_3323, plain, (k3_xcmplx_0(np__12,X1)=k3_xcmplx_0(X1,np__12)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1168, plain, (v1_xxreal_0(np__7)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_4295, plain, (r1_xxreal_0(esk7_0,np__10)|r1_xxreal_0(np__10,esk7_0))).#trainneg
cnf(i_0_1169, plain, (v1_xxreal_0(np__7)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1170, plain, (v1_xxreal_0(np__5)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1171, plain, (v1_xxreal_0(np__5)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1172, plain, (v1_xxreal_0(np__3)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_7280, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__12|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4849, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__2,X1)|X2!=k2_xcmplx_0(X1,np__10)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__2,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_8804, plain, (k3_xcmplx_0(np__19,X1)=k3_xcmplx_0(X1,np__19)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1173, plain, (v1_xxreal_0(np__3)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1573, plain, (v3_membered(esk2_1(k1_zfmisc_1(X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_3018, plain, (r1_xxreal_0(np__162,np__163))).#trainneg
cnf(i_0_3332, plain, (k3_xcmplx_0(np__14,X1)=k3_xcmplx_0(X1,np__14)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1174, plain, (v1_xxreal_0(np__2)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_4296, plain, (r1_xxreal_0(esk7_0,np__8)|r1_xxreal_0(np__8,esk7_0))).#trainneg
cnf(i_0_10233, plain, (v1_xxreal_0(np__17))).#trainneg
cnf(i_0_10238, plain, (v1_xxreal_0(np__13))).#trainneg
cnf(i_0_1175, plain, (v1_xxreal_0(np__2)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1176, plain, (v1_xxreal_0(np__841)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_10428, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(k3_xcmplx_0(np__19,X2),X1)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_10318, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__152,X1)|X2!=k2_xcmplx_0(X1,np__152)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__152,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10427, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X1,k3_xcmplx_0(np__19,X2))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_10195, plain, (esk1_1(X1)=np__17|esk1_1(X1)=np__13|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|v2_xxreal_0(X1)|v1_xboole_0(X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__163|~m1_subset_1(k3_xcmplx_0(np__23,np__23),k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10204, plain, (r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__3|~v7_ordinal1(X1)|X1!=np__5|X1!=np__12|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_10203, plain, (r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|X1!=np__5|X1!=np__12|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_10240, plain, (v1_xxreal_0(np__11))).#trainneg
cnf(i_0_10400, plain, (v1_xxreal_0(np__7))).#trainneg
cnf(i_0_10405, plain, (v1_xxreal_0(np__5))).#trainneg
cnf(i_0_2470, plain, (k3_xcmplx_0(esk11_0,X1)=k3_xcmplx_0(X1,esk11_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1177, plain, (v1_xxreal_0(np__841)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_4297, plain, (r1_xxreal_0(esk7_0,np__9)|r1_xxreal_0(np__9,esk7_0))).#trainneg
cnf(i_0_10407, plain, (v1_xxreal_0(np__3))).#trainneg
cnf(i_0_1179, plain, (v1_xxreal_0(np__1)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_4183, plain, (r1_xxreal_0(np__7,X1)|r1_xxreal_0(X1,np__7)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4218, plain, (r1_xxreal_0(np__11,X1)|r1_xxreal_0(X1,np__11)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1178, plain, (v1_xxreal_0(k5_numbers)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1180, plain, (v1_xxreal_0(np__1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_5560, plain, (k3_xcmplx_0(np__1,np__160)=np__160)).#trainneg
cnf(i_0_8061, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__2)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7115, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9702, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(k3_xcmplx_0(X2,np__2),X1)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_7250, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1181, plain, (v1_xxreal_0(np__163)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_10342, plain, (v1_xcmplx_0(esk2_1(X1))|~v5_membered(X1))).#trainneg
cnf(i_0_3019, plain, (r1_xxreal_0(np__54,np__163))).#trainneg
cnf(i_0_2471, plain, (k3_xcmplx_0(esk9_0,X1)=k3_xcmplx_0(X1,esk9_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1182, plain, (v1_xxreal_0(np__163)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_4298, plain, (r1_xxreal_0(esk7_0,np__156)|r1_xxreal_0(np__156,esk7_0))).#trainneg
cnf(i_0_10488, plain, (v1_xxreal_0(np__2))).#trainneg
cnf(i_0_10507, plain, (v1_xxreal_0(np__841))).#trainneg
cnf(i_0_10766, plain, (v1_xxreal_0(np__163))).#trainneg
cnf(i_0_4322, plain, (r1_xxreal_0(esk7_0,X1)|r1_xxreal_0(X1,esk7_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1553, plain, (k3_xcmplx_0(np__1,esk7_0)=esk7_0)).#trainneg
cnf(i_0_1554, plain, (k3_xcmplx_0(np__1,esk6_0)=esk6_0)).#trainneg
cnf(i_0_10533, plain, (r1_nat_d(np__17,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|v2_xxreal_0(X1)|v1_xboole_0(X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m1_subset_1(k3_xcmplx_0(np__23,np__23),k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_10686, plain, (v1_xxreal_0(np__1))).#trainneg
cnf(i_0_10518, plain, (X1!=np__163|~r1_nat_d(np__152,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10757, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(X1,k3_xcmplx_0(X2,np__2))|~m1_subset_1(X2,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(k3_xcmplx_0(X2,np__2)))).#trainneg
cnf(i_0_10524, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X1,k3_xcmplx_0(X2,np__19))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_9340, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__2)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__12|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_10424, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__19,X1)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(np__19,X3))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__19,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_1580, plain, (v1_relat_1(k5_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1576, plain, (v1_finset_1(esk2_1(k1_zfmisc_1(X1)))|~v1_finset_1(X1))).#trainneg
cnf(i_0_3021, plain, (r1_xxreal_0(np__160,np__163))).#trainneg
cnf(i_0_2486, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v7_ordinal1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1557, plain, (k3_xcmplx_0(np__1,X1)=X1|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4299, plain, (r1_xxreal_0(esk7_0,np__154)|r1_xxreal_0(np__154,esk7_0))).#trainneg
cnf(i_0_1583, plain, (v2_membered(k5_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_6587, plain, (v1_xcmplx_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4357, plain, (r1_xxreal_0(esk6_0,X1)|r1_xxreal_0(X1,esk6_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1559, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1563, plain, (v7_ordinal1(esk12_1(X1))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_6042, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6039, plain, (r1_nat_d(np__2,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(np__9,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_11086, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~r1_nat_d(X2,k2_nat_1(X2,X1))|~v7_ordinal1(k2_nat_1(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_11081, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(X2,X1)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_10972, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__7,X1)|X2!=k2_nat_1(X1,np__14)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__7,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10973, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__5,X1)|X2!=k2_nat_1(X1,np__10)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__5,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10980, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X1,k3_xcmplx_0(X2,np__2))|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(k3_xcmplx_0(X2,np__2))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_11079, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X2,X1)|~m1_subset_1(X2,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_11080, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X1,X2)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_10839, plain, (r1_xxreal_0(X1,np__163)|r1_xxreal_0(np__163,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_11090, plain, (v2_membered(k5_numbers))).#trainneg
cnf(i_0_10966, plain, (np__152!=np__163)).#trainneg
cnf(i_0_1564, plain, (v1_xcmplx_0(esk12_1(X1))|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4300, plain, (r1_xxreal_0(esk7_0,np__81)|r1_xxreal_0(np__81,esk7_0))).#trainneg
cnf(i_0_1565, plain, (v3_ordinal1(esk12_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1566, plain, (v1_xreal_0(esk12_1(X1))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_10883, plain, (r1_nat_d(np__13,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|v2_xxreal_0(X1)|v1_xboole_0(X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m1_subset_1(k3_xcmplx_0(np__23,np__23),k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_11138, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~r1_nat_d(X2,k2_nat_1(X1,X2))|~v7_ordinal1(k2_nat_1(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_11137, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~r1_nat_d(X2,k2_nat_1(X1,X2))|~v7_ordinal1(k2_nat_1(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_11191, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(X1,X2)|~m1_subset_1(X2,k1_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_11276, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_nat_d(X2,k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_11192, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(X1,X2)|~m1_subset_1(X2,k1_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_11195, plain, (r1_xxreal_0(np__7,X1)|r1_xxreal_0(X1,np__0)|X2!=k2_nat_1(np__14,X1)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__7,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7369, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1579, plain, (v1_relat_1(esk2_1(k1_zfmisc_1(X1)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_3022, plain, (r1_xxreal_0(np__32,np__163))).#trainneg
cnf(i_0_11218, plain, (X1!=np__163|~r1_nat_d(np__161,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1567, plain, (v1_int_1(esk12_1(X1))|v1_xboole_0(X1)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4301, plain, (r1_xxreal_0(esk7_0,np__162)|r1_xxreal_0(np__162,esk7_0))).#trainneg
cnf(i_0_10693, plain, (v1_xxreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1568, plain, (v1_xxreal_0(esk12_1(X1))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1572, plain, (v3_membered(esk12_1(X1))|v1_xboole_0(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_11199, plain, (r1_xxreal_0(np__5,X1)|r1_xxreal_0(X1,np__0)|X2!=k2_nat_1(np__10,X1)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__5,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7370, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1582, plain, (v2_membered(esk2_1(k1_zfmisc_1(X1)))|~v2_membered(X1))).#trainneg
cnf(i_0_3023, plain, (r1_xxreal_0(np__152,np__163))).#trainneg
cnf(i_0_11230, plain, (X1!=np__163|~r1_nat_d(np__156,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4302, plain, (r1_xxreal_0(esk7_0,np__54)|r1_xxreal_0(np__54,esk7_0))).#trainneg
cnf(i_0_1578, plain, (v1_relat_1(esk12_1(X1))|v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1581, plain, (v2_membered(esk12_1(X1))|v1_xboole_0(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_11220, plain, (X1!=np__11|~r1_nat_d(np__10,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11026, plain, (r1_xxreal_0(np__19,X1)|r1_xxreal_0(X1,np__0)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(X3,np__19))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X3,k5_numbers)|~r1_nat_d(np__19,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_11139, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~r1_nat_d(X2,k2_xcmplx_0(X2,X1))|~v7_ordinal1(k2_xcmplx_0(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_7371, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_11702, plain, (np__161!=np__163)).#trainneg
cnf(i_0_11725, plain, (np__156!=np__163)).#trainneg
cnf(i_0_1584, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk12_1(X2)))).#trainneg
cnf(i_0_4304, plain, (r1_xxreal_0(esk7_0,np__160)|r1_xxreal_0(np__160,esk7_0))).#trainneg
cnf(i_0_11583, plain, (~r1_nat_d(np__153,np__163))).#trainneg
cnf(i_0_11586, plain, (~r1_nat_d(np__11,np__19))).#trainneg
cnf(i_0_4432, plain, (r1_xxreal_0(X1,np__14)|r1_xxreal_0(np__14,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1585, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk2_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_1586, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_11691, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X2,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_11578, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_11470, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X2,k5_numbers)|~r1_nat_d(X2,k2_nat_1(X2,X1))|~v7_ordinal1(k2_nat_1(X2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11201, plain, (r1_xxreal_0(np__5,X1)|r1_xxreal_0(X1,np__0)|X2!=k2_xcmplx_0(X1,np__10)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__5,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_7372, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_11729, plain, (np__10!=np__11)).#trainneg
cnf(i_0_11591, plain, (~r1_nat_d(np__154,np__163))).#trainneg
cnf(i_0_4305, plain, (r1_xxreal_0(esk7_0,np__32)|r1_xxreal_0(np__32,esk7_0))).#trainneg
cnf(i_0_11609, plain, (~r1_nat_d(np__152,np__163))).#trainneg
cnf(i_0_11611, plain, (~r1_nat_d(np__156,np__163))).#trainneg
cnf(i_0_4439, plain, (r1_xxreal_0(X1,np__12)|r1_xxreal_0(np__12,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_11228, plain, (X1!=np__11|~r1_nat_d(np__9,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10018, plain, (esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|esk1_1(X1)=np__13|esk1_1(X1)=np__17|v2_xxreal_0(X1)|v1_xboole_0(X1)|r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__841,X1)|v1_int_2(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(k3_xcmplx_0(np__23,np__23))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11231, plain, (X1!=np__163|~r1_nat_d(np__154,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9321, plain, (k3_xcmplx_0(np__3,X1)=k3_xcmplx_0(X1,np__3)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_7373, plain, (v1_int_1(X1)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_12008, plain, (np__9!=np__11)).#trainneg
cnf(i_0_12096, plain, (np__154!=np__163)).#trainneg
cnf(i_0_4306, plain, (r1_xxreal_0(esk7_0,np__152)|r1_xxreal_0(np__152,esk7_0))).#trainneg
cnf(i_0_11614, plain, (~r1_nat_d(np__8,np__13))).#trainneg
cnf(i_0_11616, plain, (~r1_nat_d(np__8,np__11))).#trainneg
cnf(i_0_4446, plain, (r1_xxreal_0(X1,np__8)|r1_xxreal_0(np__8,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_11905, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(X2,k2_nat_1(X2,X1))|~v7_ordinal1(X2)|~v7_ordinal1(k2_nat_1(X2,X1)))).#trainneg
cnf(i_0_12115, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X1,k3_xcmplx_0(np__3,X2))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_12116, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(k3_xcmplx_0(np__3,X2),X1)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_12016, plain, (r1_nat_d(np__17,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__13|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|v2_xxreal_0(X1)|v1_xboole_0(X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(k3_xcmplx_0(np__23,np__23)))).#trainneg
cnf(i_0_12132, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(k3_xcmplx_0(X2,np__3),X1)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_12134, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__3,X1)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(X3,np__3))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__3,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_7374, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3024, plain, (r1_xxreal_0(np__153,np__163))).#trainneg
cnf(i_0_11617, plain, (~r1_nat_d(np__160,np__163))).#trainneg
cnf(i_0_4307, plain, (r1_xxreal_0(esk7_0,np__153)|r1_xxreal_0(np__153,esk7_0))).#trainneg
cnf(i_0_11619, plain, (~r1_nat_d(np__5,np__8))).#trainneg
cnf(i_0_11621, plain, (~r1_nat_d(np__12,np__14))).#trainneg
cnf(i_0_4485, plain, (r1_xxreal_0(np__0,X1)|r1_xxreal_0(X1,np__0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_12137, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~r1_nat_d(X2,k2_nat_1(k3_xcmplx_0(X2,np__3),X1))|~v7_ordinal1(k2_nat_1(k3_xcmplx_0(X2,np__3),X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_12244, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__3,X1)|X2!=k2_nat_1(np__9,X1)|~r1_nat_d(np__3,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_12259, plain, (r1_nat_d(np__13,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__5|esk1_1(X1)=np__7|esk1_1(X1)=np__11|v2_xxreal_0(X1)|v1_xboole_0(X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(k3_xcmplx_0(np__23,np__23)))).#trainneg
cnf(i_0_12352, plain, (r1_xxreal_0(np__3,X1)|r1_xxreal_0(X1,np__0)|X2!=k2_xcmplx_0(X1,np__9)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__3,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_11698, plain, (v7_ordinal1(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3026, plain, (r1_xxreal_0(np__0,np__163))).#trainneg
cnf(i_0_11622, plain, (~r1_nat_d(np__161,np__163))).#trainneg
cnf(i_0_11623, plain, (~r1_nat_d(np__8,np__10))).#trainneg
cnf(i_0_11624, plain, (~r1_nat_d(np__9,np__11))).#trainneg
cnf(i_0_4596, plain, (r1_xxreal_0(X1,np__9)|r1_xxreal_0(np__9,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_12376, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__3,X1)|~r1_nat_d(np__3,k2_nat_1(np__9,X1))|~v7_ordinal1(k2_nat_1(np__9,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_12241, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X1,k3_xcmplx_0(X2,np__3))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_12123, plain, (v7_ordinal1(k3_xcmplx_0(X1,np__3))|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_12382, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k5_numbers)|~r1_nat_d(X2,k2_nat_1(X1,k3_xcmplx_0(X2,np__3)))|~v7_ordinal1(k2_nat_1(X1,k3_xcmplx_0(X2,np__3)))|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(k3_xcmplx_0(X2,np__3)))).#trainneg
cnf(i_0_12665, plain, (v7_ordinal1(k3_xcmplx_0(np__3,np__19)))).#trainneg
cnf(i_0_3027, plain, (r1_xxreal_0(np__23,np__163))).#trainneg
cnf(i_0_11632, plain, (~r1_nat_d(np__162,np__163))).#trainneg
cnf(i_0_4310, plain, (r1_xxreal_0(esk7_0,np__23)|r1_xxreal_0(np__23,esk7_0))).#trainneg
cnf(i_0_11626, plain, (~r1_nat_d(np__3,np__5))).#trainneg
cnf(i_0_11628, plain, (~r1_nat_d(np__12,np__13))).#trainneg
cnf(i_0_4603, plain, (r1_xxreal_0(X1,np__81)|r1_xxreal_0(np__81,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_12582, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__19,X1)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(np__3,np__19))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__19,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_12516, plain, (X1!=np__11|~r1_nat_d(np__3,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_12437, plain, (r1_xxreal_0(np__3,X1)|r1_xxreal_0(X1,np__0)|X2!=k2_nat_1(X1,np__9)|~r1_nat_d(np__3,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_11400, plain, (r1_nat_d(np__5,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__7|esk1_1(X1)=np__11|v2_xxreal_0(X1)|v1_xboole_0(X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__163|~m1_subset_1(k3_xcmplx_0(np__23,np__23),k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9676, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~r1_nat_d(X2,k2_nat_1(k3_xcmplx_0(np__2,X2),X1))|~v7_ordinal1(k2_nat_1(k3_xcmplx_0(np__2,X2),X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_12884, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12887, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|X1!=np__12|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12874, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|X1!=np__19|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12886, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|X1!=np__12|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12876, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12873, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|X1!=np__12|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_11721, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_12803, plain, (r1_xxreal_0(k3_xcmplx_0(np__3,np__19),k3_xcmplx_0(np__3,np__19)))).#trainneg
cnf(i_0_12870, plain, (np__11!=np__3)).#trainneg
cnf(i_0_4311, plain, (r1_xxreal_0(esk7_0,np__19)|r1_xxreal_0(np__19,esk7_0))).#trainneg
cnf(i_0_11629, plain, (~r1_nat_d(np__10,np__11))).#trainneg
cnf(i_0_11630, plain, (~r1_nat_d(np__8,np__9))).#trainneg
cnf(i_0_4610, plain, (r1_xxreal_0(X1,np__54)|r1_xxreal_0(np__54,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_12880, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|X1!=np__12|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12885, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|X1!=np__19|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12872, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|X1!=np__19|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12882, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|X1!=np__19|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12877, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_11849, plain, (v3_ordinal1(X1)|~v7_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3028, plain, (r1_xxreal_0(np__19,np__163))).#trainneg
cnf(i_0_4312, plain, (r1_xxreal_0(esk7_0,np__17)|r1_xxreal_0(np__17,esk7_0))).#trainneg
cnf(i_0_11631, plain, (~r1_nat_d(np__9,np__10))).#trainneg
cnf(i_0_11634, plain, (~r1_nat_d(np__13,np__14))).#trainneg
cnf(i_0_4645, plain, (r1_xxreal_0(esk11_0,X1)|r1_xxreal_0(X1,esk11_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4680, plain, (r1_xxreal_0(esk9_0,X1)|r1_xxreal_0(X1,esk9_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5874, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__11)|v1_int_2(X1)|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8083, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__162,X1)|X2!=k2_xcmplx_0(X1,np__162)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__162,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_7097, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__3)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_11850, plain, (v3_ordinal1(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3029, plain, (r1_xxreal_0(np__17,np__163))).#trainneg
cnf(i_0_4313, plain, (r1_xxreal_0(esk7_0,np__13)|r1_xxreal_0(np__13,esk7_0))).#trainneg
cnf(i_0_11635, plain, (~r1_nat_d(np__11,np__12))).#trainneg
cnf(i_0_11636, plain, (~r1_nat_d(np__7,np__8))).#trainneg
cnf(i_0_4716, plain, (r1_xxreal_0(X1,np__32)|r1_xxreal_0(np__32,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_11233, plain, (X1!=np__163|~r1_nat_d(np__160,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_12579, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~r1_nat_d(X2,k2_xcmplx_0(X1,k3_xcmplx_0(X2,np__3)))|~v7_ordinal1(k2_xcmplx_0(X1,k3_xcmplx_0(X2,np__3))))).#trainneg
cnf(i_0_13236, plain, (X1!=np__163|~r1_nat_d(np__162,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13140, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__19|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13142, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__19|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13162, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__19|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13044, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__19|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13046, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__19|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13160, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__19|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_11966, plain, (v1_xreal_0(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_13177, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_13302, plain, (np__160!=np__163)).#trainneg
cnf(i_0_11637, plain, (~r1_nat_d(np__2,np__3))).#trainneg
cnf(i_0_4723, plain, (r1_xxreal_0(X1,np__152)|r1_xxreal_0(np__152,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_13034, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__12|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13026, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13054, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__12|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13152, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__19|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13024, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13036, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__12|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13305, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__19,X1)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__19,k2_xcmplx_0(X1,k3_xcmplx_0(np__3,np__19)))|~v7_ordinal1(k2_xcmplx_0(X1,k3_xcmplx_0(np__3,np__19))))).#trainneg
cnf(i_0_13150, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__19|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13056, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__12|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12140, plain, (v1_int_1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3030, plain, (r1_xxreal_0(np__13,np__163))).#trainneg
cnf(i_0_4762, plain, (r1_xxreal_0(np__161,X1)|r1_xxreal_0(X1,np__161)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_13074, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__12|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13076, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__12|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13170, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13130, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__12|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13172, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4229, plain, (r1_nat_d(np__13,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(np__841,X1)|v1_int_2(np__23)|v1_int_2(np__19)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12359, plain, (v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3031, plain, (r1_xxreal_0(np__11,np__163))).#trainneg
cnf(i_0_4316, plain, (r1_xxreal_0(esk7_0,np__5)|r1_xxreal_0(np__5,esk7_0))).#trainneg
cnf(i_0_4957, plain, (r1_xxreal_0(X1,np__23)|r1_xxreal_0(np__23,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6873, plain, (k3_xcmplx_0(np__1,np__841)=np__841)).#trainneg
cnf(i_0_6259, plain, (r1_xxreal_0(k4_nat_1(np__3,np__3),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__8|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4382, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__7)|v1_int_2(X1)|esk1_1(X1)=np__2|esk1_1(X1)=np__3|esk1_1(X1)=np__11|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9013, plain, (k3_xcmplx_0(np__13,X1)=k3_xcmplx_0(X1,np__13)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_13081, plain, (v1_xcmplx_0(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_13687, plain, (v7_ordinal1(k3_xcmplx_0(np__3,np__13)))).#trainneg
cnf(i_0_13357, plain, (np__162!=np__163)).#trainneg
cnf(i_0_4317, plain, (r1_xxreal_0(esk7_0,np__3)|r1_xxreal_0(np__3,esk7_0))).#trainneg
cnf(i_0_4964, plain, (r1_xxreal_0(X1,np__19)|r1_xxreal_0(np__19,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_13678, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(k3_xcmplx_0(np__13,X2),X1)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_13662, plain, (r1_xxreal_0(np__19,X1)|r1_xxreal_0(X1,np__0)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(np__13,np__19))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__19,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13663, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__13,X1)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(np__13,np__19))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__13,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13677, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X1,k3_xcmplx_0(np__13,X2))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_13586, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__841,X1)|X2!=k2_xcmplx_0(X1,np__841)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__841,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13563, plain, (r1_xxreal_0(np__19,np__23))).#trainneg
cnf(i_0_13585, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__841,X1)|X2!=k2_xcmplx_0(np__841,X1)|~r1_nat_d(np__841,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13587, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__841,X1)|X2!=k2_nat_1(np__841,X1)|~r1_nat_d(np__841,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3032, plain, (r1_xxreal_0(np__7,np__163))).#trainneg
cnf(i_0_4318, plain, (r1_xxreal_0(esk7_0,np__2)|r1_xxreal_0(np__2,esk7_0))).#trainneg
cnf(i_0_4971, plain, (r1_xxreal_0(X1,np__17)|r1_xxreal_0(np__17,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5006, plain, (r1_xxreal_0(np__156,X1)|r1_xxreal_0(X1,np__156)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_13509, plain, (r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|X1!=np__12|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13700, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X1,k3_xcmplx_0(X2,np__13))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_13674, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__13,X1)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(np__13,X3))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__13,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_13479, plain, (v1_int_1(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3033, plain, (r1_xxreal_0(np__5,np__163))).#trainneg
cnf(i_0_4319, plain, (r1_xxreal_0(esk7_0,np__841)|r1_xxreal_0(np__841,esk7_0))).#trainneg
cnf(i_0_5041, plain, (r1_xxreal_0(np__154,X1)|r1_xxreal_0(X1,np__154)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7234, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2839, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_13532, plain, (v1_xxreal_0(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_13781, plain, (r1_xxreal_0(k3_xcmplx_0(np__3,np__13),k3_xcmplx_0(np__3,np__13)))).#trainneg
cnf(i_0_13708, plain, (np__23!=np__163)).#trainneg
cnf(i_0_5079, plain, (r1_xxreal_0(X1,np__11)|r1_xxreal_0(np__11,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_13854, plain, (r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__0)|X2!=k2_nat_1(X1,np__841)|~r1_nat_d(np__841,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_13953, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|~v7_ordinal1(X1)|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13855, plain, (r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__0)|X2!=k2_nat_1(X1,np__841)|~r1_nat_d(np__841,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3034, plain, (r1_xxreal_0(np__3,np__163))).#trainneg
cnf(i_0_4324, plain, (r1_xxreal_0(esk6_0,esk11_0)|r1_xxreal_0(esk11_0,esk6_0))).#trainneg
cnf(i_0_5086, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__7,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5124, plain, (r1_xxreal_0(np__81,X1)|r1_xxreal_0(X1,np__81)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_13952, plain, (esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__5|~v7_ordinal1(X1)|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3036, plain, (r1_xxreal_0(np__163,np__841))).#trainneg
cnf(i_0_4325, plain, (r1_xxreal_0(esk6_0,esk9_0)|r1_xxreal_0(esk9_0,esk6_0))).#trainneg
cnf(i_0_5159, plain, (r1_xxreal_0(np__162,X1)|r1_xxreal_0(X1,np__162)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_13978, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(k3_xcmplx_0(X2,np__19),X1)|~m1_subset_1(X2,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_13832, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(k3_xcmplx_0(X2,np__13),X1)|~m1_subset_1(X2,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_13997, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X1,k3_xcmplx_0(X2,np__2))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_3179, plain, (r1_xxreal_0(np__1,np__163))).#trainneg
cnf(i_0_5260, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5295, plain, (r1_xxreal_0(np__54,X1)|r1_xxreal_0(X1,np__54)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3188, plain, (r1_xxreal_0(np__9,np__163))).#trainneg
cnf(i_0_4327, plain, (r1_xxreal_0(esk6_0,np__14)|r1_xxreal_0(np__14,esk6_0))).#trainneg
cnf(i_0_5330, plain, (r1_xxreal_0(np__160,X1)|r1_xxreal_0(X1,np__160)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_13833, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(X1,k3_xcmplx_0(np__13,X2))|~m1_subset_1(X2,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(k3_xcmplx_0(np__13,X2)))).#trainneg
cnf(i_0_3197, plain, (r1_xxreal_0(np__8,np__163))).#trainneg
cnf(i_0_11216, plain, (X1!=np__14|~r1_nat_d(np__12,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1759, plain, (k3_xcmplx_0(np__1,k2_xcmplx_0(X1,X2))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4328, plain, (r1_xxreal_0(esk6_0,np__12)|r1_xxreal_0(np__12,esk6_0))).#trainneg
cnf(i_0_5411, plain, (r1_xxreal_0(X1,np__841)|r1_xxreal_0(np__841,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_14242, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__841,X1)|X2!=k2_nat_1(np__841,X1)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__841,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_14356, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__13,X1)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(np__2,np__13))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__13,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_14423, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__3,X1)|X2!=k2_nat_1(X1,k3_xcmplx_0(np__3,np__13))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__3,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_14360, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__2,X1)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(np__2,np__2))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__2,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_14245, plain, (r1_nat_d(np__2,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_14007, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k5_numbers)|~r1_nat_d(X2,k2_nat_1(k3_xcmplx_0(np__3,X2),X1))|~v7_ordinal1(k2_nat_1(k3_xcmplx_0(np__3,X2),X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3310, plain, (r1_xxreal_0(np__10,np__163))).#trainneg
cnf(i_0_11217, plain, (X1!=np__13|~r1_nat_d(np__12,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4329, plain, (r1_xxreal_0(esk6_0,np__161)|r1_xxreal_0(np__161,esk6_0))).#trainneg
cnf(i_0_5418, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5457, plain, (r1_xxreal_0(np__32,X1)|r1_xxreal_0(X1,np__32)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_14611, plain, (r1_xxreal_0(np__3,X1)|r1_xxreal_0(X1,np__0)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(np__3,np__13))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__3,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_14602, plain, (r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__0)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__841,k2_nat_1(np__841,X1))|~v7_ordinal1(k2_nat_1(np__841,X1)))).#trainneg
cnf(i_0_14306, plain, (r1_nat_d(np__2,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13834, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(X1,k3_xcmplx_0(np__13,X2))|~m1_subset_1(X2,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(k3_xcmplx_0(np__13,X2),k5_numbers))).#trainneg
cnf(i_0_14522, plain, (np__12!=np__14)).#trainneg
cnf(i_0_14680, plain, (np__12!=np__13)).#trainneg
cnf(i_0_4330, plain, (r1_xxreal_0(esk6_0,np__10)|r1_xxreal_0(np__10,esk6_0))).#trainneg
cnf(i_0_5492, plain, (r1_xxreal_0(np__152,X1)|r1_xxreal_0(X1,np__152)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5611, plain, (r1_xxreal_0(np__153,X1)|r1_xxreal_0(X1,np__153)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_8799, plain, (k3_xcmplx_0(np__19,np__2)=k3_xcmplx_0(np__2,np__19))).#trainneg
cnf(i_0_14000, plain, (k3_xcmplx_0(X1,np__1)=X1|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_8723, plain, (k3_xcmplx_0(np__23,X1)=k3_xcmplx_0(X1,np__23)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5658, plain, (r1_nat_d(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|v1_int_2(np__19)|v1_int_2(np__23)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3319, plain, (r1_xxreal_0(np__12,np__163))).#trainneg
cnf(i_0_11223, plain, (X1!=np__13|~r1_nat_d(np__8,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4331, plain, (r1_xxreal_0(esk6_0,np__8)|r1_xxreal_0(np__8,esk6_0))).#trainneg
cnf(i_0_5821, plain, (r1_xxreal_0(np__23,X1)|r1_xxreal_0(X1,np__23)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5856, plain, (r1_xxreal_0(np__19,X1)|r1_xxreal_0(X1,np__19)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3555, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__3,X1)|X2!=k2_xcmplx_0(np__9,X1)|~r1_nat_d(np__3,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3365, plain, (r1_xxreal_0(k4_nat_1(np__13,np__13),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__11|esk1_1(X1)=np__7|esk1_1(X1)=np__5|esk1_1(X1)=np__3|esk1_1(X1)=np__2|r1_xxreal_0(np__841,X1)|v1_int_2(np__23)|v1_int_2(np__19)|v1_int_2(np__17)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13909, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(k3_xcmplx_0(X2,X3),X1)|X4!=k2_xcmplx_0(X1,k3_xcmplx_0(X2,k3_xcmplx_0(X3,np__13)))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(k3_xcmplx_0(X2,X3),k5_numbers)|~r1_nat_d(k3_xcmplx_0(X2,X3),X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_14960, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__2)|v1_int_2(X1)|r1_xxreal_0(np__841,X1)|v1_int_2(np__23)|v1_int_2(np__19)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_14851, plain, (v7_ordinal1(k3_xcmplx_0(np__2,np__19)))).#trainneg
cnf(i_0_14973, plain, (np__13!=np__8)).#trainneg
cnf(i_0_14855, plain, (~v1_xboole_0(k3_xcmplx_0(np__2,np__19)))).#trainneg
cnf(i_0_4332, plain, (r1_xxreal_0(esk6_0,np__9)|r1_xxreal_0(np__9,esk6_0))).#trainneg
cnf(i_0_5999, plain, (r1_xxreal_0(np__17,X1)|r1_xxreal_0(X1,np__17)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6034, plain, (r1_xxreal_0(np__13,X1)|r1_xxreal_0(X1,np__13)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_14922, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__23,X1)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(np__23,X3))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__23,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_14926, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(k3_xcmplx_0(np__23,X2),X1)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_14940, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__23,X1)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(np__2,np__23))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__23,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_14961, plain, (r1_nat_d(np__2,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(np__841,X1)|v1_int_2(np__23)|v1_int_2(np__19)|~v7_ordinal1(X1)|X1!=np__163|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_14925, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X1,k3_xcmplx_0(np__23,X2))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_14952, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(k3_xcmplx_0(X2,np__23),X1)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_15198, plain, (r1_xxreal_0(np__23,X1)|r1_xxreal_0(X1,np__0)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(np__13,np__23))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__23,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_14852, plain, (v1_xcmplx_0(k3_xcmplx_0(np__2,np__19)))).#trainneg
cnf(i_0_15255, plain, (v1_int_2(np__19)|v1_int_2(np__23))).#trainneg
cnf(i_0_11224, plain, (X1!=np__11|~r1_nat_d(np__8,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4333, plain, (r1_xxreal_0(esk6_0,np__156)|r1_xxreal_0(np__156,esk6_0))).#trainneg
cnf(i_0_6171, plain, (r1_xxreal_0(np__5,X1)|r1_xxreal_0(X1,np__5)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6206, plain, (r1_xxreal_0(np__3,X1)|r1_xxreal_0(X1,np__3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_15202, plain, (r1_xxreal_0(np__23,X1)|r1_xxreal_0(X1,np__0)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(X3,np__23))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X3,k5_numbers)|~r1_nat_d(np__23,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_15199, plain, (r1_xxreal_0(np__23,X1)|r1_xxreal_0(X1,np__0)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(np__3,np__23))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__23,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_15260, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__2,X1)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(np__2,np__23))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__2,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_14853, plain, (v1_xreal_0(k3_xcmplx_0(np__2,np__19)))).#trainneg
cnf(i_0_3328, plain, (r1_xxreal_0(np__14,np__163))).#trainneg
cnf(i_0_11225, plain, (X1!=np__10|~r1_nat_d(np__8,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4334, plain, (r1_xxreal_0(esk6_0,np__154)|r1_xxreal_0(np__154,esk6_0))).#trainneg
cnf(i_0_6339, plain, (r1_xxreal_0(np__2,X1)|r1_xxreal_0(X1,np__2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6374, plain, (r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__841)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_15258, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__13,X1)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(np__13,np__23))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__13,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_15262, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(X1,k3_xcmplx_0(X2,np__23))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_15279, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(X1,k3_xcmplx_0(X2,np__23))|~m1_subset_1(X2,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(k3_xcmplx_0(X2,np__23)))).#trainneg
cnf(i_0_15257, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__23,X1)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(np__23,np__23))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__23,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_15259, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__3,X1)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(np__3,np__23))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__3,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_15220, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(X1,k3_xcmplx_0(np__23,X2))|~m1_subset_1(X2,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(k3_xcmplx_0(np__23,X2),k5_numbers))).#trainneg
cnf(i_0_15281, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_xcmplx_0(k3_xcmplx_0(X2,np__23),X1)|~m1_subset_1(X2,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(k3_xcmplx_0(X2,np__23),k5_numbers))).#trainneg
cnf(i_0_14854, plain, (v1_int_1(k3_xcmplx_0(np__2,np__19)))).#trainneg
cnf(i_0_11226, plain, (X1!=np__9|~r1_nat_d(np__8,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1820, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4335, plain, (r1_xxreal_0(esk6_0,np__81)|r1_xxreal_0(np__81,esk6_0))).#trainneg
cnf(i_0_6549, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6584, plain, (r1_xxreal_0(np__9,X1)|r1_xxreal_0(X1,np__9)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1821, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k3_xcmplx_0(X2,X3))|r1_xxreal_0(k3_xcmplx_0(X2,X3),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_15474, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(X1,k3_xcmplx_0(X2,np__23))|~m1_subset_1(k3_xcmplx_0(X2,np__23),k5_numbers)|~m1_subset_1(X2,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_15268, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k5_numbers)|~r1_nat_d(X2,k2_nat_1(k3_xcmplx_0(X2,np__23),X1))|~v7_ordinal1(X1)|~v7_ordinal1(k2_nat_1(k3_xcmplx_0(X2,np__23),X1)))).#trainneg
cnf(i_0_14937, plain, (v7_ordinal1(k3_xcmplx_0(np__3,np__23)))).#trainneg
cnf(i_0_15357, plain, (np__11!=np__8)).#trainneg
cnf(i_0_15422, plain, (np__10!=np__8)).#trainneg
cnf(i_0_4336, plain, (r1_xxreal_0(esk6_0,np__162)|r1_xxreal_0(np__162,esk6_0))).#trainneg
cnf(i_0_15495, plain, (np__9!=np__8)).#trainneg
cnf(i_0_6770, plain, (r1_xxreal_0(np__8,X1)|r1_xxreal_0(X1,np__8)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_15219, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(X1,k3_xcmplx_0(np__23,X2))|~m1_subset_1(X2,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(k3_xcmplx_0(np__23,X2)))).#trainneg
cnf(i_0_14862, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__19,X1)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(np__2,np__19))|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__19,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13064, plain, (r1_xxreal_0(k4_nat_1(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__11|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5772, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),np__14)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__14))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13688, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__13,X1)|X2!=k2_nat_1(k3_xcmplx_0(np__2,np__13),X1)|~r1_nat_d(np__13,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_7050, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__3)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__19|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_15183, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__19),k3_xcmplx_0(np__2,np__19)))).#trainneg
cnf(i_0_11229, plain, (X1!=np__10|~r1_nat_d(np__9,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4337, plain, (r1_xxreal_0(esk6_0,np__54)|r1_xxreal_0(np__54,esk6_0))).#trainneg
cnf(i_0_15692, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__7,X1)|X2!=k2_nat_1(X1,np__161)|~r1_nat_d(np__7,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_15857, plain, (r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(np__841,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|X1!=np__11|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15843, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__7,X1)|X2!=k2_nat_1(X1,np__161)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__7,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_7324, plain, (k3_xcmplx_0(np__156,X1)=k3_xcmplx_0(X1,np__156)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4339, plain, (r1_xxreal_0(esk6_0,np__160)|r1_xxreal_0(np__160,esk6_0))).#trainneg
cnf(i_0_7230, plain, (r1_xxreal_0(X1,np__10)|r1_xxreal_0(np__10,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_9668, plain, (r1_xxreal_0(k3_xcmplx_0(np__2,np__2),X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__13,np__13),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__5|X1!=np__19|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7244, plain, (r1_xxreal_0(np__9,X1)|r1_xxreal_0(X1,np__1)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__11,np__11),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__8|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7428, plain, (r1_xxreal_0(np__17,X1)|r1_xxreal_0(X1,np__0)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(np__17,k2_xcmplx_0(X1,np__17))|~v7_ordinal1(k2_xcmplx_0(X1,np__17)))).#trainneg
cnf(i_0_7974, plain, (k3_xcmplx_0(np__32,np__1)=np__32)).#trainneg
cnf(i_0_14014, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|X3!=k2_nat_1(X1,k3_xcmplx_0(np__2,X2))|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~r1_nat_d(X2,X3)|~v7_ordinal1(k3_xcmplx_0(np__2,X2))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_5955, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),np__8)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__8))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_15825, plain, (r1_xxreal_0(k3_xcmplx_0(np__3,np__23),k3_xcmplx_0(np__3,np__23)))).#trainneg
cnf(i_0_4340, plain, (r1_xxreal_0(esk6_0,np__32)|r1_xxreal_0(np__32,esk6_0))).#trainneg
cnf(i_0_7806, plain, (r1_xxreal_0(np__12,X1)|r1_xxreal_0(X1,np__12)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7916, plain, (r1_xxreal_0(np__14,X1)|r1_xxreal_0(X1,np__14)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1874, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1875, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k2_xcmplx_0(X2,X3))|r1_xxreal_0(k2_xcmplx_0(X2,X3),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_13923, plain, (r1_xxreal_0(np__13,X1)|r1_xxreal_0(X1,np__0)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(X3,np__13))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X3,k5_numbers)|~r1_nat_d(np__13,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_7140, plain, (r1_xxreal_0(X1,np__1)|v1_int_2(np__3)|v1_int_2(X1)|esk1_1(X1)=np__2|r1_xxreal_0(k3_xcmplx_0(np__5,np__5),X1)|r1_xxreal_0(k3_xcmplx_0(np__23,np__23),X1)|r1_xxreal_0(k3_xcmplx_0(np__19,np__19),X1)|r1_xxreal_0(k3_xcmplx_0(np__17,np__17),X1)|r1_xxreal_0(k3_xcmplx_0(np__7,np__7),X1)|r1_xxreal_0(np__841,X1)|~v7_ordinal1(X1)|X1!=np__12|X1!=np__14|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5760, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),np__153)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__153))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 617
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 771
# Removed in clause preprocessing      : 26
# Initial clauses in saturation        : 745
# Processed clauses                    : 3202
# ...of these trivial                  : 291
# ...subsumed                          : 822
# ...remaining for further processing  : 2089
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 137
# Backward-rewritten                   : 177
# Generated clauses                    : 15514
# ...of the previous two non-trivial   : 13074
# Contextual simplify-reflections      : 90
# Paramodulations                      : 15098
# Factorizations                       : 277
# Equation resolutions                 : 139
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1775
#    Positive orientable unit clauses  : 641
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 298
#    Non-unit-clauses                  : 836
# Current number of unprocessed clauses: 10188
# ...number of literals in the above   : 81750
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2403
# Proof object given clauses           : 139
# Proof search given clauses           : 2089
# Clause-clause subsumption calls (NU) : 866920
# Rec. Clause-clause subsumption calls : 50384
# Non-unit clause-clause subsumptions  : 858
# Unit Clause-clause subsumption calls : 76700
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 91
# BW rewrite match successes           : 85
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 578593

# -------------------------------------------------
# User time                : 37.363 s
# System time              : 0.812 s
# Total time               : 38.175 s
# Maximum resident set size: 380792 pages
37.36user 0.88system 0:40.37elapsed 94%CPU (0avgtext+0avgdata 380792maxresident)k
0inputs+664outputs (0major+247115minor)pagefaults 0swaps
