# ENIGMA: Connected to the TF server '127.0.0.1:8888'.
# Preprocessing time       : 0.022 s
#TF#SERVER: Sending data to the TF server (conj=7, context=98, query=90, total=195).
#TF#SERVER: Sent 97624 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (756 bytes expected).
#TF#SERVER: Received 756 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=154, query=128, total=289).
#TF#SERVER: Sent 211296 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1132 bytes expected).
#TF#SERVER: Received 1132 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=186, query=140, total=333).
#TF#SERVER: Sent 249660 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1308 bytes expected).
#TF#SERVER: Received 1308 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=197, query=161, total=365).
#TF#SERVER: Sent 239014 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1436 bytes expected).
#TF#SERVER: Received 1436 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=210, query=164, total=381).
#TF#SERVER: Sent 279253 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1500 bytes expected).
#TF#SERVER: Received 1500 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=231, query=129, total=367).
#TF#SERVER: Sent 251156 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1444 bytes expected).
#TF#SERVER: Received 1444 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=235, query=137, total=379).
#TF#SERVER: Sent 338096 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1492 bytes expected).
#TF#SERVER: Received 1492 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=245, query=134, total=386).
#TF#SERVER: Sent 314219 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1520 bytes expected).
#TF#SERVER: Received 1520 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=252, query=128, total=387).
#TF#SERVER: Sent 325630 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1524 bytes expected).
#TF#SERVER: Received 1524 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=252, query=65, total=324).
#TF#SERVER: Sent 237139 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1272 bytes expected).
#TF#SERVER: Received 1272 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=258, query=164, total=429).
#TF#SERVER: Sent 298837 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1692 bytes expected).
#TF#SERVER: Received 1692 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=263, query=157, total=427).
#TF#SERVER: Sent 397664 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1684 bytes expected).
#TF#SERVER: Received 1684 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=272, query=143, total=422).
#TF#SERVER: Sent 350382 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1664 bytes expected).
#TF#SERVER: Received 1664 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=278, query=182, total=467).
#TF#SERVER: Sent 435376 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1844 bytes expected).
#TF#SERVER: Received 1844 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=287, query=172, total=466).
#TF#SERVER: Sent 339041 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1840 bytes expected).
#TF#SERVER: Received 1840 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=291, query=128, total=426).
#TF#SERVER: Sent 342026 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1680 bytes expected).
#TF#SERVER: Received 1680 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=291, query=73, total=371).
#TF#SERVER: Sent 289390 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1460 bytes expected).
#TF#SERVER: Received 1460 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=295, query=128, total=430).
#TF#SERVER: Sent 243704 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1696 bytes expected).
#TF#SERVER: Received 1696 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=295, query=65, total=367).
#TF#SERVER: Sent 244683 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1444 bytes expected).
#TF#SERVER: Received 1444 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=299, query=128, total=434).
#TF#SERVER: Sent 410491 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1712 bytes expected).
#TF#SERVER: Received 1712 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=299, query=89, total=395).
#TF#SERVER: Sent 332932 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1556 bytes expected).
#TF#SERVER: Received 1556 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=302, query=178, total=487).
#TF#SERVER: Sent 457660 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1924 bytes expected).
#TF#SERVER: Received 1924 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=307, query=170, total=484).
#TF#SERVER: Sent 338398 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1912 bytes expected).
#TF#SERVER: Received 1912 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=309, query=128, total=444).
#TF#SERVER: Sent 399587 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1752 bytes expected).
#TF#SERVER: Received 1752 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=309, query=123, total=439).
#TF#SERVER: Sent 382009 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1732 bytes expected).
#TF#SERVER: Received 1732 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=311, query=136, total=454).
#TF#SERVER: Sent 354411 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1792 bytes expected).
#TF#SERVER: Received 1792 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=315, query=128, total=450).
#TF#SERVER: Sent 383840 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1776 bytes expected).
#TF#SERVER: Received 1776 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=315, query=115, total=437).
#TF#SERVER: Sent 392173 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1724 bytes expected).
#TF#SERVER: Received 1724 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=317, query=128, total=452).
#TF#SERVER: Sent 348512 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1784 bytes expected).
#TF#SERVER: Received 1784 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=317, query=128, total=452).
#TF#SERVER: Sent 401423 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1784 bytes expected).
#TF#SERVER: Received 1784 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=317, query=44, total=368).
#TF#SERVER: Sent 251321 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1448 bytes expected).
#TF#SERVER: Received 1448 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=320, query=136, total=463).
#TF#SERVER: Sent 354540 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1828 bytes expected).
#TF#SERVER: Received 1828 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=324, query=149, total=480).
#TF#SERVER: Sent 429378 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1896 bytes expected).
#TF#SERVER: Received 1896 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=332, query=128, total=467).
#TF#SERVER: Sent 358773 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1844 bytes expected).
#TF#SERVER: Received 1844 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=332, query=77, total=416).
#TF#SERVER: Sent 305299 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1640 bytes expected).
#TF#SERVER: Received 1640 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=333, query=128, total=468).
#TF#SERVER: Sent 463993 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1848 bytes expected).
#TF#SERVER: Received 1848 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=333, query=75, total=415).
#TF#SERVER: Sent 345984 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1636 bytes expected).
#TF#SERVER: Received 1636 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=337, query=157, total=501).
#TF#SERVER: Sent 438081 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1980 bytes expected).
#TF#SERVER: Received 1980 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=343, query=138, total=488).
#TF#SERVER: Sent 436506 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1928 bytes expected).
#TF#SERVER: Received 1928 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=350, query=140, total=497).
#TF#SERVER: Sent 383821 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1964 bytes expected).
#TF#SERVER: Received 1964 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=358, query=128, total=493).
#TF#SERVER: Sent 430792 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1948 bytes expected).
#TF#SERVER: Received 1948 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=358, query=80, total=445).
#TF#SERVER: Sent 380530 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1756 bytes expected).
#TF#SERVER: Received 1756 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=359, query=149, total=515).
#TF#SERVER: Sent 524306 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2036 bytes expected).
#TF#SERVER: Received 2036 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=371, query=161, total=539).
#TF#SERVER: Sent 401779 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2132 bytes expected).
#TF#SERVER: Received 2132 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=377, query=128, total=512).
#TF#SERVER: Sent 426679 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2024 bytes expected).
#TF#SERVER: Received 2024 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=377, query=128, total=512).
#TF#SERVER: Sent 487386 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2024 bytes expected).
#TF#SERVER: Received 2024 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=377, query=92, total=476).
#TF#SERVER: Sent 413360 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1880 bytes expected).
#TF#SERVER: Received 1880 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=380, query=151, total=538).
#TF#SERVER: Sent 422375 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2128 bytes expected).
#TF#SERVER: Received 2128 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=383, query=179, total=569).
#TF#SERVER: Sent 466864 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2252 bytes expected).
#TF#SERVER: Received 2252 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=388, query=128, total=523).
#TF#SERVER: Sent 455120 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2068 bytes expected).
#TF#SERVER: Received 2068 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=388, query=91, total=486).
#TF#SERVER: Sent 407759 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1920 bytes expected).
#TF#SERVER: Received 1920 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=390, query=128, total=525).
#TF#SERVER: Sent 425719 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2076 bytes expected).
#TF#SERVER: Received 2076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=390, query=78, total=475).
#TF#SERVER: Sent 337465 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1876 bytes expected).
#TF#SERVER: Received 1876 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=394, query=128, total=529).
#TF#SERVER: Sent 450960 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2092 bytes expected).
#TF#SERVER: Received 2092 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=394, query=128, total=529).
#TF#SERVER: Sent 499083 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2092 bytes expected).
#TF#SERVER: Received 2092 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=394, query=14, total=415).
#TF#SERVER: Sent 270582 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1636 bytes expected).
#TF#SERVER: Received 1636 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=396, query=128, total=531).
#TF#SERVER: Sent 426333 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2100 bytes expected).
#TF#SERVER: Received 2100 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=396, query=91, total=494).
#TF#SERVER: Sent 361070 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1952 bytes expected).
#TF#SERVER: Received 1952 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=399, query=160, total=566).
#TF#SERVER: Sent 478638 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2240 bytes expected).
#TF#SERVER: Received 2240 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=403, query=128, total=538).
#TF#SERVER: Sent 505242 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2128 bytes expected).
#TF#SERVER: Received 2128 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=403, query=128, total=538).
#TF#SERVER: Sent 514626 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2128 bytes expected).
#TF#SERVER: Received 2128 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=403, query=22, total=432).
#TF#SERVER: Sent 298990 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1704 bytes expected).
#TF#SERVER: Received 1704 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=404, query=184, total=595).
#TF#SERVER: Sent 561868 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2356 bytes expected).
#TF#SERVER: Received 2356 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=409, query=128, total=544).
#TF#SERVER: Sent 420713 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2152 bytes expected).
#TF#SERVER: Received 2152 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=409, query=100, total=516).
#TF#SERVER: Sent 421234 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2040 bytes expected).
#TF#SERVER: Received 2040 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=413, query=175, total=595).
#TF#SERVER: Sent 417064 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2356 bytes expected).
#TF#SERVER: Received 2356 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=419, query=190, total=616).
#TF#SERVER: Sent 472899 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2440 bytes expected).
#TF#SERVER: Received 2440 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=425, query=132, total=564).
#TF#SERVER: Sent 395727 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2232 bytes expected).
#TF#SERVER: Received 2232 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=426, query=176, total=609).
#TF#SERVER: Sent 604161 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2412 bytes expected).
#TF#SERVER: Received 2412 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=429, query=128, total=564).
#TF#SERVER: Sent 371086 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2232 bytes expected).
#TF#SERVER: Received 2232 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=429, query=85, total=521).
#TF#SERVER: Sent 343853 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2060 bytes expected).
#TF#SERVER: Received 2060 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=432, query=136, total=575).
#TF#SERVER: Sent 441930 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2276 bytes expected).
#TF#SERVER: Received 2276 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=436, query=175, total=618).
#TF#SERVER: Sent 550080 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2448 bytes expected).
#TF#SERVER: Received 2448 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=440, query=128, total=575).
#TF#SERVER: Sent 452938 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2276 bytes expected).
#TF#SERVER: Received 2276 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=440, query=128, total=575).
#TF#SERVER: Sent 538431 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2276 bytes expected).
#TF#SERVER: Received 2276 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=440, query=48, total=495).
#TF#SERVER: Sent 375019 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1956 bytes expected).
#TF#SERVER: Received 1956 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=441, query=128, total=576).
#TF#SERVER: Sent 529104 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2280 bytes expected).
#TF#SERVER: Received 2280 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=441, query=76, total=524).
#TF#SERVER: Sent 413512 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2072 bytes expected).
#TF#SERVER: Received 2072 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=444, query=152, total=603).
#TF#SERVER: Sent 498669 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2388 bytes expected).
#TF#SERVER: Received 2388 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=446, query=128, total=581).
#TF#SERVER: Sent 530934 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2300 bytes expected).
#TF#SERVER: Received 2300 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=446, query=128, total=581).
#TF#SERVER: Sent 497008 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2300 bytes expected).
#TF#SERVER: Received 2300 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=446, query=14, total=467).
#TF#SERVER: Sent 314896 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1844 bytes expected).
#TF#SERVER: Received 1844 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=447, query=128, total=582).
#TF#SERVER: Sent 552892 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2304 bytes expected).
#TF#SERVER: Received 2304 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=447, query=128, total=582).
#TF#SERVER: Sent 524844 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2304 bytes expected).
#TF#SERVER: Received 2304 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=447, query=88, total=542).
#TF#SERVER: Sent 463199 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2144 bytes expected).
#TF#SERVER: Received 2144 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=450, query=128, total=585).
#TF#SERVER: Sent 538527 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2316 bytes expected).
#TF#SERVER: Received 2316 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=450, query=107, total=564).
#TF#SERVER: Sent 493681 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2232 bytes expected).
#TF#SERVER: Received 2232 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=451, query=137, total=595).
#TF#SERVER: Sent 498868 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2356 bytes expected).
#TF#SERVER: Received 2356 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=457, query=128, total=592).
#TF#SERVER: Sent 527843 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2344 bytes expected).
#TF#SERVER: Received 2344 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=457, query=128, total=592).
#TF#SERVER: Sent 512938 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2344 bytes expected).
#TF#SERVER: Received 2344 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=457, query=46, total=510).
#TF#SERVER: Sent 382638 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2016 bytes expected).
#TF#SERVER: Received 2016 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=461, query=128, total=596).
#TF#SERVER: Sent 522824 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2360 bytes expected).
#TF#SERVER: Received 2360 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=461, query=128, total=596).
#TF#SERVER: Sent 558080 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2360 bytes expected).
#TF#SERVER: Received 2360 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=461, query=6, total=474).
#TF#SERVER: Sent 321125 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1872 bytes expected).
#TF#SERVER: Received 1872 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=464, query=157, total=628).
#TF#SERVER: Sent 542858 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2488 bytes expected).
#TF#SERVER: Received 2488 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=467, query=128, total=602).
#TF#SERVER: Sent 532256 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2384 bytes expected).
#TF#SERVER: Received 2384 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=467, query=128, total=602).
#TF#SERVER: Sent 556476 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2384 bytes expected).
#TF#SERVER: Received 2384 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=467, query=29, total=503).
#TF#SERVER: Sent 365565 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1988 bytes expected).
#TF#SERVER: Received 1988 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=473, query=147, total=627).
#TF#SERVER: Sent 529190 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2484 bytes expected).
#TF#SERVER: Received 2484 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=479, query=135, total=621).
#TF#SERVER: Sent 549312 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2460 bytes expected).
#TF#SERVER: Received 2460 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=484, query=128, total=619).
#TF#SERVER: Sent 422958 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2452 bytes expected).
#TF#SERVER: Received 2452 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=484, query=128, total=619).
#TF#SERVER: Sent 556237 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2452 bytes expected).
#TF#SERVER: Received 2452 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=484, query=128, total=619).
#TF#SERVER: Sent 556456 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2452 bytes expected).
#TF#SERVER: Received 2452 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=484, query=18, total=509).
#TF#SERVER: Sent 356176 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2012 bytes expected).
#TF#SERVER: Received 2012 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=491, query=145, total=643).
#TF#SERVER: Sent 585293 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2548 bytes expected).
#TF#SERVER: Received 2548 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=495, query=128, total=630).
#TF#SERVER: Sent 538112 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2496 bytes expected).
#TF#SERVER: Received 2496 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=495, query=128, total=630).
#TF#SERVER: Sent 595235 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2496 bytes expected).
#TF#SERVER: Received 2496 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=495, query=7, total=509).
#TF#SERVER: Sent 358690 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2012 bytes expected).
#TF#SERVER: Received 2012 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=496, query=128, total=631).
#TF#SERVER: Sent 588993 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2500 bytes expected).
#TF#SERVER: Received 2500 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=496, query=70, total=573).
#TF#SERVER: Sent 476155 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2268 bytes expected).
#TF#SERVER: Received 2268 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=498, query=128, total=633).
#TF#SERVER: Sent 577871 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2508 bytes expected).
#TF#SERVER: Received 2508 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=498, query=94, total=599).
#TF#SERVER: Sent 495153 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2372 bytes expected).
#TF#SERVER: Received 2372 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=500, query=128, total=635).
#TF#SERVER: Sent 585416 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2516 bytes expected).
#TF#SERVER: Received 2516 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=500, query=128, total=635).
#TF#SERVER: Sent 576972 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2516 bytes expected).
#TF#SERVER: Received 2516 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=500, query=20, total=527).
#TF#SERVER: Sent 373878 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2084 bytes expected).
#TF#SERVER: Received 2084 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=501, query=135, total=643).
#TF#SERVER: Sent 572002 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2548 bytes expected).
#TF#SERVER: Received 2548 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=504, query=128, total=639).
#TF#SERVER: Sent 566478 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2532 bytes expected).
#TF#SERVER: Received 2532 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=504, query=128, total=639).
#TF#SERVER: Sent 601310 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2532 bytes expected).
#TF#SERVER: Received 2532 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=504, query=86, total=597).
#TF#SERVER: Sent 501415 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2364 bytes expected).
#TF#SERVER: Received 2364 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=506, query=128, total=641).
#TF#SERVER: Sent 574024 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2540 bytes expected).
#TF#SERVER: Received 2540 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=506, query=104, total=617).
#TF#SERVER: Sent 528168 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2444 bytes expected).
#TF#SERVER: Received 2444 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=511, query=128, total=646).
#TF#SERVER: Sent 512407 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2560 bytes expected).
#TF#SERVER: Received 2560 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=511, query=128, total=646).
#TF#SERVER: Sent 579195 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2560 bytes expected).
#TF#SERVER: Received 2560 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=511, query=19, total=537).
#TF#SERVER: Sent 388249 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2124 bytes expected).
#TF#SERVER: Received 2124 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=170, total=689).
#TF#SERVER: Sent 623072 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2732 bytes expected).
#TF#SERVER: Received 2732 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 500521 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=87, total=606).
#TF#SERVER: Sent 482892 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2400 bytes expected).
#TF#SERVER: Received 2400 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=148, total=667).
#TF#SERVER: Sent 585815 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2644 bytes expected).
#TF#SERVER: Received 2644 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 599425 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 617979 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 625248 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=59, total=578).
#TF#SERVER: Sent 468308 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2288 bytes expected).
#TF#SERVER: Received 2288 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 621958 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 623129 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 609888 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=30, total=549).
#TF#SERVER: Sent 416238 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2172 bytes expected).
#TF#SERVER: Received 2172 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=166, total=685).
#TF#SERVER: Sent 488026 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2716 bytes expected).
#TF#SERVER: Received 2716 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 595952 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 595936 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=16, total=535).
#TF#SERVER: Sent 389625 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2116 bytes expected).
#TF#SERVER: Received 2116 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 577920 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 604126 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=59, total=578).
#TF#SERVER: Sent 453672 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2288 bytes expected).
#TF#SERVER: Received 2288 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 509575 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 594870 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 616040 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=115, total=634).
#TF#SERVER: Sent 551959 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2512 bytes expected).
#TF#SERVER: Received 2512 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=139, total=658).
#TF#SERVER: Sent 589617 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2608 bytes expected).
#TF#SERVER: Received 2608 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 513458 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 525880 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=9, total=528).
#TF#SERVER: Sent 370272 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2088 bytes expected).
#TF#SERVER: Received 2088 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=162, total=681).
#TF#SERVER: Sent 498982 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2700 bytes expected).
#TF#SERVER: Received 2700 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 595801 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=77, total=596).
#TF#SERVER: Sent 507003 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2360 bytes expected).
#TF#SERVER: Received 2360 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=163, total=682).
#TF#SERVER: Sent 573274 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2704 bytes expected).
#TF#SERVER: Received 2704 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 452319 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 590080 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 610370 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=109, total=628).
#TF#SERVER: Sent 559852 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2488 bytes expected).
#TF#SERVER: Received 2488 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=164, total=683).
#TF#SERVER: Sent 589789 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2708 bytes expected).
#TF#SERVER: Received 2708 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 446624 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 456880 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=2, total=521).
#TF#SERVER: Sent 362644 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2060 bytes expected).
#TF#SERVER: Received 2060 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 465041 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 516817 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=57, total=576).
#TF#SERVER: Sent 433594 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2280 bytes expected).
#TF#SERVER: Received 2280 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 543671 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 581166 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 607770 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 587197 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=6, total=525).
#TF#SERVER: Sent 370012 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2076 bytes expected).
#TF#SERVER: Received 2076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=163, total=682).
#TF#SERVER: Sent 604060 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2704 bytes expected).
#TF#SERVER: Received 2704 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=152, total=671).
#TF#SERVER: Sent 590372 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2660 bytes expected).
#TF#SERVER: Received 2660 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 537633 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 524565 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=144, total=663).
#TF#SERVER: Sent 545052 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2628 bytes expected).
#TF#SERVER: Received 2628 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 464643 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=158, total=677).
#TF#SERVER: Sent 561445 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2684 bytes expected).
#TF#SERVER: Received 2684 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=185, total=704).
#TF#SERVER: Sent 561188 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2792 bytes expected).
#TF#SERVER: Received 2792 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=145, total=664).
#TF#SERVER: Sent 591643 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2632 bytes expected).
#TF#SERVER: Received 2632 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=184, total=703).
#TF#SERVER: Sent 664412 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2788 bytes expected).
#TF#SERVER: Received 2788 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 537926 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 538582 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 518834 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=11, total=530).
#TF#SERVER: Sent 372152 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2096 bytes expected).
#TF#SERVER: Received 2096 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 547585 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=99, total=618).
#TF#SERVER: Sent 522673 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2448 bytes expected).
#TF#SERVER: Received 2448 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=142, total=661).
#TF#SERVER: Sent 546655 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2620 bytes expected).
#TF#SERVER: Received 2620 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 457770 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=84, total=603).
#TF#SERVER: Sent 452608 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2388 bytes expected).
#TF#SERVER: Received 2388 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 470691 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 458844 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=84, total=603).
#TF#SERVER: Sent 449147 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2388 bytes expected).
#TF#SERVER: Received 2388 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=129, total=648).
#TF#SERVER: Sent 477260 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2568 bytes expected).
#TF#SERVER: Received 2568 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 582677 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=85, total=604).
#TF#SERVER: Sent 522329 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2392 bytes expected).
#TF#SERVER: Received 2392 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 596219 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 591699 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 585782 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=16, total=535).
#TF#SERVER: Sent 382868 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2116 bytes expected).
#TF#SERVER: Received 2116 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 614681 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 625108 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 594963 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=13, total=532).
#TF#SERVER: Sent 383832 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2104 bytes expected).
#TF#SERVER: Received 2104 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 541175 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=84, total=603).
#TF#SERVER: Sent 499880 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2388 bytes expected).
#TF#SERVER: Received 2388 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 542393 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 600865 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=11, total=530).
#TF#SERVER: Sent 380092 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2096 bytes expected).
#TF#SERVER: Received 2096 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=150, total=669).
#TF#SERVER: Sent 614954 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2652 bytes expected).
#TF#SERVER: Received 2652 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 609415 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 597160 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=31, total=550).
#TF#SERVER: Sent 417488 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2176 bytes expected).
#TF#SERVER: Received 2176 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 503271 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 573523 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=21, total=540).
#TF#SERVER: Sent 398060 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2136 bytes expected).
#TF#SERVER: Received 2136 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 662333 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 671539 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=74, total=593).
#TF#SERVER: Sent 537703 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2348 bytes expected).
#TF#SERVER: Received 2348 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 621014 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 612442 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 607049 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=29, total=548).
#TF#SERVER: Sent 414193 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2168 bytes expected).
#TF#SERVER: Received 2168 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 590892 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 595040 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=43, total=562).
#TF#SERVER: Sent 422563 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2224 bytes expected).
#TF#SERVER: Received 2224 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 621635 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 654341 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 654932 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=25, total=544).
#TF#SERVER: Sent 407845 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2152 bytes expected).
#TF#SERVER: Received 2152 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 578979 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=128, total=647).
#TF#SERVER: Sent 587462 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=51, total=570).
#TF#SERVER: Sent 441071 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2256 bytes expected).
#TF#SERVER: Received 2256 bytes from the TF server.

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(free_g3_algstr_0, axiom, ![X1, X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k2_zfmisc_1(X1,X1),X1))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))))=>![X3, X4]:(g3_algstr_0(X1,X2)=g3_algstr_0(X3,X4)=>(X1=X3&X2=X4))), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', free_g3_algstr_0)).
fof(dt_u2_algstr_0, axiom, ![X1]:(l3_algstr_0(X1)=>((v1_funct_1(u2_algstr_0(X1))&v1_funct_2(u2_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)))&m1_subset_1(u2_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)))))), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', dt_u2_algstr_0)).
fof(abstractness_v15_algstr_0, axiom, ![X1]:(l3_algstr_0(X1)=>(v15_algstr_0(X1)=>X1=g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', abstractness_v15_algstr_0)).
fof(dt_g3_algstr_0, axiom, ![X1, X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k2_zfmisc_1(X1,X1),X1))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))))=>(v15_algstr_0(g3_algstr_0(X1,X2))&l3_algstr_0(g3_algstr_0(X1,X2)))), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', dt_g3_algstr_0)).
fof(dt_u1_algstr_0, axiom, ![X1]:(l1_algstr_0(X1)=>((v1_funct_1(u1_algstr_0(X1))&v1_funct_2(u1_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)))&m1_subset_1(u1_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)))))), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', dt_u1_algstr_0)).
fof(d18_algstr_0, axiom, ![X1]:(l3_algstr_0(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k6_algstr_0(X1,X2,X3)=k5_binop_1(u1_struct_0(X1),u2_algstr_0(X1),X2,X3)))), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', d18_algstr_0)).
fof(d1_algstr_0, axiom, ![X1]:(l1_algstr_0(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k1_algstr_0(X1,X2,X3)=k5_binop_1(u1_struct_0(X1),u1_algstr_0(X1),X2,X3)))), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', d1_algstr_0)).
fof(d12_group_1, axiom, ![X1]:(l3_algstr_0(X1)=>(v5_group_1(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k6_algstr_0(X1,X2,X3)=k6_algstr_0(X1,X3,X2))))), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', d12_group_1)).
fof(redefinition_k3_rlvect_1, axiom, ![X1, X2, X3]:((((v2_rlvect_1(X1)&l1_algstr_0(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>k3_rlvect_1(X1,X2,X3)=k1_algstr_0(X1,X2,X3)), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', redefinition_k3_rlvect_1)).
fof(commutativity_k3_rlvect_1, axiom, ![X1, X2, X3]:((((v2_rlvect_1(X1)&l1_algstr_0(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>k3_rlvect_1(X1,X2,X3)=k3_rlvect_1(X1,X3,X2)), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', commutativity_k3_rlvect_1)).
fof(t35_vectsp_1, conjecture, ![X1]:((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&l2_algstr_0(X1))=>((((~(v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))&v2_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))&v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))&v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))&l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', t35_vectsp_1)).
fof(redefinition_k5_binop_1, axiom, ![X1, X2, X3, X4]:(((((v1_funct_1(X2)&v1_funct_2(X2,k2_zfmisc_1(X1,X1),X1))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))))&m1_subset_1(X3,X1))&m1_subset_1(X4,X1))=>k5_binop_1(X1,X2,X3,X4)=k1_binop_1(X2,X3,X4)), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', redefinition_k5_binop_1)).
fof(dt_l2_algstr_0, axiom, ![X1]:(l2_algstr_0(X1)=>(l2_struct_0(X1)&l1_algstr_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', dt_l2_algstr_0)).
fof(d3_rlvect_1, axiom, ![X1]:(l1_algstr_0(X1)=>(v3_rlvect_1(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))))))), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', d3_rlvect_1)).
fof(t5_rlvect_1, axiom, ![X1]:(((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&l2_algstr_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))=k4_struct_0(X1)&k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)=k4_struct_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', t5_rlvect_1)).
fof(dt_k1_algstr_0, axiom, ![X1, X2, X3]:(((l1_algstr_0(X1)&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', dt_k1_algstr_0)).
fof(d3_group_1, axiom, ![X1]:(l3_algstr_0(X1)=>(v3_group_1(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>k6_algstr_0(X1,k6_algstr_0(X1,X2,X3),X4)=k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4))))))), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', d3_group_1)).
fof(dt_k4_algstr_0, axiom, ![X1, X2]:((l2_algstr_0(X1)&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', dt_k4_algstr_0)).
fof(t4_rlvect_1, axiom, ![X1]:(((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&l2_algstr_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_algstr_0(X1,X2,k4_struct_0(X1))=X2&k1_algstr_0(X1,k4_struct_0(X1),X2)=X2))), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', t4_rlvect_1)).
fof(existence_m1_subset_1, axiom, ![X1]:?[X2]:m1_subset_1(X2,X1), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', existence_m1_subset_1)).
fof(d2_group_1, axiom, ![X1]:(l3_algstr_0(X1)=>(v2_group_1(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((k6_algstr_0(X1,X3,X2)=X3&k6_algstr_0(X1,X2,X3)=X3)&?[X4]:((m1_subset_1(X4,u1_struct_0(X1))&k6_algstr_0(X1,X3,X4)=X2)&k6_algstr_0(X1,X4,X3)=X2)))))), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', d2_group_1)).
fof(fc12_algstr_0, axiom, ![X1, X2]:((((~(v1_xboole_0(X1))&v1_funct_1(X2))&v1_funct_2(X2,k2_zfmisc_1(X1,X1),X1))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))))=>(~(v2_struct_0(g3_algstr_0(X1,X2)))&v15_algstr_0(g3_algstr_0(X1,X2)))), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', fc12_algstr_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', fc2_struct_0)).
fof(dt_l1_algstr_0, axiom, ![X1]:(l1_algstr_0(X1)=>l1_struct_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t35_vectsp_1', dt_l1_algstr_0)).
fof(c_0_24, plain, ![X53, X54, X55, X56]:((X53=X55|g3_algstr_0(X53,X54)!=g3_algstr_0(X55,X56)|(~v1_funct_1(X54)|~v1_funct_2(X54,k2_zfmisc_1(X53,X53),X53)|~m1_subset_1(X54,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X53,X53),X53)))))&(X54=X56|g3_algstr_0(X53,X54)!=g3_algstr_0(X55,X56)|(~v1_funct_1(X54)|~v1_funct_2(X54,k2_zfmisc_1(X53,X53),X53)|~m1_subset_1(X54,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X53,X53),X53)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[free_g3_algstr_0])])])])).
fof(c_0_25, plain, ![X45]:(((v1_funct_1(u2_algstr_0(X45))|~l3_algstr_0(X45))&(v1_funct_2(u2_algstr_0(X45),k2_zfmisc_1(u1_struct_0(X45),u1_struct_0(X45)),u1_struct_0(X45))|~l3_algstr_0(X45)))&(m1_subset_1(u2_algstr_0(X45),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X45),u1_struct_0(X45)),u1_struct_0(X45))))|~l3_algstr_0(X45))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u2_algstr_0])])])).
cnf(c_0_26, plain, (X1=X2|g3_algstr_0(X1,X3)!=g3_algstr_0(X2,X4)|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_27, plain, (v1_funct_2(u2_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_28, plain, (X1=X2|g3_algstr_0(X1,X3)!=g3_algstr_0(X2,X4)|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))), inference(evalgc,[status(thm)],[c_0_26])).
cnf(c_0_29, plain, (v1_funct_2(u2_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~l3_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_27])).
cnf(c_0_30, plain, (m1_subset_1(u2_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31, plain, (v1_funct_1(u2_algstr_0(X1))|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_32, plain, ![X36]:(~l3_algstr_0(X36)|(~v15_algstr_0(X36)|X36=g3_algstr_0(u1_struct_0(X36),u2_algstr_0(X36)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[abstractness_v15_algstr_0])])).
fof(c_0_33, plain, ![X70, X71]:((v15_algstr_0(g3_algstr_0(X70,X71))|(~v1_funct_1(X71)|~v1_funct_2(X71,k2_zfmisc_1(X70,X70),X70)|~m1_subset_1(X71,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X70,X70),X70)))))&(l3_algstr_0(g3_algstr_0(X70,X71))|(~v1_funct_1(X71)|~v1_funct_2(X71,k2_zfmisc_1(X70,X70),X70)|~m1_subset_1(X71,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X70,X70),X70)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_g3_algstr_0])])])).
fof(c_0_34, plain, ![X41]:(((v1_funct_1(u1_algstr_0(X41))|~l1_algstr_0(X41))&(v1_funct_2(u1_algstr_0(X41),k2_zfmisc_1(u1_struct_0(X41),u1_struct_0(X41)),u1_struct_0(X41))|~l1_algstr_0(X41)))&(m1_subset_1(u1_algstr_0(X41),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X41),u1_struct_0(X41)),u1_struct_0(X41))))|~l1_algstr_0(X41))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_algstr_0])])])).
cnf(c_0_35, plain, (u1_struct_0(X1)=X2|g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))!=g3_algstr_0(X2,X3)|~v1_funct_1(u2_algstr_0(X1))|~m1_subset_1(u2_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l3_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29])])).
cnf(c_0_36, plain, (m1_subset_1(u2_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l3_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_30])).
cnf(c_0_37, plain, (v1_funct_1(u2_algstr_0(X1))|~l3_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_31])).
cnf(c_0_38, plain, (X1=g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))|~l3_algstr_0(X1)|~v15_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39, plain, (v15_algstr_0(g3_algstr_0(X1,X2))|~v1_funct_1(X2)|~v1_funct_2(X2,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_40, plain, (v1_funct_2(u1_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~l1_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_41, plain, (X1=X2|g3_algstr_0(X3,X1)!=g3_algstr_0(X4,X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k2_zfmisc_1(X3,X3),X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X3,X3),X3)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_42, plain, (u1_struct_0(X1)=X2|g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))!=g3_algstr_0(X2,X3)|~l3_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_35, c_0_36]), c_0_37])).
cnf(c_0_43, plain, (g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))=X1|~l3_algstr_0(X1)|~v15_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_38])).
cnf(c_0_44, plain, (v15_algstr_0(g3_algstr_0(X1,X2))|~v1_funct_1(X2)|~v1_funct_2(X2,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))), inference(evalgc,[status(thm)],[c_0_39])).
cnf(c_0_45, plain, (v1_funct_2(u1_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_40])).
cnf(c_0_46, plain, (m1_subset_1(u1_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_47, plain, (v1_funct_1(u1_algstr_0(X1))|~l1_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_48, plain, (l3_algstr_0(g3_algstr_0(X1,X2))|~v1_funct_1(X2)|~v1_funct_2(X2,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_49, plain, (X1=X2|g3_algstr_0(X3,X1)!=g3_algstr_0(X4,X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k2_zfmisc_1(X3,X3),X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X3,X3),X3)))), inference(evalgc,[status(thm)],[c_0_41])).
cnf(c_0_50, plain, (u1_struct_0(X1)=X2|X1!=g3_algstr_0(X2,X3)|~l3_algstr_0(X1)|~v15_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43])])).
cnf(c_0_51, plain, (v15_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v1_funct_1(u1_algstr_0(X1))|~m1_subset_1(u1_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45])])).
cnf(c_0_52, plain, (m1_subset_1(u1_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_46])).
cnf(c_0_53, plain, (v1_funct_1(u1_algstr_0(X1))|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_47])).
cnf(c_0_54, plain, (l3_algstr_0(g3_algstr_0(X1,X2))|~v1_funct_1(X2)|~v1_funct_2(X2,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))), inference(evalgc,[status(thm)],[c_0_48])).
cnf(c_0_55, plain, (u2_algstr_0(X1)=X2|g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))!=g3_algstr_0(X3,X2)|~v1_funct_1(u2_algstr_0(X1))|~m1_subset_1(u2_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l3_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_29])])).
fof(c_0_56, plain, ![X67, X68, X69]:(~l3_algstr_0(X67)|(~m1_subset_1(X68,u1_struct_0(X67))|(~m1_subset_1(X69,u1_struct_0(X67))|k6_algstr_0(X67,X68,X69)=k5_binop_1(u1_struct_0(X67),u2_algstr_0(X67),X68,X69)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d18_algstr_0])])])).
cnf(c_0_57, plain, (u1_struct_0(g3_algstr_0(X1,X2))=X1|~v15_algstr_0(g3_algstr_0(X1,X2))|~l3_algstr_0(g3_algstr_0(X1,X2))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_50])])).
cnf(c_0_58, plain, (v15_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_51, c_0_52]), c_0_53])).
cnf(c_0_59, plain, (l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v1_funct_1(u1_algstr_0(X1))|~m1_subset_1(u1_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_45])])).
cnf(c_0_60, plain, (u2_algstr_0(X1)=X2|g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))!=g3_algstr_0(X3,X2)|~l3_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_55, c_0_36]), c_0_37])).
cnf(c_0_61, plain, (k6_algstr_0(X1,X2,X3)=k5_binop_1(u1_struct_0(X1),u2_algstr_0(X1),X2,X3)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_62, plain, (u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))=u1_struct_0(X1)|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58])])).
cnf(c_0_63, plain, (l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_59, c_0_52]), c_0_53])).
cnf(c_0_64, plain, (u2_algstr_0(X1)=X2|X1!=g3_algstr_0(X3,X2)|~l3_algstr_0(X1)|~v15_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_43])])).
cnf(c_0_65, plain, (k5_binop_1(u1_struct_0(X1),u2_algstr_0(X1),X2,X3)=k6_algstr_0(X1,X2,X3)|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_61])).
cnf(c_0_66, plain, (u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))=u1_struct_0(X1)|~l1_algstr_0(X1)), inference(csr,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_67, plain, (u2_algstr_0(g3_algstr_0(X1,X2))=X2|~v15_algstr_0(g3_algstr_0(X1,X2))|~l3_algstr_0(g3_algstr_0(X1,X2))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_64])])).
fof(c_0_68, plain, ![X72, X73, X74]:(~l1_algstr_0(X72)|(~m1_subset_1(X73,u1_struct_0(X72))|(~m1_subset_1(X74,u1_struct_0(X72))|k1_algstr_0(X72,X73,X74)=k5_binop_1(u1_struct_0(X72),u1_algstr_0(X72),X73,X74)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_algstr_0])])])).
cnf(c_0_69, plain, (k5_binop_1(u1_struct_0(X1),u2_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),X2,X3)=k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66])])).
cnf(c_0_70, plain, (u2_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))=u1_algstr_0(X1)|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_58])])).
fof(c_0_71, plain, ![X10, X11, X12]:((~v5_group_1(X10)|(~m1_subset_1(X11,u1_struct_0(X10))|(~m1_subset_1(X12,u1_struct_0(X10))|k6_algstr_0(X10,X11,X12)=k6_algstr_0(X10,X12,X11)))|~l3_algstr_0(X10))&((m1_subset_1(esk2_1(X10),u1_struct_0(X10))|v5_group_1(X10)|~l3_algstr_0(X10))&((m1_subset_1(esk3_1(X10),u1_struct_0(X10))|v5_group_1(X10)|~l3_algstr_0(X10))&(k6_algstr_0(X10,esk2_1(X10),esk3_1(X10))!=k6_algstr_0(X10,esk3_1(X10),esk2_1(X10))|v5_group_1(X10)|~l3_algstr_0(X10))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_group_1])])])])])).
cnf(c_0_72, plain, (k1_algstr_0(X1,X2,X3)=k5_binop_1(u1_struct_0(X1),u1_algstr_0(X1),X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_73, plain, (k5_binop_1(u1_struct_0(X1),u2_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),X2,X3)=k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[c_0_69, c_0_63])).
cnf(c_0_74, plain, (u2_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))=u1_algstr_0(X1)|~l1_algstr_0(X1)), inference(csr,[status(thm)],[c_0_70, c_0_63])).
cnf(c_0_75, plain, (m1_subset_1(esk3_1(X1),u1_struct_0(X1))|v5_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_76, plain, (m1_subset_1(esk2_1(X1),u1_struct_0(X1))|v5_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_77, plain, (v5_group_1(X1)|k6_algstr_0(X1,esk2_1(X1),esk3_1(X1))!=k6_algstr_0(X1,esk3_1(X1),esk2_1(X1))|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_78, plain, (k5_binop_1(u1_struct_0(X1),u1_algstr_0(X1),X2,X3)=k1_algstr_0(X1,X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_72])).
cnf(c_0_79, plain, (k5_binop_1(u1_struct_0(X1),u1_algstr_0(X1),X2,X3)=k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74])])).
cnf(c_0_80, plain, (v5_group_1(X1)|m1_subset_1(esk3_1(X1),u1_struct_0(X1))|~l3_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_75])).
cnf(c_0_81, plain, (v5_group_1(X1)|m1_subset_1(esk2_1(X1),u1_struct_0(X1))|~l3_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_76])).
fof(c_0_82, plain, ![X37, X38, X39]:(~v2_rlvect_1(X37)|~l1_algstr_0(X37)|~m1_subset_1(X38,u1_struct_0(X37))|~m1_subset_1(X39,u1_struct_0(X37))|k3_rlvect_1(X37,X38,X39)=k1_algstr_0(X37,X38,X39)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k3_rlvect_1])])).
fof(c_0_83, plain, ![X42, X43, X44]:(~v2_rlvect_1(X42)|~l1_algstr_0(X42)|~m1_subset_1(X43,u1_struct_0(X42))|~m1_subset_1(X44,u1_struct_0(X42))|k3_rlvect_1(X42,X43,X44)=k3_rlvect_1(X42,X44,X43)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k3_rlvect_1])])).
fof(c_0_84, negated_conjecture, ~(![X1]:((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&l2_algstr_0(X1))=>((((~(v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))&v2_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))&v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))&v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))&l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))), inference(assume_negation,[status(cth)],[t35_vectsp_1])).
fof(c_0_85, plain, ![X111, X112, X113, X114]:(~v1_funct_1(X112)|~v1_funct_2(X112,k2_zfmisc_1(X111,X111),X111)|~m1_subset_1(X112,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X111,X111),X111)))|~m1_subset_1(X113,X111)|~m1_subset_1(X114,X111)|k5_binop_1(X111,X112,X113,X114)=k1_binop_1(X112,X113,X114)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k5_binop_1])])).
cnf(c_0_86, plain, (v5_group_1(X1)|k6_algstr_0(X1,esk3_1(X1),esk2_1(X1))!=k6_algstr_0(X1,esk2_1(X1),esk3_1(X1))|~l3_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_77])).
cnf(c_0_87, plain, (k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2,X3)=k1_algstr_0(X1,X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79])])).
cnf(c_0_88, plain, (m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_66])])).
cnf(c_0_89, plain, (m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_66])])).
cnf(c_0_90, plain, (k3_rlvect_1(X1,X2,X3)=k1_algstr_0(X1,X2,X3)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_91, plain, (k3_rlvect_1(X1,X2,X3)=k3_rlvect_1(X1,X3,X2)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_83])).
fof(c_0_92, plain, ![X40]:((l2_struct_0(X40)|~l2_algstr_0(X40))&(l1_algstr_0(X40)|~l2_algstr_0(X40))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l2_algstr_0])])])).
fof(c_0_93, negated_conjecture, ((((((~v2_struct_0(esk1_0)&v13_algstr_0(esk1_0))&v2_rlvect_1(esk1_0))&v3_rlvect_1(esk1_0))&v4_rlvect_1(esk1_0))&l2_algstr_0(esk1_0))&(v2_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~v5_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_84])])])])).
cnf(c_0_94, plain, (k5_binop_1(X2,X1,X3,X4)=k1_binop_1(X1,X3,X4)|~v1_funct_1(X1)|~v1_funct_2(X1,k2_zfmisc_1(X2,X2),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))|~m1_subset_1(X3,X2)|~m1_subset_1(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_95, plain, (v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k1_algstr_0(X1,esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)|~m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87])])).
cnf(c_0_96, plain, (m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)), inference(csr,[status(thm)],[c_0_88, c_0_63])).
cnf(c_0_97, plain, (m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)), inference(csr,[status(thm)],[c_0_89, c_0_63])).
cnf(c_0_98, plain, (k3_rlvect_1(X1,X2,X3)=k1_algstr_0(X1,X2,X3)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_90])).
cnf(c_0_99, plain, (k3_rlvect_1(X1,X2,X3)=k3_rlvect_1(X1,X3,X2)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_91])).
cnf(c_0_100, plain, (l1_algstr_0(X1)|~l2_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_101, negated_conjecture, (l2_algstr_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
fof(c_0_102, plain, ![X29, X30, X31, X32]:((~v3_rlvect_1(X29)|(~m1_subset_1(X30,u1_struct_0(X29))|(~m1_subset_1(X31,u1_struct_0(X29))|(~m1_subset_1(X32,u1_struct_0(X29))|k1_algstr_0(X29,k1_algstr_0(X29,X30,X31),X32)=k1_algstr_0(X29,X30,k1_algstr_0(X29,X31,X32)))))|~l1_algstr_0(X29))&((m1_subset_1(esk10_1(X29),u1_struct_0(X29))|v3_rlvect_1(X29)|~l1_algstr_0(X29))&((m1_subset_1(esk11_1(X29),u1_struct_0(X29))|v3_rlvect_1(X29)|~l1_algstr_0(X29))&((m1_subset_1(esk12_1(X29),u1_struct_0(X29))|v3_rlvect_1(X29)|~l1_algstr_0(X29))&(k1_algstr_0(X29,k1_algstr_0(X29,esk10_1(X29),esk11_1(X29)),esk12_1(X29))!=k1_algstr_0(X29,esk10_1(X29),k1_algstr_0(X29,esk11_1(X29),esk12_1(X29)))|v3_rlvect_1(X29)|~l1_algstr_0(X29)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_rlvect_1])])])])])).
fof(c_0_103, plain, ![X6, X7]:((k1_algstr_0(X6,X7,k4_algstr_0(X6,X7))=k4_struct_0(X6)|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|~v13_algstr_0(X6)|~v3_rlvect_1(X6)|~v4_rlvect_1(X6)|~l2_algstr_0(X6)))&(k1_algstr_0(X6,k4_algstr_0(X6,X7),X7)=k4_struct_0(X6)|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|~v13_algstr_0(X6)|~v3_rlvect_1(X6)|~v4_rlvect_1(X6)|~l2_algstr_0(X6)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t5_rlvect_1])])])])])).
fof(c_0_104, plain, ![X63, X64, X65]:(~l1_algstr_0(X63)|~m1_subset_1(X64,u1_struct_0(X63))|~m1_subset_1(X65,u1_struct_0(X63))|m1_subset_1(k1_algstr_0(X63,X64,X65),u1_struct_0(X63))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_algstr_0])])).
fof(c_0_105, plain, ![X22, X23, X24, X25]:((~v3_group_1(X22)|(~m1_subset_1(X23,u1_struct_0(X22))|(~m1_subset_1(X24,u1_struct_0(X22))|(~m1_subset_1(X25,u1_struct_0(X22))|k6_algstr_0(X22,k6_algstr_0(X22,X23,X24),X25)=k6_algstr_0(X22,X23,k6_algstr_0(X22,X24,X25)))))|~l3_algstr_0(X22))&((m1_subset_1(esk7_1(X22),u1_struct_0(X22))|v3_group_1(X22)|~l3_algstr_0(X22))&((m1_subset_1(esk8_1(X22),u1_struct_0(X22))|v3_group_1(X22)|~l3_algstr_0(X22))&((m1_subset_1(esk9_1(X22),u1_struct_0(X22))|v3_group_1(X22)|~l3_algstr_0(X22))&(k6_algstr_0(X22,k6_algstr_0(X22,esk7_1(X22),esk8_1(X22)),esk9_1(X22))!=k6_algstr_0(X22,esk7_1(X22),k6_algstr_0(X22,esk8_1(X22),esk9_1(X22)))|v3_group_1(X22)|~l3_algstr_0(X22)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_group_1])])])])])).
cnf(c_0_106, plain, (k5_binop_1(X2,X1,X3,X4)=k1_binop_1(X1,X3,X4)|~v1_funct_1(X1)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v1_funct_2(X1,k2_zfmisc_1(X2,X2),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2)))), inference(evalgc,[status(thm)],[c_0_94])).
cnf(c_0_107, plain, (v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k1_algstr_0(X1,esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_95, c_0_63]), c_0_96]), c_0_97])).
cnf(c_0_108, plain, (k3_rlvect_1(X1,X3,X2)=k1_algstr_0(X1,X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99])])).
cnf(c_0_109, negated_conjecture, (v2_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_110, plain, (l1_algstr_0(X1)|~l2_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_100])).
cnf(c_0_111, negated_conjecture, (l2_algstr_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_101])).
cnf(c_0_112, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|~v3_rlvect_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_113, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_103])).
fof(c_0_114, plain, ![X57, X58]:(~l2_algstr_0(X57)|~m1_subset_1(X58,u1_struct_0(X57))|m1_subset_1(k4_algstr_0(X57,X58),u1_struct_0(X57))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_algstr_0])])).
cnf(c_0_115, plain, (m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_116, plain, (m1_subset_1(esk9_1(X1),u1_struct_0(X1))|v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_117, plain, (k5_binop_1(u1_struct_0(X1),u2_algstr_0(X1),X2,X3)=k1_binop_1(u2_algstr_0(X1),X2,X3)|~v1_funct_1(u2_algstr_0(X1))|~m1_subset_1(u2_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_29])])).
cnf(c_0_118, plain, (v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_algstr_0(X1,esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k1_algstr_0(X1,esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_87])])).
cnf(c_0_119, plain, (k1_algstr_0(X1,X3,X2)=k1_algstr_0(X1,X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_108])])).
cnf(c_0_120, negated_conjecture, (v2_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_109])).
cnf(c_0_121, negated_conjecture, (l1_algstr_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111])])).
cnf(c_0_122, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|~v3_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_112])).
cnf(c_0_123, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))=k4_struct_0(X1)|v2_struct_0(X1)|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_113])).
cnf(c_0_124, plain, (m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_114])).
fof(c_0_125, plain, ![X8, X9]:((k1_algstr_0(X8,X9,k4_struct_0(X8))=X9|~m1_subset_1(X9,u1_struct_0(X8))|(v2_struct_0(X8)|~v13_algstr_0(X8)|~v3_rlvect_1(X8)|~v4_rlvect_1(X8)|~l2_algstr_0(X8)))&(k1_algstr_0(X8,k4_struct_0(X8),X9)=X9|~m1_subset_1(X9,u1_struct_0(X8))|(v2_struct_0(X8)|~v13_algstr_0(X8)|~v3_rlvect_1(X8)|~v4_rlvect_1(X8)|~l2_algstr_0(X8)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_rlvect_1])])])])])).
cnf(c_0_126, plain, (m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_115])).
fof(c_0_127, plain, ![X76]:m1_subset_1(esk13_1(X76),X76), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[existence_m1_subset_1])])).
cnf(c_0_128, plain, (v3_group_1(X1)|k6_algstr_0(X1,k6_algstr_0(X1,esk7_1(X1),esk8_1(X1)),esk9_1(X1))!=k6_algstr_0(X1,esk7_1(X1),k6_algstr_0(X1,esk8_1(X1),esk9_1(X1)))|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_129, plain, (v3_group_1(X1)|m1_subset_1(esk9_1(X1),u1_struct_0(X1))|~l3_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_116])).
cnf(c_0_130, plain, (k5_binop_1(u1_struct_0(X1),u2_algstr_0(X1),X2,X3)=k1_binop_1(u2_algstr_0(X1),X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_117, c_0_36]), c_0_37])).
cnf(c_0_131, plain, (v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_algstr_0(X1,esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k1_algstr_0(X1,esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_118, c_0_97]), c_0_96])).
cnf(c_0_132, negated_conjecture, (k1_algstr_0(esk1_0,X1,X2)=k1_algstr_0(esk1_0,X2,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120]), c_0_121])])])).
cnf(c_0_133, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_algstr_0(X1,X2),X3))=k1_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123])])).
cnf(c_0_134, plain, (m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_124])).
cnf(c_0_135, plain, (k1_algstr_0(X1,X2,k4_struct_0(X1))=X2|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_125])).
cnf(c_0_136, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_123])])).
cnf(c_0_137, plain, (m1_subset_1(esk13_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_138, plain, (v3_group_1(X1)|k6_algstr_0(X1,k6_algstr_0(X1,esk7_1(X1),esk8_1(X1)),esk9_1(X1))!=k6_algstr_0(X1,esk7_1(X1),k6_algstr_0(X1,esk8_1(X1),esk9_1(X1)))|~l3_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_128])).
cnf(c_0_139, plain, (m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_66])])).
cnf(c_0_140, plain, (k1_binop_1(u2_algstr_0(X1),X2,X3)=k6_algstr_0(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_130])])).
cnf(c_0_141, negated_conjecture, (v5_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_132]), c_0_121])])])).
cnf(c_0_142, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_algstr_0(X1,X2),X3))=k1_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_133, c_0_134]), c_0_110])).
cnf(c_0_143, plain, (k1_algstr_0(X1,X2,k4_struct_0(X1))=X2|v2_struct_0(X1)|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_135])).
cnf(c_0_144, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_136, c_0_134]), c_0_110])).
cnf(c_0_145, plain, (m1_subset_1(esk13_1(X1),X1)), inference(evalgc,[status(thm)],[c_0_137])).
cnf(c_0_146, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))!=k1_algstr_0(X1,k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_87])])).
cnf(c_0_147, plain, (m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)), inference(csr,[status(thm)],[c_0_139, c_0_63])).
cnf(c_0_148, plain, (k1_binop_1(u1_algstr_0(X1),X2,X3)=k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2,X3)|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_74])])).
fof(c_0_149, plain, ![X15, X17, X19, X21]:(((m1_subset_1(esk4_1(X15),u1_struct_0(X15))|~v2_group_1(X15)|~l3_algstr_0(X15))&(((k6_algstr_0(X15,X17,esk4_1(X15))=X17|~m1_subset_1(X17,u1_struct_0(X15))|~v2_group_1(X15)|~l3_algstr_0(X15))&(k6_algstr_0(X15,esk4_1(X15),X17)=X17|~m1_subset_1(X17,u1_struct_0(X15))|~v2_group_1(X15)|~l3_algstr_0(X15)))&(((m1_subset_1(esk5_2(X15,X17),u1_struct_0(X15))|~m1_subset_1(X17,u1_struct_0(X15))|~v2_group_1(X15)|~l3_algstr_0(X15))&(k6_algstr_0(X15,X17,esk5_2(X15,X17))=esk4_1(X15)|~m1_subset_1(X17,u1_struct_0(X15))|~v2_group_1(X15)|~l3_algstr_0(X15)))&(k6_algstr_0(X15,esk5_2(X15,X17),X17)=esk4_1(X15)|~m1_subset_1(X17,u1_struct_0(X15))|~v2_group_1(X15)|~l3_algstr_0(X15)))))&((m1_subset_1(esk6_2(X15,X19),u1_struct_0(X15))|~m1_subset_1(X19,u1_struct_0(X15))|v2_group_1(X15)|~l3_algstr_0(X15))&(k6_algstr_0(X15,esk6_2(X15,X19),X19)!=esk6_2(X15,X19)|k6_algstr_0(X15,X19,esk6_2(X15,X19))!=esk6_2(X15,X19)|(~m1_subset_1(X21,u1_struct_0(X15))|k6_algstr_0(X15,esk6_2(X15,X19),X21)!=X19|k6_algstr_0(X15,X21,esk6_2(X15,X19))!=X19)|~m1_subset_1(X19,u1_struct_0(X15))|v2_group_1(X15)|~l3_algstr_0(X15)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_group_1])])])])])).
cnf(c_0_150, plain, (k6_algstr_0(X1,X2,X3)=k6_algstr_0(X1,X3,X2)|~v5_group_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_151, negated_conjecture, (v5_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_96]), c_0_121])])])).
cnf(c_0_152, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_153, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))=k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_143])])).
cnf(c_0_154, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_145])])).
cnf(c_0_155, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))!=k1_algstr_0(X1,k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_146, c_0_63]), c_0_147])).
cnf(c_0_156, plain, (k1_binop_1(u1_algstr_0(X1),X2,X3)=k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))), inference(csr,[status(thm)],[c_0_148, c_0_63])).
cnf(c_0_157, plain, (k5_binop_1(u1_struct_0(X1),u1_algstr_0(X1),X2,X3)=k1_binop_1(u1_algstr_0(X1),X2,X3)|~v1_funct_1(u1_algstr_0(X1))|~m1_subset_1(u1_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_45])])).
cnf(c_0_158, plain, (m1_subset_1(esk8_1(X1),u1_struct_0(X1))|v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_159, plain, (v2_group_1(X1)|k6_algstr_0(X1,esk6_2(X1,X2),X2)!=esk6_2(X1,X2)|k6_algstr_0(X1,X2,esk6_2(X1,X2))!=esk6_2(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|k6_algstr_0(X1,esk6_2(X1,X2),X3)!=X2|k6_algstr_0(X1,X3,esk6_2(X1,X2))!=X2|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_149])).
cnf(c_0_160, plain, (k6_algstr_0(X1,X2,X3)=k6_algstr_0(X1,X3,X2)|~v5_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_150])).
cnf(c_0_161, negated_conjecture, (v5_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_97]), c_0_121])])])).
cnf(c_0_162, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)=k4_struct_0(X1)|v2_struct_0(X1)|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_152])).
cnf(c_0_163, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))=k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_153, c_0_154]), c_0_134])).
cnf(c_0_164, negated_conjecture, (v4_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_165, negated_conjecture, (v3_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_166, negated_conjecture, (v13_algstr_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_167, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_168, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_binop_1(u1_algstr_0(X1),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))!=k1_algstr_0(X1,k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_156])])).
cnf(c_0_169, plain, (k5_binop_1(u1_struct_0(X1),u1_algstr_0(X1),X2,X3)=k1_binop_1(u1_algstr_0(X1),X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_157, c_0_52]), c_0_53])).
cnf(c_0_170, plain, (v3_group_1(X1)|m1_subset_1(esk8_1(X1),u1_struct_0(X1))|~l3_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_158])).
cnf(c_0_171, plain, (v2_group_1(X1)|k6_algstr_0(X1,X3,esk6_2(X1,X2))!=X2|k6_algstr_0(X1,esk6_2(X1,X2),X3)!=X2|k6_algstr_0(X1,X2,esk6_2(X1,X2))!=esk6_2(X1,X2)|k6_algstr_0(X1,esk6_2(X1,X2),X2)!=esk6_2(X1,X2)|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_159])).
cnf(c_0_172, negated_conjecture, (k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X2)=k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X1)|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_161])])).
cnf(c_0_173, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~v3_rlvect_1(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_162])])).
cnf(c_0_174, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k4_struct_0(X1)))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_122])])).
cnf(c_0_175, plain, (k1_algstr_0(X1,k4_struct_0(X1),X2)=X2|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_125])).
cnf(c_0_176, plain, (m1_subset_1(k1_algstr_0(X1,X2,k4_algstr_0(X1,X2)),u1_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_163])])).
cnf(c_0_177, negated_conjecture, (v4_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_164])).
cnf(c_0_178, negated_conjecture, (v3_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_165])).
cnf(c_0_179, negated_conjecture, (v13_algstr_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_166])).
cnf(c_0_180, negated_conjecture, (~v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_167])).
cnf(c_0_181, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_binop_1(u1_algstr_0(X1),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_binop_1(u1_algstr_0(X1),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))!=k1_algstr_0(X1,k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(k1_binop_1(u1_algstr_0(X1),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_156])])).
cnf(c_0_182, plain, (k1_binop_1(u1_algstr_0(X1),X2,X3)=k1_algstr_0(X1,X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_169])])).
cnf(c_0_183, plain, (m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_66])])).
cnf(c_0_184, plain, (m1_subset_1(esk7_1(X1),u1_struct_0(X1))|v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_185, plain, (v2_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2,esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2))!=esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2)|k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2),X2)!=esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2)|k1_algstr_0(X1,X3,esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2))!=X2|k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2),X3)!=X2|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_87])])).
cnf(c_0_186, negated_conjecture, (k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X1)=k1_algstr_0(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_172]), c_0_121])])])).
cnf(c_0_187, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_173, c_0_134]), c_0_110])).
cnf(c_0_188, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k4_struct_0(X1)))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_174, c_0_144]), c_0_126]), c_0_110])).
cnf(c_0_189, plain, (m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)),u1_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_163])])).
cnf(c_0_190, plain, (k1_algstr_0(X1,k4_struct_0(X1),X2)=X2|v2_struct_0(X1)|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_175])).
cnf(c_0_191, plain, (m1_subset_1(k1_algstr_0(X1,X2,k4_algstr_0(X1,X2)),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_176, c_0_154]), c_0_110])).
cnf(c_0_192, negated_conjecture, (k1_algstr_0(esk1_0,X1,k4_algstr_0(esk1_0,X1))=k4_struct_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_132]), c_0_111]), c_0_177]), c_0_178]), c_0_179])]), c_0_180])])).
cnf(c_0_193, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_binop_1(u1_algstr_0(X1),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))!=k1_algstr_0(X1,k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_182])])).
cnf(c_0_194, plain, (m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)), inference(csr,[status(thm)],[c_0_183, c_0_63])).
cnf(c_0_195, plain, (v3_group_1(X1)|m1_subset_1(esk7_1(X1),u1_struct_0(X1))|~l3_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_184])).
cnf(c_0_196, plain, (v2_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2,esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2))!=esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2)|k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2),X2)!=esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2)|k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2),X3)!=X2|k1_algstr_0(X1,X3,esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2))!=X2|~l1_algstr_0(X1)|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))), inference(csr,[status(thm)],[c_0_185, c_0_63])).
cnf(c_0_197, negated_conjecture, (k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X2)=k1_algstr_0(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_186])])).
cnf(c_0_198, plain, (m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1))|v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_149])).
cnf(c_0_199, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),X3))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_143])])).
cnf(c_0_200, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)=k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_188])])).
cnf(c_0_201, plain, (m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_187])])).
cnf(c_0_202, plain, (m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_189, c_0_134]), c_0_110])).
cnf(c_0_203, plain, (k4_algstr_0(X1,k4_struct_0(X1))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k4_struct_0(X1)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_190, c_0_123])])).
cnf(c_0_204, negated_conjecture, (m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_192]), c_0_111]), c_0_177]), c_0_178]), c_0_179])]), c_0_180])])).
cnf(c_0_205, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_binop_1(u1_algstr_0(X1),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))!=k1_algstr_0(X1,k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_193, c_0_194]), c_0_147])).
cnf(c_0_206, plain, (m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_66])])).
cnf(c_0_207, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),X1)!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k1_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),X2)!=X1|k1_algstr_0(esk1_0,X2,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=X1|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_197]), c_0_121])])])).
cnf(c_0_208, plain, (v2_group_1(X1)|m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1))|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_198])).
cnf(c_0_209, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),X3))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_199, c_0_144]), c_0_110])).
cnf(c_0_210, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)=k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_200, c_0_154]), c_0_134])).
cnf(c_0_211, plain, (m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_201, c_0_134]), c_0_126]), c_0_110])).
cnf(c_0_212, plain, (m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_145])])).
cnf(c_0_213, plain, (k4_algstr_0(X1,k4_struct_0(X1))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_203, c_0_134])).
cnf(c_0_214, negated_conjecture, (m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_134]), c_0_111])])])).
cnf(c_0_215, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_algstr_0(X1,k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_182])])).
cnf(c_0_216, plain, (m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)), inference(csr,[status(thm)],[c_0_206, c_0_63])).
cnf(c_0_217, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~v3_rlvect_1(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_190])])).
cnf(c_0_218, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),X1)!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k1_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),X2)!=X1|k1_algstr_0(esk1_0,X2,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=X1|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))), inference(csr,[status(thm)],[c_0_207, c_0_208])).
cnf(c_0_219, plain, (k1_algstr_0(X1,k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2)),X2)=k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_210])])).
cnf(c_0_220, plain, (m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_212])])).
cnf(c_0_221, negated_conjecture, (k1_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk1_0))=k4_struct_0(esk1_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_213]), c_0_111]), c_0_177]), c_0_178]), c_0_179])]), c_0_180])])).
cnf(c_0_222, negated_conjecture, (m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_145])])).
cnf(c_0_223, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_algstr_0(X1,k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))), inference(csr,[status(thm)],[c_0_215, c_0_216])).
cnf(c_0_224, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k1_algstr_0(X1,k4_struct_0(X1),X4)))=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_209])])).
cnf(c_0_225, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_217, c_0_154]), c_0_110])).
cnf(c_0_226, negated_conjecture, (k1_algstr_0(esk1_0,X1,k4_struct_0(esk1_0))=X1|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_190, c_0_132]), c_0_111]), c_0_177]), c_0_178]), c_0_179])]), c_0_180])])).
cnf(c_0_227, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),X1)!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k1_algstr_0(esk1_0,X2,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=X1|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_132])])).
cnf(c_0_228, plain, (k1_algstr_0(X1,k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2)),X2)=k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_219, c_0_220])).
cnf(c_0_229, negated_conjecture, (k1_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk1_0))=k4_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_221, c_0_222])])).
cnf(c_0_230, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_algstr_0(X1,k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_87])])).
cnf(c_0_231, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k1_algstr_0(X1,k4_struct_0(X1),X4)))=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_224, c_0_126]), c_0_211]), c_0_110])).
cnf(c_0_232, negated_conjecture, (k1_algstr_0(esk1_0,k4_struct_0(esk1_0),X1)=X1|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_226]), c_0_111]), c_0_177]), c_0_178]), c_0_179])]), c_0_180])])).
cnf(c_0_233, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),X1)!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k4_struct_0(esk1_0)!=X1|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(k4_algstr_0(esk1_0,k1_algstr_0(esk1_0,k4_struct_0(esk1_0),esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(k4_algstr_0(esk1_0,k1_algstr_0(esk1_0,k4_struct_0(esk1_0),esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))),u1_struct_0(esk1_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_228]), c_0_229]), c_0_111]), c_0_177]), c_0_178]), c_0_179])]), c_0_180])])).
cnf(c_0_234, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_algstr_0(X1,k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_230, c_0_216]), c_0_194])).
cnf(c_0_235, negated_conjecture, (k1_algstr_0(esk1_0,k1_algstr_0(esk1_0,X1,X2),X3)=k1_algstr_0(esk1_0,X1,k1_algstr_0(esk1_0,X2,X3))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_232]), c_0_111]), c_0_177]), c_0_178]), c_0_179])]), c_0_180])])).
cnf(c_0_236, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),X1)!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k4_struct_0(esk1_0)!=X1|~m1_subset_1(k4_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(k4_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)),u1_struct_0(esk1_0))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_190]), c_0_111]), c_0_177]), c_0_178]), c_0_179])]), c_0_180])])).
cnf(c_0_237, negated_conjecture, (v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(k1_algstr_0(esk1_0,esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),esk9_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(k1_algstr_0(esk1_0,esk7_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))),u1_struct_0(esk1_0))|~m1_subset_1(k1_algstr_0(esk1_0,esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),esk9_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))),u1_struct_0(esk1_0))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_235]), c_0_121])])])).
cnf(c_0_238, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),X1)!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k4_struct_0(esk1_0)!=X1|~m1_subset_1(k4_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)),u1_struct_0(esk1_0))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_236, c_0_66]), c_0_121])])])).
cnf(c_0_239, negated_conjecture, (v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(k1_algstr_0(esk1_0,esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),esk9_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))),u1_struct_0(esk1_0))|~m1_subset_1(k1_algstr_0(esk1_0,esk7_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))),u1_struct_0(esk1_0))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_66]), c_0_121])])])).
cnf(c_0_240, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k4_struct_0(esk1_0)!=X1|~m1_subset_1(k4_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)),u1_struct_0(esk1_0))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_172])])).
cnf(c_0_241, negated_conjecture, (v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(k1_algstr_0(esk1_0,esk7_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))),u1_struct_0(esk1_0))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_126]), c_0_121])])])).
cnf(c_0_242, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k4_struct_0(esk1_0)!=X1|~m1_subset_1(k4_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)),u1_struct_0(esk1_0))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))), inference(csr,[status(thm)],[c_0_240, c_0_208])).
cnf(c_0_243, negated_conjecture, (v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_126]), c_0_121])])])).
cnf(c_0_244, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k1_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),X1)!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k4_struct_0(esk1_0)!=X1|~m1_subset_1(k4_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)),u1_struct_0(esk1_0))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_186])])).
cnf(c_0_245, negated_conjecture, (v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_147]), c_0_121])])])).
cnf(c_0_246, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k1_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),X1)!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k4_struct_0(esk1_0)!=X1|~m1_subset_1(k4_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)),u1_struct_0(esk1_0))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))), inference(csr,[status(thm)],[c_0_244, c_0_208])).
fof(c_0_247, plain, ![X78, X79]:((~v2_struct_0(g3_algstr_0(X78,X79))|(v1_xboole_0(X78)|~v1_funct_1(X79)|~v1_funct_2(X79,k2_zfmisc_1(X78,X78),X78)|~m1_subset_1(X79,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X78,X78),X78)))))&(v15_algstr_0(g3_algstr_0(X78,X79))|(v1_xboole_0(X78)|~v1_funct_1(X79)|~v1_funct_2(X79,k2_zfmisc_1(X78,X78),X78)|~m1_subset_1(X79,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X78,X78),X78)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc12_algstr_0])])])])).
cnf(c_0_248, negated_conjecture, (v2_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~v5_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_249, negated_conjecture, (v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_194]), c_0_121])])])).
cnf(c_0_250, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k1_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),X1)!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k4_struct_0(esk1_0)!=X1|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_134]), c_0_111])])])).
cnf(c_0_251, plain, (m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2),u1_struct_0(X1))|v2_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_66])])).
fof(c_0_252, plain, ![X75]:(v2_struct_0(X75)|~l1_struct_0(X75)|~v1_xboole_0(u1_struct_0(X75))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_253, plain, ![X47]:(~l1_algstr_0(X47)|l1_struct_0(X47)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_algstr_0])])).
cnf(c_0_254, plain, (v1_xboole_0(X1)|~v2_struct_0(g3_algstr_0(X1,X2))|~v1_funct_1(X2)|~v1_funct_2(X2,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))), inference(split_conjunct,[status(thm)],[c_0_247])).
cnf(c_0_255, negated_conjecture, (v2_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~v5_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))), inference(evalgc,[status(thm)],[c_0_248])).
cnf(c_0_256, negated_conjecture, (v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_249, c_0_216]), c_0_121])])])).
cnf(c_0_257, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),k4_struct_0(esk1_0)),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_143]), c_0_222]), c_0_111]), c_0_177]), c_0_178]), c_0_179])]), c_0_180])])).
cnf(c_0_258, plain, (m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2),u1_struct_0(X1))|v2_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[c_0_251, c_0_63])).
cnf(c_0_259, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_252])).
cnf(c_0_260, plain, (l1_struct_0(X1)|~l1_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_253])).
cnf(c_0_261, plain, (v1_xboole_0(X1)|~v1_funct_1(X2)|~v2_struct_0(g3_algstr_0(X1,X2))|~v1_funct_2(X2,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))), inference(evalgc,[status(thm)],[c_0_254])).
cnf(c_0_262, negated_conjecture, (v2_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_255, c_0_161])])])).
cnf(c_0_263, negated_conjecture, (v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_256, c_0_222])])])).
cnf(c_0_264, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_258]), c_0_121]), c_0_222])])])).
cnf(c_0_265, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_259])).
cnf(c_0_266, plain, (l1_struct_0(X1)|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_260])).
cnf(c_0_267, plain, (v1_xboole_0(u1_struct_0(X1))|~v1_funct_1(u1_algstr_0(X1))|~m1_subset_1(u1_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_261, c_0_45])])).
cnf(c_0_268, negated_conjecture, (v2_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_262, c_0_263])])])).
cnf(c_0_269, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_66]), c_0_222]), c_0_121])])])).
cnf(c_0_270, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_266])])).
cnf(c_0_271, plain, (v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(X1)|~v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_267, c_0_52]), c_0_53])).
cnf(c_0_272, negated_conjecture, (v2_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_269])])).
cnf(c_0_273, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_270, c_0_121]), c_0_180])])).
cnf(c_0_274, negated_conjecture, (~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_272]), c_0_121])]), c_0_273])])).
cnf(c_0_275, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_274, c_0_63]), c_0_121])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 276
# Proof object clause steps            : 227
# Proof object formula steps           : 49
# Proof object conjectures             : 62
# Proof object clause conjectures      : 59
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 45
# Proof object initial formulas used   : 24
# Proof object generating inferences   : 92
# Proof object simplifying inferences  : 177
# Training examples: 139 positive, 505 negative
# Training: Positive examples begin
cnf(c_0_179, negated_conjecture, (v13_algstr_0(esk1_0))).# trainpos
cnf(c_0_120, negated_conjecture, (v2_rlvect_1(esk1_0))).# trainpos
cnf(c_0_178, negated_conjecture, (v3_rlvect_1(esk1_0))).# trainpos
cnf(c_0_177, negated_conjecture, (v4_rlvect_1(esk1_0))).# trainpos
cnf(c_0_111, negated_conjecture, (l2_algstr_0(esk1_0))).# trainpos
cnf(c_0_180, negated_conjecture, (~v2_struct_0(esk1_0))).# trainpos
cnf(c_0_255, negated_conjecture, (v2_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~v5_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).# trainpos
cnf(c_0_86, plain, (v5_group_1(X1)|k6_algstr_0(X1,esk3_1(X1),esk2_1(X1))!=k6_algstr_0(X1,esk2_1(X1),esk3_1(X1))|~l3_algstr_0(X1))).# trainpos
cnf(c_0_110, plain, (l1_algstr_0(X1)|~l2_algstr_0(X1))).# trainpos
cnf(c_0_266, plain, (l1_struct_0(X1)|~l1_algstr_0(X1))).# trainpos
cnf(c_0_145, plain, (m1_subset_1(esk13_1(X1),X1))).# trainpos
cnf(c_0_53, plain, (v1_funct_1(u1_algstr_0(X1))|~l1_algstr_0(X1))).# trainpos
cnf(c_0_37, plain, (v1_funct_1(u2_algstr_0(X1))|~l3_algstr_0(X1))).# trainpos
cnf(c_0_265, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).# trainpos
cnf(c_0_138, plain, (v3_group_1(X1)|k6_algstr_0(X1,k6_algstr_0(X1,esk7_1(X1),esk8_1(X1)),esk9_1(X1))!=k6_algstr_0(X1,esk7_1(X1),k6_algstr_0(X1,esk8_1(X1),esk9_1(X1)))|~l3_algstr_0(X1))).# trainpos
cnf(c_0_195, plain, (v3_group_1(X1)|m1_subset_1(esk7_1(X1),u1_struct_0(X1))|~l3_algstr_0(X1))).# trainpos
cnf(c_0_170, plain, (v3_group_1(X1)|m1_subset_1(esk8_1(X1),u1_struct_0(X1))|~l3_algstr_0(X1))).# trainpos
cnf(c_0_129, plain, (v3_group_1(X1)|m1_subset_1(esk9_1(X1),u1_struct_0(X1))|~l3_algstr_0(X1))).# trainpos
cnf(c_0_81, plain, (v5_group_1(X1)|m1_subset_1(esk2_1(X1),u1_struct_0(X1))|~l3_algstr_0(X1))).# trainpos
cnf(c_0_80, plain, (v5_group_1(X1)|m1_subset_1(esk3_1(X1),u1_struct_0(X1))|~l3_algstr_0(X1))).# trainpos
cnf(c_0_43, plain, (g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))=X1|~l3_algstr_0(X1)|~v15_algstr_0(X1))).# trainpos
cnf(c_0_134, plain, (m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_208, plain, (v2_group_1(X1)|m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1))|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_45, plain, (v1_funct_2(u1_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~l1_algstr_0(X1))).# trainpos
cnf(c_0_29, plain, (v1_funct_2(u2_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~l3_algstr_0(X1))).# trainpos
cnf(c_0_52, plain, (m1_subset_1(u1_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_algstr_0(X1))).# trainpos
cnf(c_0_36, plain, (m1_subset_1(u2_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l3_algstr_0(X1))).# trainpos
cnf(c_0_126, plain, (m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_54, plain, (l3_algstr_0(g3_algstr_0(X1,X2))|~v1_funct_1(X2)|~v1_funct_2(X2,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))))).# trainpos
cnf(c_0_143, plain, (k1_algstr_0(X1,X2,k4_struct_0(X1))=X2|v2_struct_0(X1)|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_171, plain, (v2_group_1(X1)|k6_algstr_0(X1,X3,esk6_2(X1,X2))!=X2|k6_algstr_0(X1,esk6_2(X1,X2),X3)!=X2|k6_algstr_0(X1,X2,esk6_2(X1,X2))!=esk6_2(X1,X2)|k6_algstr_0(X1,esk6_2(X1,X2),X2)!=esk6_2(X1,X2)|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_190, plain, (k1_algstr_0(X1,k4_struct_0(X1),X2)=X2|v2_struct_0(X1)|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_98, plain, (k3_rlvect_1(X1,X2,X3)=k1_algstr_0(X1,X2,X3)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_261, plain, (v1_xboole_0(X1)|~v1_funct_1(X2)|~v2_struct_0(g3_algstr_0(X1,X2))|~v1_funct_2(X2,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))))).# trainpos
cnf(c_0_44, plain, (v15_algstr_0(g3_algstr_0(X1,X2))|~v1_funct_1(X2)|~v1_funct_2(X2,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))))).# trainpos
cnf(c_0_123, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))=k4_struct_0(X1)|v2_struct_0(X1)|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_162, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)=k4_struct_0(X1)|v2_struct_0(X1)|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_160, plain, (k6_algstr_0(X1,X2,X3)=k6_algstr_0(X1,X3,X2)|~v5_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_99, plain, (k3_rlvect_1(X1,X2,X3)=k3_rlvect_1(X1,X3,X2)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_78, plain, (k5_binop_1(u1_struct_0(X1),u1_algstr_0(X1),X2,X3)=k1_algstr_0(X1,X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_65, plain, (k5_binop_1(u1_struct_0(X1),u2_algstr_0(X1),X2,X3)=k6_algstr_0(X1,X2,X3)|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_49, plain, (X1=X2|g3_algstr_0(X3,X1)!=g3_algstr_0(X4,X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k2_zfmisc_1(X3,X3),X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X3,X3),X3))))).# trainpos
cnf(c_0_28, plain, (X1=X2|g3_algstr_0(X1,X3)!=g3_algstr_0(X2,X4)|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))))).# trainpos
cnf(c_0_122, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|~v3_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_106, plain, (k5_binop_1(X2,X1,X3,X4)=k1_binop_1(X1,X3,X4)|~v1_funct_1(X1)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v1_funct_2(X1,k2_zfmisc_1(X2,X2),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X2),X2))))).# trainpos
cnf(c_0_35, plain, (u1_struct_0(X1)=X2|g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))!=g3_algstr_0(X2,X3)|~v1_funct_1(u2_algstr_0(X1))|~m1_subset_1(u2_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l3_algstr_0(X1))).# trainpos
cnf(c_0_121, negated_conjecture, (l1_algstr_0(esk1_0))).# trainpos
cnf(c_0_270, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(X1))).# trainpos
cnf(c_0_108, plain, (k3_rlvect_1(X1,X3,X2)=k1_algstr_0(X1,X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_267, plain, (v1_xboole_0(u1_struct_0(X1))|~v1_funct_1(u1_algstr_0(X1))|~m1_subset_1(u1_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).# trainpos
cnf(c_0_203, plain, (k4_algstr_0(X1,k4_struct_0(X1))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k4_struct_0(X1)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1)))).# trainpos
cnf(c_0_59, plain, (l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v1_funct_1(u1_algstr_0(X1))|~m1_subset_1(u1_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_algstr_0(X1))).# trainpos
cnf(c_0_136, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_51, plain, (v15_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v1_funct_1(u1_algstr_0(X1))|~m1_subset_1(u1_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_algstr_0(X1))).# trainpos
cnf(c_0_133, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_algstr_0(X1,X2),X3))=k1_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_173, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~v3_rlvect_1(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_157, plain, (k5_binop_1(u1_struct_0(X1),u1_algstr_0(X1),X2,X3)=k1_binop_1(u1_algstr_0(X1),X2,X3)|~v1_funct_1(u1_algstr_0(X1))|~m1_subset_1(u1_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_algstr_0(X1))).# trainpos
cnf(c_0_55, plain, (u2_algstr_0(X1)=X2|g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))!=g3_algstr_0(X3,X2)|~v1_funct_1(u2_algstr_0(X1))|~m1_subset_1(u2_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l3_algstr_0(X1))).# trainpos
cnf(c_0_174, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k4_struct_0(X1)))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_199, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),X3))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_117, plain, (k5_binop_1(u1_struct_0(X1),u2_algstr_0(X1),X2,X3)=k1_binop_1(u2_algstr_0(X1),X2,X3)|~v1_funct_1(u2_algstr_0(X1))|~m1_subset_1(u2_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1))).# trainpos
cnf(c_0_182, plain, (k1_binop_1(u1_algstr_0(X1),X2,X3)=k1_algstr_0(X1,X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_273, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_50, plain, (u1_struct_0(X1)=X2|X1!=g3_algstr_0(X2,X3)|~l3_algstr_0(X1)|~v15_algstr_0(X1))).# trainpos
cnf(c_0_64, plain, (u2_algstr_0(X1)=X2|X1!=g3_algstr_0(X3,X2)|~l3_algstr_0(X1)|~v15_algstr_0(X1))).# trainpos
cnf(c_0_154, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_119, plain, (k1_algstr_0(X1,X3,X2)=k1_algstr_0(X1,X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_140, plain, (k1_binop_1(u2_algstr_0(X1),X2,X3)=k6_algstr_0(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1))).# trainpos
cnf(c_0_153, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))=k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1)))).# trainpos
cnf(c_0_200, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)=k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1)))).# trainpos
cnf(c_0_201, plain, (m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_57, plain, (u1_struct_0(g3_algstr_0(X1,X2))=X1|~v15_algstr_0(g3_algstr_0(X1,X2))|~l3_algstr_0(g3_algstr_0(X1,X2)))).# trainpos
cnf(c_0_67, plain, (u2_algstr_0(g3_algstr_0(X1,X2))=X2|~v15_algstr_0(g3_algstr_0(X1,X2))|~l3_algstr_0(g3_algstr_0(X1,X2)))).# trainpos
cnf(c_0_132, negated_conjecture, (k1_algstr_0(esk1_0,X1,X2)=k1_algstr_0(esk1_0,X2,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_62, plain, (u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))=u1_struct_0(X1)|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).# trainpos
cnf(c_0_217, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~v3_rlvect_1(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_183, plain, (m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).# trainpos
cnf(c_0_206, plain, (m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).# trainpos
cnf(c_0_89, plain, (m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).# trainpos
cnf(c_0_88, plain, (m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).# trainpos
cnf(c_0_139, plain, (m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).# trainpos
cnf(c_0_251, plain, (m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2),u1_struct_0(X1))|v2_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).# trainpos
cnf(c_0_69, plain, (k5_binop_1(u1_struct_0(X1),u2_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),X2,X3)=k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).# trainpos
cnf(c_0_226, negated_conjecture, (k1_algstr_0(esk1_0,X1,k4_struct_0(esk1_0))=X1|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_232, negated_conjecture, (k1_algstr_0(esk1_0,k4_struct_0(esk1_0),X1)=X1|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_224, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k1_algstr_0(X1,k4_struct_0(X1),X4)))=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_235, negated_conjecture, (k1_algstr_0(esk1_0,k1_algstr_0(esk1_0,X1,X2),X3)=k1_algstr_0(esk1_0,X1,k1_algstr_0(esk1_0,X2,X3))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_70, plain, (u2_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))=u1_algstr_0(X1)|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).# trainpos
cnf(c_0_79, plain, (k5_binop_1(u1_struct_0(X1),u1_algstr_0(X1),X2,X3)=k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_148, plain, (k1_binop_1(u1_algstr_0(X1),X2,X3)=k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2,X3)|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).# trainpos
cnf(c_0_87, plain, (k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2,X3)=k1_algstr_0(X1,X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_95, plain, (v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k1_algstr_0(X1,esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)|~m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).# trainpos
cnf(c_0_118, plain, (v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_algstr_0(X1,esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k1_algstr_0(X1,esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).# trainpos
cnf(c_0_146, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))!=k1_algstr_0(X1,k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1)))).# trainpos
cnf(c_0_141, negated_conjecture, (v5_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_151, negated_conjecture, (v5_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_161, negated_conjecture, (v5_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).# trainpos
cnf(c_0_262, negated_conjecture, (v2_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).# trainpos
cnf(c_0_172, negated_conjecture, (k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X2)=k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X1)|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).# trainpos
cnf(c_0_176, plain, (m1_subset_1(k1_algstr_0(X1,X2,k4_algstr_0(X1,X2)),u1_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_186, negated_conjecture, (k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X1)=k1_algstr_0(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).# trainpos
cnf(c_0_197, negated_conjecture, (k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X2)=k1_algstr_0(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_189, plain, (m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)),u1_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_185, plain, (v2_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2,esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2))!=esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2)|k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2),X2)!=esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2)|k1_algstr_0(X1,X3,esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2))!=X2|k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2),X3)!=X2|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).# trainpos
cnf(c_0_168, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_binop_1(u1_algstr_0(X1),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))!=k1_algstr_0(X1,k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))).# trainpos
cnf(c_0_212, plain, (m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_220, plain, (m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_181, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_binop_1(u1_algstr_0(X1),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_binop_1(u1_algstr_0(X1),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))!=k1_algstr_0(X1,k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(k1_binop_1(u1_algstr_0(X1),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))).# trainpos
cnf(c_0_193, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_binop_1(u1_algstr_0(X1),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))!=k1_algstr_0(X1,k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).# trainpos
cnf(c_0_215, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_algstr_0(X1,k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).# trainpos
cnf(c_0_230, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_algstr_0(X1,k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).# trainpos
cnf(c_0_237, negated_conjecture, (v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(k1_algstr_0(esk1_0,esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),esk9_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(k1_algstr_0(esk1_0,esk7_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))),u1_struct_0(esk1_0))|~m1_subset_1(k1_algstr_0(esk1_0,esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),esk9_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))),u1_struct_0(esk1_0))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_239, negated_conjecture, (v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(k1_algstr_0(esk1_0,esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),esk9_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))),u1_struct_0(esk1_0))|~m1_subset_1(k1_algstr_0(esk1_0,esk7_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))),u1_struct_0(esk1_0))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_241, negated_conjecture, (v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(k1_algstr_0(esk1_0,esk7_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))),u1_struct_0(esk1_0))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_243, negated_conjecture, (v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_245, negated_conjecture, (v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_249, negated_conjecture, (v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_256, negated_conjecture, (v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_207, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),X1)!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k1_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),X2)!=X1|k1_algstr_0(esk1_0,X2,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=X1|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).# trainpos
cnf(c_0_227, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),X1)!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k1_algstr_0(esk1_0,X2,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=X1|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).# trainpos
cnf(c_0_192, negated_conjecture, (k1_algstr_0(esk1_0,X1,k4_algstr_0(esk1_0,X1))=k4_struct_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_204, negated_conjecture, (m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_214, negated_conjecture, (m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_222, negated_conjecture, (m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_221, negated_conjecture, (k1_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk1_0))=k4_struct_0(esk1_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_263, negated_conjecture, (v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).# trainpos
cnf(c_0_268, negated_conjecture, (v2_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).# trainpos
cnf(c_0_219, plain, (k1_algstr_0(X1,k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2)),X2)=k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1)))).# trainpos
cnf(c_0_233, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),X1)!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k4_struct_0(esk1_0)!=X1|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(k4_algstr_0(esk1_0,k1_algstr_0(esk1_0,k4_struct_0(esk1_0),esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(k4_algstr_0(esk1_0,k1_algstr_0(esk1_0,k4_struct_0(esk1_0),esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))),u1_struct_0(esk1_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).# trainpos
cnf(c_0_236, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),X1)!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k4_struct_0(esk1_0)!=X1|~m1_subset_1(k4_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(k4_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)),u1_struct_0(esk1_0))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).# trainpos
cnf(c_0_238, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),X1)!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k4_struct_0(esk1_0)!=X1|~m1_subset_1(k4_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)),u1_struct_0(esk1_0))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).# trainpos
cnf(c_0_240, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k4_struct_0(esk1_0)!=X1|~m1_subset_1(k4_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)),u1_struct_0(esk1_0))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))))).# trainpos
cnf(c_0_244, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k1_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),X1)!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k4_struct_0(esk1_0)!=X1|~m1_subset_1(k4_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)),u1_struct_0(esk1_0))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))))).# trainpos
cnf(c_0_250, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k1_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),X1)!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k4_struct_0(esk1_0)!=X1|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).# trainpos
cnf(c_0_257, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),k4_struct_0(esk1_0)),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).# trainpos
cnf(c_0_264, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).# trainpos
cnf(c_0_269, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).# trainpos
cnf(c_0_272, negated_conjecture, (v2_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).# trainpos
cnf(c_0_274, negated_conjecture, (~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_101, plain, (v1_xboole_0(esk15_0))).#trainneg
cnf(i_0_88, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_98, plain, (v1_xreal_0(esk15_0))).#trainneg
cnf(i_0_100, plain, (v1_xcmplx_0(esk15_0))).#trainneg
cnf(i_0_99, plain, (v1_xxreal_0(esk15_0))).#trainneg
cnf(i_0_106, plain, (v3_membered(k1_numbers))).#trainneg
cnf(i_0_81, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_85, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_62, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_87, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_86, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_78, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_37, plain, (l2_struct_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_45, plain, (l1_struct_0(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_84, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_53, plain, (l1_struct_0(X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_48, plain, (u2_struct_0(X1)=k4_struct_0(X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_89, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_95, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_79, plain, (v9_struct_0(esk14_1(X1),X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_77, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_49, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_76, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_82, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_105, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_71, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_72, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_97, plain, (v10_algstr_0(X2,X1)|~v13_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_60, plain, (m1_subset_1(u2_struct_0(X1),u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_80, plain, (m1_subset_1(esk14_1(X1),u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_83, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_94, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_93, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_32, plain, (v3_rlvect_1(X1)|m1_subset_1(esk10_1(X1),u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_31, plain, (v3_rlvect_1(X1)|m1_subset_1(esk11_1(X1),u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_30, plain, (v3_rlvect_1(X1)|m1_subset_1(esk12_1(X1),u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_23, plain, (m1_subset_1(esk4_1(X1),u1_struct_0(X1))|~v2_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_54, plain, (v2_struct_0(X1)|v2_binop_1(u2_algstr_0(X1),u1_struct_0(X1))|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_29, plain, (v3_rlvect_1(X1)|k1_algstr_0(X1,k1_algstr_0(X1,esk10_1(X1),esk11_1(X1)),esk12_1(X1))!=k1_algstr_0(X1,esk10_1(X1),k1_algstr_0(X1,esk11_1(X1),esk12_1(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_57, plain, (v2_struct_0(X1)|v1_setwiseo(u2_algstr_0(X1),u1_struct_0(X1))|~l3_algstr_0(X1)|~v1_group_1(X1))).#trainneg
cnf(i_0_75, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_103, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_104, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_74, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_90, plain, (v2_struct_0(X1)|v1_finseqop(u2_algstr_0(X1),u1_struct_0(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_22, plain, (k6_algstr_0(X1,X2,esk4_1(X1))=X2|~v2_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_21, plain, (k6_algstr_0(X1,esk4_1(X1),X2)=X2|~v2_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_20, plain, (m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|~v2_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_19, plain, (k6_algstr_0(X1,X2,esk5_2(X1,X2))=esk4_1(X1)|~v2_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_18, plain, (k6_algstr_0(X1,esk5_2(X1,X2),X2)=esk4_1(X1)|~v2_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_47, plain, (m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_73, plain, (m1_subset_1(k3_rlvect_1(X1,X2,X3),u1_struct_0(X1))|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_28, plain, (k6_algstr_0(X1,k6_algstr_0(X1,X2,X3),X4)=k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4))|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_188, plain, (u1_struct_0(X1)=X2|g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))!=g3_algstr_0(X2,X3)|~v1_funct_1(u1_algstr_0(X1))|~m1_subset_1(u1_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_109, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_107, plain, (k1_xboole_0=esk15_0)).#trainneg
cnf(i_0_200, plain, (v2_group_1(X1)|k6_algstr_0(X1,X2,esk6_2(X1,X2))!=esk6_2(X1,X2)|k6_algstr_0(X1,esk6_2(X1,X2),X2)!=esk6_2(X1,X2)|k6_algstr_0(X1,X3,k6_algstr_0(X1,X4,esk6_2(X1,X2)))!=X2|k6_algstr_0(X1,esk6_2(X1,X2),k6_algstr_0(X1,X3,X4))!=X2|~m1_subset_1(k6_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_group_1(X1))).#trainneg
cnf(i_0_118, plain, (v1_xreal_0(np__1))).#trainneg
cnf(i_0_116, plain, (v1_xreal_0(esk13_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_145, plain, (v1_relat_1(u1_algstr_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_112, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_111, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_113, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_algstr_0(X1))).#trainneg
cnf(i_0_186, plain, (u1_algstr_0(X1)=X2|g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))!=g3_algstr_0(X3,X2)|~v1_funct_1(u1_algstr_0(X1))|~m1_subset_1(u1_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_114, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_115, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_151, plain, (v1_relat_1(u2_algstr_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_155, plain, (v5_relat_1(u2_algstr_0(X1),u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_119, plain, (v1_xreal_0(k4_struct_0(X1))|~v3_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_157, plain, (esk5_2(X1,esk4_1(X1))=esk4_1(X1)|~m1_subset_1(esk5_2(X1,esk4_1(X1)),u1_struct_0(X1))|~l3_algstr_0(X1)|~v2_group_1(X1)|~m1_subset_1(esk4_1(X1),u1_struct_0(X1)))).#trainneg
cnf(i_0_120, plain, (v1_relat_1(esk13_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_123, plain, (v1_xreal_0(u2_struct_0(X1))|~v3_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_125, plain, (v1_xreal_0(esk14_1(X1))|~v3_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_148, plain, (v1_xboole_0(u1_algstr_0(X1))|~v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_149, plain, (v5_relat_1(u1_algstr_0(X1),u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_154, plain, (v1_xboole_0(u2_algstr_0(X1))|~v1_xboole_0(u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_127, plain, (v4_relat_1(esk13_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_128, plain, (v5_relat_1(esk13_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_189, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k4_algstr_0(X1,k1_algstr_0(X1,X2,X3))))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,X2,X3)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_168, plain, (v15_algstr_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))|~v1_funct_1(u2_algstr_0(X1))|~m1_subset_1(u2_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_191, plain, (m1_subset_1(k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4)),u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1))).#trainneg
cnf(i_0_137, plain, (v1_xreal_0(esk4_1(X1))|~v3_membered(u1_struct_0(X1))|~l3_algstr_0(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_132, plain, (v1_xreal_0(esk7_1(X1))|v3_group_1(X1)|~v3_membered(u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_129, plain, (v1_xreal_0(esk10_1(X1))|v3_rlvect_1(X1)|~v3_membered(u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_130, plain, (v1_xreal_0(esk11_1(X1))|v3_rlvect_1(X1)|~v3_membered(u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_131, plain, (v1_xreal_0(esk12_1(X1))|v3_rlvect_1(X1)|~v3_membered(u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_146, plain, (v1_xboole_0(u1_algstr_0(X1))|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_152, plain, (v1_xboole_0(u2_algstr_0(X1))|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_133, plain, (v1_xreal_0(esk8_1(X1))|v3_group_1(X1)|~v3_membered(u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_134, plain, (v1_xreal_0(esk9_1(X1))|v3_group_1(X1)|~v3_membered(u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_135, plain, (v1_xreal_0(esk2_1(X1))|v5_group_1(X1)|~v3_membered(u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_136, plain, (v1_xreal_0(esk3_1(X1))|v5_group_1(X1)|~v3_membered(u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_166, plain, (v1_xboole_0(u1_struct_0(X1))|~v1_funct_1(u2_algstr_0(X1))|~m1_subset_1(u2_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~v2_struct_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_227, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_223, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1))).#trainneg
cnf(i_0_62, plain, (X1=esk15_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_229, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_228, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_139, plain, (v1_xboole_0(esk13_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_144, plain, (v1_xreal_0(u1_algstr_0(X1))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_140, plain, (v1_xboole_0(esk13_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_142, plain, (v1_xreal_0(k4_algstr_0(X1,X2))|~v3_membered(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1))).#trainneg
cnf(i_0_232, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_231, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_230, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_143, plain, (v1_xreal_0(esk6_2(X1,X2))|v2_group_1(X1)|~v3_membered(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_233, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_150, plain, (v1_xreal_0(u2_algstr_0(X1))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_156, plain, (v1_xreal_0(esk5_2(X1,X2))|~v3_membered(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_147, plain, (v4_relat_1(u1_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_170, plain, (v1_xreal_0(k6_algstr_0(X1,X2,X3))|~v3_membered(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_153, plain, (v4_relat_1(u2_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_159, plain, (v1_xreal_0(k1_algstr_0(X1,X2,X3))|~v3_membered(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_268, plain, (k1_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,k4_algstr_0(X1,X2)))=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k4_algstr_0(X1,X2)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_290, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X3,k4_algstr_0(X1,k1_algstr_0(X1,X2,X3))))=k1_algstr_0(X1,k4_algstr_0(X1,X2),k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,k4_algstr_0(X1,k1_algstr_0(X1,X2,X3))),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,X2,X3)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_368, plain, (v1_xreal_0(k4_struct_0(X1))|v2_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_160, plain, (l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))|~v1_funct_1(u2_algstr_0(X1))|~m1_subset_1(u2_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_371, plain, (k2_zfmisc_1(X1,X2)=esk15_0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_361, plain, (u1_struct_0(X1)=esk15_0|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_359, plain, (u2_algstr_0(X1)=esk15_0|~v1_xboole_0(u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_360, plain, (u1_algstr_0(X1)=esk15_0|~v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_403, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))=k1_algstr_0(X1,k4_algstr_0(X1,X3),X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_404, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)=k1_algstr_0(X1,k4_algstr_0(X1,X3),X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_502, plain, (k2_zfmisc_1(X1,esk15_0)=esk15_0)).#trainneg
cnf(i_0_395, plain, (m1_subset_1(k1_algstr_0(X1,k4_algstr_0(X1,X2),X2),u1_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_516, plain, (l3_algstr_0(g3_algstr_0(u1_struct_0(X1),esk15_0))|~l3_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_538, plain, (l3_algstr_0(g3_algstr_0(u1_struct_0(X1),esk15_0))|~l1_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_519, plain, (esk15_0=X2|g3_algstr_0(u1_struct_0(X1),esk15_0)!=g3_algstr_0(X3,X2)|~l3_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_175, plain, (v1_xreal_0(k3_rlvect_1(X1,X2,X3))|~v3_membered(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_530, plain, (v1_funct_1(esk15_0)|~l1_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_510, plain, (v1_funct_1(esk15_0)|~l3_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_520, plain, (v1_relat_1(esk15_0)|~l3_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_534, plain, (v1_relat_1(esk15_0)|~l1_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_511, plain, (g3_algstr_0(u1_struct_0(X1),esk15_0)=X1|~v15_algstr_0(X1)|~l3_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_192, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4)),X5)=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),k1_algstr_0(X1,X4,X5))|~l1_algstr_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~v3_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_690, plain, (m1_subset_1(k4_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~l2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_692, plain, (m1_subset_1(esk14_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~l2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_686, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk15_0)))).#trainneg
cnf(i_0_563, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)=k4_algstr_0(X1,k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k4_struct_0(X1)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_688, plain, (v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v1_xboole_0(u1_struct_0(X1))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_677, plain, (v1_relat_1(esk15_0))).#trainneg
cnf(i_0_681, plain, (v5_relat_1(esk15_0,esk15_0))).#trainneg
cnf(i_0_198, plain, (m1_subset_1(k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_group_1(X1))).#trainneg
cnf(i_0_679, plain, (v4_relat_1(esk15_0,X1))).#trainneg
cnf(i_0_672, plain, (v1_relat_1(esk13_1(k1_zfmisc_1(esk15_0))))).#trainneg
cnf(i_0_723, plain, (v5_relat_1(u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_729, plain, (l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_199, plain, (k6_algstr_0(X1,k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4)),X5)=k6_algstr_0(X1,k6_algstr_0(X1,X2,X3),k6_algstr_0(X1,X4,X5))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_201, plain, (k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,esk5_2(X1,k6_algstr_0(X1,X2,X3))))=esk4_1(X1)|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~v2_group_1(X1)|~m1_subset_1(esk5_2(X1,k6_algstr_0(X1,X2,X3)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_group_1(X1))).#trainneg
cnf(i_0_924, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)=k4_algstr_0(X1,k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1)))).#trainneg
cnf(i_0_691, plain, (m1_subset_1(u2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~l2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_798, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),k1_algstr_0(X1,X4,X5))=k1_algstr_0(X1,X2,k1_algstr_0(X1,k1_algstr_0(X1,X3,X4),X5))|~l1_algstr_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_338, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_algstr_0(X1,X2),X3))=k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,k4_struct_0(X1),X3))|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_769, plain, (v1_funct_1(esk15_0)|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_676, plain, (v1_xboole_0(esk13_1(k1_zfmisc_1(esk15_0))))).#trainneg
cnf(i_0_541, plain, (esk15_0=X2|g3_algstr_0(u1_struct_0(X1),esk15_0)!=g3_algstr_0(X3,X2)|~l1_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_202, plain, (k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,esk4_1(X1)))=k6_algstr_0(X1,X2,X3)|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~v2_group_1(X1)|~m1_subset_1(esk4_1(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_group_1(X1))).#trainneg
cnf(i_0_687, plain, (v5_relat_1(X1,esk15_0)|~m1_subset_1(X1,k1_zfmisc_1(esk15_0)))).#trainneg
cnf(i_0_673, plain, (v5_relat_1(esk13_1(k1_zfmisc_1(esk15_0)),esk15_0))).#trainneg
cnf(i_0_733, plain, (v15_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_928, plain, (v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_204, plain, (k6_algstr_0(X1,esk5_2(X1,X2),k6_algstr_0(X1,X2,X3))=k6_algstr_0(X1,esk4_1(X1),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_205, plain, (k6_algstr_0(X1,esk4_1(X1),k6_algstr_0(X1,X2,X3))=k6_algstr_0(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(esk4_1(X1),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1190, plain, (k2_zfmisc_1(X1,esk13_1(k1_zfmisc_1(esk15_0)))=esk15_0)).#trainneg
cnf(i_0_1189, plain, (esk13_1(k1_zfmisc_1(esk15_0))=esk15_0)).#trainneg
cnf(i_0_1083, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4))=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_547, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))=k4_algstr_0(X1,k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k4_struct_0(X1)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1)))).#trainneg
cnf(i_0_445, plain, (k4_algstr_0(X1,k4_algstr_0(X1,X2))=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k4_algstr_0(X1,X2)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_325, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_algstr_0(X1,X2),X3))=k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X3,k4_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,k4_struct_0(X1)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_1351, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))=k4_algstr_0(X1,k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1)))).#trainneg
cnf(i_0_927, plain, (v2_struct_0(g3_algstr_0(u1_struct_0(X1),esk15_0))|~v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_683, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk15_0)))).#trainneg
cnf(i_0_206, plain, (k6_algstr_0(X1,X2,k6_algstr_0(X1,esk5_2(X1,X2),X3))=k6_algstr_0(X1,esk4_1(X1),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_417, plain, (k2_zfmisc_1(X1,X2)=esk15_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_742, plain, (u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))=u1_struct_0(X1)|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_207, plain, (k6_algstr_0(X1,X2,k6_algstr_0(X1,esk4_1(X1),X3))=k6_algstr_0(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(esk4_1(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_217, plain, (v2_group_1(X1)|k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,k6_algstr_0(X1,X4,esk6_2(X1,X5))))!=X5|k6_algstr_0(X1,esk6_2(X1,X5),k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4)))!=X5|k6_algstr_0(X1,X5,esk6_2(X1,X5))!=esk6_2(X1,X5)|k6_algstr_0(X1,esk6_2(X1,X5),X5)!=esk6_2(X1,X5)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~m1_subset_1(esk6_2(X1,X5),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_1520, plain, (k2_zfmisc_1(esk15_0,X1)=esk15_0)).#trainneg
cnf(i_0_1304, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|v2_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1303, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1355, plain, (k4_algstr_0(X1,k4_algstr_0(X1,X2))=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_526, plain, (u1_struct_0(X1)=X2|g3_algstr_0(u1_struct_0(X1),esk15_0)!=g3_algstr_0(X2,X3)|~l3_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_1659, plain, (v5_relat_1(esk15_0,X1))).#trainneg
cnf(i_0_542, plain, (u1_struct_0(X1)=X2|g3_algstr_0(u1_struct_0(X1),esk15_0)!=g3_algstr_0(X2,X3)|~l1_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_1671, plain, (v4_relat_1(X1,esk15_0)|~m1_subset_1(X1,k1_zfmisc_1(esk15_0)))).#trainneg
cnf(i_0_940, plain, (v5_relat_1(u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))),u1_struct_0(X1))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_220, plain, (v2_group_1(X1)|k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,k6_algstr_0(X1,X4,esk6_2(X1,X5))))!=X5|k6_algstr_0(X1,esk6_2(X1,X5),k6_algstr_0(X1,k6_algstr_0(X1,X2,X3),X4))!=X5|k6_algstr_0(X1,X5,esk6_2(X1,X5))!=esk6_2(X1,X5)|k6_algstr_0(X1,esk6_2(X1,X5),X5)!=esk6_2(X1,X5)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~m1_subset_1(k6_algstr_0(X1,X4,esk6_2(X1,X5)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_297, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X2)))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_400, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_algstr_0(X1,X3),X3))=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_241, plain, (esk5_2(X1,esk4_1(X1))=esk4_1(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~m1_subset_1(esk4_1(X1),u1_struct_0(X1)))).#trainneg
cnf(i_0_685, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(esk15_0)))).#trainneg
cnf(i_0_730, plain, (v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_262, plain, (k1_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),X2)=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1)))).#trainneg
cnf(i_0_263, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k1_algstr_0(X1,k4_algstr_0(X1,k1_algstr_0(X1,X2,X3)),X4)))=k1_algstr_0(X1,k4_struct_0(X1),X4)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_algstr_0(X1,k1_algstr_0(X1,X2,X3)),X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_308, plain, (m1_subset_1(k1_algstr_0(X1,X2,k4_struct_0(X1)),u1_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1527, plain, (k2_zfmisc_1(u1_struct_0(X1),X2)=esk15_0|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1673, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(esk15_0)))).#trainneg
cnf(i_0_1965, negated_conjecture, (k1_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk1_0))=k4_algstr_0(esk1_0,k4_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(k4_algstr_0(esk1_0,k4_struct_0(esk1_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_509, plain, (k2_zfmisc_1(X1,u1_struct_0(X2))=esk15_0|~l3_algstr_0(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_1963, negated_conjecture, (k4_algstr_0(esk1_0,k4_struct_0(esk1_0))=k4_struct_0(esk1_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(k4_algstr_0(esk1_0,k4_struct_0(esk1_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_945, plain, (l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_264, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,k4_algstr_0(X1,k4_algstr_0(X1,X2)),X3))=k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),X3))|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_algstr_0(X1,k4_algstr_0(X1,X2)),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_2202, plain, (m1_subset_1(k1_algstr_0(X1,X2,k4_struct_0(X1)),u1_struct_0(X1))|v2_struct_0(X1)|v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~l1_algstr_0(X1))).#trainneg
cnf(i_0_2192, plain, (m1_subset_1(k1_algstr_0(X1,X2,k4_struct_0(X1)),u1_struct_0(X1))|v2_struct_0(X1)|v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~l1_algstr_0(X1))).#trainneg
cnf(i_0_1914, plain, (k1_algstr_0(X1,k4_algstr_0(X1,k4_algstr_0(X1,X2)),k4_struct_0(X1))=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,k4_algstr_0(X1,X2)),u1_struct_0(X1)))).#trainneg
cnf(i_0_735, plain, (u1_struct_0(X1)=X2|g3_algstr_0(u1_struct_0(X1),u2_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=g3_algstr_0(X2,X3)|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_2253, negated_conjecture, (k4_algstr_0(esk1_0,k4_struct_0(esk1_0))=k4_struct_0(esk1_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1493, plain, (u1_struct_0(g3_algstr_0(u1_struct_0(X1),esk15_0))=esk15_0|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),esk15_0))|~v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_270, plain, (k1_algstr_0(X1,X2,k4_struct_0(X1))=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_669, plain, (esk15_0=X1|g3_algstr_0(esk15_0,X2)!=g3_algstr_0(X1,X3)|~v1_funct_2(X2,esk15_0,esk15_0)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(esk15_0)))).#trainneg
cnf(i_0_1213, plain, (v15_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_266, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X2),X3),X4))=k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3),X4)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,k4_algstr_0(X1,X2),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_272, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4))=k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3),X4)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~v3_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2493, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2391, plain, (k4_algstr_0(X1,k4_algstr_0(X1,X2))=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k4_algstr_0(X1,X2)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2526, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),k4_algstr_0(X1,X3))=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3046, plain, (k4_algstr_0(X1,k4_algstr_0(X1,X2))=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_2555, negated_conjecture, (k1_algstr_0(esk1_0,k4_algstr_0(esk1_0,X1),X1)=k4_struct_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2536, negated_conjecture, (k1_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk1_0))=k4_struct_0(esk1_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2605, plain, (esk15_0=X2|g3_algstr_0(esk15_0,u1_algstr_0(g3_algstr_0(u1_struct_0(X1),esk15_0)))!=g3_algstr_0(X2,X3)|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),esk15_0))|~v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_275, plain, (k1_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),X2)=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_273, plain, (k1_algstr_0(X1,k4_algstr_0(X1,k4_algstr_0(X1,X2)),k1_algstr_0(X1,k4_struct_0(X1),X3))=k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_274, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),X3))=k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_2422, plain, (k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),X3)=k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),X3))|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_3226, negated_conjecture, (k1_algstr_0(esk1_0,X1,k4_algstr_0(esk1_0,X1))=k4_struct_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_711, plain, (l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_277, plain, (k1_algstr_0(X1,k4_algstr_0(X1,k4_algstr_0(X1,X2)),k4_struct_0(X1))=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1219, plain, (v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))))|~v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_276, plain, (k1_algstr_0(X1,k4_algstr_0(X1,k1_algstr_0(X1,X2,X3)),k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4)))=k1_algstr_0(X1,k4_struct_0(X1),X4)|v2_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1318, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,k4_struct_0(X1),X3))=k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_4067, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)))=k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1)))).#trainneg
cnf(i_0_3689, plain, (l3_algstr_0(g3_algstr_0(u1_struct_0(X1),esk15_0))|~l1_algstr_0(X1)|~v1_xboole_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))).#trainneg
cnf(i_0_3234, plain, (esk15_0=X1|g3_algstr_0(esk15_0,u1_algstr_0(g3_algstr_0(esk15_0,esk15_0)))!=g3_algstr_0(X1,X3)|~l1_algstr_0(g3_algstr_0(esk15_0,esk15_0))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X2),esk15_0))|~v1_xboole_0(u1_struct_0(X2))|~l1_algstr_0(X2))).#trainneg
cnf(i_0_682, plain, (X1=X2|g3_algstr_0(esk15_0,X1)!=g3_algstr_0(X3,X2)|~v1_funct_2(X1,esk15_0,esk15_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk15_0)))).#trainneg
cnf(i_0_280, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,k4_struct_0(X1),X3))=k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,k4_algstr_0(X1,X2),X3))|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_algstr_0(X1,X2),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_283, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2)))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2))),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1)))).#trainneg
cnf(i_0_4346, plain, (k5_binop_1(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(X1),X2,X3)=k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2,X3)|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_4350, plain, (v1_xboole_0(u1_algstr_0(X1))|~v1_xboole_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_4355, plain, (k5_binop_1(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(X1),X2,X3)=k1_binop_1(u1_algstr_0(X1),X2,X3)|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_713, plain, (m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_4337, plain, (l3_algstr_0(g3_algstr_0(u1_struct_0(X1),esk15_0))|~l1_algstr_0(X1)|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))).#trainneg
cnf(i_0_302, plain, (v1_xreal_0(k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4)))|~v3_membered(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1))).#trainneg
cnf(i_0_284, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k1_algstr_0(X1,X4,k4_algstr_0(X1,k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)))))=k4_struct_0(X1)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X4,k4_algstr_0(X1,k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4))),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_285, plain, (k1_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_4577, plain, (v1_xboole_0(u1_algstr_0(X1))|~l1_algstr_0(X1)|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))).#trainneg
cnf(i_0_4516, plain, (k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),X2,X3)=k1_binop_1(u1_algstr_0(X1),X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_525, plain, (v15_algstr_0(g3_algstr_0(u1_struct_0(X1),esk15_0))|~l3_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_365, plain, (esk13_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=esk15_0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_4351, plain, (v5_relat_1(u1_algstr_0(X1),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_4344, plain, (l3_algstr_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_286, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X3,k4_algstr_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X2),X3))))=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,k4_algstr_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X2),X3))),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X2),X3)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_5231, plain, (v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_binop_1(u1_algstr_0(X1),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k1_algstr_0(X1,esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))).#trainneg
cnf(i_0_5233, plain, (v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_binop_1(u1_algstr_0(X1),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k1_algstr_0(X1,esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).#trainneg
cnf(i_0_4534, plain, (v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_binop_1(u1_algstr_0(X1),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)|~m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))).#trainneg
cnf(i_0_4953, plain, (k4_algstr_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X2),X2))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_4990, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))!=k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).#trainneg
cnf(i_0_2604, plain, (esk15_0=X2|g3_algstr_0(esk15_0,u2_algstr_0(g3_algstr_0(u1_struct_0(X1),esk15_0)))!=g3_algstr_0(X2,X3)|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),esk15_0))|~v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_366, plain, (esk13_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=esk15_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_667, plain, (l3_algstr_0(g3_algstr_0(esk15_0,X1))|~v1_funct_2(X1,esk15_0,esk15_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk15_0)))).#trainneg
cnf(i_0_2547, negated_conjecture, (m1_subset_1(k1_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk1_0)),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4353, plain, (v15_algstr_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_4340, plain, (v2_binop_1(u1_algstr_0(X1),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_288, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k1_algstr_0(X1,X3,k4_algstr_0(X1,k1_algstr_0(X1,X2,X3))),X4))=k1_algstr_0(X1,k4_struct_0(X1),X4)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,X3,k4_algstr_0(X1,k1_algstr_0(X1,X2,X3))),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,X2,X3)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_289, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k1_algstr_0(X1,X3,k4_algstr_0(X1,k1_algstr_0(X1,X2,X3))),k4_algstr_0(X1,k4_struct_0(X1))))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k4_struct_0(X1)),u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,X3,k4_algstr_0(X1,k1_algstr_0(X1,X2,X3))),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,X2,X3)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_5283, plain, (v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_binop_1(u1_algstr_0(X1),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k1_algstr_0(X1,esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).#trainneg
cnf(i_0_5290, plain, (v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))).#trainneg
cnf(i_0_5282, plain, (v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_binop_1(u1_algstr_0(X1),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k1_binop_1(u1_algstr_0(X1),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))).#trainneg
cnf(i_0_5266, plain, (v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k5_binop_1(u1_struct_0(X1),u1_algstr_0(X1),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k1_algstr_0(X1,esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).#trainneg
cnf(i_0_2616, plain, (l3_algstr_0(g3_algstr_0(esk15_0,esk15_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),esk15_0))|~v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_5421, plain, (v2_binop_1(u1_algstr_0(X1),u1_struct_0(X1))|v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)|~v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))).#trainneg
cnf(i_0_668, plain, (v15_algstr_0(g3_algstr_0(esk15_0,X1))|~v1_funct_2(X1,esk15_0,esk15_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk15_0)))).#trainneg
cnf(i_0_5391, negated_conjecture, (v1_xreal_0(k1_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk1_0)))|~v3_membered(u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4341, plain, (v1_setwiseo(u1_algstr_0(X1),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v1_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_4339, plain, (g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(X1))=g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))|~v15_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_291, plain, (k1_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_292, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X2,k4_algstr_0(X1,X2)))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_5759, plain, (v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k1_algstr_0(X1,esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))).#trainneg
cnf(i_0_5378, plain, (esk15_0=X1|g3_algstr_0(esk15_0,u2_algstr_0(g3_algstr_0(esk15_0,esk15_0)))!=g3_algstr_0(X1,X3)|~l1_algstr_0(g3_algstr_0(u1_struct_0(X2),esk15_0))|~v1_xboole_0(u1_struct_0(X2))|~l1_algstr_0(X2))).#trainneg
cnf(i_0_5828, plain, (v1_setwiseo(u1_algstr_0(X1),u1_struct_0(X1))|v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v1_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_293, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),k1_algstr_0(X1,X4,k4_algstr_0(X1,k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4)))))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_294, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,X2,k4_algstr_0(X1,k4_struct_0(X1))))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k4_struct_0(X1)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_6350, negated_conjecture, (k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X1),X3)=k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X3))|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).#trainneg
cnf(i_0_5787, plain, (v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k1_binop_1(u1_algstr_0(X1),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))).#trainneg
cnf(i_0_6347, negated_conjecture, (k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X3,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X2))=k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X3))|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X2),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))))).#trainneg
cnf(i_0_5092, plain, (m1_subset_1(esk15_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_5103, plain, (v5_relat_1(u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_4342, plain, (v1_finseqop(u1_algstr_0(X1),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v2_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_295, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3))))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3)),u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_296, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_algstr_0(X1,X2),k4_algstr_0(X1,k4_struct_0(X1))))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k4_struct_0(X1)),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_6443, negated_conjecture, (k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X3))=k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X3))|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).#trainneg
cnf(i_0_6648, plain, (v1_finseqop(u1_algstr_0(X1),u1_struct_0(X1))|v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)|~v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v2_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))).#trainneg
cnf(i_0_378, plain, (k1_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),k1_algstr_0(X1,X2,k4_algstr_0(X1,X2)))=k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_5110, plain, (l3_algstr_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_298, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X2),X3),k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3))))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3)),u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,k4_algstr_0(X1,X2),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_299, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,k1_algstr_0(X1,X2,k4_struct_0(X1))))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,X2,k4_struct_0(X1))),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_6447, negated_conjecture, (k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X3))=k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X3,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X1))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X3),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).#trainneg
cnf(i_0_6947, negated_conjecture, (k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X3),X1)=k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X3))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X3),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).#trainneg
cnf(i_0_5757, plain, (v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k5_binop_1(u1_struct_0(X1),u1_algstr_0(X1),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k1_algstr_0(X1,esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).#trainneg
cnf(i_0_6945, negated_conjecture, (k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X3))=k1_algstr_0(esk1_0,X1,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X3))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X3),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~v3_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).#trainneg
cnf(i_0_4347, plain, (u1_algstr_0(X1)=X2|g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(X1))!=g3_algstr_0(X3,X2)|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_5382, plain, (m1_subset_1(esk15_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5806, plain, (v2_binop_1(u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_300, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k1_algstr_0(X1,X4,k4_algstr_0(X1,k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))))))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))),u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_305, plain, (m1_subset_1(k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k1_algstr_0(X1,X4,X5))),u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_6629, plain, (m1_subset_1(k4_algstr_0(X1,k4_struct_0(X1)),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2529, plain, (u1_struct_0(X1)=X2|g3_algstr_0(u1_struct_0(X1),u2_algstr_0(g3_algstr_0(u1_struct_0(X1),esk15_0)))!=g3_algstr_0(X2,X3)|~l1_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_2349, plain, (m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)),u1_struct_0(X1))|v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_7688, plain, (esk15_0=X2|g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),esk15_0)),esk15_0)!=g3_algstr_0(X3,X2)|~l1_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_383, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))=k1_algstr_0(X1,X3,k4_algstr_0(X1,X3))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_8235, plain, (esk15_0=X1|g3_algstr_0(esk15_0,esk15_0)!=g3_algstr_0(X3,X1)|~v1_xboole_0(u1_struct_0(X2))|~l1_algstr_0(X2))).#trainneg
cnf(i_0_8191, negated_conjecture, (m1_subset_1(k4_algstr_0(esk1_0,k4_struct_0(esk1_0)),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5414, plain, (v15_algstr_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_5409, plain, (u1_struct_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(X1)))=u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_307, plain, (m1_subset_1(k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),X4)),u1_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_8164, plain, (m1_subset_1(k4_algstr_0(X1,k4_struct_0(X1)),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_8403, negated_conjecture, (v1_xreal_0(k4_algstr_0(esk1_0,k4_struct_0(esk1_0)))|~v3_membered(u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_384, plain, (k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))=k4_algstr_0(X1,k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k4_struct_0(X1)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1)))).#trainneg
cnf(i_0_8430, plain, (v2_struct_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(X1)))|~v1_xboole_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_8709, plain, (v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v1_xboole_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_8621, plain, (v1_xreal_0(k4_algstr_0(X1,k4_struct_0(X1)))|v2_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_388, plain, (k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2637, plain, (v2_struct_0(g3_algstr_0(esk15_0,esk15_0))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),esk15_0))|~v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_8465, plain, (v5_relat_1(u1_algstr_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_322, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k1_algstr_0(X1,X4,k4_struct_0(X1))))=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X4,k4_struct_0(X1)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_8625, plain, (m1_subset_1(k1_algstr_0(X1,X2,k4_struct_0(X1)),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_8757, negated_conjecture, (m1_subset_1(k1_algstr_0(esk1_0,X2,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).#trainneg
cnf(i_0_743, plain, (u1_struct_0(g3_algstr_0(u1_struct_0(X1),esk15_0))=u1_struct_0(X1)|~l1_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_8232, plain, (esk15_0=X1|g3_algstr_0(u1_struct_0(g3_algstr_0(esk15_0,esk15_0)),esk15_0)!=g3_algstr_0(X3,X1)|~l1_algstr_0(g3_algstr_0(u1_struct_0(X2),esk15_0))|~v1_xboole_0(u1_struct_0(X2))|~l1_algstr_0(X2))).#trainneg
cnf(i_0_323, plain, (k1_algstr_0(X1,k4_algstr_0(X1,k1_algstr_0(X1,X2,k4_struct_0(X1))),X2)=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,k4_struct_0(X1)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,X2,k4_struct_0(X1))),u1_struct_0(X1)))).#trainneg
cnf(i_0_5288, plain, (v5_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k5_binop_1(u1_struct_0(X1),u1_algstr_0(X1),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).#trainneg
cnf(i_0_334, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,k4_struct_0(X1)))=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k4_struct_0(X1)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1)))).#trainneg
cnf(i_0_398, plain, (k1_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),k1_algstr_0(X1,k4_algstr_0(X1,X2),X2))=k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_8470, plain, (l3_algstr_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(X1)))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_327, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k1_algstr_0(X1,X3,k4_struct_0(X1)),X4))=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,X3,k4_struct_0(X1)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_328, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k1_algstr_0(X1,X3,k4_struct_0(X1)),k4_algstr_0(X1,k1_algstr_0(X1,X2,X3))))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,X2,X3)),u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,X3,k4_struct_0(X1)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_9648, plain, (k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_5089, plain, (v15_algstr_0(g3_algstr_0(esk15_0,esk15_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),esk15_0))|~v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_329, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X3,k4_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,k4_struct_0(X1)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_10408, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,k4_algstr_0(X1,X4),k1_algstr_0(X1,X4,X3))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_4352, plain, (v1_xboole_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v2_struct_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_333, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k1_algstr_0(X1,X4,k4_struct_0(X1))))=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_336, plain, (k1_algstr_0(X1,k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2)),X2)=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2)),u1_struct_0(X1)))).#trainneg
cnf(i_0_10981, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k5_binop_1(u1_struct_0(X1),u1_algstr_0(X1),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))!=k1_algstr_0(X1,k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).#trainneg
cnf(i_0_10993, plain, (v1_xboole_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))).#trainneg
cnf(i_0_10975, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_binop_1(u1_algstr_0(X1),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))!=k1_algstr_0(X1,k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).#trainneg
cnf(i_0_10983, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_algstr_0(X1,k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))).#trainneg
cnf(i_0_10395, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X3),k1_algstr_0(X1,X3,X2))=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_9646, plain, (v1_xreal_0(k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)))|v2_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_5810, plain, (u1_struct_0(g3_algstr_0(esk15_0,X1))=esk15_0|~l3_algstr_0(g3_algstr_0(esk15_0,X1))|~v1_funct_2(X1,esk15_0,esk15_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk15_0)))).#trainneg
cnf(i_0_11362, plain, (v1_xboole_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_340, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3),X4))=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_11352, plain, (u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))=esk15_0|~l1_algstr_0(X1)|~v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))).#trainneg
cnf(i_0_416, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)))=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_540, plain, (v15_algstr_0(g3_algstr_0(u1_struct_0(X1),esk15_0))|~l1_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_5840, plain, (u1_struct_0(X1)=u1_struct_0(g3_algstr_0(u1_struct_0(X2),u1_algstr_0(X2)))|g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))!=g3_algstr_0(u1_struct_0(X2),u1_algstr_0(X2))|~l1_algstr_0(X1)|~l1_algstr_0(X2))).#trainneg
cnf(i_0_341, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3),k4_algstr_0(X1,k1_algstr_0(X1,X2,X3))))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,X2,X3)),u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_11920, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~v3_rlvect_1(X1))).#trainneg
cnf(i_0_9516, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_binop_1(u1_algstr_0(X1),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))!=k1_algstr_0(X1,k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).#trainneg
cnf(i_0_5850, plain, (g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))=g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_8474, plain, (v15_algstr_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(X1)))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_11385, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k5_binop_1(u1_struct_0(X1),u1_algstr_0(X1),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))!=k1_algstr_0(X1,k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).#trainneg
cnf(i_0_12342, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~v3_rlvect_1(X1))).#trainneg
cnf(i_0_12384, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_binop_1(u1_algstr_0(X1),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))!=k1_algstr_0(X1,k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).#trainneg
cnf(i_0_12482, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))!=k1_algstr_0(X1,k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).#trainneg
cnf(i_0_12393, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_algstr_0(X1,k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).#trainneg
cnf(i_0_12392, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_algstr_0(X1,k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))).#trainneg
cnf(i_0_12500, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~v3_rlvect_1(X1)|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).#trainneg
cnf(i_0_5107, plain, (l3_algstr_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),esk15_0)),esk15_0))|~l1_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_2548, negated_conjecture, (k1_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk1_0))=k4_algstr_0(esk1_0,k4_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_347, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),X3))=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k4_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,k4_struct_0(X1)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1)))).#trainneg
cnf(i_0_12635, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)|~v3_rlvect_1(X1)|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).#trainneg
cnf(i_0_12521, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_algstr_0(X1,k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).#trainneg
cnf(i_0_12390, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k5_binop_1(u1_struct_0(X1),u1_algstr_0(X1),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))!=k1_algstr_0(X1,k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).#trainneg
cnf(i_0_12612, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))!=k1_algstr_0(X1,k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).#trainneg
cnf(i_0_12520, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k5_binop_1(u1_struct_0(X1),u1_algstr_0(X1),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))!=k1_algstr_0(X1,k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).#trainneg
cnf(i_0_12391, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_algstr_0(X1,k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k1_algstr_0(X1,esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(X1))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1)))).#trainneg
cnf(i_0_13020, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),esk15_0))|~l1_algstr_0(X1)|~v3_rlvect_1(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_422, plain, (v1_xreal_0(k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4)))|~v3_membered(u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_group_1(X1))).#trainneg
cnf(i_0_348, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3))))=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3))),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1)))).#trainneg
cnf(i_0_349, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),X3))=k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),X3))|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1)))).#trainneg
cnf(i_0_10645, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X3),k1_algstr_0(X1,X3,k4_struct_0(X1)))=k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1)))).#trainneg
cnf(i_0_9080, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)))=k1_algstr_0(X1,k1_algstr_0(X1,X2,k4_algstr_0(X1,X3)),X3)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1)))).#trainneg
cnf(i_0_14056, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,k4_algstr_0(X1,X3)),X3)=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_2642, plain, (u1_struct_0(g3_algstr_0(esk15_0,esk15_0))=esk15_0|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),esk15_0))|~v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_13023, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v3_rlvect_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_376, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X3,k4_algstr_0(X1,X3)),X2)=k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,k4_struct_0(X1),X2))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~v3_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_377, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,k1_algstr_0(X1,X2,k4_algstr_0(X1,X2)))))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_14068, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,k4_algstr_0(X1,X3)),X3)=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_2532, plain, (u1_struct_0(X1)=X2|g3_algstr_0(u1_struct_0(X1),esk15_0)!=g3_algstr_0(X2,X3)|~l1_algstr_0(X1)|~v1_xboole_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))).#trainneg
cnf(i_0_3685, plain, (l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(g3_algstr_0(u1_struct_0(X1),esk15_0))))|~l1_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_5036, plain, (v1_xboole_0(u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_379, plain, (m1_subset_1(k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k4_algstr_0(X1,X3))),u1_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_380, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k4_algstr_0(X1,X3)))=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_4069, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,k4_struct_0(X1),X2))=k1_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,k4_struct_0(X1)),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1)))).#trainneg
cnf(i_0_14843, plain, (u1_struct_0(X1)=X2|g3_algstr_0(u1_struct_0(X1),esk15_0)!=g3_algstr_0(X2,X3)|~l1_algstr_0(X1)|~v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))).#trainneg
cnf(i_0_14867, plain, (v1_xboole_0(u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_1585, plain, (u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))))=u1_struct_0(X1)|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_385, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)))=k1_algstr_0(X1,X2,k4_algstr_0(X1,k4_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k4_struct_0(X1)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1)))).#trainneg
cnf(i_0_4535, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_binop_1(u1_algstr_0(X1),k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)|~m1_subset_1(esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))).#trainneg
cnf(i_0_9486, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),X1)!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=X1|k1_algstr_0(esk1_0,X2,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=X1|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).#trainneg
cnf(i_0_8298, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X4,k4_algstr_0(X1,X4)),X3)=k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_algstr_0(X1,X2),X3))|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_15987, plain, (k1_algstr_0(X1,X4,k1_algstr_0(X1,k4_algstr_0(X1,X4),X3))=k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_algstr_0(X1,X2),X3))|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_727, plain, (v5_relat_1(u2_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_2635, plain, (esk15_0=X2|g3_algstr_0(esk15_0,u2_algstr_0(g3_algstr_0(esk15_0,u1_algstr_0(g3_algstr_0(u1_struct_0(X1),esk15_0)))))!=g3_algstr_0(X2,X3)|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),esk15_0))|~v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_426, plain, (v1_xreal_0(k1_algstr_0(X1,X2,k4_algstr_0(X1,X2)))|v2_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_631, negated_conjecture, (v1_xreal_0(k1_algstr_0(esk1_0,X2,X1))|~v3_membered(u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_15304, plain, (v1_xboole_0(u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))))|~v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_386, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k4_algstr_0(X1,k1_algstr_0(X1,X2,X3))))=k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,X2,X3)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_387, plain, (k1_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,k4_algstr_0(X1,X2)))=k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k4_algstr_0(X1,X2)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_9060, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,k4_struct_0(X1)),X3)=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k4_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_6374, negated_conjecture, (m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).#trainneg
cnf(i_0_701, plain, (m1_subset_1(esk4_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v2_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_427, plain, (v1_xreal_0(k1_algstr_0(X1,k4_algstr_0(X1,X2),X2))|v2_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_390, plain, (k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)),X3)=k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_algstr_0(X1,X2),X3))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_391, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_algstr_0(X1,X2),k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)))))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_5371, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|k1_binop_1(u1_algstr_0(X1),esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))!=k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),k1_algstr_0(X1,esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(X1)|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)),esk8_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),esk9_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))).#trainneg
cnf(i_0_703, plain, (v2_binop_1(u2_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_8576, plain, (u1_struct_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))=u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_393, plain, (m1_subset_1(k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))),u1_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_10224, plain, (m1_subset_1(k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4)),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_5001, plain, (m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_705, plain, (v1_finseqop(u2_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v2_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_16534, plain, (esk15_0=X1|g3_algstr_0(esk15_0,u2_algstr_0(g3_algstr_0(esk15_0,u1_algstr_0(g3_algstr_0(esk15_0,esk15_0)))))!=g3_algstr_0(X1,X3)|~l1_algstr_0(g3_algstr_0(esk15_0,esk15_0))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X2),esk15_0))|~v1_xboole_0(u1_struct_0(X2))|~l1_algstr_0(X2))).#trainneg
cnf(i_0_18228, negated_conjecture, (k4_algstr_0(esk1_0,k4_struct_0(esk1_0))=k4_struct_0(esk1_0)|~m1_subset_1(k4_algstr_0(esk1_0,k4_struct_0(esk1_0)),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_396, plain, (k1_algstr_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X3),X3),X2)=k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,k4_struct_0(X1),X2))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~v3_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_397, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X2),X2))))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_8191, negated_conjecture, (m1_subset_1(k4_algstr_0(esk1_0,k4_struct_0(esk1_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2253, negated_conjecture, (k4_algstr_0(esk1_0,k4_struct_0(esk1_0))=k4_struct_0(esk1_0))).#trainneg
cnf(i_0_617, negated_conjecture, (k1_algstr_0(esk1_0,X1,k4_struct_0(esk1_0))=X1|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1977, negated_conjecture, (k1_algstr_0(esk1_0,k4_struct_0(esk1_0),X1)=X1|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_19544, negated_conjecture, (m1_subset_1(k1_algstr_0(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_18738, negated_conjecture, (v1_xreal_0(k4_struct_0(esk1_0))|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_505, plain, (k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3))=esk15_0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_8489, plain, (u1_struct_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(X1)))))=u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_399, plain, (m1_subset_1(k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_algstr_0(X1,X3),X3)),u1_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_407, plain, (k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)),X3)=k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~v3_rlvect_1(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_18671, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_506, plain, (k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3))=esk15_0|~v1_xboole_0(X3))).#trainneg
cnf(i_0_408, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,X2,k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)))))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_409, plain, (k1_algstr_0(X1,k4_algstr_0(X1,k4_algstr_0(X1,X2)),k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)))=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_6350, negated_conjecture, (k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X2),X3)=k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X3))|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).#trainneg
cnf(i_0_20979, negated_conjecture, (k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X1),X3)=k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X3))|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).#trainneg
cnf(i_0_6947, negated_conjecture, (k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X2),X3)=k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X3,X2))|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).#trainneg
cnf(i_0_6443, negated_conjecture, (k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X3))=k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X3))|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).#trainneg
cnf(i_0_6447, negated_conjecture, (k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X3))=k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X3,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X1))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).#trainneg
cnf(i_0_1523, plain, (k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)=esk15_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3226, negated_conjecture, (k1_algstr_0(esk1_0,X1,k4_algstr_0(esk1_0,X1))=k4_struct_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4574, plain, (v1_xboole_0(u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~v1_xboole_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_6347, negated_conjecture, (k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X3))=k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X3,X1))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).#trainneg
cnf(i_0_6945, negated_conjecture, (k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X3))=k1_algstr_0(esk1_0,X2,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X3))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X3),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).#trainneg
cnf(i_0_21415, plain, (m1_subset_1(u2_algstr_0(X1),k1_zfmisc_1(esk15_0))|~l3_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_2555, negated_conjecture, (k1_algstr_0(esk1_0,k4_algstr_0(esk1_0,X1),X1)=k4_struct_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_412, plain, (k1_algstr_0(X1,k4_algstr_0(X1,k1_algstr_0(X1,X2,k4_struct_0(X1))),X2)=k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,X2,k4_struct_0(X1))),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,k4_struct_0(X1)),u1_struct_0(X1)))).#trainneg
cnf(i_0_20210, plain, (m1_subset_1(k1_algstr_0(X1,k4_algstr_0(X1,X3),k1_algstr_0(X1,X3,X2)),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k4_algstr_0(X1,X2)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_20958, negated_conjecture, (k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X3))=k1_algstr_0(esk1_0,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X2),X3)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).#trainneg
cnf(i_0_21417, plain, (v1_xreal_0(u2_algstr_0(X1))|~v3_membered(k1_zfmisc_1(esk15_0))|~l3_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_2478, negated_conjecture, (k1_algstr_0(esk1_0,k1_algstr_0(esk1_0,X1,X2),X3)=k1_algstr_0(esk1_0,X1,k1_algstr_0(esk1_0,X2,X3))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7034, plain, (v1_finseqop(u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~v2_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_8848, plain, (v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~v1_xboole_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_431, plain, (v1_xreal_0(k1_algstr_0(X1,k4_struct_0(X1),X3))|v2_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_22338, negated_conjecture, (k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,X3))=k1_algstr_0(esk1_0,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X2),X3)|~m1_subset_1(X3,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_499, plain, (u1_struct_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))=u1_struct_0(X1)|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_23076, plain, (m1_subset_1(esk3_1(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))),u1_struct_0(X1))|v5_group_1(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_23075, plain, (m1_subset_1(esk2_1(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))),u1_struct_0(X1))|v5_group_1(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_23083, plain, (m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)),X2),u1_struct_0(X1))|v2_group_1(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_23073, plain, (m1_subset_1(esk8_1(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))),u1_struct_0(X1))|v3_group_1(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_23072, plain, (m1_subset_1(esk7_1(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))),u1_struct_0(X1))|v3_group_1(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_23267, plain, (u1_struct_0(g3_algstr_0(u1_struct_0(X1),esk15_0))=u1_struct_0(X1)|~l3_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_503, plain, (k2_zfmisc_1(X1,esk13_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3))))=esk15_0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_8497, plain, (v2_struct_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(X1)))))|~v1_xboole_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_8705, plain, (v2_struct_0(g3_algstr_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~v1_xboole_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_23401, plain, (v2_struct_0(g3_algstr_0(u1_struct_0(X1),esk15_0))|~v1_xboole_0(u1_struct_0(X1))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),esk15_0))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_504, plain, (k2_zfmisc_1(X1,esk13_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3))))=esk15_0|~v1_xboole_0(X3))).#trainneg
cnf(i_0_6005, plain, (v1_setwiseo(u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))),u1_struct_0(X1))|v2_struct_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~v1_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_23064, plain, (v2_struct_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))|~v1_xboole_0(u1_struct_0(X1))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_507, plain, (k2_zfmisc_1(X1,u2_algstr_0(X2))=esk15_0|~v1_xboole_0(u1_struct_0(X2))|~l3_algstr_0(X2))).#trainneg
cnf(i_0_441, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,k4_algstr_0(X1,k4_algstr_0(X1,X2)),X3))=k1_algstr_0(X1,k1_algstr_0(X1,X2,k4_struct_0(X1)),X3)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,k4_algstr_0(X1,X2)),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~v3_rlvect_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_442, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,k4_algstr_0(X1,k4_algstr_0(X1,X2)),k4_algstr_0(X1,k1_algstr_0(X1,X2,k4_struct_0(X1)))))=k4_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,X2,k4_struct_0(X1))),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,k4_algstr_0(X1,X2)),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_13970, plain, (k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X3)),X3)=k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_24414, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_algstr_0(X1,X4),X4))=k1_algstr_0(X1,k1_algstr_0(X1,X2,k4_algstr_0(X1,X3)),X3)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_23087, plain, (l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_23605, plain, (esk15_0=X1|g3_algstr_0(esk15_0,u1_algstr_0(g3_algstr_0(u1_struct_0(X2),esk15_0)))!=g3_algstr_0(X1,X3)|~v1_xboole_0(u1_struct_0(X2))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X2),esk15_0))|~l3_algstr_0(X2))).#trainneg
cnf(i_0_508, plain, (k2_zfmisc_1(X1,u1_algstr_0(X2))=esk15_0|~v1_xboole_0(u1_struct_0(X2))|~l1_algstr_0(X2))).#trainneg
cnf(i_0_443, plain, (k1_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),k1_algstr_0(X1,X2,k4_struct_0(X1)))=k1_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,k4_algstr_0(X1,X2)))|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k4_algstr_0(X1,X2)),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_24846, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k1_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),X1)!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k4_struct_0(esk1_0)!=X1|~m1_subset_1(k4_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)),u1_struct_0(esk1_0))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).#trainneg
cnf(i_0_24853, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X2),esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X2)))!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X2))|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X2,k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X2))))!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X2))|k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X2)!=k4_struct_0(esk1_0)|~m1_subset_1(k4_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X2))),u1_struct_0(esk1_0))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X2)),u1_struct_0(esk1_0))|~m1_subset_1(k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X2),u1_struct_0(esk1_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),k6_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1,X2)),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))))).#trainneg
cnf(i_0_515, plain, (v1_funct_2(esk15_0,k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~l3_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_446, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,k4_algstr_0(X1,X3)))=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k4_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,k4_algstr_0(X1,X3)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_448, plain, (v1_xreal_0(k1_algstr_0(X1,k4_algstr_0(X1,X3),k4_struct_0(X1)))|v2_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,k4_algstr_0(X1,k1_algstr_0(X1,X3,X2))),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,X3,X2)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_25099, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k1_algstr_0(esk1_0,X1,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k1_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),X1)!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k4_struct_0(esk1_0)!=X1|~m1_subset_1(k4_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)),u1_struct_0(esk1_0))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))).#trainneg
cnf(i_0_25090, negated_conjecture, (v2_group_1(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|k1_algstr_0(esk1_0,X1,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1))!=esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)|k4_struct_0(esk1_0)!=X1|~m1_subset_1(k4_algstr_0(esk1_0,esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1)),u1_struct_0(esk1_0))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0))))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))|~m1_subset_1(esk6_2(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)),X1),u1_struct_0(g3_algstr_0(u1_struct_0(esk1_0),u1_algstr_0(esk1_0)))))).#trainneg
cnf(i_0_25306, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,X2,k4_struct_0(X1)))=k1_algstr_0(X1,X3,k4_algstr_0(X1,X3))|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,k4_algstr_0(X1,X2)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_25717, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X4,k4_algstr_0(X1,X4)))=k1_algstr_0(X1,k1_algstr_0(X1,X2,k4_algstr_0(X1,X3)),X3)|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,k4_algstr_0(X1,X3)),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_24861, plain, (l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(g3_algstr_0(u1_struct_0(X1),esk15_0))))|~l3_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_24889, plain, (esk15_0=X1|g3_algstr_0(esk15_0,u1_algstr_0(g3_algstr_0(esk15_0,esk15_0)))!=g3_algstr_0(X1,X3)|~l1_algstr_0(g3_algstr_0(esk15_0,esk15_0))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X2),esk15_0))|~v1_xboole_0(u1_struct_0(X2))|~l1_algstr_0(X2))).#trainneg
cnf(i_0_736, plain, (u1_struct_0(X1)=X2|g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1))))!=g3_algstr_0(X2,X3)|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_13024, plain, (v3_group_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~v3_rlvect_1(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))))|~l1_algstr_0(g3_algstr_0(u1_struct_0(X1),u1_algstr_0(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_449, plain, (m1_subset_1(k1_algstr_0(X1,k4_algstr_0(X1,X3),k4_struct_0(X1)),u1_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,k4_algstr_0(X1,k1_algstr_0(X1,X3,X2))),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,X3,X2)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_450, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,k1_algstr_0(X1,X2,k4_algstr_0(X1,k1_algstr_0(X1,X3,X2))),X4))=k1_algstr_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X3),k4_struct_0(X1)),X4)|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,X2,k4_algstr_0(X1,k1_algstr_0(X1,X3,X2))),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~v3_rlvect_1(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,X3,X2)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_25852, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k4_algstr_0(X1,X3)))=k1_algstr_0(X1,k1_algstr_0(X1,X2,k4_algstr_0(X1,X4)),X4)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X4),u1_struct_0(X1)))).#trainneg
cnf(i_0_26571, plain, (k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X3)),X3)=k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))|v2_struct_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_24873, plain, (l3_algstr_0(g3_algstr_0(u1_struct_0(X1),esk15_0))|~l3_algstr_0(X1)|~v1_xboole_0(u1_struct_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))))|~l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))))).#trainneg
cnf(i_0_518, plain, (k5_binop_1(u1_struct_0(X1),esk15_0,X2,X3)=k6_algstr_0(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_452, plain, (k1_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),k1_algstr_0(X1,k4_algstr_0(X1,X3),k4_struct_0(X1)))=k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X2,k4_algstr_0(X1,k1_algstr_0(X1,X3,X2))))|v2_struct_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,k4_algstr_0(X1,k1_algstr_0(X1,X3,X2))),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,k1_algstr_0(X1,X3,X2)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_26758, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X2,k4_algstr_0(X1,X2)))=k1_algstr_0(X1,k4_algstr_0(X1,X3),X3)|v2_struct_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1)))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 56
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 106
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 106
# Processed clauses                    : 1469
# ...of these trivial                  : 38
# ...subsumed                          : 787
# ...remaining for further processing  : 644
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 63
# Backward-rewritten                   : 33
# Generated clauses                    : 27747
# ...of the previous two non-trivial   : 26936
# Contextual simplify-reflections      : 350
# Paramodulations                      : 27719
# Factorizations                       : 0
# Equation resolutions                 : 28
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 548
#    Positive orientable unit clauses  : 37
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 507
# Current number of unprocessed clauses: 25270
# ...number of literals in the above   : 231876
# Current number of archived formulas  : 0
# Current number of archived clauses   : 740
# Proof object given clauses           : 139
# Proof search given clauses           : 644
# Clause-clause subsumption calls (NU) : 180458
# Rec. Clause-clause subsumption calls : 15811
# Non-unit clause-clause subsumptions  : 1186
# Unit Clause-clause subsumption calls : 1506
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 24
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1354340

# -------------------------------------------------
# User time                : 14.530 s
# System time              : 0.492 s
# Total time               : 15.022 s
# Maximum resident set size: 99044 pages
