2020-06-04 10:54:10.960846: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_69
2020-06-04 10:54:11.064495: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-06-04 10:54:11.213560: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-06-04 10:54:11.745130: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-06-04 10:54:13.853050: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 2892195 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_69' loaded (query=128; context=256; weigths=real; len_mult=0.000000).
# Preprocessing time       : 3.971 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', d5_xboole_0)).
fof(d4_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,X2)=k4_xboole_0(X1,X2)), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', d4_subset_1)).
fof(involutiveness_k3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,k3_subset_1(X1,X2))=X2), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', involutiveness_k3_subset_1)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', t6_boole)).
fof(rc2_subset_1, axiom, ![X1]:?[X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))&v1_xboole_0(X2)), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', rc2_subset_1)).
fof(dt_k3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>m1_subset_1(k3_subset_1(X1,X2),k1_zfmisc_1(X1))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', dt_k3_subset_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', t3_boole)).
fof(t4_boole, axiom, ![X1]:k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0, file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', t4_boole)).
fof(t36_yellow19, conjecture, ![X1]:(((~(v2_struct_0(X1))&v2_pre_topc(X1))&l1_pre_topc(X1))=>(v1_compts_1(X1)<=>![X2]:((((~(v2_struct_0(X2))&v4_orders_2(X2))&v7_waybel_0(X2))&l1_waybel_0(X2,X1))=>(v1_yellow19(X2,X1)=>v3_yellow_6(X2,X1))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', t36_yellow19)).
fof(t17_yellow19, axiom, ![X1]:(((~(v2_struct_0(X1))&v2_pre_topc(X1))&l1_pre_topc(X1))=>![X2]:(((((~(v1_xboole_0(X2))&v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))&v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1))))&v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1))))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1))))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2)))<=>r2_waybel_7(X1,X2,X3))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', t17_yellow19)).
fof(d16_yellow_6, axiom, ![X1]:(((~(v2_struct_0(X1))&v2_pre_topc(X1))&l1_pre_topc(X1))=>![X2]:((((~(v2_struct_0(X2))&v4_orders_2(X2))&v7_waybel_0(X2))&l1_waybel_0(X2,X1))=>(v3_yellow_6(X2,X1)<=>k10_yellow_6(X1,X2)!=k1_xboole_0))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', d16_yellow_6)).
fof(t31_yellow19, axiom, ![X1]:(((~(v2_struct_0(X1))&v2_pre_topc(X1))&l1_pre_topc(X1))=>(v1_compts_1(X1)<=>![X2]:(((((~(v1_xboole_0(X2))&v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1))))&v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1))))&v3_waybel_7(X2,k3_yellow_1(k2_struct_0(X1))))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1))))))=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&r2_waybel_7(X1,X2,X3))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', t31_yellow19)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', t4_subset)).
fof(dt_k10_yellow_6, axiom, ![X1, X2]:(((((((~(v2_struct_0(X1))&v2_pre_topc(X1))&l1_pre_topc(X1))&~(v2_struct_0(X2)))&v4_orders_2(X2))&v7_waybel_0(X2))&l1_waybel_0(X2,X1))=>m1_subset_1(k10_yellow_6(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', dt_k10_yellow_6)).
fof(fc6_yellow19, axiom, ![X1, X2]:(((((((~(v2_struct_0(X1))&l1_struct_0(X1))&~(v1_xboole_0(X2)))&v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1))))&v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1))))&v3_waybel_7(X2,k3_yellow_1(k2_struct_0(X1))))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1))))))=>((~(v2_struct_0(k3_yellow19(X1,k2_struct_0(X1),X2)))&v6_waybel_0(k3_yellow19(X1,k2_struct_0(X1),X2),X1))&v1_yellow19(k3_yellow19(X1,k2_struct_0(X1),X2),X1))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', fc6_yellow19)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', dt_l1_pre_topc)).
fof(d11_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r1_waybel_0(X1,X2,X3)<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(r1_orders_2(X2,X4,X5)=>r2_hidden(k2_waybel_0(X1,X2,X5),X3))))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', d11_waybel_0)).
fof(d12_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r2_waybel_0(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>?[X5]:((m1_subset_1(X5,u1_struct_0(X2))&r1_orders_2(X2,X4,X5))&r2_hidden(k2_waybel_0(X1,X2,X5),X3)))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', d12_waybel_0)).
fof(fc5_yellow19, axiom, ![X1, X2, X3]:(((((((((~(v2_struct_0(X1))&l1_struct_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&~(v1_xboole_0(X3)))&v1_subset_1(X3,u1_struct_0(k3_yellow_1(X2))))&v2_waybel_0(X3,k3_yellow_1(X2)))&v13_waybel_0(X3,k3_yellow_1(X2)))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2)))))=>((~(v2_struct_0(k3_yellow19(X1,X2,X3)))&v6_waybel_0(k3_yellow19(X1,X2,X3),X1))&v7_waybel_0(k3_yellow19(X1,X2,X3)))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', fc5_yellow19)).
fof(fc4_yellow19, axiom, ![X1, X2, X3]:((((((((~(v2_struct_0(X1))&l1_struct_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&~(v1_xboole_0(X3)))&v2_waybel_0(X3,k3_yellow_1(X2)))&v13_waybel_0(X3,k3_yellow_1(X2)))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2)))))=>(((~(v2_struct_0(k3_yellow19(X1,X2,X3)))&v3_orders_2(k3_yellow19(X1,X2,X3)))&v4_orders_2(k3_yellow19(X1,X2,X3)))&v6_waybel_0(k3_yellow19(X1,X2,X3),X1))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', fc4_yellow19)).
fof(cc1_waybel_7, axiom, ![X1]:(((((~(v2_struct_0(X1))&v3_orders_2(X1))&v4_orders_2(X1))&v5_orders_2(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((((~(v1_xboole_0(X2))&v2_waybel_0(X2,X1))&v13_waybel_0(X2,X1))&v3_waybel_7(X2,X1))=>(((~(v1_xboole_0(X2))&v1_subset_1(X2,u1_struct_0(X1)))&v2_waybel_0(X2,X1))&v13_waybel_0(X2,X1))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', cc1_waybel_7)).
fof(dt_k3_yellow_1, axiom, ![X1]:(v1_orders_2(k3_yellow_1(X1))&l1_orders_2(k3_yellow_1(X1))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', dt_k3_yellow_1)).
fof(fc7_yellow_1, axiom, ![X1]:((((~(v2_struct_0(k3_yellow_1(X1)))&v1_orders_2(k3_yellow_1(X1)))&v3_orders_2(k3_yellow_1(X1)))&v4_orders_2(k3_yellow_1(X1)))&v5_orders_2(k3_yellow_1(X1))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', fc7_yellow_1)).
fof(dt_k3_yellow19, axiom, ![X1, X2, X3]:((((((((~(v2_struct_0(X1))&l1_struct_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&~(v1_xboole_0(X3)))&v2_waybel_0(X3,k3_yellow_1(X2)))&v13_waybel_0(X3,k3_yellow_1(X2)))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2)))))=>((~(v2_struct_0(k3_yellow19(X1,X2,X3)))&v6_waybel_0(k3_yellow19(X1,X2,X3),X1))&l1_waybel_0(k3_yellow19(X1,X2,X3),X1))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', dt_k3_yellow19)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', t3_subset)).
fof(reflexivity_r1_tarski, axiom, ![X1, X2]:r1_tarski(X1,X1), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', reflexivity_r1_tarski)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', d3_struct_0)).
fof(d5_yellow19, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>(v1_yellow19(X2,X1)<=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r1_waybel_0(X1,X2,X3)|r1_waybel_0(X1,X2,k3_subset_1(u1_struct_0(X1),X3))))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', d5_yellow19)).
fof(d9_waybel_9, axiom, ![X1]:(((~(v2_struct_0(X1))&v2_pre_topc(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_waybel_9(X1,X2,X3)<=>![X4]:(m1_connsp_2(X4,X1,X3)=>r2_waybel_0(X1,X2,X4)))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', d9_waybel_9)).
fof(d15_yellow_6, axiom, ![X1]:(((~(v2_struct_0(X1))&v2_pre_topc(X1))&l1_pre_topc(X1))=>![X2]:((((~(v2_struct_0(X2))&v4_orders_2(X2))&v7_waybel_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=k10_yellow_6(X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)<=>![X5]:(m1_connsp_2(X5,X1,X4)=>r1_waybel_0(X1,X2,X5)))))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', d15_yellow_6)).
fof(dt_m1_connsp_2, axiom, ![X1, X2]:((((~(v2_struct_0(X1))&v2_pre_topc(X1))&l1_pre_topc(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>![X3]:(m1_connsp_2(X3,X1,X2)=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', dt_m1_connsp_2)).
fof(l35_yellow19, axiom, ![X1]:(((~(v2_struct_0(X1))&v2_pre_topc(X1))&l1_pre_topc(X1))=>(v1_compts_1(X1)=>![X2]:((((~(v2_struct_0(X2))&v4_orders_2(X2))&v7_waybel_0(X2))&l1_waybel_0(X2,X1))=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&r3_waybel_9(X1,X2,X3))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', l35_yellow19)).
fof(dt_k2_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', dt_k2_struct_0)).
fof(fc5_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(k2_struct_0(X1)))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', fc5_struct_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('/local1/mptp/loop2/hardpred/preds__0.05/t36_yellow19', fc2_struct_0)).
fof(c_0_35, plain, ![X67, X68, X69, X70, X71, X72, X73, X74]:((((r2_hidden(X70,X67)|~r2_hidden(X70,X69)|X69!=k4_xboole_0(X67,X68))&(~r2_hidden(X70,X68)|~r2_hidden(X70,X69)|X69!=k4_xboole_0(X67,X68)))&(~r2_hidden(X71,X67)|r2_hidden(X71,X68)|r2_hidden(X71,X69)|X69!=k4_xboole_0(X67,X68)))&((~r2_hidden(esk15_3(X72,X73,X74),X74)|(~r2_hidden(esk15_3(X72,X73,X74),X72)|r2_hidden(esk15_3(X72,X73,X74),X73))|X74=k4_xboole_0(X72,X73))&((r2_hidden(esk15_3(X72,X73,X74),X72)|r2_hidden(esk15_3(X72,X73,X74),X74)|X74=k4_xboole_0(X72,X73))&(~r2_hidden(esk15_3(X72,X73,X74),X73)|r2_hidden(esk15_3(X72,X73,X74),X74)|X74=k4_xboole_0(X72,X73))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_36, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_37, plain, ![X91, X92]:(~m1_subset_1(X92,k1_zfmisc_1(X91))|k3_subset_1(X91,X92)=k4_xboole_0(X91,X92)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_subset_1])])).
cnf(c_0_38, plain, (r2_hidden(X1,X2)|X3!=k4_xboole_0(X2,X4)|~r2_hidden(X1,X3)), inference(evalgc,[status(thm)],[c_0_36])).
cnf(c_0_39, plain, (k3_subset_1(X2,X1)=k4_xboole_0(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_40, plain, ![X61, X62]:(~m1_subset_1(X62,k1_zfmisc_1(X61))|k3_subset_1(X61,k3_subset_1(X61,X62))=X62), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k3_subset_1])])).
fof(c_0_41, plain, ![X106]:(~v1_xboole_0(X106)|X106=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
fof(c_0_42, plain, ![X138]:(m1_subset_1(esk17_1(X138),k1_zfmisc_1(X138))&v1_xboole_0(esk17_1(X138))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc2_subset_1])])).
cnf(c_0_43, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_38])])).
cnf(c_0_44, plain, (k4_xboole_0(X2,X1)=k3_subset_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_39])).
cnf(c_0_45, plain, (k3_subset_1(X2,k3_subset_1(X2,X1))=X1|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_46, plain, ![X59, X60]:(~m1_subset_1(X60,k1_zfmisc_1(X59))|m1_subset_1(k3_subset_1(X59,X60),k1_zfmisc_1(X59))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_subset_1])])).
cnf(c_0_47, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_48, plain, (v1_xboole_0(esk17_1(X1))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_49, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_subset_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44])])).
cnf(c_0_50, plain, (k3_subset_1(X2,k3_subset_1(X2,X1))=X1|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_45])).
cnf(c_0_51, plain, (m1_subset_1(k3_subset_1(X2,X1),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_52, plain, (m1_subset_1(esk17_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_53, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_47])).
cnf(c_0_54, plain, (v1_xboole_0(esk17_1(X1))), inference(evalgc,[status(thm)],[c_0_48])).
fof(c_0_55, plain, ![X105]:k4_xboole_0(X105,k1_xboole_0)=X105, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_56, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~m1_subset_1(k3_subset_1(X2,X3),k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50])])).
cnf(c_0_57, plain, (m1_subset_1(k3_subset_1(X2,X1),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_51])).
cnf(c_0_58, plain, (m1_subset_1(esk17_1(X1),k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_52])).
cnf(c_0_59, plain, (esk17_1(X1)=k1_xboole_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54])])).
cnf(c_0_60, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_61, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_62, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(csr,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_63, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_59])])).
cnf(c_0_64, plain, (X3!=k4_xboole_0(X4,X2)|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_60])).
cnf(c_0_65, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(evalgc,[status(thm)],[c_0_61])).
fof(c_0_66, plain, ![X130]:k4_xboole_0(k1_xboole_0,X130)=k1_xboole_0, inference(variable_rename,[status(thm)],[t4_boole])).
fof(c_0_67, negated_conjecture, ~(![X1]:(((~(v2_struct_0(X1))&v2_pre_topc(X1))&l1_pre_topc(X1))=>(v1_compts_1(X1)<=>![X2]:((((~(v2_struct_0(X2))&v4_orders_2(X2))&v7_waybel_0(X2))&l1_waybel_0(X2,X1))=>(v1_yellow19(X2,X1)=>v3_yellow_6(X2,X1)))))), inference(assume_negation,[status(cth)],[t36_yellow19])).
fof(c_0_68, plain, ![X28, X29, X30]:((~r2_hidden(X30,k10_yellow_6(X28,k3_yellow19(X28,k2_struct_0(X28),X29)))|r2_waybel_7(X28,X29,X30)|~m1_subset_1(X30,u1_struct_0(X28))|(v1_xboole_0(X29)|~v1_subset_1(X29,u1_struct_0(k3_yellow_1(k2_struct_0(X28))))|~v2_waybel_0(X29,k3_yellow_1(k2_struct_0(X28)))|~v13_waybel_0(X29,k3_yellow_1(k2_struct_0(X28)))|~m1_subset_1(X29,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X28))))))|(v2_struct_0(X28)|~v2_pre_topc(X28)|~l1_pre_topc(X28)))&(~r2_waybel_7(X28,X29,X30)|r2_hidden(X30,k10_yellow_6(X28,k3_yellow19(X28,k2_struct_0(X28),X29)))|~m1_subset_1(X30,u1_struct_0(X28))|(v1_xboole_0(X29)|~v1_subset_1(X29,u1_struct_0(k3_yellow_1(k2_struct_0(X28))))|~v2_waybel_0(X29,k3_yellow_1(k2_struct_0(X28)))|~v13_waybel_0(X29,k3_yellow_1(k2_struct_0(X28)))|~m1_subset_1(X29,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X28))))))|(v2_struct_0(X28)|~v2_pre_topc(X28)|~l1_pre_topc(X28)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t17_yellow19])])])])])).
cnf(c_0_69, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63])])).
cnf(c_0_70, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65])])])).
cnf(c_0_71, plain, (r2_hidden(esk15_3(X1,X2,X3),X1)|r2_hidden(esk15_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_72, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_66])).
fof(c_0_73, plain, ![X12, X13]:((~v3_yellow_6(X13,X12)|k10_yellow_6(X12,X13)!=k1_xboole_0|(v2_struct_0(X13)|~v4_orders_2(X13)|~v7_waybel_0(X13)|~l1_waybel_0(X13,X12))|(v2_struct_0(X12)|~v2_pre_topc(X12)|~l1_pre_topc(X12)))&(k10_yellow_6(X12,X13)=k1_xboole_0|v3_yellow_6(X13,X12)|(v2_struct_0(X13)|~v4_orders_2(X13)|~v7_waybel_0(X13)|~l1_waybel_0(X13,X12))|(v2_struct_0(X12)|~v2_pre_topc(X12)|~l1_pre_topc(X12)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_yellow_6])])])])])).
fof(c_0_74, negated_conjecture, ![X11]:(((~v2_struct_0(esk1_0)&v2_pre_topc(esk1_0))&l1_pre_topc(esk1_0))&((((((~v2_struct_0(esk2_0)|~v1_compts_1(esk1_0))&(v4_orders_2(esk2_0)|~v1_compts_1(esk1_0)))&(v7_waybel_0(esk2_0)|~v1_compts_1(esk1_0)))&(l1_waybel_0(esk2_0,esk1_0)|~v1_compts_1(esk1_0)))&((v1_yellow19(esk2_0,esk1_0)|~v1_compts_1(esk1_0))&(~v3_yellow_6(esk2_0,esk1_0)|~v1_compts_1(esk1_0))))&(v1_compts_1(esk1_0)|(v2_struct_0(X11)|~v4_orders_2(X11)|~v7_waybel_0(X11)|~l1_waybel_0(X11,esk1_0)|(~v1_yellow19(X11,esk1_0)|v3_yellow_6(X11,esk1_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_67])])])])])])).
fof(c_0_75, plain, ![X14, X15, X18]:(((m1_subset_1(esk3_2(X14,X15),u1_struct_0(X14))|(v1_xboole_0(X15)|~v2_waybel_0(X15,k3_yellow_1(k2_struct_0(X14)))|~v13_waybel_0(X15,k3_yellow_1(k2_struct_0(X14)))|~v3_waybel_7(X15,k3_yellow_1(k2_struct_0(X14)))|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X14))))))|~v1_compts_1(X14)|(v2_struct_0(X14)|~v2_pre_topc(X14)|~l1_pre_topc(X14)))&(r2_waybel_7(X14,X15,esk3_2(X14,X15))|(v1_xboole_0(X15)|~v2_waybel_0(X15,k3_yellow_1(k2_struct_0(X14)))|~v13_waybel_0(X15,k3_yellow_1(k2_struct_0(X14)))|~v3_waybel_7(X15,k3_yellow_1(k2_struct_0(X14)))|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X14))))))|~v1_compts_1(X14)|(v2_struct_0(X14)|~v2_pre_topc(X14)|~l1_pre_topc(X14))))&((((((~v1_xboole_0(esk4_1(X14))|v1_compts_1(X14)|(v2_struct_0(X14)|~v2_pre_topc(X14)|~l1_pre_topc(X14)))&(v2_waybel_0(esk4_1(X14),k3_yellow_1(k2_struct_0(X14)))|v1_compts_1(X14)|(v2_struct_0(X14)|~v2_pre_topc(X14)|~l1_pre_topc(X14))))&(v13_waybel_0(esk4_1(X14),k3_yellow_1(k2_struct_0(X14)))|v1_compts_1(X14)|(v2_struct_0(X14)|~v2_pre_topc(X14)|~l1_pre_topc(X14))))&(v3_waybel_7(esk4_1(X14),k3_yellow_1(k2_struct_0(X14)))|v1_compts_1(X14)|(v2_struct_0(X14)|~v2_pre_topc(X14)|~l1_pre_topc(X14))))&(m1_subset_1(esk4_1(X14),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X14)))))|v1_compts_1(X14)|(v2_struct_0(X14)|~v2_pre_topc(X14)|~l1_pre_topc(X14))))&(~m1_subset_1(X18,u1_struct_0(X14))|~r2_waybel_7(X14,esk4_1(X14),X18)|v1_compts_1(X14)|(v2_struct_0(X14)|~v2_pre_topc(X14)|~l1_pre_topc(X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t31_yellow19])])])])])])).
cnf(c_0_76, plain, (r2_waybel_7(X2,X3,X1)|v1_xboole_0(X3)|v2_struct_0(X2)|~r2_hidden(X1,k10_yellow_6(X2,k3_yellow19(X2,k2_struct_0(X2),X3)))|~m1_subset_1(X1,u1_struct_0(X2))|~v1_subset_1(X3,u1_struct_0(k3_yellow_1(k2_struct_0(X2))))|~v2_waybel_0(X3,k3_yellow_1(k2_struct_0(X2)))|~v13_waybel_0(X3,k3_yellow_1(k2_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X2)))))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_77, plain, (~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_78, plain, (X3=k4_xboole_0(X1,X2)|r2_hidden(esk15_3(X1,X2,X3),X3)|r2_hidden(esk15_3(X1,X2,X3),X1)), inference(evalgc,[status(thm)],[c_0_71])).
cnf(c_0_79, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(evalgc,[status(thm)],[c_0_72])).
fof(c_0_80, plain, ![X135, X136, X137]:(~r2_hidden(X135,X136)|~m1_subset_1(X136,k1_zfmisc_1(X137))|m1_subset_1(X135,X137)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_81, plain, ![X86, X87]:(v2_struct_0(X86)|~v2_pre_topc(X86)|~l1_pre_topc(X86)|v2_struct_0(X87)|~v4_orders_2(X87)|~v7_waybel_0(X87)|~l1_waybel_0(X87,X86)|m1_subset_1(k10_yellow_6(X86,X87),k1_zfmisc_1(u1_struct_0(X86)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k10_yellow_6])])])).
cnf(c_0_82, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~v3_yellow_6(X1,X2)|k10_yellow_6(X2,X1)!=k1_xboole_0|~v4_orders_2(X1)|~v7_waybel_0(X1)|~l1_waybel_0(X1,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_83, negated_conjecture, (v1_compts_1(esk1_0)|v2_struct_0(X1)|v3_yellow_6(X1,esk1_0)|~v4_orders_2(X1)|~v7_waybel_0(X1)|~l1_waybel_0(X1,esk1_0)|~v1_yellow19(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_84, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_85, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_86, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_74])).
fof(c_0_87, plain, ![X77, X78]:(((~v2_struct_0(k3_yellow19(X77,k2_struct_0(X77),X78))|(v2_struct_0(X77)|~l1_struct_0(X77)|v1_xboole_0(X78)|~v2_waybel_0(X78,k3_yellow_1(k2_struct_0(X77)))|~v13_waybel_0(X78,k3_yellow_1(k2_struct_0(X77)))|~v3_waybel_7(X78,k3_yellow_1(k2_struct_0(X77)))|~m1_subset_1(X78,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X77)))))))&(v6_waybel_0(k3_yellow19(X77,k2_struct_0(X77),X78),X77)|(v2_struct_0(X77)|~l1_struct_0(X77)|v1_xboole_0(X78)|~v2_waybel_0(X78,k3_yellow_1(k2_struct_0(X77)))|~v13_waybel_0(X78,k3_yellow_1(k2_struct_0(X77)))|~v3_waybel_7(X78,k3_yellow_1(k2_struct_0(X77)))|~m1_subset_1(X78,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X77))))))))&(v1_yellow19(k3_yellow19(X77,k2_struct_0(X77),X78),X77)|(v2_struct_0(X77)|~l1_struct_0(X77)|v1_xboole_0(X78)|~v2_waybel_0(X78,k3_yellow_1(k2_struct_0(X77)))|~v13_waybel_0(X78,k3_yellow_1(k2_struct_0(X77)))|~v3_waybel_7(X78,k3_yellow_1(k2_struct_0(X77)))|~m1_subset_1(X78,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X77)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc6_yellow19])])])])).
fof(c_0_88, plain, ![X76]:(~l1_pre_topc(X76)|l1_struct_0(X76)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_89, plain, (v1_compts_1(X2)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_waybel_7(X2,esk4_1(X2),X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_90, plain, (v2_struct_0(X2)|v1_xboole_0(X3)|r2_waybel_7(X2,X3,X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~v2_waybel_0(X3,k3_yellow_1(k2_struct_0(X2)))|~v13_waybel_0(X3,k3_yellow_1(k2_struct_0(X2)))|~v1_subset_1(X3,u1_struct_0(k3_yellow_1(k2_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X2)))))|~r2_hidden(X1,k10_yellow_6(X2,k3_yellow19(X2,k2_struct_0(X2),X3)))), inference(evalgc,[status(thm)],[c_0_76])).
cnf(c_0_91, plain, (X2=k1_xboole_0|r2_hidden(esk15_3(k1_xboole_0,X1,X2),X2)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_79])])).
cnf(c_0_92, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_93, plain, (v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(k10_yellow_6(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
fof(c_0_94, plain, ![X39, X40, X41, X43, X44, X45]:(((m1_subset_1(esk9_3(X39,X40,X41),u1_struct_0(X40))|~r1_waybel_0(X39,X40,X41)|(v2_struct_0(X40)|~l1_waybel_0(X40,X39))|(v2_struct_0(X39)|~l1_struct_0(X39)))&(~m1_subset_1(X43,u1_struct_0(X40))|(~r1_orders_2(X40,esk9_3(X39,X40,X41),X43)|r2_hidden(k2_waybel_0(X39,X40,X43),X41))|~r1_waybel_0(X39,X40,X41)|(v2_struct_0(X40)|~l1_waybel_0(X40,X39))|(v2_struct_0(X39)|~l1_struct_0(X39))))&((m1_subset_1(esk10_4(X39,X40,X44,X45),u1_struct_0(X40))|~m1_subset_1(X45,u1_struct_0(X40))|r1_waybel_0(X39,X40,X44)|(v2_struct_0(X40)|~l1_waybel_0(X40,X39))|(v2_struct_0(X39)|~l1_struct_0(X39)))&((r1_orders_2(X40,X45,esk10_4(X39,X40,X44,X45))|~m1_subset_1(X45,u1_struct_0(X40))|r1_waybel_0(X39,X40,X44)|(v2_struct_0(X40)|~l1_waybel_0(X40,X39))|(v2_struct_0(X39)|~l1_struct_0(X39)))&(~r2_hidden(k2_waybel_0(X39,X40,esk10_4(X39,X40,X44,X45)),X44)|~m1_subset_1(X45,u1_struct_0(X40))|r1_waybel_0(X39,X40,X44)|(v2_struct_0(X40)|~l1_waybel_0(X40,X39))|(v2_struct_0(X39)|~l1_struct_0(X39)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d11_waybel_0])])])])])])])).
fof(c_0_95, plain, ![X31, X32, X33, X34, X36, X38]:((((m1_subset_1(esk7_4(X31,X32,X33,X34),u1_struct_0(X32))|~m1_subset_1(X34,u1_struct_0(X32))|~r2_waybel_0(X31,X32,X33)|(v2_struct_0(X32)|~l1_waybel_0(X32,X31))|(v2_struct_0(X31)|~l1_struct_0(X31)))&(r1_orders_2(X32,X34,esk7_4(X31,X32,X33,X34))|~m1_subset_1(X34,u1_struct_0(X32))|~r2_waybel_0(X31,X32,X33)|(v2_struct_0(X32)|~l1_waybel_0(X32,X31))|(v2_struct_0(X31)|~l1_struct_0(X31))))&(r2_hidden(k2_waybel_0(X31,X32,esk7_4(X31,X32,X33,X34)),X33)|~m1_subset_1(X34,u1_struct_0(X32))|~r2_waybel_0(X31,X32,X33)|(v2_struct_0(X32)|~l1_waybel_0(X32,X31))|(v2_struct_0(X31)|~l1_struct_0(X31))))&((m1_subset_1(esk8_3(X31,X32,X36),u1_struct_0(X32))|r2_waybel_0(X31,X32,X36)|(v2_struct_0(X32)|~l1_waybel_0(X32,X31))|(v2_struct_0(X31)|~l1_struct_0(X31)))&(~m1_subset_1(X38,u1_struct_0(X32))|~r1_orders_2(X32,esk8_3(X31,X32,X36),X38)|~r2_hidden(k2_waybel_0(X31,X32,X38),X36)|r2_waybel_0(X31,X32,X36)|(v2_struct_0(X32)|~l1_waybel_0(X32,X31))|(v2_struct_0(X31)|~l1_struct_0(X31))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d12_waybel_0])])])])])])])).
cnf(c_0_96, plain, (v2_struct_0(X2)|v2_struct_0(X1)|k10_yellow_6(X2,X1)!=k1_xboole_0|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v4_orders_2(X1)|~v7_waybel_0(X1)|~l1_waybel_0(X1,X2)|~v3_yellow_6(X1,X2)), inference(evalgc,[status(thm)],[c_0_82])).
cnf(c_0_97, negated_conjecture, (v2_struct_0(X1)|v1_compts_1(esk1_0)|v3_yellow_6(X1,esk1_0)|~v4_orders_2(X1)|~v7_waybel_0(X1)|~l1_waybel_0(X1,esk1_0)|~v1_yellow19(X1,esk1_0)), inference(evalgc,[status(thm)],[c_0_83])).
cnf(c_0_98, negated_conjecture, (l1_pre_topc(esk1_0)), inference(evalgc,[status(thm)],[c_0_84])).
cnf(c_0_99, negated_conjecture, (v2_pre_topc(esk1_0)), inference(evalgc,[status(thm)],[c_0_85])).
cnf(c_0_100, negated_conjecture, (~v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_86])).
cnf(c_0_101, plain, (v1_yellow19(k3_yellow19(X1,k2_struct_0(X1),X2),X1)|v2_struct_0(X1)|v1_xboole_0(X2)|~l1_struct_0(X1)|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v3_waybel_7(X2,k3_yellow_1(k2_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_102, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_103, plain, (v2_struct_0(X2)|v1_compts_1(X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_waybel_7(X2,esk4_1(X2),X1)), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_104, plain, (r2_waybel_7(X1,X2,esk15_3(k1_xboole_0,X3,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2))))|v1_xboole_0(X2)|v2_struct_0(X1)|k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2))=k1_xboole_0|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))|~m1_subset_1(esk15_3(k1_xboole_0,X3,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2))),u1_struct_0(X1))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91])])).
cnf(c_0_105, plain, (v2_waybel_0(esk4_1(X1),k3_yellow_1(k2_struct_0(X1)))|v1_compts_1(X1)|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_106, plain, (v13_waybel_0(esk4_1(X1),k3_yellow_1(k2_struct_0(X1)))|v1_compts_1(X1)|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_107, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))|v1_compts_1(X1)|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_108, plain, (v1_compts_1(X1)|v2_struct_0(X1)|~v1_xboole_0(esk4_1(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_109, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_92])).
cnf(c_0_110, plain, (v2_struct_0(X2)|v2_struct_0(X1)|m1_subset_1(k10_yellow_6(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)), inference(evalgc,[status(thm)],[c_0_93])).
cnf(c_0_111, plain, (X1!=k3_subset_1(X2,X3)|~r2_hidden(X4,X1)|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_44])])).
cnf(c_0_112, plain, (r2_hidden(k2_waybel_0(X3,X2,X1),X4)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_orders_2(X2,esk9_3(X3,X2,X4),X1)|~r1_waybel_0(X3,X2,X4)|~l1_waybel_0(X2,X3)|~l1_struct_0(X3)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_113, plain, (r1_orders_2(X1,X2,esk7_4(X3,X1,X4,X2))|v2_struct_0(X1)|v2_struct_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_waybel_0(X3,X1,X4)|~l1_waybel_0(X1,X3)|~l1_struct_0(X3)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_114, negated_conjecture, (v2_struct_0(X1)|v1_compts_1(esk1_0)|k10_yellow_6(esk1_0,X1)!=k1_xboole_0|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~v1_yellow19(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_98]), c_0_99])]), c_0_100])])).
cnf(c_0_115, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|v1_yellow19(k3_yellow19(X1,k2_struct_0(X1),X2),X1)|~l1_struct_0(X1)|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v3_waybel_7(X2,k3_yellow_1(k2_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))), inference(evalgc,[status(thm)],[c_0_101])).
cnf(c_0_116, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[c_0_102])).
cnf(c_0_117, plain, (v1_compts_1(X1)|v2_struct_0(X1)|k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)))=k1_xboole_0|v1_xboole_0(esk4_1(X1))|~m1_subset_1(esk15_3(k1_xboole_0,X2,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)))),u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))|~v13_waybel_0(esk4_1(X1),k3_yellow_1(k2_struct_0(X1)))|~v2_waybel_0(esk4_1(X1),k3_yellow_1(k2_struct_0(X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104])])).
cnf(c_0_118, plain, (v2_struct_0(X1)|v1_compts_1(X1)|v2_waybel_0(esk4_1(X1),k3_yellow_1(k2_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_119, plain, (v2_struct_0(X1)|v1_compts_1(X1)|v13_waybel_0(esk4_1(X1),k3_yellow_1(k2_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[c_0_106])).
cnf(c_0_120, plain, (v2_struct_0(X1)|v1_compts_1(X1)|m1_subset_1(esk4_1(X1),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[c_0_107])).
cnf(c_0_121, plain, (v2_struct_0(X1)|v1_compts_1(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_xboole_0(esk4_1(X1))), inference(evalgc,[status(thm)],[c_0_108])).
cnf(c_0_122, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X3)|~r2_hidden(X1,k10_yellow_6(X2,X3))|~l1_waybel_0(X3,X2)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110])])).
cnf(c_0_123, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,k3_subset_1(X3,X2))|~m1_subset_1(k3_subset_1(X3,X2),k1_zfmisc_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_50])])])).
cnf(c_0_124, plain, (r2_hidden(k2_waybel_0(X1,X2,esk7_4(X1,X2,X3,X4)),X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X2))|~r2_waybel_0(X1,X2,X3)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_125, plain, (v2_struct_0(X3)|v2_struct_0(X2)|r2_hidden(k2_waybel_0(X3,X2,X1),X4)|~l1_struct_0(X3)|~l1_waybel_0(X2,X3)|~r1_waybel_0(X3,X2,X4)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_orders_2(X2,esk9_3(X3,X2,X4),X1)), inference(evalgc,[status(thm)],[c_0_112])).
cnf(c_0_126, plain, (v2_struct_0(X3)|v2_struct_0(X1)|r1_orders_2(X1,X2,esk7_4(X3,X1,X4,X2))|~l1_struct_0(X3)|~l1_waybel_0(X1,X3)|~r2_waybel_0(X3,X1,X4)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_113])).
cnf(c_0_127, plain, (m1_subset_1(esk9_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~r1_waybel_0(X1,X2,X3)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_128, negated_conjecture, (v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1))|v1_xboole_0(X1)|k10_yellow_6(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1))!=k1_xboole_0|~l1_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1),esk1_0)|~v7_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1))|~v4_orders_2(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1))|~l1_struct_0(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~v3_waybel_7(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~v13_waybel_0(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(X1,k3_yellow_1(k2_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_100])])).
cnf(c_0_129, negated_conjecture, (l1_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_98])])).
cnf(c_0_130, plain, (k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)))=k1_xboole_0|v1_compts_1(X1)|v2_struct_0(X1)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~m1_subset_1(esk15_3(k1_xboole_0,X2,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)))),u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_117, c_0_118]), c_0_119]), c_0_120]), c_0_121])).
cnf(c_0_131, plain, (m1_subset_1(esk15_3(k1_xboole_0,X1,k10_yellow_6(X2,X3)),u1_struct_0(X2))|v2_struct_0(X3)|v2_struct_0(X2)|k10_yellow_6(X2,X3)=k1_xboole_0|~l1_waybel_0(X3,X2)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_91])])).
fof(c_0_132, plain, ![X101, X102, X103]:(((~v2_struct_0(k3_yellow19(X101,X102,X103))|(v2_struct_0(X101)|~l1_struct_0(X101)|v1_xboole_0(X102)|~m1_subset_1(X102,k1_zfmisc_1(u1_struct_0(X101)))|v1_xboole_0(X103)|~v1_subset_1(X103,u1_struct_0(k3_yellow_1(X102)))|~v2_waybel_0(X103,k3_yellow_1(X102))|~v13_waybel_0(X103,k3_yellow_1(X102))|~m1_subset_1(X103,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X102))))))&(v6_waybel_0(k3_yellow19(X101,X102,X103),X101)|(v2_struct_0(X101)|~l1_struct_0(X101)|v1_xboole_0(X102)|~m1_subset_1(X102,k1_zfmisc_1(u1_struct_0(X101)))|v1_xboole_0(X103)|~v1_subset_1(X103,u1_struct_0(k3_yellow_1(X102)))|~v2_waybel_0(X103,k3_yellow_1(X102))|~v13_waybel_0(X103,k3_yellow_1(X102))|~m1_subset_1(X103,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X102)))))))&(v7_waybel_0(k3_yellow19(X101,X102,X103))|(v2_struct_0(X101)|~l1_struct_0(X101)|v1_xboole_0(X102)|~m1_subset_1(X102,k1_zfmisc_1(u1_struct_0(X101)))|v1_xboole_0(X103)|~v1_subset_1(X103,u1_struct_0(k3_yellow_1(X102)))|~v2_waybel_0(X103,k3_yellow_1(X102))|~v13_waybel_0(X103,k3_yellow_1(X102))|~m1_subset_1(X103,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X102))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc5_yellow19])])])])).
cnf(c_0_133, plain, (~r2_hidden(X1,k3_subset_1(X2,X3))|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(csr,[status(thm)],[c_0_123, c_0_57])).
cnf(c_0_134, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r2_hidden(k2_waybel_0(X1,X2,esk7_4(X1,X2,X3,X4)),X3)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~r2_waybel_0(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))), inference(evalgc,[status(thm)],[c_0_124])).
cnf(c_0_135, plain, (r2_hidden(k2_waybel_0(X1,X2,esk7_4(X3,X2,X4,esk9_3(X1,X2,X5))),X5)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|~r1_waybel_0(X1,X2,X5)|~l1_struct_0(X1)|~m1_subset_1(esk7_4(X3,X2,X4,esk9_3(X1,X2,X5)),u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~r2_waybel_0(X3,X2,X4)|~l1_struct_0(X3)|~m1_subset_1(esk9_3(X1,X2,X5),u1_struct_0(X2))|~l1_waybel_0(X2,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126])])).
cnf(c_0_136, plain, (v2_struct_0(X2)|v2_struct_0(X1)|m1_subset_1(esk9_3(X1,X2,X3),u1_struct_0(X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~r1_waybel_0(X1,X2,X3)), inference(evalgc,[status(thm)],[c_0_127])).
cnf(c_0_137, negated_conjecture, (v1_xboole_0(X1)|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1))|k10_yellow_6(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1))!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~v3_waybel_7(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~v13_waybel_0(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1),esk1_0)|~v7_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1))|~v4_orders_2(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_128, c_0_129])])).
cnf(c_0_138, plain, (k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)))=k1_xboole_0|v1_compts_1(X1)|v2_struct_0(X1)|v2_struct_0(k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)))|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)),X1)|~v7_waybel_0(k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)))|~v4_orders_2(k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_131])])).
cnf(c_0_139, plain, (v7_waybel_0(k3_yellow19(X1,X2,X3))|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_subset_1(X3,u1_struct_0(k3_yellow_1(X2)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_132])).
fof(c_0_140, plain, ![X83, X84, X85]:((((~v2_struct_0(k3_yellow19(X83,X84,X85))|(v2_struct_0(X83)|~l1_struct_0(X83)|v1_xboole_0(X84)|~m1_subset_1(X84,k1_zfmisc_1(u1_struct_0(X83)))|v1_xboole_0(X85)|~v2_waybel_0(X85,k3_yellow_1(X84))|~v13_waybel_0(X85,k3_yellow_1(X84))|~m1_subset_1(X85,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X84))))))&(v3_orders_2(k3_yellow19(X83,X84,X85))|(v2_struct_0(X83)|~l1_struct_0(X83)|v1_xboole_0(X84)|~m1_subset_1(X84,k1_zfmisc_1(u1_struct_0(X83)))|v1_xboole_0(X85)|~v2_waybel_0(X85,k3_yellow_1(X84))|~v13_waybel_0(X85,k3_yellow_1(X84))|~m1_subset_1(X85,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X84)))))))&(v4_orders_2(k3_yellow19(X83,X84,X85))|(v2_struct_0(X83)|~l1_struct_0(X83)|v1_xboole_0(X84)|~m1_subset_1(X84,k1_zfmisc_1(u1_struct_0(X83)))|v1_xboole_0(X85)|~v2_waybel_0(X85,k3_yellow_1(X84))|~v13_waybel_0(X85,k3_yellow_1(X84))|~m1_subset_1(X85,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X84)))))))&(v6_waybel_0(k3_yellow19(X83,X84,X85),X83)|(v2_struct_0(X83)|~l1_struct_0(X83)|v1_xboole_0(X84)|~m1_subset_1(X84,k1_zfmisc_1(u1_struct_0(X83)))|v1_xboole_0(X85)|~v2_waybel_0(X85,k3_yellow_1(X84))|~v13_waybel_0(X85,k3_yellow_1(X84))|~m1_subset_1(X85,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X84))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc4_yellow19])])])])).
cnf(c_0_141, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_hidden(k2_waybel_0(X1,X2,esk7_4(X1,X2,k3_subset_1(X3,X4),X5)),X4)|~m1_subset_1(X4,k1_zfmisc_1(X3))|~r2_waybel_0(X1,X2,k3_subset_1(X3,X4))|~l1_struct_0(X1)|~m1_subset_1(X5,u1_struct_0(X2))|~l1_waybel_0(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_134])])).
cnf(c_0_142, plain, (r2_hidden(k2_waybel_0(X1,X2,esk7_4(X3,X2,X4,esk9_3(X1,X2,X5))),X5)|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X3,X2,X4)|~r1_waybel_0(X1,X2,X5)|~l1_struct_0(X1)|~l1_struct_0(X3)|~m1_subset_1(esk7_4(X3,X2,X4,esk9_3(X1,X2,X5)),u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_waybel_0(X2,X3)), inference(csr,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_143, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)))|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~v3_waybel_7(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)),esk1_0)|~v7_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)))|~v4_orders_2(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)))|~v1_subset_1(esk4_1(esk1_0),u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0))))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_138]), c_0_98]), c_0_99])]), c_0_100])])).
cnf(c_0_144, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|v7_waybel_0(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_subset_1(X3,u1_struct_0(k3_yellow_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(evalgc,[status(thm)],[c_0_139])).
cnf(c_0_145, plain, (v4_orders_2(k3_yellow19(X1,X2,X3))|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_140])).
fof(c_0_146, plain, ![X188, X189]:((((~v1_xboole_0(X189)|(v1_xboole_0(X189)|~v2_waybel_0(X189,X188)|~v13_waybel_0(X189,X188)|~v3_waybel_7(X189,X188))|~m1_subset_1(X189,k1_zfmisc_1(u1_struct_0(X188)))|(v2_struct_0(X188)|~v3_orders_2(X188)|~v4_orders_2(X188)|~v5_orders_2(X188)|~l1_orders_2(X188)))&(v1_subset_1(X189,u1_struct_0(X188))|(v1_xboole_0(X189)|~v2_waybel_0(X189,X188)|~v13_waybel_0(X189,X188)|~v3_waybel_7(X189,X188))|~m1_subset_1(X189,k1_zfmisc_1(u1_struct_0(X188)))|(v2_struct_0(X188)|~v3_orders_2(X188)|~v4_orders_2(X188)|~v5_orders_2(X188)|~l1_orders_2(X188))))&(v2_waybel_0(X189,X188)|(v1_xboole_0(X189)|~v2_waybel_0(X189,X188)|~v13_waybel_0(X189,X188)|~v3_waybel_7(X189,X188))|~m1_subset_1(X189,k1_zfmisc_1(u1_struct_0(X188)))|(v2_struct_0(X188)|~v3_orders_2(X188)|~v4_orders_2(X188)|~v5_orders_2(X188)|~l1_orders_2(X188))))&(v13_waybel_0(X189,X188)|(v1_xboole_0(X189)|~v2_waybel_0(X189,X188)|~v13_waybel_0(X189,X188)|~v3_waybel_7(X189,X188))|~m1_subset_1(X189,k1_zfmisc_1(u1_struct_0(X188)))|(v2_struct_0(X188)|~v3_orders_2(X188)|~v4_orders_2(X188)|~v5_orders_2(X188)|~l1_orders_2(X188)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_waybel_7])])])])])).
fof(c_0_147, plain, ![X134]:(v1_orders_2(k3_yellow_1(X134))&l1_orders_2(k3_yellow_1(X134))), inference(variable_rename,[status(thm)],[dt_k3_yellow_1])).
fof(c_0_148, plain, ![X88]:((((~v2_struct_0(k3_yellow_1(X88))&v1_orders_2(k3_yellow_1(X88)))&v3_orders_2(k3_yellow_1(X88)))&v4_orders_2(k3_yellow_1(X88)))&v5_orders_2(k3_yellow_1(X88))), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc7_yellow_1])])).
cnf(c_0_149, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X2,X1,k3_subset_1(X3,X4))|~l1_struct_0(X2)|~m1_subset_1(X4,k1_zfmisc_1(X3))|~m1_subset_1(esk9_3(X2,X1,X4),u1_struct_0(X1))|~l1_waybel_0(X1,X2)|~r1_waybel_0(X2,X1,X4)|~m1_subset_1(esk7_4(X2,X1,k3_subset_1(X3,X4),esk9_3(X2,X1,X4)),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_142])])).
cnf(c_0_150, plain, (m1_subset_1(esk7_4(X1,X2,X3,X4),u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X2))|~r2_waybel_0(X1,X2,X3)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_151, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)))|v1_xboole_0(k2_struct_0(esk1_0))|~v1_subset_1(esk4_1(esk1_0),u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0))))|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~v3_waybel_7(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)),esk1_0)|~v4_orders_2(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144]), c_0_129])]), c_0_100])])).
cnf(c_0_152, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|v4_orders_2(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(evalgc,[status(thm)],[c_0_145])).
cnf(c_0_153, plain, (v1_subset_1(X1,u1_struct_0(X2))|v1_xboole_0(X1)|v2_struct_0(X2)|~v2_waybel_0(X1,X2)|~v13_waybel_0(X1,X2)|~v3_waybel_7(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_146])).
cnf(c_0_154, plain, (l1_orders_2(k3_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_147])).
cnf(c_0_155, plain, (v5_orders_2(k3_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_148])).
cnf(c_0_156, plain, (v3_orders_2(k3_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_148])).
cnf(c_0_157, plain, (v4_orders_2(k3_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_148])).
cnf(c_0_158, plain, (~v2_struct_0(k3_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_148])).
fof(c_0_159, plain, ![X80, X81, X82]:(((~v2_struct_0(k3_yellow19(X80,X81,X82))|(v2_struct_0(X80)|~l1_struct_0(X80)|v1_xboole_0(X81)|~m1_subset_1(X81,k1_zfmisc_1(u1_struct_0(X80)))|v1_xboole_0(X82)|~v2_waybel_0(X82,k3_yellow_1(X81))|~v13_waybel_0(X82,k3_yellow_1(X81))|~m1_subset_1(X82,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X81))))))&(v6_waybel_0(k3_yellow19(X80,X81,X82),X80)|(v2_struct_0(X80)|~l1_struct_0(X80)|v1_xboole_0(X81)|~m1_subset_1(X81,k1_zfmisc_1(u1_struct_0(X80)))|v1_xboole_0(X82)|~v2_waybel_0(X82,k3_yellow_1(X81))|~v13_waybel_0(X82,k3_yellow_1(X81))|~m1_subset_1(X82,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X81)))))))&(l1_waybel_0(k3_yellow19(X80,X81,X82),X80)|(v2_struct_0(X80)|~l1_struct_0(X80)|v1_xboole_0(X81)|~m1_subset_1(X81,k1_zfmisc_1(u1_struct_0(X80)))|v1_xboole_0(X82)|~v2_waybel_0(X82,k3_yellow_1(X81))|~v13_waybel_0(X82,k3_yellow_1(X81))|~m1_subset_1(X82,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X81))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_yellow19])])])])).
fof(c_0_160, plain, ![X93, X94]:((~m1_subset_1(X93,k1_zfmisc_1(X94))|r1_tarski(X93,X94))&(~r1_tarski(X93,X94)|m1_subset_1(X93,k1_zfmisc_1(X94)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_161, plain, ![X104]:r1_tarski(X104,X104), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r1_tarski])])).
fof(c_0_162, plain, ![X66]:(~l1_struct_0(X66)|k2_struct_0(X66)=u1_struct_0(X66)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_163, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X1,X2,k3_subset_1(X3,X4))|~r1_waybel_0(X1,X2,X4)|~l1_struct_0(X1)|~m1_subset_1(esk7_4(X1,X2,k3_subset_1(X3,X4),esk9_3(X1,X2,X4)),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(X3))|~l1_waybel_0(X2,X1)), inference(csr,[status(thm)],[c_0_149, c_0_136])).
cnf(c_0_164, plain, (v2_struct_0(X2)|v2_struct_0(X1)|m1_subset_1(esk7_4(X1,X2,X3,X4),u1_struct_0(X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~r2_waybel_0(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))), inference(evalgc,[status(thm)],[c_0_150])).
cnf(c_0_165, negated_conjecture, (v1_xboole_0(k2_struct_0(esk1_0))|v1_xboole_0(esk4_1(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)))|~v1_subset_1(esk4_1(esk1_0),u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0))))|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v3_waybel_7(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)),esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152]), c_0_129])]), c_0_100])])).
cnf(c_0_166, plain, (v2_struct_0(X2)|v1_xboole_0(X1)|v1_subset_1(X1,u1_struct_0(X2))|~v4_orders_2(X2)|~v3_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~v2_waybel_0(X1,X2)|~v13_waybel_0(X1,X2)|~v3_waybel_7(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))), inference(evalgc,[status(thm)],[c_0_153])).
cnf(c_0_167, plain, (l1_orders_2(k3_yellow_1(X1))), inference(evalgc,[status(thm)],[c_0_154])).
cnf(c_0_168, plain, (v5_orders_2(k3_yellow_1(X1))), inference(evalgc,[status(thm)],[c_0_155])).
cnf(c_0_169, plain, (v3_orders_2(k3_yellow_1(X1))), inference(evalgc,[status(thm)],[c_0_156])).
cnf(c_0_170, plain, (v4_orders_2(k3_yellow_1(X1))), inference(evalgc,[status(thm)],[c_0_157])).
cnf(c_0_171, plain, (~v2_struct_0(k3_yellow_1(X1))), inference(evalgc,[status(thm)],[c_0_158])).
cnf(c_0_172, plain, (l1_waybel_0(k3_yellow19(X1,X2,X3),X1)|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_159])).
cnf(c_0_173, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_160])).
cnf(c_0_174, plain, (r1_tarski(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_161])).
cnf(c_0_175, plain, (v3_waybel_7(esk4_1(X1),k3_yellow_1(k2_struct_0(X1)))|v1_compts_1(X1)|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_176, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_162])).
cnf(c_0_177, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X1,X2,k3_subset_1(X3,X4))|~r1_waybel_0(X1,X2,X4)|~l1_struct_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(X3))|~l1_waybel_0(X2,X1)|~m1_subset_1(esk9_3(X1,X2,X4),u1_struct_0(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_164])])).
cnf(c_0_178, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_xboole_0(k2_struct_0(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)))|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v3_waybel_7(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)),esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_166]), c_0_167]), c_0_168]), c_0_169]), c_0_170])]), c_0_171])])).
cnf(c_0_179, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|l1_waybel_0(k3_yellow19(X1,X2,X3),X1)|~l1_struct_0(X1)|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(evalgc,[status(thm)],[c_0_172])).
cnf(c_0_180, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_173])).
cnf(c_0_181, plain, (r1_tarski(X1,X1)), inference(evalgc,[status(thm)],[c_0_174])).
cnf(c_0_182, plain, (v2_struct_0(X1)|v1_compts_1(X1)|v3_waybel_7(esk4_1(X1),k3_yellow_1(k2_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[c_0_175])).
cnf(c_0_183, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_176])).
cnf(c_0_184, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X2,X1,k3_subset_1(X3,X4))|~r1_waybel_0(X2,X1,X4)|~l1_struct_0(X2)|~m1_subset_1(X4,k1_zfmisc_1(X3))|~l1_waybel_0(X1,X2)), inference(csr,[status(thm)],[c_0_177, c_0_136])).
fof(c_0_185, plain, ![X19, X20, X21]:((~v1_yellow19(X20,X19)|(~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X19)))|(r1_waybel_0(X19,X20,X21)|r1_waybel_0(X19,X20,k3_subset_1(u1_struct_0(X19),X21))))|(v2_struct_0(X20)|~l1_waybel_0(X20,X19))|(v2_struct_0(X19)|~l1_struct_0(X19)))&((m1_subset_1(esk5_2(X19,X20),k1_zfmisc_1(u1_struct_0(X19)))|v1_yellow19(X20,X19)|(v2_struct_0(X20)|~l1_waybel_0(X20,X19))|(v2_struct_0(X19)|~l1_struct_0(X19)))&((~r1_waybel_0(X19,X20,esk5_2(X19,X20))|v1_yellow19(X20,X19)|(v2_struct_0(X20)|~l1_waybel_0(X20,X19))|(v2_struct_0(X19)|~l1_struct_0(X19)))&(~r1_waybel_0(X19,X20,k3_subset_1(u1_struct_0(X19),esk5_2(X19,X20)))|v1_yellow19(X20,X19)|(v2_struct_0(X20)|~l1_waybel_0(X20,X19))|(v2_struct_0(X19)|~l1_struct_0(X19)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_yellow19])])])])])])).
fof(c_0_186, plain, ![X23, X24, X25, X26]:((~r3_waybel_9(X23,X24,X25)|(~m1_connsp_2(X26,X23,X25)|r2_waybel_0(X23,X24,X26))|~m1_subset_1(X25,u1_struct_0(X23))|(v2_struct_0(X24)|~l1_waybel_0(X24,X23))|(v2_struct_0(X23)|~v2_pre_topc(X23)|~l1_pre_topc(X23)))&((m1_connsp_2(esk6_3(X23,X24,X25),X23,X25)|r3_waybel_9(X23,X24,X25)|~m1_subset_1(X25,u1_struct_0(X23))|(v2_struct_0(X24)|~l1_waybel_0(X24,X23))|(v2_struct_0(X23)|~v2_pre_topc(X23)|~l1_pre_topc(X23)))&(~r2_waybel_0(X23,X24,esk6_3(X23,X24,X25))|r3_waybel_9(X23,X24,X25)|~m1_subset_1(X25,u1_struct_0(X23))|(v2_struct_0(X24)|~l1_waybel_0(X24,X23))|(v2_struct_0(X23)|~v2_pre_topc(X23)|~l1_pre_topc(X23))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d9_waybel_9])])])])])])).
fof(c_0_187, plain, ![X47, X48, X49, X50, X51, X55]:(((~r2_hidden(X50,X49)|(~m1_connsp_2(X51,X47,X50)|r1_waybel_0(X47,X48,X51))|~m1_subset_1(X50,u1_struct_0(X47))|X49!=k10_yellow_6(X47,X48)|~m1_subset_1(X49,k1_zfmisc_1(u1_struct_0(X47)))|(v2_struct_0(X48)|~v4_orders_2(X48)|~v7_waybel_0(X48)|~l1_waybel_0(X48,X47))|(v2_struct_0(X47)|~v2_pre_topc(X47)|~l1_pre_topc(X47)))&((m1_connsp_2(esk11_4(X47,X48,X49,X50),X47,X50)|r2_hidden(X50,X49)|~m1_subset_1(X50,u1_struct_0(X47))|X49!=k10_yellow_6(X47,X48)|~m1_subset_1(X49,k1_zfmisc_1(u1_struct_0(X47)))|(v2_struct_0(X48)|~v4_orders_2(X48)|~v7_waybel_0(X48)|~l1_waybel_0(X48,X47))|(v2_struct_0(X47)|~v2_pre_topc(X47)|~l1_pre_topc(X47)))&(~r1_waybel_0(X47,X48,esk11_4(X47,X48,X49,X50))|r2_hidden(X50,X49)|~m1_subset_1(X50,u1_struct_0(X47))|X49!=k10_yellow_6(X47,X48)|~m1_subset_1(X49,k1_zfmisc_1(u1_struct_0(X47)))|(v2_struct_0(X48)|~v4_orders_2(X48)|~v7_waybel_0(X48)|~l1_waybel_0(X48,X47))|(v2_struct_0(X47)|~v2_pre_topc(X47)|~l1_pre_topc(X47)))))&((m1_subset_1(esk12_3(X47,X48,X49),u1_struct_0(X47))|X49=k10_yellow_6(X47,X48)|~m1_subset_1(X49,k1_zfmisc_1(u1_struct_0(X47)))|(v2_struct_0(X48)|~v4_orders_2(X48)|~v7_waybel_0(X48)|~l1_waybel_0(X48,X47))|(v2_struct_0(X47)|~v2_pre_topc(X47)|~l1_pre_topc(X47)))&(((m1_connsp_2(esk13_3(X47,X48,X49),X47,esk12_3(X47,X48,X49))|~r2_hidden(esk12_3(X47,X48,X49),X49)|X49=k10_yellow_6(X47,X48)|~m1_subset_1(X49,k1_zfmisc_1(u1_struct_0(X47)))|(v2_struct_0(X48)|~v4_orders_2(X48)|~v7_waybel_0(X48)|~l1_waybel_0(X48,X47))|(v2_struct_0(X47)|~v2_pre_topc(X47)|~l1_pre_topc(X47)))&(~r1_waybel_0(X47,X48,esk13_3(X47,X48,X49))|~r2_hidden(esk12_3(X47,X48,X49),X49)|X49=k10_yellow_6(X47,X48)|~m1_subset_1(X49,k1_zfmisc_1(u1_struct_0(X47)))|(v2_struct_0(X48)|~v4_orders_2(X48)|~v7_waybel_0(X48)|~l1_waybel_0(X48,X47))|(v2_struct_0(X47)|~v2_pre_topc(X47)|~l1_pre_topc(X47))))&(r2_hidden(esk12_3(X47,X48,X49),X49)|(~m1_connsp_2(X55,X47,esk12_3(X47,X48,X49))|r1_waybel_0(X47,X48,X55))|X49=k10_yellow_6(X47,X48)|~m1_subset_1(X49,k1_zfmisc_1(u1_struct_0(X47)))|(v2_struct_0(X48)|~v4_orders_2(X48)|~v7_waybel_0(X48)|~l1_waybel_0(X48,X47))|(v2_struct_0(X47)|~v2_pre_topc(X47)|~l1_pre_topc(X47)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d15_yellow_6])])])])])])).
cnf(c_0_188, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~v2_struct_0(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_159])).
cnf(c_0_189, negated_conjecture, (v1_xboole_0(k2_struct_0(esk1_0))|v1_xboole_0(esk4_1(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)))|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v3_waybel_7(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_179]), c_0_129])]), c_0_100])])).
cnf(c_0_190, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_181])])).
cnf(c_0_191, plain, (v3_waybel_7(esk4_1(X1),k3_yellow_1(u1_struct_0(X1)))|v1_compts_1(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_183])])).
cnf(c_0_192, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X2,X1,X4)|~r1_waybel_0(X2,X1,k3_subset_1(X3,X4))|~l1_struct_0(X2)|~m1_subset_1(k3_subset_1(X3,X4),k1_zfmisc_1(X3))|~l1_waybel_0(X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_50])])).
cnf(c_0_193, plain, (r1_waybel_0(X2,X1,X3)|r1_waybel_0(X2,X1,k3_subset_1(u1_struct_0(X2),X3))|v2_struct_0(X1)|v2_struct_0(X2)|~v1_yellow19(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)), inference(split_conjunct,[status(thm)],[c_0_185])).
cnf(c_0_194, plain, (r2_waybel_0(X1,X2,X4)|v2_struct_0(X2)|v2_struct_0(X1)|~r3_waybel_9(X1,X2,X3)|~m1_connsp_2(X4,X1,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_186])).
cnf(c_0_195, plain, (m1_connsp_2(esk11_4(X1,X2,X3,X4),X1,X4)|r2_hidden(X4,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k10_yellow_6(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_187])).
fof(c_0_196, plain, ![X63, X64, X65]:(v2_struct_0(X63)|~v2_pre_topc(X63)|~l1_pre_topc(X63)|~m1_subset_1(X64,u1_struct_0(X63))|(~m1_connsp_2(X65,X63,X64)|m1_subset_1(X65,k1_zfmisc_1(u1_struct_0(X63))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m1_connsp_2])])])])).
cnf(c_0_197, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|~l1_struct_0(X1)|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~v2_struct_0(k3_yellow19(X1,X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(evalgc,[status(thm)],[c_0_188])).
cnf(c_0_198, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_xboole_0(k2_struct_0(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)))|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_182]), c_0_98]), c_0_99])]), c_0_100])])).
cnf(c_0_199, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_xboole_0(u1_struct_0(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk4_1(esk1_0)))|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))|~v3_waybel_7(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_183]), c_0_190]), c_0_129])])])).
cnf(c_0_200, plain, (v3_waybel_7(esk4_1(X1),k3_yellow_1(u1_struct_0(X1)))|v1_compts_1(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_191, c_0_116])).
cnf(c_0_201, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X1,X2,X3)|~r1_waybel_0(X1,X2,k3_subset_1(X4,X3))|~l1_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(X4))|~l1_waybel_0(X2,X1)), inference(csr,[status(thm)],[c_0_192, c_0_57])).
cnf(c_0_202, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r1_waybel_0(X2,X1,X3)|r1_waybel_0(X2,X1,k3_subset_1(u1_struct_0(X2),X3))|~l1_struct_0(X2)|~l1_waybel_0(X1,X2)|~v1_yellow19(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(evalgc,[status(thm)],[c_0_193])).
cnf(c_0_203, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r2_waybel_0(X1,X2,X4)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~l1_waybel_0(X2,X1)|~r3_waybel_9(X1,X2,X3)|~m1_connsp_2(X4,X1,X3)|~m1_subset_1(X3,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_194])).
cnf(c_0_204, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r2_hidden(X4,X3)|m1_connsp_2(esk11_4(X1,X2,X3,X4),X1,X4)|X3!=k10_yellow_6(X1,X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[c_0_195])).
cnf(c_0_205, plain, (v2_struct_0(X1)|m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_connsp_2(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_196])).
cnf(c_0_206, negated_conjecture, (v1_xboole_0(k2_struct_0(esk1_0))|v1_xboole_0(esk4_1(esk1_0))|v1_compts_1(esk1_0)|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_198]), c_0_129])]), c_0_100])])).
cnf(c_0_207, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(esk4_1(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk4_1(esk1_0)))|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_200]), c_0_98]), c_0_99])]), c_0_100])])).
cnf(c_0_208, plain, (v2_waybel_0(esk4_1(X1),k3_yellow_1(u1_struct_0(X1)))|v1_compts_1(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_183])])).
cnf(c_0_209, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r1_waybel_0(X1,X2,X3)|~r2_waybel_0(X1,X2,X3)|~l1_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_waybel_0(X2,X1)|~v1_yellow19(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_202])])).
cnf(c_0_210, plain, (r2_waybel_0(X1,X2,esk11_4(X1,X3,X4,X5))|v2_struct_0(X1)|v2_struct_0(X2)|r2_hidden(X5,X4)|v2_struct_0(X3)|~r3_waybel_9(X1,X2,X5)|~m1_subset_1(X5,u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|X4!=k10_yellow_6(X1,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~l1_waybel_0(X3,X1)|~v7_waybel_0(X3)|~v4_orders_2(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_204])])).
cnf(c_0_211, plain, (v2_struct_0(X1)|m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_connsp_2(X3,X1,X2)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_205])).
cnf(c_0_212, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_xboole_0(k2_struct_0(esk1_0))|v1_compts_1(esk1_0)|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_118]), c_0_98]), c_0_99])]), c_0_100])])).
cnf(c_0_213, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(esk4_1(esk1_0))|v1_compts_1(esk1_0)|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_207]), c_0_129]), c_0_190])]), c_0_100])])).
cnf(c_0_214, plain, (v2_waybel_0(esk4_1(X1),k3_yellow_1(u1_struct_0(X1)))|v1_compts_1(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_208, c_0_116])).
cnf(c_0_215, plain, (v13_waybel_0(esk4_1(X1),k3_yellow_1(u1_struct_0(X1)))|v1_compts_1(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_183])])).
cnf(c_0_216, plain, (r2_hidden(X4,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_waybel_0(X1,X2,esk11_4(X1,X2,X3,X4))|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k10_yellow_6(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_187])).
cnf(c_0_217, plain, (r1_waybel_0(X1,X2,esk11_4(X1,X3,X4,X5))|v2_struct_0(X2)|v2_struct_0(X1)|r2_hidden(X5,X4)|v2_struct_0(X3)|~l1_struct_0(X1)|~m1_subset_1(esk11_4(X1,X3,X4,X5),k1_zfmisc_1(u1_struct_0(X1)))|~v1_yellow19(X2,X1)|~l1_waybel_0(X2,X1)|X4!=k10_yellow_6(X1,X3)|~r3_waybel_9(X1,X2,X5)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,u1_struct_0(X1))|~l1_waybel_0(X3,X1)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_210])])).
cnf(c_0_218, plain, (m1_subset_1(esk11_4(X1,X2,X3,X4),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|r2_hidden(X4,X3)|v2_struct_0(X2)|~m1_subset_1(X4,u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|X3!=k10_yellow_6(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_204])])).
cnf(c_0_219, negated_conjecture, (v1_xboole_0(k2_struct_0(esk1_0))|v1_xboole_0(esk4_1(esk1_0))|v1_compts_1(esk1_0)|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_119]), c_0_98]), c_0_99])]), c_0_100])])).
cnf(c_0_220, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_xboole_0(u1_struct_0(esk1_0))|v1_compts_1(esk1_0)|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_214]), c_0_98]), c_0_99])]), c_0_100])])).
cnf(c_0_221, plain, (v13_waybel_0(esk4_1(X1),k3_yellow_1(u1_struct_0(X1)))|v1_compts_1(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_215, c_0_116])).
cnf(c_0_222, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X1)))))|v1_compts_1(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_183])])).
cnf(c_0_223, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r2_hidden(X4,X3)|X3!=k10_yellow_6(X1,X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_waybel_0(X1,X2,esk11_4(X1,X2,X3,X4))), inference(evalgc,[status(thm)],[c_0_216])).
cnf(c_0_224, plain, (r2_hidden(X1,X2)|r1_waybel_0(X3,X4,esk11_4(X3,X5,X2,X1))|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X4)|X2!=k10_yellow_6(X3,X5)|~r3_waybel_9(X3,X4,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~v1_yellow19(X4,X3)|~l1_waybel_0(X4,X3)|~l1_waybel_0(X5,X3)|~v7_waybel_0(X5)|~v4_orders_2(X5)|~l1_pre_topc(X3)|~v2_pre_topc(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_217, c_0_218]), c_0_116])).
fof(c_0_225, plain, ![X56, X57]:((m1_subset_1(esk14_2(X56,X57),u1_struct_0(X56))|(v2_struct_0(X57)|~v4_orders_2(X57)|~v7_waybel_0(X57)|~l1_waybel_0(X57,X56))|~v1_compts_1(X56)|(v2_struct_0(X56)|~v2_pre_topc(X56)|~l1_pre_topc(X56)))&(r3_waybel_9(X56,X57,esk14_2(X56,X57))|(v2_struct_0(X57)|~v4_orders_2(X57)|~v7_waybel_0(X57)|~l1_waybel_0(X57,X56))|~v1_compts_1(X56)|(v2_struct_0(X56)|~v2_pre_topc(X56)|~l1_pre_topc(X56)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l35_yellow19])])])])])])).
cnf(c_0_226, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_227, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_xboole_0(k2_struct_0(esk1_0))|v1_compts_1(esk1_0)|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_120]), c_0_98]), c_0_99])]), c_0_100])])).
fof(c_0_228, plain, ![X79]:(~l1_struct_0(X79)|m1_subset_1(k2_struct_0(X79),k1_zfmisc_1(u1_struct_0(X79)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_struct_0])])).
cnf(c_0_229, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(esk4_1(esk1_0))|v1_compts_1(esk1_0)|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_221]), c_0_98]), c_0_99])]), c_0_100])])).
cnf(c_0_230, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X1)))))|v1_compts_1(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_222, c_0_116])).
cnf(c_0_231, plain, (r2_hidden(X1,X2)|v2_struct_0(X3)|v2_struct_0(X4)|X2!=k10_yellow_6(X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_waybel_0(X4,X3)|~v7_waybel_0(X4)|~v4_orders_2(X4)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~r3_waybel_9(X3,X4,X1)|~v1_yellow19(X4,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_224])])).
cnf(c_0_232, plain, (r3_waybel_9(X1,X2,esk14_2(X1,X2))|v2_struct_0(X2)|v2_struct_0(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v1_compts_1(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_225])).
fof(c_0_233, plain, ![X151]:(v2_struct_0(X151)|~l1_struct_0(X151)|~v1_xboole_0(k2_struct_0(X151))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc5_struct_0])])])).
cnf(c_0_234, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0)|~v1_compts_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_226])).
cnf(c_0_235, negated_conjecture, (v1_compts_1(esk1_0)|v1_xboole_0(k2_struct_0(esk1_0))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_227]), c_0_98]), c_0_99])]), c_0_100])])).
cnf(c_0_236, plain, (m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_228])).
fof(c_0_237, plain, ![X187]:(v2_struct_0(X187)|~l1_struct_0(X187)|~v1_xboole_0(u1_struct_0(X187))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_238, negated_conjecture, (~v2_struct_0(esk2_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_239, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_xboole_0(u1_struct_0(esk1_0))|v1_compts_1(esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_230]), c_0_98]), c_0_99])]), c_0_100])])).
cnf(c_0_240, plain, (v2_struct_0(X2)|v2_struct_0(X3)|k10_yellow_6(X3,X2)!=k1_xboole_0|~r3_waybel_9(X3,X2,X1)|~m1_subset_1(X1,u1_struct_0(X3))|~v1_yellow19(X2,X3)|~l1_waybel_0(X2,X3)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X3)|~v2_pre_topc(X3)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_63]), c_0_77])])).
cnf(c_0_241, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r3_waybel_9(X1,X2,esk14_2(X1,X2))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_compts_1(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)), inference(evalgc,[status(thm)],[c_0_232])).
cnf(c_0_242, plain, (m1_subset_1(esk14_2(X1,X2),u1_struct_0(X1))|v2_struct_0(X2)|v2_struct_0(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v1_compts_1(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_225])).
cnf(c_0_243, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(k2_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_233])).
cnf(c_0_244, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0)|v1_xboole_0(k2_struct_0(esk1_0))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_235])])).
cnf(c_0_245, plain, (m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_236])).
cnf(c_0_246, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_237])).
cnf(c_0_247, negated_conjecture, (~v2_struct_0(esk2_0)|~v1_compts_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_238])).
cnf(c_0_248, negated_conjecture, (v1_compts_1(esk1_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_239]), c_0_98]), c_0_99])]), c_0_100])])).
cnf(c_0_249, negated_conjecture, (~v3_yellow_6(esk2_0,esk1_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_250, plain, (k10_yellow_6(X1,X2)=k1_xboole_0|v3_yellow_6(X2,X1)|v2_struct_0(X2)|v2_struct_0(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_251, plain, (v2_struct_0(X1)|v2_struct_0(X2)|k10_yellow_6(X1,X2)!=k1_xboole_0|~m1_subset_1(esk14_2(X1,X2),u1_struct_0(X1))|~v1_yellow19(X2,X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_compts_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_241])])).
cnf(c_0_252, plain, (v2_struct_0(X2)|v2_struct_0(X1)|m1_subset_1(esk14_2(X1,X2),u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_compts_1(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)), inference(evalgc,[status(thm)],[c_0_242])).
cnf(c_0_253, negated_conjecture, (v1_yellow19(esk2_0,esk1_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_254, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(k2_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_243])).
cnf(c_0_255, negated_conjecture, (v1_xboole_0(k2_struct_0(esk1_0))|l1_waybel_0(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_244, c_0_245]), c_0_129])])])).
cnf(c_0_256, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_246])).
cnf(c_0_257, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~v2_struct_0(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_248])])).
cnf(c_0_258, negated_conjecture, (~v1_compts_1(esk1_0)|~v3_yellow_6(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_249])).
cnf(c_0_259, plain, (k10_yellow_6(X1,X2)=k1_xboole_0|v2_struct_0(X2)|v2_struct_0(X1)|v3_yellow_6(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)), inference(evalgc,[status(thm)],[c_0_250])).
cnf(c_0_260, negated_conjecture, (v4_orders_2(esk2_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_261, negated_conjecture, (v7_waybel_0(esk2_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_262, plain, (v2_struct_0(X1)|v2_struct_0(X2)|k10_yellow_6(X2,X1)!=k1_xboole_0|~v1_yellow19(X1,X2)|~l1_waybel_0(X1,X2)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~v1_compts_1(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[c_0_251, c_0_252])).
cnf(c_0_263, negated_conjecture, (v1_yellow19(esk2_0,esk1_0)|~v1_compts_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_253])).
cnf(c_0_264, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_255]), c_0_129])]), c_0_100])])).
cnf(c_0_265, negated_conjecture, (~v2_struct_0(esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_256, c_0_257]), c_0_129])]), c_0_100])])).
cnf(c_0_266, negated_conjecture, (k10_yellow_6(esk1_0,esk2_0)=k1_xboole_0|v2_struct_0(esk2_0)|~v1_compts_1(esk1_0)|~l1_waybel_0(esk2_0,esk1_0)|~v7_waybel_0(esk2_0)|~v4_orders_2(esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_259]), c_0_98]), c_0_99])]), c_0_100])])).
cnf(c_0_267, negated_conjecture, (v4_orders_2(esk2_0)|~v1_compts_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_260])).
cnf(c_0_268, negated_conjecture, (v7_waybel_0(esk2_0)|~v1_compts_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_261])).
cnf(c_0_269, negated_conjecture, (k10_yellow_6(esk1_0,esk2_0)!=k1_xboole_0|~v7_waybel_0(esk2_0)|~v4_orders_2(esk2_0)|~v1_compts_1(esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_263]), c_0_264]), c_0_98]), c_0_99])]), c_0_265]), c_0_100])])).
cnf(c_0_270, negated_conjecture, (k10_yellow_6(esk1_0,esk2_0)=k1_xboole_0|~v1_compts_1(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_266, c_0_267]), c_0_268]), c_0_234]), c_0_247])).
cnf(c_0_271, negated_conjecture, (~v1_compts_1(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_269, c_0_267]), c_0_268]), c_0_270])).
cnf(c_0_272, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[c_0_248, c_0_271])])).
cnf(c_0_273, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_256, c_0_272]), c_0_129])]), c_0_100]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 274
# Proof object clause steps            : 203
# Proof object formula steps           : 71
# Proof object conjectures             : 56
# Proof object clause conjectures      : 53
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 62
# Proof object initial formulas used   : 35
# Proof object generating inferences   : 60
# Proof object simplifying inferences  : 128
# Training examples: 123 positive, 792 negative
# Training: Positive examples begin
cnf(c_0_99, negated_conjecture, (v2_pre_topc(esk1_0))).# trainpos
cnf(c_0_98, negated_conjecture, (l1_pre_topc(esk1_0))).# trainpos
cnf(c_0_100, negated_conjecture, (~v2_struct_0(esk1_0))).# trainpos
cnf(c_0_170, plain, (v4_orders_2(k3_yellow_1(X1)))).# trainpos
cnf(c_0_54, plain, (v1_xboole_0(esk17_1(X1)))).# trainpos
cnf(c_0_169, plain, (v3_orders_2(k3_yellow_1(X1)))).# trainpos
cnf(c_0_247, negated_conjecture, (~v2_struct_0(esk2_0)|~v1_compts_1(esk1_0))).# trainpos
cnf(c_0_258, negated_conjecture, (~v1_compts_1(esk1_0)|~v3_yellow_6(esk2_0,esk1_0))).# trainpos
cnf(c_0_267, negated_conjecture, (v4_orders_2(esk2_0)|~v1_compts_1(esk1_0))).# trainpos
cnf(c_0_171, plain, (~v2_struct_0(k3_yellow_1(X1)))).# trainpos
cnf(c_0_268, negated_conjecture, (v7_waybel_0(esk2_0)|~v1_compts_1(esk1_0))).# trainpos
cnf(c_0_263, negated_conjecture, (v1_yellow19(esk2_0,esk1_0)|~v1_compts_1(esk1_0))).# trainpos
cnf(c_0_234, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0)|~v1_compts_1(esk1_0))).# trainpos
cnf(c_0_53, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_168, plain, (v5_orders_2(k3_yellow_1(X1)))).# trainpos
cnf(c_0_167, plain, (l1_orders_2(k3_yellow_1(X1)))).# trainpos
cnf(c_0_79, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0)).# trainpos
cnf(c_0_65, plain, (k4_xboole_0(X1,k1_xboole_0)=X1)).# trainpos
cnf(c_0_181, plain, (r1_tarski(X1,X1))).# trainpos
cnf(c_0_116, plain, (l1_struct_0(X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_97, negated_conjecture, (v2_struct_0(X1)|v1_compts_1(esk1_0)|v3_yellow_6(X1,esk1_0)|~v4_orders_2(X1)|~v7_waybel_0(X1)|~l1_waybel_0(X1,esk1_0)|~v1_yellow19(X1,esk1_0))).# trainpos
cnf(c_0_121, plain, (v2_struct_0(X1)|v1_compts_1(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_xboole_0(esk4_1(X1)))).# trainpos
cnf(c_0_183, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1))).# trainpos
cnf(c_0_58, plain, (m1_subset_1(esk17_1(X1),k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_180, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_96, plain, (v2_struct_0(X2)|v2_struct_0(X1)|k10_yellow_6(X2,X1)!=k1_xboole_0|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v4_orders_2(X1)|~v7_waybel_0(X1)|~l1_waybel_0(X1,X2)|~v3_yellow_6(X1,X2))).# trainpos
cnf(c_0_254, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(k2_struct_0(X1)))).# trainpos
cnf(c_0_256, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).# trainpos
cnf(c_0_245, plain, (m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).# trainpos
cnf(c_0_64, plain, (X3!=k4_xboole_0(X4,X2)|~r2_hidden(X1,X3)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_259, plain, (k10_yellow_6(X1,X2)=k1_xboole_0|v2_struct_0(X2)|v2_struct_0(X1)|v3_yellow_6(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1))).# trainpos
cnf(c_0_203, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r2_waybel_0(X1,X2,X4)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~l1_waybel_0(X2,X1)|~r3_waybel_9(X1,X2,X3)|~m1_connsp_2(X4,X1,X3)|~m1_subset_1(X3,u1_struct_0(X1)))).# trainpos
cnf(c_0_44, plain, (k4_xboole_0(X2,X1)=k3_subset_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_50, plain, (k3_subset_1(X2,k3_subset_1(X2,X1))=X1|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_38, plain, (r2_hidden(X1,X2)|X3!=k4_xboole_0(X2,X4)|~r2_hidden(X1,X3))).# trainpos
cnf(c_0_57, plain, (m1_subset_1(k3_subset_1(X2,X1),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_109, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_118, plain, (v2_struct_0(X1)|v1_compts_1(X1)|v2_waybel_0(esk4_1(X1),k3_yellow_1(k2_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_119, plain, (v2_struct_0(X1)|v1_compts_1(X1)|v13_waybel_0(esk4_1(X1),k3_yellow_1(k2_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_182, plain, (v2_struct_0(X1)|v1_compts_1(X1)|v3_waybel_7(esk4_1(X1),k3_yellow_1(k2_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_241, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r3_waybel_9(X1,X2,esk14_2(X1,X2))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_compts_1(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1))).# trainpos
cnf(c_0_78, plain, (X3=k4_xboole_0(X1,X2)|r2_hidden(esk15_3(X1,X2,X3),X3)|r2_hidden(esk15_3(X1,X2,X3),X1))).# trainpos
cnf(c_0_120, plain, (v2_struct_0(X1)|v1_compts_1(X1)|m1_subset_1(esk4_1(X1),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_103, plain, (v2_struct_0(X2)|v1_compts_1(X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_waybel_7(X2,esk4_1(X2),X1))).# trainpos
cnf(c_0_110, plain, (v2_struct_0(X2)|v2_struct_0(X1)|m1_subset_1(k10_yellow_6(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1))).# trainpos
cnf(c_0_252, plain, (v2_struct_0(X2)|v2_struct_0(X1)|m1_subset_1(esk14_2(X1,X2),u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_compts_1(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1))).# trainpos
cnf(c_0_211, plain, (v2_struct_0(X1)|m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_connsp_2(X3,X1,X2)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_166, plain, (v2_struct_0(X2)|v1_xboole_0(X1)|v1_subset_1(X1,u1_struct_0(X2))|~v4_orders_2(X2)|~v3_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~v2_waybel_0(X1,X2)|~v13_waybel_0(X1,X2)|~v3_waybel_7(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).# trainpos
cnf(c_0_136, plain, (v2_struct_0(X2)|v2_struct_0(X1)|m1_subset_1(esk9_3(X1,X2,X3),u1_struct_0(X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~r1_waybel_0(X1,X2,X3))).# trainpos
cnf(c_0_164, plain, (v2_struct_0(X2)|v2_struct_0(X1)|m1_subset_1(esk7_4(X1,X2,X3,X4),u1_struct_0(X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~r2_waybel_0(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X2)))).# trainpos
cnf(c_0_126, plain, (v2_struct_0(X3)|v2_struct_0(X1)|r1_orders_2(X1,X2,esk7_4(X3,X1,X4,X2))|~l1_struct_0(X3)|~l1_waybel_0(X1,X3)|~r2_waybel_0(X3,X1,X4)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_134, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r2_hidden(k2_waybel_0(X1,X2,esk7_4(X1,X2,X3,X4)),X3)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~r2_waybel_0(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X2)))).# trainpos
cnf(c_0_152, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|v4_orders_2(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2)))))).# trainpos
cnf(c_0_197, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|~l1_struct_0(X1)|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~v2_struct_0(k3_yellow19(X1,X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2)))))).# trainpos
cnf(c_0_125, plain, (v2_struct_0(X3)|v2_struct_0(X2)|r2_hidden(k2_waybel_0(X3,X2,X1),X4)|~l1_struct_0(X3)|~l1_waybel_0(X2,X3)|~r1_waybel_0(X3,X2,X4)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_orders_2(X2,esk9_3(X3,X2,X4),X1))).# trainpos
cnf(c_0_202, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r1_waybel_0(X2,X1,X3)|r1_waybel_0(X2,X1,k3_subset_1(u1_struct_0(X2),X3))|~l1_struct_0(X2)|~l1_waybel_0(X1,X2)|~v1_yellow19(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))).# trainpos
cnf(c_0_179, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|l1_waybel_0(k3_yellow19(X1,X2,X3),X1)|~l1_struct_0(X1)|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2)))))).# trainpos
cnf(c_0_115, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|v1_yellow19(k3_yellow19(X1,k2_struct_0(X1),X2),X1)|~l1_struct_0(X1)|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v3_waybel_7(X2,k3_yellow_1(k2_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1))))))).# trainpos
cnf(c_0_144, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|v7_waybel_0(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_subset_1(X3,u1_struct_0(k3_yellow_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2)))))).# trainpos
cnf(c_0_223, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r2_hidden(X4,X3)|X3!=k10_yellow_6(X1,X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_waybel_0(X1,X2,esk11_4(X1,X2,X3,X4)))).# trainpos
cnf(c_0_204, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r2_hidden(X4,X3)|m1_connsp_2(esk11_4(X1,X2,X3,X4),X1,X4)|X3!=k10_yellow_6(X1,X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_90, plain, (v2_struct_0(X2)|v1_xboole_0(X3)|r2_waybel_7(X2,X3,X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~v2_waybel_0(X3,k3_yellow_1(k2_struct_0(X2)))|~v13_waybel_0(X3,k3_yellow_1(k2_struct_0(X2)))|~v1_subset_1(X3,u1_struct_0(k3_yellow_1(k2_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X2)))))|~r2_hidden(X1,k10_yellow_6(X2,k3_yellow19(X2,k2_struct_0(X2),X3))))).# trainpos
cnf(c_0_59, plain, (esk17_1(X1)=k1_xboole_0)).# trainpos
cnf(c_0_114, negated_conjecture, (v2_struct_0(X1)|v1_compts_1(esk1_0)|k10_yellow_6(esk1_0,X1)!=k1_xboole_0|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~v1_yellow19(X1,esk1_0))).# trainpos
cnf(c_0_266, negated_conjecture, (k10_yellow_6(esk1_0,esk2_0)=k1_xboole_0|v2_struct_0(esk2_0)|~v1_compts_1(esk1_0)|~l1_waybel_0(esk2_0,esk1_0)|~v7_waybel_0(esk2_0)|~v4_orders_2(esk2_0))).# trainpos
cnf(c_0_43, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3)))).# trainpos
cnf(c_0_122, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X3)|~r2_hidden(X1,k10_yellow_6(X2,X3))|~l1_waybel_0(X3,X2)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).# trainpos
cnf(c_0_222, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X1)))))|v1_compts_1(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).# trainpos
cnf(c_0_135, plain, (r2_hidden(k2_waybel_0(X1,X2,esk7_4(X3,X2,X4,esk9_3(X1,X2,X5))),X5)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|~r1_waybel_0(X1,X2,X5)|~l1_struct_0(X1)|~m1_subset_1(esk7_4(X3,X2,X4,esk9_3(X1,X2,X5)),u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~r2_waybel_0(X3,X2,X4)|~l1_struct_0(X3)|~m1_subset_1(esk9_3(X1,X2,X5),u1_struct_0(X2))|~l1_waybel_0(X2,X3))).# trainpos
cnf(c_0_70, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0))).# trainpos
cnf(c_0_129, negated_conjecture, (l1_struct_0(esk1_0))).# trainpos
cnf(c_0_63, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_190, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_128, negated_conjecture, (v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1))|v1_xboole_0(X1)|k10_yellow_6(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1))!=k1_xboole_0|~l1_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1),esk1_0)|~v7_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1))|~v4_orders_2(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1))|~l1_struct_0(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~v3_waybel_7(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~v13_waybel_0(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(X1,k3_yellow_1(k2_struct_0(esk1_0))))).# trainpos
cnf(c_0_49, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_subset_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_111, plain, (X1!=k3_subset_1(X2,X3)|~r2_hidden(X4,X1)|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_56, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~m1_subset_1(k3_subset_1(X2,X3),k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_69, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0))).# trainpos
cnf(c_0_123, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,k3_subset_1(X3,X2))|~m1_subset_1(k3_subset_1(X3,X2),k1_zfmisc_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_141, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_hidden(k2_waybel_0(X1,X2,esk7_4(X1,X2,k3_subset_1(X3,X4),X5)),X4)|~m1_subset_1(X4,k1_zfmisc_1(X3))|~r2_waybel_0(X1,X2,k3_subset_1(X3,X4))|~l1_struct_0(X1)|~m1_subset_1(X5,u1_struct_0(X2))|~l1_waybel_0(X2,X1))).# trainpos
cnf(c_0_91, plain, (X2=k1_xboole_0|r2_hidden(esk15_3(k1_xboole_0,X1,X2),X2))).# trainpos
cnf(c_0_149, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X2,X1,k3_subset_1(X3,X4))|~l1_struct_0(X2)|~m1_subset_1(X4,k1_zfmisc_1(X3))|~m1_subset_1(esk9_3(X2,X1,X4),u1_struct_0(X1))|~l1_waybel_0(X1,X2)|~r1_waybel_0(X2,X1,X4)|~m1_subset_1(esk7_4(X2,X1,k3_subset_1(X3,X4),esk9_3(X2,X1,X4)),u1_struct_0(X1)))).# trainpos
cnf(c_0_177, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X1,X2,k3_subset_1(X3,X4))|~r1_waybel_0(X1,X2,X4)|~l1_struct_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(X3))|~l1_waybel_0(X2,X1)|~m1_subset_1(esk9_3(X1,X2,X4),u1_struct_0(X2)))).# trainpos
cnf(c_0_131, plain, (m1_subset_1(esk15_3(k1_xboole_0,X1,k10_yellow_6(X2,X3)),u1_struct_0(X2))|v2_struct_0(X3)|v2_struct_0(X2)|k10_yellow_6(X2,X3)=k1_xboole_0|~l1_waybel_0(X3,X2)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).# trainpos
cnf(c_0_104, plain, (r2_waybel_7(X1,X2,esk15_3(k1_xboole_0,X3,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2))))|v1_xboole_0(X2)|v2_struct_0(X1)|k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2))=k1_xboole_0|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))|~m1_subset_1(esk15_3(k1_xboole_0,X3,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2))),u1_struct_0(X1))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_117, plain, (v1_compts_1(X1)|v2_struct_0(X1)|k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)))=k1_xboole_0|v1_xboole_0(esk4_1(X1))|~m1_subset_1(esk15_3(k1_xboole_0,X2,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)))),u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))|~v13_waybel_0(esk4_1(X1),k3_yellow_1(k2_struct_0(X1)))|~v2_waybel_0(esk4_1(X1),k3_yellow_1(k2_struct_0(X1))))).# trainpos
cnf(c_0_138, plain, (k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)))=k1_xboole_0|v1_compts_1(X1)|v2_struct_0(X1)|v2_struct_0(k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)))|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)),X1)|~v7_waybel_0(k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)))|~v4_orders_2(k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1))))).# trainpos
cnf(c_0_143, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)))|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~v3_waybel_7(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)),esk1_0)|~v7_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)))|~v4_orders_2(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)))|~v1_subset_1(esk4_1(esk1_0),u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))).# trainpos
cnf(c_0_208, plain, (v2_waybel_0(esk4_1(X1),k3_yellow_1(u1_struct_0(X1)))|v1_compts_1(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).# trainpos
cnf(c_0_215, plain, (v13_waybel_0(esk4_1(X1),k3_yellow_1(u1_struct_0(X1)))|v1_compts_1(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).# trainpos
cnf(c_0_192, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X2,X1,X4)|~r1_waybel_0(X2,X1,k3_subset_1(X3,X4))|~l1_struct_0(X2)|~m1_subset_1(k3_subset_1(X3,X4),k1_zfmisc_1(X3))|~l1_waybel_0(X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_191, plain, (v3_waybel_7(esk4_1(X1),k3_yellow_1(u1_struct_0(X1)))|v1_compts_1(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).# trainpos
cnf(c_0_209, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r1_waybel_0(X1,X2,X3)|~r2_waybel_0(X1,X2,X3)|~l1_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_waybel_0(X2,X1)|~v1_yellow19(X2,X1))).# trainpos
cnf(c_0_151, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)))|v1_xboole_0(k2_struct_0(esk1_0))|~v1_subset_1(esk4_1(esk1_0),u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0))))|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~v3_waybel_7(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)),esk1_0)|~v4_orders_2(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_165, negated_conjecture, (v1_xboole_0(k2_struct_0(esk1_0))|v1_xboole_0(esk4_1(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)))|~v1_subset_1(esk4_1(esk1_0),u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0))))|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v3_waybel_7(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)),esk1_0))).# trainpos
cnf(c_0_178, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_xboole_0(k2_struct_0(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)))|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v3_waybel_7(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)),esk1_0))).# trainpos
cnf(c_0_189, negated_conjecture, (v1_xboole_0(k2_struct_0(esk1_0))|v1_xboole_0(esk4_1(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)))|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v3_waybel_7(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0))))).# trainpos
cnf(c_0_199, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_xboole_0(u1_struct_0(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk4_1(esk1_0)))|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))|~v3_waybel_7(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_198, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_xboole_0(k2_struct_0(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk4_1(esk1_0)))|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0))))).# trainpos
cnf(c_0_207, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(esk4_1(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk4_1(esk1_0)))|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_206, negated_conjecture, (v1_xboole_0(k2_struct_0(esk1_0))|v1_xboole_0(esk4_1(esk1_0))|v1_compts_1(esk1_0)|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0))))).# trainpos
cnf(c_0_213, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(esk4_1(esk1_0))|v1_compts_1(esk1_0)|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_212, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_xboole_0(k2_struct_0(esk1_0))|v1_compts_1(esk1_0)|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0))))).# trainpos
cnf(c_0_220, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_xboole_0(u1_struct_0(esk1_0))|v1_compts_1(esk1_0)|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_219, negated_conjecture, (v1_xboole_0(k2_struct_0(esk1_0))|v1_xboole_0(esk4_1(esk1_0))|v1_compts_1(esk1_0)|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_229, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(esk4_1(esk1_0))|v1_compts_1(esk1_0)|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0))))))).# trainpos
cnf(c_0_227, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_xboole_0(k2_struct_0(esk1_0))|v1_compts_1(esk1_0)|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_235, negated_conjecture, (v1_compts_1(esk1_0)|v1_xboole_0(k2_struct_0(esk1_0))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_239, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_xboole_0(u1_struct_0(esk1_0))|v1_compts_1(esk1_0))).# trainpos
cnf(c_0_248, negated_conjecture, (v1_compts_1(esk1_0)|v1_xboole_0(u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_244, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0)|v1_xboole_0(k2_struct_0(esk1_0))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_255, negated_conjecture, (v1_xboole_0(k2_struct_0(esk1_0))|l1_waybel_0(esk2_0,esk1_0))).# trainpos
cnf(c_0_257, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~v2_struct_0(esk2_0))).# trainpos
cnf(c_0_264, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0))).# trainpos
cnf(c_0_265, negated_conjecture, (~v2_struct_0(esk2_0))).# trainpos
cnf(c_0_210, plain, (r2_waybel_0(X1,X2,esk11_4(X1,X3,X4,X5))|v2_struct_0(X1)|v2_struct_0(X2)|r2_hidden(X5,X4)|v2_struct_0(X3)|~r3_waybel_9(X1,X2,X5)|~m1_subset_1(X5,u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|X4!=k10_yellow_6(X1,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~l1_waybel_0(X3,X1)|~v7_waybel_0(X3)|~v4_orders_2(X3))).# trainpos
cnf(c_0_218, plain, (m1_subset_1(esk11_4(X1,X2,X3,X4),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|r2_hidden(X4,X3)|v2_struct_0(X2)|~m1_subset_1(X4,u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|X3!=k10_yellow_6(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2))).# trainpos
cnf(c_0_217, plain, (r1_waybel_0(X1,X2,esk11_4(X1,X3,X4,X5))|v2_struct_0(X2)|v2_struct_0(X1)|r2_hidden(X5,X4)|v2_struct_0(X3)|~l1_struct_0(X1)|~m1_subset_1(esk11_4(X1,X3,X4,X5),k1_zfmisc_1(u1_struct_0(X1)))|~v1_yellow19(X2,X1)|~l1_waybel_0(X2,X1)|X4!=k10_yellow_6(X1,X3)|~r3_waybel_9(X1,X2,X5)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,u1_struct_0(X1))|~l1_waybel_0(X3,X1)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_231, plain, (r2_hidden(X1,X2)|v2_struct_0(X3)|v2_struct_0(X4)|X2!=k10_yellow_6(X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_waybel_0(X4,X3)|~v7_waybel_0(X4)|~v4_orders_2(X4)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~r3_waybel_9(X3,X4,X1)|~v1_yellow19(X4,X3))).# trainpos
cnf(c_0_240, plain, (v2_struct_0(X2)|v2_struct_0(X3)|k10_yellow_6(X3,X2)!=k1_xboole_0|~r3_waybel_9(X3,X2,X1)|~m1_subset_1(X1,u1_struct_0(X3))|~v1_yellow19(X2,X3)|~l1_waybel_0(X2,X3)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X3)|~v2_pre_topc(X3))).# trainpos
cnf(c_0_251, plain, (v2_struct_0(X1)|v2_struct_0(X2)|k10_yellow_6(X1,X2)!=k1_xboole_0|~m1_subset_1(esk14_2(X1,X2),u1_struct_0(X1))|~v1_yellow19(X2,X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_compts_1(X1))).# trainpos
cnf(c_0_269, negated_conjecture, (k10_yellow_6(esk1_0,esk2_0)!=k1_xboole_0|~v7_waybel_0(esk2_0)|~v4_orders_2(esk2_0)|~v1_compts_1(esk1_0))).# trainpos
cnf(c_0_272, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0)))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_97, plain, (o_0_0_xboole_0=k1_xboole_0)).#trainneg
cnf(i_0_127, plain, (v1_xboole_0(o_0_0_xboole_0))).#trainneg
cnf(i_0_157, plain, (v7_waybel_0(k3_yellow_1(X1)))).#trainneg
cnf(i_0_75, plain, (v1_orders_2(k3_yellow_1(X1)))).#trainneg
cnf(i_0_181, plain, (~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_192, plain, (v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_173, plain, (v1_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_84, plain, (v11_waybel_1(k3_yellow_1(X1)))).#trainneg
cnf(i_0_95, plain, (v3_lattice3(k3_yellow_1(X1)))).#trainneg
cnf(i_0_110, plain, (m1_subset_1(esk16_1(X1),X1))).#trainneg
cnf(i_0_117, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_94, plain, (l1_struct_0(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_198, plain, (~l1_struct_0(X1)|~v2_struct_0(esk19_1(X1)))).#trainneg
cnf(i_0_128, plain, (~v2_struct_0(X1)|~l1_orders_2(X1)|~v1_lattice3(X1))).#trainneg
cnf(i_0_104, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_125, plain, (~v2_struct_0(X1)|~l1_orders_2(X1)|~v2_lattice3(X1))).#trainneg
cnf(i_0_149, plain, (v2_yellow_0(X1)|~l1_orders_2(X1)|~v3_yellow_0(X1))).#trainneg
cnf(i_0_120, plain, (v2_struct_0(X1)|v1_lattice3(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_168, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_196, plain, (v4_orders_2(esk19_1(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_193, plain, (v7_waybel_0(esk19_1(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_197, plain, (v3_orders_2(esk19_1(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_195, plain, (v5_orders_2(esk19_1(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_163, plain, (v1_xboole_0(X1)|~v7_struct_0(k3_yellow_1(X1)))).#trainneg
cnf(i_0_199, plain, (l1_waybel_0(esk19_1(X1),X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_194, plain, (v6_waybel_0(esk19_1(X1),X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_169, plain, (v1_tops_1(k2_struct_0(X1),X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_107, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_150, plain, (v1_yellow_0(X1)|~l1_orders_2(X1)|~v3_yellow_0(X1))).#trainneg
cnf(i_0_82, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_101, plain, (v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_tops_1(k2_struct_0(X1),X1))).#trainneg
cnf(i_0_133, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_132, plain, (v1_xboole_0(k2_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_119, plain, (v2_struct_0(X1)|v2_lattice3(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_179, plain, (v2_struct_0(X1)|v3_yellow_0(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_86, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_200, plain, (v2_struct_0(X1)|v25_waybel_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_140, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_156, plain, (v1_finset_1(u1_struct_0(X1))|~l1_struct_0(X1)|~v8_struct_0(X1))).#trainneg
cnf(i_0_142, plain, (v2_struct_0(X1)|~l1_orders_2(X1)|~v1_xboole_0(k2_struct_0(X1)))).#trainneg
cnf(i_0_177, plain, (l1_orders_2(X2)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_135, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_138, plain, (v2_struct_0(X1)|v13_waybel_0(k2_struct_0(X1),X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_139, plain, (v2_struct_0(X1)|v12_waybel_0(k2_struct_0(X1),X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_178, plain, (v2_waybel_0(k2_struct_0(X1),X1)|~l1_orders_2(X1)|~v2_lattice3(X1))).#trainneg
cnf(i_0_159, plain, (v4_pre_topc(k2_struct_0(X1),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_109, plain, (v3_pre_topc(k2_struct_0(X1),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_146, plain, (v1_waybel_0(k2_struct_0(X1),X1)|~l1_orders_2(X1)|~v1_lattice3(X1))).#trainneg
cnf(i_0_100, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_134, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_204, plain, (v1_finset_1(X2)|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_201, plain, (v2_struct_0(X1)|v24_waybel_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_147, plain, (X1=X2|g1_orders_2(X3,X1)!=g1_orders_2(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X3))))).#trainneg
cnf(i_0_188, plain, (v2_struct_0(X1)|v1_yellow_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v25_waybel_0(X1))).#trainneg
cnf(i_0_145, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_176, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).#trainneg
cnf(i_0_160, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_191, plain, (v2_struct_0(X1)|v2_waybel_0(k2_struct_0(X1),X1)|~l1_orders_2(X1)|~v1_yellow_0(X1))).#trainneg
cnf(i_0_170, plain, (v2_struct_0(X1)|v1_waybel_0(k2_struct_0(X1),X1)|~l1_orders_2(X1)|~v2_yellow_0(X1))).#trainneg
cnf(i_0_184, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_130, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_129, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_87, plain, (g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=X1|~v1_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_79, plain, (v1_funct_1(u1_waybel_0(X1,X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_172, plain, (v1_orders_2(g1_orders_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_171, plain, (l1_orders_2(g1_orders_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_123, plain, (v2_struct_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~v1_xboole_0(u1_waybel_0(X1,X2)))).#trainneg
cnf(i_0_148, plain, (X1=X2|g1_orders_2(X1,X3)!=g1_orders_2(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_185, plain, (v2_tops_1(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v3_tops_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_161, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_98, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1))).#trainneg
cnf(i_0_56, plain, (r2_hidden(X1,X4)|r2_hidden(X1,X3)|X4!=k4_xboole_0(X2,X3)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_23, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v1_yellow19(X2,X1)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_78, plain, (v1_funct_2(u1_waybel_0(X1,X2),u1_struct_0(X2),u1_struct_0(X1))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_77, plain, (m1_subset_1(u1_waybel_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_53, plain, (X3=k4_xboole_0(X1,X2)|r2_hidden(esk15_3(X1,X2,X3),X3)|~r2_hidden(esk15_3(X1,X2,X3),X2))).#trainneg
cnf(i_0_118, plain, (v2_struct_0(X1)|m1_connsp_2(esk18_2(X1,X2),X1,X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_116, plain, (g1_waybel_0(X1,u1_struct_0(X2),u1_orders_2(X2),u1_waybel_0(X1,X2))=X2|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~v6_waybel_0(X2,X1))).#trainneg
cnf(i_0_106, plain, (v2_struct_0(X1)|v13_waybel_0(k3_subset_1(u1_struct_0(X1),X2),X1)|~l1_orders_2(X1)|~v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_141, plain, (v2_struct_0(X1)|v12_waybel_0(k3_subset_1(u1_struct_0(X1),X2),X1)|~l1_orders_2(X1)|~v13_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_99, plain, (v4_pre_topc(k3_subset_1(u1_struct_0(X1),X2),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_83, plain, (v3_pre_topc(k3_subset_1(u1_struct_0(X1),X2),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_pre_topc(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_137, plain, (v1_tops_1(k3_subset_1(u1_struct_0(X1),X2),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v3_tops_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_22, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v1_yellow19(X2,X1)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~r1_waybel_0(X1,X2,esk5_2(X1,X2)))).#trainneg
cnf(i_0_183, plain, (v1_xboole_0(X2)|~l1_struct_0(X1)|~v1_funct_1(X4)|~v1_funct_2(X4,X2,u1_struct_0(X1))|~v2_struct_0(g1_waybel_0(X1,X2,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X1)))))).#trainneg
cnf(i_0_62, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|~l1_struct_0(X1)|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v3_waybel_7(X2,k3_yellow_1(k2_struct_0(X1)))|~v2_struct_0(k3_yellow19(X1,k2_struct_0(X1),X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1))))))).#trainneg
cnf(i_0_26, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r3_waybel_9(X1,X2,X3)|m1_connsp_2(esk6_3(X1,X2,X3),X1,X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_36, plain, (v2_struct_0(X3)|v2_struct_0(X1)|r1_waybel_0(X3,X1,X4)|r1_orders_2(X1,X2,esk10_4(X3,X1,X4,X2))|~l1_struct_0(X3)|~l1_waybel_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_102, plain, (v2_struct_0(X2)|v2_struct_0(X1)|m1_subset_1(k2_waybel_0(X1,X2,X3),u1_struct_0(X1))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2)))).#trainneg
cnf(i_0_55, plain, (X3=k4_xboole_0(X1,X2)|r2_hidden(esk15_3(X1,X2,X3),X2)|~r2_hidden(esk15_3(X1,X2,X3),X3)|~r2_hidden(esk15_3(X1,X2,X3),X1))).#trainneg
cnf(i_0_31, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r2_waybel_0(X1,X2,X3)|m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_21, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v1_yellow19(X2,X1)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~r1_waybel_0(X1,X2,k3_subset_1(u1_struct_0(X1),esk5_2(X1,X2))))).#trainneg
cnf(i_0_105, plain, (k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_136, plain, (v1_xboole_0(X1)|m1_subset_1(k3_funct_2(X1,X3,X2,X4),X3)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_131, plain, (k3_funct_2(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2),X3)=k2_waybel_0(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2)))).#trainneg
cnf(i_0_37, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r1_waybel_0(X1,X2,X3)|m1_subset_1(esk10_4(X1,X2,X3,X4),u1_struct_0(X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X4,u1_struct_0(X2)))).#trainneg
cnf(i_0_25, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r3_waybel_9(X1,X2,X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_waybel_0(X1,X2,esk6_3(X1,X2,X3)))).#trainneg
cnf(i_0_164, plain, (X1=X2|g1_waybel_0(X3,X4,X5,X1)!=g1_waybel_0(X6,X7,X8,X2)|~l1_struct_0(X3)|~v1_funct_1(X1)|~v1_funct_2(X1,X4,u1_struct_0(X3))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,u1_struct_0(X3)))))).#trainneg
cnf(i_0_186, plain, (l1_waybel_0(g1_waybel_0(X1,X2,X3,X4),X1)|~l1_struct_0(X1)|~v1_funct_1(X4)|~v1_funct_2(X4,X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X1)))))).#trainneg
cnf(i_0_43, plain, (X3=k10_yellow_6(X1,X2)|v2_struct_0(X2)|v2_struct_0(X1)|m1_subset_1(esk12_3(X1,X2,X3),u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_30, plain, (v2_struct_0(X3)|v2_struct_0(X2)|r2_waybel_0(X3,X2,X4)|~l1_struct_0(X3)|~l1_waybel_0(X2,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(k2_waybel_0(X3,X2,X1),X4)|~r1_orders_2(X2,esk8_3(X3,X2,X4),X1))).#trainneg
cnf(i_0_46, plain, (v2_struct_0(X5)|v2_struct_0(X4)|r1_waybel_0(X4,X5,X3)|X2!=k10_yellow_6(X4,X5)|~v2_pre_topc(X4)|~l1_pre_topc(X4)|~v4_orders_2(X5)|~v7_waybel_0(X5)|~l1_waybel_0(X5,X4)|~r2_hidden(X1,X2)|~m1_connsp_2(X3,X4,X1)|~m1_subset_1(X1,u1_struct_0(X4))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X4))))).#trainneg
cnf(i_0_69, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|v3_orders_2(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2)))))).#trainneg
cnf(i_0_35, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r1_waybel_0(X1,X2,X3)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(k2_waybel_0(X1,X2,esk10_4(X1,X2,X3,X4)),X3))).#trainneg
cnf(i_0_187, plain, (v6_waybel_0(g1_waybel_0(X1,X2,X3,X4),X1)|~l1_struct_0(X1)|~v1_funct_1(X4)|~v1_funct_2(X4,X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X1)))))).#trainneg
cnf(i_0_165, plain, (X1=X2|g1_waybel_0(X3,X4,X1,X5)!=g1_waybel_0(X6,X7,X2,X8)|~l1_struct_0(X3)|~v1_funct_1(X5)|~v1_funct_2(X5,X4,u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,u1_struct_0(X3)))))).#trainneg
cnf(i_0_166, plain, (X1=X2|g1_waybel_0(X3,X1,X4,X5)!=g1_waybel_0(X6,X2,X7,X8)|~l1_struct_0(X3)|~v1_funct_1(X5)|~v1_funct_2(X5,X1,u1_struct_0(X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X3)))))).#trainneg
cnf(i_0_167, plain, (X1=X2|g1_waybel_0(X1,X3,X4,X5)!=g1_waybel_0(X2,X6,X7,X8)|~l1_struct_0(X1)|~v1_funct_1(X5)|~v1_funct_2(X5,X3,u1_struct_0(X1))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X3)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(X1)))))).#trainneg
cnf(i_0_65, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|v6_waybel_0(k3_yellow19(X1,X2,X3),X1)|~l1_struct_0(X1)|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2)))))).#trainneg
cnf(i_0_61, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|v6_waybel_0(k3_yellow19(X1,k2_struct_0(X1),X2),X1)|~l1_struct_0(X1)|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v3_waybel_7(X2,k3_yellow_1(k2_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1))))))).#trainneg
cnf(i_0_20, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|m1_subset_1(esk3_2(X1,X2),u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_compts_1(X1)|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v3_waybel_7(X2,k3_yellow_1(k2_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1))))))).#trainneg
cnf(i_0_19, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|r2_waybel_7(X1,X2,esk3_2(X1,X2))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_compts_1(X1)|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v3_waybel_7(X2,k3_yellow_1(k2_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1))))))).#trainneg
cnf(i_0_41, plain, (X3=k10_yellow_6(X1,X2)|v2_struct_0(X2)|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(esk12_3(X1,X2,X3),X3)|~r1_waybel_0(X1,X2,esk13_3(X1,X2,X3)))).#trainneg
cnf(i_0_103, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|m1_subset_1(o_2_13_yellow19(X1,X2),k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v3_waybel_7(X2,k3_yellow_1(k2_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1))))))).#trainneg
cnf(i_0_42, plain, (X3=k10_yellow_6(X1,X2)|v2_struct_0(X2)|v2_struct_0(X1)|m1_connsp_2(esk13_3(X1,X2,X3),X1,esk12_3(X1,X2,X3))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(esk12_3(X1,X2,X3),X3))).#trainneg
cnf(i_0_40, plain, (X3=k10_yellow_6(X1,X2)|v2_struct_0(X2)|v2_struct_0(X1)|r1_waybel_0(X1,X2,X4)|r2_hidden(esk12_3(X1,X2,X3),X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_connsp_2(X4,X1,esk12_3(X1,X2,X3)))).#trainneg
cnf(i_0_28, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|r2_hidden(X3,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~r2_waybel_7(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1))))))).#trainneg
cnf(i_0_482, negated_conjecture, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|v2_struct_0(esk2_0)|~l1_waybel_0(esk2_0,esk1_0)|~v7_waybel_0(esk2_0)|~v4_orders_2(esk2_0)|~v1_compts_1(esk1_0))).#trainneg
cnf(i_0_427, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X4)|k10_yellow_6(X1,X4)!=k10_yellow_6(X1,X2)|~r2_hidden(X5,k10_yellow_6(X1,X4))|~m1_connsp_2(X3,X1,X5)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X4,X1)|~v7_waybel_0(X4)|~v4_orders_2(X4))).#trainneg
cnf(i_0_461, plain, (r2_waybel_7(X1,X2,X3)|v1_xboole_0(X2)|v2_struct_0(X1)|~r2_hidden(X3,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2)))|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X1)))))|~m1_subset_1(X3,u1_struct_0(X1))|~v13_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~v2_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_428, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|v1_xboole_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~m1_subset_1(k2_waybel_0(X1,X2,esk10_4(X1,X2,X3,X4)),X3))).#trainneg
cnf(i_0_207, plain, (l1_struct_0(k3_yellow_1(X1)))).#trainneg
cnf(i_0_217, plain, (~r2_hidden(X1,esk16_1(k1_zfmisc_1(X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_234, plain, (~r2_hidden(X1,k2_struct_0(X2))|~v1_xboole_0(u1_struct_0(X2))|~l1_struct_0(X2))).#trainneg
cnf(i_0_223, plain, (v1_yellow_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_212, plain, (v1_tops_1(u1_struct_0(X1),X1)|~l1_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_416, plain, (r2_hidden(k2_waybel_0(X1,X2,esk10_4(X3,X2,X4,esk9_3(X1,X2,X5))),X5)|v2_struct_0(X2)|v2_struct_0(X1)|r1_waybel_0(X3,X2,X4)|v2_struct_0(X3)|~r1_waybel_0(X1,X2,X5)|~l1_struct_0(X1)|~m1_subset_1(esk10_4(X3,X2,X4,esk9_3(X1,X2,X5)),u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_struct_0(X3)|~m1_subset_1(esk9_3(X1,X2,X5),u1_struct_0(X2))|~l1_waybel_0(X2,X3))).#trainneg
cnf(i_0_323, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|X3!=k3_subset_1(X4,X2)|~r2_hidden(X1,X4)|~m1_subset_1(X2,k1_zfmisc_1(X4)))).#trainneg
cnf(i_0_320, plain, (r2_hidden(X1,X2)|r2_hidden(X1,k4_xboole_0(X3,X2))|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_251, plain, (X1!=k1_xboole_0|~r2_hidden(X3,X1)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_489, plain, (r2_hidden(X1,X2)|v1_xboole_0(k4_xboole_0(X2,X3))|~m1_subset_1(X1,k4_xboole_0(X2,X3)))).#trainneg
cnf(i_0_464, plain, (r2_waybel_7(X1,X2,k2_waybel_0(X3,X4,esk7_4(X3,X4,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2)),X5)))|v1_xboole_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X4)|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))|~m1_subset_1(k2_waybel_0(X3,X4,esk7_4(X3,X4,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2)),X5)),u1_struct_0(X1))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~r2_waybel_0(X3,X4,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2)))|~l1_struct_0(X3)|~m1_subset_1(X5,u1_struct_0(X4))|~l1_waybel_0(X4,X3))).#trainneg
cnf(i_0_425, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|k3_subset_1(u1_struct_0(X1),X4)!=k10_yellow_6(X1,X2)|~r2_hidden(X5,k3_subset_1(u1_struct_0(X1),X4))|~m1_connsp_2(X3,X1,X5)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_625, plain, (v1_compts_1(X1)|v2_struct_0(X1)|v1_xboole_0(esk4_1(X1))|v2_struct_0(X3)|v2_struct_0(X2)|~m1_subset_1(k2_waybel_0(X2,X3,esk7_4(X2,X3,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1))),X4)),u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~r2_waybel_0(X2,X3,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1))))|~l1_struct_0(X2)|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))|~m1_subset_1(X4,u1_struct_0(X3))|~v13_waybel_0(esk4_1(X1),k3_yellow_1(k2_struct_0(X1)))|~v2_waybel_0(esk4_1(X1),k3_yellow_1(k2_struct_0(X1)))|~l1_waybel_0(X3,X2))).#trainneg
cnf(i_0_592, negated_conjecture, (v1_xboole_0(X1)|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),X1))|k10_yellow_6(esk1_0,k3_yellow19(esk1_0,u1_struct_0(esk1_0),X1))!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))|~v3_waybel_7(X1,k3_yellow_1(u1_struct_0(esk1_0)))|~v13_waybel_0(X1,k3_yellow_1(u1_struct_0(esk1_0)))|~v2_waybel_0(X1,k3_yellow_1(u1_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),X1),esk1_0)|~v7_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),X1))|~v4_orders_2(k3_yellow19(esk1_0,u1_struct_0(esk1_0),X1)))).#trainneg
cnf(i_0_215, plain, (v2_struct_0(X1)|~v2_tops_1(u1_struct_0(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_531, plain, (v1_relat_1(k1_xboole_0))).#trainneg
cnf(i_0_490, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|v2_struct_0(esk2_0)|~r2_hidden(X1,k1_xboole_0)|~l1_waybel_0(esk2_0,esk1_0)|~v7_waybel_0(esk2_0)|~v4_orders_2(esk2_0)|~v1_compts_1(esk1_0))).#trainneg
cnf(i_0_218, plain, (X1=k2_struct_0(X2)|~v1_xboole_0(X1)|~l1_struct_0(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_545, plain, (v1_finset_1(k1_xboole_0)|~v1_finset_1(X1))).#trainneg
cnf(i_0_542, plain, (~r2_hidden(X1,k1_xboole_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_240, plain, (v2_waybel_0(k2_struct_0(X1),X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_219, plain, (k2_struct_0(X1)=k1_xboole_0|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_600, plain, (r2_hidden(X1,k3_subset_1(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_496, negated_conjecture, (r1_waybel_0(esk1_0,X1,X2)|v2_struct_0(esk2_0)|v2_struct_0(X1)|k10_yellow_6(esk1_0,X1)!=k1_xboole_0|~r2_hidden(X3,k1_xboole_0)|~m1_connsp_2(X2,esk1_0,X3)|~l1_waybel_0(X1,esk1_0)|~l1_waybel_0(esk2_0,esk1_0)|~v7_waybel_0(X1)|~v7_waybel_0(esk2_0)|~v4_orders_2(X1)|~v4_orders_2(esk2_0)|~v1_compts_1(esk1_0))).#trainneg
cnf(i_0_512, plain, (r1_waybel_0(X1,X2,u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~m1_subset_1(esk10_4(X1,X2,u1_struct_0(X1),X3),u1_struct_0(X2)))).#trainneg
cnf(i_0_614, plain, (v2_struct_0(X2)|v2_struct_0(X3)|X1!=k1_xboole_0|~r2_hidden(k2_waybel_0(X2,X3,esk7_4(X2,X3,X1,X4)),X5)|~r2_waybel_0(X2,X3,X1)|~l1_struct_0(X2)|~m1_subset_1(X4,u1_struct_0(X3))|~l1_waybel_0(X3,X2))).#trainneg
cnf(i_0_510, plain, (m1_subset_1(X1,u1_struct_0(k3_yellow_1(u1_struct_0(X2))))|v1_compts_1(X2)|v2_struct_0(X2)|~r2_hidden(X1,esk4_1(X2))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_504, plain, (r2_waybel_7(X1,X2,k2_waybel_0(X3,X4,esk7_4(X3,X4,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2)),X5)))|v1_xboole_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X4)|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X1)))))|~m1_subset_1(k2_waybel_0(X3,X4,esk7_4(X3,X4,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2)),X5)),u1_struct_0(X1))|~v13_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~v2_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~r2_waybel_0(X3,X4,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2)))|~l1_struct_0(X3)|~m1_subset_1(X5,u1_struct_0(X4))|~l1_waybel_0(X4,X3))).#trainneg
cnf(i_0_530, plain, (v5_relat_1(k1_xboole_0,X1))).#trainneg
cnf(i_0_534, plain, (v4_relat_1(k1_xboole_0,X1))).#trainneg
cnf(i_0_527, plain, (k1_xboole_0=X1|g1_orders_2(X2,k1_xboole_0)!=g1_orders_2(X3,X1))).#trainneg
cnf(i_0_228, plain, (X1=u1_struct_0(X2)|~v1_xboole_0(X1)|~l1_struct_0(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_546, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_525, plain, (v1_orders_2(g1_orders_2(X1,k1_xboole_0)))).#trainneg
cnf(i_0_229, plain, (u1_struct_0(X1)=k1_xboole_0|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_233, plain, (l1_orders_2(esk19_1(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_237, plain, (v1_relat_1(esk16_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_505, plain, (r2_waybel_7(X1,X2,X3)|v1_xboole_0(X2)|v2_struct_0(X1)|v1_xboole_0(k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2)))|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X1)))))|~m1_subset_1(X3,u1_struct_0(X1))|~v13_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~v2_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2))))).#trainneg
cnf(i_0_465, plain, (r2_waybel_7(X1,X2,X3)|v1_xboole_0(X2)|v2_struct_0(X1)|v1_xboole_0(k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2)))|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))|~m1_subset_1(X3,u1_struct_0(X1))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2))))).#trainneg
cnf(i_0_401, plain, (m1_subset_1(k2_waybel_0(X1,X2,esk7_4(X1,X2,X3,X4)),X3)|v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X1,X2,X3)|~l1_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_615, plain, (r1_waybel_0(X4,X3,X5)|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X3)|X1!=k1_xboole_0|~r2_hidden(k2_waybel_0(X2,X3,esk10_4(X4,X3,X5,esk9_3(X2,X3,X1))),X6)|~r1_waybel_0(X2,X3,X1)|~l1_struct_0(X2)|~l1_struct_0(X4)|~m1_subset_1(esk10_4(X4,X3,X5,esk9_3(X2,X3,X1)),u1_struct_0(X3))|~l1_waybel_0(X3,X2)|~l1_waybel_0(X3,X4))).#trainneg
cnf(i_0_250, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_586, plain, (r2_waybel_7(X1,X2,k2_waybel_0(X3,X4,esk10_4(X5,X4,X6,esk9_3(X3,X4,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2))))))|v1_xboole_0(X2)|v2_struct_0(X1)|r1_waybel_0(X5,X4,X6)|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X4)|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X1)))))|~m1_subset_1(k2_waybel_0(X3,X4,esk10_4(X5,X4,X6,esk9_3(X3,X4,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2))))),u1_struct_0(X1))|~v13_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~v2_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~r1_waybel_0(X3,X4,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2)))|~l1_struct_0(X3)|~l1_struct_0(X5)|~m1_subset_1(esk10_4(X5,X4,X6,esk9_3(X3,X4,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2)))),u1_struct_0(X4))|~l1_waybel_0(X4,X3)|~l1_waybel_0(X4,X5))).#trainneg
cnf(i_0_584, plain, (r1_waybel_0(X3,X2,X4)|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|~r2_hidden(k2_waybel_0(X1,X2,esk10_4(X3,X2,X4,esk9_3(X1,X2,k1_xboole_0))),X5)|~r1_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X1)|~l1_struct_0(X3)|~m1_subset_1(esk10_4(X3,X2,X4,esk9_3(X1,X2,k1_xboole_0)),u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_waybel_0(X2,X3))).#trainneg
cnf(i_0_705, plain, (v2_struct_0(X1)|v2_struct_0(X2)|X3!=k1_xboole_0|~r2_waybel_0(X2,X1,X3)|~l1_struct_0(X2)|~m1_subset_1(esk9_3(X2,X1,X4),u1_struct_0(X1))|~l1_waybel_0(X1,X2)|~r1_waybel_0(X2,X1,X4)|~m1_subset_1(esk7_4(X2,X1,X3,esk9_3(X2,X1,X4)),u1_struct_0(X1)))).#trainneg
cnf(i_0_642, plain, (v1_finset_1(k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_526, plain, (l1_orders_2(g1_orders_2(X1,k1_xboole_0)))).#trainneg
cnf(i_0_308, plain, (v1_relat_1(u1_orders_2(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_716, plain, (u1_orders_2(X1)=k1_xboole_0|g1_orders_2(X2,k1_xboole_0)!=X1|~l1_orders_2(X1)|~v1_orders_2(X1))).#trainneg
cnf(i_0_238, plain, (v13_waybel_0(u1_struct_0(X1),X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_306, plain, (~r2_hidden(X1,u1_orders_2(X2))|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2)))|~l1_orders_2(X2))).#trainneg
cnf(i_0_239, plain, (v12_waybel_0(u1_struct_0(X1),X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_241, plain, (v4_pre_topc(u1_struct_0(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_696, plain, (r1_waybel_0(X1,X2,u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X2)|r2_waybel_0(X3,X2,X4)|v2_struct_0(X3)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_struct_0(X3)|~l1_waybel_0(X2,X3))).#trainneg
cnf(i_0_691, plain, (r1_waybel_0(X1,X2,u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X2)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_794, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X4)|v2_struct_0(X1)|~r1_waybel_0(X4,X2,k1_xboole_0)|~l1_struct_0(X4)|~l1_struct_0(X1)|~m1_subset_1(esk10_4(X1,X2,X3,esk9_3(X4,X2,k1_xboole_0)),u1_struct_0(X2))|~l1_waybel_0(X2,X4)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_635, plain, (v1_compts_1(X1)|v2_struct_0(X1)|v2_struct_0(X2)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~r2_waybel_0(X1,X2,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1))))|~l1_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk7_4(X1,X2,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1))),X3),u1_struct_0(X2)))).#trainneg
cnf(i_0_712, plain, (v1_compts_1(X1)|v2_struct_0(X1)|v1_xboole_0(esk4_1(X1))|v2_struct_0(X3)|v2_struct_0(X2)|~m1_subset_1(k2_waybel_0(X2,X3,esk7_4(X2,X3,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))),X4)),u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~r2_waybel_0(X2,X3,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))))|~l1_struct_0(X2)|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X1)))))|~m1_subset_1(X4,u1_struct_0(X3))|~v13_waybel_0(esk4_1(X1),k3_yellow_1(u1_struct_0(X1)))|~v2_waybel_0(esk4_1(X1),k3_yellow_1(u1_struct_0(X1)))|~l1_waybel_0(X3,X2))).#trainneg
cnf(i_0_634, plain, (v1_compts_1(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X1)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~r2_waybel_0(X2,X3,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))))|~l1_struct_0(X2)|~m1_subset_1(k2_waybel_0(X2,X3,esk7_4(X2,X3,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))),X4)),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_waybel_0(X3,X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_732, plain, (r2_hidden(X1,u1_struct_0(k3_yellow_1(k2_struct_0(X2))))|v1_compts_1(X2)|v2_struct_0(X2)|~r2_hidden(X1,esk4_1(X2))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_740, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X3)|~r2_hidden(X1,k10_yellow_6(X2,X3))|~l1_waybel_0(X3,X2)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_733, plain, (r2_hidden(X1,u1_struct_0(k3_yellow_1(u1_struct_0(X2))))|v1_compts_1(X2)|v2_struct_0(X2)|~r2_hidden(X1,esk4_1(X2))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_738, plain, (r2_hidden(X1,u1_struct_0(X2))|~r2_hidden(X1,k2_struct_0(X2))|~l1_struct_0(X2))).#trainneg
cnf(i_0_535, plain, (X1=X2|g1_orders_2(X1,k1_xboole_0)!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_564, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_313, plain, (u1_orders_2(X1)=X2|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(X3,X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_242, plain, (v3_pre_topc(u1_struct_0(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_652, plain, (v1_tops_1(k1_xboole_0,X1)|~l1_pre_topc(X1)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_727, plain, (l1_struct_0(esk19_1(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_245, plain, (v1_xboole_0(esk16_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_639, plain, (u1_struct_0(X1)=k2_struct_0(X2)|~l1_struct_0(X2)|~v2_struct_0(X2)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_566, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_243, plain, (v1_waybel_0(u1_struct_0(X1),X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_246, plain, (v1_xboole_0(k2_struct_0(X1))|~v1_xboole_0(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_778, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X4)|v2_struct_0(X1)|X5!=k1_xboole_0|~r1_waybel_0(X4,X2,X5)|~l1_struct_0(X4)|~l1_struct_0(X1)|~m1_subset_1(esk10_4(X1,X2,X3,esk9_3(X4,X2,X5)),u1_struct_0(X2))|~l1_waybel_0(X2,X4)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_789, plain, (v1_compts_1(X1)|v2_struct_0(X1)|r1_waybel_0(X4,X3,X5)|v1_xboole_0(esk4_1(X1))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k2_waybel_0(X2,X3,esk10_4(X4,X3,X5,esk9_3(X2,X3,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))))),u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~r1_waybel_0(X2,X3,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))))|~l1_struct_0(X2)|~l1_struct_0(X4)|~m1_subset_1(esk10_4(X4,X3,X5,esk9_3(X2,X3,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))))),u1_struct_0(X3))|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X1)))))|~v13_waybel_0(esk4_1(X1),k3_yellow_1(u1_struct_0(X1)))|~v2_waybel_0(esk4_1(X1),k3_yellow_1(u1_struct_0(X1)))|~l1_waybel_0(X3,X2)|~l1_waybel_0(X3,X4))).#trainneg
cnf(i_0_358, plain, (m1_subset_1(X1,u1_struct_0(k3_yellow_1(k2_struct_0(X2))))|v1_compts_1(X2)|v2_struct_0(X2)|~r2_hidden(X1,esk4_1(X2))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_499, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X5)|v2_struct_0(X6)|k10_yellow_6(X1,X4)!=k10_yellow_6(X1,X2)|~m1_connsp_2(X3,X1,k2_waybel_0(X5,X6,esk7_4(X5,X6,k10_yellow_6(X1,X4),X7)))|~l1_waybel_0(X2,X1)|~l1_waybel_0(X4,X1)|~v7_waybel_0(X2)|~v7_waybel_0(X4)|~v4_orders_2(X2)|~v4_orders_2(X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~r2_waybel_0(X5,X6,k10_yellow_6(X1,X4))|~l1_struct_0(X5)|~m1_subset_1(X7,u1_struct_0(X6))|~l1_waybel_0(X6,X5))).#trainneg
cnf(i_0_494, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X3)|v2_struct_0(X2)|v1_xboole_0(k10_yellow_6(X2,X3))|~l1_waybel_0(X3,X2)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,k10_yellow_6(X2,X3)))).#trainneg
cnf(i_0_784, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_hidden(k2_waybel_0(X1,X2,esk7_4(X1,X2,k4_xboole_0(X3,X4),X5)),X4)|~r2_waybel_0(X1,X2,k4_xboole_0(X3,X4))|~l1_struct_0(X1)|~m1_subset_1(X5,u1_struct_0(X2))|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_523, plain, (v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|~r2_hidden(k2_waybel_0(X1,X2,esk7_4(X3,X2,X4,esk9_3(X1,X2,k1_xboole_0))),X5)|~r2_waybel_0(X3,X2,X4)|~r1_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X1)|~l1_struct_0(X3)|~m1_subset_1(esk7_4(X3,X2,X4,esk9_3(X1,X2,k1_xboole_0)),u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_waybel_0(X2,X3))).#trainneg
cnf(i_0_353, plain, (r2_waybel_0(X1,X2,esk18_2(X1,X3))|v2_struct_0(X1)|v2_struct_0(X2)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_657, plain, (v4_pre_topc(k1_xboole_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_568, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_319, plain, (u1_struct_0(X1)=X2|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(X2,X3)|~l1_orders_2(X1))).#trainneg
cnf(i_0_248, plain, (v1_relat_1(esk16_1(k1_zfmisc_1(X1)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_658, plain, (v3_pre_topc(k1_xboole_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_255, plain, (v1_finset_1(esk16_1(k1_zfmisc_1(X1)))|~v1_finset_1(X1))).#trainneg
cnf(i_0_265, plain, (k3_subset_1(k1_xboole_0,X1)=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_567, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_332, plain, (v1_relat_1(u1_waybel_0(X1,X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_249, plain, (v1_relat_1(k2_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_256, plain, (v1_finset_1(k2_struct_0(X1))|~v1_finset_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_619, plain, (v1_xboole_0(X1)|X1!=k1_xboole_0|~r2_hidden(X2,X3)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_900, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X2,X1,X3)|~r1_waybel_0(X2,X1,k1_xboole_0)|~l1_struct_0(X2)|~m1_subset_1(esk7_4(X2,X1,X3,esk9_3(X2,X1,k1_xboole_0)),u1_struct_0(X1))|~l1_waybel_0(X1,X2)|~m1_subset_1(esk9_3(X2,X1,k1_xboole_0),u1_struct_0(X1)))).#trainneg
cnf(i_0_810, plain, (v1_compts_1(X1)|v2_struct_0(X2)|v2_struct_0(X1)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~r2_waybel_0(X1,X2,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_895, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X2,X1,k4_xboole_0(X3,X4))|~l1_struct_0(X2)|~m1_subset_1(esk9_3(X2,X1,X4),u1_struct_0(X1))|~l1_waybel_0(X1,X2)|~r1_waybel_0(X2,X1,X4)|~m1_subset_1(esk7_4(X2,X1,k4_xboole_0(X3,X4),esk9_3(X2,X1,X4)),u1_struct_0(X1)))).#trainneg
cnf(i_0_882, plain, (r1_waybel_0(X1,X2,esk6_3(X1,X3,k2_waybel_0(X4,X5,esk7_4(X4,X5,k10_yellow_6(X1,X6),X7))))|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X6)|r3_waybel_9(X1,X3,k2_waybel_0(X4,X5,esk7_4(X4,X5,k10_yellow_6(X1,X6),X7)))|v2_struct_0(X3)|k10_yellow_6(X1,X6)!=k10_yellow_6(X1,X2)|~r2_waybel_0(X4,X5,k10_yellow_6(X1,X6))|~l1_struct_0(X4)|~m1_subset_1(X7,u1_struct_0(X5))|~l1_waybel_0(X2,X1)|~l1_waybel_0(X6,X1)|~l1_waybel_0(X5,X4)|~v7_waybel_0(X2)|~v7_waybel_0(X6)|~v4_orders_2(X2)|~v4_orders_2(X6)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(k2_waybel_0(X4,X5,esk7_4(X4,X5,k10_yellow_6(X1,X6),X7)),u1_struct_0(X1))|~l1_waybel_0(X3,X1))).#trainneg
cnf(i_0_798, plain, (v1_finset_1(k1_xboole_0))).#trainneg
cnf(i_0_799, plain, (l1_struct_0(g1_orders_2(X1,k1_xboole_0)))).#trainneg
cnf(i_0_843, plain, (X1=u1_struct_0(X2)|g1_orders_2(X1,k1_xboole_0)!=X2|~l1_orders_2(X2)|~v1_orders_2(X2))).#trainneg
cnf(i_0_258, plain, (esk16_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1)))=X2|g1_orders_2(X1,esk16_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))!=g1_orders_2(X3,X2))).#trainneg
cnf(i_0_861, plain, (~l1_orders_2(X1)|~r2_hidden(X2,u1_orders_2(X1))|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_857, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_914, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_261, plain, (v1_xboole_0(esk16_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_264, plain, (v1_waybel_0(u1_struct_0(X1),X1)|v2_struct_0(X1)|~v2_yellow_0(X1)|~l1_orders_2(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_936, plain, (v1_xboole_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|X1!=k1_xboole_0|~m1_subset_1(k2_waybel_0(X2,X3,esk7_4(X2,X3,X4,X5)),X1)|~r2_waybel_0(X2,X3,X4)|~l1_struct_0(X2)|~m1_subset_1(X5,u1_struct_0(X3))|~l1_waybel_0(X3,X2))).#trainneg
cnf(i_0_838, plain, (r2_hidden(k2_waybel_0(X1,X2,esk7_4(X1,X2,k2_struct_0(X3),X4)),u1_struct_0(X3))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_struct_0(X3)|~r2_waybel_0(X1,X2,k2_struct_0(X3))|~l1_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_872, plain, (r1_waybel_0(X1,X2,X3)|v1_compts_1(X4)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X2)|~v1_subset_1(esk4_1(X4),u1_struct_0(k3_yellow_1(u1_struct_0(X4))))|~r1_waybel_0(X4,X2,k10_yellow_6(X4,k3_yellow19(X4,u1_struct_0(X4),esk4_1(X4))))|~l1_struct_0(X4)|~l1_struct_0(X1)|~m1_subset_1(esk10_4(X1,X2,X3,esk9_3(X4,X2,k10_yellow_6(X4,k3_yellow19(X4,u1_struct_0(X4),esk4_1(X4))))),u1_struct_0(X2))|~v13_waybel_0(esk4_1(X4),k3_yellow_1(u1_struct_0(X4)))|~v2_waybel_0(esk4_1(X4),k3_yellow_1(u1_struct_0(X4)))|~l1_waybel_0(X2,X4)|~l1_waybel_0(X2,X1)|~l1_pre_topc(X4)|~v2_pre_topc(X4))).#trainneg
cnf(i_0_899, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v1_xboole_0(X5)|~r2_waybel_0(X3,X1,X4)|~r1_waybel_0(X2,X1,k1_xboole_0)|~l1_struct_0(X2)|~l1_struct_0(X3)|~m1_subset_1(esk7_4(X3,X1,X4,esk9_3(X2,X1,k1_xboole_0)),u1_struct_0(X1))|~l1_waybel_0(X1,X2)|~l1_waybel_0(X1,X3)|~m1_subset_1(k2_waybel_0(X2,X1,esk7_4(X3,X1,X4,esk9_3(X2,X1,k1_xboole_0))),X5))).#trainneg
cnf(i_0_805, plain, (v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X3,X2,X4)|~r1_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X1)|~l1_struct_0(X3)|~m1_subset_1(esk7_4(X3,X2,X4,esk9_3(X1,X2,k1_xboole_0)),u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_waybel_0(X2,X3))).#trainneg
cnf(i_0_879, negated_conjecture, (r1_waybel_0(esk1_0,X1,X2)|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(esk2_0)|k10_yellow_6(esk1_0,X1)!=k1_xboole_0|~r2_waybel_0(X4,X3,k1_xboole_0)|~m1_connsp_2(X2,esk1_0,k2_waybel_0(X4,X3,esk7_4(X4,X3,k1_xboole_0,X5)))|~l1_struct_0(X4)|~m1_subset_1(X5,u1_struct_0(X3))|~l1_waybel_0(X1,esk1_0)|~l1_waybel_0(esk2_0,esk1_0)|~l1_waybel_0(X3,X4)|~v7_waybel_0(X1)|~v7_waybel_0(esk2_0)|~v4_orders_2(X1)|~v4_orders_2(esk2_0)|~v1_compts_1(esk1_0))).#trainneg
cnf(i_0_809, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X2)|~r1_waybel_0(X4,X2,k1_xboole_0)|~l1_struct_0(X4)|~l1_struct_0(X1)|~l1_waybel_0(X2,X4)|~l1_waybel_0(X2,X1)|~m1_subset_1(esk9_3(X4,X2,k1_xboole_0),u1_struct_0(X2)))).#trainneg
cnf(i_0_823, plain, (r2_hidden(k2_waybel_0(X1,X2,esk7_4(X1,X2,k10_yellow_6(X3,X4),X5)),u1_struct_0(X3))|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|~l1_waybel_0(X4,X3)|~v7_waybel_0(X4)|~v4_orders_2(X4)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~r2_waybel_0(X1,X2,k10_yellow_6(X3,X4))|~l1_struct_0(X1)|~m1_subset_1(X5,u1_struct_0(X2))|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_816, plain, (r2_hidden(k2_waybel_0(X1,X2,esk7_4(X1,X2,esk4_1(X3),X4)),u1_struct_0(k3_yellow_1(k2_struct_0(X3))))|v1_compts_1(X3)|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~r2_waybel_0(X1,X2,esk4_1(X3))|~l1_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_942, plain, (v1_xboole_0(X1)|v1_xboole_0(X3)|X1!=k1_xboole_0|~m1_subset_1(X2,X1)|~m1_subset_1(X2,X3))).#trainneg
cnf(i_0_569, plain, (X1=X2|g1_orders_2(X1,k2_zfmisc_1(X1,X1))!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_559, plain, (v1_orders_2(g1_orders_2(X1,k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_800, plain, (u1_orders_2(g1_orders_2(X1,k1_xboole_0))=k1_xboole_0)).#trainneg
cnf(i_0_847, plain, (u1_orders_2(X1)=X2|X1!=g1_orders_2(X3,X2)|~l1_orders_2(X1)|~v1_orders_2(X1))).#trainneg
cnf(i_0_266, plain, (k3_subset_1(X1,k1_xboole_0)=X1|~m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_850, plain, (X1=esk16_1(k1_zfmisc_1(X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_269, plain, (v4_relat_1(esk16_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_271, plain, (v5_relat_1(esk16_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_944, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X1,X2,k1_xboole_0)|~r1_waybel_0(X1,X2,X3)|~l1_struct_0(X1)|~m1_subset_1(esk7_4(X1,X2,k1_xboole_0,esk9_3(X1,X2,X3)),u1_struct_0(X2))|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_1057, plain, (r1_waybel_0(X1,X2,u1_struct_0(X1))|r1_waybel_0(X1,X2,k1_xboole_0)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~v1_yellow19(X2,X1)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_1059, plain, (k3_subset_1(X1,X1)=k1_xboole_0)).#trainneg
cnf(i_0_1094, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X1)|k10_yellow_6(X1,X4)=k1_xboole_0|k10_yellow_6(X1,X4)!=k10_yellow_6(X1,X2)|~m1_connsp_2(X3,X1,esk15_3(k1_xboole_0,X5,k10_yellow_6(X1,X4)))|~l1_waybel_0(X2,X1)|~l1_waybel_0(X4,X1)|~v7_waybel_0(X2)|~v7_waybel_0(X4)|~v4_orders_2(X2)|~v4_orders_2(X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_560, plain, (l1_orders_2(g1_orders_2(X1,k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_952, plain, (u1_struct_0(g1_orders_2(X1,k1_xboole_0))=X1)).#trainneg
cnf(i_0_905, plain, (u1_struct_0(X1)=X2|X1!=g1_orders_2(X2,X3)|~l1_orders_2(X1)|~v1_orders_2(X1))).#trainneg
cnf(i_0_275, plain, (r2_hidden(X1,X2)|X3!=k3_subset_1(X2,X4)|~r2_hidden(X1,X3)|~m1_subset_1(X4,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_908, plain, (k3_subset_1(k1_xboole_0,k1_xboole_0)=X1|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_851, plain, (esk16_1(k1_zfmisc_1(X1))=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_964, plain, (u1_struct_0(X1)=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X3)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_971, plain, (u1_struct_0(X1)=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X2)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_277, plain, (k3_subset_1(k2_zfmisc_1(X1,X1),X2)=X3|g1_orders_2(X1,k3_subset_1(k2_zfmisc_1(X1,X1),X2))!=g1_orders_2(X4,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_278, plain, (v5_relat_1(k3_subset_1(k2_zfmisc_1(X1,X2),X3),X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_943, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X1,X2,X3)|~r1_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(esk9_3(X1,X2,k1_xboole_0),u1_struct_0(X2)))).#trainneg
cnf(i_0_1085, plain, (k4_xboole_0(X2,X3)=k1_xboole_0|~r2_hidden(esk15_3(k1_xboole_0,X1,k4_xboole_0(X2,X3)),X3))).#trainneg
cnf(i_0_1082, plain, (r2_hidden(esk15_3(k1_xboole_0,X1,esk4_1(X2)),u1_struct_0(k3_yellow_1(k2_struct_0(X2))))|v1_compts_1(X2)|v2_struct_0(X2)|esk4_1(X2)=k1_xboole_0|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_1083, plain, (r2_hidden(esk15_3(k1_xboole_0,X1,esk4_1(X2)),u1_struct_0(k3_yellow_1(u1_struct_0(X2))))|v1_compts_1(X2)|v2_struct_0(X2)|esk4_1(X2)=k1_xboole_0|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_1104, negated_conjecture, (r1_waybel_0(esk1_0,esk2_0,u1_struct_0(esk1_0))|r1_waybel_0(esk1_0,esk2_0,k1_xboole_0)|v2_struct_0(esk2_0)|~l1_waybel_0(esk2_0,esk1_0)|~v1_compts_1(esk1_0))).#trainneg
cnf(i_0_1228, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r1_waybel_0(X2,X1,k1_xboole_0)|~l1_struct_0(X2)|~l1_waybel_0(X1,X2)|~r3_waybel_9(X2,X1,X3)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_1103, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X2,X1,k1_xboole_0)|~r1_waybel_0(X2,X1,X3)|~l1_struct_0(X2)|~l1_waybel_0(X1,X2)|~m1_subset_1(esk9_3(X2,X1,X3),u1_struct_0(X1)))).#trainneg
cnf(i_0_1141, plain, (v1_xboole_0(X1)|~v2_struct_0(g1_orders_2(X1,k1_xboole_0)))).#trainneg
cnf(i_0_1142, plain, (v2_struct_0(g1_orders_2(X1,k1_xboole_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1197, plain, (esk16_1(k1_zfmisc_1(k1_xboole_0))=k1_xboole_0)).#trainneg
cnf(i_0_279, plain, (v1_relat_1(k3_subset_1(k2_zfmisc_1(X1,X2),X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1194, plain, (r2_hidden(X1,X2)|X3!=k1_xboole_0|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_280, plain, (v1_xboole_0(k3_subset_1(k2_zfmisc_1(X1,X2),X3))|~v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_281, plain, (v4_relat_1(k3_subset_1(k2_zfmisc_1(X1,X2),X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1096, plain, (r2_hidden(esk15_3(k1_xboole_0,X1,k10_yellow_6(X2,X3)),u1_struct_0(X2))|v2_struct_0(X3)|v2_struct_0(X2)|k10_yellow_6(X2,X3)=k1_xboole_0|~l1_waybel_0(X3,X2)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_947, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X1,X2,k4_xboole_0(X3,X4))|~r1_waybel_0(X1,X2,X4)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(esk9_3(X1,X2,X4),u1_struct_0(X2)))).#trainneg
cnf(i_0_1230, plain, (k3_subset_1(X1,X2)=k1_xboole_0|~r2_hidden(esk15_3(k1_xboole_0,X3,k3_subset_1(X1,X2)),X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1034, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_hidden(k2_waybel_0(X2,X1,esk7_4(X2,X1,X4,X5)),k3_subset_1(X3,X4))|~r2_waybel_0(X2,X1,X4)|~l1_struct_0(X2)|~m1_subset_1(k3_subset_1(X3,X4),k1_zfmisc_1(X3))|~m1_subset_1(X5,u1_struct_0(X1))|~l1_waybel_0(X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_1086, plain, (r2_hidden(esk15_3(k1_xboole_0,X1,k4_xboole_0(X2,X3)),X2)|k4_xboole_0(X2,X3)=k1_xboole_0)).#trainneg
cnf(i_0_991, plain, (r1_waybel_0(X1,X2,X3)|v1_compts_1(X4)|v2_struct_0(X2)|v2_struct_0(X4)|v2_struct_0(X1)|~v1_subset_1(esk4_1(X4),u1_struct_0(k3_yellow_1(u1_struct_0(X4))))|~r1_waybel_0(X4,X2,k10_yellow_6(X4,k3_yellow19(X4,u1_struct_0(X4),esk4_1(X4))))|~l1_struct_0(X1)|~v13_waybel_0(esk4_1(X4),k3_yellow_1(u1_struct_0(X4)))|~v2_waybel_0(esk4_1(X4),k3_yellow_1(u1_struct_0(X4)))|~l1_waybel_0(X2,X4)|~l1_waybel_0(X2,X1)|~l1_pre_topc(X4)|~v2_pre_topc(X4)|~m1_subset_1(esk9_3(X4,X2,k10_yellow_6(X4,k3_yellow19(X4,u1_struct_0(X4),esk4_1(X4)))),u1_struct_0(X2)))).#trainneg
cnf(i_0_1027, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|v2_struct_0(X4)|v2_struct_0(X3)|X2!=k1_xboole_0|~m1_subset_1(k2_waybel_0(X3,X4,esk7_4(X3,X4,X2,X5)),X1)|~r2_waybel_0(X3,X4,X2)|~l1_struct_0(X3)|~m1_subset_1(X5,u1_struct_0(X4))|~l1_waybel_0(X4,X3))).#trainneg
cnf(i_0_997, negated_conjecture, (r1_waybel_0(esk1_0,X1,esk11_4(esk1_0,X2,X3,k2_waybel_0(X4,X5,esk7_4(X4,X5,k1_xboole_0,X6))))|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X5)|r2_hidden(k2_waybel_0(X4,X5,esk7_4(X4,X5,k1_xboole_0,X6)),X3)|v2_struct_0(X2)|k10_yellow_6(esk1_0,X1)!=k1_xboole_0|~r2_waybel_0(X4,X5,k1_xboole_0)|~l1_struct_0(X4)|~m1_subset_1(X6,u1_struct_0(X5))|~l1_waybel_0(X1,esk1_0)|~l1_waybel_0(X5,X4)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~v1_compts_1(esk1_0)|X3!=k10_yellow_6(esk1_0,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_waybel_0(X4,X5,esk7_4(X4,X5,k1_xboole_0,X6)),u1_struct_0(esk1_0))|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(X2)|~v4_orders_2(X2))).#trainneg
cnf(i_0_1124, plain, (l1_struct_0(g1_orders_2(X1,k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_283, plain, (v1_xboole_0(k3_subset_1(X1,X2))|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_284, plain, (v1_relat_1(k3_subset_1(X1,X2))|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_285, plain, (v1_finset_1(k3_subset_1(X1,X2))|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1004, negated_conjecture, (r2_hidden(k2_waybel_0(X1,X2,esk7_4(X1,X2,k1_xboole_0,X3)),u1_struct_0(esk1_0))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(esk2_0)|~r2_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_waybel_0(esk2_0,esk1_0)|~l1_waybel_0(X2,X1)|~v7_waybel_0(esk2_0)|~v4_orders_2(esk2_0)|~v1_compts_1(esk1_0))).#trainneg
cnf(i_0_1087, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|k3_subset_1(u1_struct_0(X1),X4)=k1_xboole_0|k3_subset_1(u1_struct_0(X1),X4)!=k10_yellow_6(X1,X2)|~m1_connsp_2(X3,X1,esk15_3(k1_xboole_0,X5,k3_subset_1(u1_struct_0(X1),X4)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_995, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~r2_waybel_0(X3,X1,X4)|~r1_waybel_0(X2,X1,k1_xboole_0)|~l1_struct_0(X2)|~l1_struct_0(X3)|~l1_waybel_0(X1,X2)|~l1_waybel_0(X1,X3)|~m1_subset_1(esk9_3(X2,X1,k1_xboole_0),u1_struct_0(X1)))).#trainneg
cnf(i_0_1023, plain, (v1_xboole_0(X1)|v1_xboole_0(k10_yellow_6(X2,k3_yellow19(X2,k2_struct_0(X2),X3)))|v1_xboole_0(X3)|v2_struct_0(X2)|k10_yellow_6(X2,k3_yellow19(X2,k2_struct_0(X2),X3))!=k1_xboole_0|~m1_subset_1(o_2_13_yellow19(X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X2)))))|~v3_waybel_7(X3,k3_yellow_1(k2_struct_0(X2)))|~v13_waybel_0(X3,k3_yellow_1(k2_struct_0(X2)))|~v2_waybel_0(X3,k3_yellow_1(k2_struct_0(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_948, negated_conjecture, (r3_waybel_9(esk1_0,X1,k2_waybel_0(X2,X3,esk7_4(X2,X3,k1_xboole_0,X4)))|r1_waybel_0(esk1_0,X5,esk6_3(esk1_0,X1,k2_waybel_0(X2,X3,esk7_4(X2,X3,k1_xboole_0,X4))))|v2_struct_0(X1)|v2_struct_0(esk2_0)|v2_struct_0(X5)|v2_struct_0(X2)|v2_struct_0(X3)|k10_yellow_6(esk1_0,X5)!=k1_xboole_0|~r2_waybel_0(X2,X3,k1_xboole_0)|~l1_struct_0(X2)|~m1_subset_1(k2_waybel_0(X2,X3,esk7_4(X2,X3,k1_xboole_0,X4)),u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_waybel_0(X5,esk1_0)|~l1_waybel_0(esk2_0,esk1_0)|~l1_waybel_0(X3,X2)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X5)|~v7_waybel_0(esk2_0)|~v4_orders_2(X5)|~v4_orders_2(esk2_0)|~v1_compts_1(esk1_0))).#trainneg
cnf(i_0_1081, plain, (m1_subset_1(esk15_3(k1_xboole_0,X1,esk4_1(X2)),u1_struct_0(k3_yellow_1(k2_struct_0(X2))))|v1_compts_1(X2)|v2_struct_0(X2)|esk4_1(X2)=k1_xboole_0|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_1080, plain, (m1_subset_1(esk15_3(k1_xboole_0,X1,esk4_1(X2)),u1_struct_0(k3_yellow_1(u1_struct_0(X2))))|v1_compts_1(X2)|v2_struct_0(X2)|esk4_1(X2)=k1_xboole_0|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_1089, plain, (r2_hidden(esk15_3(k1_xboole_0,X1,k3_subset_1(X2,X3)),X2)|k3_subset_1(X2,X3)=k1_xboole_0|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1108, plain, (v4_pre_topc(k1_xboole_0,X1)|~v3_pre_topc(u1_struct_0(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_288, plain, (v1_orders_2(g1_orders_2(X1,esk16_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))).#trainneg
cnf(i_0_975, plain, (X1=esk16_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1073, plain, (u1_struct_0(X1)=esk16_1(k1_zfmisc_1(X2))|~v1_xboole_0(X2)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_289, plain, (v1_orders_2(g1_orders_2(X1,k3_subset_1(k2_zfmisc_1(X1,X1),X2)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_291, plain, (l1_orders_2(g1_orders_2(X1,esk16_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))).#trainneg
cnf(i_0_1272, plain, (X1=k4_xboole_0(X2,X3)|r2_hidden(esk15_3(X2,X3,X1),X2)|X1!=k1_xboole_0)).#trainneg
cnf(i_0_1254, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r1_waybel_0(X1,X2,k1_xboole_0)|~m1_subset_1(esk14_2(X1,X2),u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~v1_compts_1(X1))).#trainneg
cnf(i_0_1350, plain, (r1_waybel_0(X1,X2,X3)|v1_compts_1(X4)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X2)|~v1_subset_1(esk4_1(X4),u1_struct_0(k3_yellow_1(u1_struct_0(X4))))|~r1_waybel_0(X4,X2,k10_yellow_6(X4,k3_yellow19(X4,u1_struct_0(X4),esk4_1(X4))))|~l1_struct_0(X1)|~v13_waybel_0(esk4_1(X4),k3_yellow_1(u1_struct_0(X4)))|~v2_waybel_0(esk4_1(X4),k3_yellow_1(u1_struct_0(X4)))|~l1_waybel_0(X2,X4)|~l1_waybel_0(X2,X1)|~l1_pre_topc(X4)|~v2_pre_topc(X4)|~l1_struct_0(X4))).#trainneg
cnf(i_0_1271, plain, (X3=k4_xboole_0(X1,X2)|r2_hidden(esk15_3(X1,X2,X3),X3)|X1!=k1_xboole_0)).#trainneg
cnf(i_0_1317, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X1,X2,X3)|~l1_struct_0(X1)|~m1_subset_1(esk9_3(X1,X2,k3_subset_1(X4,X3)),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(X4))|~l1_waybel_0(X2,X1)|~r1_waybel_0(X1,X2,k3_subset_1(X4,X3))|~m1_subset_1(esk7_4(X1,X2,X3,esk9_3(X1,X2,k3_subset_1(X4,X3))),u1_struct_0(X2)))).#trainneg
cnf(i_0_1382, negated_conjecture, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X2,X1,k1_xboole_0)|~l1_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X1,X2)|~v1_compts_1(esk1_0))).#trainneg
cnf(i_0_1109, plain, (v3_pre_topc(k1_xboole_0,X1)|~v4_pre_topc(u1_struct_0(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1144, plain, (m1_subset_1(k2_struct_0(g1_orders_2(X1,k1_xboole_0)),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_292, plain, (l1_orders_2(g1_orders_2(X1,k3_subset_1(k2_zfmisc_1(X1,X1),X2)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_858, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1539, plain, (k2_struct_0(g1_orders_2(k1_xboole_0,k1_xboole_0))=k1_xboole_0)).#trainneg
cnf(i_0_1540, plain, (~r2_hidden(X1,k2_struct_0(g1_orders_2(X2,k1_xboole_0)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_607, plain, (m1_subset_1(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_1360, plain, (X1=k3_subset_1(X2,X3)|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_294, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk16_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_295, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k3_subset_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1345, plain, (k4_xboole_0(X1,X1)=k1_xboole_0)).#trainneg
cnf(i_0_1429, plain, (X2=k1_xboole_0|r2_hidden(esk15_3(k1_xboole_0,X3,X2),X1)|~m1_subset_1(k3_subset_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1714, plain, (k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))=k1_xboole_0|v1_compts_1(X1)|v2_struct_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))|v2_struct_0(X1)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~l1_waybel_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)),X1)|~v7_waybel_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))|~v4_orders_2(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_1597, plain, (v2_struct_0(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_1712, plain, (v1_compts_1(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)))|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~r2_waybel_0(X2,X3,k1_xboole_0)|~l1_struct_0(X2)|~m1_subset_1(k2_waybel_0(X2,X3,esk7_4(X2,X3,k1_xboole_0,X4)),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_waybel_0(X3,X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)),X1)|~v7_waybel_0(k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)))|~v4_orders_2(k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1))))).#trainneg
cnf(i_0_1589, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1528, plain, (v1_relat_1(k2_struct_0(g1_orders_2(k2_zfmisc_1(X1,X2),k1_xboole_0))))).#trainneg
cnf(i_0_1480, plain, (X3=k4_xboole_0(X1,X2)|X1!=k1_xboole_0|X3!=k1_xboole_0)).#trainneg
cnf(i_0_296, plain, (m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,k2_struct_0(X2))|~l1_struct_0(X2))).#trainneg
cnf(i_0_1596, plain, (v1_tops_1(k1_xboole_0,g1_orders_2(k1_xboole_0,k1_xboole_0))|~l1_pre_topc(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_1603, plain, (v1_waybel_0(k1_xboole_0,g1_orders_2(k1_xboole_0,k1_xboole_0))|~v1_lattice3(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_915, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1711, plain, (v1_compts_1(X1)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)))|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~r2_waybel_0(X1,X2,k1_xboole_0)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)),X1)|~v7_waybel_0(k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)))|~v4_orders_2(k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1))))).#trainneg
cnf(i_0_1752, plain, (X2=k1_xboole_0|~v1_xboole_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_1590, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_314, plain, (v5_relat_1(u1_orders_2(X1),u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_300, plain, (X1=X2|g1_orders_2(X1,esk16_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_301, plain, (X1=X2|g1_orders_2(X1,k3_subset_1(k2_zfmisc_1(X1,X1),X3))!=g1_orders_2(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_1453, plain, (k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))=k1_xboole_0|v1_compts_1(X1)|v2_struct_0(X1)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~m1_subset_1(esk15_3(k1_xboole_0,X2,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))),u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_1278, plain, (v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X3)|X1!=k1_xboole_0|~r2_waybel_0(X4,X3,X5)|~r1_waybel_0(X2,X3,X1)|~l1_struct_0(X2)|~l1_struct_0(X4)|~m1_subset_1(esk7_4(X4,X3,X5,esk9_3(X2,X3,X1)),u1_struct_0(X3))|~l1_waybel_0(X3,X2)|~l1_waybel_0(X3,X4))).#trainneg
cnf(i_0_1397, plain, (k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2))=k1_xboole_0|r2_waybel_7(X1,X2,esk15_3(k1_xboole_0,X3,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2))))|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~m1_subset_1(esk15_3(k1_xboole_0,X3,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2))),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X1)))))|~v13_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~v2_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_1305, plain, (X2=k1_xboole_0|~r2_hidden(esk15_3(k1_xboole_0,X3,X2),k3_subset_1(X1,X2))|~m1_subset_1(k3_subset_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1760, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk4_1(esk1_0)))|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))|~v3_waybel_7(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk4_1(esk1_0)),esk1_0)|~v7_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk4_1(esk1_0)))|~v4_orders_2(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk4_1(esk1_0)))|~v1_subset_1(esk4_1(esk1_0),u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1793, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_1890, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_1478, plain, (X3=k4_xboole_0(X1,X2)|~v1_xboole_0(X1)|X3!=k1_xboole_0)).#trainneg
cnf(i_0_304, plain, (v1_xboole_0(esk16_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_305, plain, (v1_xboole_0(k3_subset_1(k2_zfmisc_1(X1,X2),X3))|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_307, plain, (v1_xboole_0(u1_orders_2(X1))|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1776, plain, (r1_waybel_0(X1,X2,X3)|v1_compts_1(X4)|v2_struct_0(X2)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(k3_yellow19(X4,u1_struct_0(X4),esk4_1(X4)))|~v1_subset_1(esk4_1(X4),u1_struct_0(k3_yellow_1(u1_struct_0(X4))))|~r1_waybel_0(X4,X2,k1_xboole_0)|~l1_struct_0(X1)|~v13_waybel_0(esk4_1(X4),k3_yellow_1(u1_struct_0(X4)))|~v2_waybel_0(esk4_1(X4),k3_yellow_1(u1_struct_0(X4)))|~l1_waybel_0(X2,X4)|~l1_waybel_0(X2,X1)|~l1_pre_topc(X4)|~v2_pre_topc(X4)|~l1_waybel_0(k3_yellow19(X4,u1_struct_0(X4),esk4_1(X4)),X4)|~v7_waybel_0(k3_yellow19(X4,u1_struct_0(X4),esk4_1(X4)))|~v4_orders_2(k3_yellow19(X4,u1_struct_0(X4),esk4_1(X4))))).#trainneg
cnf(i_0_1769, plain, (v1_compts_1(X1)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~r2_waybel_0(X1,X2,k1_xboole_0)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)),X1)|~v7_waybel_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))|~v4_orders_2(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))))).#trainneg
cnf(i_0_318, plain, (v4_relat_1(u1_orders_2(X1),u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1461, plain, (X3=k4_xboole_0(u1_orders_2(X1),X2)|~l1_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1))|X3!=k1_xboole_0)).#trainneg
cnf(i_0_309, plain, (v1_finset_1(u1_orders_2(X1))|~v1_finset_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_310, plain, (m1_subset_1(X1,k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2)))|~r2_hidden(X1,u1_orders_2(X2))|~l1_orders_2(X2))).#trainneg
cnf(i_0_311, plain, (v1_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1705, plain, (r2_hidden(k2_waybel_0(X1,X2,esk7_4(X1,X2,k1_xboole_0,X4)),u1_struct_0(X3))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(k3_yellow19(X3,k2_struct_0(X3),esk4_1(X3)))|v1_compts_1(X3)|~r2_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_waybel_0(k3_yellow19(X3,k2_struct_0(X3),esk4_1(X3)),X3)|~l1_waybel_0(X2,X1)|~v7_waybel_0(k3_yellow19(X3,k2_struct_0(X3),esk4_1(X3)))|~v4_orders_2(k3_yellow19(X3,k2_struct_0(X3),esk4_1(X3)))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~v1_subset_1(esk4_1(X3),u1_struct_0(k3_yellow_1(k2_struct_0(X3)))))).#trainneg
cnf(i_0_1633, plain, (v2_struct_0(X2)|v2_struct_0(X1)|~r2_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_1337, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X1)|k4_xboole_0(k10_yellow_6(X1,X4),X6)=k1_xboole_0|k10_yellow_6(X1,X4)!=k10_yellow_6(X1,X2)|~m1_connsp_2(X3,X1,esk15_3(k1_xboole_0,X5,k4_xboole_0(k10_yellow_6(X1,X4),X6)))|~l1_waybel_0(X2,X1)|~l1_waybel_0(X4,X1)|~v7_waybel_0(X2)|~v7_waybel_0(X4)|~v4_orders_2(X2)|~v4_orders_2(X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1324, plain, (k4_xboole_0(k4_xboole_0(X2,X3),X4)=k1_xboole_0|~r2_hidden(esk15_3(k1_xboole_0,X1,k4_xboole_0(k4_xboole_0(X2,X3),X4)),X3))).#trainneg
cnf(i_0_1342, plain, (k4_xboole_0(X1,X3)=k1_xboole_0|X1!=k1_xboole_0)).#trainneg
cnf(i_0_1541, plain, (v1_xboole_0(k2_struct_0(g1_orders_2(X1,k1_xboole_0)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1467, plain, (X3=k4_xboole_0(k2_struct_0(X1),X2)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|X3!=k1_xboole_0)).#trainneg
cnf(i_0_312, plain, (l1_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1588, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1854, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_316, plain, (v1_xboole_0(u1_orders_2(X1))|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1279, plain, (r2_hidden(X3,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0|~r2_hidden(X3,X1))).#trainneg
cnf(i_0_1425, plain, (k3_subset_1(X1,X3)=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X3,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1921, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk4_1(esk1_0)))|v1_xboole_0(u1_struct_0(esk1_0))|~v1_subset_1(esk4_1(esk1_0),u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))|~v3_waybel_7(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk4_1(esk1_0)),esk1_0)|~v4_orders_2(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk4_1(esk1_0))))).#trainneg
cnf(i_0_1920, plain, (r1_waybel_0(X1,X2,esk18_2(X1,X3))|v2_struct_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(esk18_2(X1,X3),k1_zfmisc_1(u1_struct_0(X1)))|~v1_yellow19(X2,X1)|~l1_waybel_0(X2,X1)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1280, plain, (r2_hidden(X3,X2)|k3_subset_1(X1,X2)!=k1_xboole_0|~r2_hidden(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1922, plain, (v1_compts_1(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_1927, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_1456, plain, (X3=k4_xboole_0(esk16_1(k1_zfmisc_1(X1)),X2)|~v1_xboole_0(X1)|X3!=k1_xboole_0)).#trainneg
cnf(i_0_1926, plain, (~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_2039, plain, (X1=k2_struct_0(g1_orders_2(X2,k1_xboole_0))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_327, plain, (v1_relat_1(esk5_2(X1,X2))|v1_yellow19(X2,X1)|v2_struct_0(X1)|v2_struct_0(X2)|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_328, plain, (v1_finset_1(esk5_2(X1,X2))|v1_yellow19(X2,X1)|v2_struct_0(X1)|v2_struct_0(X2)|~v1_finset_1(u1_struct_0(X1))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_1298, plain, (k10_yellow_6(X1,X3)=k1_xboole_0|v2_struct_0(X1)|v2_struct_0(X3)|u1_struct_0(X1)!=k1_xboole_0|~l1_waybel_0(X3,X1)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1754, plain, (X3=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X3,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1310, plain, (k3_subset_1(X1,X2)=k1_xboole_0|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(esk15_3(k1_xboole_0,X3,k3_subset_1(X1,X2)),X2))).#trainneg
cnf(i_0_1327, plain, (k4_xboole_0(k3_subset_1(X2,X3),X4)=k1_xboole_0|~r2_hidden(esk15_3(k1_xboole_0,X1,k4_xboole_0(k3_subset_1(X2,X3),X4)),X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1928, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_1929, plain, (v2_struct_0(X1)|~l1_orders_2(X1)|~m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_329, plain, (m1_subset_1(X1,u1_struct_0(X2))|v1_yellow19(X3,X2)|v2_struct_0(X2)|v2_struct_0(X3)|~r2_hidden(X1,esk5_2(X2,X3))|~l1_struct_0(X2)|~l1_waybel_0(X3,X2))).#trainneg
cnf(i_0_1592, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1939, plain, (X1=esk16_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2059, plain, (X1=u1_orders_2(X2)|~v1_xboole_0(X1)|~l1_orders_2(X2)|~v1_xboole_0(u1_struct_0(X2)))).#trainneg
cnf(i_0_331, plain, (v1_xboole_0(u1_waybel_0(X1,X2))|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1)))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_333, plain, (v1_finset_1(u1_waybel_0(X1,X2))|~v1_finset_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1)))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_2112, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(esk4_1(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk4_1(esk1_0)))|~v1_subset_1(esk4_1(esk1_0),u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))|~v3_waybel_7(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk4_1(esk1_0)),esk1_0))).#trainneg
cnf(i_0_2136, negated_conjecture, (v1_xboole_0(X1)|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~v3_waybel_7(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~v13_waybel_0(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1),esk1_0)|~v7_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1))|~v4_orders_2(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1))|u1_struct_0(esk1_0)!=k1_xboole_0)).#trainneg
cnf(i_0_2135, negated_conjecture, (v1_xboole_0(X1)|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))|~v3_waybel_7(X1,k3_yellow_1(u1_struct_0(esk1_0)))|~v13_waybel_0(X1,k3_yellow_1(u1_struct_0(esk1_0)))|~v2_waybel_0(X1,k3_yellow_1(u1_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),X1),esk1_0)|~v7_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),X1))|~v4_orders_2(k3_yellow19(esk1_0,u1_struct_0(esk1_0),X1))|u1_struct_0(esk1_0)!=k1_xboole_0)).#trainneg
cnf(i_0_2114, plain, (r2_hidden(X1,k3_subset_1(X2,X3))|X3!=k1_xboole_0|~r2_hidden(X1,X2)|~m1_subset_1(k3_subset_1(X2,X3),k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1017, plain, (v1_xboole_0(X1)|v1_xboole_0(u1_struct_0(X2))|r2_waybel_0(X3,X2,X4)|v2_struct_0(X3)|v2_struct_0(X2)|u1_struct_0(X2)!=k1_xboole_0|~m1_subset_1(esk8_3(X3,X2,X4),X1)|~l1_struct_0(X3)|~l1_waybel_0(X2,X3))).#trainneg
cnf(i_0_2166, plain, (u1_orders_2(X1)=k1_xboole_0|k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))!=k1_xboole_0|~l1_orders_2(X1))).#trainneg
cnf(i_0_334, plain, (m1_subset_1(X1,k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)))|~r2_hidden(X1,u1_waybel_0(X3,X2))|~l1_struct_0(X3)|~l1_waybel_0(X2,X3))).#trainneg
cnf(i_0_1850, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1923, plain, (~l1_orders_2(X1)|~r2_hidden(X2,u1_orders_2(X1))|~m1_subset_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_2223, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_335, plain, (v1_orders_2(g1_orders_2(u1_struct_0(X1),u1_waybel_0(X1,X1)))|~l1_struct_0(X1)|~l1_waybel_0(X1,X1))).#trainneg
cnf(i_0_336, plain, (l1_orders_2(g1_orders_2(u1_struct_0(X1),u1_waybel_0(X1,X1)))|~l1_struct_0(X1)|~l1_waybel_0(X1,X1))).#trainneg
cnf(i_0_1296, plain, (k10_yellow_6(X2,X3)=k1_xboole_0|v2_struct_0(X2)|v2_struct_0(X3)|~v1_xboole_0(u1_struct_0(X2))|~l1_waybel_0(X3,X2)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_1384, plain, (X2=k1_xboole_0|r1_waybel_0(X1,X3,X4)|v2_struct_0(X1)|v2_struct_0(X3)|X2!=k10_yellow_6(X1,X3)|~m1_connsp_2(X4,X1,esk15_3(k1_xboole_0,X5,X2))|~m1_subset_1(k3_subset_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_waybel_0(X3,X1)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_937, plain, (v1_xboole_0(X1)|r1_waybel_0(X4,X3,X5)|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X3)|X1!=k1_xboole_0|~m1_subset_1(k2_waybel_0(X2,X3,esk10_4(X4,X3,X5,esk9_3(X2,X3,X6))),X1)|~r1_waybel_0(X2,X3,X6)|~l1_struct_0(X2)|~l1_struct_0(X4)|~m1_subset_1(esk10_4(X4,X3,X5,esk9_3(X2,X3,X6)),u1_struct_0(X3))|~l1_waybel_0(X3,X2)|~l1_waybel_0(X3,X4))).#trainneg
cnf(i_0_1308, plain, (k3_subset_1(X1,k3_subset_1(X2,X3))=k1_xboole_0|r2_hidden(esk15_3(k1_xboole_0,X4,k3_subset_1(X1,k3_subset_1(X2,X3))),X3)|~m1_subset_1(k3_subset_1(X2,X3),k1_zfmisc_1(X1))|~r2_hidden(esk15_3(k1_xboole_0,X4,k3_subset_1(X1,k3_subset_1(X2,X3))),X2)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1091, plain, (r2_hidden(esk15_3(k1_xboole_0,X1,k2_struct_0(X2)),u1_struct_0(X2))|k2_struct_0(X2)=k1_xboole_0|~l1_struct_0(X2))).#trainneg
cnf(i_0_795, plain, (v2_struct_0(X1)|v2_struct_0(X2)|X3!=k1_xboole_0|~r2_waybel_0(X1,X2,X3)|~r1_waybel_0(X1,X2,X4)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(esk9_3(X1,X2,X4),u1_struct_0(X2)))).#trainneg
cnf(i_0_2025, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k1_xboole_0|~m1_subset_1(k4_xboole_0(k4_xboole_0(X1,X2),X3),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_981, plain, (v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X2)|v2_struct_0(X1)|u1_struct_0(X1)!=k1_xboole_0|~r2_waybel_0(X1,X2,X3)|~l1_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~m1_subset_1(esk7_4(X1,X2,X3,X4),u1_struct_0(X2)))).#trainneg
cnf(i_0_2290, plain, (k3_subset_1(X1,X2)!=k1_xboole_0|X2!=k1_xboole_0|~r2_hidden(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_337, plain, (u1_waybel_0(X1,X1)=X2|g1_orders_2(u1_struct_0(X1),u1_waybel_0(X1,X1))!=g1_orders_2(X3,X2)|~l1_struct_0(X1)|~l1_waybel_0(X1,X1))).#trainneg
cnf(i_0_2271, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_xboole_0(u1_struct_0(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk4_1(esk1_0)))|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))|~v3_waybel_7(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk4_1(esk1_0)),esk1_0))).#trainneg
cnf(i_0_2334, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_2129, plain, (u1_struct_0(X1)=k2_struct_0(g1_orders_2(X2,k1_xboole_0))|~v1_xboole_0(X2)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_338, plain, (v5_relat_1(u1_waybel_0(X1,X2),u1_struct_0(X1))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_340, plain, (v1_xboole_0(u1_waybel_0(X1,X2))|~v1_xboole_0(u1_struct_0(X1))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_2345, negated_conjecture, (v1_xboole_0(X1)|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~v3_waybel_7(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~v13_waybel_0(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1),esk1_0)|~v7_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1))|~v4_orders_2(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1))|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2344, negated_conjecture, (v1_xboole_0(X1)|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))|~v3_waybel_7(X1,k3_yellow_1(u1_struct_0(esk1_0)))|~v13_waybel_0(X1,k3_yellow_1(u1_struct_0(esk1_0)))|~v2_waybel_0(X1,k3_yellow_1(u1_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),X1),esk1_0)|~v7_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),X1))|~v4_orders_2(k3_yellow19(esk1_0,u1_struct_0(esk1_0),X1))|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2274, negated_conjecture, (v1_xboole_0(X1)|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1))|v1_xboole_0(k2_struct_0(esk1_0))|u1_struct_0(esk1_0)!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~v3_waybel_7(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~v13_waybel_0(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1),esk1_0)|~v4_orders_2(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1))|~v1_subset_1(X1,u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0))))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2275, negated_conjecture, (v1_xboole_0(X1)|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),X1))|v1_xboole_0(u1_struct_0(esk1_0))|u1_struct_0(esk1_0)!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))|~v3_waybel_7(X1,k3_yellow_1(u1_struct_0(esk1_0)))|~v13_waybel_0(X1,k3_yellow_1(u1_struct_0(esk1_0)))|~v2_waybel_0(X1,k3_yellow_1(u1_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),X1),esk1_0)|~v4_orders_2(k3_yellow19(esk1_0,u1_struct_0(esk1_0),X1))|~v1_subset_1(X1,u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2302, plain, (r2_waybel_0(X1,X2,X3)|v1_xboole_0(u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X1)|u1_struct_0(X2)!=k1_xboole_0|~l1_struct_0(X1)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_2413, plain, (v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X2)|v2_struct_0(X4)|u1_struct_0(X1)!=k1_xboole_0|~r2_waybel_0(X1,X2,X3)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~r1_waybel_0(X4,X2,X5)|~l1_struct_0(X4)|~l1_waybel_0(X2,X4))).#trainneg
cnf(i_0_1232, plain, (k4_xboole_0(X1,k3_subset_1(X2,X3))=k1_xboole_0|r2_hidden(esk15_3(k1_xboole_0,X4,k4_xboole_0(X1,k3_subset_1(X2,X3))),X3)|~r2_hidden(esk15_3(k1_xboole_0,X4,k4_xboole_0(X1,k3_subset_1(X2,X3))),X2)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1930, plain, (X1=X2|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_1542, plain, (v1_relat_1(k2_struct_0(g1_orders_2(X1,k1_xboole_0)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_2161, plain, (esk16_1(k1_zfmisc_1(X1))=k1_xboole_0|X1!=k1_xboole_0)).#trainneg
cnf(i_0_341, plain, (v1_xboole_0(u1_waybel_0(X1,X2))|~v1_xboole_0(u1_struct_0(X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_342, plain, (v4_relat_1(u1_waybel_0(X1,X2),u1_struct_0(X2))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_343, plain, (u1_struct_0(X1)=X2|g1_orders_2(u1_struct_0(X1),u1_waybel_0(X1,X1))!=g1_orders_2(X2,X3)|~l1_struct_0(X1)|~l1_waybel_0(X1,X1))).#trainneg
cnf(i_0_1341, plain, (m1_subset_1(esk15_3(k1_xboole_0,X1,k4_xboole_0(X2,X3)),X2)|k4_xboole_0(X2,X3)=k1_xboole_0)).#trainneg
cnf(i_0_982, plain, (v1_xboole_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|X1!=k1_xboole_0|~r2_waybel_0(X3,X2,X1)|~l1_struct_0(X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_waybel_0(X2,X3))).#trainneg
cnf(i_0_2521, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~v1_xboole_0(u1_struct_0(X2)))).#trainneg
cnf(i_0_1543, plain, (v1_finset_1(k2_struct_0(g1_orders_2(X1,k1_xboole_0)))|~v1_finset_1(X1))).#trainneg
cnf(i_0_344, plain, (k4_xboole_0(X4,X5)=X4|r2_hidden(esk15_3(X4,X5,X4),X4))).#trainneg
cnf(i_0_2439, plain, (X1=u1_waybel_0(X2,X3)|~v1_xboole_0(X1)|~l1_struct_0(X2)|~v1_xboole_0(u1_struct_0(X2))|~l1_waybel_0(X3,X2))).#trainneg
cnf(i_0_2510, plain, (X1=u1_waybel_0(X2,X3)|~v1_xboole_0(X1)|~l1_struct_0(X2)|~v1_xboole_0(u1_struct_0(X3))|~l1_waybel_0(X3,X2))).#trainneg
cnf(i_0_347, plain, (m1_subset_1(esk15_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)|r2_hidden(esk15_3(X1,X2,X3),X1))).#trainneg
cnf(i_0_2579, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_xboole_0(k2_struct_0(esk1_0))|v1_compts_1(esk1_0)|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~v3_waybel_7(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(k2_struct_0(esk1_0)))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2463, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v1_xboole_0(u1_struct_0(X1))|~r1_waybel_0(X2,X1,k1_xboole_0)|~l1_struct_0(X2)|~l1_waybel_0(X1,X2)|u1_struct_0(X1)!=k1_xboole_0)).#trainneg
cnf(i_0_2465, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v1_xboole_0(u1_struct_0(X1))|~r1_waybel_0(X2,X1,X4)|~l1_struct_0(X2)|~l1_waybel_0(X1,X2)|u1_struct_0(X1)!=k1_xboole_0)).#trainneg
cnf(i_0_2621, plain, (X1=k1_xboole_0|r2_hidden(esk15_3(X1,X1,X1),X1))).#trainneg
cnf(i_0_2460, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X2))|~l1_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_yellow19(X2,X1)|~l1_waybel_0(X2,X1)|u1_struct_0(X2)!=k1_xboole_0)).#trainneg
cnf(i_0_2457, negated_conjecture, (v1_xboole_0(k2_struct_0(esk1_0))|v1_xboole_0(X1)|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1))|u1_struct_0(esk1_0)!=k1_xboole_0|~v1_subset_1(X1,u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v3_waybel_7(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~v13_waybel_0(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1),esk1_0))).#trainneg
cnf(i_0_1640, plain, (k2_struct_0(g1_orders_2(X2,k1_xboole_0))=k1_xboole_0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_349, plain, (m1_subset_1(esk15_3(X1,X2,X3),X1)|X3=k4_xboole_0(X1,X2)|r2_hidden(esk15_3(X1,X2,X3),X3))).#trainneg
cnf(i_0_2732, negated_conjecture, (v2_struct_0(esk2_0)|r1_waybel_0(esk1_0,esk2_0,k1_xboole_0)|u1_struct_0(esk2_0)!=k1_xboole_0|~l1_waybel_0(esk2_0,esk1_0)|~v1_compts_1(esk1_0))).#trainneg
cnf(i_0_2617, plain, (m1_subset_1(esk15_3(X1,X2,X1),X1)|k4_xboole_0(X1,X2)=X1)).#trainneg
cnf(i_0_1438, plain, (u1_struct_0(X1)=esk16_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xboole_0(X3)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_350, plain, (X1=k4_xboole_0(X2,X3)|r2_hidden(esk15_3(X2,X3,X1),X1)|v1_xboole_0(X3)|~m1_subset_1(esk15_3(X2,X3,X1),X3))).#trainneg
cnf(i_0_351, plain, (X1=k4_xboole_0(X2,X2)|r2_hidden(esk15_3(X2,X2,X1),X1))).#trainneg
cnf(i_0_2769, plain, (v2_struct_0(X1)|v2_struct_0(X2)|u1_struct_0(X2)!=k1_xboole_0|~l1_struct_0(X1)|~v1_yellow19(X2,X1)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_2785, negated_conjecture, (v1_xboole_0(k2_struct_0(esk1_0))|v1_xboole_0(X1)|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1))|u1_struct_0(esk1_0)!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v3_waybel_7(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~v13_waybel_0(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~v2_waybel_0(X1,k3_yellow_1(k2_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),X1),esk1_0))).#trainneg
cnf(i_0_2616, plain, (k4_xboole_0(X1,X2)=X1|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2568, plain, (v1_xboole_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|X1!=k1_xboole_0|~r2_waybel_0(X2,X3,X1)|~l1_struct_0(X2)|~l1_waybel_0(X3,X2)|~l1_waybel_0(X4,X3)|~v7_waybel_0(X4)|~v4_orders_2(X4)|~v1_compts_1(X3)|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_2572, plain, (v1_xboole_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|r2_waybel_0(X4,X3,X5)|v2_struct_0(X4)|X1!=k1_xboole_0|~r2_waybel_0(X2,X3,X1)|~l1_struct_0(X2)|~l1_waybel_0(X3,X2)|~l1_struct_0(X4)|~l1_waybel_0(X3,X4))).#trainneg
cnf(i_0_785, plain, (r1_waybel_0(X3,X2,X4)|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|~r2_hidden(k2_waybel_0(X1,X2,esk10_4(X3,X2,X4,esk9_3(X1,X2,k4_xboole_0(X5,X6)))),X6)|~r1_waybel_0(X1,X2,k4_xboole_0(X5,X6))|~l1_struct_0(X1)|~l1_struct_0(X3)|~m1_subset_1(esk10_4(X3,X2,X4,esk9_3(X1,X2,k4_xboole_0(X5,X6))),u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_waybel_0(X2,X3))).#trainneg
cnf(i_0_1642, plain, (X3=k4_xboole_0(k2_struct_0(g1_orders_2(X1,k1_xboole_0)),X2)|~v1_xboole_0(X1)|X3!=k1_xboole_0)).#trainneg
cnf(i_0_1562, plain, (v1_relat_1(k3_subset_1(k1_xboole_0,X3))|~m1_subset_1(X3,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2904, negated_conjecture, (v2_struct_0(esk2_0)|u1_struct_0(esk2_0)!=k1_xboole_0|~l1_waybel_0(esk2_0,esk1_0)|~v1_compts_1(esk1_0))).#trainneg
cnf(i_0_2234, plain, (u1_struct_0(X1)=esk16_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xboole_0(X2)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_2249, plain, (u1_struct_0(X1)=u1_orders_2(X2)|~l1_orders_2(X2)|~v1_xboole_0(u1_struct_0(X2))|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_355, plain, (v1_xboole_0(esk4_1(X1))|v1_compts_1(X1)|v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_356, plain, (v1_relat_1(esk4_1(X1))|v1_compts_1(X1)|v2_struct_0(X1)|~v1_relat_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_2963, plain, (r2_hidden(esk15_3(X1,X1,esk4_1(X2)),u1_struct_0(k3_yellow_1(u1_struct_0(X2))))|v1_compts_1(X2)|v2_struct_0(X2)|esk4_1(X2)=k1_xboole_0|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_357, plain, (v1_finset_1(esk4_1(X1))|v1_compts_1(X1)|v2_struct_0(X1)|~v1_finset_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1563, plain, (v1_xboole_0(k3_subset_1(k1_xboole_0,X3))|~m1_subset_1(X3,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1664, plain, (u1_struct_0(X1)=k3_subset_1(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v1_xboole_0(X2)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_2634, plain, (u1_struct_0(X1)=u1_waybel_0(X2,X3)|~l1_struct_0(X2)|~v1_xboole_0(u1_struct_0(X2))|~l1_waybel_0(X3,X2)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_360, plain, (v13_waybel_0(X2,X1)|v2_struct_0(X1)|~v12_waybel_0(k3_subset_1(u1_struct_0(X1),X2),X1)|~l1_orders_2(X1)|~m1_subset_1(k3_subset_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_361, plain, (v12_waybel_0(X2,X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~m1_subset_1(k3_subset_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~v13_waybel_0(k3_subset_1(u1_struct_0(X1),X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_2962, plain, (r2_hidden(esk15_3(X1,X1,esk4_1(X2)),u1_struct_0(k3_yellow_1(k2_struct_0(X2))))|v1_compts_1(X2)|v2_struct_0(X2)|esk4_1(X2)=k1_xboole_0|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_2968, plain, (r2_hidden(esk15_3(X1,X1,k4_xboole_0(X2,X3)),X2)|k4_xboole_0(X2,X3)=k1_xboole_0)).#trainneg
cnf(i_0_2960, plain, (m1_subset_1(esk15_3(X1,X1,esk4_1(X2)),u1_struct_0(k3_yellow_1(u1_struct_0(X2))))|v1_compts_1(X2)|v2_struct_0(X2)|esk4_1(X2)=k1_xboole_0|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_2967, plain, (k4_xboole_0(X2,X3)=k1_xboole_0|~r2_hidden(esk15_3(X1,X1,k4_xboole_0(X2,X3)),X3))).#trainneg
cnf(i_0_2981, plain, (r2_hidden(esk15_3(X1,X1,k10_yellow_6(X2,X3)),u1_struct_0(X2))|v2_struct_0(X3)|v2_struct_0(X2)|k10_yellow_6(X2,X3)=k1_xboole_0|~l1_waybel_0(X3,X2)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_2977, plain, (r2_waybel_7(X1,X2,esk15_3(X3,X3,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2))))|v1_xboole_0(X2)|v2_struct_0(X1)|k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2))=k1_xboole_0|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))|~m1_subset_1(esk15_3(X3,X3,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2))),u1_struct_0(X1))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_2961, plain, (m1_subset_1(esk15_3(X1,X1,esk4_1(X2)),u1_struct_0(k3_yellow_1(k2_struct_0(X2))))|v1_compts_1(X2)|v2_struct_0(X2)|esk4_1(X2)=k1_xboole_0|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_3194, plain, (v1_compts_1(X1)|v2_struct_0(X1)|k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)))=k1_xboole_0|v1_xboole_0(esk4_1(X1))|~m1_subset_1(esk15_3(X2,X2,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)))),u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))|~v13_waybel_0(esk4_1(X1),k3_yellow_1(k2_struct_0(X1)))|~v2_waybel_0(esk4_1(X1),k3_yellow_1(k2_struct_0(X1))))).#trainneg
cnf(i_0_3071, plain, (v1_compts_1(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1))),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_362, plain, (v4_pre_topc(X2,X1)|~v3_pre_topc(k3_subset_1(u1_struct_0(X1),X2),X1)|~m1_subset_1(k3_subset_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_2652, plain, (u1_struct_0(X1)=u1_waybel_0(X2,X3)|~l1_struct_0(X2)|~v1_xboole_0(u1_struct_0(X3))|~l1_waybel_0(X3,X2)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_363, plain, (v3_pre_topc(X2,X1)|~v4_pre_topc(k3_subset_1(u1_struct_0(X1),X2),X1)|~m1_subset_1(k3_subset_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_364, plain, (v1_tops_1(X2,X1)|~v3_tops_1(k3_subset_1(u1_struct_0(X1),X2),X1)|~m1_subset_1(k3_subset_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_2976, plain, (r2_hidden(esk15_3(X1,X1,k2_struct_0(X2)),u1_struct_0(X2))|k2_struct_0(X2)=k1_xboole_0|~l1_struct_0(X2))).#trainneg
cnf(i_0_2972, plain, (r2_hidden(esk15_3(X1,X1,k3_subset_1(X2,X3)),X2)|k3_subset_1(X2,X3)=k1_xboole_0|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_2993, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(X1)|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),X1))|u1_struct_0(esk1_0)!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))|~v3_waybel_7(X1,k3_yellow_1(u1_struct_0(esk1_0)))|~v13_waybel_0(X1,k3_yellow_1(u1_struct_0(esk1_0)))|~v2_waybel_0(X1,k3_yellow_1(u1_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),X1),esk1_0))).#trainneg
cnf(i_0_2970, plain, (k3_subset_1(X2,X3)=k1_xboole_0|~r2_hidden(esk15_3(X1,X1,k3_subset_1(X2,X3)),X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_3017, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X2,X1,X3)|~r1_waybel_0(X2,X1,X4)|~l1_struct_0(X2)|~l1_waybel_0(X1,X2)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_3069, plain, (v1_compts_1(X1)|v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_1450, plain, (l1_struct_0(g1_orders_2(X1,esk16_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))).#trainneg
cnf(i_0_2929, plain, (esk15_3(k1_zfmisc_1(k1_xboole_0),X1,k1_zfmisc_1(k1_xboole_0))=k1_xboole_0|k4_xboole_0(k1_zfmisc_1(k1_xboole_0),X1)=k1_zfmisc_1(k1_xboole_0))).#trainneg
cnf(i_0_366, plain, (v1_xboole_0(k10_yellow_6(X1,X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v1_xboole_0(u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3315, negated_conjecture, (~v2_struct_0(esk2_0)|~v1_xboole_0(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_3319, negated_conjecture, (u1_struct_0(esk2_0)!=k1_xboole_0|~v1_xboole_0(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_3347, plain, (X1=k10_yellow_6(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|~v1_xboole_0(X1)|~v1_xboole_0(u1_struct_0(X2))|~l1_waybel_0(X3,X2)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_367, plain, (v1_relat_1(k10_yellow_6(X1,X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v1_relat_1(u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_368, plain, (v1_finset_1(k10_yellow_6(X1,X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v1_finset_1(u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3211, plain, (k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))=k1_xboole_0|v1_compts_1(X1)|v2_struct_0(X1)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~m1_subset_1(esk15_3(X2,X2,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))),u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_3299, plain, (X2=k1_xboole_0|r2_hidden(esk15_3(X3,X3,X2),X1)|~m1_subset_1(k3_subset_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3302, plain, (X2=k1_xboole_0|~r2_hidden(esk15_3(X3,X3,X2),k3_subset_1(X1,X2))|~m1_subset_1(k3_subset_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3257, plain, (m1_subset_1(esk15_3(X1,X1,k2_struct_0(X2)),u1_struct_0(X2))|k2_struct_0(X2)=k1_xboole_0|~l1_struct_0(X2))).#trainneg
cnf(i_0_3197, plain, (k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2))=k1_xboole_0|r2_waybel_7(X1,X2,esk15_3(X3,X3,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2))))|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~m1_subset_1(esk15_3(X3,X3,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2))),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X1)))))|~v13_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~v2_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_3216, plain, (v1_compts_1(X1)|v2_struct_0(X1)|~m1_subset_1(u1_struct_0(k3_yellow_1(u1_struct_0(X1))),k1_zfmisc_1(k1_xboole_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_3316, negated_conjecture, (v7_waybel_0(esk2_0)|~v1_xboole_0(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1571, plain, (X1=X2|g1_orders_2(X1,X3)!=g1_orders_2(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_370, plain, (m1_subset_1(esk18_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3384, plain, (u1_struct_0(X1)=k10_yellow_6(X2,X3)|v2_struct_0(X2)|v2_struct_0(X3)|~v1_xboole_0(u1_struct_0(X2))|~l1_waybel_0(X3,X2)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_575, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|k10_yellow_6(X1,X2)!=u1_struct_0(X1)|~r2_hidden(X4,u1_struct_0(X1))|~m1_connsp_2(X3,X1,X4)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_372, plain, (v1_xboole_0(X1)|~v1_funct_2(X2,X1,u1_struct_0(X3))|~v1_funct_1(X2)|~l1_struct_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X3))))|~v2_struct_0(g1_waybel_0(X3,X1,esk16_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))),X2)))).#trainneg
cnf(i_0_373, plain, (v1_xboole_0(u1_struct_0(X1))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(X3))|~v1_funct_1(X2)|~l1_struct_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))|~v2_struct_0(g1_waybel_0(X3,u1_struct_0(X1),u1_orders_2(X1),X2))|~l1_orders_2(X1))).#trainneg
cnf(i_0_3402, negated_conjecture, (v7_waybel_0(esk2_0)|v1_xboole_0(k2_struct_0(esk1_0))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3403, negated_conjecture, (v4_orders_2(esk2_0)|v1_xboole_0(k2_struct_0(esk1_0))|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3401, negated_conjecture, (v1_xboole_0(k2_struct_0(esk1_0))|~v2_struct_0(esk2_0)|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2582, plain, (v2_struct_0(esk19_1(X1))|v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(esk19_1(X1))))).#trainneg
cnf(i_0_3493, negated_conjecture, (v7_waybel_0(esk2_0)|~m1_subset_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0))),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_374, plain, (v1_xboole_0(u1_struct_0(X1))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(X3))|~v1_funct_1(X2)|~l1_struct_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))|~v2_struct_0(g1_waybel_0(X3,u1_struct_0(X1),u1_waybel_0(X1,X1),X2))|~l1_struct_0(X1)|~l1_waybel_0(X1,X1))).#trainneg
cnf(i_0_3367, negated_conjecture, (~v2_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0))),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_426, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|k10_yellow_6(X1,X2)!=k2_struct_0(X1)|~r2_hidden(X4,k2_struct_0(X1))|~m1_connsp_2(X3,X1,X4)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_375, plain, (v1_xboole_0(X1)|~v1_funct_2(X2,X1,u1_struct_0(X3))|~v1_funct_1(X2)|~l1_struct_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X3))))|~v2_struct_0(g1_waybel_0(X3,X1,k3_subset_1(k2_zfmisc_1(X1,X1),X4),X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_376, plain, (v1_xboole_0(X1)|v2_struct_0(X2)|~l1_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X2)))))|~v3_waybel_7(X1,k3_yellow_1(u1_struct_0(X2)))|~v13_waybel_0(X1,k3_yellow_1(u1_struct_0(X2)))|~v2_waybel_0(X1,k3_yellow_1(u1_struct_0(X2)))|~v2_struct_0(k3_yellow19(X2,u1_struct_0(X2),X1)))).#trainneg
cnf(i_0_3451, plain, (r2_waybel_7(X1,X2,esk15_3(X3,X3,X4))|v1_xboole_0(X2)|v2_struct_0(X1)|X4=k1_xboole_0|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X1)))))|~m1_subset_1(esk15_3(X3,X3,X4),u1_struct_0(X1))|~v13_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~v2_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2)))))).#trainneg
cnf(i_0_3436, plain, (X2=k1_xboole_0|~r2_hidden(esk15_3(X1,X1,X2),X4)|~m1_subset_1(X2,k1_zfmisc_1(k4_xboole_0(X3,X4))))).#trainneg
cnf(i_0_3461, plain, (k3_subset_1(X1,X2)=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(k3_subset_1(X1,X2),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_3437, plain, (r2_hidden(esk15_3(X1,X1,X2),X3)|X2=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(k4_xboole_0(X3,X4))))).#trainneg
cnf(i_0_3456, plain, (m1_subset_1(esk15_3(X1,X1,X2),X3)|X2=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_3531, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0)|v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1925, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(u1_waybel_0(X1,X2),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_3529, negated_conjecture, (v7_waybel_0(esk2_0)|v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1584, plain, (X1=X2|g1_orders_2(X3,X1)!=g1_orders_2(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_377, plain, (r2_waybel_0(X1,X2,esk6_3(X1,X3,X4))|v2_struct_0(X1)|v2_struct_0(X2)|r3_waybel_9(X1,X3,X4)|v2_struct_0(X3)|~r3_waybel_9(X1,X2,X4)|~m1_subset_1(X4,u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X3,X1))).#trainneg
cnf(i_0_3369, negated_conjecture, (u1_struct_0(esk2_0)!=k1_xboole_0|~m1_subset_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0))),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_378, plain, (m1_subset_1(esk6_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|r3_waybel_9(X1,X2,X3)|v2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_379, plain, (X1=k4_xboole_0(X2,X3)|r2_hidden(esk15_3(X2,X3,X1),X3)|v1_xboole_0(X1)|~r2_hidden(esk15_3(X2,X3,X1),X2)|~m1_subset_1(esk15_3(X2,X3,X1),X1))).#trainneg
cnf(i_0_3530, negated_conjecture, (v4_orders_2(esk2_0)|v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3541, negated_conjecture, (v1_xboole_0(k2_struct_0(esk1_0))|~v2_struct_0(esk2_0))).#trainneg
cnf(i_0_3538, negated_conjecture, (v1_xboole_0(k2_struct_0(esk1_0))|v4_orders_2(esk2_0))).#trainneg
cnf(i_0_3642, plain, (k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)))=k1_xboole_0|v1_compts_1(X1)|v2_struct_0(X1)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1))),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_2267, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))).#trainneg
cnf(i_0_1527, plain, (v5_relat_1(k2_struct_0(g1_orders_2(k2_zfmisc_1(X1,X2),k1_xboole_0)),X2))).#trainneg
cnf(i_0_3532, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|u1_struct_0(esk2_0)!=k1_xboole_0)).#trainneg
cnf(i_0_381, plain, (k3_funct_2(X1,X2,X3,X4)=k3_funct_2(X5,X6,X3,X4)|v1_xboole_0(X1)|v1_xboole_0(X5)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X5,X6)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X4,X5))).#trainneg
cnf(i_0_3317, negated_conjecture, (v4_orders_2(esk2_0)|~v1_xboole_0(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_3494, negated_conjecture, (v4_orders_2(esk2_0)|~m1_subset_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0))),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_382, plain, (v1_orders_2(g1_orders_2(X1,k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)),X3,X4)))|v1_xboole_0(X2)|~v1_funct_2(X3,X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_383, plain, (l1_orders_2(g1_orders_2(X1,k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)),X3,X4)))|v1_xboole_0(X2)|~v1_funct_2(X3,X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_3656, plain, (r1_waybel_0(X1,X2,esk6_3(X1,X3,X4))|v2_struct_0(X2)|v2_struct_0(X1)|r3_waybel_9(X1,X3,X4)|v2_struct_0(X3)|~l1_struct_0(X1)|~m1_subset_1(esk6_3(X1,X3,X4),k1_zfmisc_1(u1_struct_0(X1)))|~v1_yellow19(X2,X1)|~l1_waybel_0(X2,X1)|~r3_waybel_9(X1,X2,X4)|~m1_subset_1(X4,u1_struct_0(X1))|~l1_waybel_0(X3,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3657, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r3_waybel_9(X2,X3,X4)|v2_struct_0(X3)|esk6_3(X2,X3,X4)!=k1_xboole_0|~r1_waybel_0(X2,X1,X5)|~l1_struct_0(X2)|~l1_waybel_0(X1,X2)|~r3_waybel_9(X2,X1,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_waybel_0(X3,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_3745, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~l1_struct_0(X2)|~l1_waybel_0(X1,X2)|~m1_subset_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_1531, plain, (v4_relat_1(k2_struct_0(g1_orders_2(k2_zfmisc_1(X1,X2),k1_xboole_0)),X1))).#trainneg
cnf(i_0_3770, negated_conjecture, (u1_struct_0(esk2_0)!=k1_xboole_0)).#trainneg
cnf(i_0_384, plain, (k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)),X3,X4)=X5|v1_xboole_0(X1)|g1_orders_2(X2,k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)),X3,X4))!=g1_orders_2(X6,X5)|~v1_funct_2(X3,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))))|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_1983, plain, (v1_finset_1(u1_orders_2(X1))|~l1_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_385, plain, (v5_relat_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5),X3)|v1_xboole_0(X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1))).#trainneg
cnf(i_0_386, plain, (v1_relat_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|v1_xboole_0(X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1))).#trainneg
cnf(i_0_3643, plain, (k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))=k1_xboole_0|v1_compts_1(X1)|v2_struct_0(X1)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_3586, plain, (X2=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(k3_subset_1(X1,X2)))|~m1_subset_1(k3_subset_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3661, plain, (v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X2)|v2_struct_0(X3)|r3_waybel_9(X1,X4,X5)|v2_struct_0(X4)|v2_struct_0(X1)|u1_struct_0(X1)!=k1_xboole_0|~r1_waybel_0(X2,X3,X6)|~l1_struct_0(X1)|~l1_struct_0(X2)|~l1_waybel_0(X3,X1)|~l1_waybel_0(X3,X2)|~r3_waybel_9(X1,X3,X5)|~m1_subset_1(X5,u1_struct_0(X1))|~l1_waybel_0(X4,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3660, plain, (r2_waybel_0(X1,X2,X3)|v1_xboole_0(esk6_3(X4,X5,X6))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X4)|r3_waybel_9(X4,X5,X6)|v2_struct_0(X5)|esk6_3(X4,X5,X6)!=k1_xboole_0|~l1_struct_0(X4)|~l1_struct_0(X1)|~l1_waybel_0(X2,X4)|~l1_waybel_0(X2,X1)|~r3_waybel_9(X4,X2,X6)|~m1_subset_1(X6,u1_struct_0(X4))|~l1_waybel_0(X5,X4)|~l1_pre_topc(X4)|~v2_pre_topc(X4))).#trainneg
cnf(i_0_3655, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r3_waybel_9(X1,X3,X4)|v2_struct_0(X3)|~r1_waybel_0(X1,X2,X5)|~l1_struct_0(X1)|~v1_xboole_0(esk6_3(X1,X3,X4))|~l1_waybel_0(X2,X1)|~r3_waybel_9(X1,X2,X4)|~m1_subset_1(X4,u1_struct_0(X1))|~l1_waybel_0(X3,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_2118, plain, (v2_struct_0(g1_orders_2(X1,k1_xboole_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_387, plain, (v1_xboole_0(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|v1_xboole_0(X1)|~v1_xboole_0(X3)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1))).#trainneg
cnf(i_0_388, plain, (v1_xboole_0(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1))).#trainneg
cnf(i_0_389, plain, (v4_relat_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5),X2)|v1_xboole_0(X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1))).#trainneg
cnf(i_0_3429, plain, (m1_subset_1(esk15_3(X1,X1,X2),u1_struct_0(k3_yellow_1(u1_struct_0(X3))))|v1_compts_1(X3)|v2_struct_0(X3)|X2=k1_xboole_0|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(X2,k1_zfmisc_1(esk4_1(X3))))).#trainneg
cnf(i_0_3441, plain, (r2_hidden(esk15_3(X1,X1,X2),X3)|X2=k1_xboole_0|~m1_subset_1(X4,k1_zfmisc_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(k3_subset_1(X3,X4))))).#trainneg
cnf(i_0_3450, plain, (r2_waybel_7(X1,X2,esk15_3(X3,X3,X4))|v1_xboole_0(X2)|v2_struct_0(X1)|X4=k1_xboole_0|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))|~m1_subset_1(esk15_3(X3,X3,X4),u1_struct_0(X1))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2)))))).#trainneg
cnf(i_0_3795, negated_conjecture, (v2_struct_0(esk2_0)|~v1_xboole_0(u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_3794, negated_conjecture, (l1_orders_2(esk2_0))).#trainneg
cnf(i_0_3346, plain, (k1_zfmisc_1(k1_xboole_0)=k1_xboole_0|r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_390, plain, (X1=X2|v1_xboole_0(X3)|g1_orders_2(X1,k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)),X4,X5))!=g1_orders_2(X2,X6)|~v1_funct_2(X4,X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))|~m1_subset_1(X5,X3))).#trainneg
cnf(i_0_3475, plain, (k1_zfmisc_1(k1_xboole_0)=k1_xboole_0|~r2_hidden(k1_xboole_0,k3_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))|~m1_subset_1(k1_zfmisc_1(k1_xboole_0),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_391, plain, (v1_xboole_0(X1)|v1_xboole_0(X4)|~v1_funct_2(X2,X1,u1_struct_0(X3))|~v1_funct_1(X2)|~l1_struct_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X3))))|~v2_struct_0(g1_waybel_0(X3,X1,k3_funct_2(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X1)),X5,X6),X2))|~v1_funct_2(X5,X4,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))|~m1_subset_1(X6,X4))).#trainneg
cnf(i_0_393, plain, (v1_xboole_0(k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_3833, plain, (esk15_3(X1,X1,X2)=k1_xboole_0|X2=k1_xboole_0|~m1_subset_1(esk15_3(X1,X1,X2),k1_zfmisc_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(k3_subset_1(X3,esk15_3(X1,X1,X2))))))).#trainneg
cnf(i_0_3841, plain, (r3_waybel_9(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X4)|v2_struct_0(X1)|~r3_waybel_9(X1,X4,X3)|~r1_waybel_0(X1,X4,X5)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X4,X1)|~l1_waybel_0(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk6_3(X1,X2,X3),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_3142, plain, (k4_xboole_0(k4_xboole_0(X2,X3),X4)=k1_xboole_0|~r2_hidden(esk15_3(X1,X1,k4_xboole_0(k4_xboole_0(X2,X3),X4)),X3))).#trainneg
cnf(i_0_3430, plain, (m1_subset_1(esk15_3(X1,X1,X2),u1_struct_0(k3_yellow_1(k2_struct_0(X3))))|v1_compts_1(X3)|v2_struct_0(X3)|X2=k1_xboole_0|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(X2,k1_zfmisc_1(esk4_1(X3))))).#trainneg
cnf(i_0_3440, plain, (m1_subset_1(esk15_3(X1,X1,X2),X3)|X2=k1_xboole_0|~m1_subset_1(X4,k1_zfmisc_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(k3_subset_1(X3,X4))))).#trainneg
cnf(i_0_1110, plain, (v1_tops_1(k1_xboole_0,X1)|~v3_tops_1(u1_struct_0(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_2053, plain, (l1_struct_0(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_3338, plain, (k4_xboole_0(k1_zfmisc_1(k1_xboole_0),X1)=k1_zfmisc_1(k1_xboole_0)|r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_394, plain, (v1_relat_1(k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|v1_xboole_0(X1)|~v1_relat_1(X2)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_395, plain, (v1_finset_1(k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|v1_xboole_0(X1)|~v1_finset_1(X2)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_396, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X3)|~r2_hidden(X1,k3_funct_2(X3,k1_zfmisc_1(X2),X4,X5))|~v1_funct_2(X4,X3,k1_zfmisc_1(X2))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,k1_zfmisc_1(X2))))|~m1_subset_1(X5,X3))).#trainneg
cnf(i_0_3939, negated_conjecture, (l1_struct_0(esk2_0))).#trainneg
cnf(i_0_3664, plain, (v1_xboole_0(esk6_3(X1,X2,X3))|r3_waybel_9(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4005, plain, (k3_subset_1(X1,u1_struct_0(k3_yellow_1(k2_struct_0(X2))))=k1_xboole_0|v1_xboole_0(u1_struct_0(k3_yellow_1(k2_struct_0(X2))))|v1_compts_1(X2)|v2_struct_0(X2)|~m1_subset_1(u1_struct_0(k3_yellow_1(k2_struct_0(X2))),k1_zfmisc_1(X1))|~m1_subset_1(k3_subset_1(X1,u1_struct_0(k3_yellow_1(k2_struct_0(X2)))),k1_zfmisc_1(esk4_1(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_1160, plain, (v1_tops_1(X1,g1_orders_2(X1,k1_xboole_0))|~l1_pre_topc(g1_orders_2(X1,k1_xboole_0)))).#trainneg
cnf(i_0_1901, plain, (X1=X2|g1_orders_2(X1,k3_subset_1(k1_xboole_0,X3))!=g1_orders_2(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_397, plain, (m1_subset_1(k1_funct_1(X3,X4),X2)|v1_xboole_0(X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_3937, negated_conjecture, (~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_682, plain, (m1_subset_1(X1,k3_subset_1(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_398, plain, (k1_funct_1(u1_waybel_0(X2,X1),X3)=k2_waybel_0(X2,X1,X3)|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_funct_2(u1_waybel_0(X2,X1),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(u1_waybel_0(X2,X1))|~m1_subset_1(u1_waybel_0(X2,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_struct_0(X2)|~l1_waybel_0(X1,X2))).#trainneg
cnf(i_0_3445, plain, (r2_hidden(esk15_3(X1,X1,X2),u1_struct_0(X3))|X2=k1_xboole_0|~l1_struct_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_struct_0(X3))))).#trainneg
cnf(i_0_3987, plain, (k4_xboole_0(k3_subset_1(X1,X2),X3)=k1_xboole_0|~r2_hidden(esk15_3(X4,X4,k4_xboole_0(k3_subset_1(X1,X2),X3)),X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2980, plain, (m1_subset_1(esk15_3(X1,X1,k10_yellow_6(X2,X3)),u1_struct_0(X2))|v2_struct_0(X3)|v2_struct_0(X2)|k10_yellow_6(X2,X3)=k1_xboole_0|~l1_waybel_0(X3,X2)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_3894, plain, (k3_subset_1(X1,u1_struct_0(k3_yellow_1(u1_struct_0(X2))))=k1_xboole_0|v1_xboole_0(u1_struct_0(k3_yellow_1(u1_struct_0(X2))))|v1_compts_1(X2)|v2_struct_0(X2)|~m1_subset_1(u1_struct_0(k3_yellow_1(u1_struct_0(X2))),k1_zfmisc_1(X1))|~m1_subset_1(k3_subset_1(X1,u1_struct_0(k3_yellow_1(u1_struct_0(X2)))),k1_zfmisc_1(esk4_1(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_1522, plain, (v1_orders_2(g1_orders_2(X1,k2_struct_0(g1_orders_2(k2_zfmisc_1(X1,X1),k1_xboole_0)))))).#trainneg
cnf(i_0_403, plain, (X1=X2|g1_waybel_0(X3,X4,esk16_1(k1_zfmisc_1(k2_zfmisc_1(X4,X4))),X1)!=g1_waybel_0(X5,X6,X7,X2)|~v1_funct_2(X1,X4,u1_struct_0(X3))|~v1_funct_1(X1)|~l1_struct_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,u1_struct_0(X3)))))).#trainneg
cnf(i_0_2060, plain, (u1_orders_2(X1)=k1_xboole_0|~l1_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_404, plain, (X1=X2|v1_xboole_0(X5)|g1_waybel_0(X3,X4,k3_funct_2(X5,k1_zfmisc_1(k2_zfmisc_1(X4,X4)),X6,X7),X1)!=g1_waybel_0(X8,X9,X10,X2)|~v1_funct_2(X1,X4,u1_struct_0(X3))|~v1_funct_1(X1)|~l1_struct_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,u1_struct_0(X3))))|~v1_funct_2(X6,X5,k1_zfmisc_1(k2_zfmisc_1(X4,X4)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,X4)))))|~m1_subset_1(X7,X5))).#trainneg
cnf(i_0_405, plain, (X1=X2|g1_waybel_0(X3,u1_struct_0(X4),u1_orders_2(X4),X1)!=g1_waybel_0(X5,X6,X7,X2)|~v1_funct_2(X1,u1_struct_0(X4),u1_struct_0(X3))|~v1_funct_1(X1)|~l1_struct_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X3))))|~l1_orders_2(X4))).#trainneg
cnf(i_0_4140, plain, (k3_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(k2_struct_0(X1))))=k1_xboole_0|v1_compts_1(X1)|v2_struct_0(X1)|~m1_subset_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1))),k1_zfmisc_1(esk4_1(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1523, plain, (l1_orders_2(g1_orders_2(X1,k2_struct_0(g1_orders_2(k2_zfmisc_1(X1,X1),k1_xboole_0)))))).#trainneg
cnf(i_0_406, plain, (X1=X2|g1_waybel_0(X3,u1_struct_0(X4),u1_waybel_0(X4,X4),X1)!=g1_waybel_0(X5,X6,X7,X2)|~v1_funct_2(X1,u1_struct_0(X4),u1_struct_0(X3))|~v1_funct_1(X1)|~l1_struct_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X3))))|~l1_struct_0(X4)|~l1_waybel_0(X4,X4))).#trainneg
cnf(i_0_4119, plain, (X1=esk6_3(X2,X3,X4)|r3_waybel_9(X2,X3,X4)|v2_struct_0(X2)|v2_struct_0(X3)|~v1_xboole_0(X1)|~m1_subset_1(X4,u1_struct_0(X2))|~v1_xboole_0(u1_struct_0(X2))|~l1_waybel_0(X3,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_407, plain, (X1=X2|g1_waybel_0(X3,X4,k3_subset_1(k2_zfmisc_1(X4,X4),X5),X1)!=g1_waybel_0(X6,X7,X8,X2)|~v1_funct_2(X1,X4,u1_struct_0(X3))|~v1_funct_1(X1)|~l1_struct_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,u1_struct_0(X3))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,X4))))).#trainneg
cnf(i_0_409, plain, (l1_waybel_0(g1_waybel_0(X1,X2,esk16_1(k1_zfmisc_1(k2_zfmisc_1(X2,X2))),X3),X1)|~v1_funct_2(X3,X2,u1_struct_0(X1))|~v1_funct_1(X3)|~l1_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X1)))))).#trainneg
cnf(i_0_4207, plain, (k3_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))=k1_xboole_0|v1_compts_1(X1)|v2_struct_0(X1)|~m1_subset_1(u1_struct_0(k3_yellow_1(u1_struct_0(X1))),k1_zfmisc_1(esk4_1(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_2979, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X1)|k10_yellow_6(X1,X4)=k1_xboole_0|k10_yellow_6(X1,X4)!=k10_yellow_6(X1,X2)|~m1_connsp_2(X3,X1,esk15_3(X5,X5,k10_yellow_6(X1,X4)))|~l1_waybel_0(X2,X1)|~l1_waybel_0(X4,X1)|~v7_waybel_0(X2)|~v7_waybel_0(X4)|~v4_orders_2(X2)|~v4_orders_2(X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4198, plain, (v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X2)|k10_yellow_6(X2,X5)=k1_xboole_0|v2_struct_0(X5)|u1_struct_0(X1)!=k1_xboole_0|~r2_waybel_0(X1,X2,X3)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_waybel_0(X5,X2)|~v7_waybel_0(X5)|~v4_orders_2(X5)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_3836, plain, (r3_waybel_9(X1,X2,X3)|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X4)|v2_struct_0(X5)|r3_waybel_9(X5,X6,X7)|v2_struct_0(X6)|u1_struct_0(X1)!=k1_xboole_0|~r3_waybel_9(X1,X4,X3)|~l1_struct_0(X5)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X4,X1)|~l1_waybel_0(X4,X5)|~l1_waybel_0(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~r3_waybel_9(X5,X4,X7)|~m1_subset_1(X7,u1_struct_0(X5))|~v1_yellow19(X4,X5)|~l1_waybel_0(X6,X5)|~l1_pre_topc(X5)|~v2_pre_topc(X5))).#trainneg
cnf(i_0_3029, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X4)|v2_struct_0(X1)|~r2_hidden(k2_waybel_0(X4,X2,esk10_4(X1,X2,X3,esk9_3(X4,X2,k3_subset_1(X5,X6)))),X6)|~r1_waybel_0(X4,X2,k3_subset_1(X5,X6))|~l1_struct_0(X4)|~l1_struct_0(X1)|~m1_subset_1(esk10_4(X1,X2,X3,esk9_3(X4,X2,k3_subset_1(X5,X6))),u1_struct_0(X2))|~l1_waybel_0(X2,X4)|~l1_waybel_0(X2,X1)|~m1_subset_1(X6,k1_zfmisc_1(X5)))).#trainneg
cnf(i_0_3001, plain, (k3_subset_1(X1,X2)=X1|~m1_subset_1(X2,k1_zfmisc_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4281, plain, (l1_struct_0(g1_orders_2(X1,k2_struct_0(g1_orders_2(k2_zfmisc_1(X1,X1),k1_xboole_0)))))).#trainneg
cnf(i_0_411, plain, (l1_waybel_0(g1_waybel_0(X1,u1_struct_0(X2),u1_orders_2(X2),X3),X1)|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~l1_orders_2(X2))).#trainneg
cnf(i_0_4065, plain, (r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0))|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_2488, plain, (X1=k3_subset_1(k1_xboole_0,X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_410, plain, (l1_waybel_0(g1_waybel_0(X1,X2,k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X2)),X4,X5),X6),X1)|v1_xboole_0(X3)|~v1_funct_2(X6,X2,u1_struct_0(X1))|~v1_funct_1(X6)|~l1_struct_0(X1)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X1))))|~v1_funct_2(X4,X3,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))))|~m1_subset_1(X5,X3))).#trainneg
cnf(i_0_412, plain, (l1_waybel_0(g1_waybel_0(X1,u1_struct_0(X2),u1_waybel_0(X2,X2),X3),X1)|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~l1_struct_0(X2)|~l1_waybel_0(X2,X2))).#trainneg
cnf(i_0_4370, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X2)|~r2_hidden(k2_waybel_0(X4,X2,esk10_4(X1,X2,X3,esk9_3(X4,X2,X6))),k3_subset_1(X5,X6))|~r1_waybel_0(X4,X2,X6)|~l1_struct_0(X4)|~l1_struct_0(X1)|~m1_subset_1(esk10_4(X1,X2,X3,esk9_3(X4,X2,X6)),u1_struct_0(X2))|~m1_subset_1(k3_subset_1(X5,X6),k1_zfmisc_1(X5))|~l1_waybel_0(X2,X4)|~l1_waybel_0(X2,X1)|~m1_subset_1(X6,k1_zfmisc_1(X5)))).#trainneg
cnf(i_0_4366, plain, (r3_waybel_9(X1,X2,X3)|r3_waybel_9(X4,X5,esk14_2(X4,X6))|v1_xboole_0(u1_struct_0(X4))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X4)|u1_struct_0(X4)!=k1_xboole_0|~r3_waybel_9(X1,X6,X3)|~m1_subset_1(esk14_2(X4,X6),u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_yellow19(X6,X1)|~l1_waybel_0(X6,X4)|~l1_waybel_0(X6,X1)|~l1_waybel_0(X5,X4)|~l1_waybel_0(X2,X1)|~l1_pre_topc(X4)|~l1_pre_topc(X1)|~v2_pre_topc(X4)|~v2_pre_topc(X1)|~v7_waybel_0(X6)|~v4_orders_2(X6)|~v1_compts_1(X4))).#trainneg
cnf(i_0_4249, plain, (esk4_1(X1)=u1_struct_0(k3_yellow_1(k2_struct_0(X1)))|v1_compts_1(X1)|v2_struct_0(X1)|~m1_subset_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1))),k1_zfmisc_1(esk4_1(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3055, plain, (v1_relat_1(k3_subset_1(k1_xboole_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_3340, plain, (k4_xboole_0(k1_zfmisc_1(k1_xboole_0),X1)=k1_zfmisc_1(k1_xboole_0)|r2_hidden(k1_xboole_0,X1)|~r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_413, plain, (l1_waybel_0(g1_waybel_0(X1,X2,k3_subset_1(k2_zfmisc_1(X2,X2),X3),X4),X1)|~v1_funct_2(X4,X2,u1_struct_0(X1))|~v1_funct_1(X4)|~l1_struct_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X2))))).#trainneg
cnf(i_0_1564, plain, (v4_relat_1(k3_subset_1(k1_xboole_0,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1823, plain, (v5_relat_1(k3_subset_1(k1_xboole_0,X3),X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_414, plain, (r2_waybel_0(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|r1_waybel_0(X4,X2,X5)|v2_struct_0(X4)|~r2_hidden(k2_waybel_0(X1,X2,esk10_4(X4,X2,X5,esk8_3(X1,X2,X3))),X3)|~l1_struct_0(X1)|~m1_subset_1(esk10_4(X4,X2,X5,esk8_3(X1,X2,X3)),u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_struct_0(X4)|~m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X2))|~l1_waybel_0(X2,X4))).#trainneg
cnf(i_0_415, plain, (r2_waybel_0(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X4)|~r2_hidden(k2_waybel_0(X1,X2,esk7_4(X4,X2,X5,esk8_3(X1,X2,X3))),X3)|~l1_struct_0(X1)|~m1_subset_1(esk7_4(X4,X2,X5,esk8_3(X1,X2,X3)),u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~r2_waybel_0(X4,X2,X5)|~l1_struct_0(X4)|~m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X2))|~l1_waybel_0(X2,X4))).#trainneg
cnf(i_0_4549, plain, (esk4_1(X1)=u1_struct_0(k3_yellow_1(u1_struct_0(X1)))|v1_compts_1(X1)|v2_struct_0(X1)|~m1_subset_1(u1_struct_0(k3_yellow_1(u1_struct_0(X1))),k1_zfmisc_1(esk4_1(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_3444, plain, (m1_subset_1(esk15_3(X1,X1,X2),u1_struct_0(X3))|X2=k1_xboole_0|~l1_struct_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_struct_0(X3))))).#trainneg
cnf(i_0_3158, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X1)|k4_xboole_0(k10_yellow_6(X1,X4),X6)=k1_xboole_0|k10_yellow_6(X1,X4)!=k10_yellow_6(X1,X2)|~m1_connsp_2(X3,X1,esk15_3(X5,X5,k4_xboole_0(k10_yellow_6(X1,X4),X6)))|~l1_waybel_0(X2,X1)|~l1_waybel_0(X4,X1)|~v7_waybel_0(X2)|~v7_waybel_0(X4)|~v4_orders_2(X2)|~v4_orders_2(X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_730, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,esk16_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_3099, plain, (v1_xboole_0(k3_subset_1(k1_xboole_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_421, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|esk16_1(k1_zfmisc_1(u1_struct_0(X1)))!=k10_yellow_6(X1,X2)|~r2_hidden(X4,esk16_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_connsp_2(X3,X1,X4)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_422, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|v1_yellow19(X4,X1)|v2_struct_0(X4)|esk5_2(X1,X4)!=k10_yellow_6(X1,X2)|~r2_hidden(X5,esk5_2(X1,X4))|~m1_connsp_2(X3,X1,X5)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1)|~l1_waybel_0(X4,X1))).#trainneg
cnf(i_0_424, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|v1_xboole_0(X4)|k3_funct_2(X4,k1_zfmisc_1(u1_struct_0(X1)),X5,X6)!=k10_yellow_6(X1,X2)|~r2_hidden(X7,k3_funct_2(X4,k1_zfmisc_1(u1_struct_0(X1)),X5,X6))|~m1_connsp_2(X3,X1,X7)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_funct_2(X5,X4,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,k1_zfmisc_1(u1_struct_0(X1)))))|~m1_subset_1(X6,X4))).#trainneg
cnf(i_0_4666, plain, (k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))=k1_xboole_0|v1_compts_1(X1)|v2_struct_0(X1)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1)|~m1_subset_1(k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))),k1_zfmisc_1(k2_struct_0(X1))))).#trainneg
cnf(i_0_4665, plain, (k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1)))=k1_xboole_0|v1_compts_1(X1)|v2_struct_0(X1)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1)|~m1_subset_1(k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1))),k1_zfmisc_1(k2_struct_0(X1))))).#trainneg
cnf(i_0_4680, plain, (k4_xboole_0(k10_yellow_6(X1,X2),X3)=k1_xboole_0|r1_waybel_0(X1,X4,esk11_4(X1,X5,X6,esk15_3(X7,X7,k4_xboole_0(k10_yellow_6(X1,X2),X3))))|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X2)|r2_hidden(esk15_3(X7,X7,k4_xboole_0(k10_yellow_6(X1,X2),X3)),X6)|v2_struct_0(X5)|k10_yellow_6(X1,X2)!=k10_yellow_6(X1,X4)|~l1_waybel_0(X4,X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X4)|~v7_waybel_0(X2)|~v4_orders_2(X4)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|X6!=k10_yellow_6(X1,X5)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk15_3(X7,X7,k4_xboole_0(k10_yellow_6(X1,X2),X3)),u1_struct_0(X1))|~l1_waybel_0(X5,X1)|~v7_waybel_0(X5)|~v4_orders_2(X5))).#trainneg
cnf(i_0_4654, plain, (r2_waybel_0(X1,X2,u1_struct_0(k3_yellow_1(k2_struct_0(X3))))|v2_struct_0(X1)|v2_struct_0(X2)|v1_compts_1(X3)|v2_struct_0(X3)|~r2_waybel_0(X1,X2,esk4_1(X3))|~l1_struct_0(X1)|~m1_subset_1(esk7_4(X1,X2,esk4_1(X3),esk8_3(X1,X2,u1_struct_0(k3_yellow_1(k2_struct_0(X3))))),u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~m1_subset_1(esk8_3(X1,X2,u1_struct_0(k3_yellow_1(k2_struct_0(X3)))),u1_struct_0(X2))|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_4655, plain, (r2_waybel_0(X1,X2,u1_struct_0(X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~r2_waybel_0(X1,X2,k10_yellow_6(X3,X4))|~l1_struct_0(X1)|~m1_subset_1(esk7_4(X1,X2,k10_yellow_6(X3,X4),esk8_3(X1,X2,u1_struct_0(X3))),u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~m1_subset_1(esk8_3(X1,X2,u1_struct_0(X3)),u1_struct_0(X2))|~l1_waybel_0(X4,X3)|~v7_waybel_0(X4)|~v4_orders_2(X4)|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_1529, plain, (v1_xboole_0(k2_struct_0(g1_orders_2(k2_zfmisc_1(X1,X2),k1_xboole_0)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1577, plain, (k3_subset_1(k1_xboole_0,X2)=X3|g1_orders_2(X1,k3_subset_1(k1_xboole_0,X2))!=g1_orders_2(X4,X3)|~m1_subset_1(X2,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_430, plain, (v6_waybel_0(g1_waybel_0(X1,X2,esk16_1(k1_zfmisc_1(k2_zfmisc_1(X2,X2))),X3),X1)|~v1_funct_2(X3,X2,u1_struct_0(X1))|~v1_funct_1(X3)|~l1_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X1)))))).#trainneg
cnf(i_0_4498, plain, (r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))|~r2_hidden(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_431, plain, (v6_waybel_0(g1_waybel_0(X1,X2,k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X2)),X4,X5),X6),X1)|v1_xboole_0(X3)|~v1_funct_2(X6,X2,u1_struct_0(X1))|~v1_funct_1(X6)|~l1_struct_0(X1)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X1))))|~v1_funct_2(X4,X3,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))))|~m1_subset_1(X5,X3))).#trainneg
cnf(i_0_432, plain, (v6_waybel_0(g1_waybel_0(X1,u1_struct_0(X2),u1_orders_2(X2),X3),X1)|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~l1_orders_2(X2))).#trainneg
cnf(i_0_3668, plain, (esk6_3(X1,X2,X3)=k1_xboole_0|r3_waybel_9(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|u1_struct_0(X1)!=k1_xboole_0|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4829, plain, (r2_waybel_0(X1,X2,u1_struct_0(X3))|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r2_waybel_0(X1,X2,k10_yellow_6(X3,X4))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_waybel_0(X4,X3)|~v7_waybel_0(X4)|~v4_orders_2(X4)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(esk8_3(X1,X2,u1_struct_0(X3)),u1_struct_0(X2)))).#trainneg
cnf(i_0_4821, plain, (r2_waybel_0(X1,X2,u1_struct_0(k3_yellow_1(k2_struct_0(X3))))|v1_compts_1(X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r2_waybel_0(X1,X2,esk4_1(X3))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(esk8_3(X1,X2,u1_struct_0(k3_yellow_1(k2_struct_0(X3)))),u1_struct_0(X2)))).#trainneg
cnf(i_0_4916, plain, (r3_waybel_9(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X4)|~r3_waybel_9(X1,X4,X3)|~r1_waybel_0(X1,X4,X5)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X4,X1)|~l1_waybel_0(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|u1_struct_0(X1)!=k1_xboole_0)).#trainneg
cnf(i_0_4920, plain, (r2_waybel_0(X1,X2,k1_xboole_0)|r3_waybel_9(X1,X3,X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r3_waybel_9(X1,X2,X4)|~m1_subset_1(X4,u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~l1_waybel_0(X3,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|u1_struct_0(X1)!=k1_xboole_0)).#trainneg
cnf(i_0_433, plain, (v6_waybel_0(g1_waybel_0(X1,u1_struct_0(X2),u1_waybel_0(X2,X2),X3),X1)|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~l1_struct_0(X2)|~l1_waybel_0(X2,X2))).#trainneg
cnf(i_0_4830, plain, (X1=k2_struct_0(g1_orders_2(k2_zfmisc_1(X2,X3),k1_xboole_0))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_434, plain, (v6_waybel_0(g1_waybel_0(X1,X2,k3_subset_1(k2_zfmisc_1(X2,X2),X3),X4),X1)|~v1_funct_2(X4,X2,u1_struct_0(X1))|~v1_funct_1(X4)|~l1_struct_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X2))))).#trainneg
cnf(i_0_436, plain, (esk16_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1)))=X2|g1_waybel_0(X3,X1,esk16_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))),X4)!=g1_waybel_0(X5,X6,X2,X7)|~v1_funct_2(X4,X1,u1_struct_0(X3))|~v1_funct_1(X4)|~l1_struct_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X3)))))).#trainneg
cnf(i_0_4915, plain, (r3_waybel_9(X1,X2,X3)|r1_waybel_0(X1,X4,k1_xboole_0)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X4)|~r3_waybel_9(X1,X4,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_yellow19(X4,X1)|~l1_waybel_0(X4,X1)|~l1_waybel_0(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|u1_struct_0(X1)!=k1_xboole_0)).#trainneg
cnf(i_0_4813, plain, (k10_yellow_6(X1,X2)=k1_xboole_0|r2_hidden(esk15_3(X3,X3,k10_yellow_6(X1,X2)),X4)|r1_waybel_0(X1,X5,esk11_4(X1,X6,X4,esk15_3(X3,X3,k10_yellow_6(X1,X2))))|v2_struct_0(X6)|v2_struct_0(X2)|v2_struct_0(X5)|v2_struct_0(X1)|k10_yellow_6(X1,X2)!=k10_yellow_6(X1,X5)|X4!=k10_yellow_6(X1,X6)|~m1_subset_1(esk15_3(X3,X3,k10_yellow_6(X1,X2)),u1_struct_0(X1))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~l1_waybel_0(X5,X1)|~l1_waybel_0(X2,X1)|~l1_waybel_0(X6,X1)|~v7_waybel_0(X5)|~v7_waybel_0(X2)|~v7_waybel_0(X6)|~v4_orders_2(X5)|~v4_orders_2(X2)|~v4_orders_2(X6)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4912, plain, (m1_connsp_2(k1_xboole_0,X1,X3)|r3_waybel_9(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|u1_struct_0(X1)!=k1_xboole_0)).#trainneg
cnf(i_0_4662, plain, (k3_subset_1(X1,u1_struct_0(X2))=k1_xboole_0|v1_xboole_0(u1_struct_0(X2))|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(X1))|~l1_struct_0(X2)|~m1_subset_1(k3_subset_1(X1,u1_struct_0(X2)),k1_zfmisc_1(k2_struct_0(X2))))).#trainneg
cnf(i_0_2971, plain, (m1_subset_1(esk15_3(X1,X1,k3_subset_1(X2,X3)),X2)|k3_subset_1(X2,X3)=k1_xboole_0|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_475, plain, (r1_waybel_0(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,k2_struct_0(X3),X4)))|v2_struct_0(X1)|v2_struct_0(X2)|v1_xboole_0(X4)|v2_struct_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X5,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~v1_subset_1(X4,u1_struct_0(k3_yellow_1(k2_struct_0(X3))))|~r2_waybel_7(X3,X4,k2_waybel_0(X1,X2,esk10_4(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,k2_struct_0(X3),X4)),X5)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X3)))))|~m1_subset_1(k2_waybel_0(X1,X2,esk10_4(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,k2_struct_0(X3),X4)),X5)),u1_struct_0(X3))|~v13_waybel_0(X4,k3_yellow_1(k2_struct_0(X3)))|~v2_waybel_0(X4,k3_yellow_1(k2_struct_0(X3)))|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_999, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X4)|~v1_xboole_0(u1_struct_0(X3))|~r2_waybel_0(X1,X2,k10_yellow_6(X3,X4))|~l1_struct_0(X1)|~m1_subset_1(X5,u1_struct_0(X2))|~l1_waybel_0(X4,X3)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X4)|~v4_orders_2(X4)|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_4848, plain, (v2_struct_0(g1_orders_2(k2_zfmisc_1(X1,X2),k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_438, plain, (u1_orders_2(X1)=X2|g1_waybel_0(X3,u1_struct_0(X1),u1_orders_2(X1),X4)!=g1_waybel_0(X5,X6,X2,X7)|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X3))|~v1_funct_1(X4)|~l1_struct_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))|~l1_orders_2(X1))).#trainneg
cnf(i_0_4957, plain, (u1_struct_0(X1)=k2_struct_0(g1_orders_2(k2_zfmisc_1(X2,X3),k1_xboole_0))|~v1_xboole_0(X3)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_437, plain, (k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)),X3,X4)=X5|v1_xboole_0(X1)|g1_waybel_0(X6,X2,k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)),X3,X4),X7)!=g1_waybel_0(X8,X9,X5,X10)|~v1_funct_2(X7,X2,u1_struct_0(X6))|~v1_funct_1(X7)|~l1_struct_0(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X6))))|~v1_funct_2(X3,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))))|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_439, plain, (u1_waybel_0(X1,X1)=X2|g1_waybel_0(X3,u1_struct_0(X1),u1_waybel_0(X1,X1),X4)!=g1_waybel_0(X5,X6,X2,X7)|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X3))|~v1_funct_1(X4)|~l1_struct_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))|~l1_struct_0(X1)|~l1_waybel_0(X1,X1))).#trainneg
cnf(i_0_4922, negated_conjecture, (r2_waybel_0(X1,X2,u1_struct_0(esk1_0))|v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(esk2_0)|~v4_orders_2(esk2_0)|~v1_compts_1(esk1_0))).#trainneg
cnf(i_0_5049, plain, (r1_waybel_0(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,k2_struct_0(X3),X4)))|v1_xboole_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~v1_subset_1(X4,u1_struct_0(k3_yellow_1(k2_struct_0(X3))))|~l1_struct_0(X1)|~m1_subset_1(k2_waybel_0(X1,X2,esk10_4(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,k2_struct_0(X3),X4)),esk9_3(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,u1_struct_0(X3),X4))))),u1_struct_0(X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X3)))))|~m1_subset_1(esk9_3(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,u1_struct_0(X3),X4))),u1_struct_0(X2))|~v13_waybel_0(X4,k3_yellow_1(k2_struct_0(X3)))|~v2_waybel_0(X4,k3_yellow_1(k2_struct_0(X3)))|~l1_waybel_0(X2,X1)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~v1_subset_1(X4,u1_struct_0(k3_yellow_1(u1_struct_0(X3))))|~r1_waybel_0(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,u1_struct_0(X3),X4)))|~m1_subset_1(esk10_4(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,k2_struct_0(X3),X4)),esk9_3(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,u1_struct_0(X3),X4)))),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X3)))))|~v13_waybel_0(X4,k3_yellow_1(u1_struct_0(X3)))|~v2_waybel_0(X4,k3_yellow_1(u1_struct_0(X3))))).#trainneg
cnf(i_0_4996, plain, (k3_subset_1(k2_struct_0(X1),u1_struct_0(X1))=k1_xboole_0|v1_xboole_0(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(k2_struct_0(X1))))).#trainneg
cnf(i_0_4979, plain, (r3_waybel_9(X1,X2,esk14_2(X1,X3))|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|u1_struct_0(X1)!=k1_xboole_0|~m1_subset_1(esk14_2(X1,X3),u1_struct_0(X1))|~v1_yellow19(X3,X1)|~l1_waybel_0(X3,X1)|~l1_waybel_0(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~v1_compts_1(X1))).#trainneg
cnf(i_0_4983, plain, (r2_hidden(esk15_3(X1,X1,k10_yellow_6(X2,X3)),X4)|v2_struct_0(X2)|v2_struct_0(X5)|k10_yellow_6(X2,X3)=k1_xboole_0|v2_struct_0(X3)|X4!=k10_yellow_6(X2,X5)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(esk15_3(X1,X1,k10_yellow_6(X2,X3)),u1_struct_0(X2))|~l1_waybel_0(X5,X2)|~v7_waybel_0(X5)|~v4_orders_2(X5)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|k10_yellow_6(X2,X3)!=k10_yellow_6(X2,X5)|~l1_waybel_0(X3,X2)|~v7_waybel_0(X3)|~v4_orders_2(X3))).#trainneg
cnf(i_0_4935, plain, (r3_waybel_9(X1,X2,X3)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X2)|v1_xboole_0(u1_struct_0(X1))|u1_struct_0(X1)!=k1_xboole_0|~r3_waybel_9(X1,X4,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X4,X1)|~l1_waybel_0(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_5044, plain, (r1_waybel_0(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,u1_struct_0(X3),X4)))|v1_xboole_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~v1_subset_1(X4,u1_struct_0(k3_yellow_1(u1_struct_0(X3))))|~l1_struct_0(X1)|~r2_waybel_7(X3,X4,k2_waybel_0(X1,X2,esk10_4(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,u1_struct_0(X3),X4)),X5)))|~m1_subset_1(k2_waybel_0(X1,X2,esk10_4(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,u1_struct_0(X3),X4)),X5)),u1_struct_0(X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X3)))))|~m1_subset_1(X5,u1_struct_0(X2))|~v13_waybel_0(X4,k3_yellow_1(u1_struct_0(X3)))|~v2_waybel_0(X4,k3_yellow_1(u1_struct_0(X3)))|~l1_waybel_0(X2,X1)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~l1_struct_0(X3))).#trainneg
cnf(i_0_4924, plain, (r2_waybel_0(X1,X2,u1_struct_0(X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(k3_yellow19(X3,k2_struct_0(X3),esk4_1(X3)))|v1_compts_1(X3)|~r2_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_waybel_0(k3_yellow19(X3,k2_struct_0(X3),esk4_1(X3)),X3)|~v7_waybel_0(k3_yellow19(X3,k2_struct_0(X3),esk4_1(X3)))|~v4_orders_2(k3_yellow19(X3,k2_struct_0(X3),esk4_1(X3)))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~v1_subset_1(esk4_1(X3),u1_struct_0(k3_yellow_1(k2_struct_0(X3)))))).#trainneg
cnf(i_0_4923, plain, (r2_waybel_0(X1,X2,u1_struct_0(X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3)))|v1_compts_1(X3)|~r2_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_waybel_0(k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3)),X3)|~v7_waybel_0(k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3)))|~v4_orders_2(k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3)))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~v1_subset_1(esk4_1(X3),u1_struct_0(k3_yellow_1(u1_struct_0(X3)))))).#trainneg
cnf(i_0_440, plain, (k3_subset_1(k2_zfmisc_1(X1,X1),X2)=X3|g1_waybel_0(X4,X1,k3_subset_1(k2_zfmisc_1(X1,X1),X2),X5)!=g1_waybel_0(X6,X7,X3,X8)|~v1_funct_2(X5,X1,u1_struct_0(X4))|~v1_funct_1(X5)|~l1_struct_0(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X4))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_442, plain, (X1=X2|g1_waybel_0(X3,X1,esk16_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))),X4)!=g1_waybel_0(X5,X2,X6,X7)|~v1_funct_2(X4,X1,u1_struct_0(X3))|~v1_funct_1(X4)|~l1_struct_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X3)))))).#trainneg
cnf(i_0_443, plain, (X1=X2|v1_xboole_0(X4)|g1_waybel_0(X3,X1,k3_funct_2(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X1)),X5,X6),X7)!=g1_waybel_0(X8,X2,X9,X10)|~v1_funct_2(X7,X1,u1_struct_0(X3))|~v1_funct_1(X7)|~l1_struct_0(X3)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X3))))|~v1_funct_2(X5,X4,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))|~m1_subset_1(X6,X4))).#trainneg
cnf(i_0_2375, plain, (r1_waybel_0(X1,X2,X3)|v1_xboole_0(u1_struct_0(X4))|v2_struct_0(X2)|v2_struct_0(X4)|v2_struct_0(X1)|u1_struct_0(X4)!=k1_xboole_0|~r1_waybel_0(X4,X2,X5)|~l1_struct_0(X4)|~l1_struct_0(X1)|~m1_subset_1(esk10_4(X1,X2,X3,esk9_3(X4,X2,X5)),u1_struct_0(X2))|~l1_waybel_0(X2,X4)|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_830, plain, (r2_hidden(k2_waybel_0(X1,X2,esk7_4(X1,X2,esk4_1(X3),X4)),u1_struct_0(k3_yellow_1(u1_struct_0(X3))))|v1_compts_1(X3)|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~r2_waybel_0(X1,X2,esk4_1(X3))|~l1_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_1003, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|k10_yellow_6(X3,X4)!=k1_xboole_0|~r2_waybel_0(X2,X1,k10_yellow_6(X3,X4))|~l1_struct_0(X2)|~m1_subset_1(X5,u1_struct_0(X1))|~l1_waybel_0(X1,X2)|~l1_waybel_0(X4,X3)|~v7_waybel_0(X4)|~v4_orders_2(X4)|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_515, plain, (r2_waybel_7(X1,X2,k2_waybel_0(X3,X4,esk7_4(X5,X4,X6,esk9_3(X3,X4,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2))))))|v1_xboole_0(X2)|v2_struct_0(X1)|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X4)|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X1)))))|~m1_subset_1(k2_waybel_0(X3,X4,esk7_4(X5,X4,X6,esk9_3(X3,X4,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2))))),u1_struct_0(X1))|~v13_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~v2_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~r2_waybel_0(X5,X4,X6)|~r1_waybel_0(X3,X4,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2)))|~l1_struct_0(X3)|~l1_struct_0(X5)|~m1_subset_1(esk7_4(X5,X4,X6,esk9_3(X3,X4,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2)))),u1_struct_0(X4))|~l1_waybel_0(X4,X3)|~l1_waybel_0(X4,X5))).#trainneg
cnf(i_0_5168, plain, (r2_waybel_0(X1,X2,u1_struct_0(X3))|v1_compts_1(X3)|v2_struct_0(k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3)))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X3))|v1_xboole_0(esk4_1(X3))|~v1_subset_1(esk4_1(X3),u1_struct_0(k3_yellow_1(u1_struct_0(X3))))|~r2_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3)))|~v4_orders_2(k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3)))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~l1_struct_0(X3)|~m1_subset_1(esk4_1(X3),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X3)))))|~v13_waybel_0(esk4_1(X3),k3_yellow_1(u1_struct_0(X3)))|~v2_waybel_0(esk4_1(X3),k3_yellow_1(u1_struct_0(X3))))).#trainneg
cnf(i_0_1532, plain, (X1=X2|g1_orders_2(X1,k2_struct_0(g1_orders_2(k2_zfmisc_1(X1,X1),k1_xboole_0)))!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_5060, plain, (u1_orders_2(X1)=u1_orders_2(X2)|g1_waybel_0(X3,u1_struct_0(X1),u1_orders_2(X1),X4)!=X2|~l1_orders_2(X1)|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X3))|~v1_funct_1(X4)|~l1_struct_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))|~v6_waybel_0(X2,X5)|~l1_struct_0(X5)|~l1_waybel_0(X2,X5))).#trainneg
cnf(i_0_444, plain, (u1_struct_0(X1)=X2|g1_waybel_0(X3,u1_struct_0(X1),u1_orders_2(X1),X4)!=g1_waybel_0(X5,X2,X6,X7)|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X3))|~v1_funct_1(X4)|~l1_struct_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))|~l1_orders_2(X1))).#trainneg
cnf(i_0_445, plain, (u1_struct_0(X1)=X2|g1_waybel_0(X3,u1_struct_0(X1),u1_waybel_0(X1,X1),X4)!=g1_waybel_0(X5,X2,X6,X7)|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X3))|~v1_funct_1(X4)|~l1_struct_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))|~l1_struct_0(X1)|~l1_waybel_0(X1,X1))).#trainneg
cnf(i_0_446, plain, (X1=X2|g1_waybel_0(X3,X1,k3_subset_1(k2_zfmisc_1(X1,X1),X4),X5)!=g1_waybel_0(X6,X2,X7,X8)|~v1_funct_2(X5,X1,u1_struct_0(X3))|~v1_funct_1(X5)|~l1_struct_0(X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_5166, plain, (r2_waybel_0(X1,X2,u1_struct_0(X3))|v1_compts_1(X3)|v2_struct_0(k3_yellow19(X3,k2_struct_0(X3),esk4_1(X3)))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v1_xboole_0(k2_struct_0(X3))|v1_xboole_0(esk4_1(X3))|~v1_subset_1(esk4_1(X3),u1_struct_0(k3_yellow_1(k2_struct_0(X3))))|~r2_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(k3_yellow19(X3,k2_struct_0(X3),esk4_1(X3)))|~v4_orders_2(k3_yellow19(X3,k2_struct_0(X3),esk4_1(X3)))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~l1_struct_0(X3)|~m1_subset_1(esk4_1(X3),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X3)))))|~m1_subset_1(k2_struct_0(X3),k1_zfmisc_1(u1_struct_0(X3)))|~v13_waybel_0(esk4_1(X3),k3_yellow_1(k2_struct_0(X3)))|~v2_waybel_0(esk4_1(X3),k3_yellow_1(k2_struct_0(X3))))).#trainneg
cnf(i_0_5085, negated_conjecture, (r1_waybel_0(X1,X2,u1_struct_0(esk1_0))|v2_struct_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(X1)))|~v1_yellow19(X2,X1)|~l1_waybel_0(X2,X1)|~r2_waybel_0(X1,X2,k1_xboole_0)|~v1_compts_1(esk1_0))).#trainneg
cnf(i_0_5096, plain, (r1_waybel_0(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,k2_struct_0(X3),esk4_1(X3))))|v1_xboole_0(esk4_1(X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v1_compts_1(X3)|v2_struct_0(k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3)))|~v1_subset_1(esk4_1(X3),u1_struct_0(k3_yellow_1(k2_struct_0(X3))))|~v1_subset_1(esk4_1(X3),u1_struct_0(k3_yellow_1(u1_struct_0(X3))))|~r1_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X1)|~m1_subset_1(k2_waybel_0(X1,X2,esk10_4(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,k2_struct_0(X3),esk4_1(X3))),esk9_3(X1,X2,k1_xboole_0))),u1_struct_0(X3))|~m1_subset_1(esk10_4(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,k2_struct_0(X3),esk4_1(X3))),esk9_3(X1,X2,k1_xboole_0)),u1_struct_0(X2))|~m1_subset_1(esk4_1(X3),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X3)))))|~m1_subset_1(esk4_1(X3),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X3)))))|~v13_waybel_0(esk4_1(X3),k3_yellow_1(k2_struct_0(X3)))|~v13_waybel_0(esk4_1(X3),k3_yellow_1(u1_struct_0(X3)))|~v2_waybel_0(esk4_1(X3),k3_yellow_1(k2_struct_0(X3)))|~v2_waybel_0(esk4_1(X3),k3_yellow_1(u1_struct_0(X3)))|~l1_waybel_0(X2,X1)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~l1_waybel_0(k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3)),X3)|~v7_waybel_0(k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3)))|~v4_orders_2(k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3))))).#trainneg
cnf(i_0_5207, plain, (v1_compts_1(X1)|v2_struct_0(X1)|v1_xboole_0(esk4_1(X1))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k2_waybel_0(X2,X3,esk7_4(X4,X3,X5,esk9_3(X2,X3,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))))),u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~r2_waybel_0(X4,X3,X5)|~r1_waybel_0(X2,X3,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))))|~l1_struct_0(X2)|~l1_struct_0(X4)|~m1_subset_1(esk7_4(X4,X3,X5,esk9_3(X2,X3,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))))),u1_struct_0(X3))|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X1)))))|~v13_waybel_0(esk4_1(X1),k3_yellow_1(u1_struct_0(X1)))|~v2_waybel_0(esk4_1(X1),k3_yellow_1(u1_struct_0(X1)))|~l1_waybel_0(X3,X2)|~l1_waybel_0(X3,X4))).#trainneg
cnf(i_0_5194, plain, (m1_subset_1(k2_waybel_0(X1,X2,esk7_4(X1,X2,esk4_1(X3),X4)),u1_struct_0(k3_yellow_1(u1_struct_0(X3))))|v1_compts_1(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|~r2_waybel_0(X1,X2,esk4_1(X3))|~l1_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_5103, plain, (r1_waybel_0(X1,X2,k1_xboole_0)|v1_xboole_0(esk4_1(X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v1_compts_1(X3)|v2_struct_0(k3_yellow19(X3,k2_struct_0(X3),esk4_1(X3)))|~v1_subset_1(esk4_1(X3),u1_struct_0(k3_yellow_1(k2_struct_0(X3))))|~v1_subset_1(esk4_1(X3),u1_struct_0(k3_yellow_1(u1_struct_0(X3))))|~r1_waybel_0(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3))))|~l1_struct_0(X1)|~m1_subset_1(k2_waybel_0(X1,X2,esk10_4(X1,X2,k1_xboole_0,esk9_3(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3)))))),u1_struct_0(X3))|~m1_subset_1(esk10_4(X1,X2,k1_xboole_0,esk9_3(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3))))),u1_struct_0(X2))|~m1_subset_1(esk4_1(X3),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X3)))))|~m1_subset_1(esk4_1(X3),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X3)))))|~v13_waybel_0(esk4_1(X3),k3_yellow_1(k2_struct_0(X3)))|~v13_waybel_0(esk4_1(X3),k3_yellow_1(u1_struct_0(X3)))|~v2_waybel_0(esk4_1(X3),k3_yellow_1(k2_struct_0(X3)))|~v2_waybel_0(esk4_1(X3),k3_yellow_1(u1_struct_0(X3)))|~l1_waybel_0(X2,X1)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~l1_waybel_0(k3_yellow19(X3,k2_struct_0(X3),esk4_1(X3)),X3)|~v7_waybel_0(k3_yellow19(X3,k2_struct_0(X3),esk4_1(X3)))|~v4_orders_2(k3_yellow19(X3,k2_struct_0(X3),esk4_1(X3))))).#trainneg
cnf(i_0_5167, plain, (r2_waybel_0(X1,X2,X3)|v1_compts_1(g1_orders_2(X3,k1_xboole_0))|v2_struct_0(k3_yellow19(g1_orders_2(X3,k1_xboole_0),X3,esk4_1(g1_orders_2(X3,k1_xboole_0))))|v2_struct_0(g1_orders_2(X3,k1_xboole_0))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_subset_1(esk4_1(g1_orders_2(X3,k1_xboole_0)),u1_struct_0(k3_yellow_1(X3)))|~r2_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X1)|~l1_waybel_0(k3_yellow19(g1_orders_2(X3,k1_xboole_0),X3,esk4_1(g1_orders_2(X3,k1_xboole_0))),g1_orders_2(X3,k1_xboole_0))|~l1_waybel_0(X2,X1)|~v7_waybel_0(k3_yellow19(g1_orders_2(X3,k1_xboole_0),X3,esk4_1(g1_orders_2(X3,k1_xboole_0))))|~v4_orders_2(k3_yellow19(g1_orders_2(X3,k1_xboole_0),X3,esk4_1(g1_orders_2(X3,k1_xboole_0))))|~l1_pre_topc(g1_orders_2(X3,k1_xboole_0))|~v2_pre_topc(g1_orders_2(X3,k1_xboole_0)))).#trainneg
cnf(i_0_1544, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k2_struct_0(g1_orders_2(X2,k1_xboole_0))))).#trainneg
cnf(i_0_1530, plain, (v1_xboole_0(k2_struct_0(g1_orders_2(k2_zfmisc_1(X1,X2),k1_xboole_0)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_448, plain, (X1=X2|g1_waybel_0(X1,X3,esk16_1(k1_zfmisc_1(k2_zfmisc_1(X3,X3))),X4)!=g1_waybel_0(X2,X5,X6,X7)|~v1_funct_2(X4,X3,u1_struct_0(X1))|~v1_funct_1(X4)|~l1_struct_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1361, plain, (k3_subset_1(X1,X2)=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_449, plain, (X1=X2|v1_xboole_0(X4)|g1_waybel_0(X1,X3,k3_funct_2(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X3)),X5,X6),X7)!=g1_waybel_0(X2,X8,X9,X10)|~v1_funct_2(X7,X3,u1_struct_0(X1))|~v1_funct_1(X7)|~l1_struct_0(X1)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(X1))))|~v1_funct_2(X5,X4,k1_zfmisc_1(k2_zfmisc_1(X3,X3)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X3)))))|~m1_subset_1(X6,X4))).#trainneg
cnf(i_0_450, plain, (X1=X2|g1_waybel_0(X1,u1_struct_0(X3),u1_orders_2(X3),X4)!=g1_waybel_0(X2,X5,X6,X7)|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X4)|~l1_struct_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~l1_orders_2(X3))).#trainneg
cnf(i_0_5254, plain, (r1_waybel_0(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3))))|v1_compts_1(X3)|v2_struct_0(k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3)))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~v1_subset_1(esk4_1(X3),u1_struct_0(k3_yellow_1(u1_struct_0(X3))))|~r1_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X1)|~m1_subset_1(k2_waybel_0(X1,X2,esk10_4(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3))),esk9_3(X1,X2,k1_xboole_0))),u1_struct_0(X3))|~m1_subset_1(esk10_4(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3))),esk9_3(X1,X2,k1_xboole_0)),u1_struct_0(X2))|~l1_waybel_0(k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3)),X3)|~l1_waybel_0(X2,X1)|~v7_waybel_0(k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3)))|~v4_orders_2(k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3)))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~l1_struct_0(X3))).#trainneg
cnf(i_0_5265, plain, (v1_compts_1(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X1)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~r2_waybel_0(X2,X3,u1_struct_0(X1))|~r1_waybel_0(X2,X3,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))))|~l1_struct_0(X2)|~m1_subset_1(esk7_4(X2,X3,u1_struct_0(X1),esk9_3(X2,X3,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))))),u1_struct_0(X3))|~l1_waybel_0(X3,X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk9_3(X2,X3,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))),u1_struct_0(X3)))).#trainneg
cnf(i_0_5266, plain, (v1_compts_1(X1)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~r2_waybel_0(X2,X3,X4)|~r1_waybel_0(X1,X3,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))))|~l1_struct_0(X1)|~l1_struct_0(X2)|~m1_subset_1(esk7_4(X2,X3,X4,esk9_3(X1,X3,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))))),u1_struct_0(X3))|~l1_waybel_0(X3,X1)|~l1_waybel_0(X3,X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_5259, plain, (r1_waybel_0(X1,X2,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1))))|v1_compts_1(X1)|v2_struct_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))|v2_struct_0(X1)|v2_struct_0(X2)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~r1_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X1)|~m1_subset_1(esk10_4(X1,X2,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1))),esk9_3(X1,X2,k1_xboole_0)),u1_struct_0(X2))|~l1_waybel_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)),X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))|~v4_orders_2(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_5213, plain, (r2_waybel_0(X1,X2,u1_struct_0(X3))|v1_xboole_0(u1_struct_0(X3))|v1_compts_1(X3)|v2_struct_0(k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3)))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v1_xboole_0(esk4_1(X3))|~v1_subset_1(esk4_1(X3),u1_struct_0(k3_yellow_1(u1_struct_0(X3))))|~r2_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~v4_orders_2(k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3)))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~l1_struct_0(X3)|~m1_subset_1(esk4_1(X3),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X3)))))|~v13_waybel_0(esk4_1(X3),k3_yellow_1(u1_struct_0(X3)))|~v2_waybel_0(esk4_1(X3),k3_yellow_1(u1_struct_0(X3))))).#trainneg
cnf(i_0_5249, plain, (r2_waybel_0(X1,X2,u1_struct_0(X3))|v1_xboole_0(k2_struct_0(X3))|v1_compts_1(X3)|v2_struct_0(k3_yellow19(X3,k2_struct_0(X3),esk4_1(X3)))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v1_xboole_0(esk4_1(X3))|~r2_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(esk4_1(X3),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X3)))))|~v3_waybel_7(esk4_1(X3),k3_yellow_1(k2_struct_0(X3)))|~v13_waybel_0(esk4_1(X3),k3_yellow_1(k2_struct_0(X3)))|~v2_waybel_0(esk4_1(X3),k3_yellow_1(k2_struct_0(X3))))).#trainneg
cnf(i_0_5222, plain, (u1_orders_2(X2)=k1_xboole_0|g1_waybel_0(X3,X1,k1_xboole_0,X4)!=X2|~v1_funct_2(X4,X1,u1_struct_0(X3))|~v1_funct_1(X4)|~v6_waybel_0(X2,X5)|~l1_struct_0(X3)|~l1_struct_0(X5)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X3))))|~l1_waybel_0(X2,X5))).#trainneg
cnf(i_0_451, plain, (X1=X2|g1_waybel_0(X1,u1_struct_0(X3),u1_waybel_0(X3,X3),X4)!=g1_waybel_0(X2,X5,X6,X7)|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X4)|~l1_struct_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~l1_struct_0(X3)|~l1_waybel_0(X3,X3))).#trainneg
cnf(i_0_5303, plain, (X1=k2_struct_0(g1_orders_2(k2_zfmisc_1(X2,X3),k1_xboole_0))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_452, plain, (X1=X2|g1_waybel_0(X1,X3,k3_subset_1(k2_zfmisc_1(X3,X3),X4),X5)!=g1_waybel_0(X2,X6,X7,X8)|~v1_funct_2(X5,X3,u1_struct_0(X1))|~v1_funct_1(X5)|~l1_struct_0(X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X3))))).#trainneg
cnf(i_0_453, plain, (l1_orders_2(k3_yellow19(X1,X2,X3))|v1_xboole_0(X2)|v1_xboole_0(X3)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v13_waybel_0(X3,k3_yellow_1(X2))|~v2_waybel_0(X3,k3_yellow_1(X2)))).#trainneg
cnf(i_0_5199, plain, (r2_waybel_0(X1,X2,u1_struct_0(k3_yellow_1(u1_struct_0(X3))))|v2_struct_0(X1)|v2_struct_0(X2)|v1_compts_1(X3)|v2_struct_0(X3)|~r2_waybel_0(X1,X2,esk4_1(X3))|~l1_struct_0(X1)|~m1_subset_1(esk7_4(X1,X2,esk4_1(X3),esk8_3(X1,X2,u1_struct_0(k3_yellow_1(u1_struct_0(X3))))),u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~m1_subset_1(esk8_3(X1,X2,u1_struct_0(k3_yellow_1(u1_struct_0(X3)))),u1_struct_0(X2))|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_5372, plain, (X1=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5322, plain, (v2_struct_0(g1_orders_2(k2_zfmisc_1(X1,X2),k1_xboole_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_454, plain, (v1_xboole_0(X1)|v1_yellow19(k3_yellow19(X2,u1_struct_0(X2),X1),X2)|v2_struct_0(X2)|~l1_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X2)))))|~v3_waybel_7(X1,k3_yellow_1(u1_struct_0(X2)))|~v13_waybel_0(X1,k3_yellow_1(u1_struct_0(X2)))|~v2_waybel_0(X1,k3_yellow_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_4063, plain, (r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))|X1!=k1_zfmisc_1(k1_xboole_0)|~r2_hidden(X3,X1)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_455, plain, (v6_waybel_0(k3_yellow19(X1,u1_struct_0(X1),X2),X1)|v1_xboole_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X1)))))|~v3_waybel_7(X2,k3_yellow_1(u1_struct_0(X1)))|~v13_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~v2_waybel_0(X2,k3_yellow_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_456, plain, (m1_subset_1(esk3_2(X1,X2),u1_struct_0(X1))|v1_xboole_0(X2)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X1)))))|~v3_waybel_7(X2,k3_yellow_1(u1_struct_0(X1)))|~v13_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~v2_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~v1_compts_1(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_5445, plain, (r2_waybel_0(X1,X2,u1_struct_0(X3))|v1_xboole_0(u1_struct_0(X3))|v1_compts_1(X3)|v2_struct_0(k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3)))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v1_xboole_0(esk4_1(X3))|~v1_subset_1(esk4_1(X3),u1_struct_0(k3_yellow_1(u1_struct_0(X3))))|~r2_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~l1_struct_0(X3)|~m1_subset_1(esk4_1(X3),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X3)))))|~v13_waybel_0(esk4_1(X3),k3_yellow_1(u1_struct_0(X3)))|~v2_waybel_0(esk4_1(X3),k3_yellow_1(u1_struct_0(X3))))).#trainneg
cnf(i_0_5428, plain, (v1_compts_1(X1)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~r2_waybel_0(X3,X2,u1_struct_0(X1))|~r1_waybel_0(X3,X2,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))))|~l1_struct_0(X3)|~l1_waybel_0(X2,X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk9_3(X3,X2,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))),u1_struct_0(X2)))).#trainneg
cnf(i_0_5519, plain, (r1_waybel_0(X1,k3_yellow19(X1,u1_struct_0(X1),X2),u1_struct_0(X1))|r1_waybel_0(X1,k3_yellow19(X1,u1_struct_0(X1),X2),k1_xboole_0)|v2_struct_0(X1)|v2_struct_0(k3_yellow19(X1,u1_struct_0(X1),X2))|v1_xboole_0(X2)|~l1_struct_0(X1)|~l1_waybel_0(k3_yellow19(X1,u1_struct_0(X1),X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X1)))))|~v3_waybel_7(X2,k3_yellow_1(u1_struct_0(X1)))|~v13_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~v2_waybel_0(X2,k3_yellow_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_5433, plain, (v1_compts_1(X1)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~r2_waybel_0(X3,X2,X4)|~r1_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X3)|~m1_subset_1(esk7_4(X3,X2,X4,esk9_3(X1,X2,k1_xboole_0)),u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_waybel_0(X2,X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)),X1)|~v7_waybel_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))|~v4_orders_2(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))))).#trainneg
cnf(i_0_5521, plain, (v1_xboole_0(X1)|v1_yellow19(k3_yellow19(g1_orders_2(X2,k1_xboole_0),X2,X1),g1_orders_2(X2,k1_xboole_0))|v2_struct_0(g1_orders_2(X2,k1_xboole_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))|~v3_waybel_7(X1,k3_yellow_1(X2))|~v13_waybel_0(X1,k3_yellow_1(X2))|~v2_waybel_0(X1,k3_yellow_1(X2)))).#trainneg
cnf(i_0_5443, plain, (r1_waybel_0(X1,X2,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),esk4_1(X1))))|v1_compts_1(X1)|v2_struct_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~r1_waybel_0(X1,X2,k1_xboole_0)|~l1_waybel_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)),X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))|~v4_orders_2(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1)|~m1_subset_1(esk9_3(X1,X2,k1_xboole_0),u1_struct_0(X2)))).#trainneg
cnf(i_0_5493, plain, (r2_waybel_0(X1,X2,u1_struct_0(k3_yellow_1(u1_struct_0(X3))))|v1_compts_1(X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r2_waybel_0(X1,X2,esk4_1(X3))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(esk8_3(X1,X2,u1_struct_0(k3_yellow_1(u1_struct_0(X3)))),u1_struct_0(X2)))).#trainneg
cnf(i_0_5434, plain, (v1_compts_1(X1)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~r2_waybel_0(X3,X2,X4)|~r1_waybel_0(X1,X2,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))))|~l1_struct_0(X3)|~l1_waybel_0(X2,X1)|~l1_waybel_0(X2,X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk9_3(X1,X2,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))),u1_struct_0(X2)))).#trainneg
cnf(i_0_5421, plain, (r1_waybel_0(X1,X2,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))))|v1_compts_1(X1)|v2_struct_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))|v2_struct_0(X1)|v2_struct_0(X2)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~r1_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X1)|~m1_subset_1(esk10_4(X1,X2,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))),esk9_3(X1,X2,k1_xboole_0)),u1_struct_0(X2))|~l1_waybel_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)),X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))|~v4_orders_2(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4681, plain, (k4_xboole_0(k10_yellow_6(X1,X2),X3)=k1_xboole_0|r1_waybel_0(X1,X4,esk6_3(X1,X5,esk15_3(X6,X6,k4_xboole_0(k10_yellow_6(X1,X2),X3))))|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X2)|r3_waybel_9(X1,X5,esk15_3(X6,X6,k4_xboole_0(k10_yellow_6(X1,X2),X3)))|v2_struct_0(X5)|k10_yellow_6(X1,X2)!=k10_yellow_6(X1,X4)|~l1_waybel_0(X4,X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X4)|~v7_waybel_0(X2)|~v4_orders_2(X4)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk15_3(X6,X6,k4_xboole_0(k10_yellow_6(X1,X2),X3)),u1_struct_0(X1))|~l1_waybel_0(X5,X1))).#trainneg
cnf(i_0_3476, plain, (X1=k1_xboole_0|v1_xboole_0(k3_subset_1(X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(X3))|~m1_subset_1(esk15_3(X2,X2,X1),k3_subset_1(X3,X1)))).#trainneg
cnf(i_0_2170, plain, (k2_struct_0(g1_orders_2(X1,k1_xboole_0))=k1_xboole_0|X1!=k1_xboole_0)).#trainneg
cnf(i_0_950, plain, (u1_struct_0(g1_orders_2(X1,k1_xboole_0))=k1_xboole_0|~v2_struct_0(g1_orders_2(X1,k1_xboole_0)))).#trainneg
cnf(i_0_5470, plain, (u1_struct_0(X1)=k2_struct_0(g1_orders_2(k2_zfmisc_1(X2,X3),k1_xboole_0))|~v1_xboole_0(X2)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_462, plain, (r2_waybel_7(X1,X2,esk15_3(k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2)),X3,X4))|v1_xboole_0(X2)|v2_struct_0(X1)|X4=k4_xboole_0(k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2)),X3)|r2_hidden(esk15_3(k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2)),X3,X4),X4)|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))|~m1_subset_1(esk15_3(k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2)),X3,X4),u1_struct_0(X1))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_5599, plain, (r1_waybel_0(X1,X2,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))))|v1_compts_1(X1)|v2_struct_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~r1_waybel_0(X1,X2,k1_xboole_0)|~l1_waybel_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)),X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))|~v4_orders_2(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1)|~m1_subset_1(esk9_3(X1,X2,k1_xboole_0),u1_struct_0(X2)))).#trainneg
cnf(i_0_5735, plain, (r1_waybel_0(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|r2_hidden(X6,X5)|v2_struct_0(X4)|esk11_4(X1,X4,X5,X6)!=k10_yellow_6(X1,X2)|~r2_hidden(X7,esk11_4(X1,X4,X5,X6))|~m1_connsp_2(X3,X1,X7)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|X5!=k10_yellow_6(X1,X4)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X6,u1_struct_0(X1))|~l1_waybel_0(X4,X1)|~v7_waybel_0(X4)|~v4_orders_2(X4))).#trainneg
cnf(i_0_5565, plain, (r2_waybel_0(X1,X2,u1_struct_0(X3))|v1_xboole_0(u1_struct_0(X3))|v1_compts_1(X3)|v2_struct_0(k3_yellow19(X3,u1_struct_0(X3),esk4_1(X3)))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v1_xboole_0(esk4_1(X3))|~r2_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(esk4_1(X3),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X3)))))|~v3_waybel_7(esk4_1(X3),k3_yellow_1(u1_struct_0(X3)))|~v13_waybel_0(esk4_1(X3),k3_yellow_1(u1_struct_0(X3)))|~v2_waybel_0(esk4_1(X3),k3_yellow_1(u1_struct_0(X3))))).#trainneg
cnf(i_0_5593, plain, (v1_compts_1(X1)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~r2_waybel_0(X2,X3,X4)|~r1_waybel_0(X1,X3,k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))))|~l1_struct_0(X2)|~l1_waybel_0(X3,X1)|~l1_waybel_0(X3,X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_5704, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r2_hidden(X5,X4)|v2_struct_0(X3)|~r1_waybel_0(X1,X2,X6)|~l1_struct_0(X1)|~v1_xboole_0(esk11_4(X1,X3,X4,X5))|~l1_waybel_0(X2,X1)|X4!=k10_yellow_6(X1,X3)|~r3_waybel_9(X1,X2,X5)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,u1_struct_0(X1))|~l1_waybel_0(X3,X1)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_5734, plain, (r2_hidden(X1,u1_struct_0(X2))|r2_hidden(X5,X4)|v2_struct_0(X3)|v2_struct_0(X2)|~r2_hidden(X1,esk11_4(X2,X3,X4,X5))|X4!=k10_yellow_6(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_waybel_0(X3,X2)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_5626, plain, (v2_struct_0(g1_orders_2(X1,k1_xboole_0))|X1!=k1_xboole_0)).#trainneg
cnf(i_0_466, plain, (m1_subset_1(o_2_13_yellow19(X1,X2),k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2)))|v1_xboole_0(X2)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X1)))))|~v3_waybel_7(X2,k3_yellow_1(u1_struct_0(X1)))|~v13_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~v2_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_463, plain, (r2_waybel_7(X1,X2,esk15_3(X3,X4,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2))))|v1_xboole_0(X2)|v2_struct_0(X1)|k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2))=k4_xboole_0(X3,X4)|r2_hidden(esk15_3(X3,X4,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2))),X3)|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))|~m1_subset_1(esk15_3(X3,X4,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2))),u1_struct_0(X1))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_467, plain, (r2_waybel_0(X1,X2,esk13_3(X1,X3,X4))|v2_struct_0(X1)|v2_struct_0(X2)|X4=k10_yellow_6(X1,X3)|v2_struct_0(X3)|~r3_waybel_9(X1,X2,esk12_3(X1,X3,X4))|~m1_subset_1(esk12_3(X1,X3,X4),u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(esk12_3(X1,X3,X4),X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~l1_waybel_0(X3,X1)|~v7_waybel_0(X3)|~v4_orders_2(X3))).#trainneg
cnf(i_0_5733, plain, (m1_subset_1(X1,u1_struct_0(X2))|r2_hidden(X5,X4)|v2_struct_0(X3)|v2_struct_0(X2)|~r2_hidden(X1,esk11_4(X2,X3,X4,X5))|X4!=k10_yellow_6(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_waybel_0(X3,X2)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_3135, plain, (m1_subset_1(esk15_3(X1,X1,k4_xboole_0(esk4_1(X2),X3)),u1_struct_0(k3_yellow_1(u1_struct_0(X2))))|v1_compts_1(X2)|v2_struct_0(X2)|k4_xboole_0(esk4_1(X2),X3)=k1_xboole_0|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_4929, negated_conjecture, (r3_waybel_9(esk1_0,X1,X2)|v2_struct_0(X1)|r1_waybel_0(esk1_0,esk2_0,k1_xboole_0)|u1_struct_0(esk1_0)!=k1_xboole_0|~r3_waybel_9(esk1_0,esk2_0,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l1_waybel_0(X1,esk1_0)|~v1_compts_1(esk1_0))).#trainneg
cnf(i_0_938, plain, (v1_xboole_0(X1)|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X3)|X1!=k1_xboole_0|~m1_subset_1(k2_waybel_0(X2,X3,esk7_4(X4,X3,X5,esk9_3(X2,X3,X6))),X1)|~r2_waybel_0(X4,X3,X5)|~r1_waybel_0(X2,X3,X6)|~l1_struct_0(X2)|~l1_struct_0(X4)|~m1_subset_1(esk7_4(X4,X3,X5,esk9_3(X2,X3,X6)),u1_struct_0(X3))|~l1_waybel_0(X3,X2)|~l1_waybel_0(X3,X4))).#trainneg
cnf(i_0_5811, plain, (v1_xboole_0(X1)|X1!=k1_xboole_0)).#trainneg
cnf(i_0_468, plain, (m1_subset_1(esk13_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|X3=k10_yellow_6(X1,X2)|v2_struct_0(X2)|~m1_subset_1(esk12_3(X1,X2,X3),u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(esk12_3(X1,X2,X3),X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2))).#trainneg
cnf(i_0_5495, plain, (esk16_1(k1_zfmisc_1(X1))=X1|~v1_xboole_0(X1))).#trainneg
cnf(i_0_469, plain, (X1=k10_yellow_6(X2,X3)|r2_hidden(esk12_3(X2,X3,X1),X1)|r1_waybel_0(X2,X3,esk18_2(X2,esk12_3(X2,X3,X1)))|v2_struct_0(X2)|v2_struct_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_waybel_0(X3,X2)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(esk12_3(X2,X3,X1),u1_struct_0(X2)))).#trainneg
cnf(i_0_471, plain, (X1=k10_yellow_6(X2,X3)|r2_hidden(esk12_3(X2,X3,X1),X1)|r1_waybel_0(X2,X3,esk11_4(X2,X4,X5,esk12_3(X2,X3,X1)))|v2_struct_0(X2)|v2_struct_0(X3)|r2_hidden(esk12_3(X2,X3,X1),X5)|v2_struct_0(X4)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_waybel_0(X3,X2)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|X5!=k10_yellow_6(X2,X4)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(esk12_3(X2,X3,X1),u1_struct_0(X2))|~l1_waybel_0(X4,X2)|~v7_waybel_0(X4)|~v4_orders_2(X4))).#trainneg
cnf(i_0_5869, negated_conjecture, (r3_waybel_9(esk1_0,X1,esk14_2(esk1_0,esk2_0))|r1_waybel_0(esk1_0,esk2_0,k1_xboole_0)|v2_struct_0(X1)|u1_struct_0(esk1_0)!=k1_xboole_0|~m1_subset_1(esk14_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~l1_waybel_0(X1,esk1_0)|~v1_compts_1(esk1_0)|~v7_waybel_0(esk2_0)|~v4_orders_2(esk2_0))).#trainneg
cnf(i_0_5815, plain, (r2_waybel_7(X1,X2,o_2_13_yellow19(X1,X2))|v1_xboole_0(k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),X2)))|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X1)))))|~m1_subset_1(o_2_13_yellow19(X1,X2),u1_struct_0(X1))|~v13_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~v2_waybel_0(X2,k3_yellow_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v3_waybel_7(X2,k3_yellow_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_6030, plain, (r2_waybel_7(X1,esk4_1(X1),o_2_13_yellow19(X1,esk4_1(X1)))|v1_xboole_0(k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))))|v1_xboole_0(esk4_1(X1))|v2_struct_0(X1)|v1_compts_1(X1)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~m1_subset_1(o_2_13_yellow19(X1,esk4_1(X1)),u1_struct_0(X1))|~v3_waybel_7(esk4_1(X1),k3_yellow_1(u1_struct_0(X1)))|~v13_waybel_0(esk4_1(X1),k3_yellow_1(u1_struct_0(X1)))|~v2_waybel_0(esk4_1(X1),k3_yellow_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_5898, plain, (r2_hidden(X1,k10_yellow_6(X2,X3))|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X3)|k10_yellow_6(X2,X3)!=k10_yellow_6(X2,X4)|~r3_waybel_9(X2,X4,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v1_yellow19(X4,X2)|~l1_waybel_0(X4,X2)|~v7_waybel_0(X4)|~v4_orders_2(X4)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~l1_waybel_0(X3,X2)|~v7_waybel_0(X3)|~v4_orders_2(X3))).#trainneg
cnf(i_0_5912, negated_conjecture, (v4_orders_2(esk2_0)|u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))!=k1_xboole_0)).#trainneg
cnf(i_0_486, plain, (r2_hidden(esk15_3(k4_xboole_0(X1,X2),X3,X4),X1)|X4=k4_xboole_0(k4_xboole_0(X1,X2),X3)|r2_hidden(esk15_3(k4_xboole_0(X1,X2),X3,X4),X4))).#trainneg
cnf(i_0_5908, plain, (k1_zfmisc_1(X1)!=k1_xboole_0)).#trainneg
cnf(i_0_640, plain, (k2_struct_0(X1)=k2_struct_0(X2)|~l1_struct_0(X2)|~v2_struct_0(X2)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_4610, plain, (v4_relat_1(k3_subset_1(k1_xboole_0,X1),X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~l1_struct_0(X3)|~v2_struct_0(X3))).#trainneg
cnf(i_0_472, plain, (X1=k10_yellow_6(X2,X3)|r2_hidden(esk12_3(X2,X3,X1),X1)|r1_waybel_0(X2,X3,esk6_3(X2,X4,esk12_3(X2,X3,X1)))|v2_struct_0(X2)|v2_struct_0(X3)|r3_waybel_9(X2,X4,esk12_3(X2,X3,X1))|v2_struct_0(X4)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_waybel_0(X3,X2)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(esk12_3(X2,X3,X1),u1_struct_0(X2))|~l1_waybel_0(X4,X2))).#trainneg
cnf(i_0_473, plain, (k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2))=k4_xboole_0(X3,X4)|r2_hidden(esk15_3(X3,X4,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2))),X4)|v1_xboole_0(X2)|v2_struct_0(X1)|~r2_hidden(esk15_3(X3,X4,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2))),X3)|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~r2_waybel_7(X1,X2,esk15_3(X3,X4,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))|~m1_subset_1(esk15_3(X3,X4,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2))),u1_struct_0(X1))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_6040, plain, (k10_yellow_6(X1,k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))=k1_xboole_0|v1_compts_1(X1)|v2_struct_0(X1)|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~m1_subset_1(o_2_13_yellow19(X1,esk4_1(X1)),u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_6064, plain, (r2_hidden(esk14_2(X1,X2),k10_yellow_6(X1,X3))|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|k10_yellow_6(X1,X3)!=k10_yellow_6(X1,X2)|~m1_subset_1(esk14_2(X1,X2),u1_struct_0(X1))|~v1_yellow19(X2,X1)|~l1_waybel_0(X2,X1)|~l1_waybel_0(X3,X1)|~v7_waybel_0(X2)|~v7_waybel_0(X3)|~v4_orders_2(X2)|~v4_orders_2(X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_compts_1(X1))).#trainneg
cnf(i_0_5902, plain, (v1_compts_1(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|esk4_1(X1)!=k1_xboole_0)).#trainneg
cnf(i_0_6011, plain, (k10_yellow_6(X1,X2)=k1_xboole_0|r2_hidden(esk12_3(X1,X2,k1_xboole_0),X3)|r1_waybel_0(X1,X2,esk11_4(X1,X4,X3,esk12_3(X1,X2,k1_xboole_0)))|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X1)|X3!=k10_yellow_6(X1,X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_waybel_0(X2,X1)|~l1_waybel_0(X4,X1)|~v7_waybel_0(X2)|~v7_waybel_0(X4)|~v4_orders_2(X2)|~v4_orders_2(X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_5900, plain, (r2_hidden(X1,X2)|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X5)|X2!=k10_yellow_6(X5,X3)|~r3_waybel_9(X5,X4,X1)|~r1_waybel_0(X5,X4,X6)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X5)))|~m1_subset_1(X1,u1_struct_0(X5))|~l1_waybel_0(X4,X5)|~l1_waybel_0(X3,X5)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X5)|~v2_pre_topc(X5)|esk11_4(X5,X3,X2,X1)!=k1_xboole_0)).#trainneg
cnf(i_0_5870, negated_conjecture, (r3_waybel_9(esk1_0,X1,esk14_2(esk1_0,X2))|r1_waybel_0(esk1_0,esk2_0,k1_xboole_0)|v2_struct_0(X1)|v2_struct_0(X2)|u1_struct_0(esk1_0)!=k1_xboole_0|~m1_subset_1(esk14_2(esk1_0,X2),u1_struct_0(esk1_0))|~l1_waybel_0(X1,esk1_0)|~v1_compts_1(esk1_0)|~v1_yellow19(X2,esk1_0)|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(X2)|~v4_orders_2(X2))).#trainneg
cnf(i_0_5887, plain, (r2_hidden(X1,esk11_4(X2,X3,X4,X5))|v2_struct_0(X6)|v2_struct_0(X2)|r2_hidden(X5,X4)|v2_struct_0(X3)|esk11_4(X2,X3,X4,X5)!=k10_yellow_6(X2,X6)|~r3_waybel_9(X2,X6,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v1_yellow19(X6,X2)|~l1_waybel_0(X6,X2)|~v7_waybel_0(X6)|~v4_orders_2(X6)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|X4!=k10_yellow_6(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_waybel_0(X3,X2)|~v7_waybel_0(X3)|~v4_orders_2(X3))).#trainneg
cnf(i_0_5790, plain, (v1_compts_1(X1)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))|~v1_subset_1(esk4_1(X1),u1_struct_0(k3_yellow_1(u1_struct_0(X1))))|~r2_waybel_0(X3,X2,X4)|~l1_struct_0(X3)|~l1_waybel_0(X2,X1)|~l1_waybel_0(X2,X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~r1_waybel_0(X1,X2,k1_xboole_0)|~l1_waybel_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)),X1)|~v7_waybel_0(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1)))|~v4_orders_2(k3_yellow19(X1,u1_struct_0(X1),esk4_1(X1))))).#trainneg
cnf(i_0_6012, plain, (k10_yellow_6(X1,X2)=k10_yellow_6(X1,X3)|r2_hidden(esk12_3(X1,X3,k10_yellow_6(X1,X2)),X4)|r2_hidden(esk12_3(X1,X3,k10_yellow_6(X1,X2)),k10_yellow_6(X1,X2))|r1_waybel_0(X1,X3,esk11_4(X1,X5,X4,esk12_3(X1,X3,k10_yellow_6(X1,X2))))|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|X4!=k10_yellow_6(X1,X5)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~l1_waybel_0(X3,X1)|~l1_waybel_0(X5,X1)|~v7_waybel_0(X3)|~v7_waybel_0(X5)|~v4_orders_2(X3)|~v4_orders_2(X5)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2))).#trainneg
cnf(i_0_5917, plain, (v2_struct_0(X1)|~l1_orders_2(X1)|k2_struct_0(X1)!=k1_xboole_0)).#trainneg
cnf(i_0_3346, plain, (r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_5904, plain, (~l1_orders_2(X1)|~r2_hidden(X2,u1_orders_2(X1))|k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))!=k1_xboole_0)).#trainneg
cnf(i_0_487, plain, (r2_hidden(esk15_3(X1,X2,k4_xboole_0(X3,X4)),X3)|k4_xboole_0(X3,X4)=k4_xboole_0(X1,X2)|r2_hidden(esk15_3(X1,X2,k4_xboole_0(X3,X4)),X1))).#trainneg
cnf(i_0_6516, negated_conjecture, (k2_struct_0(esk2_0)!=k1_xboole_0)).#trainneg
cnf(i_0_474, plain, (X1=k4_xboole_0(X2,k10_yellow_6(X3,k3_yellow19(X3,k2_struct_0(X3),X4)))|r2_hidden(esk15_3(X2,k10_yellow_6(X3,k3_yellow19(X3,k2_struct_0(X3),X4)),X1),X1)|v1_xboole_0(X4)|v2_struct_0(X3)|~v1_subset_1(X4,u1_struct_0(k3_yellow_1(k2_struct_0(X3))))|~r2_waybel_7(X3,X4,esk15_3(X2,k10_yellow_6(X3,k3_yellow19(X3,k2_struct_0(X3),X4)),X1))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X3)))))|~m1_subset_1(esk15_3(X2,k10_yellow_6(X3,k3_yellow19(X3,k2_struct_0(X3),X4)),X1),u1_struct_0(X3))|~v13_waybel_0(X4,k3_yellow_1(k2_struct_0(X3)))|~v2_waybel_0(X4,k3_yellow_1(k2_struct_0(X3)))|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_477, plain, (m1_subset_1(X1,k10_yellow_6(X2,k3_yellow19(X2,k2_struct_0(X2),X3)))|v1_xboole_0(X3)|v2_struct_0(X2)|~v1_subset_1(X3,u1_struct_0(k3_yellow_1(k2_struct_0(X2))))|~r2_waybel_7(X2,X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X2)))))|~m1_subset_1(X1,u1_struct_0(X2))|~v13_waybel_0(X3,k3_yellow_1(k2_struct_0(X2)))|~v2_waybel_0(X3,k3_yellow_1(k2_struct_0(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_6466, plain, (r2_hidden(esk14_2(X1,X2),u1_struct_0(X1))|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|~l1_waybel_0(X3,X1)|~v7_waybel_0(X3)|~v4_orders_2(X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|k10_yellow_6(X1,X3)!=k10_yellow_6(X1,X2)|~v1_yellow19(X2,X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~v1_compts_1(X1))).#trainneg
cnf(i_0_6416, negated_conjecture, (v1_xboole_0(esk4_1(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk4_1(esk1_0)))|~m1_subset_1(esk4_1(esk1_0),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0)))))|~v3_waybel_7(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~v13_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~v2_waybel_0(esk4_1(esk1_0),k3_yellow_1(u1_struct_0(esk1_0)))|~l1_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk4_1(esk1_0)),esk1_0)|~v7_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk4_1(esk1_0)))|~v4_orders_2(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk4_1(esk1_0)))|~v1_subset_1(esk4_1(esk1_0),u1_struct_0(k3_yellow_1(u1_struct_0(esk1_0))))|~m1_subset_1(o_2_13_yellow19(esk1_0,esk4_1(esk1_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6526, plain, (k2_struct_0(k3_yellow_1(X1))!=k1_xboole_0)).#trainneg
cnf(i_0_6445, plain, (r1_waybel_0(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,k2_struct_0(X3),esk4_1(X3))))|v1_xboole_0(esk4_1(X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v1_compts_1(X3)|~v1_subset_1(esk4_1(X3),u1_struct_0(k3_yellow_1(k2_struct_0(X3))))|~v1_subset_1(esk4_1(X3),u1_struct_0(k3_yellow_1(u1_struct_0(X3))))|~r1_waybel_0(X1,X2,k1_xboole_0)|~l1_struct_0(X1)|~m1_subset_1(k2_waybel_0(X1,X2,esk10_4(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,k2_struct_0(X3),esk4_1(X3))),esk9_3(X1,X2,k1_xboole_0))),u1_struct_0(X3))|~m1_subset_1(esk10_4(X1,X2,k10_yellow_6(X3,k3_yellow19(X3,k2_struct_0(X3),esk4_1(X3))),esk9_3(X1,X2,k1_xboole_0)),u1_struct_0(X2))|~m1_subset_1(esk4_1(X3),k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X3)))))|~m1_subset_1(esk4_1(X3),k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X3)))))|~v13_waybel_0(esk4_1(X3),k3_yellow_1(k2_struct_0(X3)))|~v13_waybel_0(esk4_1(X3),k3_yellow_1(u1_struct_0(X3)))|~v2_waybel_0(esk4_1(X3),k3_yellow_1(k2_struct_0(X3)))|~v2_waybel_0(esk4_1(X3),k3_yellow_1(u1_struct_0(X3)))|~l1_waybel_0(X2,X1)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(o_2_13_yellow19(X3,esk4_1(X3)),u1_struct_0(X3)))).#trainneg
cnf(i_0_6488, plain, (r2_hidden(esk14_2(X1,X2),esk11_4(X1,X3,X4,X5))|r2_hidden(X5,X4)|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|esk11_4(X1,X3,X4,X5)!=k10_yellow_6(X1,X2)|X4!=k10_yellow_6(X1,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk14_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~v1_yellow19(X2,X1)|~l1_waybel_0(X2,X1)|~l1_waybel_0(X3,X1)|~v7_waybel_0(X2)|~v7_waybel_0(X3)|~v4_orders_2(X2)|~v4_orders_2(X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_compts_1(X1))).#trainneg
cnf(i_0_1545, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k2_struct_0(g1_orders_2(X2,k1_xboole_0))))).#trainneg
cnf(i_0_5920, plain, (X1=k2_struct_0(X2)|~l1_struct_0(X2)|~v2_struct_0(X2)|X1!=k1_xboole_0)).#trainneg
cnf(i_0_488, plain, (r2_hidden(k2_waybel_0(X1,X2,esk7_4(X1,X2,k4_xboole_0(X3,X4),X5)),X3)|v2_struct_0(X1)|v2_struct_0(X2)|~r2_waybel_0(X1,X2,k4_xboole_0(X3,X4))|~l1_struct_0(X1)|~m1_subset_1(X5,u1_struct_0(X2))|~l1_waybel_0(X2,X1))).#trainneg
cnf(i_0_479, plain, (r2_hidden(X1,k10_yellow_6(X2,k3_yellow19(X2,u1_struct_0(X2),X3)))|v1_xboole_0(X3)|v2_struct_0(X2)|~v1_subset_1(X3,u1_struct_0(k3_yellow_1(u1_struct_0(X2))))|~r2_waybel_7(X2,X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(u1_struct_0(X2)))))|~m1_subset_1(X1,u1_struct_0(X2))|~v13_waybel_0(X3,k3_yellow_1(u1_struct_0(X2)))|~v2_waybel_0(X3,k3_yellow_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~l1_struct_0(X2))).#trainneg
cnf(i_0_491, plain, (m1_subset_1(esk15_3(k10_yellow_6(X1,X2),X3,X4),u1_struct_0(X1))|v2_struct_0(X2)|v2_struct_0(X1)|X4=k4_xboole_0(k10_yellow_6(X1,X2),X3)|r2_hidden(esk15_3(k10_yellow_6(X1,X2),X3,X4),X4)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 106
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 204
# Removed in clause preprocessing      : 13
# Initial clauses in saturation        : 191
# Processed clauses                    : 1561
# ...of these trivial                  : 26
# ...subsumed                          : 620
# ...remaining for further processing  : 915
# Other redundant clauses eliminated   : 18
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 107
# Backward-rewritten                   : 15
# Generated clauses                    : 6659
# ...of the previous two non-trivial   : 5759
# Contextual simplify-reflections      : 263
# Paramodulations                      : 6574
# Factorizations                       : 6
# Equation resolutions                 : 77
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 791
#    Positive orientable unit clauses  : 60
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 13
#    Non-unit-clauses                  : 718
# Current number of unprocessed clauses: 4185
# ...number of literals in the above   : 33758
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1039
# Proof object given clauses           : 123
# Proof search given clauses           : 915
# Clause-clause subsumption calls (NU) : 460266
# Rec. Clause-clause subsumption calls : 12841
# Non-unit clause-clause subsumptions  : 874
# Unit Clause-clause subsumption calls : 6178
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 67
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 207948

# -------------------------------------------------
# User time                : 17.906 s
# System time              : 0.497 s
# Total time               : 18.403 s
# Maximum resident set size: 358700 pages
17.90user 0.56system 0:18.85elapsed 97%CPU (0avgtext+0avgdata 359132maxresident)k
0inputs+600outputs (0major+176441minor)pagefaults 0swaps
