2020-06-02 12:59:34.966546: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_5_bb1/premsel_enigma_01_2020_T30_loop02_epoch_137
2020-06-02 12:59:34.997099: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-06-02 12:59:35.058146: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-06-02 12:59:35.399086: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-06-02 12:59:36.040872: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 1074283 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_5_bb1/premsel_enigma_01_2020_T30_loop02_epoch_137' loaded (query=256; context=768; weigths=real; len_mult=0.000000).
# Preprocessing time       : 1.621 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(redefinition_k1_real_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>k1_real_1(X1)=k4_xcmplx_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', redefinition_k1_real_1)).
fof(involutiveness_k1_real_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>k1_real_1(k1_real_1(X1))=X1), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', involutiveness_k1_real_1)).
fof(dt_k8_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>m1_subset_1(k8_real_1(X1,X2),k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', dt_k8_real_1)).
fof(redefinition_k8_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', redefinition_k8_real_1)).
fof(cc17_membered, axiom, ![X1]:(v6_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v7_ordinal1(X2))), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', cc17_membered)).
fof(fc6_membered, axiom, v6_membered(k4_ordinal1), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', fc6_membered)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', redefinition_k5_numbers)).
fof(dt_k1_real_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>m1_subset_1(k1_real_1(X1),k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', dt_k1_real_1)).
fof(t3_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k3_xcmplx_0(np__1,X1)=X1), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', t3_arithm)).
fof(fc4_int_1, axiom, ![X1, X2]:((v1_int_1(X1)&v1_int_1(X2))=>v1_int_1(k6_xcmplx_0(X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', fc4_int_1)).
fof(cc2_int_1, axiom, ![X1]:(v7_ordinal1(X1)=>v1_int_1(X1)), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', cc2_int_1)).
fof(spc2_numerals, axiom, (((v2_xxreal_0(np__2)&m2_subset_1(np__2,k1_numbers,k5_numbers))&m1_subset_1(np__2,k5_numbers))&m1_subset_1(np__2,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', spc2_numerals)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', spc1_numerals)).
fof(dt_k4_xcmplx_0, axiom, ![X1]:(v1_xcmplx_0(X1)=>v1_xcmplx_0(k4_xcmplx_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', dt_k4_xcmplx_0)).
fof(cc3_int_1, axiom, ![X1]:(v1_int_1(X1)=>v1_xreal_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', cc3_int_1)).
fof(rqRealDiff__k6_xcmplx_0__r1_r2_rm1, axiom, k6_xcmplx_0(np__1,np__2)=k4_xcmplx_0(np__1), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', rqRealDiff__k6_xcmplx_0__r1_r2_rm1)).
fof(cc12_membered, axiom, ![X1]:(v1_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xcmplx_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', cc12_membered)).
fof(cc21_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v3_membered(X2))), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', cc21_membered)).
fof(cc5_membered, axiom, ![X1]:(v3_membered(X1)=>v1_membered(X1)), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', cc5_membered)).
fof(dt_k5_numbers, axiom, m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', dt_k5_numbers)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', fc3_membered)).
fof(rqRealNeg__k4_xcmplx_0__rm1_r1, axiom, k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1, file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', rqRealNeg__k4_xcmplx_0__rm1_r1)).
fof(spc2_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', spc2_arithm)).
fof(commutativity_k3_xcmplx_0, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', commutativity_k3_xcmplx_0)).
fof(spc7_arithm, axiom, ![X1, X2, X3]:(((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))&v1_xcmplx_0(X3))=>k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', spc7_arithm)).
fof(involutiveness_k4_xcmplx_0, axiom, ![X1]:(v1_xcmplx_0(X1)=>k4_xcmplx_0(k4_xcmplx_0(X1))=X1), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', involutiveness_k4_xcmplx_0)).
fof(t9_fdiff_1, axiom, ![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>![X2]:((v1_funct_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(r2_fdiff_1(X2,X1)<=>(r1_tarski(X1,k1_relset_1(k1_numbers,X2))&![X3]:(m1_subset_1(X3,k1_numbers)=>(r2_hidden(X3,X1)=>r1_fdiff_1(X2,X3))))))), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', t9_fdiff_1)).
fof(t13_fdiff_2, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>![X2]:((v1_funct_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>((r1_fdiff_1(X2,X1)&r1_fdiff_1(X3,k1_seq_1(X2,X1)))=>(r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),X1)&k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),X1)=k8_real_1(k1_fdiff_1(X3,k1_seq_1(X2,X1)),k1_fdiff_1(X2,X1))))))), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', t13_fdiff_2)).
fof(cc14_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xreal_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', cc14_membered)).
fof(dt_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:((((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&v1_funct_1(X6))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))=>(v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))&m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4))))), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', dt_k1_partfun1)).
fof(t36_fdiff_7, conjecture, ![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>(r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))=>(r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),X1)&![X2]:(m1_subset_1(X2,k1_numbers)=>(r2_hidden(X2,X1)=>k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),X1),X2)=k1_real_1(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,X2)),k1_seq_1(k16_sin_cos,X2)))))))), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', t36_fdiff_7)).
fof(t65_sin_cos, axiom, ![X1]:(v1_xreal_0(X1)=>(r1_fdiff_1(k24_sin_cos,X1)&k1_fdiff_1(k24_sin_cos,X1)=k1_seq_1(k24_sin_cos,X1))), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', t65_sin_cos)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', t4_subset)).
fof(dt_k24_sin_cos, axiom, ((v1_funct_1(k24_sin_cos)&v1_funct_2(k24_sin_cos,k1_numbers,k1_numbers))&m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', dt_k24_sin_cos)).
fof(t63_sin_cos, axiom, ![X1]:(v1_xreal_0(X1)=>(r1_fdiff_1(k19_sin_cos,X1)&k1_fdiff_1(k19_sin_cos,X1)=k1_real_1(k1_seq_1(k16_sin_cos,X1)))), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', t63_sin_cos)).
fof(cc18_valued_0, axiom, ![X1, X2]:(v3_membered(X2)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v3_valued_0(X3))), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', cc18_valued_0)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', cc1_relset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', t3_subset)).
fof(redefinition_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:((((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&v1_funct_1(X6))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))=>k1_partfun1(X1,X2,X3,X4,X5,X6)=k3_relat_1(X5,X6)), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', redefinition_k1_partfun1)).
fof(dt_k19_sin_cos, axiom, ((v1_funct_1(k19_sin_cos)&v1_funct_2(k19_sin_cos,k1_numbers,k1_numbers))&m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', dt_k19_sin_cos)).
fof(dt_k1_seq_1, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v3_valued_0(X1))=>m1_subset_1(k1_seq_1(X1,X2),k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', dt_k1_seq_1)).
fof(dt_k1_fdiff_1, axiom, ![X1, X2]:(((v1_funct_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))&v1_xreal_0(X2))=>m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', dt_k1_fdiff_1)).
fof(d7_fdiff_1, axiom, ![X1]:((v1_funct_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>![X2]:(r2_fdiff_1(X1,X2)=>![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(X3=k2_fdiff_1(X1,X2)<=>(k1_relset_1(k1_numbers,X3)=X2&![X4]:(m1_subset_1(X4,k1_numbers)=>(r2_hidden(X4,X2)=>k1_seq_1(X3,X4)=k1_fdiff_1(X1,X4)))))))), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', d7_fdiff_1)).
fof(dt_k2_fdiff_1, axiom, ![X1, X2]:((v1_funct_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(v1_funct_1(k2_fdiff_1(X1,X2))&m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', dt_k2_fdiff_1)).
fof(dt_k16_sin_cos, axiom, ((v1_funct_1(k16_sin_cos)&v1_funct_2(k16_sin_cos,k1_numbers,k1_numbers))&m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), file('/scratch/mptp/allpred/preds__0.005/t36_fdiff_7', dt_k16_sin_cos)).
fof(c_0_45, plain, ![X34]:(~m1_subset_1(X34,k1_numbers)|k1_real_1(X34)=k4_xcmplx_0(X34)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_real_1])])).
fof(c_0_46, plain, ![X72]:(~m1_subset_1(X72,k1_numbers)|k1_real_1(k1_real_1(X72))=X72), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k1_real_1])])).
fof(c_0_47, plain, ![X48, X49]:(~m1_subset_1(X48,k1_numbers)|~v1_xreal_0(X49)|m1_subset_1(k8_real_1(X48,X49),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_real_1])])).
fof(c_0_48, plain, ![X51, X52]:(~m1_subset_1(X51,k1_numbers)|~v1_xreal_0(X52)|k8_real_1(X51,X52)=k3_xcmplx_0(X51,X52)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_real_1])])).
fof(c_0_49, plain, ![X110, X111]:(~v6_membered(X110)|(~m1_subset_1(X111,X110)|v7_ordinal1(X111))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc17_membered])])])).
cnf(c_0_50, plain, (v6_membered(k4_ordinal1)), inference(split_conjunct,[status(thm)],[fc6_membered])).
cnf(c_0_51, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
cnf(c_0_52, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_53, plain, (k1_real_1(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_54, plain, ![X83]:(~m1_subset_1(X83,k1_numbers)|m1_subset_1(k1_real_1(X83),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_real_1])])).
cnf(c_0_55, plain, (m1_subset_1(k8_real_1(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_56, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_57, plain, ![X104]:(~v1_xcmplx_0(X104)|k3_xcmplx_0(np__1,X104)=X104), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_arithm])])).
fof(c_0_58, plain, ![X153, X154]:(~v1_int_1(X153)|~v1_int_1(X154)|v1_int_1(k6_xcmplx_0(X153,X154))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc4_int_1])])).
fof(c_0_59, plain, ![X90]:(~v7_ordinal1(X90)|v1_int_1(X90)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_int_1])])).
cnf(c_0_60, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_61, plain, (m1_subset_1(np__2,k5_numbers)), inference(split_conjunct,[status(thm)],[spc2_numerals])).
cnf(c_0_62, plain, (v6_membered(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_50])).
cnf(c_0_63, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_51])).
cnf(c_0_64, plain, (m1_subset_1(np__1,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
fof(c_0_65, plain, ![X86]:(~v1_xcmplx_0(X86)|v1_xcmplx_0(k4_xcmplx_0(X86))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_xcmplx_0])])).
cnf(c_0_66, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_52])).
cnf(c_0_67, plain, (k1_real_1(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_53])).
cnf(c_0_68, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_69, plain, (m1_subset_1(k8_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_55])).
cnf(c_0_70, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_56])).
cnf(c_0_71, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_72, plain, (m1_subset_1(np__1,k1_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
fof(c_0_73, plain, ![X103]:(~v1_int_1(X103)|v1_xreal_0(X103)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_int_1])])).
cnf(c_0_74, plain, (v1_int_1(k6_xcmplx_0(X1,X2))|~v1_int_1(X1)|~v1_int_1(X2)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_75, plain, (k6_xcmplx_0(np__1,np__2)=k4_xcmplx_0(np__1)), inference(split_conjunct,[status(thm)],[rqRealDiff__k6_xcmplx_0__r1_r2_rm1])).
cnf(c_0_76, plain, (v1_int_1(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_77, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_60])).
cnf(c_0_78, plain, (m1_subset_1(np__2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_61])).
cnf(c_0_79, plain, (v6_membered(k5_numbers)), inference(rw,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_80, plain, (m1_subset_1(np__1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_64])).
fof(c_0_81, plain, ![X57, X58]:(~v1_membered(X57)|(~m1_subset_1(X58,X57)|v1_xcmplx_0(X58))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc12_membered])])])).
fof(c_0_82, plain, ![X91, X92]:(~v3_membered(X91)|(~m1_subset_1(X92,k1_zfmisc_1(X91))|v3_membered(X92))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc21_membered])])])).
cnf(c_0_83, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_84, plain, (k4_xcmplx_0(k1_real_1(X1))=X1|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67])])).
cnf(c_0_85, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_68])).
cnf(c_0_86, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70])])).
cnf(c_0_87, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_71])).
cnf(c_0_88, plain, (m1_subset_1(np__1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_72])).
cnf(c_0_89, plain, (v1_xreal_0(X1)|~v1_int_1(X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_90, plain, (v1_int_1(k6_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[c_0_74])).
cnf(c_0_91, plain, (k6_xcmplx_0(np__1,np__2)=k4_xcmplx_0(np__1)), inference(evalgc,[status(thm)],[c_0_75])).
cnf(c_0_92, plain, (v1_int_1(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_76])).
cnf(c_0_93, plain, (v7_ordinal1(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_79])])])).
cnf(c_0_94, plain, (v7_ordinal1(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_80]), c_0_79])])])).
cnf(c_0_95, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
fof(c_0_96, plain, ![X50]:(~v3_membered(X50)|v1_membered(X50)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc5_membered])])).
cnf(c_0_97, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_98, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[dt_k5_numbers])).
cnf(c_0_99, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
cnf(c_0_100, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_83])).
cnf(c_0_101, plain, (k4_xcmplx_0(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_102, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_88])])])).
cnf(c_0_103, plain, (v1_xreal_0(X1)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_104, plain, (v1_int_1(k4_xcmplx_0(np__1))|~v1_int_1(np__2)|~v1_int_1(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91])])).
cnf(c_0_105, plain, (v1_int_1(np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93])])).
cnf(c_0_106, plain, (v1_int_1(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_94])])).
cnf(c_0_107, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_95])).
cnf(c_0_108, plain, (v1_membered(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_109, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_97])).
cnf(c_0_110, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_98])).
cnf(c_0_111, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_99])).
cnf(c_0_112, plain, (v1_xcmplx_0(X1)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101])])).
cnf(c_0_113, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1), inference(split_conjunct,[status(thm)],[rqRealNeg__k4_xcmplx_0__rm1_r1])).
cnf(c_0_114, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103])])).
cnf(c_0_115, plain, (v1_int_1(k4_xcmplx_0(np__1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_104, c_0_105]), c_0_106])])).
cnf(c_0_116, plain, (v1_xcmplx_0(np__1)|~v1_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_80])])).
cnf(c_0_117, plain, (v1_membered(X1)|~v3_membered(X1)), inference(evalgc,[status(thm)],[c_0_108])).
cnf(c_0_118, plain, (v3_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110]), c_0_111])])])).
fof(c_0_119, plain, ![X125]:(~v1_xcmplx_0(X125)|k3_xcmplx_0(X125,k4_xcmplx_0(np__1))=k4_xcmplx_0(X125)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc2_arithm])])).
fof(c_0_120, plain, ![X96, X97]:(~v1_xcmplx_0(X96)|~v1_xcmplx_0(X97)|k3_xcmplx_0(X96,X97)=k3_xcmplx_0(X97,X96)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k3_xcmplx_0])])).
cnf(c_0_121, plain, (v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_66])])).
cnf(c_0_122, plain, (k1_real_1(k4_xcmplx_0(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_66])])).
cnf(c_0_123, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1), inference(evalgc,[status(thm)],[c_0_113])).
cnf(c_0_124, plain, (m1_subset_1(k4_xcmplx_0(np__1),k1_numbers)|~v1_xcmplx_0(k4_xcmplx_0(np__1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115])])).
cnf(c_0_125, plain, (v1_xcmplx_0(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_118])])])).
cnf(c_0_126, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_127, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_128, plain, (v1_xcmplx_0(k1_real_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_101])])).
cnf(c_0_129, plain, (k1_real_1(np__1)=k4_xcmplx_0(np__1)|~m1_subset_1(k4_xcmplx_0(np__1),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123])])).
cnf(c_0_130, plain, (m1_subset_1(k4_xcmplx_0(np__1),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_100]), c_0_125])])])).
fof(c_0_131, plain, ![X132, X133, X134]:(~v1_xcmplx_0(X132)|~v1_xcmplx_0(X133)|~v1_xcmplx_0(X134)|k3_xcmplx_0(k3_xcmplx_0(X132,X133),X134)=k3_xcmplx_0(X132,k3_xcmplx_0(X133,X134))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc7_arithm])])).
cnf(c_0_132, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_126])).
cnf(c_0_133, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_127])).
cnf(c_0_134, plain, (v1_xcmplx_0(k1_real_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_128, c_0_85])).
cnf(c_0_135, plain, (k1_real_1(np__1)=k4_xcmplx_0(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_129, c_0_130])])])).
cnf(c_0_136, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)), inference(split_conjunct,[status(thm)],[c_0_131])).
cnf(c_0_137, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_133])])).
cnf(c_0_138, plain, (v1_xcmplx_0(k4_xcmplx_0(np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135]), c_0_125]), c_0_88])])])).
cnf(c_0_139, plain, (m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_66])])).
cnf(c_0_140, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_136])).
cnf(c_0_141, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_137, c_0_138])])])).
cnf(c_0_142, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_139])])).
cnf(c_0_143, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),k3_xcmplx_0(X1,X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_141]), c_0_138])])])).
cnf(c_0_144, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_117]), c_0_111])])])).
fof(c_0_145, plain, ![X93]:(~v1_xcmplx_0(X93)|k4_xcmplx_0(k4_xcmplx_0(X93))=X93), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k4_xcmplx_0])])).
cnf(c_0_146, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_143])])).
cnf(c_0_147, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_144])])).
cnf(c_0_148, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_145])).
fof(c_0_149, plain, ![X17, X18, X19]:(((r1_tarski(X17,k1_relset_1(k1_numbers,X18))|~r2_fdiff_1(X18,X17)|(~v1_funct_1(X18)|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X17)|~m1_subset_1(X17,k1_zfmisc_1(k1_numbers))))&(~m1_subset_1(X19,k1_numbers)|(~r2_hidden(X19,X17)|r1_fdiff_1(X18,X19))|~r2_fdiff_1(X18,X17)|(~v1_funct_1(X18)|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X17)|~m1_subset_1(X17,k1_zfmisc_1(k1_numbers)))))&((m1_subset_1(esk4_2(X17,X18),k1_numbers)|~r1_tarski(X17,k1_relset_1(k1_numbers,X18))|r2_fdiff_1(X18,X17)|(~v1_funct_1(X18)|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X17)|~m1_subset_1(X17,k1_zfmisc_1(k1_numbers))))&((r2_hidden(esk4_2(X17,X18),X17)|~r1_tarski(X17,k1_relset_1(k1_numbers,X18))|r2_fdiff_1(X18,X17)|(~v1_funct_1(X18)|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X17)|~m1_subset_1(X17,k1_zfmisc_1(k1_numbers))))&(~r1_fdiff_1(X18,esk4_2(X17,X18))|~r1_tarski(X17,k1_relset_1(k1_numbers,X18))|r2_fdiff_1(X18,X17)|(~v1_funct_1(X18)|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X17)|~m1_subset_1(X17,k1_zfmisc_1(k1_numbers))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t9_fdiff_1])])])])])).
fof(c_0_150, plain, ![X14, X15, X16]:((r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X15,X16),X14)|(~r1_fdiff_1(X15,X14)|~r1_fdiff_1(X16,k1_seq_1(X15,X14)))|(~v1_funct_1(X16)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v1_funct_1(X15)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~m1_subset_1(X14,k1_numbers))&(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X15,X16),X14)=k8_real_1(k1_fdiff_1(X16,k1_seq_1(X15,X14)),k1_fdiff_1(X15,X14))|(~r1_fdiff_1(X15,X14)|~r1_fdiff_1(X16,k1_seq_1(X15,X14)))|(~v1_funct_1(X16)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v1_funct_1(X15)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~m1_subset_1(X14,k1_numbers))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_fdiff_2])])])])).
fof(c_0_151, plain, ![X55, X56]:(~v3_membered(X55)|(~m1_subset_1(X56,X55)|v1_xreal_0(X56))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc14_membered])])])).
cnf(c_0_152, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_147])])).
cnf(c_0_153, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_148])).
cnf(c_0_154, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(k4_xcmplx_0(np__1),X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_132])])).
cnf(c_0_155, plain, (r2_fdiff_1(X1,X2)|~r1_fdiff_1(X1,esk4_2(X2,X1))|~r1_tarski(X2,k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_149])).
cnf(c_0_156, plain, (r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~r1_fdiff_1(X1,X3)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_150])).
fof(c_0_157, plain, ![X64, X65, X66, X67, X68, X69]:((v1_funct_1(k1_partfun1(X64,X65,X66,X67,X68,X69))|(~v1_funct_1(X68)|~m1_subset_1(X68,k1_zfmisc_1(k2_zfmisc_1(X64,X65)))|~v1_funct_1(X69)|~m1_subset_1(X69,k1_zfmisc_1(k2_zfmisc_1(X66,X67)))))&(m1_subset_1(k1_partfun1(X64,X65,X66,X67,X68,X69),k1_zfmisc_1(k2_zfmisc_1(X64,X67)))|(~v1_funct_1(X68)|~m1_subset_1(X68,k1_zfmisc_1(k2_zfmisc_1(X64,X65)))|~v1_funct_1(X69)|~m1_subset_1(X69,k1_zfmisc_1(k2_zfmisc_1(X66,X67)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_partfun1])])])).
cnf(c_0_158, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_151])).
fof(c_0_159, negated_conjecture, ~(![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>(r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))=>(r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),X1)&![X2]:(m1_subset_1(X2,k1_numbers)=>(r2_hidden(X2,X1)=>k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),X1),X2)=k1_real_1(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,X2)),k1_seq_1(k16_sin_cos,X2))))))))), inference(assume_negation,[status(cth)],[t36_fdiff_7])).
cnf(c_0_160, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(k3_xcmplx_0(k4_xcmplx_0(X1),X2),k1_numbers)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_153])])).
cnf(c_0_161, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(k4_xcmplx_0(np__1),X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_154, c_0_138])])])).
cnf(c_0_162, plain, (r2_fdiff_1(X1,X2)|~v3_rcomp_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r1_fdiff_1(X1,esk4_2(X2,X1))|~r1_tarski(X2,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_155])).
cnf(c_0_163, plain, (r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_fdiff_1(X1,X3)|~m1_subset_1(X3,k1_numbers)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_156])).
cnf(c_0_164, plain, (m1_subset_1(esk4_2(X1,X2),k1_numbers)|r2_fdiff_1(X2,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_149])).
cnf(c_0_165, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_157])).
cnf(c_0_166, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_157])).
fof(c_0_167, plain, ![X21]:((r1_fdiff_1(k24_sin_cos,X21)|~v1_xreal_0(X21))&(k1_fdiff_1(k24_sin_cos,X21)=k1_seq_1(k24_sin_cos,X21)|~v1_xreal_0(X21))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t65_sin_cos])])])).
cnf(c_0_168, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_158])).
fof(c_0_169, plain, ![X38, X39, X40]:(~r2_hidden(X38,X39)|~m1_subset_1(X39,k1_zfmisc_1(X40))|m1_subset_1(X38,X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_170, negated_conjecture, ((v3_rcomp_1(esk1_0)&m1_subset_1(esk1_0,k1_zfmisc_1(k1_numbers)))&(r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))&((m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0))&((r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0))&(k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)!=k1_real_1(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0)))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_159])])])])).
cnf(c_0_171, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k3_xcmplx_0(k4_xcmplx_0(X1),X2),k1_numbers)), inference(csr,[status(thm)],[c_0_160, c_0_100])).
cnf(c_0_172, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_141])])).
cnf(c_0_173, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))))|~r1_fdiff_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_163])])).
cnf(c_0_174, plain, (r2_fdiff_1(X2,X1)|m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_164])).
cnf(c_0_175, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_165])).
cnf(c_0_176, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_166])).
cnf(c_0_177, plain, (r1_fdiff_1(k24_sin_cos,X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_167])).
cnf(c_0_178, plain, (v1_funct_1(k24_sin_cos)), inference(split_conjunct,[status(thm)],[dt_k24_sin_cos])).
cnf(c_0_179, plain, (m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[dt_k24_sin_cos])).
cnf(c_0_180, plain, (v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_85]), c_0_111])])])).
cnf(c_0_181, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_169])).
cnf(c_0_182, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_170])).
cnf(c_0_183, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k3_xcmplx_0(X1,k4_xcmplx_0(X2)),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_172])])).
cnf(c_0_184, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_66])])).
cnf(c_0_185, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))))|~r1_fdiff_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_173, c_0_174]), c_0_175]), c_0_176])).
cnf(c_0_186, plain, (r1_fdiff_1(k24_sin_cos,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_177])).
cnf(c_0_187, plain, (v1_funct_1(k24_sin_cos)), inference(evalgc,[status(thm)],[c_0_178])).
cnf(c_0_188, plain, (m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_179])).
cnf(c_0_189, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_67])])).
fof(c_0_190, plain, ![X22]:((r1_fdiff_1(k19_sin_cos,X22)|~v1_xreal_0(X22))&(k1_fdiff_1(k19_sin_cos,X22)=k1_real_1(k1_seq_1(k16_sin_cos,X22))|~v1_xreal_0(X22))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_sin_cos])])])).
fof(c_0_191, plain, ![X41, X42, X43]:(~v3_membered(X42)|(~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(X41,X42)))|v3_valued_0(X43))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc18_valued_0])])])).
fof(c_0_192, plain, ![X29, X30, X31]:(~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(X29,X30)))|v1_relat_1(X31)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_193, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_181])).
cnf(c_0_194, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_182])).
cnf(c_0_195, plain, (r2_hidden(esk4_2(X1,X2),X1)|r2_fdiff_1(X2,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_149])).
cnf(c_0_196, negated_conjecture, (v3_rcomp_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_170])).
fof(c_0_197, plain, ![X32, X33]:((~m1_subset_1(X32,k1_zfmisc_1(X33))|r1_tarski(X32,X33))&(~r1_tarski(X32,X33)|m1_subset_1(X32,k1_zfmisc_1(X33)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_198, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_184])])).
fof(c_0_199, plain, ![X23, X24, X25, X26, X27, X28]:(~v1_funct_1(X27)|~m1_subset_1(X27,k1_zfmisc_1(k2_zfmisc_1(X23,X24)))|~v1_funct_1(X28)|~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(X25,X26)))|k1_partfun1(X23,X24,X25,X26,X27,X28)=k3_relat_1(X27,X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_partfun1])])).
cnf(c_0_200, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),X2)|~r1_fdiff_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_xreal_0(k1_seq_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_186]), c_0_187]), c_0_188])])])).
cnf(c_0_201, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_189, c_0_85])).
cnf(c_0_202, plain, (r1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_190])).
cnf(c_0_203, plain, (v1_funct_1(k19_sin_cos)), inference(split_conjunct,[status(thm)],[dt_k19_sin_cos])).
cnf(c_0_204, plain, (m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[dt_k19_sin_cos])).
fof(c_0_205, plain, ![X75, X76]:(~v1_relat_1(X75)|~v1_funct_1(X75)|~v3_valued_0(X75)|m1_subset_1(k1_seq_1(X75,X76),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_seq_1])])).
cnf(c_0_206, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_191])).
cnf(c_0_207, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_192])).
cnf(c_0_208, negated_conjecture, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_194])])).
cnf(c_0_209, plain, (r2_fdiff_1(X2,X1)|r2_hidden(esk4_2(X1,X2),X1)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_195])).
cnf(c_0_210, negated_conjecture, (v3_rcomp_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_196])).
cnf(c_0_211, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_197])).
cnf(c_0_212, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_198, c_0_144])).
cnf(c_0_213, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k3_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(split_conjunct,[status(thm)],[c_0_199])).
cnf(c_0_214, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),X2)|~r1_fdiff_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k1_seq_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos))),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_201])])).
cnf(c_0_215, plain, (r1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_202])).
cnf(c_0_216, plain, (v1_funct_1(k19_sin_cos)), inference(evalgc,[status(thm)],[c_0_203])).
cnf(c_0_217, plain, (m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_204])).
cnf(c_0_218, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_205])).
cnf(c_0_219, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_206])).
cnf(c_0_220, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_207])).
cnf(c_0_221, negated_conjecture, (m1_subset_1(esk4_2(esk1_0,X1),k1_numbers)|r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~r1_tarski(esk1_0,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_209]), c_0_194]), c_0_210])])])).
cnf(c_0_222, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_211])).
cnf(c_0_223, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_197])).
cnf(c_0_224, negated_conjecture, (r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))), inference(split_conjunct,[status(thm)],[c_0_170])).
cnf(c_0_225, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,k1_real_1(X2)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k3_xcmplx_0(X1,k1_real_1(X2)),k1_numbers)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_101])])).
cnf(c_0_226, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_170])).
cnf(c_0_227, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k3_relat_1(X1,X4)|~v1_funct_1(X4)|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_213])).
cnf(c_0_228, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~m1_subset_1(k1_seq_1(k19_sin_cos,esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xreal_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_215]), c_0_216]), c_0_217])])])).
cnf(c_0_229, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_218])).
cnf(c_0_230, plain, (v3_valued_0(k19_sin_cos)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_217]), c_0_111])])])).
cnf(c_0_231, plain, (v1_relat_1(k19_sin_cos)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_217])])).
cnf(c_0_232, negated_conjecture, (r2_fdiff_1(X1,esk1_0)|m1_subset_1(esk4_2(esk1_0,X1),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_222])])).
cnf(c_0_233, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_223])).
cnf(c_0_234, negated_conjecture, (r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))), inference(evalgc,[status(thm)],[c_0_224])).
cnf(c_0_235, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,k1_real_1(X2)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k3_xcmplx_0(X1,k1_real_1(X2)),k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_225, c_0_85])).
cnf(c_0_236, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)), inference(evalgc,[status(thm)],[c_0_226])).
cnf(c_0_237, plain, (k1_partfun1(X1,X2,X3,X4,X5,X6)=k1_partfun1(X7,X8,X9,X10,X5,X6)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_227])])).
cnf(c_0_238, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),X1)|~v1_xreal_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_229]), c_0_230]), c_0_231]), c_0_216])])])).
cnf(c_0_239, negated_conjecture, (v1_xreal_0(esk4_2(esk1_0,X1))|r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_232]), c_0_111])])])).
cnf(c_0_240, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_234])])).
cnf(c_0_241, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)!=k1_real_1(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0)))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_170])).
cnf(c_0_242, plain, (k1_real_1(k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X1,k1_real_1(X2))|~m1_subset_1(k3_xcmplx_0(X1,k1_real_1(X2)),k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_235])])).
cnf(c_0_243, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_236, c_0_237]), c_0_187]), c_0_216]), c_0_188]), c_0_217])])])).
cnf(c_0_244, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_239]), c_0_234]), c_0_194]), c_0_210]), c_0_240])])])).
cnf(c_0_245, plain, (k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)=k8_real_1(k1_fdiff_1(X2,k1_seq_1(X1,X3)),k1_fdiff_1(X1,X3))|~r1_fdiff_1(X1,X3)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_150])).
cnf(c_0_246, plain, (k1_fdiff_1(k24_sin_cos,X1)=k1_seq_1(k24_sin_cos,X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_167])).
cnf(c_0_247, negated_conjecture, (k1_real_1(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)), inference(evalgc,[status(thm)],[c_0_241])).
cnf(c_0_248, plain, (k1_real_1(k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X1,k1_real_1(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_86])])).
cnf(c_0_249, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_244]), c_0_188]), c_0_217])])])).
cnf(c_0_250, plain, (k8_real_1(k1_fdiff_1(X2,k1_seq_1(X1,X3)),k1_fdiff_1(X1,X3))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_fdiff_1(X1,X3)|~m1_subset_1(X3,k1_numbers)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_245])).
cnf(c_0_251, plain, (k1_fdiff_1(k24_sin_cos,X1)=k1_seq_1(k24_sin_cos,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_246])).
fof(c_0_252, plain, ![X70, X71]:(~v1_funct_1(X70)|~m1_subset_1(X70,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X71)|m1_subset_1(k1_fdiff_1(X70,X71),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_fdiff_1])])).
cnf(c_0_253, negated_conjecture, (k1_real_1(k3_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~v1_xreal_0(k1_seq_1(k16_sin_cos,esk2_0))|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_70])])).
cnf(c_0_254, plain, (k1_real_1(k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X1,k1_real_1(X2))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_248, c_0_180]), c_0_147])).
cnf(c_0_255, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_249, c_0_175]), c_0_187]), c_0_216]), c_0_188]), c_0_217])])])).
cnf(c_0_256, plain, (k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(X1,X2)),k1_fdiff_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),X2)|~r1_fdiff_1(k24_sin_cos,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(k1_seq_1(X1,X2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_251]), c_0_187]), c_0_188])])])).
cnf(c_0_257, plain, (m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_252])).
cnf(c_0_258, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_170])).
cnf(c_0_259, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_real_1(k1_seq_1(k16_sin_cos,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(k1_seq_1(k16_sin_cos,esk2_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_254])])).
cnf(c_0_260, plain, (k1_fdiff_1(k19_sin_cos,X1)=k1_real_1(k1_seq_1(k16_sin_cos,X1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_190])).
cnf(c_0_261, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_176]), c_0_187]), c_0_216]), c_0_188]), c_0_217])])])).
cnf(c_0_262, plain, (k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(X1,X2)),k1_fdiff_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),X2)|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_256, c_0_186])).
cnf(c_0_263, plain, (m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_257])).
fof(c_0_264, plain, ![X9, X10, X11, X12]:(((k1_relset_1(k1_numbers,X11)=X10|X11!=k2_fdiff_1(X9,X10)|(~v1_funct_1(X11)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X9,X10)|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(~m1_subset_1(X12,k1_numbers)|(~r2_hidden(X12,X10)|k1_seq_1(X11,X12)=k1_fdiff_1(X9,X12))|X11!=k2_fdiff_1(X9,X10)|(~v1_funct_1(X11)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X9,X10)|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))))&((m1_subset_1(esk3_3(X9,X10,X11),k1_numbers)|k1_relset_1(k1_numbers,X11)!=X10|X11=k2_fdiff_1(X9,X10)|(~v1_funct_1(X11)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X9,X10)|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&((r2_hidden(esk3_3(X9,X10,X11),X10)|k1_relset_1(k1_numbers,X11)!=X10|X11=k2_fdiff_1(X9,X10)|(~v1_funct_1(X11)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X9,X10)|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(k1_seq_1(X11,esk3_3(X9,X10,X11))!=k1_fdiff_1(X9,esk3_3(X9,X10,X11))|k1_relset_1(k1_numbers,X11)!=X10|X11=k2_fdiff_1(X9,X10)|(~v1_funct_1(X11)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X9,X10)|(~v1_funct_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_fdiff_1])])])])])).
cnf(c_0_265, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)), inference(evalgc,[status(thm)],[c_0_258])).
cnf(c_0_266, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_real_1(k1_seq_1(k16_sin_cos,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)), inference(csr,[status(thm)],[c_0_259, c_0_201])).
cnf(c_0_267, plain, (k1_real_1(k1_seq_1(k16_sin_cos,X1))=k1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_260])).
cnf(c_0_268, negated_conjecture, (v1_xreal_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_261]), c_0_111])])])).
cnf(c_0_269, plain, (k3_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(X1,X2)),k1_fdiff_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),X2)|~v1_xreal_0(k1_fdiff_1(X1,X2))|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(X1,X2)),k1_numbers)|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_262])])).
cnf(c_0_270, plain, (m1_subset_1(k1_seq_1(k24_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_251]), c_0_187]), c_0_188])])])).
cnf(c_0_271, plain, (k1_seq_1(X3,X1)=k1_fdiff_1(X4,X1)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,X2)|X3!=k2_fdiff_1(X4,X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_fdiff_1(X4,X2)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_264])).
fof(c_0_272, plain, ![X46, X47]:((v1_funct_1(k2_fdiff_1(X46,X47))|(~v1_funct_1(X46)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(m1_subset_1(k2_fdiff_1(X46,X47),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|(~v1_funct_1(X46)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_fdiff_1])])])).
cnf(c_0_273, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_237]), c_0_187]), c_0_216]), c_0_188]), c_0_217])])])).
cnf(c_0_274, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_fdiff_1(k19_sin_cos,esk2_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_267]), c_0_268])])])).
cnf(c_0_275, plain, (k3_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(X1,X2)),k1_fdiff_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),X2)|~v1_xreal_0(k1_fdiff_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_269, c_0_270])).
cnf(c_0_276, plain, (k1_fdiff_1(X4,X1)=k1_seq_1(X3,X1)|X3!=k2_fdiff_1(X4,X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~r2_fdiff_1(X4,X2)|~r2_hidden(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_271])).
cnf(c_0_277, plain, (m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_272])).
cnf(c_0_278, plain, (v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_272])).
cnf(c_0_279, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_244]), c_0_188]), c_0_217])])])).
cnf(c_0_280, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)|~v1_xreal_0(k1_fdiff_1(k19_sin_cos,esk2_0))|~v1_xreal_0(k1_seq_1(k19_sin_cos,esk2_0))|~r1_fdiff_1(k19_sin_cos,esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_274, c_0_275]), c_0_216]), c_0_217]), c_0_261])])])).
cnf(c_0_281, plain, (k1_seq_1(k2_fdiff_1(X1,X3),X2)=k1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(X1,X3))|~r2_hidden(X2,X3)|~r2_fdiff_1(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_276])])).
cnf(c_0_282, plain, (m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_277])).
cnf(c_0_283, plain, (v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_278])).
cnf(c_0_284, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_279, c_0_175]), c_0_187]), c_0_216]), c_0_188]), c_0_217])])])).
cnf(c_0_285, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk2_0)|~v1_xreal_0(k1_fdiff_1(k19_sin_cos,esk2_0))|~v1_xreal_0(k1_seq_1(k19_sin_cos,esk2_0))|~r1_fdiff_1(k19_sin_cos,esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)), inference(csr,[status(thm)],[c_0_280, c_0_270])).
cnf(c_0_286, plain, (k1_seq_1(k2_fdiff_1(X1,X2),X3)=k1_fdiff_1(X1,X3)|~v1_funct_1(X1)|~r2_hidden(X3,X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_281, c_0_282]), c_0_283])).
cnf(c_0_287, negated_conjecture, (r2_hidden(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_284, c_0_176]), c_0_187]), c_0_216]), c_0_188]), c_0_217])])])).
cnf(c_0_288, negated_conjecture, (~v1_xreal_0(k1_fdiff_1(k19_sin_cos,esk2_0))|~v1_xreal_0(k1_seq_1(k19_sin_cos,esk2_0))|~r1_fdiff_1(k19_sin_cos,esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_285, c_0_286]), c_0_287]), c_0_261])])])).
cnf(c_0_289, negated_conjecture, (~v1_xreal_0(k1_fdiff_1(k19_sin_cos,esk2_0))|~v1_xreal_0(k1_seq_1(k19_sin_cos,esk2_0))|~r1_fdiff_1(k19_sin_cos,esk2_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)), inference(csr,[status(thm)],[c_0_288, c_0_244])).
cnf(c_0_290, plain, (v1_xreal_0(k1_fdiff_1(X1,X2))|~v1_xreal_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_263]), c_0_111])])])).
cnf(c_0_291, negated_conjecture, (~v1_xreal_0(k1_seq_1(k19_sin_cos,esk2_0))|~r1_fdiff_1(k19_sin_cos,esk2_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_289, c_0_290]), c_0_268]), c_0_216]), c_0_217])])])).
cnf(c_0_292, negated_conjecture, (~v1_xreal_0(k1_seq_1(k19_sin_cos,esk2_0))|~r1_fdiff_1(k19_sin_cos,esk2_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_175]), c_0_187]), c_0_216]), c_0_188]), c_0_217])])])).
cnf(c_0_293, plain, (v1_xreal_0(k1_seq_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_229]), c_0_111])])])).
cnf(c_0_294, negated_conjecture, (~r1_fdiff_1(k19_sin_cos,esk2_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_293]), c_0_230]), c_0_231]), c_0_216])])])).
cnf(c_0_295, plain, (m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[dt_k16_sin_cos])).
cnf(c_0_296, negated_conjecture, (~r1_fdiff_1(k19_sin_cos,esk2_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_294, c_0_176]), c_0_187]), c_0_216]), c_0_188]), c_0_217])])])).
cnf(c_0_297, plain, (m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_295])).
cnf(c_0_298, plain, (v1_funct_1(k16_sin_cos)), inference(split_conjunct,[status(thm)],[dt_k16_sin_cos])).
cnf(c_0_299, negated_conjecture, (~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_296, c_0_215]), c_0_268])])])).
cnf(c_0_300, plain, (v3_valued_0(k16_sin_cos)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_297]), c_0_111])])])).
cnf(c_0_301, plain, (v1_relat_1(k16_sin_cos)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_297])])).
cnf(c_0_302, plain, (v1_funct_1(k16_sin_cos)), inference(evalgc,[status(thm)],[c_0_298])).
cnf(c_0_303, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_299, c_0_229]), c_0_300]), c_0_301]), c_0_302])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 304
# Proof object clause steps            : 224
# Proof object formula steps           : 80
# Proof object conjectures             : 44
# Proof object clause conjectures      : 41
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 62
# Proof object initial formulas used   : 45
# Proof object generating inferences   : 81
# Proof object simplifying inferences  : 154
# Training examples: 145 positive, 1139 negative
# Training: Positive examples begin
cnf(c_0_63, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_210, negated_conjecture, (v3_rcomp_1(esk1_0))).# trainpos
cnf(c_0_216, plain, (v1_funct_1(k19_sin_cos))).# trainpos
cnf(c_0_187, plain, (v1_funct_1(k24_sin_cos))).# trainpos
cnf(c_0_302, plain, (v1_funct_1(k16_sin_cos))).# trainpos
cnf(c_0_111, plain, (v3_membered(k1_numbers))).# trainpos
cnf(c_0_194, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_222, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_247, negated_conjecture, (k1_real_1(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0))).# trainpos
cnf(c_0_103, plain, (v1_xreal_0(X1)|~v1_int_1(X1))).# trainpos
cnf(c_0_62, plain, (v6_membered(k4_ordinal1))).# trainpos
cnf(c_0_123, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1)).# trainpos
cnf(c_0_88, plain, (m1_subset_1(np__1,k1_numbers))).# trainpos
cnf(c_0_80, plain, (m1_subset_1(np__1,k5_numbers))).# trainpos
cnf(c_0_78, plain, (m1_subset_1(np__2,k5_numbers))).# trainpos
cnf(c_0_236, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0))).# trainpos
cnf(c_0_110, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_117, plain, (v1_membered(X1)|~v3_membered(X1))).# trainpos
cnf(c_0_92, plain, (v1_int_1(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_91, plain, (k6_xcmplx_0(np__1,np__2)=k4_xcmplx_0(np__1))).# trainpos
cnf(c_0_100, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_233, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_265, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0))).# trainpos
cnf(c_0_217, plain, (m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_153, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_215, plain, (r1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_188, plain, (m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_297, plain, (m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_186, plain, (r1_fdiff_1(k24_sin_cos,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_87, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_66, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_67, plain, (k1_real_1(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_132, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_251, plain, (k1_fdiff_1(k24_sin_cos,X1)=k1_seq_1(k24_sin_cos,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_85, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_168, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_107, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_77, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_193, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_234, negated_conjecture, (r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))))).# trainpos
cnf(c_0_174, plain, (r2_fdiff_1(X2,X1)|m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_109, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_69, plain, (m1_subset_1(k8_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_229, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).# trainpos
cnf(c_0_133, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_267, plain, (k1_real_1(k1_seq_1(k16_sin_cos,X1))=k1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_220, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_90, plain, (v1_int_1(k6_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).# trainpos
cnf(c_0_70, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_219, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).# trainpos
cnf(c_0_283, plain, (v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_162, plain, (r2_fdiff_1(X1,X2)|~v3_rcomp_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r1_fdiff_1(X1,esk4_2(X2,X1))|~r1_tarski(X2,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_140, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_263, plain, (m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_282, plain, (m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_176, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_209, plain, (r2_fdiff_1(X2,X1)|r2_hidden(esk4_2(X1,X2),X1)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_276, plain, (k1_fdiff_1(X4,X1)=k1_seq_1(X3,X1)|X3!=k2_fdiff_1(X4,X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~r2_fdiff_1(X4,X2)|~r2_hidden(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_163, plain, (r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_fdiff_1(X1,X3)|~m1_subset_1(X3,k1_numbers)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_227, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k3_relat_1(X1,X4)|~v1_funct_1(X4)|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_175, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_250, plain, (k8_real_1(k1_fdiff_1(X2,k1_seq_1(X1,X3)),k1_fdiff_1(X1,X3))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_fdiff_1(X1,X3)|~m1_subset_1(X3,k1_numbers)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_301, plain, (v1_relat_1(k16_sin_cos))).# trainpos
cnf(c_0_231, plain, (v1_relat_1(k19_sin_cos))).# trainpos
cnf(c_0_270, plain, (m1_subset_1(k1_seq_1(k24_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_208, negated_conjecture, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk1_0))).# trainpos
cnf(c_0_300, plain, (v3_valued_0(k16_sin_cos))).# trainpos
cnf(c_0_253, negated_conjecture, (k1_real_1(k3_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~v1_xreal_0(k1_seq_1(k16_sin_cos,esk2_0))|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers))).# trainpos
cnf(c_0_180, plain, (v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_139, plain, (m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_93, plain, (v7_ordinal1(np__2))).# trainpos
cnf(c_0_94, plain, (v7_ordinal1(np__1))).# trainpos
cnf(c_0_230, plain, (v3_valued_0(k19_sin_cos))).# trainpos
cnf(c_0_84, plain, (k4_xcmplx_0(k1_real_1(X1))=X1|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_86, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_293, plain, (v1_xreal_0(k1_seq_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).# trainpos
cnf(c_0_118, plain, (v3_membered(k5_numbers))).# trainpos
cnf(c_0_122, plain, (k1_real_1(k4_xcmplx_0(X1))=X1|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_240, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))))).# trainpos
cnf(c_0_116, plain, (v1_xcmplx_0(np__1)|~v1_membered(k5_numbers))).# trainpos
cnf(c_0_102, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_281, plain, (k1_seq_1(k2_fdiff_1(X1,X3),X2)=k1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(X1,X3))|~r2_hidden(X2,X3)|~r2_fdiff_1(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_221, negated_conjecture, (m1_subset_1(esk4_2(esk1_0,X1),k1_numbers)|r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~r1_tarski(esk1_0,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_105, plain, (v1_int_1(np__2))).# trainpos
cnf(c_0_106, plain, (v1_int_1(np__1))).# trainpos
cnf(c_0_125, plain, (v1_xcmplx_0(np__1))).# trainpos
cnf(c_0_237, plain, (k1_partfun1(X1,X2,X3,X4,X5,X6)=k1_partfun1(X7,X8,X9,X10,X5,X6)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).# trainpos
cnf(c_0_173, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))))|~r1_fdiff_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)),k1_numbers))).# trainpos
cnf(c_0_256, plain, (k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(X1,X2)),k1_fdiff_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),X2)|~r1_fdiff_1(k24_sin_cos,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(k1_seq_1(X1,X2)))).# trainpos
cnf(c_0_189, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_114, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~v1_int_1(X1))).# trainpos
cnf(c_0_232, negated_conjecture, (r2_fdiff_1(X1,esk1_0)|m1_subset_1(esk4_2(esk1_0,X1),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).# trainpos
cnf(c_0_243, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_273, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_200, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),X2)|~r1_fdiff_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_xreal_0(k1_seq_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))))).# trainpos
cnf(c_0_129, plain, (k1_real_1(np__1)=k4_xcmplx_0(np__1)|~m1_subset_1(k4_xcmplx_0(np__1),k1_numbers))).# trainpos
cnf(c_0_239, negated_conjecture, (v1_xreal_0(esk4_2(esk1_0,X1))|r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).# trainpos
cnf(c_0_137, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).# trainpos
cnf(c_0_214, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),X2)|~r1_fdiff_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k1_seq_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos))),k1_numbers))).# trainpos
cnf(c_0_269, plain, (k3_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(X1,X2)),k1_fdiff_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),X2)|~v1_xreal_0(k1_fdiff_1(X1,X2))|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(X1,X2)),k1_numbers)|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_228, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~m1_subset_1(k1_seq_1(k19_sin_cos,esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xreal_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))))).# trainpos
cnf(c_0_104, plain, (v1_int_1(k4_xcmplx_0(np__1))|~v1_int_1(np__2)|~v1_int_1(np__1))).# trainpos
cnf(c_0_238, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),X1)|~v1_xreal_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).# trainpos
cnf(c_0_112, plain, (v1_xcmplx_0(X1)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_244, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_121, plain, (v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_279, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_249, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_154, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(k4_xcmplx_0(np__1),X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_284, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))).# trainpos
cnf(c_0_255, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))).# trainpos
cnf(c_0_290, plain, (v1_xreal_0(k1_fdiff_1(X1,X2))|~v1_xreal_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_128, plain, (v1_xcmplx_0(k1_real_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_142, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_184, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_287, negated_conjecture, (r2_hidden(esk2_0,esk1_0))).# trainpos
cnf(c_0_261, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers))).# trainpos
cnf(c_0_144, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_124, plain, (m1_subset_1(k4_xcmplx_0(np__1),k1_numbers)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).# trainpos
cnf(c_0_268, negated_conjecture, (v1_xreal_0(esk2_0))).# trainpos
cnf(c_0_147, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_130, plain, (m1_subset_1(k4_xcmplx_0(np__1),k1_numbers))).# trainpos
cnf(c_0_135, plain, (k1_real_1(np__1)=k4_xcmplx_0(np__1))).# trainpos
cnf(c_0_138, plain, (v1_xcmplx_0(k4_xcmplx_0(np__1)))).# trainpos
cnf(c_0_141, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_161, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(k4_xcmplx_0(np__1),X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_143, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),k3_xcmplx_0(X1,X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_172, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_146, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_152, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers))).# trainpos
cnf(c_0_160, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(k3_xcmplx_0(k4_xcmplx_0(X1),X2),k1_numbers)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_183, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k3_xcmplx_0(X1,k4_xcmplx_0(X2)),k1_numbers))).# trainpos
cnf(c_0_198, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_225, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,k1_real_1(X2)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k3_xcmplx_0(X1,k1_real_1(X2)),k1_numbers)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_242, plain, (k1_real_1(k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X1,k1_real_1(X2))|~m1_subset_1(k3_xcmplx_0(X1,k1_real_1(X2)),k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_248, plain, (k1_real_1(k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X1,k1_real_1(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_259, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_real_1(k1_seq_1(k16_sin_cos,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(k1_seq_1(k16_sin_cos,esk2_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers))).# trainpos
cnf(c_0_274, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_fdiff_1(k19_sin_cos,esk2_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers))).# trainpos
cnf(c_0_280, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)|~v1_xreal_0(k1_fdiff_1(k19_sin_cos,esk2_0))|~v1_xreal_0(k1_seq_1(k19_sin_cos,esk2_0))|~r1_fdiff_1(k19_sin_cos,esk2_0))).# trainpos
cnf(c_0_288, negated_conjecture, (~v1_xreal_0(k1_fdiff_1(k19_sin_cos,esk2_0))|~v1_xreal_0(k1_seq_1(k19_sin_cos,esk2_0))|~r1_fdiff_1(k19_sin_cos,esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_291, negated_conjecture, (~v1_xreal_0(k1_seq_1(k19_sin_cos,esk2_0))|~r1_fdiff_1(k19_sin_cos,esk2_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers))).# trainpos
cnf(c_0_292, negated_conjecture, (~v1_xreal_0(k1_seq_1(k19_sin_cos,esk2_0))|~r1_fdiff_1(k19_sin_cos,esk2_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers))).# trainpos
cnf(c_0_294, negated_conjecture, (~r1_fdiff_1(k19_sin_cos,esk2_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers))).# trainpos
cnf(c_0_296, negated_conjecture, (~r1_fdiff_1(k19_sin_cos,esk2_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers))).# trainpos
cnf(c_0_299, negated_conjecture, (~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_138, plain, (k4_xcmplx_0(np__0)=np__0)).#trainneg
cnf(i_0_73, plain, (v1_xreal_0(esk5_0))).#trainneg
cnf(i_0_44, plain, (v1_xboole_0(np__0))).#trainneg
cnf(i_0_74, plain, (v1_xcmplx_0(esk5_0))).#trainneg
cnf(i_0_72, plain, (v1_int_1(esk5_0))).#trainneg
cnf(i_0_108, plain, (v1_int_1(esk6_0))).#trainneg
cnf(i_0_75, plain, (v1_xxreal_0(esk5_0))).#trainneg
cnf(i_0_88, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_58, plain, (v1_funct_2(k19_sin_cos,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_115, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_140, plain, (~v1_xboole_0(np__2))).#trainneg
cnf(i_0_105, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_99, plain, (v2_xxreal_0(np__2))).#trainneg
cnf(i_0_132, plain, (k4_xcmplx_0(k4_xcmplx_0(np__2))=np__2)).#trainneg
cnf(i_0_120, plain, (k6_xcmplx_0(np__1,np__1)=np__0)).#trainneg
cnf(i_0_133, plain, (k6_xcmplx_0(np__2,np__1)=np__1)).#trainneg
cnf(i_0_137, plain, (k6_xcmplx_0(np__2,np__2)=np__0)).#trainneg
cnf(i_0_100, plain, (m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_101, plain, (m1_subset_1(np__0,k5_numbers))).#trainneg
cnf(i_0_124, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_96, plain, (m1_subset_1(np__2,k1_numbers))).#trainneg
cnf(i_0_60, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_76, plain, (m1_subset_1(esk5_0,k1_numbers))).#trainneg
cnf(i_0_64, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_91, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_40, plain, (v1_funct_2(k24_sin_cos,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_35, plain, (v1_funct_2(k16_sin_cos,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_106, plain, (v5_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_110, plain, (v2_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_119, plain, (k6_xcmplx_0(np__0,np__1)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_142, plain, (k6_xcmplx_0(np__0,np__2)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_144, plain, (k6_xcmplx_0(np__1,k4_xcmplx_0(np__1))=np__2)).#trainneg
cnf(i_0_130, plain, (m1_subset_1(esk7_1(X1),X1))).#trainneg
cnf(i_0_46, plain, (v1_relat_1(k3_relat_1(X1,X2)))).#trainneg
cnf(i_0_102, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_87, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_98, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_82, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_81, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_7, plain, (X1=k2_fdiff_1(X2,X3)|k1_relset_1(k1_numbers,X1)!=X3|k1_fdiff_1(X2,esk3_3(X2,X3,X1))!=k1_seq_1(X1,esk3_3(X2,X3,X1))|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_83, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_135, plain, (v2_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_77, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_116, plain, (v3_valued_0(k2_zfmisc_1(X2,X1))|~v3_membered(X1))).#trainneg
cnf(i_0_125, plain, (v1_valued_0(k2_zfmisc_1(X2,X1))|~v1_membered(X1))).#trainneg
cnf(i_0_66, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainneg
cnf(i_0_94, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_107, plain, (v1_int_1(X2)|~v5_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_61, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_139, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_80, plain, (v1_membered(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_143, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_141, plain, (v1_int_1(k3_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_30, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_29, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_103, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_49, plain, (k8_real_1(X1,X2)=k8_real_1(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_126, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_127, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_50, plain, (v1_valued_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_33, plain, (k1_funct_1(X1,X2)=k1_seq_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_111, plain, (v3_valued_0(X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_117, plain, (v1_valued_0(X2)|~v1_relat_1(X1)|~v1_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_131, plain, (v2_valued_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_112, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_113, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_55, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_128, plain, (v3_valued_0(k3_relat_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v3_valued_0(X2))).#trainneg
cnf(i_0_62, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_145, plain, (v7_ordinal1(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_109, plain, (v1_xxreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_valued_0(X1))).#trainneg
cnf(i_0_121, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_68, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_123, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_18, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_17, plain, (r1_fdiff_1(X3,X1)|~v3_rcomp_1(X2)|~v1_funct_1(X3)|~r2_fdiff_1(X3,X2)|~r2_hidden(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_11, plain, (k1_relset_1(k1_numbers,X1)=X2|X1!=k2_fdiff_1(X3,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_9, plain, (X3=k2_fdiff_1(X1,X2)|m1_subset_1(esk3_3(X1,X2,X3),k1_numbers)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_8, plain, (X3=k2_fdiff_1(X1,X2)|r2_hidden(esk3_3(X1,X2,X3),X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_146, plain, (k1_xboole_0=np__0)).#trainneg
cnf(i_0_148, plain, (v1_xcmplx_0(np__2)|~v1_xcmplx_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_424, plain, (r1_fdiff_1(X1,esk4_2(X2,X3))|r2_fdiff_1(X3,X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_funct_1(X3)|~r1_tarski(X2,k1_relset_1(k1_numbers,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_343, plain, (r2_fdiff_1(k24_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k24_sin_cos))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xreal_0(esk4_2(X1,k24_sin_cos)))).#trainneg
cnf(i_0_422, plain, (r2_fdiff_1(X2,X1)|~v1_xboole_0(X1)|~v1_funct_1(X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_344, plain, (r2_fdiff_1(k19_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xreal_0(esk4_2(X1,k19_sin_cos)))).#trainneg
cnf(i_0_162, negated_conjecture, (k4_xcmplx_0(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_292, plain, (v1_relat_1(k24_sin_cos))).#trainneg
cnf(i_0_228, plain, (r2_fdiff_1(X1,X2)|m1_subset_1(esk4_2(X2,X1),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_314, plain, (v4_relat_1(k19_sin_cos,k1_numbers))).#trainneg
cnf(i_0_421, plain, (k1_relset_1(k1_numbers,k2_fdiff_1(X1,X2))=X2|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_221, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_172, plain, (v1_xreal_0(np__2))).#trainneg
cnf(i_0_319, negated_conjecture, (k1_real_1(k8_real_1(k1_seq_1(k16_sin_cos,esk2_0),k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~v1_xreal_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)))|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers))).#trainneg
cnf(i_0_149, plain, (v1_xcmplx_0(np__1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_312, plain, (v4_relat_1(k16_sin_cos,k1_numbers))).#trainneg
cnf(i_0_176, plain, (v1_xreal_0(np__1))).#trainneg
cnf(i_0_178, plain, (v1_xreal_0(np__0))).#trainneg
cnf(i_0_262, negated_conjecture, (v3_membered(esk1_0))).#trainneg
cnf(i_0_150, plain, (v1_xcmplx_0(np__0)|~v1_int_1(np__0))).#trainneg
cnf(i_0_151, plain, (v1_xcmplx_0(np__2)|~v1_int_1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_423, plain, (m1_subset_1(esk4_2(X1,X2),X1)|r2_fdiff_1(X2,X1)|~v1_funct_1(X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_288, plain, (k1_real_1(k1_fdiff_1(k19_sin_cos,X1))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_313, plain, (v4_relat_1(k24_sin_cos,k1_numbers))).#trainneg
cnf(i_0_321, negated_conjecture, (k1_real_1(k8_real_1(k1_seq_1(k16_sin_cos,esk2_0),k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~v1_xreal_0(k1_seq_1(k16_sin_cos,esk2_0))|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_287, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1))=k1_fdiff_1(k19_sin_cos,X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_325, plain, (k8_real_1(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_219, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk7_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_156, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_216, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k16_sin_cos))).#trainneg
cnf(i_0_152, plain, (v1_xcmplx_0(np__1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_316, plain, (v5_relat_1(k16_sin_cos,k1_numbers))).#trainneg
cnf(i_0_207, plain, (v7_ordinal1(np__0))).#trainneg
cnf(i_0_329, plain, (v3_valued_0(k24_sin_cos))).#trainneg
cnf(i_0_154, plain, (v1_int_1(np__2)|~v1_int_1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_155, plain, (v1_int_1(np__1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_322, plain, (m1_subset_1(k8_real_1(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_411, plain, (v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_413, plain, (v4_relat_1(k2_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_318, plain, (v5_relat_1(k19_sin_cos,k1_numbers))).#trainneg
cnf(i_0_317, plain, (v5_relat_1(k24_sin_cos,k1_numbers))).#trainneg
cnf(i_0_199, plain, (v7_ordinal1(esk5_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_217, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k24_sin_cos))).#trainneg
cnf(i_0_160, plain, (v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_282, negated_conjecture, (v1_membered(esk1_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_295, negated_conjecture, (v2_membered(esk1_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_165, plain, (X1=k2_fdiff_1(k24_sin_cos,X2)|k1_seq_1(k24_sin_cos,esk3_3(k24_sin_cos,X2,X1))!=k1_seq_1(X1,esk3_3(k24_sin_cos,X2,X1))|k1_relset_1(k1_numbers,X1)!=X2|~v1_funct_1(X1)|~r2_fdiff_1(k24_sin_cos,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(esk3_3(k24_sin_cos,X2,X1)))).#trainneg
cnf(i_0_267, plain, (v1_xreal_0(k8_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_420, plain, (r1_fdiff_1(X1,X2)|v1_xboole_0(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~m1_subset_1(X2,X3))).#trainneg
cnf(i_0_226, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|m1_subset_1(esk4_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)),k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_410, plain, (v5_relat_1(k2_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_289, plain, (m1_subset_1(k1_fdiff_1(k19_sin_cos,X1),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_218, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k19_sin_cos))).#trainneg
cnf(i_0_227, plain, (r2_fdiff_1(X1,k1_relset_1(k1_numbers,X1))|m1_subset_1(esk4_2(k1_relset_1(k1_numbers,X1),X1),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,X1)))).#trainneg
cnf(i_0_418, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_186, plain, (v1_xcmplx_0(np__2)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_213, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,esk7_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_168, plain, (v1_xreal_0(esk7_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_187, plain, (v1_xcmplx_0(np__2)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_191, plain, (v1_xcmplx_0(np__1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_340, plain, (v1_valued_0(k16_sin_cos)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_341, plain, (v1_valued_0(k24_sin_cos)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_170, negated_conjecture, (v1_xreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_174, plain, (v1_xreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_277, plain, (k3_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_194, plain, (v1_xcmplx_0(k1_real_1(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_668, negated_conjecture, (m1_subset_1(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_478, plain, (v1_xreal_0(esk4_2(X1,X2))|r2_fdiff_1(X2,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_469, plain, (r2_fdiff_1(k24_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k24_sin_cos))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_int_1(esk4_2(X1,k24_sin_cos)))).#trainneg
cnf(i_0_91, plain, (X1=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_180, plain, (v1_xreal_0(k16_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_580, plain, (v1_int_1(np__0))).#trainneg
cnf(i_0_689, plain, (v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_342, plain, (v1_valued_0(k19_sin_cos)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_358, plain, (v2_valued_0(k16_sin_cos)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_181, plain, (v1_xreal_0(k24_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_182, plain, (v1_xreal_0(k19_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_474, plain, (r2_fdiff_1(k19_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_int_1(esk4_2(X1,k19_sin_cos)))).#trainneg
cnf(i_0_488, plain, (v1_xreal_0(k1_seq_1(k24_sin_cos,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_509, plain, (r2_fdiff_1(X1,X2)|m1_subset_1(esk4_2(X2,X1),X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_472, plain, (r2_fdiff_1(X1,X2)|~v1_xboole_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_460, plain, (k8_real_1(k1_fdiff_1(X1,X2),k1_fdiff_1(X3,k1_seq_1(X1,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X3),X2)|~v1_xreal_0(k1_fdiff_1(X3,k1_seq_1(X1,X2)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~r1_fdiff_1(X3,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_539, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_548, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_563, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_633, plain, (k1_real_1(np__0)=np__0)).#trainneg
cnf(i_0_499, negated_conjecture, (k1_real_1(k3_xcmplx_0(k1_seq_1(k16_sin_cos,esk2_0),k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers))).#trainneg
cnf(i_0_183, plain, (v1_xcmplx_0(esk7_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_192, plain, (v1_xcmplx_0(np__0)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_193, plain, (v1_xcmplx_0(np__0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_359, plain, (v2_valued_0(k24_sin_cos)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_360, plain, (v2_valued_0(k19_sin_cos)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_185, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_189, plain, (v1_xcmplx_0(k5_numbers)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_470, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~v1_xboole_0(esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_476, negated_conjecture, (k4_xcmplx_0(k8_real_1(k1_seq_1(k16_sin_cos,esk2_0),k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k8_real_1(k1_seq_1(k16_sin_cos,esk2_0),k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0))),k1_numbers)|~v1_xreal_0(k1_seq_1(k16_sin_cos,esk2_0))|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_534, plain, (k3_xcmplx_0(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_527, plain, (k3_xcmplx_0(k1_fdiff_1(X2,X3),k1_fdiff_1(X1,k1_seq_1(X2,X3)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3)|~r1_fdiff_1(X1,k1_seq_1(X2,X3))|~r1_fdiff_1(X2,X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~v1_xreal_0(k1_fdiff_1(X1,k1_seq_1(X2,X3)))|~m1_subset_1(k1_fdiff_1(X2,X3),k1_numbers))).#trainneg
cnf(i_0_511, plain, (k4_xcmplx_0(k1_fdiff_1(k19_sin_cos,X1))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(k1_fdiff_1(k19_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_475, negated_conjecture, (k4_xcmplx_0(k8_real_1(k1_seq_1(k16_sin_cos,esk2_0),k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k8_real_1(k1_seq_1(k16_sin_cos,esk2_0),k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0))),k1_numbers)|~v1_xreal_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)))|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers))).#trainneg
cnf(i_0_249, plain, (v1_xxreal_0(np__2)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_250, plain, (v1_xxreal_0(np__2)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_195, plain, (v1_xcmplx_0(k16_sin_cos)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_346, negated_conjecture, (v3_valued_0(esk1_0)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_352, negated_conjecture, (v1_valued_0(esk1_0)|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_196, plain, (v1_xcmplx_0(k24_sin_cos)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_197, plain, (v1_xcmplx_0(k19_sin_cos)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_477, negated_conjecture, (k4_xcmplx_0(k3_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k3_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0)),k1_numbers)|~v1_xreal_0(k1_seq_1(k16_sin_cos,esk2_0))|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_614, plain, (k9_xtuple_0(k2_fdiff_1(X1,X2))=k1_relset_1(k1_numbers,k2_fdiff_1(X1,X2))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_531, negated_conjecture, (k4_xcmplx_0(k3_xcmplx_0(k1_seq_1(k16_sin_cos,esk2_0),k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k3_xcmplx_0(k1_seq_1(k16_sin_cos,esk2_0),k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0))),k1_numbers)|~v1_xreal_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)))|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers))).#trainneg
cnf(i_0_564, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_559, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_464, plain, (k8_real_1(k1_fdiff_1(X1,k1_seq_1(k24_sin_cos,X2)),k1_seq_1(k24_sin_cos,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,X1),X2)|~r1_fdiff_1(X1,k1_seq_1(k24_sin_cos,X2))|~r1_fdiff_1(k24_sin_cos,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_487, plain, (r2_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_xboole_0(X3)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r1_tarski(X3,X2)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_435, plain, (v1_relat_1(k1_partfun1(X3,X4,X5,X6,X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_253, plain, (v1_xxreal_0(np__1)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_254, plain, (v1_xxreal_0(np__1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_524, negated_conjecture, (k1_real_1(k3_xcmplx_0(k1_seq_1(k16_sin_cos,esk2_0),k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(k1_seq_1(k16_sin_cos,esk2_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers)|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,esk2_0))|~v1_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0))))).#trainneg
cnf(i_0_198, plain, (v7_ordinal1(esk7_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_255, plain, (v1_xxreal_0(np__0)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_256, plain, (v1_xxreal_0(np__0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_696, negated_conjecture, (v1_xreal_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_209, plain, (v7_ordinal1(k1_real_1(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_200, negated_conjecture, (v7_ordinal1(esk1_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_204, plain, (v7_ordinal1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_452, plain, (v1_xboole_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_xboole_0(X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_514, plain, (k9_xtuple_0(k24_sin_cos)=k1_relset_1(k1_numbers,k24_sin_cos))).#trainneg
cnf(i_0_503, plain, (k9_xtuple_0(k16_sin_cos)=k1_relset_1(k1_numbers,k16_sin_cos))).#trainneg
cnf(i_0_599, plain, (m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3),k1_numbers)|~v1_xreal_0(k1_fdiff_1(X1,k1_seq_1(X2,X3)))|~m1_subset_1(k1_fdiff_1(X2,X3),k1_numbers)|~r1_fdiff_1(X1,k1_seq_1(X2,X3))|~r1_fdiff_1(X2,X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_483, plain, (k9_xtuple_0(k19_sin_cos)=k1_relset_1(k1_numbers,k19_sin_cos))).#trainneg
cnf(i_0_283, plain, (v1_membered(k5_numbers)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_296, plain, (v2_membered(k5_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_575, plain, (X2=k2_fdiff_1(X1,k16_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|k1_relset_1(k1_numbers,X2)!=k16_sin_cos|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k16_sin_cos)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_210, plain, (v7_ordinal1(k16_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_231, negated_conjecture, (v1_int_1(esk1_0)|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_235, plain, (v1_int_1(k5_numbers)|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_240, plain, (v1_int_1(k1_real_1(X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_257, plain, (v1_xxreal_0(k1_real_1(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_211, plain, (v7_ordinal1(k24_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_212, plain, (v7_ordinal1(k19_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_461, plain, (k8_real_1(k1_fdiff_1(X2,X3),k1_fdiff_1(X1,k1_seq_1(X2,X3)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3)|~v1_xreal_0(k1_fdiff_1(X2,X3))|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(X2,X3)),k1_numbers)|~r1_fdiff_1(X1,k1_seq_1(X2,X3))|~r1_fdiff_1(X2,X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_636, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_669, negated_conjecture, (~v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_457, plain, (v1_xboole_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_xboole_0(X1)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_594, plain, (v1_xreal_0(k8_real_1(X1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_644, negated_conjecture, (v1_xreal_0(esk4_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_462, plain, (k3_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(X2,X3)),k1_fdiff_1(X2,X3))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3)|~v1_xreal_0(k1_fdiff_1(X2,X3))|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(X2,X3)),k1_numbers)|~r1_fdiff_1(X1,k1_seq_1(X2,X3))|~r1_fdiff_1(X2,X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_150, plain, (v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_277, plain, (k3_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_222, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k16_sin_cos))).#trainneg
cnf(i_0_248, negated_conjecture, (v1_xxreal_0(esk1_0)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_252, plain, (v1_xxreal_0(k5_numbers)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_898, plain, (v1_int_1(k1_real_1(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_484, plain, (m1_subset_1(k1_relset_1(k1_numbers,k19_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_223, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k24_sin_cos))).#trainneg
cnf(i_0_224, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k19_sin_cos))).#trainneg
cnf(i_0_725, plain, (r2_fdiff_1(k24_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k24_sin_cos))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k24_sin_cos))))).#trainneg
cnf(i_0_747, plain, (r1_fdiff_1(k1_partfun1(X4,X5,X6,X7,X1,X2),X3)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~r1_fdiff_1(X1,X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_750, plain, (m1_subset_1(k1_partfun1(X7,X8,X9,X10,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_726, plain, (r2_fdiff_1(k19_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k19_sin_cos))))).#trainneg
cnf(i_0_495, negated_conjecture, (m1_subset_1(X1,k1_numbers)|v1_xboole_0(esk1_0)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_562, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_347, plain, (v3_valued_0(k5_numbers)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_229, plain, (v1_int_1(esk7_1(X1))|~v5_membered(X1))).#trainneg
cnf(i_0_353, plain, (v1_valued_0(k5_numbers)|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_833, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_834, plain, (v1_xcmplx_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_504, plain, (m1_subset_1(k1_relset_1(k1_numbers,k16_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_515, plain, (m1_subset_1(k1_relset_1(k1_numbers,k24_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_241, plain, (v1_int_1(k16_sin_cos)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_242, plain, (v1_int_1(k24_sin_cos)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_768, plain, (r2_fdiff_1(k1_partfun1(X4,X5,X6,X7,X1,X2),X3)|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X3,k1_partfun1(X4,X5,X6,X7,X1,X2))))|~r1_fdiff_1(X1,esk4_2(X3,k1_partfun1(X4,X5,X6,X7,X1,X2)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(X4,X5,X6,X7,X1,X2)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_720, negated_conjecture, (v1_xcmplx_0(X1)|~v1_membered(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_791, plain, (v1_xcmplx_0(esk4_2(X1,X2))|r2_fdiff_1(X2,X1)|~v1_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_498, plain, (m1_subset_1(X1,k1_numbers)|v1_xboole_0(k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_565, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_697, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_243, plain, (v1_int_1(k19_sin_cos)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_264, plain, (v3_membered(k16_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_928, negated_conjecture, (v1_int_1(esk1_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_929, plain, (v1_int_1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_489, plain, (v1_xcmplx_0(k1_seq_1(k24_sin_cos,X1))|~v1_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_490, plain, (v7_ordinal1(k1_seq_1(k24_sin_cos,X1))|~v6_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_246, plain, (v1_xxreal_0(esk7_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_258, plain, (v1_xxreal_0(k16_sin_cos)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_773, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),X2)|~r1_fdiff_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_xreal_0(k1_seq_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))))).#trainneg
cnf(i_0_993, plain, (k3_xcmplx_0(k1_fdiff_1(X1,X2),k1_fdiff_1(X3,k1_seq_1(X1,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X3),X2)|~v1_xcmplx_0(k1_fdiff_1(X3,k1_seq_1(X1,X2)))|~v1_xcmplx_0(k1_fdiff_1(X1,X2))|~v1_xreal_0(k1_fdiff_1(X1,X2))|~r1_fdiff_1(X3,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(k1_fdiff_1(X3,k1_seq_1(X1,X2)),k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_897, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers)|~v1_xcmplx_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_896, plain, (v1_int_1(esk7_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_634, plain, (k1_real_1(np__2)=k4_xcmplx_0(np__2)|~m1_subset_1(k4_xcmplx_0(np__2),k1_numbers))).#trainneg
cnf(i_0_259, plain, (v1_xxreal_0(k24_sin_cos)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_265, plain, (v3_membered(k24_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_266, plain, (v3_membered(k19_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_284, plain, (v1_membered(k16_sin_cos)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_491, plain, (v1_int_1(k1_seq_1(k24_sin_cos,X1))|~v5_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_492, plain, (v1_xxreal_0(k1_seq_1(k24_sin_cos,X1))|~v2_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_260, plain, (v1_xxreal_0(k19_sin_cos)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_261, plain, (v3_membered(esk7_1(k1_zfmisc_1(X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_960, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1))=k1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1)),k1_numbers))).#trainneg
cnf(i_0_771, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4)|~r1_fdiff_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))|~v1_funct_1(X1)|~r1_tarski(X4,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X4)|~r1_fdiff_1(X3,k1_seq_1(X2,k1_seq_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))))))|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)))))|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)))),k1_numbers))).#trainneg
cnf(i_0_815, plain, (v1_xreal_0(k1_fdiff_1(k19_sin_cos,X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_862, plain, (v1_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~v1_xcmplx_0(k1_fdiff_1(k19_sin_cos,X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_813, plain, (v1_xreal_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_719, negated_conjecture, (v1_xreal_0(X1)|~v3_membered(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_553, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_542, plain, (v7_ordinal1(k4_xcmplx_0(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_268, plain, (v1_xcmplx_0(k8_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_285, plain, (v1_membered(k24_sin_cos)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_286, plain, (v1_membered(k19_sin_cos)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_297, plain, (v2_membered(k16_sin_cos)|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_554, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_555, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_269, plain, (v7_ordinal1(k8_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_270, plain, (v1_int_1(k8_real_1(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_730, plain, (k8_real_1(k1_fdiff_1(X1,k1_fdiff_1(X2,X4)),k1_fdiff_1(k2_fdiff_1(X2,X3),X4))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X2,X3),X1),X4)|~r1_fdiff_1(X1,k1_fdiff_1(X2,X4))|~r1_fdiff_1(k2_fdiff_1(X2,X3),X4)|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_numbers)|~v1_funct_1(X2)|~r2_hidden(X4,X3)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_863, plain, (v1_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~v1_int_1(k1_fdiff_1(k19_sin_cos,X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_984, plain, (v1_xreal_0(k3_xcmplx_0(X2,X1))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_924, plain, (m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k1_numbers)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~r1_fdiff_1(X1,X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X3),k1_numbers)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(k1_fdiff_1(X2,k1_seq_1(X1,X3)),k1_numbers))).#trainneg
cnf(i_0_1124, plain, (k1_relset_1(X1,k2_zfmisc_1(X1,X2))=k9_xtuple_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1039, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k19_sin_cos)))).#trainneg
cnf(i_0_1082, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k16_sin_cos)))).#trainneg
cnf(i_0_1019, negated_conjecture, (v1_xxreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_271, plain, (v1_xxreal_0(k8_real_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_998, plain, (m1_subset_1(esk6_0,k1_numbers)|~v1_xcmplx_0(esk6_0))).#trainneg
cnf(i_0_1020, plain, (v1_xxreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_298, plain, (v2_membered(k24_sin_cos)|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1034, plain, (v3_membered(k1_relset_1(k1_numbers,k19_sin_cos)))).#trainneg
cnf(i_0_1077, plain, (v3_membered(k1_relset_1(k1_numbers,k16_sin_cos)))).#trainneg
cnf(i_0_273, plain, (v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_274, plain, (v7_ordinal1(k1_seq_1(X1,X2))|~v6_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1125, plain, (m1_subset_1(k1_relset_1(X1,k2_zfmisc_1(X1,X2)),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1064, plain, (r2_fdiff_1(k19_sin_cos,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k19_sin_cos))))).#trainneg
cnf(i_0_1048, plain, (r2_fdiff_1(k24_sin_cos,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k24_sin_cos))))).#trainneg
cnf(i_0_1147, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),X2)|~r1_fdiff_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k1_seq_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_1135, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_xboole_0(esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1094, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k24_sin_cos)))).#trainneg
cnf(i_0_1068, plain, (v3_valued_0(k2_zfmisc_1(X1,k1_numbers)))).#trainneg
cnf(i_0_1070, negated_conjecture, (v1_xcmplx_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_275, plain, (v1_int_1(k1_seq_1(X1,X2))|~v5_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_299, plain, (v2_membered(k19_sin_cos)|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1169, plain, (v1_membered(k16_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1220, plain, (v1_membered(k24_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1089, plain, (v3_membered(k1_relset_1(k1_numbers,k24_sin_cos)))).#trainneg
cnf(i_0_1192, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_276, plain, (v1_xxreal_0(k1_seq_1(X1,X2))|~v2_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1052, plain, (r1_fdiff_1(k3_relat_1(X5,X6),X7)|~r1_fdiff_1(X6,k1_seq_1(X5,X7))|~r1_fdiff_1(X5,X7)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X7,k1_numbers))).#trainneg
cnf(i_0_1187, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),X2)|~r1_fdiff_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_xcmplx_0(k1_seq_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos))))|~m1_subset_1(k4_xcmplx_0(k1_seq_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))),k1_numbers))).#trainneg
cnf(i_0_1146, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),X2)|~r1_fdiff_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_int_1(k1_seq_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))))).#trainneg
cnf(i_0_1098, plain, (r2_fdiff_1(k3_relat_1(X5,X6),X7)|~r1_fdiff_1(X6,k1_seq_1(X5,esk4_2(X7,k3_relat_1(X5,X6))))|~r1_fdiff_1(X5,esk4_2(X7,k3_relat_1(X5,X6)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~r1_tarski(X7,k1_relset_1(k1_numbers,k3_relat_1(X5,X6)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X7))).#trainneg
cnf(i_0_1244, plain, (k8_real_1(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_fdiff_1(X4,k1_fdiff_1(X1,X3)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X4),X3)|~v1_xreal_0(k1_fdiff_1(X4,k1_fdiff_1(X1,X3)))|~m1_subset_1(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_numbers)|~r1_fdiff_1(X4,k1_fdiff_1(X1,X3))|~r1_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(X4)|~v1_funct_1(X1)|~r2_hidden(X3,X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_1186, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),X2)|~r1_fdiff_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_xcmplx_0(k1_seq_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos))))|~m1_subset_1(k4_xcmplx_0(k1_seq_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))),k1_numbers))).#trainneg
cnf(i_0_1253, plain, (k3_xcmplx_0(k1_fdiff_1(k2_fdiff_1(X2,X4),X3),k1_fdiff_1(X1,k1_fdiff_1(X2,X3)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X2,X4),X1),X3)|~v1_xreal_0(k1_fdiff_1(X1,k1_fdiff_1(X2,X3)))|~m1_subset_1(k1_fdiff_1(k2_fdiff_1(X2,X4),X3),k1_numbers)|~r1_fdiff_1(X1,k1_fdiff_1(X2,X3))|~r1_fdiff_1(k2_fdiff_1(X2,X4),X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~r2_hidden(X3,X4)|~r2_fdiff_1(X2,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_1132, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_544, plain, (v1_xxreal_0(k4_xcmplx_0(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1071, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1221, plain, (v1_membered(k19_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1222, plain, (v2_membered(k16_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1035, plain, (v1_membered(k1_relset_1(k1_numbers,k19_sin_cos))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1036, plain, (v2_membered(k1_relset_1(k1_numbers,k19_sin_cos))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_281, plain, (v1_membered(esk7_1(k1_zfmisc_1(X1)))|~v1_membered(X1))).#trainneg
cnf(i_0_290, plain, (v1_relat_1(esk7_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_1141, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k24_sin_cos),X1)|~r1_fdiff_1(k24_sin_cos,esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k24_sin_cos)))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k24_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xreal_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_1139, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),X2)|~r1_fdiff_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_int_1(k1_seq_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))))).#trainneg
cnf(i_0_463, plain, (m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3),k1_numbers)|~v1_xreal_0(k1_fdiff_1(X2,X3))|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(X2,X3)),k1_numbers)|~r1_fdiff_1(X1,k1_seq_1(X2,X3))|~r1_fdiff_1(X2,X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_1148, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos),X1)|~r1_fdiff_1(k24_sin_cos,esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos)))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xreal_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_731, negated_conjecture, (r2_fdiff_1(k2_fdiff_1(X1,X2),esk1_0)|m1_subset_1(esk4_2(esk1_0,k2_fdiff_1(X1,X2)),k1_numbers)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r1_tarski(esk1_0,X2)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_910, plain, (k8_real_1(k1_fdiff_1(X1,X2),k1_seq_1(k24_sin_cos,k1_seq_1(X1,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),X2)|~v1_xreal_0(k1_fdiff_1(X1,X2))|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(X1,X2)),k1_numbers)|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_451, plain, (v3_valued_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v3_membered(X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_556, plain, (v1_int_1(X1)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1128, plain, (v3_membered(k16_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_294, plain, (v2_membered(esk7_1(k1_zfmisc_1(X1)))|~v2_membered(X1))).#trainneg
cnf(i_0_1167, plain, (v3_membered(k24_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_874, plain, (v1_xcmplx_0(k16_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_875, plain, (v1_xcmplx_0(k24_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1078, plain, (v1_membered(k1_relset_1(k1_numbers,k16_sin_cos))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1079, plain, (v2_membered(k1_relset_1(k1_numbers,k16_sin_cos))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1245, plain, (k8_real_1(k1_fdiff_1(k2_fdiff_1(X2,X4),X3),k1_fdiff_1(X1,k1_fdiff_1(X2,X3)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X2,X4),X1),X3)|~v1_xreal_0(k1_fdiff_1(k2_fdiff_1(X2,X4),X3))|~m1_subset_1(k1_fdiff_1(X1,k1_fdiff_1(X2,X3)),k1_numbers)|~r1_fdiff_1(X1,k1_fdiff_1(X2,X3))|~r1_fdiff_1(k2_fdiff_1(X2,X4),X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~r2_hidden(X3,X4)|~r2_fdiff_1(X2,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_909, plain, (k8_real_1(k1_fdiff_1(X1,X2),k1_seq_1(k24_sin_cos,k1_seq_1(X1,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),X2)|~v1_xreal_0(k1_seq_1(k24_sin_cos,k1_seq_1(X1,X2)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1111, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_890, plain, (k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k24_sin_cos,X1)),k1_seq_1(k24_sin_cos,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k24_sin_cos),X1)|~v1_xreal_0(X1)|~r1_fdiff_1(k24_sin_cos,k1_seq_1(k24_sin_cos,X1))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(k1_seq_1(k24_sin_cos,X1)))).#trainneg
cnf(i_0_1341, negated_conjecture, (~v1_xboole_0(esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))).#trainneg
cnf(i_0_1320, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k1_relset_1(X1,k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_1338, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),X2)|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~m1_subset_1(k1_seq_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos))),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos))|~r2_fdiff_1(X1,X2)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1319, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1168, plain, (v3_membered(k19_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1292, plain, (v2_membered(k24_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1349, plain, (v2_membered(k19_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1090, plain, (v1_membered(k1_relset_1(k1_numbers,k24_sin_cos))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1091, plain, (v2_membered(k1_relset_1(k1_numbers,k24_sin_cos))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_309, plain, (v1_int_1(k4_xcmplx_0(np__2))|~v1_int_1(np__2)|~v1_int_1(np__0))).#trainneg
cnf(i_0_1461, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_1379, plain, (r1_fdiff_1(k3_relat_1(X1,k24_sin_cos),X2)|~r1_fdiff_1(k24_sin_cos,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1372, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),X2)|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))|~m1_subset_1(k1_seq_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos))),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos))|~r2_fdiff_1(X1,X2)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1380, plain, (r1_fdiff_1(k3_relat_1(X1,k19_sin_cos),X2)|~r1_fdiff_1(k19_sin_cos,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1476, negated_conjecture, (v1_xreal_0(esk4_2(esk1_0,k2_fdiff_1(X1,X2)))|r2_fdiff_1(k2_fdiff_1(X1,X2),esk1_0)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r1_tarski(esk1_0,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_996, plain, (k3_xcmplx_0(k1_fdiff_1(X1,k1_fdiff_1(X2,X4)),k1_fdiff_1(k2_fdiff_1(X2,X3),X4))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X2,X3),X1),X4)|~v1_xreal_0(k1_fdiff_1(k2_fdiff_1(X2,X3),X4))|~r1_fdiff_1(X1,k1_fdiff_1(X2,X4))|~r1_fdiff_1(k2_fdiff_1(X2,X3),X4)|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(X2,X3))|~m1_subset_1(k1_fdiff_1(X1,k1_fdiff_1(X2,X4)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_numbers)|~v1_funct_1(X2)|~r2_hidden(X4,X3)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_746, negated_conjecture, (m1_subset_1(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k24_sin_cos)))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1356, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1350, plain, (v1_membered(k16_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_877, negated_conjecture, (k4_xcmplx_0(k3_xcmplx_0(k1_seq_1(k16_sin_cos,esk2_0),k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(k1_seq_1(k16_sin_cos,esk2_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers)|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,esk2_0))|~v1_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0))))).#trainneg
cnf(i_0_876, plain, (v1_xcmplx_0(k19_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1096, plain, (v1_int_1(k16_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1029, plain, (v1_xreal_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1030, plain, (v1_xcmplx_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_311, plain, (v4_relat_1(esk7_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_315, plain, (v5_relat_1(esk7_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_915, plain, (k3_xcmplx_0(k1_fdiff_1(X1,X2),k1_seq_1(k24_sin_cos,k1_seq_1(X1,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),X2)|~v1_xreal_0(k1_seq_1(k24_sin_cos,k1_seq_1(X1,X2)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1441, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k24_sin_cos),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k24_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_884, plain, (k8_real_1(k1_seq_1(k24_sin_cos,X2),k1_fdiff_1(X1,k1_seq_1(k24_sin_cos,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,X1),X2)|~v1_xreal_0(k1_seq_1(k24_sin_cos,X2))|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k24_sin_cos,X2)),k1_numbers)|~v1_xreal_0(X2)|~r1_fdiff_1(X1,k1_seq_1(k24_sin_cos,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_770, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3),X4)|~r1_fdiff_1(X3,k1_fdiff_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3))))|~r1_fdiff_1(k2_fdiff_1(X1,X2),esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3)))|~v1_funct_1(X3)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r1_tarski(X4,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X4)|~v1_funct_1(X1)|~r2_hidden(esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3)),X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3)),k1_numbers))).#trainneg
cnf(i_0_1346, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1321, plain, (m1_subset_1(k9_xtuple_0(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1351, plain, (v1_membered(k24_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_327, plain, (v3_valued_0(esk7_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v3_membered(X2))).#trainneg
cnf(i_0_1436, plain, (v1_membered(k19_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1097, plain, (v1_int_1(k24_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1127, plain, (v1_int_1(k19_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1031, plain, (v7_ordinal1(k1_relset_1(k1_numbers,k19_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1032, plain, (v1_int_1(k1_relset_1(k1_numbers,k19_sin_cos))|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_331, plain, (v1_xboole_0(esk7_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_335, plain, (v1_xboole_0(esk7_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1117, negated_conjecture, (v1_xcmplx_0(X1)|~r2_hidden(X1,esk1_0)|~v3_membered(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_986, negated_conjecture, (v1_xreal_0(esk4_2(esk1_0,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos)))|r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_458, plain, (v4_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6),X1)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1460, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_int_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_1051, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X2,X3,X4,X5,X6,X7)),X8)|~r1_fdiff_1(X1,esk4_2(X8,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X2,X3,X4,X5,X6,X7))))|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,X6,X7))|~v1_funct_1(X1)|~r1_tarski(X8,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X2,X3,X4,X5,X6,X7))))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,X6,X7),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X8,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X8)|~r1_fdiff_1(X7,k1_seq_1(X6,k1_seq_1(X1,esk4_2(X8,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X2,X3,X4,X5,X6,X7))))))|~r1_fdiff_1(X6,k1_seq_1(X1,esk4_2(X8,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X2,X3,X4,X5,X6,X7)))))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k1_seq_1(X1,esk4_2(X8,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(X2,X3,X4,X5,X6,X7)))),k1_numbers))).#trainneg
cnf(i_0_1527, plain, (k8_real_1(k1_seq_1(k24_sin_cos,X1),k1_seq_1(k24_sin_cos,k1_seq_1(k24_sin_cos,X1)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k24_sin_cos),X1)|~v1_xreal_0(k1_seq_1(k24_sin_cos,X1))|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k24_sin_cos,X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_557, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1437, plain, (v2_membered(k16_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_339, plain, (v1_valued_0(esk7_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_membered(X2))).#trainneg
cnf(i_0_1126, plain, (m1_subset_1(k5_numbers,k1_numbers)|~v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1131, plain, (v1_xxreal_0(k16_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1166, plain, (v1_xxreal_0(k24_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1033, plain, (v1_xxreal_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1072, plain, (v1_xreal_0(k1_relset_1(k1_numbers,k16_sin_cos))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_345, plain, (v3_valued_0(esk7_1(k1_zfmisc_1(X1)))|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_885, plain, (k3_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k24_sin_cos,X2)),k1_seq_1(k24_sin_cos,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,X1),X2)|~v1_xreal_0(k1_seq_1(k24_sin_cos,X2))|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k24_sin_cos,X2)),k1_numbers)|~v1_xreal_0(X2)|~r1_fdiff_1(X1,k1_seq_1(k24_sin_cos,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_436, plain, (v3_valued_0(k1_partfun1(X3,X4,X5,X6,X1,X2))|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_957, negated_conjecture, (k4_xcmplx_0(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~v1_xcmplx_0(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0)))|~m1_subset_1(k4_xcmplx_0(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0))),k1_numbers))).#trainneg
cnf(i_0_1153, plain, (k3_xcmplx_0(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_fdiff_1(X4,k1_fdiff_1(X1,X3)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X4),X3)|~v1_xcmplx_0(k1_fdiff_1(X4,k1_fdiff_1(X1,X3)))|~v1_xcmplx_0(k1_fdiff_1(k2_fdiff_1(X1,X2),X3))|~v1_xreal_0(k1_fdiff_1(k2_fdiff_1(X1,X2),X3))|~r1_fdiff_1(X4,k1_fdiff_1(X1,X3))|~r1_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(X4)|~v1_funct_1(k2_fdiff_1(X1,X2))|~m1_subset_1(k1_fdiff_1(X4,k1_fdiff_1(X1,X3)),k1_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~v1_funct_1(X1)|~r2_hidden(X3,X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_847, plain, (k3_xcmplx_0(k3_xcmplx_0(X2,X1),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1223, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_561, plain, (v2_valued_0(k2_zfmisc_1(X1,X2))|~v2_membered(X2))).#trainneg
cnf(i_0_620, plain, (X2=k2_fdiff_1(X1,k24_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|k1_relset_1(k1_numbers,X2)!=k24_sin_cos|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k24_sin_cos)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1171, plain, (v1_xxreal_0(k19_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_738, plain, (v1_xreal_0(k16_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1073, plain, (v1_xcmplx_0(k1_relset_1(k1_numbers,k16_sin_cos))|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1074, plain, (v7_ordinal1(k1_relset_1(k1_numbers,k16_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_351, plain, (v1_valued_0(esk7_1(k1_zfmisc_1(X1)))|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1466, plain, (k3_xcmplx_0(k1_fdiff_1(X1,X2),k1_seq_1(k24_sin_cos,k1_seq_1(X1,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),X2)|~v1_xcmplx_0(k1_fdiff_1(X1,X2))|~v1_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(X1,X2)))|~v1_xreal_0(k1_fdiff_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1700, plain, (k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,X5,X6))=k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1540, negated_conjecture, (~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_1541, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k9_xtuple_0(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_1566, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos),X1)|~m1_subset_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos)))))).#trainneg
cnf(i_0_1577, plain, (r2_fdiff_1(k1_partfun1(X3,X4,X5,X6,X1,k24_sin_cos),X2)|~v1_funct_1(k1_partfun1(X3,X4,X5,X6,X1,k24_sin_cos))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k24_sin_cos)))|~m1_subset_1(k1_seq_1(X1,esk4_2(X2,k1_partfun1(X3,X4,X5,X6,X1,k24_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(X3,X4,X5,X6,X1,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_1545, plain, (r2_fdiff_1(k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos),X2)|~v1_funct_1(k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos)))|~m1_subset_1(k1_seq_1(X1,esk4_2(X2,k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_1553, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k9_xtuple_0(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1557, plain, (v2_membered(k24_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_656, plain, (X2=k2_fdiff_1(X1,k19_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|k1_relset_1(k1_numbers,X2)!=k19_sin_cos|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k19_sin_cos)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_354, plain, (v1_valued_0(k16_sin_cos)|~v1_valued_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_relat_1(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1558, plain, (v2_membered(k19_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_739, plain, (v1_xreal_0(k24_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_740, plain, (v1_xreal_0(k19_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1731, plain, (v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1792, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_355, plain, (v1_valued_0(k24_sin_cos)|~v1_valued_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_relat_1(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_356, plain, (v1_valued_0(k19_sin_cos)|~v1_valued_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_relat_1(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1703, plain, (v4_relat_1(k1_partfun1(X7,X8,X9,X10,X5,X6),X1)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_1708, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos),X1)|~v1_int_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos)))))).#trainneg
cnf(i_0_1701, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,X5,X6)),k1_zfmisc_1(X1))|~v1_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1579, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k24_sin_cos),X3)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k24_sin_cos))|~v1_funct_1(k2_fdiff_1(X1,X2))|~r2_fdiff_1(k2_fdiff_1(X1,X2),X3)|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k24_sin_cos)))|~m1_subset_1(k1_fdiff_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k24_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(X1)|~r2_hidden(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k24_sin_cos)),X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_1695, negated_conjecture, (v1_xcmplx_0(X1)|~r2_hidden(X1,esk1_0)|~m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1581, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),X2)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1549, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),X2)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1547, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos),X3)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos))|~v1_funct_1(k2_fdiff_1(X1,X2))|~r2_fdiff_1(k2_fdiff_1(X1,X2),X3)|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos)))|~m1_subset_1(k1_fdiff_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(X1)|~r2_hidden(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos)),X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_1496, plain, (v1_int_1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1631, plain, (v3_valued_0(esk7_1(k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))))).#trainneg
cnf(i_0_1751, negated_conjecture, (k4_xcmplx_0(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k4_xcmplx_0(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0))),k1_numbers)|~v1_membered(k1_numbers)|~v1_xreal_0(k1_seq_1(k16_sin_cos,esk2_0))|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_357, plain, (v2_valued_0(esk7_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v2_membered(X2))).#trainneg
cnf(i_0_1497, plain, (v1_xcmplx_0(k16_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1498, plain, (v1_xcmplx_0(k24_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1075, plain, (v1_int_1(k1_relset_1(k1_numbers,k16_sin_cos))|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1076, plain, (v1_xxreal_0(k1_relset_1(k1_numbers,k16_sin_cos))|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_366, plain, (v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_367, plain, (v1_int_1(k1_funct_1(X1,X2))|~v4_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1107, negated_conjecture, (v1_xcmplx_0(esk4_2(esk1_0,X1))|r2_fdiff_1(X1,esk1_0)|~v1_membered(k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_584, plain, (v1_xcmplx_0(X1)|~v1_int_1(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1178, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X3,X4))),X5)|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X3,X4))))))|~r1_fdiff_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X3,X4)))))|~v1_funct_1(X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X3,X4))|~v1_funct_1(X2)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X3,X4)))))|~m1_subset_1(k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X3,X4))))),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~r1_fdiff_1(X4,k1_seq_1(X3,k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X3,X4))))))))|~r1_fdiff_1(X3,k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X3,X4)))))))|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X3,X4)))))),k1_numbers))).#trainneg
cnf(i_0_1641, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k24_sin_cos),X1)|~m1_subset_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k24_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k24_sin_cos)))))).#trainneg
cnf(i_0_867, plain, (v1_int_1(k1_fdiff_1(k19_sin_cos,X1))|~v1_int_1(k1_seq_1(k16_sin_cos,X1))|~v1_xcmplx_0(k1_fdiff_1(k19_sin_cos,X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_1671, plain, (v1_relat_1(k9_xtuple_0(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_1066, negated_conjecture, (v1_xboole_0(esk1_0)|m1_subset_1(esk7_1(esk1_0),k1_numbers))).#trainneg
cnf(i_0_432, plain, (X3=k2_fdiff_1(X2,X1)|~v1_xboole_0(X1)|k1_relset_1(k1_numbers,X3)!=X1|~v1_funct_1(X3)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_368, plain, (v7_ordinal1(k1_seq_1(X1,X2))|~v4_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_1625, plain, (v1_xcmplx_0(k19_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1732, plain, (v1_xxreal_0(k16_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1084, plain, (v1_xreal_0(k1_relset_1(k1_numbers,k24_sin_cos))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1085, plain, (v1_xcmplx_0(k1_relset_1(k1_numbers,k24_sin_cos))|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_369, plain, (v1_xxreal_0(k1_seq_1(X1,X2))|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_370, plain, (v1_int_1(k6_xcmplx_0(X2,X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_808, plain, (k8_real_1(k1_seq_1(k24_sin_cos,X1),k1_fdiff_1(X2,k1_seq_1(k24_sin_cos,X1)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,X2),X1)|~v1_xreal_0(k1_fdiff_1(X2,k1_seq_1(k24_sin_cos,X1)))|~r1_fdiff_1(X2,k1_seq_1(k24_sin_cos,X1))|~r1_fdiff_1(k24_sin_cos,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(k24_sin_cos,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_728, plain, (m1_subset_1(k1_fdiff_1(X1,X3),k1_numbers)|~v3_valued_0(k2_fdiff_1(X1,X2))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~r2_hidden(X3,X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_585, plain, (v1_int_1(X1)|~v1_int_1(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1681, plain, (v3_membered(k9_xtuple_0(k2_zfmisc_1(X1,X2)))|~v3_membered(X1))).#trainneg
cnf(i_0_1119, plain, (v1_xboole_0(k5_numbers)|m1_subset_1(esk7_1(k5_numbers),k1_numbers))).#trainneg
cnf(i_0_691, plain, (X3=k2_fdiff_1(X2,esk7_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1)|k1_relset_1(k1_numbers,X3)!=esk7_1(k1_zfmisc_1(X1))|~v1_funct_1(X3)|~v1_funct_1(X2)|~r2_fdiff_1(X2,esk7_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_371, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),np__0)=k6_xcmplx_0(np__0,X1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1733, plain, (v1_xxreal_0(k24_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1794, plain, (v1_xxreal_0(k19_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1086, plain, (v7_ordinal1(k1_relset_1(k1_numbers,k24_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1087, plain, (v1_int_1(k1_relset_1(k1_numbers,k24_sin_cos))|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_372, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),np__2)=k6_xcmplx_0(k4_xcmplx_0(np__2),X1)|~v1_xcmplx_0(k4_xcmplx_0(np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_373, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),np__1)=k6_xcmplx_0(k4_xcmplx_0(np__1),X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_793, plain, (v1_int_1(esk4_2(X1,X2))|r2_fdiff_1(X2,X1)|~v5_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1462, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xcmplx_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos)))|~m1_subset_1(k4_xcmplx_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_1915, negated_conjecture, (v1_xcmplx_0(X1)|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_1820, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))))).#trainneg
cnf(i_0_1814, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_1856, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1848, plain, (v1_valued_0(k16_sin_cos))).#trainneg
cnf(i_0_374, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),X2)=k6_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1858, plain, (v1_valued_0(k24_sin_cos))).#trainneg
cnf(i_0_1860, plain, (v1_valued_0(k19_sin_cos))).#trainneg
cnf(i_0_1994, negated_conjecture, (v1_xreal_0(esk7_1(esk1_0)))).#trainneg
cnf(i_0_1959, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1960, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_375, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(X1))=k6_xcmplx_0(X1,np__0)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_376, plain, (k6_xcmplx_0(np__2,k4_xcmplx_0(X1))=k6_xcmplx_0(X1,k4_xcmplx_0(np__2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_1815, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xcmplx_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~m1_subset_1(k4_xcmplx_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_1813, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_int_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_1901, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k1_relset_1(X1,k1_partfun1(X1,X3,X4,X5,X6,X7)))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_1957, plain, (v1_xcmplx_0(X1)|~v1_int_1(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1945, plain, (v1_xcmplx_0(k1_fdiff_1(X1,X3))|~v3_valued_0(k2_fdiff_1(X1,X2))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~r2_hidden(X3,X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_1922, plain, (r2_fdiff_1(k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos),X2)|~v3_valued_0(X1)|~v1_funct_1(k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos)))|~m1_subset_1(k1_partfun1(X3,X4,X5,X6,X1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_1916, plain, (r2_fdiff_1(k1_partfun1(X3,X4,X5,X6,X1,k24_sin_cos),X2)|~v3_valued_0(X1)|~v1_funct_1(k1_partfun1(X3,X4,X5,X6,X1,k24_sin_cos))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k24_sin_cos)))|~m1_subset_1(k1_partfun1(X3,X4,X5,X6,X1,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_1910, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k24_sin_cos),X4)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k24_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~r2_hidden(esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k24_sin_cos)),X3)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),X4)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~r1_tarski(X4,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k24_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X4)|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k24_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k24_sin_cos)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k24_sin_cos)))),k1_numbers)|~m1_subset_1(k1_fdiff_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k24_sin_cos))),k1_numbers)|~m1_subset_1(esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_1909, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k24_sin_cos),X4)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k24_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~r2_hidden(esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k24_sin_cos)),X3)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),X4)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~r1_tarski(X4,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k24_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X4)|~v1_xreal_0(k1_fdiff_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k24_sin_cos))))|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k24_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k24_sin_cos)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k24_sin_cos)))),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_1682, plain, (v1_membered(k9_xtuple_0(k2_zfmisc_1(X1,X2)))|~v1_membered(X1))).#trainneg
cnf(i_0_2040, plain, (v1_xreal_0(esk7_1(k5_numbers))|v1_xboole_0(k5_numbers))).#trainneg
cnf(i_0_1852, negated_conjecture, (k4_xcmplx_0(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k4_xcmplx_0(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0))),k1_numbers)|~m1_subset_1(k1_zfmisc_1(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0))),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_377, plain, (k6_xcmplx_0(np__1,k4_xcmplx_0(X1))=k6_xcmplx_0(X1,k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1995, negated_conjecture, (v1_xcmplx_0(esk7_1(esk1_0))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1996, negated_conjecture, (v7_ordinal1(esk7_1(esk1_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2038, plain, (v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_378, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_379, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,k4_xcmplx_0(np__1)))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1931, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k19_sin_cos),X4)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k19_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~r2_hidden(esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k19_sin_cos)),X3)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),X4)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~r1_tarski(X4,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k19_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X4)|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k19_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k19_sin_cos)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k19_sin_cos)))),k1_numbers)|~m1_subset_1(k1_fdiff_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k19_sin_cos))),k1_numbers)|~m1_subset_1(esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_1930, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k19_sin_cos),X4)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k19_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~r2_hidden(esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k19_sin_cos)),X3)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),X4)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~r1_tarski(X4,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k19_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X4)|~v1_xreal_0(k1_fdiff_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k19_sin_cos))))|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k19_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k19_sin_cos)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k19_sin_cos)))),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_1942, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1))=k1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1)|~m1_subset_1(k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1)),k1_numbers))).#trainneg
cnf(i_0_1830, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos),X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos))))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1920, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),X2)|~v3_valued_0(X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))))).#trainneg
cnf(i_0_1979, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k24_sin_cos),X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k24_sin_cos))))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1977, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k24_sin_cos))))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1828, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos),esk1_0)|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos))))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2035, plain, (v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1683, plain, (v2_membered(k9_xtuple_0(k2_zfmisc_1(X1,X2)))|~v2_membered(X1))).#trainneg
cnf(i_0_1997, negated_conjecture, (v1_int_1(esk7_1(esk1_0))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_380, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))=k3_xcmplx_0(X3,k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1998, negated_conjecture, (v1_xxreal_0(esk7_1(esk1_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_2227, negated_conjecture, (v1_int_1(esk7_1(esk1_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2175, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_381, plain, (v1_int_1(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)))|~v1_int_1(X3)|~v1_int_1(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_382, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)),X4)=k3_xcmplx_0(k3_xcmplx_0(X1,X2),k3_xcmplx_0(X3,X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2092, negated_conjecture, (v1_xcmplx_0(esk4_2(esk1_0,X1))|r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~r1_tarski(esk1_0,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2099, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))))).#trainneg
cnf(i_0_2094, negated_conjecture, (v1_xcmplx_0(X1)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_2103, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(esk4_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_2155, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~v1_int_1(esk4_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_2105, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),X1)|~m1_subset_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))))).#trainneg
cnf(i_0_2167, plain, (v1_xcmplx_0(k4_xcmplx_0(np__2))|~m1_subset_1(k4_xcmplx_0(np__2),k1_numbers))).#trainneg
cnf(i_0_2041, plain, (v1_xcmplx_0(esk7_1(k5_numbers))|v1_xboole_0(k5_numbers)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_2168, plain, (v1_xcmplx_0(k4_xcmplx_0(np__1))|~m1_subset_1(k4_xcmplx_0(np__1),k1_numbers))).#trainneg
cnf(i_0_1088, plain, (v1_xxreal_0(k1_relset_1(k1_numbers,k24_sin_cos))|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1689, plain, (v1_int_1(k1_relset_1(k1_numbers,k19_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_384, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_2085, negated_conjecture, (r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_2157, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),X1)|~v1_int_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))))).#trainneg
cnf(i_0_2180, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,X5,k19_sin_cos),X6)|~v3_valued_0(X5)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,k19_sin_cos))|~v1_funct_1(X5)|~r2_fdiff_1(X5,X6)|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,k19_sin_cos)))))).#trainneg
cnf(i_0_2322, plain, (v1_xcmplx_0(k1_fdiff_1(k19_sin_cos,X1))|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(k1_fdiff_1(k19_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_2149, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~v1_xcmplx_0(esk4_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~m1_subset_1(k4_xcmplx_0(esk4_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_2436, plain, (v1_xcmplx_0(esk6_0)|~m1_subset_1(esk6_0,k1_numbers))).#trainneg
cnf(i_0_2432, negated_conjecture, (v1_xcmplx_0(esk7_1(esk1_0))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_2575, negated_conjecture, (k4_xcmplx_0(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k4_xcmplx_0(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0))),k1_numbers)|~m1_subset_1(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0)),esk1_0))).#trainneg
cnf(i_0_386, plain, (k3_xcmplx_0(k3_xcmplx_0(X2,X1),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2433, negated_conjecture, (v1_xcmplx_0(esk7_1(esk1_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2571, negated_conjecture, (m1_subset_1(esk6_0,k1_numbers)|~m1_subset_1(esk6_0,esk1_0))).#trainneg
cnf(i_0_1797, plain, (v1_int_1(k1_relset_1(k1_numbers,k16_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2054, plain, (v1_int_1(k1_relset_1(k1_numbers,k24_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2555, negated_conjecture, (v1_xcmplx_0(esk4_2(esk1_0,X1))|r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_2438, plain, (v1_xcmplx_0(esk4_2(X1,X2))|r2_fdiff_1(X2,X1)|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_2554, negated_conjecture, (v1_xcmplx_0(esk4_2(esk1_0,k2_fdiff_1(X1,X2)))|r2_fdiff_1(k2_fdiff_1(X1,X2),esk1_0)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r1_tarski(esk1_0,X2)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2187, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,X5,k24_sin_cos),X6)|~v3_valued_0(X5)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,k24_sin_cos))|~v1_funct_1(X5)|~r2_fdiff_1(X5,X6)|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,k24_sin_cos)))))).#trainneg
cnf(i_0_2042, plain, (v7_ordinal1(esk7_1(k5_numbers))|v1_xboole_0(k5_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_389, plain, (v1_xcmplx_0(k1_fdiff_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_664, negated_conjecture, (v1_membered(esk1_0)|~v1_membered(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_1037, plain, (v3_valued_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1038, plain, (v1_valued_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_390, plain, (v7_ordinal1(k1_fdiff_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_391, plain, (v1_int_1(k1_fdiff_1(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2095, plain, (r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k24_sin_cos),X1)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k24_sin_cos))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k24_sin_cos)))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1991, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_real_1(X1)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1670, plain, (v5_relat_1(k9_xtuple_0(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)),X2))).#trainneg
cnf(i_0_2043, plain, (v1_int_1(esk7_1(k5_numbers))|v1_xboole_0(k5_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_392, plain, (v1_xxreal_0(k1_fdiff_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_665, negated_conjecture, (v2_membered(esk1_0)|~v2_membered(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_1080, plain, (v3_valued_0(k1_relset_1(k1_numbers,k16_sin_cos))|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1081, plain, (v1_valued_0(k1_relset_1(k1_numbers,k16_sin_cos))|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_394, plain, (v1_xreal_0(k2_fdiff_1(X1,X2))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_395, plain, (v1_xcmplx_0(k2_fdiff_1(X1,X2))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2090, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1925, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),X2)|~v3_valued_0(X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))))).#trainneg
cnf(i_0_2151, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),X1)|~v1_xcmplx_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~m1_subset_1(k4_xcmplx_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))))).#trainneg
cnf(i_0_2574, negated_conjecture, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(X1),esk1_0))).#trainneg
cnf(i_0_2340, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k19_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k19_sin_cos))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k19_sin_cos))))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2337, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k24_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k24_sin_cos))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k24_sin_cos))))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2325, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_zfmisc_1(k4_xcmplx_0(X1)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1673, plain, (v4_relat_1(k9_xtuple_0(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)),X1))).#trainneg
cnf(i_0_2044, plain, (v1_xxreal_0(esk7_1(k5_numbers))|v1_xboole_0(k5_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_396, plain, (v7_ordinal1(k2_fdiff_1(X1,X2))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1092, plain, (v3_valued_0(k1_relset_1(k1_numbers,k24_sin_cos))|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1093, plain, (v1_valued_0(k1_relset_1(k1_numbers,k24_sin_cos))|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_397, plain, (v1_int_1(k2_fdiff_1(X1,X2))|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_398, plain, (v1_xxreal_0(k2_fdiff_1(X1,X2))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2623, plain, (k3_xcmplx_0(np__1,k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X3),X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X3),X2)|~v1_xcmplx_0(k1_fdiff_1(X3,k1_seq_1(X1,X2)))|~v1_xcmplx_0(k1_fdiff_1(X1,X2))|~v1_xreal_0(k1_fdiff_1(X3,k1_seq_1(X1,X2)))|~r1_fdiff_1(X3,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2299, plain, (r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k24_sin_cos,k19_sin_cos),X1)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k24_sin_cos,k19_sin_cos))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k24_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k24_sin_cos,k19_sin_cos))))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1929, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos),X2)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos))|~r2_hidden(esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos)),X1)|~r2_fdiff_1(k2_fdiff_1(k24_sin_cos,X1),X2)|~r2_fdiff_1(k24_sin_cos,X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos)))|~m1_subset_1(k1_seq_1(k24_sin_cos,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_xreal_0(esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos))))).#trainneg
cnf(i_0_1908, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos),X2)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos))|~r2_hidden(esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos)),X1)|~r2_fdiff_1(k2_fdiff_1(k24_sin_cos,X1),X2)|~r2_fdiff_1(k24_sin_cos,X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos)))|~m1_subset_1(k1_seq_1(k24_sin_cos,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_xreal_0(esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos))))).#trainneg
cnf(i_0_2017, plain, (k3_xcmplx_0(k1_seq_1(k24_sin_cos,X1),k1_fdiff_1(X2,k1_seq_1(k24_sin_cos,X1)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,X2),X1)|~v1_xreal_0(k1_fdiff_1(X2,k1_seq_1(k24_sin_cos,X1)))|~m1_subset_1(k1_seq_1(k24_sin_cos,X1),k1_numbers)|~r1_fdiff_1(X2,k1_seq_1(k24_sin_cos,X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2328, plain, (r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k24_sin_cos,k24_sin_cos),X1)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k24_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k24_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k24_sin_cos,k24_sin_cos))))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2622, plain, (k3_xcmplx_0(np__1,k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3)|~v1_xcmplx_0(k1_fdiff_1(X2,X3))|~v1_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(X2,X3)))|~v1_xreal_0(k1_fdiff_1(X2,X3))|~r1_fdiff_1(X1,k1_seq_1(X2,X3))|~r1_fdiff_1(X2,X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_1157, plain, (k1_real_1(k4_xcmplx_0(np__2))=np__2|~m1_subset_1(k4_xcmplx_0(np__2),k1_numbers))).#trainneg
cnf(i_0_399, plain, (v3_membered(k2_fdiff_1(X1,X2))|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1130, plain, (v1_int_1(k1_seq_1(k24_sin_cos,X1))|~v6_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_400, plain, (v1_membered(k2_fdiff_1(X1,X2))|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_401, plain, (v2_membered(k2_fdiff_1(X1,X2))|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_402, plain, (v3_valued_0(k2_fdiff_1(X1,X2))|~v3_valued_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_relat_1(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2799, negated_conjecture, (v1_membered(esk1_0)|~v3_membered(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_403, plain, (v1_valued_0(k2_fdiff_1(X1,X2))|~v1_valued_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_relat_1(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_404, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k2_fdiff_1(X2,X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2926, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos),X1)|~v1_xreal_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos))|~r2_fdiff_1(k2_fdiff_1(k24_sin_cos,X1),X1)|~r2_fdiff_1(k24_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_2922, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos),X1)|~v1_xreal_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos))|~r2_fdiff_1(k2_fdiff_1(k24_sin_cos,X1),X1)|~r2_fdiff_1(k24_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_2924, plain, (r2_fdiff_1(k1_partfun1(X3,X4,X5,X6,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos),X2)|~v1_xreal_0(esk4_2(X2,k1_partfun1(X3,X4,X5,X6,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos)))|~v1_funct_1(k1_partfun1(X3,X4,X5,X6,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos))|~r2_hidden(esk4_2(X2,k1_partfun1(X3,X4,X5,X6,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos)),X1)|~r2_fdiff_1(k2_fdiff_1(k24_sin_cos,X1),X2)|~r2_fdiff_1(k24_sin_cos,X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(X3,X4,X5,X6,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos)))|~m1_subset_1(k1_partfun1(X3,X4,X5,X6,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_funct_1(k2_fdiff_1(k24_sin_cos,X1))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(k2_fdiff_1(k24_sin_cos,X1),k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_fdiff_1(k24_sin_cos,X1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2920, plain, (r2_fdiff_1(k1_partfun1(X3,X4,X5,X6,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos),X2)|~v1_xreal_0(esk4_2(X2,k1_partfun1(X3,X4,X5,X6,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos)))|~v1_funct_1(k1_partfun1(X3,X4,X5,X6,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos))|~r2_hidden(esk4_2(X2,k1_partfun1(X3,X4,X5,X6,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos)),X1)|~r2_fdiff_1(k2_fdiff_1(k24_sin_cos,X1),X2)|~r2_fdiff_1(k24_sin_cos,X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(X3,X4,X5,X6,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos)))|~m1_subset_1(k1_partfun1(X3,X4,X5,X6,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_funct_1(k2_fdiff_1(k24_sin_cos,X1))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(k2_fdiff_1(k24_sin_cos,X1),k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_fdiff_1(k24_sin_cos,X1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2793, plain, (v1_int_1(esk7_1(k5_numbers))|v1_xboole_0(k5_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2829, negated_conjecture, (v2_membered(esk1_0)|~v3_membered(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_408, plain, (v1_valued_0(k2_fdiff_1(X1,X2))|~v1_membered(k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2862, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),X1)|~m1_subset_1(k4_xcmplx_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)),esk1_0))).#trainneg
cnf(i_0_2937, plain, (k3_xcmplx_0(k1_seq_1(k24_sin_cos,X1),k1_seq_1(k24_sin_cos,k1_seq_1(k24_sin_cos,X1)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k24_sin_cos),X1)|~v1_xreal_0(k1_seq_1(k24_sin_cos,k1_seq_1(k24_sin_cos,X1)))|~r1_fdiff_1(k24_sin_cos,k1_seq_1(k24_sin_cos,X1))|~m1_subset_1(k1_seq_1(k24_sin_cos,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(k1_seq_1(k24_sin_cos,X1)))).#trainneg
cnf(i_0_2775, plain, (m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3),k1_numbers)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~r1_fdiff_1(X1,X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X3),k1_numbers)|~m1_subset_1(X3,k1_numbers)|~v1_xreal_0(k1_seq_1(X1,X3)))).#trainneg
cnf(i_0_2822, plain, (v3_valued_0(k9_xtuple_0(k2_zfmisc_1(k2_zfmisc_1(X1,k1_numbers),X2))))).#trainneg
cnf(i_0_1201, plain, (k1_real_1(k4_xcmplx_0(np__1))=np__1|~m1_subset_1(k4_xcmplx_0(np__1),k1_numbers))).#trainneg
cnf(i_0_409, plain, (v2_valued_0(k2_fdiff_1(X1,X2))|~v2_membered(k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_543, plain, (v1_int_1(k4_xcmplx_0(X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_417, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_2811, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos),X5)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos)))))).#trainneg
cnf(i_0_2583, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos),esk1_0)|~v1_int_1(esk4_2(esk1_0,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1932, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k19_sin_cos),X2)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k19_sin_cos))|~r2_hidden(esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k19_sin_cos)),X1)|~r2_fdiff_1(k2_fdiff_1(k19_sin_cos,X1),X2)|~r2_fdiff_1(k19_sin_cos,X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k19_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_xreal_0(esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k19_sin_cos)))|~m1_subset_1(k1_seq_1(k16_sin_cos,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k19_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_426, plain, (v1_xreal_0(esk3_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_658, negated_conjecture, (v1_xreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))))).#trainneg
cnf(i_0_659, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v1_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))))).#trainneg
cnf(i_0_427, plain, (v1_xcmplx_0(esk3_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|~v1_membered(k1_numbers)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_428, plain, (v7_ordinal1(esk3_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|~v6_membered(k1_numbers)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1911, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k24_sin_cos),X2)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k24_sin_cos))|~r2_hidden(esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k24_sin_cos)),X1)|~r2_fdiff_1(k2_fdiff_1(k19_sin_cos,X1),X2)|~r2_fdiff_1(k19_sin_cos,X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k24_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_xreal_0(esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k24_sin_cos)))|~m1_subset_1(k1_seq_1(k16_sin_cos,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k24_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_2579, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(esk4_2(esk1_0,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos)),k1_numbers)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2687, plain, (r2_fdiff_1(k3_relat_1(X5,k19_sin_cos),X6)|~v3_valued_0(X5)|~v1_funct_1(k3_relat_1(X5,k19_sin_cos))|~v1_funct_1(X5)|~r2_fdiff_1(X5,X6)|~m1_subset_1(k3_relat_1(X5,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(X5,k19_sin_cos))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6))).#trainneg
cnf(i_0_1943, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),X2)|~r1_fdiff_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))|~m1_subset_1(k4_xcmplx_0(k1_seq_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2000, plain, (k2_fdiff_1(X1,X2)=k2_fdiff_1(X3,X2)|~v1_xboole_0(X2)|~v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X3)|~r2_fdiff_1(X3,X2)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3004, negated_conjecture, (m1_subset_1(esk2_0,esk1_0))).#trainneg
cnf(i_0_3007, negated_conjecture, (v1_xcmplx_0(esk2_0))).#trainneg
cnf(i_0_429, plain, (v1_int_1(esk3_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|~v5_membered(k1_numbers)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3010, negated_conjecture, (v7_ordinal1(esk2_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_3011, negated_conjecture, (v1_int_1(esk2_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_3012, negated_conjecture, (v1_xxreal_0(esk2_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_430, plain, (v1_xxreal_0(esk3_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|~v2_membered(k1_numbers)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_433, plain, (m1_subset_1(esk3_3(X1,X2,X3),X2)|X3=k2_fdiff_1(X1,X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3005, negated_conjecture, (r1_fdiff_1(X1,esk2_0)|~v1_funct_1(X1)|~r2_fdiff_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3028, plain, (m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~r2_hidden(X2,X3)|~r2_fdiff_1(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3039, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos))|~r2_fdiff_1(k2_fdiff_1(k24_sin_cos,X1),X1)|~r2_fdiff_1(k24_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xcmplx_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos)))|~m1_subset_1(k4_xcmplx_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_3047, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos))|~r2_fdiff_1(k2_fdiff_1(k24_sin_cos,X1),X1)|~r2_fdiff_1(k24_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xcmplx_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos)))|~m1_subset_1(k4_xcmplx_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_3166, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k24_sin_cos),X1)|~v1_xreal_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k24_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k24_sin_cos))|~r2_fdiff_1(k2_fdiff_1(k19_sin_cos,X1),X1)|~r2_fdiff_1(k19_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k24_sin_cos)))|~m1_subset_1(k1_seq_1(k16_sin_cos,esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k24_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1561, plain, (m1_subset_1(k4_xcmplx_0(np__2),k1_numbers)|~v1_xcmplx_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_438, plain, (v1_funct_1(k3_relat_1(X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_660, negated_conjecture, (v7_ordinal1(esk1_0)|~v6_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))))).#trainneg
cnf(i_0_661, negated_conjecture, (v1_int_1(esk1_0)|~v5_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))))).#trainneg
cnf(i_0_626, plain, (v1_int_1(k1_real_1(X1))|~v1_int_1(X1)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_434, plain, (r1_fdiff_1(X1,esk3_3(X2,X3,X4))|X4=k2_fdiff_1(X2,X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|k1_relset_1(k1_numbers,X4)!=X3|~v1_funct_1(X4)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_439, plain, (v1_xreal_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3146, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k19_sin_cos),X1)|~v1_xreal_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k19_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k19_sin_cos))|~r2_fdiff_1(k2_fdiff_1(k19_sin_cos,X1),X1)|~r2_fdiff_1(k19_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k19_sin_cos)))|~m1_subset_1(k1_seq_1(k16_sin_cos,esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k19_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_3038, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos))|~r2_fdiff_1(k2_fdiff_1(k24_sin_cos,X1),X1)|~r2_fdiff_1(k24_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_3046, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos))|~r2_fdiff_1(k2_fdiff_1(k24_sin_cos,X1),X1)|~r2_fdiff_1(k24_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k24_sin_cos,X1),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_3027, plain, (v1_xcmplx_0(k1_fdiff_1(X1,X2))|~v1_funct_1(X1)|~r2_hidden(X2,X3)|~r2_fdiff_1(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3079, plain, (v1_xreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3))|~v1_xreal_0(k1_seq_1(X1,X3))|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~r1_fdiff_1(X1,X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X3),k1_numbers)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_440, plain, (v1_xcmplx_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1027, plain, (v1_int_1(k4_xcmplx_0(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_662, negated_conjecture, (v1_xxreal_0(esk1_0)|~v2_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))))).#trainneg
cnf(i_0_3160, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))))).#trainneg
cnf(i_0_595, plain, (v1_xcmplx_0(k8_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_441, plain, (v7_ordinal1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_442, plain, (v1_int_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3113, negated_conjecture, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_3060, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k24_sin_cos,X5),k19_sin_cos),X5)|~v1_xreal_0(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k24_sin_cos,X5),k19_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k24_sin_cos,X5),k19_sin_cos))|~v1_funct_1(k2_fdiff_1(k24_sin_cos,X5))|~r2_fdiff_1(k2_fdiff_1(k24_sin_cos,X5),X5)|~r2_fdiff_1(k24_sin_cos,X5)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k24_sin_cos,X5),k19_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k24_sin_cos,X5),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k24_sin_cos,X5),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k24_sin_cos,X5),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5))).#trainneg
cnf(i_0_2060, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),np__0)=k4_xcmplx_0(np__2)|~v1_xcmplx_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_443, plain, (v1_xxreal_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3031, negated_conjecture, (v1_membered(esk1_0)|~m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3065, negated_conjecture, (v2_membered(esk1_0)|~m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_596, plain, (v7_ordinal1(k8_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_597, plain, (v1_int_1(k8_real_1(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_444, plain, (v3_membered(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v3_membered(k2_zfmisc_1(X1,X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_445, plain, (v1_membered(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_membered(k2_zfmisc_1(X1,X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3055, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k24_sin_cos,X5),k24_sin_cos),X5)|~v1_xreal_0(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k24_sin_cos,X5),k24_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k24_sin_cos,X5),k24_sin_cos))|~v1_funct_1(k2_fdiff_1(k24_sin_cos,X5))|~r2_fdiff_1(k2_fdiff_1(k24_sin_cos,X5),X5)|~r2_fdiff_1(k24_sin_cos,X5)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k24_sin_cos,X5),k24_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k24_sin_cos,X5),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k24_sin_cos,X5),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k24_sin_cos,X5),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5))).#trainneg
cnf(i_0_3087, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos),X5)|~r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))|~v1_funct_1(X1)|~r2_hidden(esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)),X4)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),X5)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_xreal_0(k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)))))|~r1_fdiff_1(X3,k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)))))|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))))|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)))),k1_numbers)|~m1_subset_1(k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_3351, negated_conjecture, (m1_subset_1(k1_fdiff_1(X1,esk2_0),k1_numbers)|~v1_funct_1(X1)|~r2_fdiff_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3499, negated_conjecture, (v1_membered(esk1_0))).#trainneg
cnf(i_0_3502, negated_conjecture, (v2_membered(esk1_0))).#trainneg
cnf(i_0_3289, negated_conjecture, (v7_ordinal1(esk2_0)|~v6_membered(esk1_0))).#trainneg
cnf(i_0_446, plain, (v2_membered(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v2_membered(k2_zfmisc_1(X1,X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3290, negated_conjecture, (v1_int_1(esk2_0)|~v5_membered(esk1_0))).#trainneg
cnf(i_0_3291, negated_conjecture, (v1_xxreal_0(esk2_0)|~v2_membered(esk1_0))).#trainneg
cnf(i_0_3297, negated_conjecture, (v1_int_1(esk2_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_447, plain, (v3_valued_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v3_valued_0(k2_zfmisc_1(X1,X4))|~v1_relat_1(k2_zfmisc_1(X1,X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_448, plain, (v1_valued_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_valued_0(k2_zfmisc_1(X1,X4))|~v1_relat_1(k2_zfmisc_1(X1,X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3432, negated_conjecture, (v1_xcmplx_0(k1_fdiff_1(X1,esk2_0))|~v1_funct_1(X1)|~r2_fdiff_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3352, plain, (m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|v1_xboole_0(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X2,X3))).#trainneg
cnf(i_0_3487, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k24_sin_cos,X5),k19_sin_cos),X5)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k24_sin_cos,X5),k19_sin_cos))|~v1_funct_1(k2_fdiff_1(k24_sin_cos,X5))|~r2_fdiff_1(k2_fdiff_1(k24_sin_cos,X5),X5)|~r2_fdiff_1(k24_sin_cos,X5)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k24_sin_cos,X5),k19_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k24_sin_cos,X5),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k24_sin_cos,X5),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k24_sin_cos,X5),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_xcmplx_0(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k24_sin_cos,X5),k19_sin_cos)))|~m1_subset_1(k4_xcmplx_0(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k24_sin_cos,X5),k19_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_3490, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k24_sin_cos,X5),k19_sin_cos),X5)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k24_sin_cos,X5),k19_sin_cos))|~v1_funct_1(k2_fdiff_1(k24_sin_cos,X5))|~r2_fdiff_1(k2_fdiff_1(k24_sin_cos,X5),X5)|~r2_fdiff_1(k24_sin_cos,X5)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k24_sin_cos,X5),k19_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k24_sin_cos,X5),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k24_sin_cos,X5),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k24_sin_cos,X5),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k24_sin_cos,X5),k19_sin_cos)))))).#trainneg
cnf(i_0_3420, plain, (r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos),X1)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos))|~r2_fdiff_1(k2_fdiff_1(k24_sin_cos,X1),X1)|~r2_fdiff_1(k24_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos)))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k24_sin_cos,X1),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_funct_1(k2_fdiff_1(k24_sin_cos,X1))|~m1_subset_1(k2_fdiff_1(k24_sin_cos,X1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k2_fdiff_1(k24_sin_cos,X1),k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_3349, plain, (m1_subset_1(k1_fdiff_1(X1,esk4_2(X2,X3)),k1_numbers)|r2_fdiff_1(X3,X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X2,X3),k1_numbers)|~v1_funct_1(X3)|~r1_tarski(X2,k1_relset_1(k1_numbers,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2))).#trainneg
cnf(i_0_3370, plain, (v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_3378, plain, (m1_subset_1(k4_xcmplx_0(np__2),k1_numbers))).#trainneg
cnf(i_0_1201, plain, (k1_real_1(k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_3449, negated_conjecture, (k4_xcmplx_0(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k4_xcmplx_0(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0))),k1_numbers)|~v1_membered(k1_numbers)|~v1_xreal_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)))|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers))).#trainneg
cnf(i_0_449, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X1,k1_partfun1(X2,X4,X5,X3,X6,X7))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X3)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X4))))).#trainneg
cnf(i_0_3371, plain, (v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_3372, plain, (v7_ordinal1(k4_xcmplx_0(np__1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_3374, plain, (v1_xxreal_0(k4_xcmplx_0(np__1))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_598, plain, (v1_xxreal_0(k8_real_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_453, plain, (v1_valued_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_membered(X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_454, plain, (v2_valued_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v2_membered(X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3436, plain, (v1_xreal_0(k1_fdiff_1(k1_partfun1(X4,X5,X6,X7,X1,X2),X3))|~v1_xreal_0(k1_seq_1(X1,X3))|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~r1_fdiff_1(X1,X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X3),k1_numbers)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_3430, plain, (v1_xcmplx_0(k1_fdiff_1(X1,esk4_2(X2,X3)))|r2_fdiff_1(X3,X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X2,X3),k1_numbers)|~v1_funct_1(X3)|~r1_tarski(X2,k1_relset_1(k1_numbers,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2))).#trainneg
cnf(i_0_3415, plain, (r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k19_sin_cos,X1),k19_sin_cos),X1)|~v1_xreal_0(esk4_2(X1,k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k19_sin_cos,X1),k19_sin_cos)))|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k19_sin_cos,X1),k19_sin_cos))|~r2_fdiff_1(k2_fdiff_1(k19_sin_cos,X1),X1)|~r2_fdiff_1(k19_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k19_sin_cos,X1),k19_sin_cos)))|~m1_subset_1(k1_seq_1(k16_sin_cos,esk4_2(X1,k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k19_sin_cos,X1),k19_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k19_sin_cos,X1),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_funct_1(k2_fdiff_1(k19_sin_cos,X1))|~m1_subset_1(k2_fdiff_1(k19_sin_cos,X1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k2_fdiff_1(k19_sin_cos,X1),k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_3089, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos),X5)|~r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))|~v1_funct_1(X1)|~r2_hidden(esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)),X4)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),X5)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_xreal_0(k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)))))|~r1_fdiff_1(X3,k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)))))|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))))|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)))),k1_numbers)|~m1_subset_1(k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_3571, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k24_sin_cos)),X3),k24_sin_cos),X4)|~v1_xreal_0(k1_seq_1(X2,k1_seq_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k24_sin_cos)),X3),k24_sin_cos)))))|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k24_sin_cos)),X3),k24_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k24_sin_cos)),X3),k24_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k24_sin_cos)),X3),k24_sin_cos))|~v1_funct_1(X1)|~v1_funct_1(X2)|~r2_hidden(esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k24_sin_cos)),X3),k24_sin_cos)),X3)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k24_sin_cos)),X3),X4)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k24_sin_cos)),X3)|~r1_tarski(X4,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k24_sin_cos)),X3),k24_sin_cos)))|~m1_subset_1(k1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k24_sin_cos)),X3),k24_sin_cos)))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k24_sin_cos)),X3),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k24_sin_cos)),X3),k24_sin_cos))),k1_numbers)|~m1_subset_1(k1_seq_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k24_sin_cos)),X3),k24_sin_cos))),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X4))).#trainneg
cnf(i_0_3572, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k19_sin_cos)),X3),k24_sin_cos),X4)|~v1_xreal_0(k1_seq_1(X2,k1_seq_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k19_sin_cos)),X3),k24_sin_cos)))))|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k19_sin_cos)),X3),k24_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k19_sin_cos)),X3),k24_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k19_sin_cos)),X3),k24_sin_cos))|~v1_funct_1(X1)|~v1_funct_1(X2)|~r2_hidden(esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k19_sin_cos)),X3),k24_sin_cos)),X3)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k19_sin_cos)),X3),X4)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k19_sin_cos)),X3)|~r1_tarski(X4,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k19_sin_cos)),X3),k24_sin_cos)))|~m1_subset_1(k1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k19_sin_cos)),X3),k24_sin_cos)))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k19_sin_cos)),X3),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k19_sin_cos)),X3),k24_sin_cos))),k1_numbers)|~m1_subset_1(k1_seq_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k19_sin_cos)),X3),k24_sin_cos))),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X4))).#trainneg
cnf(i_0_2205, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos),X5)|~v1_xreal_0(k1_fdiff_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))))|~r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))|~v1_funct_1(X1)|~r2_hidden(esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)),X4)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),X5)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_xreal_0(k1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)))))|~r1_fdiff_1(X3,k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)))))|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))))|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(k1_fdiff_1(X3,k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))))),k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_3086, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos),X5)|~v1_xreal_0(k1_fdiff_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))))|~r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))|~v1_funct_1(X1)|~r2_hidden(esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)),X4)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),X5)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_xreal_0(k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)))))|~r1_fdiff_1(X3,k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)))))|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))))|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)))),k1_numbers)|~m1_subset_1(k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_2071, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),np__0)=k4_xcmplx_0(np__1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_3154, plain, (X1=k2_fdiff_1(k24_sin_cos,X2)|k1_seq_1(k24_sin_cos,esk3_3(k24_sin_cos,X2,X1))!=k1_seq_1(X1,esk3_3(k24_sin_cos,X2,X1))|k1_relset_1(k1_numbers,X1)!=X2|~v1_funct_1(X1)|~r2_fdiff_1(k24_sin_cos,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_455, plain, (v5_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6),X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3397, negated_conjecture, (v1_int_1(esk1_0)|~v6_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))))).#trainneg
cnf(i_0_3444, negated_conjecture, (v1_xxreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))))).#trainneg
cnf(i_0_1224, plain, (v1_int_1(k8_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_459, plain, (m1_subset_1(k3_relat_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_479, plain, (v1_xcmplx_0(esk4_2(X1,X2))|r2_fdiff_1(X2,X1)|~v1_membered(k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_2294, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos),X5)|~v1_xreal_0(k1_fdiff_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))))|~r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))|~v1_funct_1(X1)|~r2_hidden(esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)),X4)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),X5)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_xreal_0(k1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)))))|~r1_fdiff_1(X3,k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)))))|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))))|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(k1_fdiff_1(X3,k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))))),k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_729, plain, (v1_xreal_0(k1_fdiff_1(X1,X3))|~v3_valued_0(k2_fdiff_1(X1,X2))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~r2_hidden(X3,X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_2195, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos),X5)|~r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))|~v1_funct_1(X1)|~r2_hidden(esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)),X4)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),X5)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_xreal_0(k1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)))))|~r1_fdiff_1(X3,k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)))))|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))))|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(k1_fdiff_1(X3,k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))))),k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_3573, negated_conjecture, (v1_xreal_0(k1_fdiff_1(X1,esk2_0))|~v1_funct_1(X1)|~r2_fdiff_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3579, negated_conjecture, (v1_int_1(esk2_0)|~v6_membered(esk1_0))).#trainneg
cnf(i_0_634, plain, (k1_real_1(np__2)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_480, plain, (v7_ordinal1(esk4_2(X1,X2))|r2_fdiff_1(X2,X1)|~v6_membered(k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_3637, plain, (v1_xreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_3603, negated_conjecture, (v1_xcmplx_0(k1_seq_1(k24_sin_cos,esk2_0))|~r2_fdiff_1(k24_sin_cos,esk1_0))).#trainneg
cnf(i_0_3578, negated_conjecture, (m1_subset_1(k1_seq_1(k24_sin_cos,esk2_0),k1_numbers)|~r2_fdiff_1(k24_sin_cos,esk1_0))).#trainneg
cnf(i_0_481, plain, (v1_int_1(esk4_2(X1,X2))|r2_fdiff_1(X2,X1)|~v5_membered(k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_482, plain, (v1_xxreal_0(esk4_2(X1,X2))|r2_fdiff_1(X2,X1)|~v2_membered(k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_3786, plain, (k2_fdiff_1(k24_sin_cos,X1)=k24_sin_cos|k1_relset_1(k1_numbers,k24_sin_cos)!=X1|~r2_fdiff_1(k24_sin_cos,X1))).#trainneg
cnf(i_0_3628, plain, (v1_xreal_0(k1_fdiff_1(X1,esk4_2(X2,X3)))|r2_fdiff_1(X3,X2)|~v1_funct_1(X1)|~v1_funct_1(X3)|~r2_fdiff_1(X1,X2)|~r1_tarski(X2,k1_relset_1(k1_numbers,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2))).#trainneg
cnf(i_0_3694, plain, (v1_xreal_0(k1_fdiff_1(k3_relat_1(X5,X6),X7))|~v1_xreal_0(k1_seq_1(X5,X7))|~r1_fdiff_1(X6,k1_seq_1(X5,X7))|~r1_fdiff_1(X5,X7)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k1_fdiff_1(X5,X7),k1_numbers)|~m1_subset_1(X7,k1_numbers))).#trainneg
cnf(i_0_3850, plain, (v1_xreal_0(k1_fdiff_1(X1,X2))|v1_xboole_0(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X2,X3))).#trainneg
cnf(i_0_1515, plain, (v1_xreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),X2))|~v1_xreal_0(k1_seq_1(k24_sin_cos,k1_seq_1(X1,X2)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2893, plain, (v4_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)),X1))).#trainneg
cnf(i_0_2284, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos),X5)|~r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))|~v1_funct_1(X1)|~r2_hidden(esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)),X4)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),X5)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_xreal_0(k1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)))))|~r1_fdiff_1(X3,k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)))))|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))))|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(k1_fdiff_1(X3,k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))))),k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_1149, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos),X2)|~r1_fdiff_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3088, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos),X5)|~v1_xreal_0(k1_fdiff_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))))|~r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))|~v1_funct_1(X1)|~r2_hidden(esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)),X4)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),X5)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_xreal_0(k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)))))|~r1_fdiff_1(X3,k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)))))|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))))|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)))),k1_numbers)|~m1_subset_1(k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_1157, plain, (k1_real_1(k4_xcmplx_0(np__2))=np__2)).#trainneg
cnf(i_0_485, plain, (r1_tarski(X1,X3)|~v1_funct_1(k2_fdiff_1(X2,X3))|~r2_fdiff_1(k2_fdiff_1(X2,X3),X1)|~m1_subset_1(k2_fdiff_1(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3638, plain, (v1_xcmplx_0(k4_xcmplx_0(np__2))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_486, plain, (r2_fdiff_1(k2_fdiff_1(X1,X2),X3)|m1_subset_1(esk4_2(X3,k2_fdiff_1(X1,X2)),k1_numbers)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r1_tarski(X3,X2)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3367, plain, (r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k19_sin_cos,X1),k24_sin_cos),X1)|~v1_xreal_0(esk4_2(X1,k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k19_sin_cos,X1),k24_sin_cos)))|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k19_sin_cos,X1),k24_sin_cos))|~r2_fdiff_1(k2_fdiff_1(k19_sin_cos,X1),X1)|~r2_fdiff_1(k19_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k19_sin_cos,X1),k24_sin_cos)))|~m1_subset_1(k1_seq_1(k16_sin_cos,esk4_2(X1,k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k19_sin_cos,X1),k24_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k19_sin_cos,X1),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_funct_1(k2_fdiff_1(k19_sin_cos,X1))|~m1_subset_1(k2_fdiff_1(k19_sin_cos,X1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k2_fdiff_1(k19_sin_cos,X1),k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1375, plain, (r1_fdiff_1(k3_relat_1(X1,k16_sin_cos),X2)|~r1_fdiff_1(k16_sin_cos,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_627, plain, (v1_int_1(k1_seq_1(k16_sin_cos,X1))|~v1_int_1(k1_fdiff_1(k19_sin_cos,X1))|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_3639, plain, (v7_ordinal1(k4_xcmplx_0(np__2))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_3641, plain, (v1_xxreal_0(k4_xcmplx_0(np__2))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_496, plain, (m1_subset_1(esk4_2(k5_numbers,X1),k1_numbers)|r2_fdiff_1(X1,k5_numbers)|~v1_funct_1(X1)|~r1_tarski(k5_numbers,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(k5_numbers))).#trainneg
cnf(i_0_1063, plain, (r2_fdiff_1(k19_sin_cos,k1_relset_1(k1_numbers,k19_sin_cos))|~v3_rcomp_1(k1_relset_1(k1_numbers,k19_sin_cos)))).#trainneg
cnf(i_0_3157, negated_conjecture, (v1_xreal_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3503, plain, (v1_int_1(k8_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_506, plain, (r2_fdiff_1(k2_fdiff_1(X1,X2),X3)|m1_subset_1(esk4_2(X3,k2_fdiff_1(X1,X2)),X3)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r1_tarski(X3,X2)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3390, plain, (v1_funct_1(k3_relat_1(X1,k24_sin_cos))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_639, plain, (v1_xreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3))|~v1_xreal_0(k1_fdiff_1(X2,X3))|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(X2,X3)),k1_numbers)|~r1_fdiff_1(X1,k1_seq_1(X2,X3))|~r1_fdiff_1(X2,X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_3698, plain, (v1_xcmplx_0(k1_seq_1(k24_sin_cos,esk4_2(X1,X2)))|r2_fdiff_1(X2,X1)|~v1_funct_1(X2)|~r2_fdiff_1(k24_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xreal_0(esk4_2(X1,X2)))).#trainneg
cnf(i_0_743, negated_conjecture, (r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4016, plain, (r1_tarski(X1,X2)|~r2_fdiff_1(k24_sin_cos,X1)|~r2_fdiff_1(k24_sin_cos,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|k1_relset_1(k1_numbers,k24_sin_cos)!=X2)).#trainneg
cnf(i_0_2071, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),np__0)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_508, plain, (r2_fdiff_1(X1,k1_relset_1(k1_numbers,X1))|m1_subset_1(esk4_2(k1_relset_1(k1_numbers,X1),X1),k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,X1)))).#trainneg
cnf(i_0_3877, plain, (v1_xcmplx_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_3868, negated_conjecture, (v1_xreal_0(k1_seq_1(k24_sin_cos,esk2_0))|~r2_fdiff_1(k24_sin_cos,esk1_0))).#trainneg
cnf(i_0_4021, plain, (v1_xreal_0(esk4_2(X1,k2_fdiff_1(X2,X3)))|r2_fdiff_1(k2_fdiff_1(X2,X3),X1)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~r1_tarski(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_4044, plain, (r1_fdiff_1(k3_relat_1(k24_sin_cos,k16_sin_cos),X1)|~r1_fdiff_1(k16_sin_cos,k1_seq_1(k24_sin_cos,X1))|~r1_fdiff_1(k24_sin_cos,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4045, plain, (r1_fdiff_1(k3_relat_1(k19_sin_cos,k16_sin_cos),X1)|~r1_fdiff_1(k16_sin_cos,k1_seq_1(k19_sin_cos,X1))|~r1_fdiff_1(k19_sin_cos,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2196, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos),X5)|~r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))|~v1_funct_1(X1)|~r2_hidden(esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)),X4)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),X5)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~r1_fdiff_1(X3,k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)))))|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))))|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X3,k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))))),k1_numbers)|~m1_subset_1(k1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos)))),k1_numbers)|~m1_subset_1(k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k24_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_4128, plain, (v1_funct_1(k3_relat_1(k16_sin_cos,k24_sin_cos)))).#trainneg
cnf(i_0_377, plain, (k6_xcmplx_0(np__1,k4_xcmplx_0(X1))=k6_xcmplx_0(X1,k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_519, plain, (v1_xcmplx_0(k1_fdiff_1(k19_sin_cos,X1))|~v1_int_1(k1_seq_1(k16_sin_cos,X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_4133, plain, (v1_funct_1(k3_relat_1(k24_sin_cos,k24_sin_cos)))).#trainneg
cnf(i_0_4134, plain, (v1_funct_1(k3_relat_1(k19_sin_cos,k24_sin_cos)))).#trainneg
cnf(i_0_3447, negated_conjecture, (v1_xcmplx_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3887, negated_conjecture, (v7_ordinal1(k1_seq_1(k24_sin_cos,esk2_0))|~v6_membered(k1_numbers)|~r2_fdiff_1(k24_sin_cos,esk1_0))).#trainneg
cnf(i_0_520, plain, (v1_int_1(k1_fdiff_1(k19_sin_cos,X1))|~v1_int_1(k1_seq_1(k16_sin_cos,X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_522, plain, (k6_xcmplx_0(k1_fdiff_1(k19_sin_cos,X1),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,k1_seq_1(k16_sin_cos,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_3026, plain, (v1_xcmplx_0(k1_fdiff_1(k19_sin_cos,X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2285, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos),X5)|~r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))|~v1_funct_1(X1)|~r2_hidden(esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)),X4)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),X5)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~r1_fdiff_1(X3,k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)))))|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))))|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X3,k1_seq_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))))),k1_numbers)|~m1_subset_1(k1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos)))),k1_numbers)|~m1_subset_1(k1_seq_1(X1,esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4),k19_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_2160, plain, (~v1_xboole_0(X1)|~v1_funct_1(X2)|~v1_funct_1(X3)|~r2_hidden(X4,k1_relset_1(X1,k1_partfun1(X8,X9,X10,X11,X3,X2)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X5)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X10,X11)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X8,X9))))).#trainneg
cnf(i_0_1314, plain, (v3_membered(k1_relset_1(X1,k2_zfmisc_1(X1,X2)))|~v3_membered(X1))).#trainneg
cnf(i_0_3888, negated_conjecture, (v1_int_1(k1_seq_1(k24_sin_cos,esk2_0))|~v5_membered(k1_numbers)|~r2_fdiff_1(k24_sin_cos,esk1_0))).#trainneg
cnf(i_0_373, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),np__1)=k6_xcmplx_0(k4_xcmplx_0(np__1),X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_523, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k1_fdiff_1(k19_sin_cos,X2))=k6_xcmplx_0(k1_seq_1(k16_sin_cos,X2),X1)|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(k1_seq_1(k16_sin_cos,X2),k1_numbers))).#trainneg
cnf(i_0_3889, negated_conjecture, (v1_xxreal_0(k1_seq_1(k24_sin_cos,esk2_0))|~v2_membered(k1_numbers)|~r2_fdiff_1(k24_sin_cos,esk1_0))).#trainneg
cnf(i_0_3796, negated_conjecture, (v1_xxreal_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_550, plain, (m1_subset_1(esk4_2(esk7_1(k1_zfmisc_1(X1)),X2),X1)|r2_fdiff_1(X2,esk7_1(k1_zfmisc_1(X1)))|~v1_funct_1(X2)|~r1_tarski(esk7_1(k1_zfmisc_1(X1)),k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk7_1(k1_zfmisc_1(X1)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(esk7_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_551, plain, (m1_subset_1(esk3_3(X1,esk7_1(k1_zfmisc_1(X2)),X3),X2)|X3=k2_fdiff_1(X1,esk7_1(k1_zfmisc_1(X2)))|k1_relset_1(k1_numbers,X3)!=esk7_1(k1_zfmisc_1(X2))|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,esk7_1(k1_zfmisc_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4031, plain, (r2_fdiff_1(k3_relat_1(k2_fdiff_1(k19_sin_cos,X5),k24_sin_cos),X5)|~v1_xreal_0(esk4_2(X5,k3_relat_1(k2_fdiff_1(k19_sin_cos,X5),k24_sin_cos)))|~v1_funct_1(k3_relat_1(k2_fdiff_1(k19_sin_cos,X5),k24_sin_cos))|~v1_funct_1(k2_fdiff_1(k19_sin_cos,X5))|~r2_fdiff_1(k2_fdiff_1(k19_sin_cos,X5),X5)|~r2_fdiff_1(k19_sin_cos,X5)|~r1_tarski(X5,k1_relset_1(k1_numbers,k3_relat_1(k2_fdiff_1(k19_sin_cos,X5),k24_sin_cos)))|~m1_subset_1(k1_seq_1(k16_sin_cos,esk4_2(X5,k3_relat_1(k2_fdiff_1(k19_sin_cos,X5),k24_sin_cos))),k1_numbers)|~m1_subset_1(k3_relat_1(k2_fdiff_1(k19_sin_cos,X5),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k19_sin_cos,X5),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k19_sin_cos,X5),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5))).#trainneg
cnf(i_0_1058, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k2_fdiff_1(X6,X7)),k1_zfmisc_1(k2_zfmisc_1(X8,k1_numbers)))|~v1_funct_1(k2_fdiff_1(X6,X7))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X8,X9)))|~m1_subset_1(k2_fdiff_1(X6,X7),k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1142, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos),X2)|~r1_fdiff_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k24_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1315, plain, (v1_membered(k1_relset_1(X1,k2_zfmisc_1(X1,X2)))|~v1_membered(X1))).#trainneg
cnf(i_0_552, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk7_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,esk7_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_4258, negated_conjecture, (v1_int_1(k1_seq_1(k24_sin_cos,esk2_0))|~v6_membered(k1_numbers)|~r2_fdiff_1(k24_sin_cos,esk1_0))).#trainneg
cnf(i_0_1330, plain, (r2_fdiff_1(k24_sin_cos,k1_relset_1(k1_numbers,k24_sin_cos))|~v3_rcomp_1(k1_relset_1(k1_numbers,k24_sin_cos)))).#trainneg
cnf(i_0_1626, plain, (m1_subset_1(k1_relset_1(k1_numbers,k19_sin_cos),k1_numbers)|~v1_xcmplx_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_587, plain, (k6_xcmplx_0(X2,k1_real_1(X1))=k6_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_589, plain, (k6_xcmplx_0(k1_real_1(X2),X1)=k6_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k1_real_1(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4171, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4166, plain, (m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4168, plain, (k3_xcmplx_0(np__1,k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4177, negated_conjecture, (r1_tarski(esk1_0,X1)|k1_relset_1(k1_numbers,k24_sin_cos)!=X1|~r2_fdiff_1(k24_sin_cos,esk1_0)|~r2_fdiff_1(k24_sin_cos,X1))).#trainneg
cnf(i_0_2060, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),np__0)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_376, plain, (k6_xcmplx_0(np__2,k4_xcmplx_0(X1))=k6_xcmplx_0(X1,k4_xcmplx_0(np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_605, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1018, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|v1_xboole_0(k16_sin_cos)|~m1_subset_1(X1,k16_sin_cos))).#trainneg
cnf(i_0_1690, plain, (m1_subset_1(k1_relset_1(k1_numbers,k19_sin_cos),k1_numbers)|~v1_xcmplx_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1734, plain, (m1_subset_1(k1_relset_1(k1_numbers,k16_sin_cos),k1_numbers)|~v1_xcmplx_0(k1_relset_1(k1_numbers,k16_sin_cos))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2033, plain, (v1_int_1(k1_real_1(X1))|~v1_int_1(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_606, plain, (v7_ordinal1(k3_xcmplx_0(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_607, plain, (v1_int_1(k3_xcmplx_0(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4483, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4605, plain, (k3_xcmplx_0(k1_real_1(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4622, plain, (k3_xcmplx_0(np__1,k1_seq_1(k16_sin_cos,X1))=k1_seq_1(k16_sin_cos,X1)|~v1_xcmplx_0(k1_fdiff_1(k19_sin_cos,X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_4647, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(X1))|k1_relset_1(k1_numbers,k24_sin_cos)!=X1|~r2_fdiff_1(k24_sin_cos,esk1_0)|~r2_fdiff_1(k24_sin_cos,X1))).#trainneg
cnf(i_0_4491, plain, (m1_subset_1(k4_xcmplx_0(esk5_0),k1_numbers))).#trainneg
cnf(i_0_4617, plain, (k3_xcmplx_0(np__1,np__0)=np__0)).#trainneg
cnf(i_0_372, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),np__2)=k6_xcmplx_0(k4_xcmplx_0(np__2),X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_608, plain, (v1_xxreal_0(k3_xcmplx_0(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4618, plain, (k3_xcmplx_0(np__1,np__2)=np__2)).#trainneg
cnf(i_0_4497, negated_conjecture, (m1_subset_1(k4_xcmplx_0(esk2_0),k1_numbers))).#trainneg
cnf(i_0_4619, plain, (k3_xcmplx_0(np__1,np__1)=np__1)).#trainneg
cnf(i_0_1296, plain, (v1_int_1(k1_seq_1(X1,X2))|~v6_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_610, plain, (m1_subset_1(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)),k1_numbers)|~v1_xreal_0(X3)|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4793, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_4795, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k1_numbers))).#trainneg
cnf(i_0_4972, negated_conjecture, (v1_xreal_0(k4_xcmplx_0(esk2_0)))).#trainneg
cnf(i_0_4891, plain, (v1_xreal_0(k4_xcmplx_0(esk5_0)))).#trainneg
cnf(i_0_379, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,k4_xcmplx_0(np__1)))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_612, plain, (m1_subset_1(k3_xcmplx_0(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4892, plain, (v1_xcmplx_0(k4_xcmplx_0(esk5_0))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_4893, plain, (v7_ordinal1(k4_xcmplx_0(esk5_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_4724, plain, (v7_ordinal1(X1)|~v6_membered(k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4928, plain, (v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_615, plain, (m1_subset_1(k1_relset_1(k1_numbers,k2_fdiff_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4817, plain, (k3_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(X1,X2)),k1_fdiff_1(X3,k1_seq_1(X1,X2)))=k4_xcmplx_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X3),X2))|~v1_xcmplx_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X3),X2))|~v1_xcmplx_0(k1_fdiff_1(X3,k1_seq_1(X1,X2)))|~v1_xcmplx_0(k1_fdiff_1(X1,X2))|~v1_xreal_0(k1_fdiff_1(X1,X2))|~r1_fdiff_1(X3,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(k1_fdiff_1(X3,k1_seq_1(X1,X2)),k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4863, plain, (k3_xcmplx_0(k1_real_1(X1),X2)=k3_xcmplx_0(X1,k1_real_1(X2))|~v1_xcmplx_0(k1_real_1(X1))|~v1_xcmplx_0(k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4810, plain, (k3_xcmplx_0(np__1,X1)=k4_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5179, plain, (v1_xcmplx_0(k4_xcmplx_0(esk5_0)))).#trainneg
cnf(i_0_4895, plain, (v1_xxreal_0(k4_xcmplx_0(esk5_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_4894, plain, (v1_int_1(k4_xcmplx_0(esk5_0))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_4973, negated_conjecture, (v1_xcmplx_0(k4_xcmplx_0(esk2_0))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_4974, negated_conjecture, (v7_ordinal1(k4_xcmplx_0(esk2_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_4975, negated_conjecture, (v1_int_1(k4_xcmplx_0(esk2_0))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_5258, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers))).#trainneg
cnf(i_0_5277, plain, (k3_xcmplx_0(X1,k1_real_1(X2))=k3_xcmplx_0(X2,k1_real_1(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_real_1(X1))|~v1_xcmplx_0(k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_5135, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5321, plain, (k3_xcmplx_0(X1,np__1)=k4_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5136, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_5354, negated_conjecture, (v1_xcmplx_0(k4_xcmplx_0(esk2_0)))).#trainneg
cnf(i_0_4976, negated_conjecture, (v1_xxreal_0(k4_xcmplx_0(esk2_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_5180, plain, (v1_int_1(k4_xcmplx_0(esk5_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_5355, negated_conjecture, (v1_int_1(k4_xcmplx_0(esk2_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_4757, plain, (v1_int_1(X1)|~v5_membered(k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5181, plain, (v1_int_1(X1)|~v6_membered(k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_652, plain, (v7_ordinal1(k1_fdiff_1(k19_sin_cos,X1))|~v6_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_653, plain, (v1_int_1(k1_fdiff_1(k19_sin_cos,X1))|~v5_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_5307, plain, (k3_xcmplx_0(k1_real_1(X1),k1_real_1(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k1_real_1(X2))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5267, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k1_real_1(X1)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4906, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),np__0)=np__0)).#trainneg
cnf(i_0_654, plain, (v1_xxreal_0(k1_fdiff_1(k19_sin_cos,X1))|~v2_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_4643, negated_conjecture, (r2_fdiff_1(k19_sin_cos,esk1_0)|k1_relset_1(k1_numbers,k19_sin_cos)!=k1_relset_1(k1_numbers,k24_sin_cos)|~r2_fdiff_1(k24_sin_cos,esk1_0)|~r2_fdiff_1(k24_sin_cos,k1_relset_1(k1_numbers,k19_sin_cos)))).#trainneg
cnf(i_0_4488, negated_conjecture, (m1_subset_1(k4_xcmplx_0(esk7_1(esk1_0)),k1_numbers)|~v1_xcmplx_0(esk7_1(esk1_0)))).#trainneg
cnf(i_0_4498, negated_conjecture, (m1_subset_1(k4_xcmplx_0(esk1_0),k1_numbers)|~v1_xcmplx_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1947, plain, (v1_int_1(k1_seq_1(X1,X2))|~v4_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_666, negated_conjecture, (v3_valued_0(esk1_0)|~v3_valued_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~v1_relat_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_667, negated_conjecture, (v1_valued_0(esk1_0)|~v1_valued_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~v1_relat_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_5309, plain, (k3_xcmplx_0(X1,k1_real_1(X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_5301, plain, (k3_xcmplx_0(np__1,k1_real_1(X1))=k3_xcmplx_0(k4_xcmplx_0(np__1),X1)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5305, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k1_real_1(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k1_real_1(X2))|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5687, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k1_real_1(X1)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5657, plain, (k3_xcmplx_0(k1_real_1(X1),X2)=k3_xcmplx_0(X1,k1_real_1(X2))|~v1_xcmplx_0(X2)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_5737, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),X1)=k1_real_1(X1)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5554, plain, (v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5688, plain, (k3_xcmplx_0(np__0,k4_xcmplx_0(np__1))=np__0)).#trainneg
cnf(i_0_670, plain, (v1_xreal_0(esk4_2(k1_relset_1(k1_numbers,X1),X1))|r2_fdiff_1(X1,k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,X1)))).#trainneg
cnf(i_0_5723, negated_conjecture, (v1_xreal_0(k4_xcmplx_0(esk1_0))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5724, negated_conjecture, (v1_xcmplx_0(k4_xcmplx_0(esk1_0))|~v1_membered(k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4446, plain, (m1_subset_1(k1_relset_1(k1_numbers,k19_sin_cos),k1_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers))|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_671, plain, (v1_xcmplx_0(esk4_2(k1_relset_1(k1_numbers,X1),X1))|r2_fdiff_1(X1,k1_relset_1(k1_numbers,X1))|~v1_membered(k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,X1)))).#trainneg
cnf(i_0_672, plain, (v7_ordinal1(esk4_2(k1_relset_1(k1_numbers,X1),X1))|r2_fdiff_1(X1,k1_relset_1(k1_numbers,X1))|~v6_membered(k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,X1)))).#trainneg
cnf(i_0_5659, plain, (k3_xcmplx_0(k1_real_1(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5826, plain, (k3_xcmplx_0(k1_real_1(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5670, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k1_real_1(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k1_real_1(X2))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5956, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),X1)=k1_real_1(X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6075, plain, (k3_xcmplx_0(np__0,np__1)=np__0)).#trainneg
cnf(i_0_673, plain, (v1_int_1(esk4_2(k1_relset_1(k1_numbers,X1),X1))|r2_fdiff_1(X1,k1_relset_1(k1_numbers,X1))|~v5_membered(k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,X1)))).#trainneg
cnf(i_0_6094, negated_conjecture, (v1_xcmplx_0(k4_xcmplx_0(esk1_0))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5725, negated_conjecture, (v7_ordinal1(k4_xcmplx_0(esk1_0))|~v6_membered(k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_674, plain, (v1_xxreal_0(esk4_2(k1_relset_1(k1_numbers,X1),X1))|r2_fdiff_1(X1,k1_relset_1(k1_numbers,X1))|~v2_membered(k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,X1)))).#trainneg
cnf(i_0_675, plain, (r2_fdiff_1(k2_fdiff_1(X1,X2),X2)|m1_subset_1(esk4_2(X2,k2_fdiff_1(X1,X2)),k1_numbers)|~v1_funct_1(k2_fdiff_1(X1,X2))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_6027, plain, (k3_xcmplx_0(X1,k1_real_1(X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6148, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5655, plain, (k3_xcmplx_0(k1_real_1(X1),X2)=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~m1_subset_1(k4_xcmplx_0(X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_5923, plain, (k3_xcmplx_0(X1,k1_real_1(X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_676, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1043, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|v1_xboole_0(k24_sin_cos)|~m1_subset_1(X1,k24_sin_cos))).#trainneg
cnf(i_0_5726, negated_conjecture, (v1_int_1(k4_xcmplx_0(esk1_0))|~v5_membered(k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5727, negated_conjecture, (v1_xxreal_0(k4_xcmplx_0(esk1_0))|~v2_membered(k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_677, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_678, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_6431, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6236, plain, (k3_xcmplx_0(X1,k1_real_1(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k1_real_1(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6138, plain, (k3_xcmplx_0(k1_real_1(X1),k1_fdiff_1(k19_sin_cos,X2))=k3_xcmplx_0(X1,k1_seq_1(k16_sin_cos,X2))|~m1_subset_1(k1_seq_1(k16_sin_cos,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_5193, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k2_fdiff_1(X2,X3)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_405, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k2_fdiff_1(X2,X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_679, plain, (v1_int_1(X1)|~v5_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1046, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|v1_xboole_0(k19_sin_cos)|~m1_subset_1(X1,k19_sin_cos))).#trainneg
cnf(i_0_6335, negated_conjecture, (v1_int_1(k4_xcmplx_0(esk1_0))|~v6_membered(k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5716, negated_conjecture, (m1_subset_1(k4_xcmplx_0(esk7_1(esk1_0)),k1_numbers))).#trainneg
cnf(i_0_680, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_5915, plain, (k3_xcmplx_0(X1,k1_real_1(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k1_real_1(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6396, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6857, plain, (k3_xcmplx_0(np__1,k4_xcmplx_0(X1))=k1_real_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6724, plain, (m1_subset_1(X1,k1_numbers)|~v1_funct_1(X2)|~r2_hidden(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_fdiff_1(X2,X3))).#trainneg
cnf(i_0_5803, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~m1_subset_1(k3_xcmplx_0(k4_xcmplx_0(X1),X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6719, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k1_fdiff_1(k19_sin_cos,X2))=k3_xcmplx_0(X1,k1_seq_1(k16_sin_cos,X2))|~v1_xreal_0(X2)|~m1_subset_1(k1_seq_1(k16_sin_cos,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6296, plain, (k1_real_1(k3_xcmplx_0(X1,X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers))).#trainneg
cnf(i_0_6741, negated_conjecture, (v1_xreal_0(k4_xcmplx_0(esk7_1(esk1_0))))).#trainneg
cnf(i_0_7110, negated_conjecture, (k3_xcmplx_0(k4_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0))),k1_seq_1(k16_sin_cos,esk2_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(k1_seq_1(k16_sin_cos,esk2_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers)|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,esk2_0))|~v1_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)))|~m1_subset_1(k3_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_682, plain, (v1_membered(X1)|~v1_membered(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_6742, negated_conjecture, (v1_xcmplx_0(k4_xcmplx_0(esk7_1(esk1_0)))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_6743, negated_conjecture, (v7_ordinal1(k4_xcmplx_0(esk7_1(esk1_0)))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_683, plain, (v2_membered(X1)|~v2_membered(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_684, plain, (v3_valued_0(X1)|~v3_valued_0(k1_relset_1(k1_numbers,X2))|~v1_relat_1(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_7161, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),np__1)=k1_real_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7005, plain, (k3_xcmplx_0(X1,k1_real_1(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k1_real_1(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6894, plain, (k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))=k1_real_1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6904, plain, (k3_xcmplx_0(np__1,k1_fdiff_1(k19_sin_cos,X1))=k1_real_1(k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7306, plain, (k3_xcmplx_0(np__2,np__1)=np__2)).#trainneg
cnf(i_0_7190, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k4_xcmplx_0(k1_seq_1(k16_sin_cos,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,esk2_0))|~v1_xreal_0(k1_seq_1(k16_sin_cos,esk2_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers))).#trainneg
cnf(i_0_685, plain, (v1_valued_0(X1)|~v1_valued_0(k1_relset_1(k1_numbers,X2))|~v1_relat_1(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_6744, negated_conjecture, (v1_int_1(k4_xcmplx_0(esk7_1(esk1_0)))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_6745, negated_conjecture, (v1_xxreal_0(k4_xcmplx_0(esk7_1(esk1_0)))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_686, plain, (m1_subset_1(X1,k1_relset_1(k1_numbers,X2))|~r2_hidden(X1,X3)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3))).#trainneg
cnf(i_0_688, plain, (m1_subset_1(X1,k1_zfmisc_1(X3))|~v1_funct_1(k2_fdiff_1(X2,X3))|~r2_fdiff_1(k2_fdiff_1(X2,X3),X1)|~m1_subset_1(k2_fdiff_1(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_7311, plain, (k1_real_1(k1_real_1(X1))=k3_xcmplx_0(X1,np__1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7312, plain, (k3_xcmplx_0(k1_fdiff_1(k19_sin_cos,X1),np__1)=k1_real_1(k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7428, plain, (k1_real_1(k4_xcmplx_0(X1))=k4_xcmplx_0(k1_real_1(X1))|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7437, plain, (k1_real_1(k1_real_1(X1))=k4_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7623, plain, (k3_xcmplx_0(X1,k1_real_1(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_7551, plain, (k3_xcmplx_0(X1,np__1)=X1|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7513, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_fdiff_1(k24_sin_cos,X1)|~r2_fdiff_1(k24_sin_cos,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|k1_relset_1(k1_numbers,k24_sin_cos)!=X2)).#trainneg
cnf(i_0_7492, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_fdiff_1(k19_sin_cos,esk2_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,esk2_0))|~v1_xreal_0(k1_seq_1(k16_sin_cos,esk2_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers))).#trainneg
cnf(i_0_7198, negated_conjecture, (v1_int_1(k4_xcmplx_0(esk7_1(esk1_0)))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_6095, plain, (m1_subset_1(k1_relset_1(k1_numbers,k19_sin_cos),k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_709, plain, (m1_subset_1(esk7_1(X1),k1_numbers)|~v1_xcmplx_0(esk7_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_7727, plain, (k4_xcmplx_0(k1_real_1(k4_xcmplx_0(X1)))=k1_real_1(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7728, plain, (k4_xcmplx_0(k1_real_1(k4_xcmplx_0(X1)))=k1_real_1(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7872, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7831, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers))).#trainneg
cnf(i_0_8101, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_real_1(X1)))=k1_real_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7896, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk2_0)|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,esk2_0))|~v1_xreal_0(k1_seq_1(k16_sin_cos,esk2_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers)|~v1_xreal_0(k1_fdiff_1(k19_sin_cos,esk2_0))|~v1_xreal_0(k1_seq_1(k19_sin_cos,esk2_0))|~r1_fdiff_1(k19_sin_cos,esk2_0))).#trainneg
cnf(i_0_714, plain, (m1_subset_1(k1_funct_1(X1,X2),k1_numbers)|~v1_xcmplx_0(k1_funct_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7957, plain, (k1_real_1(k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X1,k1_real_1(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X2),k1_numbers))).#trainneg
cnf(i_0_721, negated_conjecture, (v7_ordinal1(X1)|~v6_membered(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_7901, plain, (v1_xreal_0(k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4500, plain, (m1_subset_1(k4_xcmplx_0(k5_numbers),k1_numbers)|~v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_722, negated_conjecture, (v1_int_1(X1)|~v5_membered(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_723, negated_conjecture, (v1_xxreal_0(X1)|~v2_membered(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_7651, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k3_xcmplx_0(k4_xcmplx_0(X1),X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8140, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k16_sin_cos,esk2_0),k1_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_8425, negated_conjecture, (~v1_xcmplx_0(k1_seq_1(k16_sin_cos,esk2_0))|~v1_xreal_0(k1_seq_1(k16_sin_cos,esk2_0))|~v1_xreal_0(k1_fdiff_1(k19_sin_cos,esk2_0))|~v1_xreal_0(k1_seq_1(k19_sin_cos,esk2_0))|~r1_fdiff_1(k19_sin_cos,esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_8526, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k1_real_1(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8680, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k16_sin_cos,esk2_0),k4_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers))).#trainneg
cnf(i_0_8256, plain, (k3_xcmplx_0(np__1,k1_real_1(X1))=k1_real_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4968, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),np__2)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_744, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos))))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_8561, plain, (v1_xreal_0(k4_xcmplx_0(k5_numbers))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8562, plain, (v1_xcmplx_0(k4_xcmplx_0(k5_numbers))|~v1_membered(k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7902, plain, (v1_xcmplx_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v1_membered(k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_727, plain, (k2_fdiff_1(X1,X2)=k2_fdiff_1(X3,X4)|k1_fdiff_1(X3,esk3_3(X3,X4,k2_fdiff_1(X1,X2)))!=k1_fdiff_1(X1,esk3_3(X3,X4,k2_fdiff_1(X1,X2)))|k1_relset_1(k1_numbers,k2_fdiff_1(X1,X2))!=X4|~v1_funct_1(X3)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r2_fdiff_1(X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~r2_hidden(esk3_3(X3,X4,k2_fdiff_1(X1,X2)),X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_3(X3,X4,k2_fdiff_1(X1,X2)),k1_numbers))).#trainneg
cnf(i_0_8548, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k4_xcmplx_0(k1_seq_1(k16_sin_cos,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers))).#trainneg
cnf(i_0_7582, plain, (k3_xcmplx_0(X1,k1_real_1(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~m1_subset_1(k3_xcmplx_0(X1,k1_real_1(X2)),k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_8102, plain, (k3_xcmplx_0(X1,k1_real_1(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8747, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8716, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),X1)=k1_real_1(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_779, plain, (v5_relat_1(esk4_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3),X2)|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_8861, plain, (v1_xcmplx_0(k4_xcmplx_0(k5_numbers))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8733, plain, (v1_xreal_0(k1_real_1(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8725, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_775, plain, (v3_valued_0(esk4_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3))|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_membered(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_776, plain, (v1_xboole_0(esk4_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3))|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_8922, plain, (k3_xcmplx_0(X1,k1_real_1(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X2),k1_numbers))).#trainneg
cnf(i_0_8921, plain, (k3_xcmplx_0(X1,k1_real_1(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9308, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_9178, plain, (k4_xcmplx_0(k3_xcmplx_0(np__1,X1))=k1_real_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4990, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),np__1)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_8471, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k16_sin_cos,esk2_0),k1_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~v1_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)))|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,esk2_0))|~v1_xreal_0(k1_seq_1(k16_sin_cos,esk2_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_780, plain, (v1_relat_1(esk4_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3))|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_881, plain, (k2_zfmisc_1(X1,X2)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_8563, plain, (v7_ordinal1(k4_xcmplx_0(k5_numbers))|~v6_membered(k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8564, plain, (v1_int_1(k4_xcmplx_0(k5_numbers))|~v5_membered(k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_777, plain, (v1_valued_0(esk4_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3))|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_membered(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_778, plain, (v2_valued_0(esk4_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3))|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v2_membered(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_9249, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_9401, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_782, plain, (v4_relat_1(esk4_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3),X1)|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_8565, plain, (v1_xxreal_0(k4_xcmplx_0(k5_numbers))|~v2_membered(k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_9569, plain, (v1_int_1(k4_xcmplx_0(k5_numbers))|~v6_membered(k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8717, plain, (v1_int_1(k1_real_1(X1))|~v1_int_1(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_781, plain, (v1_xboole_0(esk4_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3))|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X1)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_783, plain, (v3_membered(esk4_2(k1_zfmisc_1(X1),X2))|r2_fdiff_1(X2,k1_zfmisc_1(X1))|~v3_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_9277, negated_conjecture, (k4_xcmplx_0(k3_xcmplx_0(k1_seq_1(k16_sin_cos,esk2_0),k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers))).#trainneg
cnf(i_0_9278, negated_conjecture, (k4_xcmplx_0(k3_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers))).#trainneg
cnf(i_0_8214, plain, (k1_real_1(k3_xcmplx_0(X2,X1))=k3_xcmplx_0(X1,k1_real_1(X2))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_9386, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,np__1))=k1_real_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5487, plain, (k3_xcmplx_0(np__1,k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_784, plain, (v1_membered(esk4_2(k1_zfmisc_1(X1),X2))|r2_fdiff_1(X2,k1_zfmisc_1(X1))|~v1_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_9564, plain, (k2_zfmisc_1(np__0,X1)=np__0)).#trainneg
cnf(i_0_5082, negated_conjecture, (m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(esk2_0)),k1_numbers)|~v1_xcmplx_0(k4_xcmplx_0(esk2_0)))).#trainneg
cnf(i_0_9095, plain, (v1_xcmplx_0(k1_real_1(X1))|~v1_membered(k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9096, plain, (v7_ordinal1(k1_real_1(X1))|~v6_membered(k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_785, plain, (v2_membered(esk4_2(k1_zfmisc_1(X1),X2))|r2_fdiff_1(X2,k1_zfmisc_1(X1))|~v2_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_786, plain, (v3_valued_0(esk4_2(k1_zfmisc_1(X1),X2))|r2_fdiff_1(X2,k1_zfmisc_1(X1))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_9757, plain, (k3_xcmplx_0(X2,k1_real_1(X1))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_9788, plain, (k3_xcmplx_0(X2,k1_real_1(X1))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(k3_xcmplx_0(X1,X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_10068, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(np__0))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X9)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10021, plain, (m1_subset_1(X1,np__0)|~r2_hidden(X1,k9_xtuple_0(np__0)))).#trainneg
cnf(i_0_10009, plain, (v3_valued_0(np__0))).#trainneg
cnf(i_0_10047, plain, (v1_relat_1(np__0))).#trainneg
cnf(i_0_787, plain, (v1_valued_0(esk4_2(k1_zfmisc_1(X1),X2))|r2_fdiff_1(X2,k1_zfmisc_1(X1))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_10049, plain, (v4_relat_1(np__0,np__0))).#trainneg
cnf(i_0_10008, plain, (v3_valued_0(k9_xtuple_0(np__0)))).#trainneg
cnf(i_0_10054, plain, (v1_relat_1(k9_xtuple_0(np__0)))).#trainneg
cnf(i_0_9097, plain, (v1_int_1(k1_real_1(X1))|~v5_membered(k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9098, plain, (v1_xxreal_0(k1_real_1(X1))|~v2_membered(k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_788, plain, (m1_subset_1(X1,X2)|r2_fdiff_1(X3,k1_zfmisc_1(X2))|~r2_hidden(X1,esk4_2(k1_zfmisc_1(X2),X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_792, plain, (v7_ordinal1(esk4_2(X1,X2))|r2_fdiff_1(X2,X1)|~v6_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_8748, plain, (k3_xcmplx_0(X1,k1_real_1(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_9126, plain, (k3_xcmplx_0(X1,k1_real_1(X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_7964, plain, (k1_real_1(k1_real_1(X1))=k1_real_1(k4_xcmplx_0(X1))|~m1_subset_1(k1_real_1(k4_xcmplx_0(X1)),k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_10291, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(np__0))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(np__0))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(np__0)))).#trainneg
cnf(i_0_10020, plain, (m1_subset_1(X1,np__0)|~r2_hidden(X1,k1_relset_1(np__0,np__0)))).#trainneg
cnf(i_0_10046, plain, (v5_relat_1(np__0,X1))).#trainneg
cnf(i_0_794, plain, (v1_xxreal_0(esk4_2(X1,X2))|r2_fdiff_1(X2,X1)|~v2_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_803, plain, (v1_xreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X3),X2))|~v1_xreal_0(k1_fdiff_1(X3,k1_seq_1(X1,X2)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~r1_fdiff_1(X3,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_819, plain, (m1_subset_1(k1_funct_1(X1,X2),k1_numbers)|~m1_subset_1(k4_xcmplx_0(k1_funct_1(X1,X2)),k1_numbers)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_10502, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8683, negated_conjecture, (k3_xcmplx_0(k4_xcmplx_0(k1_seq_1(k16_sin_cos,esk2_0)),k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers))).#trainneg
cnf(i_0_8090, plain, (k1_real_1(k4_xcmplx_0(X1))=k4_xcmplx_0(k1_real_1(X1))|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9412, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))=k3_xcmplx_0(X1,X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10017, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(np__0)))).#trainneg
cnf(i_0_10018, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(np__0)))).#trainneg
cnf(i_0_10034, plain, (k9_xtuple_0(np__0)=k1_relset_1(np__0,np__0))).#trainneg
cnf(i_0_820, plain, (m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10056, plain, (v3_membered(k9_xtuple_0(np__0))|~v3_membered(np__0))).#trainneg
cnf(i_0_821, plain, (m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k1_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_10728, plain, (k4_xcmplx_0(k1_real_1(X1))=k4_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10700, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_10657, plain, (k3_xcmplx_0(X2,k4_xcmplx_0(X1))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_10010, plain, (v3_valued_0(esk7_1(k1_zfmisc_1(np__0))))).#trainneg
cnf(i_0_10054, plain, (v1_relat_1(k1_relset_1(np__0,np__0)))).#trainneg
cnf(i_0_830, plain, (m1_subset_1(esk7_1(X1),k1_numbers)|~m1_subset_1(k4_xcmplx_0(esk7_1(X1)),k1_numbers)|~v1_membered(X1))).#trainneg
cnf(i_0_10024, plain, (v1_valued_0(np__0)|~v1_membered(X1))).#trainneg
cnf(i_0_10932, negated_conjecture, (~m1_subset_1(esk1_0,k1_zfmisc_1(np__0)))).#trainneg
cnf(i_0_10933, plain, (~m1_subset_1(np__2,k1_zfmisc_1(np__0)))).#trainneg
cnf(i_0_10099, plain, (v1_int_1(k1_real_1(X1))|~v6_membered(k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_843, plain, (k3_xcmplx_0(X3,k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(k3_xcmplx_0(X1,X2))|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_10722, negated_conjecture, (k3_xcmplx_0(k1_fdiff_1(k19_sin_cos,esk2_0),k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers))).#trainneg
cnf(i_0_10515, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_9119, plain, (k3_xcmplx_0(np__1,k4_xcmplx_0(X1))=k1_real_1(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8553, negated_conjecture, (k3_xcmplx_0(k4_xcmplx_0(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0))),k1_seq_1(k16_sin_cos,esk2_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers))).#trainneg
cnf(i_0_9016, negated_conjecture, (k4_xcmplx_0(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~v1_xcmplx_0(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0)))|~v1_xreal_0(k8_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_seq_1(k16_sin_cos,esk2_0))))).#trainneg
cnf(i_0_11254, negated_conjecture, (v1_valued_0(np__0))).#trainneg
cnf(i_0_10025, plain, (v1_relat_1(esk7_1(k1_zfmisc_1(np__0))))).#trainneg
cnf(i_0_10057, plain, (v1_membered(k9_xtuple_0(np__0))|~v1_membered(np__0))).#trainneg
cnf(i_0_10008, plain, (v3_valued_0(k1_relset_1(np__0,np__0)))).#trainneg
cnf(i_0_10045, plain, (v2_valued_0(np__0)|~v2_membered(X1))).#trainneg
cnf(i_0_10935, plain, (~m1_subset_1(k1_numbers,k1_zfmisc_1(np__0)))).#trainneg
cnf(i_0_10058, plain, (v2_membered(k9_xtuple_0(np__0))|~v2_membered(np__0))).#trainneg
cnf(i_0_9807, plain, (k3_xcmplx_0(X2,k1_real_1(X1))=k3_xcmplx_0(X1,k1_real_1(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11595, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0),esk2_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k24_sin_cos,k1_seq_1(k19_sin_cos,esk2_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)|~v1_xreal_0(k1_seq_1(k19_sin_cos,esk2_0))|~r1_fdiff_1(k19_sin_cos,esk2_0)|~m1_subset_1(k1_fdiff_1(k19_sin_cos,esk2_0),k1_numbers))).#trainneg
cnf(i_0_11857, plain, (k3_xcmplx_0(X2,k1_real_1(X1))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_11772, negated_conjecture, (v2_valued_0(np__0))).#trainneg
cnf(i_0_10030, plain, (v1_xboole_0(esk7_1(k1_zfmisc_1(np__0))))).#trainneg
cnf(i_0_10035, plain, (v4_relat_1(k1_relset_1(np__0,np__0),np__0))).#trainneg
cnf(i_0_11819, plain, (k3_xcmplx_0(k1_real_1(X2),X1)=k3_xcmplx_0(X2,k1_real_1(X1))|~v1_xcmplx_0(k1_real_1(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_10936, plain, (X1=np__0|~m1_subset_1(X1,k1_zfmisc_1(np__0)))).#trainneg
cnf(i_0_10053, plain, (v5_relat_1(k9_xtuple_0(np__0),X1))).#trainneg
cnf(i_0_10071, negated_conjecture, (v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(esk2_0))))).#trainneg
cnf(i_0_856, plain, (v1_int_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X3),X2))|~v1_int_1(k1_fdiff_1(X3,k1_seq_1(X1,X2)))|~v1_int_1(k1_fdiff_1(X1,X2))|~v1_xreal_0(k1_fdiff_1(X3,k1_seq_1(X1,X2)))|~r1_fdiff_1(X3,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_857, plain, (k3_xcmplx_0(k1_fdiff_1(X1,X2),k3_xcmplx_0(k1_fdiff_1(X3,k1_seq_1(X1,X2)),X4))=k3_xcmplx_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X3),X2),X4)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k1_fdiff_1(X3,k1_seq_1(X1,X2)))|~v1_xcmplx_0(k1_fdiff_1(X1,X2))|~v1_xreal_0(k1_fdiff_1(X3,k1_seq_1(X1,X2)))|~r1_fdiff_1(X3,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_11978, negated_conjecture, (~v1_xreal_0(k1_seq_1(k19_sin_cos,esk2_0))|~r1_fdiff_1(k19_sin_cos,esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)|~m1_subset_1(k1_fdiff_1(k19_sin_cos,esk2_0),k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_12179, plain, (k3_xcmplx_0(X1,k1_real_1(X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12072, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_12190, plain, (k1_partfun1(X1,X2,X3,X4,X5,X6)=np__0|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(np__0))|~m1_subset_1(X6,k1_zfmisc_1(np__0)))).#trainneg
cnf(i_0_10052, plain, (m1_subset_1(k9_xtuple_0(np__0),k1_zfmisc_1(np__0)))).#trainneg
cnf(i_0_12080, plain, (esk7_1(k1_zfmisc_1(np__0))=np__0)).#trainneg
cnf(i_0_12511, negated_conjecture, (k1_seq_1(k2_fdiff_1(np__0,esk1_0),esk2_0)!=k1_fdiff_1(np__0,esk2_0)|~v1_xreal_0(k1_seq_1(k19_sin_cos,esk2_0))|~r1_fdiff_1(k19_sin_cos,esk2_0)|~r2_fdiff_1(np__0,esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)|~m1_subset_1(k1_fdiff_1(k19_sin_cos,esk2_0),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(np__0))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(np__0)))).#trainneg
cnf(i_0_866, plain, (k6_xcmplx_0(k1_seq_1(k16_sin_cos,X1),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,k1_fdiff_1(k19_sin_cos,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_fdiff_1(k19_sin_cos,X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_12187, plain, (k1_relset_1(np__0,np__0)=np__0)).#trainneg
cnf(i_0_871, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k1_seq_1(k16_sin_cos,X2))=k6_xcmplx_0(k1_fdiff_1(k19_sin_cos,X2),X1)|~v1_xcmplx_0(k1_fdiff_1(k19_sin_cos,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(k1_seq_1(k16_sin_cos,X2),k1_numbers))).#trainneg
cnf(i_0_12629, plain, (v1_funct_1(np__0)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(np__0))|~m1_subset_1(X6,k1_zfmisc_1(np__0)))).#trainneg
cnf(i_0_12387, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10518, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k1_real_1(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k1_real_1(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_10034, plain, (k9_xtuple_0(np__0)=np__0)).#trainneg
cnf(i_0_882, plain, (k2_zfmisc_1(X1,X2)=np__0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_10072, negated_conjecture, (v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(esk2_0)))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_886, plain, (m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,X1),X2),k1_numbers)|~v1_xreal_0(k1_seq_1(k24_sin_cos,X2))|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k24_sin_cos,X2)),k1_numbers)|~v1_xreal_0(X2)|~r1_fdiff_1(X1,k1_seq_1(k24_sin_cos,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_887, plain, (v1_xreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,X1),X2))|~v1_xreal_0(k1_seq_1(k24_sin_cos,X2))|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k24_sin_cos,X2)),k1_numbers)|~v1_xreal_0(X2)|~r1_fdiff_1(X1,k1_seq_1(k24_sin_cos,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_12711, plain, (~r2_hidden(X1,np__0))).#trainneg
cnf(i_0_888, plain, (m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,X1),X2),k1_numbers)|~v1_xreal_0(k1_fdiff_1(X1,k1_seq_1(k24_sin_cos,X2)))|~m1_subset_1(k1_seq_1(k24_sin_cos,X2),k1_numbers)|~v1_xreal_0(X2)|~r1_fdiff_1(X1,k1_seq_1(k24_sin_cos,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_10019, plain, (v4_relat_1(X1,np__0)|~m1_subset_1(X1,k1_zfmisc_1(np__0)))).#trainneg
cnf(i_0_10073, negated_conjecture, (v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(esk2_0)))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_10074, negated_conjecture, (v1_int_1(k4_xcmplx_0(k4_xcmplx_0(esk2_0)))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_7903, plain, (v7_ordinal1(k1_relset_1(k1_numbers,k19_sin_cos))|~v6_membered(k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7904, plain, (v1_int_1(k1_relset_1(k1_numbers,k19_sin_cos))|~v5_membered(k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_9020, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1))=k1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~v1_xreal_0(k1_seq_1(k16_sin_cos,X1)))).#trainneg
cnf(i_0_9592, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))=k3_xcmplx_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_10699, plain, (k4_xcmplx_0(k3_xcmplx_0(np__1,X1))=k1_real_1(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_13090, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1))=k1_fdiff_1(k19_sin_cos,X1)|~v1_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7905, plain, (v1_xxreal_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v2_membered(k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_13084, plain, (v1_int_1(k1_relset_1(k1_numbers,k19_sin_cos))|~v6_membered(k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_13365, plain, (k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))=k1_real_1(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_12050, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7229, plain, (k4_xcmplx_0(k3_xcmplx_0(k1_real_1(X1),X2))=k3_xcmplx_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(k3_xcmplx_0(k1_real_1(X1),X2),k1_numbers))).#trainneg
cnf(i_0_10011, plain, (v3_valued_0(X1)|~v3_membered(X2)|~m1_subset_1(X1,k1_zfmisc_1(np__0)))).#trainneg
cnf(i_0_905, plain, (k1_partfun1(X1,X2,X3,X4,X5,X6)=np__0|~v1_xboole_0(X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_13017, plain, (k2_zfmisc_1(X1,np__0)=np__0)).#trainneg
cnf(i_0_906, plain, (v1_xboole_0(k3_relat_1(X5,X6))|~v1_xboole_0(X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_907, plain, (v1_xboole_0(k1_partfun1(X7,X8,X9,X10,X5,X6))|~v1_xboole_0(X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_14027, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(np__0))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(np__0)))).#trainneg
cnf(i_0_14007, plain, (m1_subset_1(np__0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_13370, plain, (k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1))=k1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_13689, plain, (k3_xcmplx_0(k1_real_1(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 101
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 145
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 142
# Processed clauses                    : 2217
# ...of these trivial                  : 47
# ...subsumed                          : 886
# ...remaining for further processing  : 1284
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 89
# Backward-rewritten                   : 71
# Generated clauses                    : 14192
# ...of the previous two non-trivial   : 13287
# Contextual simplify-reflections      : 251
# Paramodulations                      : 14185
# Factorizations                       : 0
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1124
#    Positive orientable unit clauses  : 163
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 8
#    Non-unit-clauses                  : 952
# Current number of unprocessed clauses: 11038
# ...number of literals in the above   : 77191
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1444
# Proof object given clauses           : 145
# Proof search given clauses           : 1284
# Clause-clause subsumption calls (NU) : 535969
# Rec. Clause-clause subsumption calls : 53791
# Non-unit clause-clause subsumptions  : 1203
# Unit Clause-clause subsumption calls : 16290
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 61
# BW rewrite match successes           : 27
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 427917

# -------------------------------------------------
# User time                : 28.404 s
# System time              : 0.693 s
# Total time               : 29.097 s
# Maximum resident set size: 381764 pages
28.40user 0.76system 0:29.47elapsed 98%CPU (0avgtext+0avgdata 381764maxresident)k
0inputs+808outputs (0major+258818minor)pagefaults 0swaps
