# ENIGMATIC: Connected to the TF server '127.0.0.1:8888'.
# ENIGMATIC: LightGBM model './big03/model.lgb' loaded with 17544 features 'C(l,x[c=3],r,h[b=4096],v,c[b=256],d,a):G:P'
# Preprocessing time       : 0.750 s
# ENIGMATIC: LightGBM parental guidance model #1 './froc_gnn_model2/model.lgb' loaded with 13963 features 'C(l,x[c=3],r,h[b=2048],v,c[b=256],d,a):G:F'
#TF#SERVER: Lgb model filter skipped 7 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=109, query=114, total=227).
#TF#SERVER: Sent 50980 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (896 bytes expected).
#TF#SERVER: Received 896 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 37 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=194, query=159, total=357).
#TF#SERVER: Sent 92726 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1416 bytes expected).
#TF#SERVER: Received 1416 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 300 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=301, query=365, total=670).
#TF#SERVER: Sent 213861 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2668 bytes expected).
#TF#SERVER: Received 2668 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 418 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=390, query=614, total=1008).
#TF#SERVER: Sent 374680 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4020 bytes expected).
#TF#SERVER: Received 4020 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 439 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=500, query=588, total=1092).
#TF#SERVER: Sent 385436 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4356 bytes expected).
#TF#SERVER: Received 4356 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 409 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=593, query=615, total=1212).
#TF#SERVER: Sent 362155 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4836 bytes expected).
#TF#SERVER: Received 4836 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 523 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=658, query=503, total=1165).
#TF#SERVER: Sent 450867 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4648 bytes expected).
#TF#SERVER: Received 4648 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 672 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=702, query=354, total=1060).
#TF#SERVER: Sent 384336 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4228 bytes expected).
#TF#SERVER: Received 4228 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 618 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=757, query=476, total=1237).
#TF#SERVER: Sent 396295 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4936 bytes expected).
#TF#SERVER: Received 4936 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 733 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=788, query=325, total=1117).
#TF#SERVER: Sent 403391 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4456 bytes expected).
#TF#SERVER: Received 4456 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 650 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=833, query=397, total=1234).
#TF#SERVER: Sent 433816 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4924 bytes expected).
#TF#SERVER: Received 4924 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 628 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=865, query=421, total=1290).
#TF#SERVER: Sent 439145 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5148 bytes expected).
#TF#SERVER: Received 5148 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 526 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=891, query=508, total=1403).
#TF#SERVER: Sent 527708 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5600 bytes expected).
#TF#SERVER: Received 5600 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 556 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=924, query=485, total=1413).
#TF#SERVER: Sent 494813 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5640 bytes expected).
#TF#SERVER: Received 5640 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 595 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=948, query=431, total=1383).
#TF#SERVER: Sent 480816 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5520 bytes expected).
#TF#SERVER: Received 5520 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 777 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=994, query=278, total=1276).
#TF#SERVER: Sent 464031 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5092 bytes expected).
#TF#SERVER: Received 5092 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 546 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1054, query=508, total=1566).
#TF#SERVER: Sent 561341 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6252 bytes expected).
#TF#SERVER: Received 6252 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 641 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1084, query=390, total=1478).
#TF#SERVER: Sent 538698 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5900 bytes expected).
#TF#SERVER: Received 5900 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 653 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1111, query=371, total=1486).
#TF#SERVER: Sent 546278 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5932 bytes expected).
#TF#SERVER: Received 5932 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 541 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1143, query=500, total=1647).
#TF#SERVER: Sent 562258 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6576 bytes expected).
#TF#SERVER: Received 6576 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 662 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1177, query=374, total=1555).
#TF#SERVER: Sent 514778 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6208 bytes expected).
#TF#SERVER: Received 6208 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 758 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1204, query=287, total=1495).
#TF#SERVER: Sent 550007 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5968 bytes expected).
#TF#SERVER: Received 5968 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 503 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1248, query=541, total=1793).
#TF#SERVER: Sent 585240 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7160 bytes expected).
#TF#SERVER: Received 7160 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 609 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1272, query=422, total=1698).
#TF#SERVER: Sent 579124 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6780 bytes expected).
#TF#SERVER: Received 6780 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 545 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1312, query=489, total=1805).
#TF#SERVER: Sent 625005 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7208 bytes expected).
#TF#SERVER: Received 7208 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 668 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1354, query=439, total=1797).
#TF#SERVER: Sent 665272 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7176 bytes expected).
#TF#SERVER: Received 7176 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 717 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1392, query=376, total=1772).
#TF#SERVER: Sent 573182 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7076 bytes expected).
#TF#SERVER: Received 7076 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 686 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1428, query=445, total=1877).
#TF#SERVER: Sent 659189 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7496 bytes expected).
#TF#SERVER: Received 7496 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 638 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1476, query=528, total=2008).
#TF#SERVER: Sent 731455 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8020 bytes expected).
#TF#SERVER: Received 8020 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 495 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1526, query=558, total=2088).
#TF#SERVER: Sent 802872 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8340 bytes expected).
#TF#SERVER: Received 8340 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 676 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=410, total=1950).
#TF#SERVER: Sent 734414 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7788 bytes expected).
#TF#SERVER: Received 7788 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 777 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=324, total=1864).
#TF#SERVER: Sent 703431 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7444 bytes expected).
#TF#SERVER: Received 7444 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 582 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=476, total=2016).
#TF#SERVER: Sent 773549 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8052 bytes expected).
#TF#SERVER: Received 8052 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 654 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=534, total=2074).
#TF#SERVER: Sent 861652 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8284 bytes expected).
#TF#SERVER: Received 8284 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 621 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=532, total=2072).
#TF#SERVER: Sent 824853 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8276 bytes expected).
#TF#SERVER: Received 8276 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 719 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=359, total=1899).
#TF#SERVER: Sent 757351 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7584 bytes expected).
#TF#SERVER: Received 7584 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 618 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=425, total=1965).
#TF#SERVER: Sent 739604 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7848 bytes expected).
#TF#SERVER: Received 7848 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 637 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=444, total=1984).
#TF#SERVER: Sent 770088 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7924 bytes expected).
#TF#SERVER: Received 7924 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 694 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=357, total=1897).
#TF#SERVER: Sent 725225 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7576 bytes expected).
#TF#SERVER: Received 7576 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 708 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=338, total=1878).
#TF#SERVER: Sent 706436 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7500 bytes expected).
#TF#SERVER: Received 7500 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 637 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=387, total=1927).
#TF#SERVER: Sent 747189 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7696 bytes expected).
#TF#SERVER: Received 7696 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 787 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=296, total=1836).
#TF#SERVER: Sent 720504 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7332 bytes expected).
#TF#SERVER: Received 7332 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 643 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=422, total=1962).
#TF#SERVER: Sent 802940 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7836 bytes expected).
#TF#SERVER: Received 7836 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 628 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=443, total=1983).
#TF#SERVER: Sent 751228 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7920 bytes expected).
#TF#SERVER: Received 7920 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 565 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=496, total=2036).
#TF#SERVER: Sent 801101 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8132 bytes expected).
#TF#SERVER: Received 8132 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 654 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=471, total=2011).
#TF#SERVER: Sent 775874 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8032 bytes expected).
#TF#SERVER: Received 8032 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 561 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=470, total=2010).
#TF#SERVER: Sent 731071 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8028 bytes expected).
#TF#SERVER: Received 8028 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 773 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=319, total=1859).
#TF#SERVER: Sent 723919 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7424 bytes expected).
#TF#SERVER: Received 7424 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 694 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=385, total=1925).
#TF#SERVER: Sent 729662 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7688 bytes expected).
#TF#SERVER: Received 7688 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 718 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=419, total=1959).
#TF#SERVER: Sent 798387 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7824 bytes expected).
#TF#SERVER: Received 7824 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 550 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=477, total=2017).
#TF#SERVER: Sent 825986 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8056 bytes expected).
#TF#SERVER: Received 8056 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 1007 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=96, total=1636).
#TF#SERVER: Sent 610782 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6532 bytes expected).
#TF#SERVER: Received 6532 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 650 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=488, total=2028).
#TF#SERVER: Sent 824121 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8100 bytes expected).
#TF#SERVER: Received 8100 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 680 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=395, total=1935).
#TF#SERVER: Sent 747998 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7728 bytes expected).
#TF#SERVER: Received 7728 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 612 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=474, total=2014).
#TF#SERVER: Sent 817428 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8044 bytes expected).
#TF#SERVER: Received 8044 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 625 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=399, total=1939).
#TF#SERVER: Sent 757429 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7744 bytes expected).
#TF#SERVER: Received 7744 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 677 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=423, total=1963).
#TF#SERVER: Sent 749843 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7840 bytes expected).
#TF#SERVER: Received 7840 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 611 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=425, total=1965).
#TF#SERVER: Sent 783128 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7848 bytes expected).
#TF#SERVER: Received 7848 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 691 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=354, total=1894).
#TF#SERVER: Sent 741682 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7564 bytes expected).
#TF#SERVER: Received 7564 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 679 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=409, total=1949).
#TF#SERVER: Sent 723048 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7784 bytes expected).
#TF#SERVER: Received 7784 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 576 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=654, total=2194).
#TF#SERVER: Sent 931196 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8764 bytes expected).
#TF#SERVER: Received 8764 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 628 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=414, total=1954).
#TF#SERVER: Sent 812046 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7804 bytes expected).
#TF#SERVER: Received 7804 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 751 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=273, total=1813).
#TF#SERVER: Sent 693007 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7240 bytes expected).
#TF#SERVER: Received 7240 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 654 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=377, total=1917).
#TF#SERVER: Sent 673790 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7656 bytes expected).
#TF#SERVER: Received 7656 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 680 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=407, total=1947).
#TF#SERVER: Sent 787433 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7776 bytes expected).
#TF#SERVER: Received 7776 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 562 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=467, total=2007).
#TF#SERVER: Sent 800706 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8016 bytes expected).
#TF#SERVER: Received 8016 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 568 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=479, total=2019).
#TF#SERVER: Sent 909029 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8064 bytes expected).
#TF#SERVER: Received 8064 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 692 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=399, total=1939).
#TF#SERVER: Sent 753492 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7744 bytes expected).
#TF#SERVER: Received 7744 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 609 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=422, total=1962).
#TF#SERVER: Sent 826224 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7836 bytes expected).
#TF#SERVER: Received 7836 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 617 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=450, total=1990).
#TF#SERVER: Sent 810847 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7948 bytes expected).
#TF#SERVER: Received 7948 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 497 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=531, total=2071).
#TF#SERVER: Sent 844150 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8272 bytes expected).
#TF#SERVER: Received 8272 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 756 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=426, total=1966).
#TF#SERVER: Sent 721162 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7852 bytes expected).
#TF#SERVER: Received 7852 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 504 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=537, total=2077).
#TF#SERVER: Sent 872744 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8296 bytes expected).
#TF#SERVER: Received 8296 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 626 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=468, total=2008).
#TF#SERVER: Sent 770298 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8020 bytes expected).
#TF#SERVER: Received 8020 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 509 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=544, total=2084).
#TF#SERVER: Sent 772554 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8324 bytes expected).
#TF#SERVER: Received 8324 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 517 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=514, total=2054).
#TF#SERVER: Sent 788305 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8204 bytes expected).
#TF#SERVER: Received 8204 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 702 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=392, total=1932).
#TF#SERVER: Sent 732273 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7716 bytes expected).
#TF#SERVER: Received 7716 bytes from the TF server.

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t1_series_2, axiom, ![X1]:(m1_subset_1(X1,k5_numbers)=>k18_complex1(k1_newton(k4_xcmplx_0(np__1),X1))=np__1), file('/scratch/mptp/minsubpreds5/t36_int_5', t1_series_2)).
fof(t1_absvalue, axiom, ![X1]:(v1_xreal_0(X1)=>(k18_complex1(X1)=X1|k18_complex1(X1)=k4_xcmplx_0(X1))), file('/scratch/mptp/minsubpreds5/t36_int_5', t1_absvalue)).
fof(involutiveness_k4_xcmplx_0, axiom, ![X1]:(v1_xcmplx_0(X1)=>k4_xcmplx_0(k4_xcmplx_0(X1))=X1), file('/scratch/mptp/minsubpreds5/t36_int_5', involutiveness_k4_xcmplx_0)).
fof(fc3_newton, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v7_ordinal1(X2))=>v1_xcmplx_0(k1_newton(X1,X2))), file('/scratch/mptp/minsubpreds5/t36_int_5', fc3_newton)).
fof(cc8_ordinal1, axiom, ![X1]:(m1_subset_1(X1,k4_ordinal1)=>v7_ordinal1(X1)), file('/scratch/mptp/minsubpreds5/t36_int_5', cc8_ordinal1)).
fof(fc3_int_1, axiom, ![X1]:(v1_int_1(X1)=>(v1_xcmplx_0(k4_xcmplx_0(X1))&v1_int_1(k4_xcmplx_0(X1)))), file('/scratch/mptp/minsubpreds5/t36_int_5', fc3_int_1)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/minsubpreds5/t36_int_5', redefinition_k5_numbers)).
fof(cc14_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xreal_0(X2))), file('/scratch/mptp/minsubpreds5/t36_int_5', cc14_membered)).
fof(cc16_membered, axiom, ![X1]:(v5_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_int_1(X2))), file('/scratch/mptp/minsubpreds5/t36_int_5', cc16_membered)).
fof(t10_int_2, axiom, ![X1]:(v1_int_1(X1)=>![X2]:(v1_int_1(X2)=>((((((((r1_int_1(X1,X2)=>r1_int_1(X1,k4_xcmplx_0(X2)))&(r1_int_1(X1,k4_xcmplx_0(X2))=>r1_int_1(X1,X2)))&(r1_int_1(X1,X2)=>r1_int_1(k4_xcmplx_0(X1),X2)))&(r1_int_1(k4_xcmplx_0(X1),X2)=>r1_int_1(X1,X2)))&(r1_int_1(X1,X2)=>r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))))&(r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=>r1_int_1(X1,X2)))&(r1_int_1(X1,k4_xcmplx_0(X2))=>r1_int_1(k4_xcmplx_0(X1),X2)))&(r1_int_1(k4_xcmplx_0(X1),X2)=>r1_int_1(X1,k4_xcmplx_0(X2)))))), file('/scratch/mptp/minsubpreds5/t36_int_5', t10_int_2)).
fof(cc3_int_1, axiom, ![X1]:(v1_int_1(X1)=>v1_xreal_0(X1)), file('/scratch/mptp/minsubpreds5/t36_int_5', cc3_int_1)).
fof(t5_newton, axiom, ![X1]:(v1_xcmplx_0(X1)=>k1_newton(X1,np__1)=X1), file('/scratch/mptp/minsubpreds5/t36_int_5', t5_newton)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/scratch/mptp/minsubpreds5/t36_int_5', spc1_numerals)).
fof(fc3_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>(v1_xcmplx_0(k4_xcmplx_0(X1))&v1_xreal_0(k4_xcmplx_0(X1)))), file('/scratch/mptp/minsubpreds5/t36_int_5', fc3_xreal_0)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/scratch/mptp/minsubpreds5/t36_int_5', fc3_membered)).
fof(cc1_membered, axiom, ![X1]:(v6_membered(X1)=>v5_membered(X1)), file('/scratch/mptp/minsubpreds5/t36_int_5', cc1_membered)).
fof(fc6_membered, axiom, v6_membered(k4_ordinal1), file('/scratch/mptp/minsubpreds5/t36_int_5', fc6_membered)).
fof(cc3_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xcmplx_0(X1)), file('/scratch/mptp/minsubpreds5/t36_int_5', cc3_xreal_0)).
fof(fc1_wsierp_1, axiom, ![X1, X2]:((v1_int_1(X1)&v7_ordinal1(X2))=>v1_int_1(k1_newton(X1,X2))), file('/scratch/mptp/minsubpreds5/t36_int_5', fc1_wsierp_1)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/scratch/mptp/minsubpreds5/t36_int_5', t6_boole)).
fof(spc2_numerals, axiom, (((v2_xxreal_0(np__2)&m2_subset_1(np__2,k1_numbers,k5_numbers))&m1_subset_1(np__2,k5_numbers))&m1_subset_1(np__2,k1_numbers)), file('/scratch/mptp/minsubpreds5/t36_int_5', spc2_numerals)).
fof(rqRealNeg__k4_xcmplx_0__rm1_r1, axiom, k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1, file('/scratch/mptp/minsubpreds5/t36_int_5', rqRealNeg__k4_xcmplx_0__rm1_r1)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/scratch/mptp/minsubpreds5/t36_int_5', redefinition_k6_numbers)).
fof(t7_nat_d, axiom, ![X1]:(v7_ordinal1(X1)=>![X2]:(v7_ordinal1(X2)=>(r1_nat_d(X2,X1)=>(r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,X1))))), file('/scratch/mptp/minsubpreds5/t36_int_5', t7_nat_d)).
fof(spc0_boole, axiom, v1_xboole_0(np__0), file('/scratch/mptp/minsubpreds5/t36_int_5', spc0_boole)).
fof(t15_int_2, axiom, ![X1]:(v1_int_1(X1)=>![X2]:(v1_int_1(X2)=>![X3]:(v1_int_1(X3)=>(r2_int_1(X1,X2,X3)<=>r1_int_1(X3,k6_xcmplx_0(X1,X2)))))), file('/scratch/mptp/minsubpreds5/t36_int_5', t15_int_2)).
fof(t28_int_5, axiom, ![X1]:(v1_int_1(X1)=>![X2]:((v7_ordinal1(X2)&v1_int_2(X2))=>(k3_int_2(X1,X2)=np__1=>(r1_xxreal_0(X2,np__2)|r2_int_1(k2_int_5(X1,X2),k1_newton(X1,k3_nat_d(k7_nat_d(X2,np__1),np__2)),X2))))), file('/scratch/mptp/minsubpreds5/t36_int_5', t28_int_5)).
fof(t36_int_5, conjecture, ![X1]:((v7_ordinal1(X1)&v1_int_2(X1))=>(~(r1_xxreal_0(X1,np__2))=>k2_int_5(k4_xcmplx_0(np__1),X1)=k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(X1,np__1),np__2)))), file('/scratch/mptp/minsubpreds5/t36_int_5', t36_int_5)).
fof(redefinition_r1_nat_d, axiom, ![X1, X2]:((v7_ordinal1(X1)&v7_ordinal1(X2))=>(r1_nat_d(X1,X2)<=>r1_int_1(X1,X2))), file('/scratch/mptp/minsubpreds5/t36_int_5', redefinition_r1_nat_d)).
fof(cc2_int_1, axiom, ![X1]:(v7_ordinal1(X1)=>v1_int_1(X1)), file('/scratch/mptp/minsubpreds5/t36_int_5', cc2_int_1)).
fof(dt_k2_int_5, axiom, ![X1, X2]:(((v1_int_1(X1)&v7_ordinal1(X2))&v1_int_2(X2))=>v1_int_1(k2_int_5(X1,X2))), file('/scratch/mptp/minsubpreds5/t36_int_5', dt_k2_int_5)).
fof(rqLessOrEqual__r1_xxreal_0__r2_r0, axiom, ~(r1_xxreal_0(np__2,np__0)), file('/scratch/mptp/minsubpreds5/t36_int_5', rqLessOrEqual__r1_xxreal_0__r2_r0)).
fof(cc2_xreal_0, axiom, ![X1]:(v7_ordinal1(X1)=>v1_xreal_0(X1)), file('/scratch/mptp/minsubpreds5/t36_int_5', cc2_xreal_0)).
fof(t25_int_5, axiom, ![X1]:(v1_int_1(X1)=>![X2]:((v7_ordinal1(X2)&v1_int_2(X2))=>~(((k2_int_5(X1,X2)!=np__1&k2_int_5(X1,X2)!=k6_numbers)&k2_int_5(X1,X2)!=k4_xcmplx_0(np__1))))), file('/scratch/mptp/minsubpreds5/t36_int_5', t25_int_5)).
fof(redefinition_k1_int_2, axiom, ![X1]:(v1_int_1(X1)=>k1_int_2(X1)=k16_complex1(X1)), file('/scratch/mptp/minsubpreds5/t36_int_5', redefinition_k1_int_2)).
fof(redefinition_k18_complex1, axiom, ![X1]:(v1_xcmplx_0(X1)=>k18_complex1(X1)=k16_complex1(X1)), file('/scratch/mptp/minsubpreds5/t36_int_5', redefinition_k18_complex1)).
fof(dt_k3_nat_d, axiom, ![X1, X2]:((v7_ordinal1(X1)&v7_ordinal1(X2))=>m1_subset_1(k3_nat_d(X1,X2),k5_numbers)), file('/scratch/mptp/minsubpreds5/t36_int_5', dt_k3_nat_d)).
fof(connectedness_r1_xxreal_0, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>(r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1))), file('/scratch/mptp/minsubpreds5/t36_int_5', connectedness_r1_xxreal_0)).
fof(cc2_xxreal_0, axiom, ![X1]:(v7_ordinal1(X1)=>v1_xxreal_0(X1)), file('/scratch/mptp/minsubpreds5/t36_int_5', cc2_xxreal_0)).
fof(cc17_membered, axiom, ![X1]:(v6_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v7_ordinal1(X2))), file('/scratch/mptp/minsubpreds5/t36_int_5', cc17_membered)).
fof(dt_k7_nat_d, axiom, ![X1, X2]:((v7_ordinal1(X1)&v7_ordinal1(X2))=>m1_subset_1(k7_nat_d(X1,X2),k5_numbers)), file('/scratch/mptp/minsubpreds5/t36_int_5', dt_k7_nat_d)).
fof(t51_newton, axiom, ![X1]:(v7_ordinal1(X1)=>k6_nat_d(X1,np__1)=np__1), file('/scratch/mptp/minsubpreds5/t36_int_5', t51_newton)).
fof(commutativity_k6_nat_d, axiom, ![X1, X2]:((v7_ordinal1(X1)&v7_ordinal1(X2))=>k6_nat_d(X1,X2)=k6_nat_d(X2,X1)), file('/scratch/mptp/minsubpreds5/t36_int_5', commutativity_k6_nat_d)).
fof(rqLessOrEqual__r1_xxreal_0__r1_r0, axiom, ~(r1_xxreal_0(np__1,np__0)), file('/scratch/mptp/minsubpreds5/t36_int_5', rqLessOrEqual__r1_xxreal_0__r1_r0)).
fof(rqRealDiff__k6_xcmplx_0__rm1_r1_rm2, axiom, k6_xcmplx_0(k4_xcmplx_0(np__1),np__1)=k4_xcmplx_0(np__2), file('/scratch/mptp/minsubpreds5/t36_int_5', rqRealDiff__k6_xcmplx_0__rm1_r1_rm2)).
fof(t13_nat_1, axiom, ![X1]:(v7_ordinal1(X1)=>![X2]:(v7_ordinal1(X2)=>(~(r1_xxreal_0(k1_nat_1(X2,np__1),X1))<=>r1_xxreal_0(X1,X2)))), file('/scratch/mptp/minsubpreds5/t36_int_5', t13_nat_1)).
fof(redefinition_k1_nat_1, axiom, ![X1, X2]:((v7_ordinal1(X1)&m1_subset_1(X2,k5_numbers))=>k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)), file('/scratch/mptp/minsubpreds5/t36_int_5', redefinition_k1_nat_1)).
fof(redefinition_k6_nat_d, axiom, ![X1, X2]:((v7_ordinal1(X1)&v7_ordinal1(X2))=>k6_nat_d(X1,X2)=k3_int_2(X1,X2)), file('/scratch/mptp/minsubpreds5/t36_int_5', redefinition_k6_nat_d)).
fof(t34_int_2, axiom, ![X1]:(v1_int_1(X1)=>![X2]:(v1_int_1(X2)=>k3_int_2(X1,X2)=k3_int_2(k1_int_2(X1),k1_int_2(X2)))), file('/scratch/mptp/minsubpreds5/t36_int_5', t34_int_2)).
fof(projectivity_k1_int_2, axiom, ![X1]:(v1_int_1(X1)=>k1_int_2(k1_int_2(X1))=k1_int_2(X1)), file('/scratch/mptp/minsubpreds5/t36_int_5', projectivity_k1_int_2)).
fof(dt_k1_int_2, axiom, ![X1]:(v1_int_1(X1)=>m1_subset_1(k1_int_2(X1),k5_numbers)), file('/scratch/mptp/minsubpreds5/t36_int_5', dt_k1_int_2)).
fof(spc0_numerals, axiom, ((m2_subset_1(np__0,k1_numbers,k5_numbers)&m1_subset_1(np__0,k5_numbers))&m1_subset_1(np__0,k1_numbers)), file('/scratch/mptp/minsubpreds5/t36_int_5', spc0_numerals)).
fof(rqRealAdd__k2_xcmplx_0__r1_r1_r2, axiom, k2_xcmplx_0(np__1,np__1)=np__2, file('/scratch/mptp/minsubpreds5/t36_int_5', rqRealAdd__k2_xcmplx_0__r1_r1_r2)).
fof(rqRealDiff__k6_xcmplx_0__r0_rm1_r1, axiom, k6_xcmplx_0(np__0,k4_xcmplx_0(np__1))=np__1, file('/scratch/mptp/minsubpreds5/t36_int_5', rqRealDiff__k6_xcmplx_0__r0_rm1_r1)).
fof(rqRealDiff__k6_xcmplx_0__r0_r1_rm1, axiom, k6_xcmplx_0(np__0,np__1)=k4_xcmplx_0(np__1), file('/scratch/mptp/minsubpreds5/t36_int_5', rqRealDiff__k6_xcmplx_0__r0_r1_rm1)).
fof(rqRealDiff__k6_xcmplx_0__r1_rm1_r2, axiom, k6_xcmplx_0(np__1,k4_xcmplx_0(np__1))=np__2, file('/scratch/mptp/minsubpreds5/t36_int_5', rqRealDiff__k6_xcmplx_0__r1_rm1_r2)).
fof(c_0_56, plain, ![X80]:(~m1_subset_1(X80,k5_numbers)|k18_complex1(k1_newton(k4_xcmplx_0(np__1),X80))=np__1), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_series_2])])).
fof(c_0_57, plain, ![X78]:(~v1_xreal_0(X78)|(k18_complex1(X78)=X78|k18_complex1(X78)=k4_xcmplx_0(X78))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_absvalue])])).
fof(c_0_58, plain, ![X54]:(~v1_xcmplx_0(X54)|k4_xcmplx_0(k4_xcmplx_0(X54))=X54), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k4_xcmplx_0])])).
fof(c_0_59, plain, ![X46, X47]:(~v1_xcmplx_0(X46)|~v7_ordinal1(X47)|v1_xcmplx_0(k1_newton(X46,X47))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc3_newton])])).
fof(c_0_60, plain, ![X24]:(~m1_subset_1(X24,k4_ordinal1)|v7_ordinal1(X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_ordinal1])])).
fof(c_0_61, plain, ![X1, X2]:(epred1_2(X2,X1)<=>((((((((r1_int_1(X1,X2)=>r1_int_1(X1,k4_xcmplx_0(X2)))&(r1_int_1(X1,k4_xcmplx_0(X2))=>r1_int_1(X1,X2)))&(r1_int_1(X1,X2)=>r1_int_1(k4_xcmplx_0(X1),X2)))&(r1_int_1(k4_xcmplx_0(X1),X2)=>r1_int_1(X1,X2)))&(r1_int_1(X1,X2)=>r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))))&(r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=>r1_int_1(X1,X2)))&(r1_int_1(X1,k4_xcmplx_0(X2))=>r1_int_1(k4_xcmplx_0(X1),X2)))&(r1_int_1(k4_xcmplx_0(X1),X2)=>r1_int_1(X1,k4_xcmplx_0(X2))))), introduced(definition)).
cnf(c_0_62, plain, (k18_complex1(k1_newton(k4_xcmplx_0(np__1),X1))=np__1|~m1_subset_1(X1,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_63, plain, (k18_complex1(X1)=X1|k18_complex1(X1)=k4_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
fof(c_0_64, plain, ![X45]:((v1_xcmplx_0(k4_xcmplx_0(X45))|~v1_int_1(X45))&(v1_int_1(k4_xcmplx_0(X45))|~v1_int_1(X45))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc3_int_1])])])).
cnf(c_0_65, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_66, plain, (v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_67, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_68, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
fof(c_0_69, plain, ![X6, X7]:(~v3_membered(X6)|(~m1_subset_1(X7,X6)|v1_xreal_0(X7))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc14_membered])])])).
fof(c_0_70, plain, ![X8, X9]:(~v5_membered(X8)|(~m1_subset_1(X9,X8)|v1_int_1(X9))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc16_membered])])])).
fof(c_0_71, plain, ![X1, X2]:(epred1_2(X2,X1)=>((((((((r1_int_1(X1,X2)=>r1_int_1(X1,k4_xcmplx_0(X2)))&(r1_int_1(X1,k4_xcmplx_0(X2))=>r1_int_1(X1,X2)))&(r1_int_1(X1,X2)=>r1_int_1(k4_xcmplx_0(X1),X2)))&(r1_int_1(k4_xcmplx_0(X1),X2)=>r1_int_1(X1,X2)))&(r1_int_1(X1,X2)=>r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))))&(r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=>r1_int_1(X1,X2)))&(r1_int_1(X1,k4_xcmplx_0(X2))=>r1_int_1(k4_xcmplx_0(X1),X2)))&(r1_int_1(k4_xcmplx_0(X1),X2)=>r1_int_1(X1,k4_xcmplx_0(X2))))), inference(split_equiv,[status(thm)],[c_0_61])).
fof(c_0_72, axiom, ![X1]:(v1_int_1(X1)=>![X2]:(v1_int_1(X2)=>epred1_2(X2,X1))), inference(apply_def,[status(thm)],[t10_int_2, c_0_61])).
cnf(c_0_73, plain, (k18_complex1(k1_newton(k4_xcmplx_0(np__1),X1))=np__1|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_62])).
cnf(c_0_74, plain, (k18_complex1(X1)=X1|k18_complex1(X1)=k4_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_63])).
fof(c_0_75, plain, ![X18]:(~v1_int_1(X18)|v1_xreal_0(X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_int_1])])).
cnf(c_0_76, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_77, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_65])).
cnf(c_0_78, plain, (v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[c_0_66])).
fof(c_0_79, plain, ![X90]:(~v1_xcmplx_0(X90)|k1_newton(X90,np__1)=X90), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_newton])])).
cnf(c_0_80, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_67])).
cnf(c_0_81, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_68])).
cnf(c_0_82, plain, (m1_subset_1(np__1,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
fof(c_0_83, plain, ![X48]:((v1_xcmplx_0(k4_xcmplx_0(X48))|~v1_xreal_0(X48))&(v1_xreal_0(k4_xcmplx_0(X48))|~v1_xreal_0(X48))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc3_xreal_0])])])).
cnf(c_0_84, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_85, plain, (m1_subset_1(np__1,k1_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_86, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
cnf(c_0_87, plain, (v1_int_1(X2)|~v5_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
fof(c_0_88, plain, ![X12]:(~v6_membered(X12)|v5_membered(X12)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_membered])])).
cnf(c_0_89, plain, (v6_membered(k4_ordinal1)), inference(split_conjunct,[status(thm)],[fc6_membered])).
fof(c_0_90, plain, ![X94, X95]:((((((((~r1_int_1(X94,X95)|r1_int_1(X94,k4_xcmplx_0(X95))|~epred1_2(X95,X94))&(~r1_int_1(X94,k4_xcmplx_0(X95))|r1_int_1(X94,X95)|~epred1_2(X95,X94)))&(~r1_int_1(X94,X95)|r1_int_1(k4_xcmplx_0(X94),X95)|~epred1_2(X95,X94)))&(~r1_int_1(k4_xcmplx_0(X94),X95)|r1_int_1(X94,X95)|~epred1_2(X95,X94)))&(~r1_int_1(X94,X95)|r1_int_1(k4_xcmplx_0(X94),k4_xcmplx_0(X95))|~epred1_2(X95,X94)))&(~r1_int_1(k4_xcmplx_0(X94),k4_xcmplx_0(X95))|r1_int_1(X94,X95)|~epred1_2(X95,X94)))&(~r1_int_1(X94,k4_xcmplx_0(X95))|r1_int_1(k4_xcmplx_0(X94),X95)|~epred1_2(X95,X94)))&(~r1_int_1(k4_xcmplx_0(X94),X95)|r1_int_1(X94,k4_xcmplx_0(X95))|~epred1_2(X95,X94))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_71])])])).
fof(c_0_91, plain, ![X71, X72]:(~v1_int_1(X71)|(~v1_int_1(X72)|epred1_2(X72,X71))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_72])])])).
fof(c_0_92, plain, ![X19]:(~v1_xreal_0(X19)|v1_xcmplx_0(X19)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_xreal_0])])).
cnf(c_0_93, plain, (k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))=np__1|k18_complex1(k1_newton(k4_xcmplx_0(np__1),X1))=k1_newton(k4_xcmplx_0(np__1),X1)|~m1_subset_1(X1,k5_numbers)|~v1_xreal_0(k1_newton(k4_xcmplx_0(np__1),X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74])])).
cnf(c_0_94, plain, (v1_xreal_0(X1)|~v1_int_1(X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
fof(c_0_95, plain, ![X43, X44]:(~v1_int_1(X43)|~v7_ordinal1(X44)|v1_int_1(k1_newton(X43,X44))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_wsierp_1])])).
cnf(c_0_96, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1)), inference(evalgc,[status(thm)],[c_0_76])).
cnf(c_0_97, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2)))=k1_newton(X1,X2)|~v7_ordinal1(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78])])).
cnf(c_0_98, plain, (k1_newton(X1,np__1)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_99, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)), inference(rw,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_100, plain, (m1_subset_1(np__1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_82])).
cnf(c_0_101, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_102, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_84])).
cnf(c_0_103, plain, (m1_subset_1(np__1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_85])).
cnf(c_0_104, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_86])).
cnf(c_0_105, plain, (v1_int_1(X2)|~v5_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_87])).
cnf(c_0_106, plain, (v5_membered(X1)|~v6_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_107, plain, (v6_membered(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_89])).
fof(c_0_108, plain, ![X91]:(~v1_xboole_0(X91)|X91=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
cnf(c_0_109, plain, (r1_int_1(X1,X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~epred1_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_110, plain, (epred1_2(X2,X1)|~v1_int_1(X1)|~v1_int_1(X2)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_111, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_112, plain, (m1_subset_1(np__2,k5_numbers)), inference(split_conjunct,[status(thm)],[spc2_numerals])).
cnf(c_0_113, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))=np__1|~m1_subset_1(X1,k5_numbers)|~v1_xreal_0(k1_newton(k4_xcmplx_0(np__1),X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_93])])).
cnf(c_0_114, plain, (v1_xreal_0(X1)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[c_0_94])).
cnf(c_0_115, plain, (v1_int_1(k1_newton(X1,X2))|~v1_int_1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_116, plain, (v1_int_1(k1_newton(X1,X2))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97])])).
cnf(c_0_117, plain, (k1_newton(X1,np__1)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_98])).
cnf(c_0_118, plain, (v7_ordinal1(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100])])).
cnf(c_0_119, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_101])).
cnf(c_0_120, plain, (v1_xreal_0(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_104])])])).
cnf(c_0_121, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1), inference(split_conjunct,[status(thm)],[rqRealNeg__k4_xcmplx_0__rm1_r1])).
cnf(c_0_122, plain, (v1_int_1(np__1)|~v5_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_100])])).
cnf(c_0_123, plain, (v5_membered(X1)|~v6_membered(X1)), inference(evalgc,[status(thm)],[c_0_106])).
cnf(c_0_124, plain, (v6_membered(k5_numbers)), inference(rw,[status(thm)],[c_0_107, c_0_81])).
cnf(c_0_125, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_108])).
cnf(c_0_126, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
cnf(c_0_127, plain, (r1_int_1(X1,X2)|~epred1_2(X2,X1)|~r1_int_1(X1,k4_xcmplx_0(X2))), inference(evalgc,[status(thm)],[c_0_109])).
cnf(c_0_128, plain, (epred1_2(X2,X1)|~v1_int_1(X2)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[c_0_110])).
cnf(c_0_129, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_111])).
cnf(c_0_130, plain, (m1_subset_1(np__2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_112])).
cnf(c_0_131, plain, (m1_subset_1(np__2,k1_numbers)), inference(split_conjunct,[status(thm)],[spc2_numerals])).
cnf(c_0_132, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~epred1_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_133, plain, (k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))=np__1|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~m1_subset_1(X1,k5_numbers)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114])])).
cnf(c_0_134, plain, (v1_int_1(k1_newton(X1,X2))|~v1_int_1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[c_0_115])).
cnf(c_0_135, plain, (v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_118])])])).
cnf(c_0_136, plain, (v1_xcmplx_0(k4_xcmplx_0(np__1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120])])).
cnf(c_0_137, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1), inference(evalgc,[status(thm)],[c_0_121])).
cnf(c_0_138, plain, (v1_int_1(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_124])])])).
fof(c_0_139, plain, ![X92, X93]:(~v7_ordinal1(X92)|(~v7_ordinal1(X93)|(~r1_nat_d(X93,X92)|(r1_xxreal_0(X92,k6_numbers)|r1_xxreal_0(X93,X92))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_nat_d])])])).
cnf(c_0_140, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_125])).
cnf(c_0_141, plain, (k1_xboole_0=k6_numbers), inference(evalgc,[status(thm)],[c_0_126])).
cnf(c_0_142, plain, (v1_xboole_0(np__0)), inference(split_conjunct,[status(thm)],[spc0_boole])).
cnf(c_0_143, plain, (r1_int_1(X1,X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128])])).
cnf(c_0_144, plain, (v1_xcmplx_0(X1)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_114])])).
cnf(c_0_145, plain, (v1_int_1(np__2)|~v5_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_130])])).
cnf(c_0_146, plain, (m1_subset_1(np__2,k1_numbers)), inference(evalgc,[status(thm)],[c_0_131])).
cnf(c_0_147, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~epred1_2(X2,X1)|~r1_int_1(k4_xcmplx_0(X1),X2)), inference(evalgc,[status(thm)],[c_0_132])).
fof(c_0_148, plain, ![X75, X76, X77]:((~r2_int_1(X75,X76,X77)|r1_int_1(X77,k6_xcmplx_0(X75,X76))|~v1_int_1(X77)|~v1_int_1(X76)|~v1_int_1(X75))&(~r1_int_1(X77,k6_xcmplx_0(X75,X76))|r2_int_1(X75,X76,X77)|~v1_int_1(X77)|~v1_int_1(X76)|~v1_int_1(X75))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t15_int_2])])])])).
fof(c_0_149, plain, ![X83, X84]:(~v1_int_1(X83)|(~v7_ordinal1(X84)|~v1_int_2(X84)|(k3_int_2(X83,X84)!=np__1|(r1_xxreal_0(X84,np__2)|r2_int_1(k2_int_5(X83,X84),k1_newton(X83,k3_nat_d(k7_nat_d(X84,np__1),np__2)),X84))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_int_5])])])).
fof(c_0_150, negated_conjecture, ~(![X1]:((v7_ordinal1(X1)&v1_int_2(X1))=>(~(r1_xxreal_0(X1,np__2))=>k2_int_5(k4_xcmplx_0(np__1),X1)=k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(X1,np__1),np__2))))), inference(assume_negation,[status(cth)],[t36_int_5])).
cnf(c_0_151, plain, (k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))=np__1|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(np__1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_134])])).
cnf(c_0_152, plain, (v1_int_1(k4_xcmplx_0(np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_136]), c_0_137]), c_0_138])])])).
cnf(c_0_153, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r1_nat_d(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_139])).
cnf(c_0_154, plain, (X1=k6_numbers|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_140, c_0_141])).
cnf(c_0_155, plain, (v1_xboole_0(np__0)), inference(evalgc,[status(thm)],[c_0_142])).
fof(c_0_156, plain, ![X66, X67]:((~r1_nat_d(X66,X67)|r1_int_1(X66,X67)|(~v7_ordinal1(X66)|~v7_ordinal1(X67)))&(~r1_int_1(X66,X67)|r1_nat_d(X66,X67)|(~v7_ordinal1(X66)|~v7_ordinal1(X67)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r1_nat_d])])])).
cnf(c_0_157, plain, (r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(X1,k1_newton(X2,X3))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(X1)|~v7_ordinal1(X3)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_97])])).
cnf(c_0_158, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_144])])).
cnf(c_0_159, plain, (v1_int_1(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_123]), c_0_124])])])).
cnf(c_0_160, plain, (v1_xreal_0(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_146]), c_0_104])])])).
cnf(c_0_161, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X2)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_128])])).
cnf(c_0_162, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~epred1_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_163, plain, (r1_int_1(X3,k6_xcmplx_0(X1,X2))|~r2_int_1(X1,X2,X3)|~v1_int_1(X3)|~v1_int_1(X2)|~v1_int_1(X1)), inference(split_conjunct,[status(thm)],[c_0_148])).
cnf(c_0_164, plain, (r1_xxreal_0(X2,np__2)|r2_int_1(k2_int_5(X1,X2),k1_newton(X1,k3_nat_d(k7_nat_d(X2,np__1),np__2)),X2)|~v1_int_1(X1)|~v7_ordinal1(X2)|~v1_int_2(X2)|k3_int_2(X1,X2)!=np__1), inference(split_conjunct,[status(thm)],[c_0_149])).
fof(c_0_165, plain, ![X15]:(~v7_ordinal1(X15)|v1_int_1(X15)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_int_1])])).
fof(c_0_166, plain, ![X36, X37]:(~v1_int_1(X36)|~v7_ordinal1(X37)|~v1_int_2(X37)|v1_int_1(k2_int_5(X36,X37))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_int_5])])).
fof(c_0_167, negated_conjecture, ((v7_ordinal1(esk2_0)&v1_int_2(esk2_0))&(~r1_xxreal_0(esk2_0,np__2)&k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_150])])])])).
cnf(c_0_168, plain, (k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))=np__1|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~m1_subset_1(X1,k5_numbers)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_151, c_0_152])]), c_0_99])).
cnf(c_0_169, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_nat_d(X2,X1)), inference(evalgc,[status(thm)],[c_0_153])).
cnf(c_0_170, plain, (k6_numbers=np__0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_155])])).
cnf(c_0_171, plain, (r1_nat_d(X1,X2)|~r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_156])).
cnf(c_0_172, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_117]), c_0_118])])])).
cnf(c_0_173, plain, (k4_xcmplx_0(k4_xcmplx_0(np__2))=np__2), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_159])])).
cnf(c_0_174, plain, (v1_xcmplx_0(k4_xcmplx_0(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_160])])).
cnf(c_0_175, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k4_xcmplx_0(X3))|~r1_int_1(k1_newton(X1,X2),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_97])])).
fof(c_0_176, plain, ~r1_xxreal_0(np__2,np__0), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r2_r0])).
fof(c_0_177, plain, ![X16]:(~v7_ordinal1(X16)|v1_xreal_0(X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_xreal_0])])).
cnf(c_0_178, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~epred1_2(X2,X1)), inference(evalgc,[status(thm)],[c_0_162])).
cnf(c_0_179, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~epred1_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_180, plain, (r1_int_1(X3,k6_xcmplx_0(X1,X2))|~v1_int_1(X3)|~v1_int_1(X2)|~v1_int_1(X1)|~r2_int_1(X1,X2,X3)), inference(evalgc,[status(thm)],[c_0_163])).
cnf(c_0_181, plain, (r1_xxreal_0(X2,np__2)|r2_int_1(k2_int_5(X1,X2),k1_newton(X1,k3_nat_d(k7_nat_d(X2,np__1),np__2)),X2)|k3_int_2(X1,X2)!=np__1|~v1_int_1(X1)|~v7_ordinal1(X2)|~v1_int_2(X2)), inference(evalgc,[status(thm)],[c_0_164])).
cnf(c_0_182, plain, (v1_int_1(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_165])).
cnf(c_0_183, plain, (v1_int_1(k2_int_5(X1,X2))|~v1_int_1(X1)|~v7_ordinal1(X2)|~v1_int_2(X2)), inference(split_conjunct,[status(thm)],[c_0_166])).
cnf(c_0_184, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))), inference(split_conjunct,[status(thm)],[c_0_167])).
cnf(c_0_185, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(np__1)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_168]), c_0_136])])])).
fof(c_0_186, plain, ![X81, X82]:(~v1_int_1(X81)|(~v7_ordinal1(X82)|~v1_int_2(X82)|(k2_int_5(X81,X82)=np__1|k2_int_5(X81,X82)=k6_numbers|k2_int_5(X81,X82)=k4_xcmplx_0(np__1)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t25_int_5])])])).
cnf(c_0_187, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~r1_nat_d(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_169, c_0_170])])).
cnf(c_0_188, plain, (r1_nat_d(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_int_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_171])).
cnf(c_0_189, plain, (r1_int_1(X1,np__2)|~r1_int_1(X1,k4_xcmplx_0(np__2))|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_173]), c_0_159]), c_0_174])])])).
cnf(c_0_190, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_117]), c_0_118])])])).
cnf(c_0_191, plain, (~r1_xxreal_0(np__2,np__0)), inference(split_conjunct,[status(thm)],[c_0_176])).
cnf(c_0_192, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_177])).
cnf(c_0_193, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~v1_int_1(X2)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_128])])).
fof(c_0_194, plain, ![X59]:(~v1_int_1(X59)|k1_int_2(X59)=k16_complex1(X59)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_int_2])])).
fof(c_0_195, plain, ![X58]:(~v1_xcmplx_0(X58)|k18_complex1(X58)=k16_complex1(X58)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k18_complex1])])).
cnf(c_0_196, plain, (r1_int_1(X1,X2)|~epred1_2(X2,X1)|~r1_int_1(k4_xcmplx_0(X1),X2)), inference(evalgc,[status(thm)],[c_0_179])).
cnf(c_0_197, plain, (r1_int_1(X1,k6_xcmplx_0(k2_int_5(X2,X1),k1_newton(X2,k3_nat_d(k7_nat_d(X1,np__1),np__2))))|r1_xxreal_0(X1,np__2)|~v1_int_1(X1)|~v1_int_1(k1_newton(X2,k3_nat_d(k7_nat_d(X1,np__1),np__2)))|~v1_int_1(k2_int_5(X2,X1))|k3_int_2(X2,X1)!=np__1|~v1_int_2(X1)|~v7_ordinal1(X1)|~v1_int_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_181])])).
cnf(c_0_198, plain, (v1_int_1(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_182])).
cnf(c_0_199, plain, (v1_int_1(k2_int_5(X1,X2))|~v1_int_1(X1)|~v7_ordinal1(X2)|~v1_int_2(X2)), inference(evalgc,[status(thm)],[c_0_183])).
cnf(c_0_200, negated_conjecture, (k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k2_int_5(k4_xcmplx_0(np__1),esk2_0)), inference(evalgc,[status(thm)],[c_0_184])).
cnf(c_0_201, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(np__1)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~m1_subset_1(X1,k5_numbers)), inference(csr,[status(thm)],[c_0_185, c_0_99])).
cnf(c_0_202, negated_conjecture, (v1_int_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_167])).
cnf(c_0_203, negated_conjecture, (v7_ordinal1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_167])).
cnf(c_0_204, negated_conjecture, (~r1_xxreal_0(esk2_0,np__2)), inference(split_conjunct,[status(thm)],[c_0_167])).
fof(c_0_205, plain, ![X38, X39]:(~v7_ordinal1(X38)|~v7_ordinal1(X39)|m1_subset_1(k3_nat_d(X38,X39),k5_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_nat_d])])).
cnf(c_0_206, plain, (k2_int_5(X1,X2)=np__1|k2_int_5(X1,X2)=k6_numbers|k2_int_5(X1,X2)=k4_xcmplx_0(np__1)|~v1_int_1(X1)|~v7_ordinal1(X2)|~v1_int_2(X2)), inference(split_conjunct,[status(thm)],[c_0_186])).
cnf(c_0_207, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_int_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_188])])).
cnf(c_0_208, plain, (r1_int_1(k4_xcmplx_0(X1),np__2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,np__2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_190]), c_0_159])])])).
cnf(c_0_209, plain, (v7_ordinal1(np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_130])])).
cnf(c_0_210, plain, (~r1_xxreal_0(np__2,np__0)), inference(evalgc,[status(thm)],[c_0_191])).
cnf(c_0_211, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_192])).
cnf(c_0_212, plain, (r1_int_1(k1_newton(X1,X2),X3)|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_97])])).
cnf(c_0_213, plain, (k1_int_2(X1)=k16_complex1(X1)|~v1_int_1(X1)), inference(split_conjunct,[status(thm)],[c_0_194])).
cnf(c_0_214, plain, (k18_complex1(X1)=k16_complex1(X1)|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_195])).
fof(c_0_215, plain, ![X33, X34]:(~v1_xxreal_0(X33)|~v1_xxreal_0(X34)|(r1_xxreal_0(X33,X34)|r1_xxreal_0(X34,X33))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[connectedness_r1_xxreal_0])])).
fof(c_0_216, plain, ![X17]:(~v7_ordinal1(X17)|v1_xxreal_0(X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_xxreal_0])])).
cnf(c_0_217, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X2)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_128])])).
cnf(c_0_218, plain, (r1_int_1(X1,k6_xcmplx_0(k2_int_5(X2,X1),k1_newton(X2,k3_nat_d(k7_nat_d(X1,np__1),np__2))))|r1_xxreal_0(X1,np__2)|k3_int_2(X2,X1)!=np__1|~v1_int_2(X1)|~v7_ordinal1(X1)|~v1_int_1(k1_newton(X2,k3_nat_d(k7_nat_d(X1,np__1),np__2)))|~v1_int_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_197, c_0_198]), c_0_199])).
cnf(c_0_219, negated_conjecture, (k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_201])])).
cnf(c_0_220, negated_conjecture, (v1_int_2(esk2_0)), inference(evalgc,[status(thm)],[c_0_202])).
cnf(c_0_221, negated_conjecture, (v7_ordinal1(esk2_0)), inference(evalgc,[status(thm)],[c_0_203])).
cnf(c_0_222, negated_conjecture, (~r1_xxreal_0(esk2_0,np__2)), inference(evalgc,[status(thm)],[c_0_204])).
cnf(c_0_223, plain, (m1_subset_1(k3_nat_d(X1,X2),k5_numbers)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_205])).
cnf(c_0_224, plain, (k2_int_5(X1,X2)=k6_numbers|k2_int_5(X1,X2)=np__1|k2_int_5(X1,X2)=k4_xcmplx_0(np__1)|~v1_int_1(X1)|~v7_ordinal1(X2)|~v1_int_2(X2)), inference(evalgc,[status(thm)],[c_0_206])).
fof(c_0_225, plain, ![X10, X11]:(~v6_membered(X10)|(~m1_subset_1(X11,X10)|v7_ordinal1(X11))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc17_membered])])])).
fof(c_0_226, plain, ![X41, X42]:(~v7_ordinal1(X41)|~v7_ordinal1(X42)|m1_subset_1(k7_nat_d(X41,X42),k5_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_nat_d])])).
cnf(c_0_227, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__2)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,np__2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_208]), c_0_209])]), c_0_210])])).
cnf(c_0_228, plain, (v1_xcmplx_0(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_211])])).
cnf(c_0_229, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_117]), c_0_118])])])).
fof(c_0_230, plain, ![X89]:(~v7_ordinal1(X89)|k6_nat_d(X89,np__1)=np__1), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t51_newton])])).
fof(c_0_231, plain, ![X31, X32]:(~v7_ordinal1(X31)|~v7_ordinal1(X32)|k6_nat_d(X31,X32)=k6_nat_d(X32,X31)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k6_nat_d])])).
cnf(c_0_232, plain, (k16_complex1(X1)=k1_int_2(X1)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[c_0_213])).
cnf(c_0_233, plain, (k16_complex1(X1)=k18_complex1(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_214])).
cnf(c_0_234, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_215])).
cnf(c_0_235, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_216])).
cnf(c_0_236, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~r1_int_1(k1_newton(X1,X2),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_97])])).
fof(c_0_237, plain, ~r1_xxreal_0(np__1,np__0), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r1_r0])).
cnf(c_0_238, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~epred1_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_239, negated_conjecture, (r1_int_1(esk2_0,k6_xcmplx_0(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_219]), c_0_220]), c_0_221]), c_0_138]), c_0_152])]), c_0_222])])).
cnf(c_0_240, plain, (m1_subset_1(k3_nat_d(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_223])).
cnf(c_0_241, plain, (k2_int_5(X1,X2)=k4_xcmplx_0(np__1)|k2_int_5(X1,X2)=np__1|k2_int_5(X1,X2)=np__0|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_224, c_0_170])])).
cnf(c_0_242, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),np__1)=k4_xcmplx_0(np__2)), inference(split_conjunct,[status(thm)],[rqRealDiff__k6_xcmplx_0__rm1_r1_rm2])).
cnf(c_0_243, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_225])).
cnf(c_0_244, plain, (m1_subset_1(k7_nat_d(X1,X2),k5_numbers)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_226])).
cnf(c_0_245, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__2)|~r1_int_1(X1,np__2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_227, c_0_198])).
cnf(c_0_246, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_228])])).
cnf(c_0_247, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_172])])).
cnf(c_0_248, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_int_1(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_249, plain, (k6_nat_d(X1,np__1)=np__1|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_230])).
cnf(c_0_250, plain, (k6_nat_d(X1,X2)=k6_nat_d(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_231])).
cnf(c_0_251, plain, (k18_complex1(X1)=k1_int_2(X1)|~v1_int_1(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_233])])).
cnf(c_0_252, plain, (k18_complex1(k4_xcmplx_0(np__1))=np__1|~v1_xcmplx_0(k4_xcmplx_0(np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_117]), c_0_100])])])).
fof(c_0_253, plain, ![X73, X74]:((r1_xxreal_0(k1_nat_1(X74,np__1),X73)|r1_xxreal_0(X73,X74)|~v7_ordinal1(X74)|~v7_ordinal1(X73))&(~r1_xxreal_0(X73,X74)|~r1_xxreal_0(k1_nat_1(X74,np__1),X73)|~v7_ordinal1(X74)|~v7_ordinal1(X73))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t13_nat_1])])])])])).
fof(c_0_254, plain, ![X60, X61]:(~v7_ordinal1(X60)|~m1_subset_1(X61,k5_numbers)|k1_nat_1(X60,X61)=k2_xcmplx_0(X60,X61)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_nat_1])])).
cnf(c_0_255, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_234])).
cnf(c_0_256, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_235])).
cnf(c_0_257, plain, (r1_int_1(X1,np__1)|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_137]), c_0_138]), c_0_136])])])).
cnf(c_0_258, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_236, c_0_117]), c_0_118])])])).
cnf(c_0_259, plain, (~r1_xxreal_0(np__1,np__0)), inference(split_conjunct,[status(thm)],[c_0_237])).
cnf(c_0_260, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~epred1_2(X2,X1)), inference(evalgc,[status(thm)],[c_0_238])).
cnf(c_0_261, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~epred1_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_262, negated_conjecture, (r1_int_1(esk2_0,k6_xcmplx_0(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_240]), c_0_209])])])).
cnf(c_0_263, negated_conjecture, (k2_int_5(X1,esk2_0)=k4_xcmplx_0(np__1)|k2_int_5(X1,esk2_0)=np__0|k2_int_5(X1,esk2_0)=np__1|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_220]), c_0_221])])])).
cnf(c_0_264, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),np__1)=k4_xcmplx_0(np__2)), inference(evalgc,[status(thm)],[c_0_242])).
cnf(c_0_265, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_243])).
cnf(c_0_266, plain, (m1_subset_1(k7_nat_d(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_244])).
cnf(c_0_267, plain, (r1_xxreal_0(X1,np__2)|~r1_int_1(k4_xcmplx_0(X1),np__2)|~v7_ordinal1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_246])])).
cnf(c_0_268, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_211])])).
cnf(c_0_269, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_246])])).
cnf(c_0_270, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_int_1(X1)), inference(evalgc,[status(thm)],[c_0_248])).
fof(c_0_271, plain, ![X64, X65]:(~v7_ordinal1(X64)|~v7_ordinal1(X65)|k6_nat_d(X64,X65)=k3_int_2(X64,X65)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k6_nat_d])])).
cnf(c_0_272, plain, (k6_nat_d(X1,np__1)=np__1|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_249])).
cnf(c_0_273, plain, (k6_nat_d(X1,X2)=k6_nat_d(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_250])).
fof(c_0_274, plain, ![X85, X86]:(~v1_int_1(X85)|(~v1_int_1(X86)|k3_int_2(X85,X86)=k3_int_2(k1_int_2(X85),k1_int_2(X86)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t34_int_2])])])).
cnf(c_0_275, plain, (k18_complex1(X1)=k1_int_2(X1)|~v1_int_1(X1)), inference(csr,[status(thm)],[c_0_251, c_0_144])).
cnf(c_0_276, plain, (k18_complex1(k4_xcmplx_0(np__1))=np__1), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_252, c_0_136])])])).
cnf(c_0_277, plain, (~r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_nat_1(X2,np__1),X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_253])).
cnf(c_0_278, plain, (k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_254])).
cnf(c_0_279, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_256])])).
cnf(c_0_280, plain, (r1_int_1(k4_xcmplx_0(X1),np__1)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_258]), c_0_152])])])).
cnf(c_0_281, plain, (~r1_xxreal_0(np__1,np__0)), inference(evalgc,[status(thm)],[c_0_259])).
cnf(c_0_282, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v1_int_1(X2)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_128])])).
cnf(c_0_283, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~epred1_2(X2,X1)|~r1_int_1(X1,k4_xcmplx_0(X2))), inference(evalgc,[status(thm)],[c_0_261])).
fof(c_0_284, plain, ![X56]:(~v1_int_1(X56)|k1_int_2(k1_int_2(X56))=k1_int_2(X56)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[projectivity_k1_int_2])])).
fof(c_0_285, plain, ![X35]:(~v1_int_1(X35)|m1_subset_1(k1_int_2(X35),k5_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_int_2])])).
cnf(c_0_286, negated_conjecture, (r1_int_1(esk2_0,k4_xcmplx_0(np__2))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_263]), c_0_264]), c_0_152])])])).
cnf(c_0_287, plain, (v7_ordinal1(k7_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_266]), c_0_124])])])).
cnf(c_0_288, plain, (r1_xxreal_0(X1,np__2)|~r1_int_1(k4_xcmplx_0(X1),np__2)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_267, c_0_268])).
cnf(c_0_289, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_269, c_0_270]), c_0_198])).
cnf(c_0_290, plain, (m1_subset_1(np__0,k5_numbers)), inference(split_conjunct,[status(thm)],[spc0_numerals])).
cnf(c_0_291, plain, (k6_nat_d(X1,X2)=k3_int_2(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_271])).
cnf(c_0_292, plain, (k6_nat_d(np__1,X1)=np__1|~v7_ordinal1(X1)|~v7_ordinal1(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_272, c_0_273])])).
cnf(c_0_293, plain, (k3_int_2(X1,X2)=k3_int_2(k1_int_2(X1),k1_int_2(X2))|~v1_int_1(X1)|~v1_int_1(X2)), inference(split_conjunct,[status(thm)],[c_0_274])).
cnf(c_0_294, plain, (k1_int_2(k4_xcmplx_0(np__1))=np__1|~v1_int_1(k4_xcmplx_0(np__1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_276])])).
cnf(c_0_295, plain, (~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_nat_1(X2,np__1),X1)), inference(evalgc,[status(thm)],[c_0_277])).
cnf(c_0_296, plain, (k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_278])).
cnf(c_0_297, plain, (k2_xcmplx_0(np__1,np__1)=np__2), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r1_r2])).
cnf(c_0_298, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_279, c_0_256])])).
cnf(c_0_299, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__1)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_280]), c_0_118])]), c_0_281])])).
cnf(c_0_300, plain, (r1_int_1(X1,k1_newton(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(X1)|~v7_ordinal1(X3)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_97])])).
cnf(c_0_301, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_128])])).
cnf(c_0_302, plain, (k1_int_2(k1_int_2(X1))=k1_int_2(X1)|~v1_int_1(X1)), inference(split_conjunct,[status(thm)],[c_0_284])).
cnf(c_0_303, plain, (m1_subset_1(k1_int_2(X1),k5_numbers)|~v1_int_1(X1)), inference(split_conjunct,[status(thm)],[c_0_285])).
cnf(c_0_304, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r1_int_1(esk2_0,k4_xcmplx_0(np__2))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_287]), c_0_118]), c_0_221])])])).
cnf(c_0_305, negated_conjecture, (~r1_int_1(k4_xcmplx_0(esk2_0),np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_288]), c_0_221])])])).
cnf(c_0_306, plain, (r1_int_1(k4_xcmplx_0(X1),np__2)|~r1_int_1(X1,k4_xcmplx_0(np__2))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_289, c_0_173]), c_0_159]), c_0_174])])])).
cnf(c_0_307, plain, (m1_subset_1(np__0,k5_numbers)), inference(evalgc,[status(thm)],[c_0_290])).
cnf(c_0_308, plain, (k6_nat_d(X1,X2)=k3_int_2(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_291])).
cnf(c_0_309, plain, (k6_nat_d(np__1,X1)=np__1|~v7_ordinal1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_292, c_0_118])])).
cnf(c_0_310, plain, (k3_int_2(k1_int_2(X1),k1_int_2(X2))=k3_int_2(X1,X2)|~v1_int_1(X2)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[c_0_293])).
cnf(c_0_311, plain, (k1_int_2(k4_xcmplx_0(np__1))=np__1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_294, c_0_152])])).
cnf(c_0_312, plain, (~r1_xxreal_0(k2_xcmplx_0(X1,np__1),X2)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_295, c_0_296]), c_0_100])])])).
cnf(c_0_313, plain, (k2_xcmplx_0(np__1,np__1)=np__2), inference(evalgc,[status(thm)],[c_0_297])).
cnf(c_0_314, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|r1_xxreal_0(esk2_0,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_221])])).
cnf(c_0_315, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__1)|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_299, c_0_198])).
cnf(c_0_316, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_246])])).
cnf(c_0_317, plain, (r1_int_1(k4_xcmplx_0(X1),k1_newton(X2,X3))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k1_newton(X2,X3))))|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_300, c_0_301])])).
cnf(c_0_318, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_289])])).
cnf(c_0_319, plain, (k1_int_2(k1_int_2(X1))=k1_int_2(X1)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[c_0_302])).
cnf(c_0_320, plain, (m1_subset_1(k1_int_2(X1),k5_numbers)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[c_0_303])).
cnf(c_0_321, negated_conjecture, (r2_int_1(np__0,k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)),esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r1_int_1(esk2_0,k4_xcmplx_0(np__2))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_304]), c_0_220]), c_0_221]), c_0_152])]), c_0_222])])).
cnf(c_0_322, negated_conjecture, (~r1_int_1(esk2_0,k4_xcmplx_0(np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_306]), c_0_221])])])).
cnf(c_0_323, plain, (v1_int_1(np__0)|~v5_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_307])])).
cnf(c_0_324, plain, (k3_int_2(np__1,X1)=np__1|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_308, c_0_309]), c_0_118])])])).
cnf(c_0_325, plain, (k3_int_2(np__1,k1_int_2(X1))=k3_int_2(k4_xcmplx_0(np__1),X1)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_311]), c_0_152])])])).
cnf(c_0_326, plain, (~r1_xxreal_0(np__2,X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_312, c_0_313]), c_0_118])])])).
cnf(c_0_327, negated_conjecture, (r1_xxreal_0(np__2,esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_314, c_0_209]), c_0_222])])).
cnf(c_0_328, plain, (r1_xxreal_0(X1,np__1)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v7_ordinal1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_246])])).
cnf(c_0_329, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_316])])).
cnf(c_0_330, plain, (r1_int_1(k4_xcmplx_0(X1),k1_newton(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k1_newton(X2,X3))))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(X1)|~v1_xcmplx_0(X2)), inference(csr,[status(thm)],[c_0_317, c_0_96])).
cnf(c_0_331, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)), inference(csr,[status(thm)],[c_0_318, c_0_198])).
cnf(c_0_332, plain, (k3_int_2(k1_int_2(X1),k1_int_2(X2))=k3_int_2(k1_int_2(X1),X2)|~v1_int_1(X2)|~v1_int_1(k1_int_2(X1))|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_319])])).
cnf(c_0_333, plain, (v1_int_1(k1_int_2(X1))|~v5_membered(k5_numbers)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_320])])).
cnf(c_0_334, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(np__0,k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)),esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1), inference(sr,[status(thm)],[c_0_321, c_0_322])).
cnf(c_0_335, plain, (v1_int_1(np__0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_123]), c_0_124])])])).
cnf(c_0_336, plain, (k3_int_2(k4_xcmplx_0(np__1),X1)=np__1|~v7_ordinal1(k1_int_2(X1))|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_324, c_0_325])])).
cnf(c_0_337, plain, (v7_ordinal1(k1_int_2(X1))|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_320])])).
cnf(c_0_338, negated_conjecture, (~r1_xxreal_0(esk2_0,np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_326, c_0_327]), c_0_221])])])).
cnf(c_0_339, plain, (r1_xxreal_0(X1,np__1)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_328, c_0_268])).
cnf(c_0_340, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_246])])).
cnf(c_0_341, plain, (r1_int_1(k4_xcmplx_0(X1),k1_newton(X2,X3))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(X1)|~v1_xcmplx_0(X2)|~r1_int_1(X1,k1_newton(X2,X3))|~v7_ordinal1(X1)|~v1_xcmplx_0(k1_newton(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_330, c_0_331])])).
cnf(c_0_342, plain, (k3_int_2(k1_int_2(X1),X2)=k3_int_2(X1,X2)|~v1_int_1(X2)|~v1_int_1(X1)|~v1_int_1(k1_int_2(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_332])])).
cnf(c_0_343, plain, (v1_int_1(k1_int_2(X1))|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_333, c_0_123]), c_0_124])])])).
cnf(c_0_344, plain, (k1_int_2(X1)=k4_xcmplx_0(X1)|k1_int_2(X1)=X1|~v1_xreal_0(X1)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_275])])).
cnf(c_0_345, negated_conjecture, (r1_int_1(esk2_0,k6_xcmplx_0(np__0,k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|~v1_int_1(esk2_0)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_334]), c_0_335])])])).
cnf(c_0_346, plain, (k3_int_2(k4_xcmplx_0(np__1),X1)=np__1|~v1_int_1(X1)), inference(csr,[status(thm)],[c_0_336, c_0_337])).
cnf(c_0_347, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(np__1))=np__1), inference(split_conjunct,[status(thm)],[rqRealDiff__k6_xcmplx_0__r0_rm1_r1])).
cnf(c_0_348, negated_conjecture, (~r1_int_1(k4_xcmplx_0(esk2_0),k4_xcmplx_0(np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_338, c_0_339]), c_0_221])])])).
cnf(c_0_349, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_340, c_0_270]), c_0_198])).
cnf(c_0_350, plain, (r1_int_1(k4_xcmplx_0(X1),k1_newton(X2,X3))|~r1_int_1(X1,k1_newton(X2,X3))|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_xcmplx_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_341, c_0_78]), c_0_198])).
cnf(c_0_351, plain, (k6_nat_d(X2,X1)=k3_int_2(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_308, c_0_273])])).
cnf(c_0_352, plain, (k3_int_2(k1_int_2(X1),X2)=k3_int_2(X1,X2)|~v1_int_1(X2)|~v1_int_1(X1)), inference(csr,[status(thm)],[c_0_342, c_0_343])).
cnf(c_0_353, plain, (k1_int_2(X1)=k4_xcmplx_0(X1)|k1_int_2(X1)=X1|~v1_int_1(X1)), inference(csr,[status(thm)],[c_0_344, c_0_114])).
cnf(c_0_354, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r1_int_1(esk2_0,k6_xcmplx_0(np__0,k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|~v1_int_1(esk2_0)), inference(csr,[status(thm)],[c_0_345, c_0_346])).
cnf(c_0_355, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(np__1))=np__1), inference(evalgc,[status(thm)],[c_0_347])).
cnf(c_0_356, plain, (~r1_int_1(esk2_0,np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_348, c_0_349]), c_0_118]), c_0_221]), c_0_152])])])).
cnf(c_0_357, plain, (k6_xcmplx_0(np__0,np__1)=k4_xcmplx_0(np__1)), inference(split_conjunct,[status(thm)],[rqRealDiff__k6_xcmplx_0__r0_r1_rm1])).
cnf(c_0_358, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_350, c_0_117]), c_0_118])])])).
cnf(c_0_359, plain, (k3_int_2(X2,X1)=k3_int_2(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_308, c_0_351])])).
cnf(c_0_360, plain, (k3_int_2(k4_xcmplx_0(X1),X2)=k3_int_2(X1,X2)|k1_int_2(X1)=X1|~v1_int_1(X2)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_352, c_0_353])])).
cnf(c_0_361, plain, (k1_newton(k4_xcmplx_0(np__1),X2)=np__1|k4_xcmplx_0(np__1)!=np__1|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(ef,[status(thm)],[c_0_201])])).
cnf(c_0_362, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))=np__1|~v1_int_1(esk2_0)|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_354, c_0_201]), c_0_355]), c_0_152])]), c_0_356])])).
cnf(c_0_363, plain, (k6_xcmplx_0(np__0,np__1)=k4_xcmplx_0(np__1)), inference(evalgc,[status(thm)],[c_0_357])).
cnf(c_0_364, plain, (~r1_int_1(esk2_0,k4_xcmplx_0(np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_348, c_0_358]), c_0_221]), c_0_137]), c_0_138]), c_0_136])])])).
cnf(c_0_365, plain, (k3_int_2(X2,k4_xcmplx_0(X1))=k3_int_2(X1,X2)|k1_int_2(X1)=X1|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v1_int_1(X2)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_359, c_0_360])])).
cnf(c_0_366, plain, (v7_ordinal1(k4_xcmplx_0(X1))|k1_int_2(X1)=X1|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_337, c_0_353])])).
cnf(c_0_367, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__1|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_361])])).
cnf(c_0_368, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|~v1_int_1(esk2_0)|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_354, c_0_362]), c_0_363]), c_0_138])]), c_0_364])])).
cnf(c_0_369, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~epred1_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_370, plain, (k3_int_2(X1,k4_xcmplx_0(X2))=k3_int_2(X2,X1)|k1_int_2(X2)=X2|~v7_ordinal1(X1)|~v1_int_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_365, c_0_198]), c_0_366])).
cnf(c_0_371, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_240]), c_0_209])])])).
cnf(c_0_372, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|~v1_int_1(esk2_0)|~v7_ordinal1(k7_nat_d(esk2_0,np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_368, c_0_240]), c_0_209])])])).
cnf(c_0_373, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~epred1_2(X2,X1)), inference(evalgc,[status(thm)],[c_0_369])).
cnf(c_0_374, plain, (r2_int_1(k2_int_5(k4_xcmplx_0(np__1),X1),k4_xcmplx_0(np__1),X1)|r1_xxreal_0(X1,np__2)|k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(X1,np__1),np__2))=np__1|k3_int_2(k4_xcmplx_0(np__1),X1)!=np__1|~v1_int_2(X1)|~v7_ordinal1(X1)|~m1_subset_1(k3_nat_d(k7_nat_d(X1,np__1),np__2),k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_201]), c_0_152])])])).
cnf(c_0_375, plain, (k3_int_2(k4_xcmplx_0(X2),X1)=k3_int_2(X2,X1)|k1_int_2(X2)=X2|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_359, c_0_370])])).
cnf(c_0_376, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__1), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_371, c_0_287]), c_0_118]), c_0_221])])])).
cnf(c_0_377, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|~v1_int_1(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_372, c_0_287]), c_0_118]), c_0_221])])])).
cnf(c_0_378, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v1_int_1(X2)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_128])])).
cnf(c_0_379, plain, (k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(X1,np__1),np__2))=np__1|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),X1),k4_xcmplx_0(np__1),X1)|r1_xxreal_0(X1,np__2)|k3_int_2(k4_xcmplx_0(np__1),X1)!=np__1|~v1_int_2(X1)|~v7_ordinal1(X1)|~v7_ordinal1(k7_nat_d(X1,np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_240]), c_0_209])])])).
cnf(c_0_380, plain, (k3_int_2(k4_xcmplx_0(X1),X2)=k3_int_2(X1,X2)|k1_int_2(X1)=X1|~v7_ordinal1(X2)|~v1_int_1(X1)), inference(csr,[status(thm)],[c_0_375, c_0_366])).
cnf(c_0_381, negated_conjecture, (k4_xcmplx_0(np__1)!=np__1|~v1_int_1(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_376, c_0_377])])).
cnf(c_0_382, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(X3))|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_378, c_0_97])])).
cnf(c_0_383, plain, (k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(X1,np__1),np__2))=np__1|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),X1),k4_xcmplx_0(np__1),X1)|r1_xxreal_0(X1,np__2)|k3_int_2(k4_xcmplx_0(np__1),X1)!=np__1|~v1_int_2(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_379, c_0_287]), c_0_118])])])).
cnf(c_0_384, plain, (k3_int_2(k4_xcmplx_0(np__1),X1)=k3_int_2(np__1,X1)|k4_xcmplx_0(np__1)=np__1|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_380, c_0_137]), c_0_311]), c_0_152])])])).
cnf(c_0_385, negated_conjecture, (k4_xcmplx_0(np__1)!=np__1), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_381, c_0_198]), c_0_221])])])).
cnf(c_0_386, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_382, c_0_117]), c_0_118])])])).
cnf(c_0_387, plain, (k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(X1,np__1),np__2))=np__1|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),X1),k4_xcmplx_0(np__1),X1)|r1_xxreal_0(X1,np__2)|k3_int_2(np__1,X1)!=np__1|~v1_int_2(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_384]), c_0_385])])).
cnf(c_0_388, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_386, c_0_316])])).
cnf(c_0_389, plain, (k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(X1,np__1),np__2))=np__1|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),X1),k4_xcmplx_0(np__1),X1)|r1_xxreal_0(X1,np__2)|~v1_int_2(X1)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_387, c_0_324])).
cnf(c_0_390, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_388, c_0_246])])).
cnf(c_0_391, negated_conjecture, (k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))=np__1|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),k4_xcmplx_0(np__1),esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_389, c_0_220]), c_0_221])]), c_0_222])])).
cnf(c_0_392, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_390, c_0_246])])).
cnf(c_0_393, negated_conjecture, (r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),k4_xcmplx_0(np__1),esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=np__1), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_391])])).
cnf(c_0_394, plain, (k6_xcmplx_0(np__1,k4_xcmplx_0(np__1))=np__2), inference(split_conjunct,[status(thm)],[rqRealDiff__k6_xcmplx_0__r1_rm1_r2])).
cnf(c_0_395, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_392, c_0_270]), c_0_198])).
cnf(c_0_396, plain, (v1_int_1(k4_xcmplx_0(np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_174]), c_0_173]), c_0_159])])])).
cnf(c_0_397, negated_conjecture, (r2_int_1(np__1,k4_xcmplx_0(np__1),esk2_0)|~v1_int_1(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_377])])).
cnf(c_0_398, plain, (k6_xcmplx_0(np__1,k4_xcmplx_0(np__1))=np__2), inference(evalgc,[status(thm)],[c_0_394])).
cnf(c_0_399, negated_conjecture, (~r1_int_1(esk2_0,np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_395]), c_0_209]), c_0_221]), c_0_396])])])).
cnf(c_0_400, negated_conjecture, (~v1_int_1(esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_397]), c_0_398]), c_0_152]), c_0_138])]), c_0_399])])).
cnf(c_0_401, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_198]), c_0_221])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 402
# Proof object clause steps            : 299
# Proof object formula steps           : 103
# Proof object conjectures             : 42
# Proof object clause conjectures      : 39
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 69
# Proof object initial formulas used   : 56
# Proof object generating inferences   : 131
# Proof object simplifying inferences  : 191
# Training examples: 202 positive, 2609 negative
# Training: Positive examples begin
cnf(c_0_81, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_141, plain, (k1_xboole_0=k6_numbers)).# trainpos
cnf(c_0_104, plain, (v3_membered(k1_numbers))).# trainpos
cnf(c_0_107, plain, (v6_membered(k4_ordinal1))).# trainpos
cnf(c_0_221, negated_conjecture, (v7_ordinal1(esk2_0))).# trainpos
cnf(c_0_220, negated_conjecture, (v1_int_2(esk2_0))).# trainpos
cnf(c_0_155, plain, (v1_xboole_0(np__0))).# trainpos
cnf(c_0_222, negated_conjecture, (~r1_xxreal_0(esk2_0,np__2))).# trainpos
cnf(c_0_281, plain, (~r1_xxreal_0(np__1,np__0))).# trainpos
cnf(c_0_210, plain, (~r1_xxreal_0(np__2,np__0))).# trainpos
cnf(c_0_200, negated_conjecture, (k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k2_int_5(k4_xcmplx_0(np__1),esk2_0))).# trainpos
cnf(c_0_103, plain, (m1_subset_1(np__1,k1_numbers))).# trainpos
cnf(c_0_137, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1)).# trainpos
cnf(c_0_313, plain, (k2_xcmplx_0(np__1,np__1)=np__2)).# trainpos
cnf(c_0_100, plain, (m1_subset_1(np__1,k5_numbers))).# trainpos
cnf(c_0_307, plain, (m1_subset_1(np__0,k5_numbers))).# trainpos
cnf(c_0_146, plain, (m1_subset_1(np__2,k1_numbers))).# trainpos
cnf(c_0_130, plain, (m1_subset_1(np__2,k5_numbers))).# trainpos
cnf(c_0_140, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_363, plain, (k6_xcmplx_0(np__0,np__1)=k4_xcmplx_0(np__1))).# trainpos
cnf(c_0_398, plain, (k6_xcmplx_0(np__1,k4_xcmplx_0(np__1))=np__2)).# trainpos
cnf(c_0_355, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(np__1))=np__1)).# trainpos
cnf(c_0_114, plain, (v1_xreal_0(X1)|~v1_int_1(X1))).# trainpos
cnf(c_0_224, plain, (k2_int_5(X1,X2)=k6_numbers|k2_int_5(X1,X2)=np__1|k2_int_5(X1,X2)=k4_xcmplx_0(np__1)|~v1_int_1(X1)|~v7_ordinal1(X2)|~v1_int_2(X2))).# trainpos
cnf(c_0_129, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_211, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_198, plain, (v1_int_1(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_256, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_264, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),np__1)=k4_xcmplx_0(np__2))).# trainpos
cnf(c_0_123, plain, (v5_membered(X1)|~v6_membered(X1))).# trainpos
cnf(c_0_232, plain, (k16_complex1(X1)=k1_int_2(X1)|~v1_int_1(X1))).# trainpos
cnf(c_0_233, plain, (k16_complex1(X1)=k18_complex1(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_119, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_270, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_int_1(X1))).# trainpos
cnf(c_0_96, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1))).# trainpos
cnf(c_0_117, plain, (k1_newton(X1,np__1)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_77, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_80, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).# trainpos
cnf(c_0_295, plain, (~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_nat_1(X2,np__1),X1))).# trainpos
cnf(c_0_320, plain, (m1_subset_1(k1_int_2(X1),k5_numbers)|~v1_int_1(X1))).# trainpos
cnf(c_0_272, plain, (k6_nat_d(X1,np__1)=np__1|~v7_ordinal1(X1))).# trainpos
cnf(c_0_73, plain, (k18_complex1(k1_newton(k4_xcmplx_0(np__1),X1))=np__1|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_319, plain, (k1_int_2(k1_int_2(X1))=k1_int_2(X1)|~v1_int_1(X1))).# trainpos
cnf(c_0_74, plain, (k18_complex1(X1)=X1|k18_complex1(X1)=k4_xcmplx_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_128, plain, (epred1_2(X2,X1)|~v1_int_1(X2)|~v1_int_1(X1))).# trainpos
cnf(c_0_102, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_105, plain, (v1_int_1(X2)|~v5_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_265, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_78, plain, (v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_134, plain, (v1_int_1(k1_newton(X1,X2))|~v1_int_1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_127, plain, (r1_int_1(X1,X2)|~epred1_2(X2,X1)|~r1_int_1(X1,k4_xcmplx_0(X2)))).# trainpos
cnf(c_0_196, plain, (r1_int_1(X1,X2)|~epred1_2(X2,X1)|~r1_int_1(k4_xcmplx_0(X1),X2))).# trainpos
cnf(c_0_240, plain, (m1_subset_1(k3_nat_d(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_266, plain, (m1_subset_1(k7_nat_d(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_199, plain, (v1_int_1(k2_int_5(X1,X2))|~v1_int_1(X1)|~v7_ordinal1(X2)|~v1_int_2(X2))).# trainpos
cnf(c_0_308, plain, (k6_nat_d(X1,X2)=k3_int_2(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_273, plain, (k6_nat_d(X1,X2)=k6_nat_d(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_296, plain, (k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_260, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~epred1_2(X2,X1))).# trainpos
cnf(c_0_178, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~epred1_2(X2,X1))).# trainpos
cnf(c_0_255, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_188, plain, (r1_nat_d(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_int_1(X1,X2))).# trainpos
cnf(c_0_310, plain, (k3_int_2(k1_int_2(X1),k1_int_2(X2))=k3_int_2(X1,X2)|~v1_int_1(X2)|~v1_int_1(X1))).# trainpos
cnf(c_0_373, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~epred1_2(X2,X1))).# trainpos
cnf(c_0_169, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_nat_d(X2,X1))).# trainpos
cnf(c_0_147, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~epred1_2(X2,X1)|~r1_int_1(k4_xcmplx_0(X1),X2))).# trainpos
cnf(c_0_283, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~epred1_2(X2,X1)|~r1_int_1(X1,k4_xcmplx_0(X2)))).# trainpos
cnf(c_0_180, plain, (r1_int_1(X3,k6_xcmplx_0(X1,X2))|~v1_int_1(X3)|~v1_int_1(X2)|~v1_int_1(X1)|~r2_int_1(X1,X2,X3))).# trainpos
cnf(c_0_181, plain, (r1_xxreal_0(X2,np__2)|r2_int_1(k2_int_5(X1,X2),k1_newton(X1,k3_nat_d(k7_nat_d(X2,np__1),np__2)),X2)|k3_int_2(X1,X2)!=np__1|~v1_int_1(X1)|~v7_ordinal1(X2)|~v1_int_2(X2))).# trainpos
cnf(c_0_118, plain, (v7_ordinal1(np__1))).# trainpos
cnf(c_0_170, plain, (k6_numbers=np__0)).# trainpos
cnf(c_0_209, plain, (v7_ordinal1(np__2))).# trainpos
cnf(c_0_143, plain, (r1_int_1(X1,X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(X1))).# trainpos
cnf(c_0_292, plain, (k6_nat_d(np__1,X1)=np__1|~v7_ordinal1(X1)|~v7_ordinal1(np__1))).# trainpos
cnf(c_0_282, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v1_int_1(X2)|~v1_int_1(X1))).# trainpos
cnf(c_0_217, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X2)|~v1_int_1(X1))).# trainpos
cnf(c_0_160, plain, (v1_xreal_0(np__2))).# trainpos
cnf(c_0_120, plain, (v1_xreal_0(np__1))).# trainpos
cnf(c_0_228, plain, (v1_xcmplx_0(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_161, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X2)|~v1_int_1(X1))).# trainpos
cnf(c_0_144, plain, (v1_xcmplx_0(X1)|~v1_int_1(X1))).# trainpos
cnf(c_0_378, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v1_int_1(X2)|~v1_int_1(X1))).# trainpos
cnf(c_0_193, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~v1_int_1(X2)|~v1_int_1(X1))).# trainpos
cnf(c_0_251, plain, (k18_complex1(X1)=k1_int_2(X1)|~v1_int_1(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_122, plain, (v1_int_1(np__1)|~v5_membered(k5_numbers))).# trainpos
cnf(c_0_279, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_287, plain, (v7_ordinal1(k7_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_268, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_145, plain, (v1_int_1(np__2)|~v5_membered(k5_numbers))).# trainpos
cnf(c_0_323, plain, (v1_int_1(np__0)|~v5_membered(k5_numbers))).# trainpos
cnf(c_0_252, plain, (k18_complex1(k4_xcmplx_0(np__1))=np__1|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).# trainpos
cnf(c_0_301, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(X1))).# trainpos
cnf(c_0_351, plain, (k6_nat_d(X2,X1)=k3_int_2(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_337, plain, (v7_ordinal1(k1_int_2(X1))|~v1_int_1(X1))).# trainpos
cnf(c_0_312, plain, (~r1_xxreal_0(k2_xcmplx_0(X1,np__1),X2)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_197, plain, (r1_int_1(X1,k6_xcmplx_0(k2_int_5(X2,X1),k1_newton(X2,k3_nat_d(k7_nat_d(X1,np__1),np__2))))|r1_xxreal_0(X1,np__2)|~v1_int_1(X1)|~v1_int_1(k1_newton(X2,k3_nat_d(k7_nat_d(X1,np__1),np__2)))|~v1_int_1(k2_int_5(X2,X1))|k3_int_2(X2,X1)!=np__1|~v1_int_2(X1)|~v7_ordinal1(X1)|~v1_int_1(X2))).# trainpos
cnf(c_0_93, plain, (k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))=np__1|k18_complex1(k1_newton(k4_xcmplx_0(np__1),X1))=k1_newton(k4_xcmplx_0(np__1),X1)|~m1_subset_1(X1,k5_numbers)|~v1_xreal_0(k1_newton(k4_xcmplx_0(np__1),X1)))).# trainpos
cnf(c_0_97, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2)))=k1_newton(X1,X2)|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_333, plain, (v1_int_1(k1_int_2(X1))|~v5_membered(k5_numbers)|~v1_int_1(X1))).# trainpos
cnf(c_0_332, plain, (k3_int_2(k1_int_2(X1),k1_int_2(X2))=k3_int_2(k1_int_2(X1),X2)|~v1_int_1(X2)|~v1_int_1(k1_int_2(X1))|~v1_int_1(X1))).# trainpos
cnf(c_0_241, plain, (k2_int_5(X1,X2)=k4_xcmplx_0(np__1)|k2_int_5(X1,X2)=np__1|k2_int_5(X1,X2)=np__0|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).# trainpos
cnf(c_0_138, plain, (v1_int_1(np__1))).# trainpos
cnf(c_0_159, plain, (v1_int_1(np__2))).# trainpos
cnf(c_0_158, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_int_1(X1))).# trainpos
cnf(c_0_298, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_187, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~r1_nat_d(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_335, plain, (v1_int_1(np__0))).# trainpos
cnf(c_0_246, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v7_ordinal1(X1))).# trainpos
cnf(c_0_326, plain, (~r1_xxreal_0(np__2,X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_343, plain, (v1_int_1(k1_int_2(X1))|~v1_int_1(X1))).# trainpos
cnf(c_0_136, plain, (v1_xcmplx_0(k4_xcmplx_0(np__1)))).# trainpos
cnf(c_0_174, plain, (v1_xcmplx_0(k4_xcmplx_0(np__2)))).# trainpos
cnf(c_0_324, plain, (k3_int_2(np__1,X1)=np__1|~v7_ordinal1(X1))).# trainpos
cnf(c_0_344, plain, (k1_int_2(X1)=k4_xcmplx_0(X1)|k1_int_2(X1)=X1|~v1_xreal_0(X1)|~v1_int_1(X1))).# trainpos
cnf(c_0_113, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))=np__1|~m1_subset_1(X1,k5_numbers)|~v1_xreal_0(k1_newton(k4_xcmplx_0(np__1),X1)))).# trainpos
cnf(c_0_300, plain, (r1_int_1(X1,k1_newton(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(X1)|~v7_ordinal1(X3)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_342, plain, (k3_int_2(k1_int_2(X1),X2)=k3_int_2(X1,X2)|~v1_int_1(X2)|~v1_int_1(X1)|~v1_int_1(k1_int_2(X1)))).# trainpos
cnf(c_0_157, plain, (r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(X1,k1_newton(X2,X3))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(X1)|~v7_ordinal1(X3)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_236, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~r1_int_1(k1_newton(X1,X2),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_359, plain, (k3_int_2(X2,X1)=k3_int_2(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_116, plain, (v1_int_1(k1_newton(X1,X2))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_212, plain, (r1_int_1(k1_newton(X1,X2),X3)|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_382, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(X3))|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_175, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k4_xcmplx_0(X3))|~r1_int_1(k1_newton(X1,X2),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_263, negated_conjecture, (k2_int_5(X1,esk2_0)=k4_xcmplx_0(np__1)|k2_int_5(X1,esk2_0)=np__0|k2_int_5(X1,esk2_0)=np__1|~v1_int_1(X1))).# trainpos
cnf(c_0_173, plain, (k4_xcmplx_0(k4_xcmplx_0(np__2))=np__2)).# trainpos
cnf(c_0_314, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|r1_xxreal_0(esk2_0,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_276, plain, (k18_complex1(k4_xcmplx_0(np__1))=np__1)).# trainpos
cnf(c_0_135, plain, (v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_172, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_316, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_207, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_int_1(X2,X1))).# trainpos
cnf(c_0_366, plain, (v7_ordinal1(k4_xcmplx_0(X1))|k1_int_2(X1)=X1|~v1_int_1(X1))).# trainpos
cnf(c_0_133, plain, (k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))=np__1|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~m1_subset_1(X1,k5_numbers)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1)))).# trainpos
cnf(c_0_229, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_386, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_258, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_190, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_360, plain, (k3_int_2(k4_xcmplx_0(X1),X2)=k3_int_2(X1,X2)|k1_int_2(X1)=X1|~v1_int_1(X2)|~v1_int_1(X1))).# trainpos
cnf(c_0_152, plain, (v1_int_1(k4_xcmplx_0(np__1)))).# trainpos
cnf(c_0_396, plain, (v1_int_1(k4_xcmplx_0(np__2)))).# trainpos
cnf(c_0_327, negated_conjecture, (r1_xxreal_0(np__2,esk2_0))).# trainpos
cnf(c_0_294, plain, (k1_int_2(k4_xcmplx_0(np__1))=np__1|~v1_int_1(k4_xcmplx_0(np__1)))).# trainpos
cnf(c_0_257, plain, (r1_int_1(X1,np__1)|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v1_int_1(X1))).# trainpos
cnf(c_0_189, plain, (r1_int_1(X1,np__2)|~r1_int_1(X1,k4_xcmplx_0(np__2))|~v1_int_1(X1))).# trainpos
cnf(c_0_151, plain, (k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))=np__1|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(np__1)))).# trainpos
cnf(c_0_329, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_247, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_365, plain, (k3_int_2(X2,k4_xcmplx_0(X1))=k3_int_2(X1,X2)|k1_int_2(X1)=X1|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v1_int_1(X2)|~v1_int_1(X1))).# trainpos
cnf(c_0_338, negated_conjecture, (~r1_xxreal_0(esk2_0,np__1))).# trainpos
cnf(c_0_185, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(np__1)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_208, plain, (r1_int_1(k4_xcmplx_0(X1),np__2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,np__2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_340, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_269, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_375, plain, (k3_int_2(k4_xcmplx_0(X2),X1)=k3_int_2(X2,X1)|k1_int_2(X2)=X2|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(X2))).# trainpos
cnf(c_0_219, negated_conjecture, (k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).# trainpos
cnf(c_0_361, plain, (k1_newton(k4_xcmplx_0(np__1),X2)=np__1|k4_xcmplx_0(np__1)!=np__1|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_227, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__2)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,np__2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_318, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_325, plain, (k3_int_2(np__1,k1_int_2(X1))=k3_int_2(k4_xcmplx_0(np__1),X1)|~v1_int_1(X1))).# trainpos
cnf(c_0_367, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__1|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).# trainpos
cnf(c_0_239, negated_conjecture, (r1_int_1(esk2_0,k6_xcmplx_0(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).# trainpos
cnf(c_0_267, plain, (r1_xxreal_0(X1,np__2)|~r1_int_1(k4_xcmplx_0(X1),np__2)|~v7_ordinal1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).# trainpos
cnf(c_0_388, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_305, negated_conjecture, (~r1_int_1(k4_xcmplx_0(esk2_0),np__2))).# trainpos
cnf(c_0_390, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_371, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).# trainpos
cnf(c_0_262, negated_conjecture, (r1_int_1(esk2_0,k6_xcmplx_0(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).# trainpos
cnf(c_0_376, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__1)).# trainpos
cnf(c_0_317, plain, (r1_int_1(k4_xcmplx_0(X1),k1_newton(X2,X3))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k1_newton(X2,X3))))|~v1_int_1(X1))).# trainpos
cnf(c_0_392, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_399, negated_conjecture, (~r1_int_1(esk2_0,np__2))).# trainpos
cnf(c_0_280, plain, (r1_int_1(k4_xcmplx_0(X1),np__1)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_336, plain, (k3_int_2(k4_xcmplx_0(np__1),X1)=np__1|~v7_ordinal1(k1_int_2(X1))|~v1_int_1(X1))).# trainpos
cnf(c_0_341, plain, (r1_int_1(k4_xcmplx_0(X1),k1_newton(X2,X3))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(X1)|~v1_xcmplx_0(X2)|~r1_int_1(X1,k1_newton(X2,X3))|~v7_ordinal1(X1)|~v1_xcmplx_0(k1_newton(X2,X3)))).# trainpos
cnf(c_0_286, negated_conjecture, (r1_int_1(esk2_0,k4_xcmplx_0(np__2))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).# trainpos
cnf(c_0_304, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r1_int_1(esk2_0,k4_xcmplx_0(np__2))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1)).# trainpos
cnf(c_0_358, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_306, plain, (r1_int_1(k4_xcmplx_0(X1),np__2)|~r1_int_1(X1,k4_xcmplx_0(np__2))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_322, negated_conjecture, (~r1_int_1(esk2_0,k4_xcmplx_0(np__2)))).# trainpos
cnf(c_0_299, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__1)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_328, plain, (r1_xxreal_0(X1,np__1)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v7_ordinal1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).# trainpos
cnf(c_0_348, negated_conjecture, (~r1_int_1(k4_xcmplx_0(esk2_0),k4_xcmplx_0(np__1)))).# trainpos
cnf(c_0_356, plain, (~r1_int_1(esk2_0,np__1))).# trainpos
cnf(c_0_364, plain, (~r1_int_1(esk2_0,k4_xcmplx_0(np__1)))).# trainpos
cnf(c_0_384, plain, (k3_int_2(k4_xcmplx_0(np__1),X1)=k3_int_2(np__1,X1)|k4_xcmplx_0(np__1)=np__1|~v7_ordinal1(X1))).# trainpos
cnf(c_0_321, negated_conjecture, (r2_int_1(np__0,k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)),esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r1_int_1(esk2_0,k4_xcmplx_0(np__2))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1)).# trainpos
cnf(c_0_345, negated_conjecture, (r1_int_1(esk2_0,k6_xcmplx_0(np__0,k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|~v1_int_1(esk2_0)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1)).# trainpos
cnf(c_0_362, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))=np__1|~v1_int_1(esk2_0)|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).# trainpos
cnf(c_0_368, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|~v1_int_1(esk2_0)|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).# trainpos
cnf(c_0_372, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|~v1_int_1(esk2_0)|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).# trainpos
cnf(c_0_377, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|~v1_int_1(esk2_0))).# trainpos
cnf(c_0_381, negated_conjecture, (k4_xcmplx_0(np__1)!=np__1|~v1_int_1(esk2_0))).# trainpos
cnf(c_0_385, negated_conjecture, (k4_xcmplx_0(np__1)!=np__1)).# trainpos
cnf(c_0_374, plain, (r2_int_1(k2_int_5(k4_xcmplx_0(np__1),X1),k4_xcmplx_0(np__1),X1)|r1_xxreal_0(X1,np__2)|k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(X1,np__1),np__2))=np__1|k3_int_2(k4_xcmplx_0(np__1),X1)!=np__1|~v1_int_2(X1)|~v7_ordinal1(X1)|~m1_subset_1(k3_nat_d(k7_nat_d(X1,np__1),np__2),k5_numbers))).# trainpos
cnf(c_0_379, plain, (k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(X1,np__1),np__2))=np__1|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),X1),k4_xcmplx_0(np__1),X1)|r1_xxreal_0(X1,np__2)|k3_int_2(k4_xcmplx_0(np__1),X1)!=np__1|~v1_int_2(X1)|~v7_ordinal1(X1)|~v7_ordinal1(k7_nat_d(X1,np__1)))).# trainpos
cnf(c_0_383, plain, (k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(X1,np__1),np__2))=np__1|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),X1),k4_xcmplx_0(np__1),X1)|r1_xxreal_0(X1,np__2)|k3_int_2(k4_xcmplx_0(np__1),X1)!=np__1|~v1_int_2(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_387, plain, (k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(X1,np__1),np__2))=np__1|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),X1),k4_xcmplx_0(np__1),X1)|r1_xxreal_0(X1,np__2)|k3_int_2(np__1,X1)!=np__1|~v1_int_2(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_391, negated_conjecture, (k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))=np__1|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),k4_xcmplx_0(np__1),esk2_0))).# trainpos
cnf(c_0_393, negated_conjecture, (r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),k4_xcmplx_0(np__1),esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=np__1)).# trainpos
cnf(c_0_397, negated_conjecture, (r2_int_1(np__1,k4_xcmplx_0(np__1),esk2_0)|~v1_int_1(esk2_0))).# trainpos
cnf(c_0_400, negated_conjecture, (~v1_int_1(esk2_0))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_68, plain, (k4_xcmplx_0(np__0)=np__0)).#trainneg
cnf(i_0_49, plain, (v1_xcmplx_0(esk1_0))).#trainneg
cnf(i_0_47, plain, (v1_xreal_0(esk1_0))).#trainneg
cnf(i_0_50, plain, (v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_77, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_18, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_48, plain, (v1_xxreal_0(esk1_0))).#trainneg
cnf(i_0_42, plain, (~v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_41, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_81, plain, (v2_xxreal_0(np__2))).#trainneg
cnf(i_0_71, plain, (m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_76, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_73, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_80, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_67, plain, (k7_xcmplx_0(np__1,k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_16, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_28, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_39, plain, (v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_35, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_38, plain, (v1_xreal_0(k5_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_17, plain, (v7_ordinal1(X1)|~v1_finset_1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_8, plain, (~v7_ordinal1(X1)|~v1_int_2(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_9, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_90, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_99, plain, (k6_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_45, plain, (k18_complex1(k18_complex1(X1))=k18_complex1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_82, plain, (k7_xcmplx_0(np__1,X1)=k5_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_59, plain, (r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_nat_d(X1,X2))).#trainneg
cnf(i_0_21, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_43, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_37, plain, (v1_int_1(k6_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_22, plain, (k3_int_2(X1,X2)=k3_int_2(X2,X1)|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_20, plain, (k1_nat_1(X1,X2)=k1_nat_1(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_54, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_83, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_106, plain, (r1_int_1(X1,X2)|~epred1_2(X2,X1)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_87, plain, (r2_int_1(X2,X3,X1)|~v1_int_1(X3)|~v1_int_1(X2)|~v1_int_1(X1)|~r1_int_1(X1,k6_xcmplx_0(X2,X3)))).#trainneg
cnf(i_0_86, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(k1_nat_1(X1,np__1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_300, negated_conjecture, (~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_401, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_192, plain, (v7_ordinal1(esk1_0))).#trainneg
cnf(i_0_193, plain, (v7_ordinal1(np__0))).#trainneg
cnf(i_0_328, plain, (v1_xreal_0(np__0))).#trainneg
cnf(i_0_286, plain, (v1_xcmplx_0(np__0)|~v1_int_1(np__0))).#trainneg
cnf(i_0_318, plain, (v1_xcmplx_0(np__2)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_319, plain, (v1_xcmplx_0(np__2)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_295, plain, (v1_int_1(np__1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_374, plain, (k3_int_2(X1,np__1)=np__1|~v7_ordinal1(X1)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_287, plain, (v1_xcmplx_0(np__1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_296, plain, (k4_xcmplx_0(k4_xcmplx_0(esk1_0))=esk1_0)).#trainneg
cnf(i_0_322, plain, (v1_xcmplx_0(np__1)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_323, plain, (v1_xcmplx_0(np__1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_320, plain, (v1_xcmplx_0(np__0)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_321, plain, (v1_xcmplx_0(np__0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_406, plain, (r1_xxreal_0(k1_nat_1(np__1,X1),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_367, plain, (v7_ordinal1(k3_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_314, plain, (k18_complex1(X2)=X2|k4_xcmplx_0(X2)!=X2|~v1_xreal_0(X2))).#trainneg
cnf(i_0_382, plain, (k1_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_361, plain, (v1_int_1(np__2)|~v1_int_1(k4_xcmplx_0(np__1))|~v1_int_1(np__1))).#trainneg
cnf(i_0_362, plain, (v1_int_1(k4_xcmplx_0(np__2))|~v1_int_1(np__1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_282, plain, (v1_xcmplx_0(np__1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_333, plain, (v1_int_1(np__2)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_337, plain, (v1_int_1(np__1)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_283, plain, (v1_xcmplx_0(k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_335, plain, (v1_int_1(np__0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_313, plain, (k5_xcmplx_0(k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_407, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__1),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_312, plain, (k18_complex1(np__1)=np__1|~v1_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_289, plain, (v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_388, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_304, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_298, plain, (k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))=k4_xcmplx_0(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_299, plain, (k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_380, plain, (~r1_xxreal_0(k1_nat_1(np__1,X1),X2)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_305, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_324, plain, (v1_xcmplx_0(k1_int_2(X1))|~v1_membered(k5_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_331, plain, (v1_xreal_0(k1_int_2(X1))|~v3_membered(k5_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_317, plain, (k18_complex1(k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|k18_complex1(X1)=X1|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_355, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_399, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(X1))=k6_xcmplx_0(X1,np__0)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_386, plain, (r1_xxreal_0(X1,esk1_0)|r1_xxreal_0(esk1_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_293, plain, (v1_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_297, plain, (k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1)))=k5_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_364, plain, (v1_xreal_0(k3_nat_d(X1,X2))|~v3_membered(k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_397, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),np__0)=k6_xcmplx_0(np__0,X1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_365, plain, (v1_xcmplx_0(k3_nat_d(X1,X2))|~v1_membered(k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_366, plain, (v1_int_1(k3_nat_d(X1,X2))|~v5_membered(k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_369, plain, (v1_xreal_0(k7_nat_d(X1,X2))|~v3_membered(k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_370, plain, (v1_xcmplx_0(k7_nat_d(X1,X2))|~v1_membered(k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_390, plain, (k3_int_2(k1_int_2(X2),k1_int_2(X1))=k3_int_2(X1,X2)|~v1_int_1(k1_int_2(X2))|~v1_int_1(k1_int_2(X1))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_392, plain, (k3_int_2(k1_int_2(X1),k1_int_2(X2))=k3_int_2(X1,k1_int_2(X2))|~v1_int_1(k1_int_2(X2))|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_371, plain, (v1_int_1(k7_nat_d(X1,X2))|~v5_membered(k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_426, plain, (v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_57, plain, (k1_xboole_0=np__0)).#trainneg
cnf(i_0_428, plain, (v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_430, plain, (v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_479, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_102, plain, (X1=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_463, plain, (v1_int_1(esk1_0)|~v1_int_1(k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_625, plain, (r1_xxreal_0(esk1_0,esk1_0))).#trainneg
cnf(i_0_524, plain, (r1_xxreal_0(k2_xcmplx_0(np__1,X1),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_575, plain, (~r1_xxreal_0(k2_xcmplx_0(np__1,X1),X2)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_537, plain, (v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(np__2))))).#trainneg
cnf(i_0_477, plain, (r1_xxreal_0(X1,esk1_0)|r1_xxreal_0(esk1_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_90, plain, (k2_xcmplx_0(X1,np__0)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_612, plain, (k18_complex1(np__0)=np__0)).#trainneg
cnf(i_0_532, plain, (k18_complex1(np__1)=np__1|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_613, plain, (k18_complex1(esk1_0)=esk1_0|k18_complex1(k4_xcmplx_0(esk1_0))=k4_xcmplx_0(esk1_0)|~v1_xreal_0(k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_543, plain, (v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_544, plain, (v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_545, plain, (v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_449, plain, (r1_int_1(np__1,k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__1),X1)|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_424, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~r1_int_1(np__1,X1)|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_563, plain, (k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))=k4_xcmplx_0(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_523, plain, (r1_xxreal_0(np__2,X1)|r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_99, plain, (k6_xcmplx_0(X1,np__0)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_634, plain, (v1_int_1(k5_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_458, plain, (k1_int_2(k1_newton(k4_xcmplx_0(np__1),X1))=np__1|~m1_subset_1(X1,k5_numbers)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_548, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_455, plain, (r1_int_1(np__1,X1)|~r1_int_1(k4_xcmplx_0(np__1),X1)|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_486, plain, (k1_int_2(X1)=X1|~v1_int_1(X1)|k4_xcmplx_0(X1)!=X1|~v1_xreal_0(X1))).#trainneg
cnf(i_0_413, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~r1_int_1(np__1,k4_xcmplx_0(X1))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_487, plain, (k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))=k4_xcmplx_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_410, plain, (r1_int_1(X1,np__0)|~r1_int_1(k4_xcmplx_0(X1),np__0)|~v1_int_1(np__0)|~v1_int_1(X1))).#trainneg
cnf(i_0_546, plain, (r1_xxreal_0(X1,esk1_0)|r1_xxreal_0(esk1_0,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_412, plain, (r1_int_1(np__0,X1)|~r1_int_1(np__0,k4_xcmplx_0(X1))|~v1_int_1(X1)|~v1_int_1(np__0))).#trainneg
cnf(i_0_422, plain, (r1_int_1(X1,np__1)|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(np__1))|~v1_int_1(X1))).#trainneg
cnf(i_0_473, plain, (r1_int_1(X1,esk1_0)|~r1_int_1(X1,k4_xcmplx_0(esk1_0))|~v1_int_1(k4_xcmplx_0(esk1_0))|~v1_int_1(X1))).#trainneg
cnf(i_0_521, plain, (r1_int_1(esk1_0,X1)|~r1_int_1(k4_xcmplx_0(esk1_0),k4_xcmplx_0(X1))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_495, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_444, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~r1_int_1(X1,X2)|~v1_int_1(X2))).#trainneg
cnf(i_0_522, plain, (r1_int_1(np__1,X1)|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_434, plain, (r1_int_1(np__0,k4_xcmplx_0(X1))|~r1_int_1(np__0,X1)|~v1_int_1(X1)|~v1_int_1(np__0))).#trainneg
cnf(i_0_517, plain, (r1_int_1(X1,X2)|~v1_int_1(X2)|~v1_int_1(X1)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_435, plain, (r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))|~r1_int_1(np__1,X1)|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_565, plain, (k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))=k4_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_411, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(X1),np__1)|~v1_int_1(k4_xcmplx_0(np__1))|~v1_int_1(X1))).#trainneg
cnf(i_0_460, plain, (k18_complex1(k1_int_2(X1))=k1_int_2(X1)|~v1_xcmplx_0(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_560, plain, (k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))=k4_xcmplx_0(k5_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_446, plain, (r1_int_1(k4_xcmplx_0(X1),np__0)|~r1_int_1(X1,np__0)|~v1_int_1(np__0)|~v1_int_1(X1))).#trainneg
cnf(i_0_415, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~r1_int_1(X1,np__1)|~v1_int_1(k4_xcmplx_0(np__1))|~v1_int_1(X1))).#trainneg
cnf(i_0_466, plain, (r1_int_1(k4_xcmplx_0(esk1_0),X1)|~r1_int_1(esk1_0,X1)|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_472, plain, (r1_int_1(X1,k4_xcmplx_0(esk1_0))|~r1_int_1(X1,esk1_0)|~v1_int_1(k4_xcmplx_0(esk1_0))|~v1_int_1(X1))).#trainneg
cnf(i_0_447, plain, (r1_int_1(k4_xcmplx_0(X1),np__1)|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(np__1))|~v1_int_1(X1))).#trainneg
cnf(i_0_436, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_588, plain, (k1_int_2(k1_newton(k4_xcmplx_0(np__1),X1))=k1_newton(k4_xcmplx_0(np__1),X1)|k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))=np__1|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_xreal_0(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_484, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~m1_subset_1(X1,k5_numbers)|k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))!=k1_newton(k4_xcmplx_0(np__1),X1)|~v1_xreal_0(k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_676, plain, (k3_int_2(k1_int_2(X1),X2)=k3_int_2(X2,X1)|~v1_int_1(k1_int_2(X1))|~v1_int_1(k1_int_2(X2))|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_468, plain, (r1_int_1(esk1_0,k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(esk1_0),X1)|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_457, plain, (k1_int_2(k18_complex1(X1))=k18_complex1(X1)|~v1_xcmplx_0(X1)|~v1_int_1(k18_complex1(X1)))).#trainneg
cnf(i_0_610, plain, (k1_int_2(k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|k18_complex1(X1)=X1|~v1_int_1(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_475, plain, (r1_int_1(X1,k4_xcmplx_0(esk1_0))|~r1_int_1(k4_xcmplx_0(X1),esk1_0)|~v1_int_1(k4_xcmplx_0(esk1_0))|~v1_int_1(X1))).#trainneg
cnf(i_0_469, plain, (r1_int_1(k4_xcmplx_0(esk1_0),X1)|~r1_int_1(esk1_0,k4_xcmplx_0(X1))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_474, plain, (r1_int_1(k4_xcmplx_0(X1),esk1_0)|~r1_int_1(X1,k4_xcmplx_0(esk1_0))|~v1_int_1(k4_xcmplx_0(esk1_0))|~v1_int_1(X1))).#trainneg
cnf(i_0_470, plain, (r1_int_1(k4_xcmplx_0(esk1_0),k4_xcmplx_0(X1))|~r1_int_1(esk1_0,X1)|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_592, plain, (v1_xreal_0(k1_newton(X1,X2))|~v1_xreal_0(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_606, plain, (r1_int_1(k4_xcmplx_0(X1),k1_newton(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(X1)|~v7_ordinal1(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_516, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1))).#trainneg
cnf(i_0_616, plain, (k4_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))=k7_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_637, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),X2)|~r1_int_1(k5_xcmplx_0(X1),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_639, plain, (r1_int_1(k5_xcmplx_0(X1),X2)|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_647, plain, (r1_int_1(X1,k4_xcmplx_0(k5_xcmplx_0(X2)))|~r1_int_1(X1,k5_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_648, plain, (r1_int_1(X1,k5_xcmplx_0(X2))|~r1_int_1(X1,k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_641, plain, (r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_643, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),X2)|~r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_600, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(X3))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_645, plain, (r1_int_1(k5_xcmplx_0(X1),X2)|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_585, plain, (r1_int_1(X1,k6_xcmplx_0(k2_int_5(X2,k4_xcmplx_0(X1)),k1_newton(X2,k3_nat_d(k7_nat_d(k4_xcmplx_0(X1),np__1),np__2))))|r1_xxreal_0(k4_xcmplx_0(X1),np__2)|~v1_int_1(k6_xcmplx_0(k2_int_5(X2,k4_xcmplx_0(X1)),k1_newton(X2,k3_nat_d(k7_nat_d(k4_xcmplx_0(X1),np__1),np__2))))|~v1_int_1(X1)|k3_int_2(X2,k4_xcmplx_0(X1))!=np__1|~v1_int_2(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(k1_newton(X2,k3_nat_d(k7_nat_d(k4_xcmplx_0(X1),np__1),np__2)))|~v1_int_1(X2))).#trainneg
cnf(i_0_650, plain, (r1_int_1(X1,k4_xcmplx_0(k5_xcmplx_0(X2)))|~r1_int_1(k4_xcmplx_0(X1),k5_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_649, plain, (r1_int_1(k4_xcmplx_0(X1),k5_xcmplx_0(X2))|~r1_int_1(X1,k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_644, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k4_xcmplx_0(X2))|~r1_int_1(k5_xcmplx_0(X1),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_602, plain, (r1_int_1(k1_newton(X1,X2),X3)|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k4_xcmplx_0(X3))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_607, plain, (r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(k4_xcmplx_0(X1),k1_newton(X2,X3))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(X1)|~v7_ordinal1(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_595, plain, (k6_xcmplx_0(k1_newton(X1,X2),k4_xcmplx_0(X3))=k6_xcmplx_0(X3,k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_603, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k1_newton(X2,X3))=k6_xcmplx_0(k4_xcmplx_0(k1_newton(X2,X3)),X1)|~v1_xcmplx_0(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_399, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(X1))=k6_xcmplx_0(X1,np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_747, plain, (esk1_0=np__0)).#trainneg
cnf(i_0_779, plain, (k1_int_2(np__1)=np__1)).#trainneg
cnf(i_0_861, plain, (r1_xxreal_0(np__0,np__1))).#trainneg
cnf(i_0_777, plain, (k1_int_2(np__0)=np__0)).#trainneg
cnf(i_0_1061, plain, (r1_int_1(np__1,np__0)|~r1_int_1(k4_xcmplx_0(np__1),np__0)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1082, plain, (r1_int_1(np__0,np__1)|~r1_int_1(np__0,k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_313, plain, (k5_xcmplx_0(k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_922, plain, (v7_ordinal1(X1)|~v1_int_1(X1)|k4_xcmplx_0(X1)!=X1)).#trainneg
cnf(i_0_397, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),np__0)=k6_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1280, plain, (r1_int_1(np__1,np__1)|~v1_int_1(k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(np__1),np__1))).#trainneg
cnf(i_0_1306, plain, (r1_int_1(np__1,np__1)|~v1_int_1(k4_xcmplx_0(np__1))|~r1_int_1(np__1,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1312, plain, (r1_int_1(np__1,np__1)|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_727, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_696, plain, (k4_xcmplx_0(k4_xcmplx_0(k2_int_5(X1,X2)))=k2_int_5(X1,X2)|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_1133, plain, (r1_int_1(X1,X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_702, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1889, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_730, plain, (r1_int_1(X1,X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_739, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|r1_xxreal_0(esk2_0,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_690, plain, (k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))=k6_xcmplx_0(X1,X2)|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_970, plain, (r1_int_1(k4_xcmplx_0(np__1),np__0)|~r1_int_1(np__1,np__0)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_995, plain, (r1_int_1(np__0,k4_xcmplx_0(np__1))|~r1_int_1(np__0,np__1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1002, plain, (r1_int_1(k4_xcmplx_0(np__1),np__1)|~v1_int_1(k4_xcmplx_0(np__1))|~r1_int_1(np__1,np__1))).#trainneg
cnf(i_0_1091, plain, (k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))=np__1|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(k1_newton(k4_xcmplx_0(np__1),X1),k1_numbers))).#trainneg
cnf(i_0_882, plain, (k1_int_2(k4_xcmplx_0(np__1))=np__1|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_722, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_703, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_691, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2)))=k1_newton(X1,X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_858, negated_conjecture, (r1_xxreal_0(esk2_0,np__1)|r1_xxreal_0(np__2,esk2_0))).#trainneg
cnf(i_0_860, plain, (r1_xxreal_0(np__2,np__1)|r1_xxreal_0(np__2,np__2))).#trainneg
cnf(i_0_862, plain, (r1_xxreal_0(np__1,np__1)|r1_xxreal_0(np__2,np__1))).#trainneg
cnf(i_0_1212, plain, (r1_int_1(np__1,k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(np__1))|~r1_int_1(np__1,np__1))).#trainneg
cnf(i_0_1063, plain, (r1_int_1(np__1,k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(np__1),np__1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1127, plain, (r1_int_1(k4_xcmplx_0(np__1),np__1)|~r1_int_1(np__1,k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1204, plain, (r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))|~r1_int_1(np__1,np__1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_920, plain, (m1_subset_1(X1,k5_numbers)|~v1_int_1(X1)|k4_xcmplx_0(X1)!=X1)).#trainneg
cnf(i_0_772, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_773, plain, (~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_910, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_704, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1089, plain, (k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))=np__1|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_720, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_741, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_887, plain, (m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|k1_int_2(X1)=X1|~v1_int_1(X1))).#trainneg
cnf(i_0_724, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1922, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_726, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_707, plain, (r1_xxreal_0(X1,k1_int_2(X2))|r1_xxreal_0(k1_int_2(X2),X1)|~v7_ordinal1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_1220, plain, (k18_complex1(X1)=X1|~v1_int_1(X1)|k4_xcmplx_0(X1)!=X1)).#trainneg
cnf(i_0_774, plain, (k2_xcmplx_0(np__0,X1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1537, plain, (v1_xreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1430, plain, (v7_ordinal1(k18_complex1(X1))|~v1_int_1(k18_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_867, plain, (v1_int_1(k5_xcmplx_0(X1))|~v1_xreal_0(X1)|~v7_ordinal1(k4_xcmplx_0(k5_xcmplx_0(X1))))).#trainneg
cnf(i_0_732, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1386, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))!=k1_newton(k4_xcmplx_0(np__1),X1)|~m1_subset_1(X1,k5_numbers)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_1145, plain, (k3_int_2(X1,k1_int_2(X2))=k3_int_2(X1,X2)|~v1_int_1(X2)|~v1_int_1(X1)|~v1_int_1(k1_int_2(X2)))).#trainneg
cnf(i_0_1572, plain, (r1_int_1(k4_xcmplx_0(X1),np__1)|~r1_int_1(X1,np__1)|~v1_int_1(X1))).#trainneg
cnf(i_0_719, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_743, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_937, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_938, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1973, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1132, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~r1_int_1(np__1,k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_816, plain, (r1_int_1(np__1,k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(np__1))|~v1_int_1(X1)|~r1_int_1(np__1,X1))).#trainneg
cnf(i_0_1385, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))!=k1_newton(k4_xcmplx_0(np__1),X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_723, plain, (k18_complex1(X1)=X1|k18_complex1(k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_735, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_int_2(X1)))=k1_int_2(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_1142, plain, (k3_int_2(X1,np__1)=np__1|~v7_ordinal1(k1_int_2(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_1363, plain, (r1_int_1(X1,np__1)|~r1_int_1(k4_xcmplx_0(X1),np__1)|~v1_int_1(X1))).#trainneg
cnf(i_0_916, plain, (r1_int_1(np__1,X1)|~v1_int_1(k4_xcmplx_0(np__1))|~v1_int_1(X1)|~r1_int_1(np__1,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1068, plain, (r1_int_1(np__1,k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__1),X1)|~v1_int_1(k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_888, plain, (k3_int_2(k4_xcmplx_0(X1),k1_int_2(X2))=k3_int_2(X1,X2)|k1_int_2(X1)=X1|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_725, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_546, plain, (r1_xxreal_0(np__0,X1)|r1_xxreal_0(X1,np__0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_48, plain, (v1_xxreal_0(np__0))).#trainneg
cnf(i_0_2064, negated_conjecture, (r1_xxreal_0(esk2_0,esk2_0))).#trainneg
cnf(i_0_625, plain, (r1_xxreal_0(np__0,np__0))).#trainneg
cnf(i_0_2329, plain, (r1_xxreal_0(np__2,np__2))).#trainneg
cnf(i_0_2330, plain, (r1_xxreal_0(np__0,np__2))).#trainneg
cnf(i_0_3220, plain, (v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_2062, plain, (r1_int_1(X1,np__2)|~r1_int_1(k4_xcmplx_0(X1),np__2)|~v1_int_1(X1))).#trainneg
cnf(i_0_2992, plain, (m1_subset_1(X1,k5_numbers)|m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_2815, plain, (r1_xxreal_0(np__1,np__2))).#trainneg
cnf(i_0_2644, plain, (v7_ordinal1(X1)|v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_3219, plain, (v1_xreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_2817, plain, (r1_xxreal_0(np__1,np__1))).#trainneg
cnf(i_0_2642, plain, (m1_subset_1(X1,k5_numbers)|v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_3373, plain, (k3_int_2(X1,k4_xcmplx_0(X2))=k3_int_2(X1,X2)|k1_int_2(X2)=X2|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_3919, plain, (r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(X1)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2073, plain, (k6_xcmplx_0(np__1,np__0)=np__1)).#trainneg
cnf(i_0_2021, plain, (r1_int_1(np__0,np__2)|~r1_int_1(np__0,k4_xcmplx_0(np__2))|~v1_int_1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_2025, plain, (r1_int_1(np__2,np__0)|~r1_int_1(k4_xcmplx_0(np__2),np__0)|~v1_int_1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_3393, plain, (r1_int_1(np__2,np__1)|~r1_int_1(k4_xcmplx_0(np__2),np__1)|~v1_int_1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_3394, plain, (r1_int_1(np__1,np__1)|~r1_int_1(k4_xcmplx_0(np__1),np__1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_2060, plain, (r1_int_1(k4_xcmplx_0(X1),np__2)|~r1_int_1(X1,np__2)|~v1_int_1(X1))).#trainneg
cnf(i_0_3233, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_3745, plain, (r1_int_1(k4_xcmplx_0(np__2),np__1)|~r1_int_1(np__2,np__1)|~v1_int_1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_2023, plain, (r1_int_1(np__0,k4_xcmplx_0(np__2))|~r1_int_1(np__0,np__2)|~v1_int_1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_2024, plain, (r1_int_1(k4_xcmplx_0(np__2),np__0)|~r1_int_1(np__2,np__0)|~v1_int_1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_3746, plain, (r1_int_1(k4_xcmplx_0(np__1),np__1)|~r1_int_1(np__1,np__1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_2067, negated_conjecture, (r1_xxreal_0(esk2_0,np__0)|r1_xxreal_0(np__0,esk2_0))).#trainneg
cnf(i_0_2068, negated_conjecture, (r1_xxreal_0(esk2_0,np__1)|r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_2814, negated_conjecture, (r1_xxreal_0(esk2_0,np__0)|r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_2962, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(X1)|~r1_int_1(X1,X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3740, plain, (k3_int_2(np__1,X1)=np__1|~v1_int_1(X1))).#trainneg
cnf(i_0_2076, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),np__0)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_2075, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),np__0)=k6_xcmplx_0(np__0,np__2))).#trainneg
cnf(i_0_2148, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(np__2))=k6_xcmplx_0(np__2,np__0))).#trainneg
cnf(i_0_2016, plain, (k18_complex1(np__2)=np__2|k18_complex1(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2)|~v1_xreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_2039, plain, (k1_int_2(np__2)=np__2|k18_complex1(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2)|~v1_xreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_3232, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_3184, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(X1)|~r1_int_1(X1,X2))).#trainneg
cnf(i_0_2299, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~r1_int_1(X1,np__1)|~v1_int_1(X1))).#trainneg
cnf(i_0_2298, plain, (r1_int_1(X1,k4_xcmplx_0(np__2))|~r1_int_1(X1,np__2)|~v1_int_1(X1))).#trainneg
cnf(i_0_2494, negated_conjecture, (r1_xxreal_0(esk2_0,X1)|r1_xxreal_0(X1,esk2_0)|~v1_int_1(X1))).#trainneg
cnf(i_0_2140, plain, (v7_ordinal1(k4_xcmplx_0(X1))|k4_xcmplx_0(k4_xcmplx_0(X1))!=k4_xcmplx_0(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_3234, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_2192, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2994, plain, (v7_ordinal1(X1)|m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_2823, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~r1_int_1(np__1,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_2648, plain, (k18_complex1(X1)=X1|v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_3361, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))!=k1_newton(k4_xcmplx_0(np__1),X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_2822, plain, (r1_int_1(k4_xcmplx_0(np__2),X1)|~r1_int_1(np__2,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_2495, negated_conjecture, (r1_xxreal_0(esk2_0,X1)|r1_xxreal_0(X1,esk2_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_386, plain, (r1_xxreal_0(np__0,X1)|r1_xxreal_0(X1,np__0)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_3187, plain, (r1_int_1(np__2,X1)|~r1_int_1(k4_xcmplx_0(np__2),X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_3188, plain, (r1_int_1(np__1,X1)|~r1_int_1(k4_xcmplx_0(np__1),X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_3172, plain, (r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(X1)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_2128, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~r1_int_1(np__1,k4_xcmplx_0(X1))|~v1_int_1(X1)|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_2348, plain, (r1_int_1(k4_xcmplx_0(np__2),X1)|~r1_int_1(np__2,k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_728, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),X2)=k6_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3487, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_3488, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_3518, plain, (v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_3638, plain, (v1_xreal_0(k1_int_2(X1))|~v1_xreal_0(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_3842, plain, (k3_int_2(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_int_2(X1,X2)|k1_int_2(X1)=X1|k1_int_2(X2)=X2|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_731, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2987, plain, (r1_xxreal_0(np__2,X1)|r1_xxreal_0(X1,np__2)|~v1_int_1(X1))).#trainneg
cnf(i_0_3458, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__1)|~v1_int_1(X1))).#trainneg
cnf(i_0_3640, plain, (v1_xcmplx_0(k1_int_2(X1))|~v1_xcmplx_0(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_3489, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k1_numbers))).#trainneg
cnf(i_0_3267, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2674, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3347, plain, (r1_int_1(np__1,k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__1),X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_1428, plain, (m1_subset_1(k18_complex1(X1),k5_numbers)|~v1_int_1(k18_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3906, plain, (r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(X1)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_2998, plain, (k18_complex1(X1)=X1|m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_2675, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2988, plain, (r1_xxreal_0(np__2,X1)|r1_xxreal_0(X1,np__2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_746, plain, (r1_xxreal_0(X1,k1_int_2(X2))|r1_xxreal_0(k1_int_2(X2),X1)|~v1_xreal_0(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_2020, plain, (r1_int_1(np__2,X1)|~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(X1))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_2072, plain, (v1_int_1(k6_xcmplx_0(X1,np__0))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2145, plain, (v1_int_1(k6_xcmplx_0(np__0,X1))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_911, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3459, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3346, plain, (r1_int_1(np__2,k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__2),X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_824, plain, (v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v1_int_1(X1))).#trainneg
cnf(i_0_3061, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2003, negated_conjecture, (r2_int_1(k4_xcmplx_0(np__1),k1_newton(X1,k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)),esk2_0)|k2_int_5(X1,esk2_0)=np__1|k2_int_5(X1,esk2_0)=np__0|k3_int_2(X1,esk2_0)!=np__1|~v1_int_1(X1))).#trainneg
cnf(i_0_3244, plain, (v7_ordinal1(k4_xcmplx_0(X1))|k18_complex1(X1)=X1|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2087, plain, (k3_int_2(np__1,k1_int_2(X1))=k3_int_2(np__1,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_825, plain, (v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v1_int_1(X1))).#trainneg
cnf(i_0_1151, plain, (v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1152, plain, (v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2222, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(X2))).#trainneg
cnf(i_0_3864, plain, (k3_int_2(X2,k4_xcmplx_0(X1))=k3_int_2(X1,X2)|k1_int_2(X1)=X1|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_3547, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2362, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(X2))).#trainneg
cnf(i_0_3447, plain, (k6_xcmplx_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))=k6_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2047, plain, (r1_int_1(k4_xcmplx_0(X1),np__2)|~r1_int_1(X1,k4_xcmplx_0(np__2))|~v1_int_1(k4_xcmplx_0(np__2))|~v1_int_1(X1))).#trainneg
cnf(i_0_2095, plain, (k3_int_2(k1_int_2(X1),np__1)=k3_int_2(X1,np__1)|~v1_int_1(X1))).#trainneg
cnf(i_0_3434, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(X1))=k6_xcmplx_0(X1,np__0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1222, plain, (v1_xreal_0(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3961, plain, (~r1_xxreal_0(np__2,np__1))).#trainneg
cnf(i_0_4411, plain, (r1_int_1(X1,np__2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),np__2))).#trainneg
cnf(i_0_4156, plain, (r1_int_1(X1,np__2)|~r1_int_1(k4_xcmplx_0(X1),np__2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1312, plain, (r1_int_1(np__1,np__1)|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_3940, plain, (v7_ordinal1(k4_xcmplx_0(np__1))|k4_xcmplx_0(np__1)!=np__1)).#trainneg
cnf(i_0_1306, plain, (r1_int_1(np__1,np__1)|~r1_int_1(np__1,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1082, plain, (r1_int_1(np__0,np__1)|~r1_int_1(np__0,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1061, plain, (r1_int_1(np__1,np__0)|~r1_int_1(k4_xcmplx_0(np__1),np__0))).#trainneg
cnf(i_0_3944, plain, (v7_ordinal1(k4_xcmplx_0(np__2))|k4_xcmplx_0(np__2)!=np__2)).#trainneg
cnf(i_0_5180, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1212, plain, (r1_int_1(np__1,k4_xcmplx_0(np__1))|~r1_int_1(np__1,np__1))).#trainneg
cnf(i_0_1204, plain, (r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))|~r1_int_1(np__1,np__1))).#trainneg
cnf(i_0_4303, plain, (k6_xcmplx_0(np__0,np__2)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_4117, plain, (r1_int_1(np__2,np__1)|~r1_int_1(np__2,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_4143, plain, (r1_int_1(np__2,np__2)|~r1_int_1(k4_xcmplx_0(np__2),np__2))).#trainneg
cnf(i_0_1132, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~r1_int_1(np__1,k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_4317, plain, (v7_ordinal1(k4_xcmplx_0(np__2))|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_4304, plain, (v1_int_1(k6_xcmplx_0(np__2,np__0)))).#trainneg
cnf(i_0_4425, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|r1_xxreal_0(esk2_0,np__1))).#trainneg
cnf(i_0_4144, plain, (r1_int_1(np__1,np__2)|~r1_int_1(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_4722, plain, (r1_int_1(np__1,np__2)|~r1_int_1(np__1,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_4858, plain, (r1_int_1(np__2,np__2)|~r1_int_1(np__2,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_5266, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4291, plain, (r1_int_1(X1,X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4843, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~r1_int_1(X1,np__1)|~v7_ordinal1(X2)|~v1_int_1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4842, plain, (r1_int_1(X1,np__1)|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|~v7_ordinal1(X2)|~v1_int_1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4374, plain, (r1_int_1(X1,X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3972, plain, (r1_int_1(k4_xcmplx_0(np__2),np__2)|~r1_int_1(np__2,np__2))).#trainneg
cnf(i_0_1127, plain, (r1_int_1(k4_xcmplx_0(np__1),np__1)|~r1_int_1(np__1,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_5795, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~r1_int_1(X1,X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5281, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4516, plain, (v7_ordinal1(k4_xcmplx_0(X1))|m1_subset_1(X1,k5_numbers)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3973, plain, (r1_int_1(k4_xcmplx_0(np__1),np__2)|~r1_int_1(np__1,np__2))).#trainneg
cnf(i_0_1063, plain, (r1_int_1(np__1,k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(np__1),np__1))).#trainneg
cnf(i_0_6270, plain, (r1_int_1(np__2,np__2)|~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_6271, plain, (r1_int_1(np__1,np__2)|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_995, plain, (r1_int_1(np__0,k4_xcmplx_0(np__1))|~r1_int_1(np__0,np__1))).#trainneg
cnf(i_0_4572, plain, (r1_int_1(np__2,k4_xcmplx_0(X1))|~v1_int_1(X1)|~r1_int_1(np__2,X1))).#trainneg
cnf(i_0_4172, plain, (r1_int_1(np__2,k4_xcmplx_0(np__1))|~r1_int_1(np__2,np__1))).#trainneg
cnf(i_0_1068, plain, (r1_int_1(np__1,k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__1),X1)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_5894, plain, (k3_int_2(np__1,k4_xcmplx_0(np__1))=k3_int_2(np__1,np__1))).#trainneg
cnf(i_0_6292, plain, (k3_int_2(k4_xcmplx_0(np__1),np__1)=k3_int_2(np__1,np__1))).#trainneg
cnf(i_0_970, plain, (r1_int_1(k4_xcmplx_0(np__1),np__0)|~r1_int_1(np__1,np__0))).#trainneg
cnf(i_0_4596, plain, (r1_int_1(k4_xcmplx_0(X1),np__2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_916, plain, (r1_int_1(np__1,X1)|~r1_int_1(np__1,k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_4311, plain, (k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2)|k18_complex1(np__2)=np__2)).#trainneg
cnf(i_0_4807, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~r1_int_1(k4_xcmplx_0(X1),np__1)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~v1_int_1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_5811, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_4014, plain, (m1_subset_1(X1,k5_numbers)|m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4680, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_411, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(X1),np__1)|~v1_int_1(X1))).#trainneg
cnf(i_0_4495, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_816, plain, (r1_int_1(np__1,k4_xcmplx_0(X1))|~r1_int_1(np__1,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_4315, plain, (k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2)|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_5630, plain, (r1_int_1(np__2,k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(np__2),np__2))).#trainneg
cnf(i_0_5631, plain, (r1_int_1(np__2,k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(np__2),np__1))).#trainneg
cnf(i_0_4633, plain, (r1_int_1(np__2,X1)|~v1_int_1(X1)|~r1_int_1(np__2,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_4385, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),np__1))).#trainneg
cnf(i_0_3965, negated_conjecture, (r1_xxreal_0(esk2_0,k4_xcmplx_0(np__1))|r1_xxreal_0(k4_xcmplx_0(np__1),esk2_0))).#trainneg
cnf(i_0_522, plain, (r1_int_1(np__1,X1)|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_5420, plain, (m1_subset_1(k4_xcmplx_0(np__2),k5_numbers)|k18_complex1(np__2)=np__2)).#trainneg
cnf(i_0_5421, plain, (m1_subset_1(k4_xcmplx_0(np__2),k5_numbers)|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_4033, negated_conjecture, (r1_xxreal_0(esk2_0,k4_xcmplx_0(np__2))|r1_xxreal_0(k4_xcmplx_0(np__2),esk2_0))).#trainneg
cnf(i_0_4382, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(X1),np__1))).#trainneg
cnf(i_0_4450, plain, (v7_ordinal1(k4_xcmplx_0(X1))|k4_xcmplx_0(k4_xcmplx_0(X1))!=k4_xcmplx_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5257, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4383, plain, (r1_int_1(X1,np__1)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(X1),np__1))).#trainneg
cnf(i_0_4524, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(np__1,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_4386, plain, (r1_int_1(X1,np__1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),np__1))).#trainneg
cnf(i_0_4362, plain, (r1_int_1(X1,np__1)|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_447, plain, (r1_int_1(k4_xcmplx_0(X1),np__1)|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v1_int_1(X1))).#trainneg
cnf(i_0_4668, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~r1_int_1(X1,np__1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4279, plain, (r1_int_1(X1,np__1)|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_4810, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|~v1_int_1(X1)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4128, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__2)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,np__2)|~v1_int_1(X1))).#trainneg
cnf(i_0_4802, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~v1_int_1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_435, plain, (r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))|~r1_int_1(np__1,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_4956, plain, (v1_xreal_0(k1_int_2(X1))|~v7_ordinal1(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_4069, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~r1_int_1(X1,np__1)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_4221, plain, (k3_int_2(k4_xcmplx_0(np__1),X1)=k3_int_2(np__1,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_5778, plain, (v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))))).#trainneg
cnf(i_0_5025, plain, (v1_xreal_0(k1_int_2(X1))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(k1_int_2(X1))))).#trainneg
cnf(i_0_5223, plain, (v1_xcmplx_0(k1_int_2(X1))|~v1_int_1(X1)|~v7_ordinal1(k4_xcmplx_0(k1_int_2(X1))))).#trainneg
cnf(i_0_4289, plain, (r1_int_1(X1,k2_int_5(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_6145, plain, (r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__2),X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_3955, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_828, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))),X2)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v1_int_1(X1))).#trainneg
cnf(i_0_6665, negated_conjecture, (~r1_xxreal_0(esk2_0,np__0))).#trainneg
cnf(i_0_7273, plain, (k3_int_2(np__1,np__1)=np__1)).#trainneg
cnf(i_0_8413, plain, (r1_xxreal_0(X1,np__2)|~r1_int_1(k4_xcmplx_0(X1),np__2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2075, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),np__0)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_6645, plain, (k18_complex1(np__2)=np__2|v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_6760, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7025, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_xcmplx_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7771, plain, (r1_int_1(np__2,k4_xcmplx_0(X1))|~r1_int_1(np__2,X1)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_830, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v1_int_1(X1))).#trainneg
cnf(i_0_7956, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8431, plain, (r1_int_1(np__2,k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(np__2,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_832, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v1_int_1(X1))).#trainneg
cnf(i_0_7403, plain, (r1_int_1(np__1,k4_xcmplx_0(X1))|~r1_int_1(np__1,X1)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_833, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))),X2)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v1_int_1(X1))).#trainneg
cnf(i_0_8496, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_840, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_8433, plain, (r1_int_1(np__1,k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(np__1,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7298, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__2)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,np__2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_8009, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_841, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_4049, plain, (k3_int_2(k4_xcmplx_0(X2),X1)=k3_int_2(X1,X2)|k1_int_2(X2)=X2|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_3417, plain, (k6_xcmplx_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_842, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(X2)))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_4372, plain, (r1_int_1(X1,k2_int_5(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_xcmplx_0(X1)|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_1741, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_843, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_7220, plain, (r1_int_1(np__2,X1)|~r1_int_1(np__2,k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4497, plain, (v1_xreal_0(k4_xcmplx_0(X1))|v7_ordinal1(X1)|~v3_membered(k5_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_863, plain, (r1_xxreal_0(k7_nat_d(X1,X2),np__1)|r1_xxreal_0(np__2,k7_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6104, plain, (k6_xcmplx_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))=k6_xcmplx_0(X2,X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4745, plain, (r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(np__1,X2)|~v1_int_1(X2)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_864, plain, (r1_xxreal_0(k3_nat_d(X1,X2),np__1)|r1_xxreal_0(np__2,k3_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3524, plain, (k6_xcmplx_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2989, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|m1_subset_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)),k5_numbers)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2935, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_865, plain, (r1_xxreal_0(k1_int_2(X1),np__1)|r1_xxreal_0(np__2,k1_int_2(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_4678, plain, (r1_int_1(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|~r1_int_1(X1,k2_int_5(X2,X3))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_xcmplx_0(X1)|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_6105, plain, (k6_xcmplx_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8271, plain, (r1_int_1(np__2,np__1)|~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1349, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2))).#trainneg
cnf(i_0_1143, plain, (k3_int_2(X2,k1_int_2(X1))=k3_int_2(X1,X2)|~v1_int_1(X2)|~v1_int_1(k1_int_2(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_6917, plain, (r1_int_1(np__2,k4_xcmplx_0(np__2))|~r1_int_1(np__2,np__2))).#trainneg
cnf(i_0_2142, plain, (v7_ordinal1(k2_int_5(X1,X2))|k4_xcmplx_0(k2_int_5(X1,X2))!=k2_int_5(X1,X2)|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_1087, plain, (r1_int_1(np__0,X1)|~r1_int_1(np__0,k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2277, plain, (r1_int_1(X1,k2_int_5(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_1(X1)|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_1000, plain, (r1_int_1(np__0,k4_xcmplx_0(X1))|~r1_int_1(np__0,X1)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4834, plain, (r1_int_1(np__1,X2)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|~v7_ordinal1(X1)|~v1_int_1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_886, plain, (k1_int_2(k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|k1_int_2(X1)=X1|~v1_int_1(X1))).#trainneg
cnf(i_0_4283, plain, (r1_int_1(X1,k1_newton(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_4751, plain, (r1_int_1(np__1,k4_xcmplx_0(X2))|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|~v1_int_1(X2)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_895, plain, (k3_int_2(k1_int_2(X1),k4_xcmplx_0(X2))=k3_int_2(X1,X2)|k1_int_2(X2)=X2|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_1563, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(np__1))|~v1_int_1(X1))).#trainneg
cnf(i_0_4446, plain, (v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(X1)))|k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))!=k4_xcmplx_0(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_891, plain, (k3_int_2(k4_xcmplx_0(X1),k1_int_2(X2))=k3_int_2(X2,X1)|k1_int_2(X1)=X1|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k1_int_2(X2))|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_4079, plain, (r1_int_1(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|~r1_int_1(X1,k2_int_5(X2,X3))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_7713, plain, (r1_int_1(np__1,X1)|~r1_int_1(np__1,k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2335, plain, (r1_int_1(X1,np__0)|~r1_int_1(k4_xcmplx_0(X1),np__0)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4648, plain, (r1_int_1(np__1,k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__1),X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3025, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_896, plain, (k3_int_2(k1_int_2(X1),k4_xcmplx_0(X2))=k3_int_2(X2,X1)|k1_int_2(X2)=X2|~v1_int_1(k1_int_2(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_2067, negated_conjecture, (r1_xxreal_0(np__0,esk2_0))).#trainneg
cnf(i_0_9005, negated_conjecture, (~r1_int_1(k4_xcmplx_0(esk2_0),np__2)|~v1_int_1(k4_xcmplx_0(esk2_0)))).#trainneg
cnf(i_0_9004, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).#trainneg
cnf(i_0_10742, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11734, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_9887, plain, (r1_xxreal_0(X1,np__2)|~r1_int_1(k4_xcmplx_0(X1),np__2)|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_6292, plain, (k3_int_2(k4_xcmplx_0(np__1),np__1)=np__1)).#trainneg
cnf(i_0_5894, plain, (k3_int_2(np__1,k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_11194, plain, (r1_int_1(X1,X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_9036, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_9118, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_8997, negated_conjecture, (r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).#trainneg
cnf(i_0_7101, plain, (r1_int_1(np__1,k4_xcmplx_0(np__2))|~r1_int_1(np__1,np__2))).#trainneg
cnf(i_0_10756, plain, (r1_xxreal_0(np__2,X1)|r1_xxreal_0(X1,np__1)|~v1_int_1(X1)|k4_xcmplx_0(X1)!=X1)).#trainneg
cnf(i_0_3956, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_941, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10568, plain, (k3_int_2(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_10379, plain, (r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(X1)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_943, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3048, plain, (r1_int_1(np__0,k4_xcmplx_0(X1))|~r1_int_1(np__0,X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_945, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7827, plain, (r1_int_1(np__1,k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_2002, negated_conjecture, (k2_int_5(X3,esk2_0)=np__1|k2_int_5(X3,esk2_0)=np__0|k4_xcmplx_0(np__1)!=np__0|~v1_int_1(X3))).#trainneg
cnf(i_0_4841, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k1_newton(k4_xcmplx_0(np__1),X2))=k6_xcmplx_0(np__1,X1)|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v7_ordinal1(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_946, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2900, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_947, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6207, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))=k6_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1273, plain, (r1_int_1(X1,np__0)|~r1_int_1(k4_xcmplx_0(X1),np__0)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6158, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X1)|~v1_int_1(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4366, plain, (r1_int_1(X1,k1_newton(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_948, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2276, plain, (r1_int_1(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|~r1_int_1(X1,k2_int_5(X2,X3))|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_1(X1)|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_3344, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_954, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5178, plain, (r1_int_1(k2_int_5(X1,X2),X3)|~r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),k4_xcmplx_0(X3))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(X3))|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_955, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4238, plain, (r1_int_1(k4_xcmplx_0(X1),np__1)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5089, plain, (k3_int_2(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_int_2(X2,X1)|k1_int_2(X1)=X1|k1_int_2(X2)=X2|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_3113, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_956, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_11780, plain, (~v1_int_1(k4_xcmplx_0(esk2_0))|~r1_int_1(esk2_0,np__2))).#trainneg
cnf(i_0_11781, plain, (~v1_int_1(k4_xcmplx_0(esk2_0))|~r1_int_1(esk2_0,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_957, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13720, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_971, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),np__0)|~r1_int_1(k5_xcmplx_0(X1),np__0)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_11799, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).#trainneg
cnf(i_0_13385, plain, (r1_int_1(X1,X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_973, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__0)|~r1_int_1(k4_xcmplx_0(X1),np__0)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_975, plain, (r1_int_1(k4_xcmplx_0(X1),np__0)|~r1_int_1(X1,np__0)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_996, plain, (r1_int_1(np__0,k4_xcmplx_0(k5_xcmplx_0(X1)))|~r1_int_1(np__0,k5_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_11999, negated_conjecture, (r1_int_1(esk2_0,k6_xcmplx_0(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_13381, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8200, plain, (r1_int_1(k4_xcmplx_0(np__2),np__1)|~r1_int_1(np__2,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_8567, plain, (r1_int_1(k4_xcmplx_0(np__1),np__2)|~r1_int_1(np__1,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_7627, plain, (r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(np__1))|~r1_int_1(np__2,np__1))).#trainneg
cnf(i_0_1042, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2))).#trainneg
cnf(i_0_10561, plain, (k3_int_2(k4_xcmplx_0(np__1),np__0)=k3_int_2(np__1,np__0))).#trainneg
cnf(i_0_8858, plain, (r1_int_1(k4_xcmplx_0(np__2),np__2)|~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_8859, plain, (r1_int_1(k4_xcmplx_0(np__2),np__1)|~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_11368, plain, (r1_int_1(np__2,k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_11369, plain, (r1_int_1(np__1,k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_13855, plain, (r1_int_1(k2_int_5(X1,X2),X3)|~v1_int_2(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_1(k4_xcmplx_0(X3))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),X3))).#trainneg
cnf(i_0_1083, plain, (r1_int_1(np__0,k5_xcmplx_0(X1))|~r1_int_1(np__0,k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_12789, negated_conjecture, (k2_int_5(X2,esk2_0)=np__1|np__0!=np__1|k4_xcmplx_0(np__1)!=np__0|~v1_int_1(X2))).#trainneg
cnf(i_0_6655, plain, (k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0)))=k6_xcmplx_0(np__2,np__0))).#trainneg
cnf(i_0_1084, plain, (r1_int_1(np__0,k1_newton(X1,X2))|~r1_int_1(np__0,k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1085, plain, (r1_int_1(np__0,k4_xcmplx_0(X1))|~r1_int_1(np__0,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13100, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_12823, plain, (k6_xcmplx_0(X1,k1_newton(k4_xcmplx_0(np__1),X2))=k6_xcmplx_0(np__1,k4_xcmplx_0(X1))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13568, plain, (r1_int_1(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),k2_int_5(X2,X3))|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_13565, plain, (r1_int_1(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~r1_int_1(k4_xcmplx_0(X1),k2_int_5(X2,X3))|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_6076, plain, (k3_int_2(X2,k4_xcmplx_0(X1))=k3_int_2(X1,k4_xcmplx_0(X2))|k1_int_2(X1)=X1|k1_int_2(X2)=X2|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2))).#trainneg
cnf(i_0_4046, plain, (k3_int_2(k4_xcmplx_0(X1),X2)=k3_int_2(X1,k4_xcmplx_0(X2))|k1_int_2(X1)=X1|k1_int_2(X2)=X2|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_3026, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_4752, plain, (r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(np__1,k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1111, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1))).#trainneg
cnf(i_0_6934, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__1)|~v7_ordinal1(k4_xcmplx_0(np__1))|~r1_int_1(np__1,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_4222, plain, (k3_int_2(X1,k4_xcmplx_0(np__1))=k3_int_2(np__1,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_4228, plain, (k3_int_2(X1,k4_xcmplx_0(np__1))=k3_int_2(X1,np__1)|~v1_int_1(X1))).#trainneg
cnf(i_0_6659, plain, (v7_ordinal1(k6_xcmplx_0(np__2,np__0))|k4_xcmplx_0(k6_xcmplx_0(np__2,np__0))!=k6_xcmplx_0(np__2,np__0))).#trainneg
cnf(i_0_11792, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__1|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_4835, plain, (r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(np__1,X2)|~v7_ordinal1(X1)|~v1_int_1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5426, plain, (m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|k18_complex1(X1)=X1|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6444, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__1)|r1_xxreal_0(np__2,k4_xcmplx_0(np__1))|k4_xcmplx_0(np__1)!=np__1)).#trainneg
cnf(i_0_11796, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_6660, plain, (v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0)))|k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0)))!=k4_xcmplx_0(k6_xcmplx_0(np__2,np__0)))).#trainneg
cnf(i_0_6445, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__0)|r1_xxreal_0(np__1,k4_xcmplx_0(np__1))|k4_xcmplx_0(np__1)!=np__1)).#trainneg
cnf(i_0_6475, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),np__1)|r1_xxreal_0(np__2,k4_xcmplx_0(np__2))|k4_xcmplx_0(np__2)!=np__2)).#trainneg
cnf(i_0_2936, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2456, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_14184, plain, (~r1_int_1(esk2_0,np__2)|~v1_int_1(esk2_0))).#trainneg
cnf(i_0_14185, plain, (~r1_int_1(esk2_0,k4_xcmplx_0(np__2))|~v1_int_1(esk2_0))).#trainneg
cnf(i_0_14195, plain, (~r1_int_1(esk2_0,np__2)|~v7_ordinal1(k4_xcmplx_0(esk2_0)))).#trainneg
cnf(i_0_14653, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=np__1)).#trainneg
cnf(i_0_16725, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1))).#trainneg
cnf(i_0_14888, negated_conjecture, (r1_int_1(esk2_0,k6_xcmplx_0(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1))|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_15681, plain, (v1_xcmplx_0(k6_xcmplx_0(np__2,np__0)))).#trainneg
cnf(i_0_14626, plain, (r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X2))).#trainneg
cnf(i_0_14600, plain, (r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1))).#trainneg
cnf(i_0_15816, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1155, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))),X2)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16618, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_xcmplx_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1157, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14998, plain, (r1_int_1(X1,X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1159, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14526, plain, (r1_int_1(X1,np__0)|~r1_int_1(k4_xcmplx_0(X1),np__0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1160, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))),X2)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1176, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_16670, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1177, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_14985, plain, (r1_int_1(X1,np__2)|~r1_int_1(X1,k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1178, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(X2)))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_14986, plain, (r1_int_1(X1,np__1)|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1179, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1550, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_13563, plain, (r1_int_1(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(X1),k2_int_5(X2,X3))|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2))).#trainneg
cnf(i_0_4806, plain, (r1_int_1(k4_xcmplx_0(X1),np__1)|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~v1_int_1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2418, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_16726, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k4_xcmplx_0(np__1)!=np__1)).#trainneg
cnf(i_0_16728, negated_conjecture, (k4_xcmplx_0(np__1)!=np__1|k4_xcmplx_0(np__1)!=np__0|np__0!=np__1)).#trainneg
cnf(i_0_16820, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k4_xcmplx_0(np__1)!=np__0|np__0!=np__1)).#trainneg
cnf(i_0_1219, plain, (k18_complex1(k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|k1_int_2(X1)=X1|~v1_int_1(X1))).#trainneg
cnf(i_0_19030, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1223, plain, (v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_16825, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~r1_int_1(np__1,X1)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_17188, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1226, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))),X2)|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_18368, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1229, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),X2)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_15217, plain, (r1_int_1(k4_xcmplx_0(np__2),np__2)|~r1_int_1(np__2,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_1231, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_15202, plain, (k3_int_2(np__0,k4_xcmplx_0(np__1))=k3_int_2(np__1,np__0))).#trainneg
cnf(i_0_1232, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))),X2)|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_15452, plain, (v1_xreal_0(k6_xcmplx_0(np__2,np__0))|~v1_xreal_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0))))).#trainneg
cnf(i_0_1235, plain, (r1_int_1(np__0,k4_xcmplx_0(k5_xcmplx_0(X1)))|~r1_int_1(np__0,k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1238, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))),np__0)|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),np__0)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_17995, plain, (r1_int_1(k2_int_5(X1,X2),X3)|~r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),X3)|~v7_ordinal1(X3)|~v7_ordinal1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_1(k4_xcmplx_0(X3))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_17426, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_int_1(X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_17443, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1248, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X2))))|~r1_int_1(X1,k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X2))))|~v1_int_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_16961, plain, (r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),X3)|~r1_int_1(k2_int_5(X1,X2),X3)|~v7_ordinal1(X3)|~v7_ordinal1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_1(k4_xcmplx_0(X3))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_16836, plain, (r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),X3)|~r1_int_1(k2_int_5(X1,X2),X3)|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(X3))|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_1249, plain, (r1_int_1(X1,k4_xcmplx_0(k5_xcmplx_0(X2)))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X2))))|~v1_int_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_19174, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_12319, negated_conjecture, (r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).#trainneg
cnf(i_0_5862, plain, (k3_int_2(X2,k1_int_2(X1))=k3_int_2(k4_xcmplx_0(X2),X1)|k1_int_2(X2)=X2|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v7_ordinal1(k1_int_2(X1))|~v1_int_1(X2))).#trainneg
cnf(i_0_1250, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(k5_xcmplx_0(X2)))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X2))))|~v1_int_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_14529, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(X1),np__1)|~v7_ordinal1(k4_xcmplx_0(np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14541, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_14557, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6243, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X2,k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4278, plain, (r1_int_1(X1,np__2)|~r1_int_1(X1,k4_xcmplx_0(np__2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1251, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X2))))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X2))))|~v1_int_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_14999, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),np__1)|~v7_ordinal1(k4_xcmplx_0(np__2))|~r1_int_1(np__2,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_15028, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__2)|~v7_ordinal1(k4_xcmplx_0(np__1))|~r1_int_1(np__1,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_3157, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3835, plain, (k3_int_2(k4_xcmplx_0(X1),np__1)=k3_int_2(X1,np__1)|k1_int_2(X1)=X1|~v1_int_1(X1))).#trainneg
cnf(i_0_1267, plain, (r1_int_1(k5_xcmplx_0(X1),np__0)|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),np__0)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_10943, plain, (k3_int_2(k4_xcmplx_0(np__1),X1)=k3_int_2(X1,np__1)|~v1_int_1(X1))).#trainneg
cnf(i_0_14586, plain, (r1_int_1(np__2,X1)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_13223, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),X2)=k6_xcmplx_0(k4_xcmplx_0(X2),X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13365, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_19359, plain, (r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~r1_int_1(k4_xcmplx_0(X1),X2))).#trainneg
cnf(i_0_1269, plain, (r1_int_1(k4_xcmplx_0(X1),np__0)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__0)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_19253, negated_conjecture, (np__0!=np__1|k4_xcmplx_0(np__1)!=np__1)).#trainneg
cnf(i_0_20440, negated_conjecture, (~r1_int_1(esk2_0,k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(esk2_0)))).#trainneg
cnf(i_0_19661, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21358, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r1_int_1(X1,X2))).#trainneg
cnf(i_0_21892, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1270, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__0)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))),np__0)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v1_int_1(X1))).#trainneg
cnf(i_0_20026, plain, (k3_int_2(np__0,np__1)=k3_int_2(np__1,np__0))).#trainneg
cnf(i_0_1271, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__0)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))),np__0)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20530, plain, (r1_int_1(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~r1_int_1(X1,k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_1272, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),np__0)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))),np__0)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_20289, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_20401, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_1333, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k1_newton(X3,X4))|~v7_ordinal1(X4)|~v1_int_1(k4_xcmplx_0(k1_newton(X3,X4)))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(X3)|~r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k1_newton(X3,X4)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_20124, plain, (v1_xreal_0(k6_xcmplx_0(np__2,np__0))|~v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0))))).#trainneg
cnf(i_0_20125, plain, (v1_xreal_0(k6_xcmplx_0(np__2,np__0))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0))))).#trainneg
cnf(i_0_19757, plain, (r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(np__1))|~r1_int_1(np__2,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1334, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k4_xcmplx_0(X3)))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X3)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_19758, plain, (r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))|~r1_int_1(np__1,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_20126, plain, (v1_xreal_0(k6_xcmplx_0(np__2,np__0))|~m1_subset_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0)),k1_numbers))).#trainneg
cnf(i_0_19914, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),np__2)|~v7_ordinal1(k4_xcmplx_0(np__2))|~r1_int_1(np__2,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_16945, plain, (v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0)))|~v7_ordinal1(k6_xcmplx_0(np__2,np__0)))).#trainneg
cnf(i_0_20732, negated_conjecture, (r1_int_1(np__1,X2)|k2_int_5(X1,esk2_0)=np__1|k2_int_5(X1,esk2_0)=np__0|~r1_int_1(k4_xcmplx_0(np__1),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1))).#trainneg
cnf(i_0_1345, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2))).#trainneg
cnf(i_0_12316, negated_conjecture, (r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_19031, negated_conjecture, (r1_int_1(X1,np__1)|k2_int_5(X2,esk2_0)=np__1|k2_int_5(X2,esk2_0)=np__0|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_4244, plain, (r1_int_1(k4_xcmplx_0(X1),np__1)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1346, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2))).#trainneg
cnf(i_0_19220, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_22013, plain, (r1_int_1(np__2,X1)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__2),X1))).#trainneg
cnf(i_0_22914, plain, (k3_int_2(np__1,np__0)=np__1)).#trainneg
cnf(i_0_22723, plain, (r1_int_1(X1,X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_24287, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_22849, plain, (r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X2))).#trainneg
cnf(i_0_22815, plain, (r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1))).#trainneg
cnf(i_0_1378, plain, (k18_complex1(k1_newton(k4_xcmplx_0(np__1),X1))=k1_newton(k4_xcmplx_0(np__1),X1)|k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))=np__1|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_22727, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(X1),np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_24401, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(X1)|~r1_int_1(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_22793, plain, (r1_int_1(np__2,X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(np__2,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_23262, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1389, plain, (k3_int_2(k1_int_2(X2),k1_int_2(X1))=k3_int_2(X1,X2)|~v1_int_1(X2)|~v1_int_1(X1)|~v7_ordinal1(k1_int_2(X2))|~v7_ordinal1(k1_int_2(X1)))).#trainneg
cnf(i_0_22332, plain, (r1_int_1(np__1,X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(np__1,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_22740, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X1)|~v1_int_1(X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_22666, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_23849, plain, (r1_int_1(X1,np__1)|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_22983, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~r1_int_1(X1,np__1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_24341, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_23891, negated_conjecture, (r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_22794, plain, (r1_int_1(np__2,X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_22796, plain, (r1_int_1(np__1,X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_4237, plain, (r1_int_1(k4_xcmplx_0(X1),np__1)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1991, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),X2)=k6_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_4066, plain, (k3_int_2(X1,k4_xcmplx_0(X2))=k3_int_2(X1,X2)|k1_int_2(k4_xcmplx_0(X2))=k4_xcmplx_0(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3363, plain, (k3_int_2(k1_int_2(X1),X2)=k3_int_2(X1,k1_int_2(X2))|~v1_int_1(k1_int_2(X2))|~v1_int_1(X1)|~v1_int_1(X2)|~v1_int_1(k1_int_2(X1)))).#trainneg
cnf(i_0_6069, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_22709, plain, (r1_int_1(X1,np__2)|~r1_int_1(X1,k4_xcmplx_0(np__2))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_20026, plain, (k3_int_2(np__0,np__1)=np__1)).#trainneg
cnf(i_0_15202, plain, (k3_int_2(np__0,k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_10561, plain, (k3_int_2(k4_xcmplx_0(np__1),np__0)=np__1)).#trainneg
cnf(i_0_24878, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_24879, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_25032, plain, (r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_25509, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_25108, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~v7_ordinal1(X1)|~r1_int_1(X1,np__1))).#trainneg
cnf(i_0_24791, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_24632, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~r1_int_1(np__1,X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_25776, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_22726, plain, (r1_int_1(X1,k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(X1),np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10560, plain, (k3_int_2(k4_xcmplx_0(np__1),np__2)=k3_int_2(np__1,np__2)|v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_24880, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X2))).#trainneg
cnf(i_0_10428, plain, (r1_int_1(X1,np__2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__2))|~v1_int_1(X1))).#trainneg
cnf(i_0_10442, plain, (r1_int_1(X1,k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(X1),np__2)|~v1_int_1(X1))).#trainneg
cnf(i_0_22710, plain, (r1_int_1(X1,np__1)|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_25523, plain, (k3_int_2(k1_int_2(X1),np__1)=k3_int_2(np__1,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_26356, plain, (r1_int_1(np__1,X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(np__1),X1))).#trainneg
cnf(i_0_1540, plain, (v1_xreal_0(k1_newton(X1,X2))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(k1_newton(X1,X2)),k1_numbers))).#trainneg
cnf(i_0_25540, plain, (k3_int_2(np__1,k1_int_2(X1))=k3_int_2(X1,np__1)|~v1_int_1(X1))).#trainneg
cnf(i_0_4915, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),np__0)=k6_xcmplx_0(np__0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15551, plain, (v1_xreal_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0)))|~v1_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0)))|~m1_subset_1(k6_xcmplx_0(np__2,np__0),k1_numbers))).#trainneg
cnf(i_0_15466, plain, (r1_int_1(np__0,k6_xcmplx_0(np__2,np__0))|~r1_int_1(np__0,k4_xcmplx_0(k6_xcmplx_0(np__2,np__0)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0))))).#trainneg
cnf(i_0_15469, plain, (r1_int_1(k6_xcmplx_0(np__2,np__0),np__0)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0)),np__0)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0))))).#trainneg
cnf(i_0_1553, plain, (r1_int_1(np__1,k1_newton(X1,X2))|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4877, plain, (k6_xcmplx_0(np__0,X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15530, plain, (r1_int_1(k6_xcmplx_0(np__2,np__0),np__1)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0)),np__1)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0))))).#trainneg
cnf(i_0_15540, plain, (r1_int_1(k6_xcmplx_0(np__2,np__0),np__2)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0)),np__2)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0))))).#trainneg
cnf(i_0_15565, plain, (r1_int_1(np__1,k6_xcmplx_0(np__2,np__0))|~r1_int_1(np__1,k4_xcmplx_0(k6_xcmplx_0(np__2,np__0)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0))))).#trainneg
cnf(i_0_15579, plain, (r1_int_1(np__2,k6_xcmplx_0(np__2,np__0))|~r1_int_1(np__2,k4_xcmplx_0(k6_xcmplx_0(np__2,np__0)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0))))).#trainneg
cnf(i_0_24643, plain, (r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),X3)|~r1_int_1(k2_int_5(X1,X2),X3)|~v7_ordinal1(k4_xcmplx_0(X3))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_1554, plain, (r1_int_1(k5_xcmplx_0(X1),k1_newton(X2,X3))|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1555, plain, (r1_int_1(k1_newton(X1,X2),k1_newton(X3,X4))|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k4_xcmplx_0(k1_newton(X3,X4)))|~v7_ordinal1(X4)|~v1_int_1(k4_xcmplx_0(k1_newton(X3,X4)))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(X3)|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1556, plain, (r1_int_1(k4_xcmplx_0(X1),k1_newton(X2,X3))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xcmplx_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_27171, plain, (r1_int_1(X1,X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1560, plain, (r1_int_1(X1,k1_newton(X2,X3))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_28764, plain, (k6_xcmplx_0(np__2,np__0)=np__2|~v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_27157, plain, (r1_int_1(X1,np__2)|~r1_int_1(X1,k4_xcmplx_0(np__2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_27158, plain, (r1_int_1(X1,np__1)|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_28088, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_28310, plain, (r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(np__2))|~r1_int_1(np__2,np__2))).#trainneg
cnf(i_0_28071, plain, (k3_int_2(np__2,np__1)=k3_int_2(np__1,np__2)|v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_27169, plain, (r1_int_1(X1,k2_int_5(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|~v7_ordinal1(X1)|~v1_int_1(k2_int_5(X2,X3))|~v1_xcmplx_0(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_28311, plain, (r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2))|~r1_int_1(np__1,np__2))).#trainneg
cnf(i_0_1586, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1))).#trainneg
cnf(i_0_27929, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1587, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1))).#trainneg
cnf(i_0_27338, plain, (r1_int_1(np__0,X1)|~r1_int_1(np__0,k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1590, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1))).#trainneg
cnf(i_0_28065, plain, (k3_int_2(np__2,k4_xcmplx_0(np__1))=k3_int_2(np__1,np__2)|v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_27980, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_27535, plain, (r1_int_1(X1,np__1)|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(X1),np__1)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_28051, plain, (r1_int_1(X1,k4_xcmplx_0(np__2))|~v7_ordinal1(X1)|~r1_int_1(X1,np__2))).#trainneg
cnf(i_0_28479, plain, (k3_int_2(k4_xcmplx_0(np__2),np__1)=k3_int_2(np__1,k4_xcmplx_0(np__2))|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_28771, plain, (r1_int_1(k6_xcmplx_0(np__2,np__0),np__1)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0)))|~r1_int_1(k6_xcmplx_0(np__2,np__0),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_28700, plain, (r1_int_1(np__1,X1)|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_30609, negated_conjecture, (~r1_int_1(esk2_0,k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_30890, plain, (r1_int_1(X1,np__2)|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(X1),np__2)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_29782, plain, (k3_int_2(np__1,np__2)=np__1|v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_1596, plain, (v1_int_1(k7_xcmplx_0(X1,X2))|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_29856, plain, (r1_int_1(k4_xcmplx_0(X1),k2_int_5(X2,X3))|~v1_int_2(X3)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X3)|~v1_int_1(X2)|~r1_int_1(X1,k2_int_5(X2,X3))|~v7_ordinal1(X1)|~v7_ordinal1(k2_int_5(X2,X3)))).#trainneg
cnf(i_0_29251, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~r1_int_1(X1,X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_29454, plain, (r1_int_1(k4_xcmplx_0(X1),np__2)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1598, plain, (k6_xcmplx_0(k7_xcmplx_0(X1,X2),k4_xcmplx_0(X3))=k6_xcmplx_0(X3,k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1599, plain, (r1_int_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)),X3)|~r1_int_1(k7_xcmplx_0(X1,X2),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_29166, plain, (r1_int_1(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|~r1_int_1(X1,k2_int_5(X2,X3))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_1602, plain, (r1_int_1(k7_xcmplx_0(X1,X2),X3)|~r1_int_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_30927, plain, (k3_int_2(np__1,k4_xcmplx_0(np__2))=np__1|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_30928, plain, (k3_int_2(np__1,k4_xcmplx_0(np__2))=k3_int_2(np__2,np__1)|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_1604, plain, (r1_int_1(k7_xcmplx_0(X1,X2),k4_xcmplx_0(X3))|~r1_int_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_30930, plain, (k3_int_2(np__1,k4_xcmplx_0(np__2))=k3_int_2(np__1,np__2)|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_30749, plain, (r1_int_1(k2_int_5(X1,X2),X3)|~r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),k4_xcmplx_0(X3))|~v7_ordinal1(k4_xcmplx_0(X3))|~v7_ordinal1(X3)|~v7_ordinal1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_1605, plain, (r1_int_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)),X3)|~r1_int_1(k7_xcmplx_0(X1,X2),k4_xcmplx_0(X3))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1606, plain, (r1_int_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)),k4_xcmplx_0(X3))|~r1_int_1(k7_xcmplx_0(X1,X2),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_27156, plain, (r1_int_1(X1,np__1)|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|~v7_ordinal1(X1)|~v1_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1607, plain, (r1_int_1(k7_xcmplx_0(X1,X2),X3)|~r1_int_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)),k4_xcmplx_0(X3))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_29160, plain, (r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(X1,k1_newton(X2,X3))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_29161, plain, (r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(X1,k1_newton(X2,X3))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_27186, plain, (r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),X3)|~r1_int_1(k2_int_5(X1,X2),X3)|~v7_ordinal1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_1(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_1624, plain, (r1_int_1(k7_xcmplx_0(X1,X2),k1_newton(X3,X4))|~r1_int_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)),k4_xcmplx_0(k1_newton(X3,X4)))|~v7_ordinal1(X4)|~v1_int_1(k4_xcmplx_0(k1_newton(X3,X4)))|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_28480, plain, (k3_int_2(k4_xcmplx_0(np__2),np__1)=k3_int_2(np__1,k4_xcmplx_0(np__2))|k18_complex1(np__2)=np__2)).#trainneg
cnf(i_0_1628, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k7_xcmplx_0(X2,X3))=k6_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(X2,X3)),X1)|~v1_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_28473, plain, (k3_int_2(np__2,np__1)=k3_int_2(np__1,np__2)|m1_subset_1(k4_xcmplx_0(np__2),k5_numbers))).#trainneg
cnf(i_0_28798, plain, (r1_int_1(k6_xcmplx_0(np__2,np__0),np__2)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0)))|~r1_int_1(k6_xcmplx_0(np__2,np__0),k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_1629, plain, (r1_int_1(X1,k4_xcmplx_0(k7_xcmplx_0(X2,X3)))|~r1_int_1(X1,k7_xcmplx_0(X2,X3))|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X2,X3)))|~v1_int_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_28179, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X2))).#trainneg
cnf(i_0_26655, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1630, plain, (r1_int_1(X1,k7_xcmplx_0(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k7_xcmplx_0(X2,X3)))|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X2,X3)))|~v1_int_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_3939, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),X1)|r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_19254, negated_conjecture, (r1_int_1(esk2_0,k4_xcmplx_0(np__1))|k4_xcmplx_0(np__1)!=np__0|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1))|k4_xcmplx_0(np__1)!=np__1)).#trainneg
cnf(i_0_21037, negated_conjecture, (r2_int_1(np__0,np__1,esk2_0)|k4_xcmplx_0(np__1)!=np__0|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1))|k4_xcmplx_0(np__1)!=np__1)).#trainneg
cnf(i_0_3322, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~r1_int_1(X1,X2)|~v1_int_1(X2))).#trainneg
cnf(i_0_2146, plain, (k6_xcmplx_0(np__0,X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1631, plain, (r1_int_1(k4_xcmplx_0(X1),k7_xcmplx_0(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k7_xcmplx_0(X2,X3)))|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X2,X3)))|~v1_int_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_3943, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),X1)|r1_xxreal_0(X1,k4_xcmplx_0(np__2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_15767, plain, (r1_int_1(np__0,k4_xcmplx_0(k6_xcmplx_0(np__2,np__0)))|~r1_int_1(np__0,k6_xcmplx_0(np__2,np__0))|~v7_ordinal1(k6_xcmplx_0(np__2,np__0)))).#trainneg
cnf(i_0_22388, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0)),np__0)|~r1_int_1(k6_xcmplx_0(np__2,np__0),np__0)|~v7_ordinal1(k6_xcmplx_0(np__2,np__0)))).#trainneg
cnf(i_0_15467, plain, (r1_int_1(np__0,k4_xcmplx_0(k6_xcmplx_0(np__2,np__0)))|~r1_int_1(np__0,k6_xcmplx_0(np__2,np__0))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0))))).#trainneg
cnf(i_0_15468, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0)),np__0)|~r1_int_1(k6_xcmplx_0(np__2,np__0),np__0)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0))))).#trainneg
cnf(i_0_15526, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0)),np__1)|~r1_int_1(k6_xcmplx_0(np__2,np__0),np__1)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__0))))).#trainneg
cnf(i_0_4284, plain, (r1_int_1(X1,k1_newton(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_1632, plain, (r1_int_1(X1,k4_xcmplx_0(k7_xcmplx_0(X2,X3)))|~r1_int_1(k4_xcmplx_0(X1),k7_xcmplx_0(X2,X3))|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X2,X3)))|~v1_int_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_33091, plain, (r1_int_1(X1,np__2)|~r1_int_1(k4_xcmplx_0(X1),np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_31398, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_32493, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_33530, plain, (k6_xcmplx_0(np__2,np__0)=np__2)).#trainneg
cnf(i_0_33360, plain, (r1_xxreal_0(X1,k4_xcmplx_0(np__1))|r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_31818, plain, (k3_int_2(np__1,np__2)=np__1|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_31819, plain, (k3_int_2(np__2,np__1)=np__1|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_33629, plain, (r1_xxreal_0(X1,k4_xcmplx_0(np__2))|r1_xxreal_0(k4_xcmplx_0(np__2),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_32942, plain, (k3_int_2(X1,k4_xcmplx_0(np__1))=np__1|~v1_int_1(X1))).#trainneg
cnf(i_0_32770, plain, (k3_int_2(np__1,k4_xcmplx_0(np__2))=np__1|k18_complex1(np__2)=np__2)).#trainneg
cnf(i_0_31822, plain, (k3_int_2(np__2,np__1)=k3_int_2(np__1,np__2)|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_33365, negated_conjecture, (r1_int_1(esk2_0,k4_xcmplx_0(np__1))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__0|k4_xcmplx_0(np__1)!=np__1)).#trainneg
cnf(i_0_6476, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),np__0)|r1_xxreal_0(np__1,k4_xcmplx_0(np__2))|k4_xcmplx_0(np__2)!=np__2)).#trainneg
cnf(i_0_31989, plain, (r1_int_1(k2_int_5(X1,X2),X3)|~v1_int_2(X2)|~v7_ordinal1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v7_ordinal1(k4_xcmplx_0(X3))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),X3))).#trainneg
cnf(i_0_1647, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k1_newton(X2,X3))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xcmplx_0(X2)|~r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(k1_newton(X2,X3)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_32817, plain, (k3_int_2(np__1,np__2)=np__1|m1_subset_1(k4_xcmplx_0(np__2),k5_numbers))).#trainneg
cnf(i_0_23603, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__0)|r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(np__1))|~r1_int_1(np__1,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_6197, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_2397, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1))).#trainneg
cnf(i_0_33092, plain, (r1_int_1(X1,np__1)|~r1_int_1(k4_xcmplx_0(X1),np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1648, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_33361, plain, (r1_xxreal_0(X1,k4_xcmplx_0(np__1))|r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_33630, plain, (r1_xxreal_0(X1,k4_xcmplx_0(np__2))|r1_xxreal_0(k4_xcmplx_0(np__2),X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_15056, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__0)|r1_xxreal_0(k4_xcmplx_0(np__2),k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(np__1))|~r1_int_1(np__2,np__1))).#trainneg
cnf(i_0_29872, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),np__0)|r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(np__2))|~r1_int_1(np__1,np__2))).#trainneg
cnf(i_0_23485, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__0)|r1_xxreal_0(k4_xcmplx_0(np__2),k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(np__1))|~r1_int_1(np__2,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_21967, plain, (k3_int_2(X1,np__1)=k3_int_2(np__1,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_1649, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5224, plain, (v1_xcmplx_0(k1_int_2(X1))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(k1_int_2(X1))))).#trainneg
cnf(i_0_10559, plain, (k3_int_2(k4_xcmplx_0(np__1),np__2)=k3_int_2(np__1,np__2)|m1_subset_1(k4_xcmplx_0(np__2),k5_numbers))).#trainneg
cnf(i_0_10750, plain, (r1_xxreal_0(np__2,k4_xcmplx_0(np__2))|r1_xxreal_0(k4_xcmplx_0(np__2),np__1)|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_10751, plain, (r1_xxreal_0(np__2,k4_xcmplx_0(np__2))|r1_xxreal_0(k4_xcmplx_0(np__2),np__1)|k18_complex1(np__2)=np__2)).#trainneg
cnf(i_0_14887, negated_conjecture, (r1_int_1(esk2_0,k4_xcmplx_0(np__1))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k4_xcmplx_0(np__1)!=np__0|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).#trainneg
cnf(i_0_21035, negated_conjecture, (r2_int_1(np__0,np__1,esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k4_xcmplx_0(np__1)!=np__0|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).#trainneg
cnf(i_0_2317, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1))).#trainneg
cnf(i_0_24794, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v7_ordinal1(X1)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_31384, plain, (r1_int_1(np__2,X1)|~r1_int_1(k4_xcmplx_0(np__2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_10566, plain, (k3_int_2(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2))=k3_int_2(np__1,k4_xcmplx_0(np__2))|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_10567, plain, (k3_int_2(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2))=k3_int_2(np__1,k4_xcmplx_0(np__2))|k18_complex1(np__2)=np__2)).#trainneg
cnf(i_0_21903, plain, (k3_int_2(k4_xcmplx_0(np__2),np__1)=k3_int_2(np__2,np__1)|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_31385, plain, (r1_int_1(np__1,X1)|~r1_int_1(k4_xcmplx_0(np__1),X1)|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_12317, negated_conjecture, (r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_12000, negated_conjecture, (r1_int_1(esk2_0,k6_xcmplx_0(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1))|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_8468, plain, (r1_int_1(k4_xcmplx_0(X1),np__1)|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10156, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__2))|~r1_int_1(X1,np__2)|~v1_int_1(X1))).#trainneg
cnf(i_0_21034, negated_conjecture, (r2_int_1(k4_xcmplx_0(np__1),np__1,esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).#trainneg
cnf(i_0_12318, negated_conjecture, (r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))|~v1_int_1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))).#trainneg
cnf(i_0_12001, negated_conjecture, (r1_int_1(esk2_0,k6_xcmplx_0(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1))|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))|~v1_int_1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))).#trainneg
cnf(i_0_8999, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(k4_xcmplx_0(np__1),np__1,esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).#trainneg
cnf(i_0_10157, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~r1_int_1(X1,np__1)|~v1_int_1(X1))).#trainneg
cnf(i_0_1657, plain, (r1_int_1(k5_xcmplx_0(X1),k6_xcmplx_0(k2_int_5(X2,k4_xcmplx_0(k5_xcmplx_0(X1))),k1_newton(X2,k3_nat_d(k7_nat_d(k4_xcmplx_0(k5_xcmplx_0(X1)),np__1),np__2))))|r1_xxreal_0(k4_xcmplx_0(k5_xcmplx_0(X1)),np__2)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k6_xcmplx_0(k2_int_5(X2,k4_xcmplx_0(k5_xcmplx_0(X1))),k1_newton(X2,k3_nat_d(k7_nat_d(k4_xcmplx_0(k5_xcmplx_0(X1)),np__1),np__2))))|~v1_xreal_0(X1)|k3_int_2(X2,k4_xcmplx_0(k5_xcmplx_0(X1)))!=np__1|~v1_int_2(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v7_ordinal1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k1_newton(X2,k3_nat_d(k7_nat_d(k4_xcmplx_0(k5_xcmplx_0(X1)),np__1),np__2)))|~v1_int_1(X2))).#trainneg
cnf(i_0_4322, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_6396, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1),esk2_0)|k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))=np__1|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).#trainneg
cnf(i_0_5026, plain, (v1_xreal_0(k1_int_2(X1))|~v1_int_1(X1)|~m1_subset_1(k4_xcmplx_0(k1_int_2(X1)),k1_numbers))).#trainneg
cnf(i_0_29153, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~r1_int_1(X1,np__1)|~v7_ordinal1(X1)|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_11190, plain, (k1_int_2(X1)=X1|k1_int_2(k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2713, plain, (r1_int_1(k1_newton(X1,X2),X3)|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_2148, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(np__2))=np__2)).#trainneg
cnf(i_0_34761, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1))).#trainneg
cnf(i_0_35255, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r1_int_1(esk2_0,k4_xcmplx_0(np__1))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__0)).#trainneg
cnf(i_0_1661, plain, (r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_xreal_0(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))),X2)|~v1_int_1(X2))).#trainneg
cnf(i_0_35284, plain, (r1_int_1(k4_xcmplx_0(X1),k2_int_5(X2,X3))|~r1_int_1(X1,k2_int_5(X2,X3))|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_1(X1)|~v1_xcmplx_0(k2_int_5(X2,X3))|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_36296, plain, (m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|k1_int_2(X1)=X1|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_35606, plain, (k3_int_2(np__2,np__1)=np__1|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_35844, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_35670, plain, (r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X1)|~v1_int_1(k1_newton(X2,X3))|~r1_int_1(X1,k1_newton(X2,X3))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1678, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(X1),k5_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_34829, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_35355, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_34431, negated_conjecture, (r1_int_1(esk2_0,k4_xcmplx_0(np__1))|k4_xcmplx_0(np__1)!=np__0|k4_xcmplx_0(np__1)!=np__1|~v1_int_1(esk2_0))).#trainneg
cnf(i_0_34260, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_34001, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~r1_int_1(np__1,X1)|~v7_ordinal1(k4_xcmplx_0(np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1688, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k5_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X2)|~r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_35578, plain, (k3_int_2(k4_xcmplx_0(np__1),np__2)=k3_int_2(np__1,k4_xcmplx_0(np__2))|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_35579, plain, (k3_int_2(np__2,k4_xcmplx_0(np__1))=k3_int_2(np__1,k4_xcmplx_0(np__2))|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_34423, negated_conjecture, (r1_int_1(esk2_0,k4_xcmplx_0(np__1))|k3_int_2(esk2_0,k4_xcmplx_0(np__1))!=np__1|k4_xcmplx_0(np__1)!=np__0|k4_xcmplx_0(np__1)!=np__1|~v7_ordinal1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_34624, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_34625, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_35358, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_35213, plain, (k3_int_2(np__2,k4_xcmplx_0(np__1))=k3_int_2(np__1,np__2)|m1_subset_1(k4_xcmplx_0(np__2),k5_numbers))).#trainneg
cnf(i_0_34832, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_34887, plain, (r1_int_1(k4_xcmplx_0(X1),np__1)|~r1_int_1(X1,np__1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_35272, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_36286, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X2))|~r1_int_1(k4_xcmplx_0(X1),np__1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_35901, plain, (r1_int_1(X1,np__1)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_36697, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r1_int_1(esk2_0,k4_xcmplx_0(np__2))|np__0!=np__1|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1)).#trainneg
cnf(i_0_1691, plain, (r1_int_1(k4_xcmplx_0(X1),k5_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X2))))|~v1_int_1(X1))).#trainneg
cnf(i_0_1704, plain, (r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),X2)|~v1_int_1(X2))).#trainneg
cnf(i_0_36655, plain, (r1_int_1(np__1,X1)|~r1_int_1(k4_xcmplx_0(np__1),X1)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_37163, plain, (k3_int_2(np__1,np__2)=np__1|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_39103, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_37834, plain, (r1_int_1(esk2_0,np__1)|~v1_int_1(esk2_0)|k4_xcmplx_0(np__1)!=np__0|k4_xcmplx_0(np__1)!=np__1)).#trainneg
cnf(i_0_1716, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k1_newton(X2,X3))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xcmplx_0(X2)|~r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(k1_newton(X2,X3))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_37862, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_37392, plain, (r1_int_1(X1,np__1)|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1717, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_38983, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(X1))).#trainneg
cnf(i_0_1718, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1719, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_39119, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_37186, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1731, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k1_newton(X3,X4))|~v7_ordinal1(X4)|~v1_int_1(k4_xcmplx_0(k1_newton(X3,X4)))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(X3)|~r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k4_xcmplx_0(k1_newton(X3,X4))))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1733, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X3))))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X3)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_28759, plain, (k6_xcmplx_0(X1,np__0)=k4_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_3435, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),k4_xcmplx_0(X1))=k6_xcmplx_0(X1,np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_35601, plain, (k3_int_2(np__1,k4_xcmplx_0(np__2))=k3_int_2(np__2,np__1)|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_39038, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_1734, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k4_xcmplx_0(X3)))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X3))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3436, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))=k6_xcmplx_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_33076, plain, (r1_int_1(X1,k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_36023, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|r2_int_1(k4_xcmplx_0(np__1),np__1,esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).#trainneg
cnf(i_0_40396, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_39838, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_41451, negated_conjecture, (r1_int_1(esk2_0,k4_xcmplx_0(np__2))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|~v1_int_1(esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).#trainneg
cnf(i_0_40381, plain, (r1_int_1(np__0,X1)|~r1_int_1(np__0,k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_40375, plain, (r1_int_1(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(X1)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1744, plain, (r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),np__1)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(np__1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1747, plain, (r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k1_newton(X2,X3))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_xreal_0(X1)|~v7_ordinal1(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_40322, plain, (r1_int_1(np__2,X1)|~v7_ordinal1(k4_xcmplx_0(np__2))|~v7_ordinal1(X1)|~r1_int_1(np__2,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_39825, plain, (r1_int_1(np__2,X1)|~r1_int_1(k4_xcmplx_0(np__2),X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1748, plain, (r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(X2)))|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_xreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_41060, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),np__1)=k6_xcmplx_0(k4_xcmplx_0(np__1),np__2)|~v7_ordinal1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_41330, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),np__1)=k6_xcmplx_0(k4_xcmplx_0(np__1),np__2)|~v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_40952, plain, (r1_int_1(X1,np__2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_39826, plain, (r1_int_1(np__1,X1)|~r1_int_1(k4_xcmplx_0(np__1),X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1752, plain, (r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),X2)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_41345, plain, (r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(np__2))|~r1_int_1(np__2,k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_40328, plain, (r1_int_1(np__1,X1)|~v7_ordinal1(k4_xcmplx_0(np__1))|~v7_ordinal1(X1)|~r1_int_1(np__1,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_41346, plain, (r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2))|~r1_int_1(np__1,k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_41047, plain, (v1_int_1(k6_xcmplx_0(X1,np__2))|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_41317, plain, (v1_int_1(k6_xcmplx_0(X1,np__1))|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_35242, plain, (k3_int_2(np__2,np__1)=k3_int_2(np__1,np__2)|r1_xxreal_0(k4_xcmplx_0(np__2),np__1)|r1_xxreal_0(np__2,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_33077, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_39106, plain, (r1_int_1(np__2,k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_39107, plain, (r1_int_1(np__1,k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_35235, plain, (k3_int_2(k4_xcmplx_0(np__1),np__2)=k3_int_2(np__1,np__2)|r1_xxreal_0(k4_xcmplx_0(np__2),np__1)|r1_xxreal_0(np__2,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_40039, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_6146, plain, (r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__1),X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_4161, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_4225, plain, (k3_int_2(k1_int_2(X1),np__1)=k3_int_2(X1,k4_xcmplx_0(np__1))|~v1_int_1(X1))).#trainneg
cnf(i_0_36020, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|r2_int_1(k4_xcmplx_0(np__1),np__1,esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_39836, plain, (r1_int_1(k2_int_5(X1,X2),X3)|~r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),X3)|~v7_ordinal1(k4_xcmplx_0(X3))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_41832, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|~v1_int_1(esk2_0))).#trainneg
cnf(i_0_43661, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1796, plain, (r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(k5_xcmplx_0(X1)))|~r1_int_1(np__1,k5_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(np__1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_43770, plain, (r1_int_1(k4_xcmplx_0(np__1),np__2)|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_43771, plain, (r1_int_1(k4_xcmplx_0(np__1),np__1)|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1798, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k4_xcmplx_0(k5_xcmplx_0(X2)))|~r1_int_1(k5_xcmplx_0(X1),k5_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1799, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k4_xcmplx_0(k5_xcmplx_0(X3)))|~r1_int_1(k1_newton(X1,X2),k5_xcmplx_0(X3))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X3)))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xreal_0(X3)|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1800, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k4_xcmplx_0(k5_xcmplx_0(X2)))|~r1_int_1(k4_xcmplx_0(X1),k5_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xreal_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1804, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(k5_xcmplx_0(X2)))|~r1_int_1(X1,k5_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_41602, plain, (r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(X1))).#trainneg
cnf(i_0_42342, plain, (r1_int_1(np__2,X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~r1_int_1(np__2,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_41625, plain, (r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1))).#trainneg
cnf(i_0_1825, plain, (r1_int_1(np__1,k5_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(np__1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_42695, plain, (r1_int_1(np__1,X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~r1_int_1(np__1,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1833, plain, (r1_int_1(X1,k5_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_43828, plain, (k3_int_2(k4_xcmplx_0(X1),np__1)=k3_int_2(X1,k4_xcmplx_0(np__1))|k1_int_2(X1)=X1|~v1_int_1(X1))).#trainneg
cnf(i_0_1838, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_41919, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(X1))).#trainneg
cnf(i_0_43303, plain, (r1_int_1(np__2,k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~r1_int_1(np__2,X1)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_41358, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__2))|~r1_int_1(X1,k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_43814, plain, (k3_int_2(np__2,k4_xcmplx_0(np__1))=k3_int_2(np__2,np__1)|v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_43257, plain, (r1_int_1(np__2,k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(np__2),X1)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_43308, plain, (r1_int_1(np__2,k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~r1_int_1(np__2,X1)|~v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_43422, plain, (r1_int_1(np__1,k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~r1_int_1(np__1,X1)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1857, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),X2)|~r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_25529, plain, (k3_int_2(k1_int_2(X1),np__1)=k3_int_2(k4_xcmplx_0(np__1),X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_42834, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),np__0)|r1_xxreal_0(k4_xcmplx_0(np__2),k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(np__2))|~r1_int_1(np__2,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_1863, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k5_xcmplx_0(X3)))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X3)))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k5_xcmplx_0(X3))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_25546, plain, (k3_int_2(np__1,k1_int_2(X1))=k3_int_2(X1,k4_xcmplx_0(np__1))|~v1_int_1(X1))).#trainneg
cnf(i_0_42967, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),np__0)|r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(np__2))|~r1_int_1(np__1,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_43813, plain, (k3_int_2(np__2,k4_xcmplx_0(np__1))=k3_int_2(np__2,np__1)|m1_subset_1(k4_xcmplx_0(np__2),k5_numbers))).#trainneg
cnf(i_0_43086, plain, (k3_int_2(np__1,np__2)=np__1|r1_xxreal_0(np__2,k4_xcmplx_0(np__2))|r1_xxreal_0(k4_xcmplx_0(np__2),np__1))).#trainneg
cnf(i_0_43648, plain, (r1_int_1(np__2,X1)|~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_1865, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k4_xcmplx_0(X3)))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X3)))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2))),X3)|~v1_int_1(X3))).#trainneg
cnf(i_0_45898, plain, (k3_int_2(np__2,np__1)=np__1|v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_43966, negated_conjecture, (k4_xcmplx_0(np__1)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|~v1_int_1(esk2_0))).#trainneg
cnf(i_0_1866, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k4_xcmplx_0(X3)))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X3)))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~v1_int_1(X3))).#trainneg
cnf(i_0_44936, plain, (r1_int_1(np__2,X1)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~r1_int_1(np__2,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_44942, plain, (r1_int_1(np__1,X1)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~r1_int_1(np__1,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_44133, plain, (r1_int_1(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1))).#trainneg
cnf(i_0_46329, plain, (k3_int_2(k4_xcmplx_0(X1),np__1)=k3_int_2(k4_xcmplx_0(np__1),X1)|k1_int_2(X1)=X1|~v1_int_1(X1))).#trainneg
cnf(i_0_1878, plain, (r1_int_1(k1_newton(X1,X2),np__0)|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),np__0)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1881, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k7_xcmplx_0(X3,X4)))|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k7_xcmplx_0(X3,X4))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X3,X4)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_1883, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k1_newton(X3,X4)))|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k1_newton(X3,X4))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(k1_newton(X3,X4)))|~v1_xcmplx_0(X1)|~v7_ordinal1(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_1884, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k4_xcmplx_0(X3)))|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k4_xcmplx_0(X3))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X3)))|~v1_xcmplx_0(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_45596, plain, (r1_int_1(np__1,k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(np__1),X1)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_41550, plain, (r1_int_1(np__1,X1)|~v7_ordinal1(X1)|~r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1888, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(X3))|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(X3))|~v1_xcmplx_0(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_46400, plain, (k3_int_2(np__2,np__1)=np__1|m1_subset_1(k4_xcmplx_0(np__2),k5_numbers))).#trainneg
cnf(i_0_24789, plain, (r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),X3)|~r1_int_1(k2_int_5(X1,X2),X3)|~v7_ordinal1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_1(X3)|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_2746, plain, (r1_int_1(X1,k1_newton(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(X1)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_1894, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(X3))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(X3))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2))),X3)|~v1_int_1(X3))).#trainneg
cnf(i_0_26654, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_30744, plain, (r1_int_1(k1_newton(X1,X2),X3)|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k4_xcmplx_0(X3))|~v7_ordinal1(k4_xcmplx_0(X3))|~v7_ordinal1(X3)|~v7_ordinal1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_8480, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_34610, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_int_1(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3331, plain, (r1_int_1(k4_xcmplx_0(X1),np__0)|~r1_int_1(X1,np__0)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_44028, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(X2))).#trainneg
cnf(i_0_27163, plain, (r1_int_1(X1,k1_newton(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X1)|~v1_int_1(k1_newton(X2,X3))|~v1_xcmplx_0(k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2898, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X2))).#trainneg
cnf(i_0_34828, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(X1))).#trainneg
cnf(i_0_49128, plain, (r1_int_1(X1,X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_46672, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k4_xcmplx_0(np__1)=np__0|np__0!=np__1|~v1_int_1(esk2_0))).#trainneg
cnf(i_0_1907, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X3))))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X3)))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k5_xcmplx_0(X3))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_48821, plain, (r1_int_1(k4_xcmplx_0(X1),np__2)|~r1_int_1(X1,np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1909, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X3))))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X3)))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2))),X3)|~v1_int_1(X3))).#trainneg
cnf(i_0_48304, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1910, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X3))))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X3)))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~v1_int_1(X3))).#trainneg
cnf(i_0_48549, plain, (r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_1911, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k4_xcmplx_0(X3)))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(X3))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~v1_int_1(X3))).#trainneg
cnf(i_0_45879, plain, (r1_int_1(np__2,k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_45880, plain, (r1_int_1(np__1,k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_46315, plain, (k3_int_2(k4_xcmplx_0(np__1),np__2)=k3_int_2(np__2,np__1)|v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_30234, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__2)|~r1_int_1(X1,k4_xcmplx_0(np__2))|~v1_int_1(X1))).#trainneg
cnf(i_0_48822, plain, (r1_int_1(k4_xcmplx_0(X1),np__1)|~r1_int_1(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_46314, plain, (k3_int_2(k4_xcmplx_0(np__1),np__2)=k3_int_2(np__2,np__1)|m1_subset_1(k4_xcmplx_0(np__2),k5_numbers))).#trainneg
cnf(i_0_30235, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v1_int_1(X1))).#trainneg
cnf(i_0_1923, plain, (r1_int_1(np__0,k4_xcmplx_0(k1_newton(X1,X2)))|~r1_int_1(np__0,k1_newton(X1,X2))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_43821, plain, (k3_int_2(k4_xcmplx_0(np__2),k4_xcmplx_0(np__1))=k3_int_2(k4_xcmplx_0(np__2),np__1)|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_43822, plain, (k3_int_2(k4_xcmplx_0(np__2),k4_xcmplx_0(np__1))=k3_int_2(k4_xcmplx_0(np__2),np__1)|k18_complex1(np__2)=np__2)).#trainneg
cnf(i_0_34597, plain, (r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(X1))|~r1_int_1(np__2,k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_1925, plain, (r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(k1_newton(X1,X2)))|~r1_int_1(np__1,k1_newton(X1,X2))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7978, plain, (r1_int_1(k4_xcmplx_0(X1),np__0)|~r1_int_1(X1,np__0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_46322, plain, (k3_int_2(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2))=k3_int_2(k4_xcmplx_0(np__2),np__1)|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_46323, plain, (k3_int_2(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2))=k3_int_2(k4_xcmplx_0(np__2),np__1)|k18_complex1(np__2)=np__2)).#trainneg
cnf(i_0_46383, plain, (k3_int_2(k4_xcmplx_0(np__2),k4_xcmplx_0(np__1))=k3_int_2(np__1,k4_xcmplx_0(np__2))|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_46384, plain, (k3_int_2(k4_xcmplx_0(np__2),k4_xcmplx_0(np__1))=k3_int_2(np__1,k4_xcmplx_0(np__2))|k18_complex1(np__2)=np__2)).#trainneg
cnf(i_0_48757, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(X2))).#trainneg
cnf(i_0_47122, plain, (k3_int_2(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2))=k3_int_2(np__2,np__1)|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_34598, plain, (r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))|~r1_int_1(np__1,k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_1926, plain, (r1_int_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)),k4_xcmplx_0(k1_newton(X3,X4)))|~r1_int_1(k7_xcmplx_0(X1,X2),k1_newton(X3,X4))|~v7_ordinal1(X4)|~v1_int_1(k4_xcmplx_0(k1_newton(X3,X4)))|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_23263, plain, (r1_int_1(np__0,k4_xcmplx_0(X1))|~r1_int_1(np__0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_45321, plain, (k3_int_2(k4_xcmplx_0(np__2),k4_xcmplx_0(np__1))=k3_int_2(np__2,np__1)|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_43436, plain, (k3_int_2(np__2,k4_xcmplx_0(np__1))=k3_int_2(np__1,np__2)|r1_xxreal_0(np__2,k4_xcmplx_0(np__2))|r1_xxreal_0(k4_xcmplx_0(np__2),np__1))).#trainneg
cnf(i_0_35259, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__2))|~r1_int_1(X1,k4_xcmplx_0(np__2))|~v1_int_1(X1))).#trainneg
cnf(i_0_1927, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(k5_xcmplx_0(X1),k1_newton(X2,X3))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_47098, plain, (k3_int_2(k4_xcmplx_0(X1),np__1)=np__1|k1_int_2(X1)=X1|~v1_int_1(X1))).#trainneg
cnf(i_0_28756, plain, (k6_xcmplx_0(X1,np__0)=k4_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_43815, plain, (k3_int_2(np__2,k4_xcmplx_0(np__1))=k3_int_2(np__2,np__1)|r1_xxreal_0(k4_xcmplx_0(np__2),np__1)|r1_xxreal_0(np__2,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_46316, plain, (k3_int_2(k4_xcmplx_0(np__1),np__2)=k3_int_2(np__2,np__1)|r1_xxreal_0(k4_xcmplx_0(np__2),np__1)|r1_xxreal_0(np__2,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_45586, plain, (r1_int_1(np__2,k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(np__2),X1)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_4672, plain, (r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(X1,k1_newton(X2,X3))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_48501, plain, (r1_int_1(np__2,X1)|~v7_ordinal1(k4_xcmplx_0(np__2))|~v7_ordinal1(X1)|~r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2637, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__0)|r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X2,X1)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2))).#trainneg
cnf(i_0_1928, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k4_xcmplx_0(k1_newton(X3,X4)))|~r1_int_1(k1_newton(X1,X2),k1_newton(X3,X4))|~v7_ordinal1(X4)|~v1_int_1(k4_xcmplx_0(k1_newton(X3,X4)))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(X3)|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1929, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(k4_xcmplx_0(X1),k1_newton(X2,X3))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xcmplx_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1930, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))),k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k1_newton(X2,X3))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v1_xcmplx_0(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_1931, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))),k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k1_newton(X2,X3))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v1_xcmplx_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1932, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))),k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k1_newton(X2,X3))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_52479, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1933, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(X1,k1_newton(X2,X3))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_50987, plain, (k3_int_2(k4_xcmplx_0(np__2),np__1)=np__1|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_51003, plain, (k3_int_2(k4_xcmplx_0(np__2),np__1)=np__1|k18_complex1(np__2)=np__2)).#trainneg
cnf(i_0_49626, plain, (r1_int_1(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|~r1_int_1(X1,k2_int_5(X2,X3))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_49621, plain, (r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(X1,k1_newton(X2,X3))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_52032, plain, (r1_int_1(np__2,k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_52033, plain, (r1_int_1(np__1,k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_1949, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2))),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(X3))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X3))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_49614, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~r1_int_1(X1,np__1)|~v7_ordinal1(X1)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_52568, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~r1_int_1(X2,k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_51138, plain, (r1_int_1(k4_xcmplx_0(np__2),X1)|~r1_int_1(np__2,X1)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_49618, plain, (r1_int_1(X1,k4_xcmplx_0(k5_xcmplx_0(X2)))|~r1_int_1(X1,k5_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_50988, plain, (k3_int_2(k4_xcmplx_0(np__2),np__1)=k3_int_2(np__2,k4_xcmplx_0(np__1))|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_50990, plain, (k3_int_2(k4_xcmplx_0(np__2),np__1)=k3_int_2(k4_xcmplx_0(np__1),np__2)|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_49750, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_49908, plain, (r1_int_1(X1,k4_xcmplx_0(np__2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),np__2)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_30219, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k4_xcmplx_0(np__2))|~r1_int_1(X1,np__2)|~v1_int_1(X1))).#trainneg
cnf(i_0_1955, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k4_xcmplx_0(X3)))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(X3))|~v1_xcmplx_0(X1)|~r1_int_1(k1_newton(X1,X2),X3)|~v1_int_1(X3))).#trainneg
cnf(i_0_52176, plain, (k3_int_2(np__2,np__1)=np__1|r1_xxreal_0(np__2,k4_xcmplx_0(np__2))|r1_xxreal_0(k4_xcmplx_0(np__2),np__1))).#trainneg
cnf(i_0_54180, plain, (~v7_ordinal1(k4_xcmplx_0(np__1))|~r1_int_1(np__2,np__1))).#trainneg
cnf(i_0_54941, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(X2))).#trainneg
cnf(i_0_54320, plain, (r1_xxreal_0(X1,np__2)|~v7_ordinal1(k4_xcmplx_0(np__2))|~v7_ordinal1(X1)|~r1_int_1(X1,np__2))).#trainneg
cnf(i_0_54178, negated_conjecture, (~v7_ordinal1(k4_xcmplx_0(np__1))|k4_xcmplx_0(np__1)!=np__0|k4_xcmplx_0(np__1)!=np__1|~v1_int_1(esk2_0))).#trainneg
cnf(i_0_54685, plain, (k3_int_2(k4_xcmplx_0(np__1),np__2)=k3_int_2(np__1,np__2)|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_1968, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(X3))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(X3))|~v1_xcmplx_0(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_54670, plain, (k3_int_2(k4_xcmplx_0(np__1),np__2)=k3_int_2(np__2,k4_xcmplx_0(np__1))|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_54683, plain, (k3_int_2(k4_xcmplx_0(np__1),np__2)=k3_int_2(np__2,np__1)|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_54323, plain, (r1_xxreal_0(X1,np__1)|~v7_ordinal1(k4_xcmplx_0(np__1))|~v7_ordinal1(X1)|~r1_int_1(X1,np__1))).#trainneg
cnf(i_0_53127, plain, (r1_int_1(np__2,k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__2),X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_53486, plain, (k3_int_2(k4_xcmplx_0(np__1),np__2)=np__1|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_1984, plain, (k6_xcmplx_0(k4_xcmplx_0(X3),k4_xcmplx_0(k1_newton(X1,X2)))=k6_xcmplx_0(k1_newton(X1,X2),X3)|~v7_ordinal1(X2)|~v1_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_54184, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__2,X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~r1_int_1(np__2,X1))).#trainneg
cnf(i_0_1986, plain, (k6_xcmplx_0(np__0,k1_newton(X1,X2))=k4_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_54666, plain, (k3_int_2(np__2,k4_xcmplx_0(np__1))=k3_int_2(np__1,np__2)|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_1989, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X3)),k1_newton(X1,X2))=k6_xcmplx_0(X3,k1_newton(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k1_newton(X1,X2))|~v7_ordinal1(X2)|~v1_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_54144, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v7_ordinal1(X1)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~r1_int_1(k4_xcmplx_0(X1),np__1)|~v7_ordinal1(k4_xcmplx_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2000, plain, (k6_xcmplx_0(k4_xcmplx_0(k1_newton(X2,X3)),k4_xcmplx_0(X1))=k6_xcmplx_0(X1,k1_newton(X2,X3))|~v7_ordinal1(X3)|~v1_xcmplx_0(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2001, negated_conjecture, (k2_int_5(X2,esk2_0)=np__1|k2_int_5(X2,esk2_0)=np__0|k4_xcmplx_0(np__1)!=np__1|~v1_int_1(X2))).#trainneg
cnf(i_0_53485, plain, (k3_int_2(np__2,k4_xcmplx_0(np__1))=np__1|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_2005, negated_conjecture, (r1_int_1(esk2_0,k6_xcmplx_0(k4_xcmplx_0(np__1),k1_newton(X1,k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))|k2_int_5(X1,esk2_0)=np__1|k2_int_5(X1,esk2_0)=np__0|k3_int_2(X1,esk2_0)!=np__1|~v1_int_1(k1_newton(X1,k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|~v1_int_1(X1))).#trainneg
cnf(i_0_54665, plain, (k3_int_2(np__2,k4_xcmplx_0(np__1))=k3_int_2(np__2,np__1)|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_10371, plain, (r1_int_1(np__2,X1)|~v1_int_1(X1)|~r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_4544, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|v7_ordinal1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~m1_subset_1(X1,k5_numbers)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_24796, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v7_ordinal1(X1)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~r1_int_1(k4_xcmplx_0(X1),np__1)|~v1_int_1(k4_xcmplx_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_26510, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_55112, plain, (r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(np__2),np__2))).#trainneg
cnf(i_0_43990, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~r1_int_1(np__1,k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4717, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_55113, plain, (r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_3134, plain, (r1_int_1(k4_xcmplx_0(X1),k2_int_5(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_2019, plain, (r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(X1))|~r1_int_1(np__2,X1)|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_54266, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),np__2)|~v7_ordinal1(k4_xcmplx_0(np__2))|~r1_int_1(np__2,np__2))).#trainneg
cnf(i_0_54274, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__1)|~v7_ordinal1(k4_xcmplx_0(np__1))|~r1_int_1(np__1,np__1))).#trainneg
cnf(i_0_4367, plain, (r1_int_1(X1,k1_newton(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_2037, plain, (r1_int_1(np__2,k1_newton(X1,X2))|~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13232, plain, (k6_xcmplx_0(X1,k1_newton(k4_xcmplx_0(np__1),X2))=k6_xcmplx_0(X1,k4_xcmplx_0(np__1))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v7_ordinal1(X1)|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4839, plain, (r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(np__1,k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3888, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_2038, plain, (r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(k1_newton(X1,X2)))|~r1_int_1(np__2,k1_newton(X1,X2))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_55433, negated_conjecture, (k4_xcmplx_0(np__1)!=np__0|k4_xcmplx_0(np__1)!=np__1)).#trainneg
cnf(i_0_55619, negated_conjecture, (~r1_int_1(esk2_0,np__1)|~v7_ordinal1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_2040, plain, (r1_int_1(np__2,k5_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(np__2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2041, plain, (r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(k5_xcmplx_0(X1)))|~r1_int_1(np__2,k5_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(np__2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_56127, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).#trainneg
cnf(i_0_57109, plain, (r1_int_1(np__2,X1)|~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2049, plain, (r1_int_1(X1,k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__2))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_55432, plain, (~r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1)|~r1_int_1(k1_nat_1(X1,np__1),np__2)|~v7_ordinal1(k4_xcmplx_0(np__2))|~v7_ordinal1(k1_nat_1(X1,np__1)))).#trainneg
cnf(i_0_2052, plain, (r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),np__2)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(np__2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_55361, plain, (r1_int_1(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_int_1(X2)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_55430, plain, (~r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1)|~r1_int_1(k2_xcmplx_0(X1,np__1),np__2)|~v7_ordinal1(k4_xcmplx_0(np__2))|~v7_ordinal1(k2_xcmplx_0(X1,np__1)))).#trainneg
cnf(i_0_55371, plain, (r1_int_1(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_int_1(X2)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_2099, plain, (k3_int_2(np__0,k1_int_2(X1))=k3_int_2(np__0,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_2106, plain, (k3_int_2(np__0,X1)=k3_int_2(X1,np__0)|~v1_int_1(X1))).#trainneg
cnf(i_0_2107, plain, (k3_int_2(k1_int_2(X1),np__0)=k3_int_2(X1,np__0)|~v1_int_1(X1))).#trainneg
cnf(i_0_56360, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X2,k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_56218, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|v7_ordinal1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_56327, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2136, plain, (v7_ordinal1(k6_xcmplx_0(X1,X2))|k4_xcmplx_0(k6_xcmplx_0(X1,X2))!=k6_xcmplx_0(X1,X2)|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_2137, plain, (v7_ordinal1(k1_newton(X1,X2))|k4_xcmplx_0(k1_newton(X1,X2))!=k1_newton(X1,X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_56967, plain, (r1_int_1(k4_xcmplx_0(np__2),X1)|~r1_int_1(np__2,k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_56318, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))=k6_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_55936, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2201, plain, (r1_int_1(X1,k4_xcmplx_0(k5_xcmplx_0(X2)))|~r1_int_1(k4_xcmplx_0(X1),k5_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_56857, plain, (r1_int_1(k4_xcmplx_0(np__2),X1)|~v7_ordinal1(k4_xcmplx_0(np__2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(np__2,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2202, plain, (r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(k4_xcmplx_0(X1),k1_newton(X2,X3))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2203, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_8513, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_30157, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2204, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(X2))).#trainneg
cnf(i_0_56373, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),np__2)|~v7_ordinal1(k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(np__2),np__2))).#trainneg
cnf(i_0_46373, plain, (k3_int_2(X1,k4_xcmplx_0(np__1))=k3_int_2(k4_xcmplx_0(np__1),X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_3343, plain, (r1_int_1(k4_xcmplx_0(X1),k2_int_5(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_xcmplx_0(X1)|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_56585, plain, (r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2))|~r1_int_1(np__1,k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_2205, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_39177, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v7_ordinal1(X1)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X2))|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(k4_xcmplx_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_55626, plain, (~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~r1_int_1(k1_nat_1(X1,np__1),np__1)|~v7_ordinal1(k4_xcmplx_0(np__1))|~v7_ordinal1(k1_nat_1(X1,np__1)))).#trainneg
cnf(i_0_2206, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X2))))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(k5_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X2))))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_24059, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__1)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,np__1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_54272, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__2)|~v7_ordinal1(k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(np__1))|~r1_int_1(np__1,np__2))).#trainneg
cnf(i_0_24304, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_60477, plain, (~r1_int_1(k4_xcmplx_0(np__2),np__1))).#trainneg
cnf(i_0_60332, plain, (~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_57642, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).#trainneg
cnf(i_0_57638, plain, (r1_int_1(k4_xcmplx_0(np__2),np__1)|~r1_int_1(np__2,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_60490, plain, (r1_xxreal_0(X1,np__1)|~r1_int_1(k4_xcmplx_0(X1),np__1)|~v7_ordinal1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_57639, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_57776, plain, (r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(np__2))|~r1_int_1(np__2,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_57777, plain, (r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2))|~r1_int_1(np__1,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_57703, plain, (r1_int_1(np__2,X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(np__2),X1))).#trainneg
cnf(i_0_58875, plain, (r1_int_1(k4_xcmplx_0(np__2),np__1)|~r1_int_1(np__2,np__1)|~v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_58659, plain, (r1_int_1(np__2,X1)|~v7_ordinal1(X1)|~r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_2237, plain, (v1_xcmplx_0(k2_int_5(X1,X2))|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_57905, plain, (~r1_int_1(k2_xcmplx_0(np__1,X1),np__2)|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(k2_xcmplx_0(np__1,X1))|~v7_ordinal1(k4_xcmplx_0(np__2))|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2238, plain, (v1_xreal_0(k2_int_5(X1,X2))|~v1_xreal_0(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_55839, plain, (r1_xxreal_0(np__2,k4_xcmplx_0(np__2))|r1_xxreal_0(k4_xcmplx_0(np__2),np__0)|~v7_ordinal1(k4_xcmplx_0(np__2))|~r1_int_1(np__2,np__2))).#trainneg
cnf(i_0_30220, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k4_xcmplx_0(np__1))|~r1_int_1(X1,np__1)|~v1_int_1(X1))).#trainneg
cnf(i_0_2240, plain, (v1_xcmplx_0(k2_int_5(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_56361, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),np__0)|r1_xxreal_0(k4_xcmplx_0(np__2),k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(np__2),np__2))).#trainneg
cnf(i_0_2241, plain, (k6_xcmplx_0(k2_int_5(X1,X2),k4_xcmplx_0(X3))=k6_xcmplx_0(X3,k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_56655, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),np__0)|r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_2242, plain, (r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),X3)|~r1_int_1(k2_int_5(X1,X2),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_59538, plain, (r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5777, plain, (v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(X1)|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))))).#trainneg
cnf(i_0_2243, plain, (v1_xcmplx_0(k2_int_5(X1,X2))|~v7_ordinal1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_59542, plain, (r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5931, plain, (v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(X1)|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))))).#trainneg
cnf(i_0_2246, plain, (r1_int_1(k2_int_5(X1,X2),X3)|~r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_59865, plain, (k3_int_2(k4_xcmplx_0(np__1),k1_int_2(X1))=k3_int_2(k4_xcmplx_0(np__1),X1)|~v1_int_1(X1)|~v1_int_1(k1_int_2(X1)))).#trainneg
cnf(i_0_5970, plain, (v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))))).#trainneg
cnf(i_0_2248, plain, (r1_int_1(k2_int_5(X1,X2),k4_xcmplx_0(X3))|~r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_35354, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(X1))).#trainneg
cnf(i_0_60505, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2249, plain, (r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),X3)|~r1_int_1(k2_int_5(X1,X2),k4_xcmplx_0(X3))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_16745, plain, (r1_int_1(k4_xcmplx_0(X1),k1_newton(X2,X3))|~r1_int_1(X1,k1_newton(X2,X3))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_7675, plain, (v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2250, plain, (r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),k4_xcmplx_0(X3))|~r1_int_1(k2_int_5(X1,X2),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_29657, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(X2))).#trainneg
cnf(i_0_60844, negated_conjecture, (~r1_int_1(k4_xcmplx_0(esk2_0),np__1))).#trainneg
cnf(i_0_60820, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1)).#trainneg
cnf(i_0_2251, plain, (r1_int_1(k2_int_5(X1,X2),X3)|~r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),k4_xcmplx_0(X3))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_60719, plain, (~r1_int_1(np__2,np__1))).#trainneg
cnf(i_0_2265, plain, (r1_int_1(k2_int_5(X1,X2),k1_newton(X3,X4))|~r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),k4_xcmplx_0(k1_newton(X3,X4)))|~v7_ordinal1(X4)|~v1_int_1(k4_xcmplx_0(k1_newton(X3,X4)))|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_xcmplx_0(X3)|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_2266, plain, (r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),k4_xcmplx_0(k1_newton(X3,X4)))|~r1_int_1(k2_int_5(X1,X2),k1_newton(X3,X4))|~v7_ordinal1(X4)|~v1_int_1(k4_xcmplx_0(k1_newton(X3,X4)))|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_xcmplx_0(X3)|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_62893, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2268, plain, (r1_int_1(k2_int_5(X1,X2),k5_xcmplx_0(X3))|~r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),k4_xcmplx_0(k5_xcmplx_0(X3)))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X3)))|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_xreal_0(X3)|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_2269, plain, (r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),k4_xcmplx_0(k5_xcmplx_0(X3)))|~r1_int_1(k2_int_5(X1,X2),k5_xcmplx_0(X3))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X3)))|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_xreal_0(X3)|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_61260, plain, (~r1_int_1(k2_xcmplx_0(X1,np__1),np__2)|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__1))|~v7_ordinal1(k4_xcmplx_0(np__2))|~m1_subset_1(X1,k5_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2275, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k2_int_5(X2,X3))=k6_xcmplx_0(k4_xcmplx_0(k2_int_5(X2,X3)),X1)|~v1_xcmplx_0(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_xcmplx_0(X1)|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_2278, plain, (r1_int_1(k4_xcmplx_0(X1),k2_int_5(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_1(X1)|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_2283, plain, (r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(k2_int_5(X2,X3)))|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k2_int_5(X2,X3))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_xreal_0(X1)|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_60859, plain, (~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(k1_nat_1(X1,np__1)),k4_xcmplx_0(np__1))|~v7_ordinal1(k1_nat_1(X1,np__1)))).#trainneg
cnf(i_0_61374, plain, (r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(np__1),np__1))).#trainneg
cnf(i_0_2287, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k2_int_5(X3,X4)))|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k2_int_5(X3,X4))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(k2_int_5(X3,X4)))|~v1_xcmplx_0(X1)|~v1_int_2(X4)|~v7_ordinal1(X4)|~v1_int_1(X3))).#trainneg
cnf(i_0_62060, plain, (k3_int_2(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))=k3_int_2(k4_xcmplx_0(np__1),X1)|k1_int_2(X1)=X1|~v1_int_1(X1))).#trainneg
cnf(i_0_60857, plain, (~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(k2_xcmplx_0(X1,np__1)),k4_xcmplx_0(np__1))|~v7_ordinal1(k2_xcmplx_0(X1,np__1)))).#trainneg
cnf(i_0_43790, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_62911, plain, (k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))=k4_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_15922, plain, (r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),k4_xcmplx_0(k2_int_5(X2,X3)))|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(np__1,k2_int_5(X2,X3))|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_1(X2)|~v1_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_15819, plain, (r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(X1))|~r1_int_1(np__2,k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_55235, plain, (r1_int_1(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|~r1_int_1(k4_xcmplx_0(X1),k2_int_5(X2,X3))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_13638, plain, (r1_int_1(X1,np__2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(X1,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_61961, plain, (r1_int_1(k2_int_5(X1,X2),k4_xcmplx_0(X3))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_1(k4_xcmplx_0(X3))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),X3)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_15820, plain, (r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))|~r1_int_1(np__1,k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_33525, plain, (k6_xcmplx_0(X1,np__0)=k4_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_60327, plain, (~r1_int_1(k2_xcmplx_0(X1,np__1),np__1)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__1))|~v7_ordinal1(k4_xcmplx_0(np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_58743, plain, (k6_xcmplx_0(X1,k1_newton(k4_xcmplx_0(np__1),X2))=k6_xcmplx_0(X1,k4_xcmplx_0(np__1))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v7_ordinal1(X1)|~v1_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_5969, plain, (v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))))).#trainneg
cnf(i_0_6012, plain, (v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))))).#trainneg
cnf(i_0_13386, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~r1_int_1(X1,np__1)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_24795, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v7_ordinal1(X1)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~v1_int_1(k4_xcmplx_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_22596, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_28129, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_63401, negated_conjecture, (k4_xcmplx_0(np__1)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1)).#trainneg
cnf(i_0_2312, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xcmplx_0(X2)|~r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2313, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xcmplx_0(X2)|~r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_65545, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2314, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~v1_int_1(k4_xcmplx_0(X3))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(X3)|~r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(X3))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2315, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~v1_int_1(k4_xcmplx_0(X3))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(X3)|~r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k4_xcmplx_0(X3)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_64845, plain, (r1_int_1(np__2,X1)|~v7_ordinal1(X1)|~r1_int_1(np__2,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_65395, plain, (r1_int_1(np__1,X1)|~v7_ordinal1(X1)|~r1_int_1(np__1,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_65446, plain, (v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2336, plain, (r1_int_1(X1,k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(X1),np__2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2339, plain, (r1_int_1(X1,k4_xcmplx_0(k5_xcmplx_0(X2)))|~r1_int_1(k4_xcmplx_0(X1),k5_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2340, plain, (r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(k4_xcmplx_0(X1),k1_newton(X2,X3))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_64584, plain, (k3_int_2(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2))=k3_int_2(k4_xcmplx_0(np__1),np__2)|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_2341, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_64877, plain, (r1_int_1(k4_xcmplx_0(np__2),X1)|~r1_int_1(np__2,X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2342, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_xcmplx_0(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_2343, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_64779, plain, (r1_int_1(k4_xcmplx_0(np__2),X1)|~v7_ordinal1(k4_xcmplx_0(np__2))|~v7_ordinal1(X1)|~r1_int_1(np__2,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_64394, plain, (~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k1_nat_1(X1,np__1))|~v7_ordinal1(X1)|~r1_int_1(k1_nat_1(X1,np__1),k4_xcmplx_0(np__1))|~v1_int_1(k1_nat_1(X1,np__1)))).#trainneg
cnf(i_0_2344, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X2))))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X2))))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_63986, plain, (r1_int_1(k2_int_5(X1,X2),X3)|~r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),X3)|~v1_int_1(k4_xcmplx_0(X3))|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_xcmplx_0(X3)|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_64649, plain, (~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__1))|~v7_ordinal1(X1)|~r1_int_1(k2_xcmplx_0(X1,np__1),np__1))).#trainneg
cnf(i_0_64422, plain, (~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k1_nat_1(X1,np__1))|~v7_ordinal1(X1)|~r1_int_1(k1_nat_1(X1,np__1),np__1))).#trainneg
cnf(i_0_64621, plain, (~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__1))|~v7_ordinal1(X1)|~r1_int_1(k2_xcmplx_0(X1,np__1),k4_xcmplx_0(np__1))|~v1_int_1(k2_xcmplx_0(X1,np__1)))).#trainneg
cnf(i_0_64446, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__1)|~v7_ordinal1(k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(np__1),np__1))).#trainneg
cnf(i_0_65652, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v7_ordinal1(X1)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~r1_int_1(k4_xcmplx_0(X1),np__1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_65494, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))))).#trainneg
cnf(i_0_64434, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__0)|r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(np__1),np__1))).#trainneg
cnf(i_0_65639, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),np__1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_64331, plain, (r1_int_1(X1,k2_int_5(X2,X3))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_30535, plain, (r1_int_1(k4_xcmplx_0(np__2),X1)|~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_30536, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_25522, plain, (k3_int_2(k1_int_2(X1),np__0)=k3_int_2(np__0,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_65760, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(X2,k4_xcmplx_0(np__1))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(X2)|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)),np__1)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_25539, plain, (k3_int_2(np__0,k1_int_2(X1))=k3_int_2(X1,np__0)|~v1_int_1(X1))).#trainneg
cnf(i_0_50989, plain, (k3_int_2(k4_xcmplx_0(np__2),np__1)=k3_int_2(np__2,np__1)|k4_xcmplx_0(np__1)=np__1|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_4667, plain, (r1_int_1(X1,k4_xcmplx_0(np__2))|~r1_int_1(X1,np__2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_54743, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_36152, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_33522, plain, (k6_xcmplx_0(X1,np__0)=k4_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_6830, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),np__1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_49438, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~r1_int_1(X1,np__2)|~v7_ordinal1(k4_xcmplx_0(np__2))|~v1_int_1(X1))).#trainneg
cnf(i_0_6294, plain, (v1_int_1(k6_xcmplx_0(X1,np__0))|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_28619, plain, (v1_int_1(k6_xcmplx_0(np__0,X1))|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_65436, plain, (k6_xcmplx_0(np__0,X1)=k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_30498, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(X1))).#trainneg
cnf(i_0_66824, plain, (v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_66066, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k4_xcmplx_0(np__1)=np__0|np__0!=np__1|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1)).#trainneg
cnf(i_0_69622, plain, (v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_67349, plain, (r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1)|~r1_int_1(np__2,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_67193, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_67362, plain, (r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1)|~r1_int_1(np__1,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_67168, plain, (k3_int_2(k4_xcmplx_0(np__1),np__2)=k3_int_2(np__1,k4_xcmplx_0(np__2))|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_67170, plain, (k3_int_2(k4_xcmplx_0(np__2),np__1)=k3_int_2(k4_xcmplx_0(np__1),np__2)|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_2393, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_67169, plain, (k3_int_2(k4_xcmplx_0(np__1),np__2)=np__1|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_68569, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v7_ordinal1(X1)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~v7_ordinal1(k4_xcmplx_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_70147, plain, (r1_int_1(k4_xcmplx_0(X1),k2_int_5(X2,X3))|~r1_int_1(X1,k2_int_5(X2,X3))|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_1(X1)|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_69161, plain, (k3_int_2(np__0,np__2)=k3_int_2(np__2,np__0)|v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_67177, plain, (k3_int_2(k4_xcmplx_0(np__1),np__2)=k3_int_2(np__2,np__1)|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_67351, plain, (r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1)|~r1_int_1(np__2,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2395, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(X3))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k4_xcmplx_0(X3)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_69160, plain, (k3_int_2(np__0,np__2)=k3_int_2(np__2,np__0)|m1_subset_1(k4_xcmplx_0(np__2),k5_numbers))).#trainneg
cnf(i_0_69167, plain, (k3_int_2(k4_xcmplx_0(np__2),np__0)=k3_int_2(np__0,k4_xcmplx_0(np__2))|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_67350, plain, (r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(np__2),X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_66814, plain, (v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(X1)))|k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))!=k4_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_69168, plain, (k3_int_2(k4_xcmplx_0(np__2),np__0)=k3_int_2(np__0,k4_xcmplx_0(np__2))|k18_complex1(np__2)=np__2)).#trainneg
cnf(i_0_69162, plain, (k3_int_2(np__0,np__2)=k3_int_2(np__2,np__0)|r1_xxreal_0(k4_xcmplx_0(np__2),np__1)|r1_xxreal_0(np__2,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_67361, plain, (r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1)|~r1_int_1(np__1,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_70152, plain, (r1_int_1(np__2,k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_70153, plain, (r1_int_1(np__1,k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_68301, negated_conjecture, (r1_int_1(k4_xcmplx_0(np__1),X2)|k2_int_5(X1,esk2_0)=np__1|k2_int_5(X1,esk2_0)=np__0|~r1_int_1(np__1,X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_68131, plain, (~r1_int_1(k2_xcmplx_0(np__1,X1),k4_xcmplx_0(np__1))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k2_xcmplx_0(np__1,X1))|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2413, plain, (r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_67363, plain, (r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(np__1),X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_68462, plain, (~r1_int_1(k2_xcmplx_0(np__1,X1),np__1)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k2_xcmplx_0(np__1,X1))|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2415, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(X3))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(X3))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_68468, plain, (~r1_int_1(k1_nat_1(np__1,X1),np__1)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k1_nat_1(np__1,X1))|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_68132, plain, (~r1_int_1(k1_nat_1(np__1,X1),k4_xcmplx_0(np__1))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k1_nat_1(np__1,X1))|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2416, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k4_xcmplx_0(X3)))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(X3))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_26602, plain, (k3_int_2(k4_xcmplx_0(np__1),k1_int_2(X1))=k3_int_2(np__1,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_67404, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_69182, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(X2,k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)),np__1)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_68129, plain, (~r1_int_1(k1_nat_1(np__1,X1),k4_xcmplx_0(np__1))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k1_nat_1(np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_65437, plain, (k6_xcmplx_0(np__0,X1)=k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_35813, negated_conjecture, (r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_70121, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_67403, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_68465, plain, (~r1_int_1(k1_nat_1(np__1,X1),np__1)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k1_nat_1(np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_69612, plain, (k6_xcmplx_0(np__0,X1)=k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_68553, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~v7_ordinal1(k4_xcmplx_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_68461, plain, (~r1_int_1(k2_xcmplx_0(np__1,X1),np__1)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k2_xcmplx_0(np__1,X1))|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_68471, plain, (~r1_int_1(k2_xcmplx_0(np__1,X1),k4_xcmplx_0(np__1))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k2_xcmplx_0(np__1,X1))|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_70142, plain, (r1_int_1(k4_xcmplx_0(X1),k1_newton(X2,X3))|~r1_int_1(X1,k1_newton(X2,X3))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(X1)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_5871, plain, (k3_int_2(k4_xcmplx_0(np__1),X1)=k3_int_2(X1,np__1)|k4_xcmplx_0(np__1)=np__1|~v7_ordinal1(X1))).#trainneg
cnf(i_0_72932, plain, (r1_int_1(X1,X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_72024, plain, (r1_int_1(np__1,X1)|~r1_int_1(k4_xcmplx_0(np__1),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_71339, plain, (r1_int_1(np__2,X1)|~r1_int_1(k4_xcmplx_0(np__2),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_72918, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_70740, plain, (k3_int_2(np__1,k4_xcmplx_0(np__2))=np__1|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_70847, plain, (k3_int_2(np__2,k4_xcmplx_0(np__1))=np__1|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_2454, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k4_xcmplx_0(X3)))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(X3))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_71261, plain, (k3_int_2(np__0,k4_xcmplx_0(np__2))=k3_int_2(np__2,np__0)|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_71263, plain, (k3_int_2(np__0,k4_xcmplx_0(np__2))=k3_int_2(np__0,np__2)|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_71091, plain, (k3_int_2(np__2,np__1)=k3_int_2(np__1,np__2)|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_70739, plain, (k3_int_2(np__1,k4_xcmplx_0(np__2))=k3_int_2(np__1,np__2)|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_71084, plain, (k3_int_2(np__2,k4_xcmplx_0(np__1))=k3_int_2(np__2,np__1)|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_70174, plain, (v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(X1)))|v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_70732, plain, (k3_int_2(np__2,k4_xcmplx_0(np__1))=k3_int_2(np__1,k4_xcmplx_0(np__2))|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_72194, plain, (~r1_int_1(k1_nat_1(X1,np__1),np__1)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k1_nat_1(X1,np__1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_33362, plain, (r1_xxreal_0(X1,k4_xcmplx_0(np__1))|r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_72198, plain, (~r1_int_1(k1_nat_1(X1,np__1),k4_xcmplx_0(np__1))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k1_nat_1(X1,np__1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_33631, plain, (r1_xxreal_0(X1,k4_xcmplx_0(np__2))|r1_xxreal_0(k4_xcmplx_0(np__2),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_71769, negated_conjecture, (k2_int_5(X1,esk2_0)=np__0|k2_int_5(X1,esk2_0)=np__1|r1_int_1(k4_xcmplx_0(np__1),X2)|~r1_int_1(np__1,X2)|~v1_int_1(X1)|~v1_xcmplx_0(X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2496, plain, (v1_xcmplx_0(k6_xcmplx_0(X1,X2))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_2497, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_3995, plain, (v1_xreal_0(k4_xcmplx_0(X1))|m1_subset_1(X1,k5_numbers)|~v3_membered(k5_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_72027, plain, (~r1_int_1(k2_xcmplx_0(X1,np__1),np__1)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_10758, plain, (r1_xxreal_0(np__2,X1)|r1_xxreal_0(X1,np__1)|v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_2499, plain, (v1_xcmplx_0(k6_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_4449, plain, (v7_ordinal1(k4_xcmplx_0(k1_int_2(X1)))|k4_xcmplx_0(k4_xcmplx_0(k1_int_2(X1)))!=k4_xcmplx_0(k1_int_2(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_69591, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))))|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_71772, plain, (~r1_int_1(k2_xcmplx_0(X1,np__1),k4_xcmplx_0(np__1))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2500, plain, (k6_xcmplx_0(k6_xcmplx_0(X1,X2),k4_xcmplx_0(X3))=k6_xcmplx_0(X3,k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_2501, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)),X3)|~r1_int_1(k6_xcmplx_0(X1,X2),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_2502, plain, (v1_xcmplx_0(k6_xcmplx_0(X1,X2))|~v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_50754, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__1)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_69613, plain, (k6_xcmplx_0(np__0,X1)=k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2505, plain, (r1_int_1(k6_xcmplx_0(X1,X2),X3)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_72916, plain, (r1_int_1(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|~r1_int_1(X1,k2_int_5(X2,X3))|~v7_ordinal1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_1(X1)|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_2507, plain, (r1_int_1(k6_xcmplx_0(X1,X2),k4_xcmplx_0(X3))|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_2508, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)),X3)|~r1_int_1(k6_xcmplx_0(X1,X2),k4_xcmplx_0(X3))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_2509, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)),k4_xcmplx_0(X3))|~r1_int_1(k6_xcmplx_0(X1,X2),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_2510, plain, (r1_int_1(k6_xcmplx_0(X1,X2),X3)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)),k4_xcmplx_0(X3))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_73696, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~v7_ordinal1(X1)|~r1_int_1(np__1,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2524, plain, (r1_int_1(k6_xcmplx_0(X1,X2),k1_newton(X3,X4))|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)),k4_xcmplx_0(k1_newton(X3,X4)))|~v7_ordinal1(X4)|~v1_int_1(k4_xcmplx_0(k1_newton(X3,X4)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v1_xcmplx_0(X3)|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_73743, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(X1))).#trainneg
cnf(i_0_74649, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__2)|r1_xxreal_0(np__2,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_2525, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)),k4_xcmplx_0(k1_newton(X3,X4)))|~r1_int_1(k6_xcmplx_0(X1,X2),k1_newton(X3,X4))|~v7_ordinal1(X4)|~v1_int_1(k4_xcmplx_0(k1_newton(X3,X4)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v1_xcmplx_0(X3)|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_73686, plain, (r1_int_1(k4_xcmplx_0(np__2),X1)|~v7_ordinal1(X1)|~r1_int_1(np__2,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_74651, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__1)|r1_xxreal_0(np__1,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_74673, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),np__2)|r1_xxreal_0(np__2,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_2535, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k6_xcmplx_0(X2,X3))=k6_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(X2,X3)),X1)|~v1_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X1)|~v1_int_1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_74675, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),np__1)|r1_xxreal_0(np__1,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_74548, plain, (k3_int_2(np__0,np__2)=k3_int_2(np__2,np__0)|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_74650, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__0)|r1_xxreal_0(np__0,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_74674, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),np__0)|r1_xxreal_0(np__0,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_2536, plain, (r1_int_1(X1,k4_xcmplx_0(k6_xcmplx_0(X2,X3)))|~r1_int_1(X1,k6_xcmplx_0(X2,X3))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(X2,X3)))|~v1_int_1(X1)|~v1_int_1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_2537, plain, (r1_int_1(X1,k6_xcmplx_0(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k6_xcmplx_0(X2,X3)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(X2,X3)))|~v1_int_1(X1)|~v1_int_1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_74844, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_int_1(X1)|~v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2538, plain, (r1_int_1(k4_xcmplx_0(X1),k6_xcmplx_0(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k6_xcmplx_0(X2,X3)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(X2,X3)))|~v1_int_1(X1)|~v1_int_1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_74364, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2539, plain, (r1_int_1(X1,k4_xcmplx_0(k6_xcmplx_0(X2,X3)))|~r1_int_1(k4_xcmplx_0(X1),k6_xcmplx_0(X2,X3))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(X2,X3)))|~v1_int_1(X1)|~v1_int_1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_2552, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k6_xcmplx_0(X3,X4)))|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k6_xcmplx_0(X3,X4))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(X3,X4)))|~v1_xcmplx_0(X1)|~v1_int_1(X4)|~v1_int_1(X3))).#trainneg
cnf(i_0_77594, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~m1_subset_1(X1,k5_numbers)|~r1_int_1(np__2,k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_78436, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_int_1(X1)|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_2615, plain, (r1_xxreal_0(k5_xcmplx_0(X1),np__0)|r1_xxreal_0(k4_xcmplx_0(X2),k5_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(k5_xcmplx_0(X1))|~r1_int_1(X2,k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_78437, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_2616, plain, (r1_xxreal_0(k1_newton(X1,X2),np__0)|r1_xxreal_0(k4_xcmplx_0(X3),k1_newton(X1,X2))|~v7_ordinal1(k4_xcmplx_0(X3))|~v7_ordinal1(k1_newton(X1,X2))|~r1_int_1(X3,k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_77437, plain, (r1_int_1(np__2,X1)|~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_77438, plain, (r1_int_1(np__1,X1)|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2618, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__0)|r1_xxreal_0(k4_xcmplx_0(k5_xcmplx_0(X2)),k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(k5_xcmplx_0(X2),X1)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2619, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__0)|r1_xxreal_0(k4_xcmplx_0(k1_newton(X2,X3)),k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(k1_newton(X2,X3),X1)|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_35816, negated_conjecture, (r1_int_1(esk2_0,k6_xcmplx_0(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1))|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_74546, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2620, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__0)|r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X2,X1)|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_36727, negated_conjecture, (np__0=np__1|r1_int_1(esk2_0,k4_xcmplx_0(np__1))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r1_int_1(esk2_0,k4_xcmplx_0(np__2))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__0)).#trainneg
cnf(i_0_73665, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_2623, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(k4_xcmplx_0(k5_xcmplx_0(X2)),X1)|~v7_ordinal1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v7_ordinal1(X1)|~r1_int_1(k5_xcmplx_0(X2),X1)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2624, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(k4_xcmplx_0(k5_xcmplx_0(X2)),X1)|~v7_ordinal1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v7_ordinal1(X1)|~r1_int_1(k5_xcmplx_0(X2),k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_74276, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2625, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(k4_xcmplx_0(k1_newton(X2,X3)),X1)|~v7_ordinal1(k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X1)|~r1_int_1(k1_newton(X2,X3),X1)|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_15730, plain, (r1_int_1(np__0,X1)|~r1_int_1(np__0,k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2626, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(k4_xcmplx_0(k1_newton(X2,X3)),X1)|~v7_ordinal1(k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X1)|~r1_int_1(k1_newton(X2,X3),k4_xcmplx_0(X1))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_36201, negated_conjecture, (r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|r2_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1),esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).#trainneg
cnf(i_0_56129, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r1_int_1(esk2_0,k6_xcmplx_0(np__1,np__1))|k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))=np__1|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).#trainneg
cnf(i_0_77711, plain, (r1_int_1(np__2,X1)|~v7_ordinal1(X1)|~r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2627, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~r1_int_1(X2,X1)|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_23265, plain, (r1_int_1(np__2,k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__2),X1)|~v7_ordinal1(k4_xcmplx_0(np__2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_23266, plain, (r1_int_1(np__1,k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__1),X1)|~v7_ordinal1(k4_xcmplx_0(np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_64373, plain, (~r1_int_1(k4_xcmplx_0(k1_nat_1(np__1,X1)),k4_xcmplx_0(np__1))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k1_nat_1(np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_74199, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_2628, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~r1_int_1(X2,k4_xcmplx_0(X1))|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_4558, plain, (r1_int_1(k4_xcmplx_0(np__2),X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(np__2,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2629, plain, (r1_xxreal_0(k6_xcmplx_0(k2_int_5(X1,X2),k1_newton(X1,k3_nat_d(k7_nat_d(X2,np__1),np__2))),np__0)|r1_xxreal_0(X2,k6_xcmplx_0(k2_int_5(X1,X2),k1_newton(X1,k3_nat_d(k7_nat_d(X2,np__1),np__2))))|r1_xxreal_0(X2,np__2)|~v7_ordinal1(X2)|~v7_ordinal1(k6_xcmplx_0(k2_int_5(X1,X2),k1_newton(X1,k3_nat_d(k7_nat_d(X2,np__1),np__2))))|k3_int_2(X1,X2)!=np__1|~v1_int_2(X2)|~v1_int_1(k1_newton(X1,k3_nat_d(k7_nat_d(X2,np__1),np__2)))|~v1_int_1(X1))).#trainneg
cnf(i_0_5971, plain, (v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1)|~m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))),k1_numbers))).#trainneg
cnf(i_0_73593, plain, (r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v7_ordinal1(X2)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~r1_int_1(np__1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2632, plain, (r1_xxreal_0(k4_xcmplx_0(k5_xcmplx_0(X1)),np__0)|r1_xxreal_0(X2,k4_xcmplx_0(k5_xcmplx_0(X1)))|~v7_ordinal1(X2)|~v7_ordinal1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~r1_int_1(X2,k5_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_71212, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_70775, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X2))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4597, plain, (r1_int_1(k4_xcmplx_0(X1),np__2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2633, plain, (r1_xxreal_0(k4_xcmplx_0(k1_newton(X1,X2)),np__0)|r1_xxreal_0(X3,k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X3)|~v7_ordinal1(k4_xcmplx_0(k1_newton(X1,X2)))|~r1_int_1(X3,k1_newton(X1,X2))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_79885, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2634, plain, (r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)),np__0)|r1_xxreal_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X2)|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~r1_int_1(k4_xcmplx_0(X2),X1)|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_2635, plain, (r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)),np__0)|r1_xxreal_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X2)|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~r1_int_1(X2,X1)|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_80414, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_79183, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(np__2,k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_2636, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__0)|r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X2,X1)|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_81483, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2638, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__0)|r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X2,X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_81452, plain, (r1_xxreal_0(X1,np__1)|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2704, plain, (v1_xcmplx_0(k1_newton(X1,X2))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_79265, plain, (v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v1_int_1(X1))).#trainneg
cnf(i_0_80622, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,np__0)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2705, plain, (v1_xreal_0(k1_newton(X1,X2))|~v1_xreal_0(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_79918, negated_conjecture, (r2_int_1(np__1,k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)),esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__0)).#trainneg
cnf(i_0_80422, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(np__0,np__1,esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__1|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).#trainneg
cnf(i_0_80425, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(np__0,np__1,esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).#trainneg
cnf(i_0_80423, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(np__0,k4_xcmplx_0(np__1),esk2_0)|k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))=np__1|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).#trainneg
cnf(i_0_80424, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(np__0,np__1,esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|r2_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1),esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).#trainneg
cnf(i_0_80522, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|r2_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1),esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).#trainneg
cnf(i_0_80519, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|r2_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1),esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_80525, negated_conjecture, (k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|r1_int_1(esk2_0,k6_xcmplx_0(np__1,np__1))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers)|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_80520, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|r2_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1),esk2_0)|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_2707, plain, (v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_80521, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|r2_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1),esk2_0)|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))|~v1_int_1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))).#trainneg
cnf(i_0_2708, plain, (k6_xcmplx_0(k1_newton(X1,X2),k4_xcmplx_0(X3))=k6_xcmplx_0(X3,k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_2709, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~r1_int_1(k1_newton(X1,X2),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_81149, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,np__0)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2710, plain, (v1_xcmplx_0(k1_newton(X1,X2))|~v7_ordinal1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_83230, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r1_int_1(esk2_0,k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|~v1_int_1(esk2_0)|~v1_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))).#trainneg
cnf(i_0_83226, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|~v1_int_1(esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).#trainneg
cnf(i_0_82162, plain, (r1_xxreal_0(X1,np__2)|~r1_int_1(X1,np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2715, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(X3))|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_82814, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v1_int_1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_81828, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~r1_int_1(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_82788, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2716, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(X3))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_81840, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v1_int_1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_83006, plain, (r1_xxreal_0(X1,np__2)|~v7_ordinal1(X1)|~r1_int_1(X1,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_2717, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k4_xcmplx_0(X3))|~r1_int_1(k1_newton(X1,X2),X3)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_82163, plain, (r1_xxreal_0(X1,np__1)|~r1_int_1(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_83174, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_int_1(X2,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2718, plain, (r1_int_1(k1_newton(X1,X2),X3)|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k4_xcmplx_0(X3))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_84164, plain, (r1_int_1(esk2_0,k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|~v1_int_1(esk2_0))).#trainneg
cnf(i_0_84192, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|~v1_int_1(esk2_0)|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).#trainneg
cnf(i_0_2719, plain, (r1_int_1(np__0,k1_newton(X1,X2))|~r1_int_1(np__0,k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_84189, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|~v1_int_1(esk2_0)|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_84822, plain, (r1_int_1(np__1,k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~r1_int_1(np__1,X1))).#trainneg
cnf(i_0_85098, plain, (r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_2732, plain, (r1_int_1(k1_newton(X1,X2),k1_newton(X3,X4))|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k4_xcmplx_0(k1_newton(X3,X4)))|~v7_ordinal1(X4)|~v1_int_1(k4_xcmplx_0(k1_newton(X3,X4)))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(X3)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_2733, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k4_xcmplx_0(k1_newton(X3,X4)))|~r1_int_1(k1_newton(X1,X2),k1_newton(X3,X4))|~v7_ordinal1(X4)|~v1_int_1(k4_xcmplx_0(k1_newton(X3,X4)))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(X3)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_85929, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__2)|~v7_ordinal1(k4_xcmplx_0(np__1))|~r1_int_1(np__1,np__2))).#trainneg
cnf(i_0_81822, plain, (r1_xxreal_0(X2,np__1)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(X2,k1_newton(k4_xcmplx_0(np__1),X1))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2736, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k4_xcmplx_0(k5_xcmplx_0(X3)))|~r1_int_1(k1_newton(X1,X2),k5_xcmplx_0(X3))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X3)))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xreal_0(X3)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_83217, negated_conjecture, (r2_int_1(np__0,np__1,esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).#trainneg
cnf(i_0_83221, negated_conjecture, (r2_int_1(np__0,np__1,esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).#trainneg
cnf(i_0_85056, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_85053, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_81827, plain, (r1_int_1(k4_xcmplx_0(np__2),X1)|~r1_int_1(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2744, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k1_newton(X2,X3))=k6_xcmplx_0(k4_xcmplx_0(k1_newton(X2,X3)),X1)|~v1_xcmplx_0(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_83213, negated_conjecture, (r2_int_1(np__1,np__1,esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).#trainneg
cnf(i_0_84190, negated_conjecture, (v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|~v1_int_1(esk2_0)|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_2745, plain, (r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(X1,k1_newton(X2,X3))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(X1)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_83039, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),np__0)|r1_xxreal_0(np__2,k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(np__2),np__2))).#trainneg
cnf(i_0_83040, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),np__0)|r1_xxreal_0(np__2,k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(np__2))|~r1_int_1(np__2,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_84199, plain, (~r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1)|~r1_int_1(k1_nat_1(X1,np__1),np__2)|~v7_ordinal1(k1_nat_1(X1,np__1)))).#trainneg
cnf(i_0_84197, plain, (~r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1)|~r1_int_1(k2_xcmplx_0(X1,np__1),np__2)|~v7_ordinal1(k2_xcmplx_0(X1,np__1)))).#trainneg
cnf(i_0_2747, plain, (r1_int_1(k4_xcmplx_0(X1),k1_newton(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(X1)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_83214, negated_conjecture, (r2_int_1(np__0,np__1,esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__1|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_83218, negated_conjecture, (r2_int_1(np__0,np__1,esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_2748, plain, (r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(k4_xcmplx_0(X1),k1_newton(X2,X3))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(X1)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_84191, negated_conjecture, (v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|~v1_int_1(esk2_0)|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))|~v1_int_1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))).#trainneg
cnf(i_0_83215, negated_conjecture, (r2_int_1(np__0,np__1,esk2_0)|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__1|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_83219, negated_conjecture, (r2_int_1(np__0,np__1,esk2_0)|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_2752, plain, (r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k1_newton(X2,X3))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_xreal_0(X1)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_83216, negated_conjecture, (r2_int_1(np__0,np__1,esk2_0)|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))|~v1_int_1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))).#trainneg
cnf(i_0_83220, negated_conjecture, (r2_int_1(np__0,np__1,esk2_0)|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))|~v1_int_1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))).#trainneg
cnf(i_0_83353, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(np__0,np__1,esk2_0)|r2_int_1(np__0,k4_xcmplx_0(np__1),esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).#trainneg
cnf(i_0_83211, negated_conjecture, (r2_int_1(np__1,k4_xcmplx_0(np__1),esk2_0)|k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))=np__1|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__0|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).#trainneg
cnf(i_0_2753, plain, (r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(k4_xcmplx_0(X1),k1_newton(X2,X3))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_86209, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|~v1_int_1(esk2_0))).#trainneg
cnf(i_0_86949, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_87670, plain, (r1_int_1(np__2,k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~r1_int_1(np__2,X1))).#trainneg
cnf(i_0_86051, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|~v1_int_1(esk2_0)|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_86052, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|~v1_int_1(esk2_0)|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_87508, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1)|~r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2758, plain, (r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(k4_xcmplx_0(X1),k1_newton(X2,X3))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_2760, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k1_newton(X3,X4))|~r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k1_newton(X3,X4)))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(k1_newton(X3,X4)))|~v1_xcmplx_0(X1)|~v7_ordinal1(X4)|~v1_int_1(X3))).#trainneg
cnf(i_0_2761, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k1_newton(X3,X4)))|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k1_newton(X3,X4))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(k1_newton(X3,X4)))|~v1_xcmplx_0(X1)|~v7_ordinal1(X4)|~v1_int_1(X3))).#trainneg
cnf(i_0_87807, negated_conjecture, (r2_int_1(np__1,np__1,esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__0|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).#trainneg
cnf(i_0_88119, negated_conjecture, (r2_int_1(np__0,np__1,esk2_0)|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_87993, plain, (~r1_int_1(k2_xcmplx_0(np__1,X1),np__2)|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(k2_xcmplx_0(np__1,X1))|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_86518, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~r1_int_1(X1,X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2818, plain, (r1_xxreal_0(k7_nat_d(X1,X2),np__0)|r1_xxreal_0(np__1,k7_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_88484, negated_conjecture, (r2_int_1(np__1,k4_xcmplx_0(np__1),esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=np__1|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__0|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).#trainneg
cnf(i_0_88121, negated_conjecture, (r2_int_1(np__0,np__1,esk2_0)|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_87804, negated_conjecture, (r2_int_1(np__1,np__1,esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__0|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_86447, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_86446, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2819, plain, (r1_xxreal_0(k3_nat_d(X1,X2),np__0)|r1_xxreal_0(np__1,k3_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_87805, negated_conjecture, (r2_int_1(np__1,np__1,esk2_0)|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__0|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_86053, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|~v1_int_1(esk2_0)|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))|~v1_int_1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))).#trainneg
cnf(i_0_87994, plain, (~r1_int_1(k1_nat_1(np__1,X1),np__2)|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(k1_nat_1(np__1,X1))|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_87682, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1)|~r1_int_1(np__2,X1))).#trainneg
cnf(i_0_2820, plain, (r1_xxreal_0(k1_int_2(X1),np__0)|r1_xxreal_0(np__1,k1_int_2(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_87806, negated_conjecture, (r2_int_1(np__1,np__1,esk2_0)|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__0|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))|~v1_int_1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))).#trainneg
cnf(i_0_88347, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(np__0,k4_xcmplx_0(np__1),esk2_0)|r2_int_1(np__0,np__1,esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).#trainneg
cnf(i_0_88344, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(np__0,k4_xcmplx_0(np__1),esk2_0)|r2_int_1(np__0,np__1,esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_83342, negated_conjecture, (r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(np__0,k4_xcmplx_0(np__1),esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).#trainneg
cnf(i_0_83232, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r1_int_1(esk2_0,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|~v1_int_1(esk2_0)|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))).#trainneg
cnf(i_0_83233, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r1_int_1(esk2_0,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|~v1_int_1(esk2_0)|~v1_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))).#trainneg
cnf(i_0_89137, plain, (r1_int_1(np__1,k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(np__1),X1))).#trainneg
cnf(i_0_89123, plain, (r1_int_1(np__2,k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(np__2),X1))).#trainneg
cnf(i_0_90175, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_90491, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_90618, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__0)|~v1_int_1(X1)|k4_xcmplx_0(X1)!=X1)).#trainneg
cnf(i_0_2835, plain, (r1_int_1(X1,k6_xcmplx_0(k2_int_5(X2,k4_xcmplx_0(X1)),k1_newton(X2,k3_nat_d(k7_nat_d(k4_xcmplx_0(X1),np__1),np__2))))|r1_xxreal_0(k4_xcmplx_0(X1),np__2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k6_xcmplx_0(k2_int_5(X2,k4_xcmplx_0(X1)),k1_newton(X2,k3_nat_d(k7_nat_d(k4_xcmplx_0(X1),np__1),np__2))))|~v1_xcmplx_0(X1)|k3_int_2(X2,k4_xcmplx_0(X1))!=np__1|~v1_int_2(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(k1_newton(X2,k3_nat_d(k7_nat_d(k4_xcmplx_0(X1),np__1),np__2)))|~v1_int_1(X2))).#trainneg
cnf(i_0_89281, plain, (r1_int_1(X1,X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_90183, plain, (r1_int_1(k4_xcmplx_0(X1),np__2)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2840, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(X2))).#trainneg
cnf(i_0_89852, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_90508, plain, (~r1_int_1(k1_nat_1(X1,np__1),np__2)|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(k1_nat_1(X1,np__1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_90619, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__0)|m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_89856, negated_conjecture, (r2_int_1(np__1,k4_xcmplx_0(np__1),esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__0|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).#trainneg
cnf(i_0_90610, plain, (r1_xxreal_0(np__1,k4_xcmplx_0(np__2))|r1_xxreal_0(k4_xcmplx_0(np__2),np__0)|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_89532, plain, (~r1_int_1(k2_xcmplx_0(X1,np__1),np__2)|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__1))|~m1_subset_1(X1,k5_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_89242, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_89533, plain, (~r1_int_1(k2_xcmplx_0(X1,np__1),np__2)|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_90611, plain, (r1_xxreal_0(np__1,k4_xcmplx_0(np__2))|r1_xxreal_0(k4_xcmplx_0(np__2),np__0)|k18_complex1(np__2)=np__2)).#trainneg
cnf(i_0_89853, negated_conjecture, (r2_int_1(np__1,k4_xcmplx_0(np__1),esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__0|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_89858, negated_conjecture, (r2_int_1(np__1,np__1,esk2_0)|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=k4_xcmplx_0(np__1)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__0|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_89854, negated_conjecture, (r2_int_1(np__1,k4_xcmplx_0(np__1),esk2_0)|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__0|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_89855, negated_conjecture, (r2_int_1(np__1,k4_xcmplx_0(np__1),esk2_0)|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__0|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))|~v1_int_1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))).#trainneg
cnf(i_0_90637, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(np__0,k4_xcmplx_0(np__1),esk2_0)|r2_int_1(np__0,np__1,esk2_0)|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_87991, plain, (~r1_int_1(k1_nat_1(np__1,X1),np__2)|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(k1_nat_1(np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_90642, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|r2_int_1(np__0,k4_xcmplx_0(np__1),esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).#trainneg
cnf(i_0_90639, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|r2_int_1(np__0,k4_xcmplx_0(np__1),esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_83482, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(np__0,np__1,esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|r1_int_1(esk2_0,k6_xcmplx_0(np__1,np__1))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).#trainneg
cnf(i_0_83359, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1),esk2_0)|r2_int_1(np__0,np__1,esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).#trainneg
cnf(i_0_90620, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__0)|v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_88345, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(np__0,k4_xcmplx_0(np__1),esk2_0)|r2_int_1(np__0,np__1,esk2_0)|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_88346, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(np__0,k4_xcmplx_0(np__1),esk2_0)|r2_int_1(np__0,np__1,esk2_0)|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))|~v1_int_1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))).#trainneg
cnf(i_0_83471, negated_conjecture, (r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r1_int_1(esk2_0,k6_xcmplx_0(np__1,np__1))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~m1_subset_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2),k5_numbers))).#trainneg
cnf(i_0_83117, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1)|~r1_int_1(np__2,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_83356, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1),esk2_0)|r2_int_1(np__0,np__1,esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_90640, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|r2_int_1(np__0,k4_xcmplx_0(np__1),esk2_0)|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_90641, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|r2_int_1(np__0,k4_xcmplx_0(np__1),esk2_0)|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))|~v1_int_1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))).#trainneg
cnf(i_0_83362, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|r2_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1),esk2_0)|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_90186, plain, (r1_int_1(k4_xcmplx_0(X1),np__1)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_83357, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1),esk2_0)|r2_int_1(np__0,np__1,esk2_0)|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_83358, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1),esk2_0)|r2_int_1(np__0,np__1,esk2_0)|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))|~v1_int_1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))).#trainneg
cnf(i_0_88656, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7673, plain, (v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))))|~v1_int_1(X1))).#trainneg
cnf(i_0_85627, plain, (~r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1)|~r1_int_1(k2_xcmplx_0(X1,np__1),k4_xcmplx_0(np__2))|~v7_ordinal1(k2_xcmplx_0(X1,np__1)))).#trainneg
cnf(i_0_85629, plain, (~r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1)|~r1_int_1(k1_nat_1(X1,np__1),k4_xcmplx_0(np__2))|~v7_ordinal1(k1_nat_1(X1,np__1)))).#trainneg
cnf(i_0_88856, plain, (r1_int_1(k4_xcmplx_0(X1),np__1)|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|~v7_ordinal1(X1)|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_93057, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2894, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),k5_xcmplx_0(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2896, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(X2))).#trainneg
cnf(i_0_91498, plain, (r1_int_1(np__2,X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(np__2,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2897, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X2))).#trainneg
cnf(i_0_91563, plain, (r1_int_1(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_xcmplx_0(X2)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1))).#trainneg
cnf(i_0_2930, plain, (r1_int_1(k4_xcmplx_0(X1),k5_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~r1_int_1(X1,k4_xcmplx_0(k5_xcmplx_0(X2)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2931, plain, (r1_int_1(k4_xcmplx_0(X1),k1_newton(X2,X3))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_93069, negated_conjecture, (r2_int_1(np__1,k4_xcmplx_0(np__1),esk2_0)|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__0|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_2932, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2))))).#trainneg
cnf(i_0_91503, plain, (r1_int_1(np__1,X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(np__1,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_38593, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_92724, plain, (k3_int_2(np__2,np__1)=k3_int_2(np__1,np__2)|r1_xxreal_0(k4_xcmplx_0(np__2),np__0)|r1_xxreal_0(np__1,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_92725, plain, (k3_int_2(np__0,np__2)=k3_int_2(np__2,np__0)|r1_xxreal_0(k4_xcmplx_0(np__2),np__0)|r1_xxreal_0(np__1,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_2937, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~r1_int_1(X2,X1)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1))).#trainneg
cnf(i_0_92713, plain, (k3_int_2(k4_xcmplx_0(np__1),np__2)=k3_int_2(np__1,np__2)|r1_xxreal_0(k4_xcmplx_0(np__2),np__0)|r1_xxreal_0(np__1,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_92716, plain, (k3_int_2(np__2,k4_xcmplx_0(np__1))=k3_int_2(np__2,np__1)|r1_xxreal_0(k4_xcmplx_0(np__2),np__0)|r1_xxreal_0(np__1,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_92723, plain, (k3_int_2(np__2,k4_xcmplx_0(np__1))=k3_int_2(np__1,np__2)|r1_xxreal_0(k4_xcmplx_0(np__2),np__0)|r1_xxreal_0(np__1,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_51213, plain, (r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))|~r1_int_1(np__1,X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3374, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__1)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,np__1)|~v1_int_1(X1))).#trainneg
cnf(i_0_6320, plain, (v1_xreal_0(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1)|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1)))))).#trainneg
cnf(i_0_92726, plain, (k3_int_2(k4_xcmplx_0(np__1),np__2)=k3_int_2(np__2,np__1)|r1_xxreal_0(k4_xcmplx_0(np__2),np__0)|r1_xxreal_0(np__1,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_7666, plain, (v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_93100, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|r2_int_1(np__0,k4_xcmplx_0(np__1),esk2_0)|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_93105, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(np__0,np__1,esk2_0)|r1_int_1(esk2_0,k6_xcmplx_0(np__1,np__1))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).#trainneg
cnf(i_0_88841, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(X1,X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_80340, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_93102, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(np__0,np__1,esk2_0)|r1_int_1(esk2_0,k6_xcmplx_0(np__1,np__1))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_57508, plain, (r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(X1))|~r1_int_1(np__2,X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_95471, plain, (r1_int_1(X1,X2)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_93967, plain, (r1_int_1(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2))))).#trainneg
cnf(i_0_95857, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(np__1,X1))).#trainneg
cnf(i_0_95111, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~r1_int_1(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_95459, plain, (r1_int_1(np__1,X1)|~r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_96038, plain, (v1_int_1(k5_xcmplx_0(X1))|~v1_xreal_0(X1)|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1)))))).#trainneg
cnf(i_0_95738, plain, (k3_int_2(np__1,np__2)=np__1|r1_xxreal_0(np__1,k4_xcmplx_0(np__2))|r1_xxreal_0(k4_xcmplx_0(np__2),np__0))).#trainneg
cnf(i_0_95786, plain, (k3_int_2(np__2,np__1)=np__1|r1_xxreal_0(np__1,k4_xcmplx_0(np__2))|r1_xxreal_0(k4_xcmplx_0(np__2),np__0))).#trainneg
cnf(i_0_95233, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_3052, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k4_xcmplx_0(X2))|~r1_int_1(k5_xcmplx_0(X1),X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_95442, plain, (r1_int_1(np__0,X1)|~r1_int_1(np__0,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_3054, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k4_xcmplx_0(X3))|~r1_int_1(k1_newton(X1,X2),X3)|~v1_int_1(k4_xcmplx_0(X3))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(X3)|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3055, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k4_xcmplx_0(X3))|~r1_int_1(k1_newton(X1,X2),X3)|~v1_int_1(k4_xcmplx_0(X3))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(X3)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_3056, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xcmplx_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_96559, plain, (r1_int_1(k4_xcmplx_0(np__2),X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(np__2,X1))).#trainneg
cnf(i_0_50989, plain, (k3_int_2(k4_xcmplx_0(np__2),np__1)=k3_int_2(np__2,np__1)|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_97850, plain, (r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~r1_int_1(X1,k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_3060, plain, (r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),k4_xcmplx_0(X3))|~r1_int_1(k2_int_5(X1,X2),X3)|~v1_int_1(k4_xcmplx_0(X3))|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_xcmplx_0(X3)|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_3067, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(X1)|~v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2)))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2))).#trainneg
cnf(i_0_3068, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(X1)|~v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2)))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2))).#trainneg
cnf(i_0_98254, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_98607, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_57837, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__2))|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3093, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__0)|r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X2,X1)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1))).#trainneg
cnf(i_0_83964, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(np__2),X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_96253, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_71496, plain, (r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(np__2,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3097, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_99778, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_99769, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~r1_int_1(X1,X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3129, plain, (r1_int_1(k4_xcmplx_0(X1),k1_newton(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_3130, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_96511, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|r2_int_1(np__0,np__1,esk2_0)|r1_int_1(esk2_0,k6_xcmplx_0(np__1,np__1))|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_93150, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|r1_int_1(esk2_0,k6_xcmplx_0(np__1,np__1))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k7_nat_d(esk2_0,np__1)))).#trainneg
cnf(i_0_71639, plain, (r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(np__1,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3131, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(X2)))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(X2))).#trainneg
cnf(i_0_93239, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|r2_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1),esk2_0)|r2_int_1(np__0,np__1,esk2_0)|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_3132, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(X2)))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_93147, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|r1_int_1(esk2_0,k6_xcmplx_0(np__1,np__1))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_93103, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(np__0,np__1,esk2_0)|r1_int_1(esk2_0,k6_xcmplx_0(np__1,np__1))|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_93104, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(np__0,np__1,esk2_0)|r1_int_1(esk2_0,k6_xcmplx_0(np__1,np__1))|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))|~v1_int_1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))).#trainneg
cnf(i_0_99941, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_99842, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~r1_int_1(X1,X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_93148, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|r1_int_1(esk2_0,k6_xcmplx_0(np__1,np__1))|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v1_int_1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_93149, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|r1_int_1(esk2_0,k6_xcmplx_0(np__1,np__1))|v7_ordinal1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))|~v1_int_1(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))))).#trainneg
cnf(i_0_93246, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)),np__1)|r1_xxreal_0(np__2,k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|r2_int_1(np__0,k4_xcmplx_0(np__1),esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_93247, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)),np__0)|r1_xxreal_0(np__1,k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|r2_int_1(np__0,k4_xcmplx_0(np__1),esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_93681, plain, (~r1_int_1(k1_nat_1(np__1,X1),k4_xcmplx_0(np__2))|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(k1_nat_1(np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_98826, plain, (r1_int_1(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|~r1_int_1(X1,k2_int_5(X2,X3))|~v7_ordinal1(X1)|~v1_int_1(k2_int_5(X2,X3))|~v1_xcmplx_0(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_102085, plain, (r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2))))).#trainneg
cnf(i_0_102013, plain, (r1_int_1(np__2,X1)|~v7_ordinal1(X1)|~r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_102027, plain, (r1_int_1(np__1,X1)|~v7_ordinal1(X1)|~r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_102637, negated_conjecture, (k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|r1_int_1(esk2_0,k6_xcmplx_0(np__1,np__1))|k4_xcmplx_0(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2))!=k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|~v7_ordinal1(k3_nat_d(k7_nat_d(esk2_0,np__1),np__2)))).#trainneg
cnf(i_0_102082, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(X1)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3151, plain, (r1_int_1(k4_xcmplx_0(X1),k5_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~r1_int_1(X1,k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_102108, plain, (r1_int_1(np__0,k4_xcmplx_0(X1))|~r1_int_1(np__0,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3152, plain, (r1_int_1(k4_xcmplx_0(X1),k1_newton(X2,X3))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3153, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_93661, plain, (~r1_int_1(k2_xcmplx_0(np__1,X1),k4_xcmplx_0(np__2))|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(k2_xcmplx_0(np__1,X1))|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4615, plain, (r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~r1_int_1(np__2,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_3154, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4644, plain, (r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~r1_int_1(np__1,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_3158, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~r1_int_1(X2,X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_102091, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_int_1(X2,k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_102080, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2))))).#trainneg
cnf(i_0_104506, plain, (~r1_xxreal_0(np__2,X1)|~v7_ordinal1(k2_xcmplx_0(np__1,X1))|~v7_ordinal1(X1)|~r1_int_1(k2_xcmplx_0(np__1,X1),np__2))).#trainneg
cnf(i_0_64375, plain, (~r1_int_1(k4_xcmplx_0(k2_xcmplx_0(np__1,X1)),k4_xcmplx_0(np__1))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k2_xcmplx_0(np__1,X1))|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_101112, plain, (r1_int_1(k4_xcmplx_0(X1),np__1)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5779, plain, (v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(X1)|~m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))),k1_numbers))).#trainneg
cnf(i_0_91414, plain, (r1_int_1(np__0,X1)|~r1_int_1(np__0,k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_64377, plain, (~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k1_nat_1(X1,np__1))|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(k1_nat_1(X1,np__1)),np__1)|~v1_int_1(k4_xcmplx_0(k1_nat_1(X1,np__1))))).#trainneg
cnf(i_0_64604, plain, (~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__1))|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(k2_xcmplx_0(X1,np__1)),np__1)|~v1_int_1(k4_xcmplx_0(k2_xcmplx_0(X1,np__1))))).#trainneg
cnf(i_0_3258, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k4_xcmplx_0(X2))|~r1_int_1(k5_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_80983, plain, (~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k1_nat_1(np__1,X1))|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(k1_nat_1(np__1,X1)),np__1)|~v1_int_1(k4_xcmplx_0(k1_nat_1(np__1,X1))))).#trainneg
cnf(i_0_81845, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),np__1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3260, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k4_xcmplx_0(X3))|~r1_int_1(k1_newton(X1,X2),X3)|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(X3))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6321, plain, (v1_xreal_0(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1)))))).#trainneg
cnf(i_0_5387, plain, (r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__1),X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_3261, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k4_xcmplx_0(X3))|~r1_int_1(k1_newton(X1,X2),X3)|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(X3))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_75059, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))))|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_5718, plain, (r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__2),X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_3262, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_81920, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_20916, plain, (r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__2),X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3266, plain, (r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),k4_xcmplx_0(X3))|~r1_int_1(k2_int_5(X1,X2),X3)|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(X3))|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_20917, plain, (r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__1),X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_92136, plain, (r1_int_1(X1,k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(X1,np__2))).#trainneg
cnf(i_0_88872, plain, (r1_int_1(np__1,k4_xcmplx_0(X2))|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_30418, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~r1_int_1(X1,np__2)|~v1_int_1(X1))).#trainneg
cnf(i_0_35818, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__1)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_40863, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__2)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(np__2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_102024, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~v1_int_1(X1)|~r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_30420, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~r1_int_1(X1,np__1)|~v1_int_1(X1))).#trainneg
cnf(i_0_96675, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v7_ordinal1(X1)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_38572, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__2))|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_107942, plain, (r1_xxreal_0(X1,np__2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__2))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_106203, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3301, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__0)|r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X2,X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_106674, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_106747, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_106445, plain, (r1_int_1(np__1,k4_xcmplx_0(X2))|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3305, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_108191, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v7_ordinal1(X1)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_105428, plain, (~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k2_xcmplx_0(np__1,X1))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k4_xcmplx_0(k2_xcmplx_0(np__1,X1)),np__1))).#trainneg
cnf(i_0_6283, plain, (r1_int_1(X1,np__2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__2))|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_105376, plain, (~r1_int_1(k4_xcmplx_0(k2_xcmplx_0(X1,np__1)),k4_xcmplx_0(np__1))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_8284, plain, (r1_int_1(X1,np__1)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_81838, plain, (r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),X3)|~r1_int_1(k2_int_5(X1,X2),X3)|~v7_ordinal1(X3)|~v1_int_1(k2_int_5(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_3338, plain, (r1_int_1(k4_xcmplx_0(X1),k1_newton(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_3339, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_108397, negated_conjecture, (~r1_int_1(k4_xcmplx_0(esk2_0),k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(esk2_0)))).#trainneg
cnf(i_0_108624, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(X1,X2))).#trainneg
cnf(i_0_3340, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(X2)))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_xcmplx_0(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_3341, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(k4_xcmplx_0(X2)))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_109664, plain, (~r1_int_1(k4_xcmplx_0(k2_xcmplx_0(X1,np__1)),np__1)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_109538, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_int_1(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3364, plain, (k3_int_2(k1_int_2(X1),X2)=k3_int_2(k1_int_2(X2),X1)|~v1_int_1(X1)|~v1_int_1(k1_int_2(X2))|~v1_int_1(X2)|~v1_int_1(k1_int_2(X1)))).#trainneg
cnf(i_0_108652, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3396, plain, (r1_int_1(k5_xcmplx_0(X1),np__1)|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),np__1)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3400, plain, (r1_int_1(k4_xcmplx_0(X1),np__1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3401, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))),np__1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v1_int_1(X1))).#trainneg
cnf(i_0_111551, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(X1),np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_112242, plain, (r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_3402, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))),np__1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3403, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),np__1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))),np__1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3409, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)),k1_newton(X2,X3))=k6_xcmplx_0(X1,k4_xcmplx_0(k4_xcmplx_0(k1_newton(X2,X3))))|~v7_ordinal1(X3)|~v1_xcmplx_0(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_111823, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X1)|~v1_int_1(X2)|~r1_int_1(X1,k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_111552, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(X1),np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3411, plain, (v1_int_1(k6_xcmplx_0(X2,k4_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3412, plain, (k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(X2,k4_xcmplx_0(X1))))=k6_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3416, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)),k1_newton(X2,X3))=k6_xcmplx_0(X1,k4_xcmplx_0(k4_xcmplx_0(k1_newton(X2,X3))))|~v7_ordinal1(X3)|~v1_xcmplx_0(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v7_ordinal1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3418, plain, (v1_int_1(k6_xcmplx_0(X2,k4_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v7_ordinal1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_110576, plain, (r1_int_1(np__2,k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_3419, plain, (k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(X2,k4_xcmplx_0(X1))))=k6_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v7_ordinal1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_110590, plain, (r1_int_1(np__1,k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_3421, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),k4_xcmplx_0(X1))=k6_xcmplx_0(X1,np__2)|~v7_ordinal1(k4_xcmplx_0(np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_111908, plain, (k3_int_2(k1_int_2(X1),k4_xcmplx_0(np__1))=k3_int_2(np__1,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_3422, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))=k6_xcmplx_0(X1,np__1)|~v7_ordinal1(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_110677, plain, (r1_int_1(np__2,k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_110654, plain, (r1_int_1(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(X1)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_110678, plain, (r1_int_1(np__1,k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_109024, plain, (r1_int_1(np__2,k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(np__2,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_111821, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X1)|~v1_int_1(X2)|~r1_int_1(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_3526, plain, (k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(X2,k4_xcmplx_0(X1))))=k6_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_109031, plain, (r1_int_1(np__1,k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(np__1,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_108937, plain, (r1_int_1(k4_xcmplx_0(X1),np__1)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3542, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3619, plain, (k1_int_2(k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|k18_complex1(X1)=X1|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_116995, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(X2)|~r1_int_1(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_3642, plain, (r1_int_1(k4_xcmplx_0(k1_int_2(X1)),X2)|~r1_int_1(k1_int_2(X1),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_3646, plain, (r1_int_1(k1_int_2(X1),X2)|~r1_int_1(k4_xcmplx_0(k1_int_2(X1)),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_3648, plain, (r1_int_1(k1_int_2(X1),k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(k1_int_2(X1)),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_3649, plain, (r1_int_1(k4_xcmplx_0(k1_int_2(X1)),X2)|~r1_int_1(k1_int_2(X1),k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_3650, plain, (r1_int_1(k4_xcmplx_0(k1_int_2(X1)),k4_xcmplx_0(X2))|~r1_int_1(k1_int_2(X1),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_3651, plain, (r1_int_1(k1_int_2(X1),X2)|~r1_int_1(k4_xcmplx_0(k1_int_2(X1)),k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_3653, plain, (r1_int_1(np__0,k1_int_2(X1))|~r1_int_1(np__0,k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_113722, plain, (v1_int_1(k6_xcmplx_0(X1,np__0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3655, plain, (r1_int_1(np__0,k4_xcmplx_0(k1_int_2(X1)))|~r1_int_1(np__0,k1_int_2(X1))|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_3656, plain, (r1_int_1(k4_xcmplx_0(k1_int_2(X1)),np__0)|~r1_int_1(k1_int_2(X1),np__0)|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_3657, plain, (r1_int_1(k1_int_2(X1),np__0)|~r1_int_1(k4_xcmplx_0(k1_int_2(X1)),np__0)|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_113744, plain, (v1_int_1(k6_xcmplx_0(X1,np__0))|~v1_int_1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_113417, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X1)|~v1_int_1(X2)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_108826, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3668, plain, (r1_int_1(k1_int_2(X1),k1_newton(X2,X3))|~r1_int_1(k4_xcmplx_0(k1_int_2(X1)),k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_xcmplx_0(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_3669, plain, (r1_int_1(k4_xcmplx_0(k1_int_2(X1)),k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(k1_int_2(X1),k1_newton(X2,X3))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_xcmplx_0(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_3672, plain, (r1_int_1(k4_xcmplx_0(k1_int_2(X1)),k4_xcmplx_0(k5_xcmplx_0(X2)))|~r1_int_1(k1_int_2(X1),k5_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_xreal_0(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_3680, plain, (r1_int_1(k4_xcmplx_0(k1_int_2(X1)),k4_xcmplx_0(X2))|~r1_int_1(k1_int_2(X1),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_120143, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_120147, plain, (r1_int_1(k4_xcmplx_0(X1),k5_xcmplx_0(X2))|~r1_int_1(X1,k5_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(k5_xcmplx_0(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_120127, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__2))|~r1_int_1(X1,k4_xcmplx_0(np__2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_118311, plain, (m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|k18_complex1(X1)=X1|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_120551, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_120146, plain, (r1_int_1(k4_xcmplx_0(X1),k7_xcmplx_0(X2,X3))|~r1_int_1(X1,k7_xcmplx_0(X2,X3))|~v7_ordinal1(X1)|~v1_int_1(k7_xcmplx_0(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_120128, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3685, plain, (r1_int_1(k1_int_2(X1),np__1)|~r1_int_1(k4_xcmplx_0(k1_int_2(X1)),np__1)|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_114616, plain, (k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(X1,np__0)))=k6_xcmplx_0(X1,np__0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_120148, plain, (r1_int_1(k4_xcmplx_0(X1),k6_xcmplx_0(X2,X3))|~r1_int_1(X1,k6_xcmplx_0(X2,X3))|~v7_ordinal1(X1)|~v1_int_1(k6_xcmplx_0(X2,X3))|~v1_int_1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_38583, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_118257, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)),np__2)=k6_xcmplx_0(k4_xcmplx_0(np__2),k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_118258, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)=k6_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3689, plain, (r1_int_1(k4_xcmplx_0(k1_int_2(X1)),k4_xcmplx_0(X2))|~r1_int_1(k1_int_2(X1),X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_xcmplx_0(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_3695, plain, (r1_int_1(X1,k4_xcmplx_0(k1_int_2(X2)))|~r1_int_1(X1,k1_int_2(X2))|~v1_int_1(k4_xcmplx_0(k1_int_2(X2)))|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_3696, plain, (r1_int_1(X1,k1_int_2(X2))|~r1_int_1(X1,k4_xcmplx_0(k1_int_2(X2)))|~v1_int_1(k4_xcmplx_0(k1_int_2(X2)))|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_123334, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v1_int_1(X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_3697, plain, (r1_int_1(k4_xcmplx_0(X1),k1_int_2(X2))|~r1_int_1(X1,k4_xcmplx_0(k1_int_2(X2)))|~v1_int_1(k4_xcmplx_0(k1_int_2(X2)))|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_121097, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_3698, plain, (r1_int_1(X1,k4_xcmplx_0(k1_int_2(X2)))|~r1_int_1(k4_xcmplx_0(X1),k1_int_2(X2))|~v1_int_1(k4_xcmplx_0(k1_int_2(X2)))|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_3705, plain, (r1_int_1(X1,k4_xcmplx_0(k1_int_2(X2)))|~r1_int_1(k4_xcmplx_0(X1),k1_int_2(X2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k1_int_2(X2)))|~v1_int_1(X2))).#trainneg
cnf(i_0_123608, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),k4_xcmplx_0(np__2))=k6_xcmplx_0(np__2,np__2))).#trainneg
cnf(i_0_123609, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),k4_xcmplx_0(np__1))=k6_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_123670, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2))=k6_xcmplx_0(np__2,np__1))).#trainneg
cnf(i_0_123671, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))=k6_xcmplx_0(np__1,np__1))).#trainneg
cnf(i_0_3710, plain, (r1_int_1(X1,k4_xcmplx_0(k1_int_2(X2)))|~r1_int_1(k4_xcmplx_0(X1),k1_int_2(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k1_int_2(X2)))|~v1_xcmplx_0(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_123856, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_121381, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3715, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k1_int_2(X3)))|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k1_int_2(X3))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(k1_int_2(X3)))|~v1_xcmplx_0(X1)|~v1_int_1(X3))).#trainneg
cnf(i_0_121096, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_124546, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(X2)|~r1_int_1(X1,X2))).#trainneg
cnf(i_0_124943, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v1_int_1(X1)|~r1_int_1(X1,X2))).#trainneg
cnf(i_0_126387, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v7_ordinal1(X2)|~v1_int_1(X1)|~r1_int_1(X1,X2))).#trainneg
cnf(i_0_125244, plain, (v1_int_1(k6_xcmplx_0(np__2,np__2)))).#trainneg
cnf(i_0_125278, plain, (v1_int_1(k6_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_125312, plain, (v1_int_1(k6_xcmplx_0(np__2,np__1)))).#trainneg
cnf(i_0_125346, plain, (v1_int_1(k6_xcmplx_0(np__1,np__1)))).#trainneg
cnf(i_0_3748, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),np__1)|~r1_int_1(k5_xcmplx_0(X1),np__1)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_125248, plain, (v1_xreal_0(k6_xcmplx_0(np__2,np__2))|~v1_xreal_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_125249, plain, (v1_xcmplx_0(k6_xcmplx_0(np__2,np__2))|~v1_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_125247, plain, (v1_xcmplx_0(k6_xcmplx_0(np__2,np__2))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_125252, plain, (v1_xcmplx_0(k6_xcmplx_0(np__2,np__2))|~v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_125281, plain, (v1_xcmplx_0(k6_xcmplx_0(np__1,np__2))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_125315, plain, (v1_xcmplx_0(k6_xcmplx_0(np__2,np__1))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1))))).#trainneg
cnf(i_0_125349, plain, (v1_xcmplx_0(k6_xcmplx_0(np__1,np__1))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1))))).#trainneg
cnf(i_0_3752, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~r1_int_1(k4_xcmplx_0(X1),np__1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_125350, plain, (v1_xreal_0(k6_xcmplx_0(np__1,np__1))|~v1_xreal_0(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1))))).#trainneg
cnf(i_0_125351, plain, (v1_xcmplx_0(k6_xcmplx_0(np__1,np__1))|~v1_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1))))).#trainneg
cnf(i_0_125354, plain, (v1_xcmplx_0(k6_xcmplx_0(np__1,np__1))|~v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1))))).#trainneg
cnf(i_0_3755, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))),np__1)|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),np__1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3757, plain, (r1_int_1(k4_xcmplx_0(k1_int_2(X1)),np__1)|~r1_int_1(k1_int_2(X1),np__1)|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_124759, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_118255, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)),np__0)=k6_xcmplx_0(np__0,k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_122306, plain, (v1_xreal_0(k6_xcmplx_0(X1,np__0))|~v1_xreal_0(k4_xcmplx_0(k6_xcmplx_0(X1,np__0)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_122308, plain, (v1_xcmplx_0(k6_xcmplx_0(X1,np__0))|~v1_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(X1,np__0)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_126203, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_3834, plain, (k3_int_2(k4_xcmplx_0(X1),np__0)=k3_int_2(X1,np__0)|k1_int_2(X1)=X1|~v1_int_1(X1))).#trainneg
cnf(i_0_124279, plain, (r1_int_1(k4_xcmplx_0(X1),k2_int_5(X2,X3))|~v1_int_2(X3)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X3)|~v1_int_1(X2)|~r1_int_1(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3))))).#trainneg
cnf(i_0_127573, plain, (v1_xreal_0(k6_xcmplx_0(np__2,np__2)))).#trainneg
cnf(i_0_127576, plain, (v1_xcmplx_0(k6_xcmplx_0(np__2,np__2)))).#trainneg
cnf(i_0_127719, plain, (v1_xreal_0(k6_xcmplx_0(np__1,np__1)))).#trainneg
cnf(i_0_127722, plain, (v1_xcmplx_0(k6_xcmplx_0(np__1,np__1)))).#trainneg
cnf(i_0_3849, plain, (k3_int_2(k7_xcmplx_0(X1,X2),k1_int_2(X3))=k3_int_2(k4_xcmplx_0(k7_xcmplx_0(X1,X2)),X3)|k1_int_2(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3851, plain, (k3_int_2(k6_xcmplx_0(X1,X2),k1_int_2(X3))=k3_int_2(k4_xcmplx_0(k6_xcmplx_0(X1,X2)),X3)|k1_int_2(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))=k4_xcmplx_0(k6_xcmplx_0(X1,X2))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_127451, plain, (k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)))=k6_xcmplx_0(np__2,np__2))).#trainneg
cnf(i_0_127468, plain, (k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)))=k6_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_127485, plain, (k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)))=k6_xcmplx_0(np__2,np__1))).#trainneg
cnf(i_0_3852, plain, (k3_int_2(k1_newton(X1,X2),k1_int_2(X3))=k3_int_2(k4_xcmplx_0(k1_newton(X1,X2)),X3)|k1_int_2(k4_xcmplx_0(k1_newton(X1,X2)))=k4_xcmplx_0(k1_newton(X1,X2))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_127502, plain, (k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)))=k6_xcmplx_0(np__1,np__1))).#trainneg
cnf(i_0_127455, plain, (v7_ordinal1(k6_xcmplx_0(np__2,np__2))|k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))!=k6_xcmplx_0(np__2,np__2))).#trainneg
cnf(i_0_127472, plain, (v7_ordinal1(k6_xcmplx_0(np__1,np__2))|k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))!=k6_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_127489, plain, (v7_ordinal1(k6_xcmplx_0(np__2,np__1))|k4_xcmplx_0(k6_xcmplx_0(np__2,np__1))!=k6_xcmplx_0(np__2,np__1))).#trainneg
cnf(i_0_3853, plain, (k3_int_2(k1_newton(X1,X2),k1_int_2(X3))=k3_int_2(k4_xcmplx_0(k1_newton(X1,X2)),X3)|k1_int_2(k4_xcmplx_0(k1_newton(X1,X2)))=k4_xcmplx_0(k1_newton(X1,X2))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_127506, plain, (v7_ordinal1(k6_xcmplx_0(np__1,np__1))|k4_xcmplx_0(k6_xcmplx_0(np__1,np__1))!=k6_xcmplx_0(np__1,np__1))).#trainneg
cnf(i_0_127456, plain, (v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)))|k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)))!=k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)))).#trainneg
cnf(i_0_3858, plain, (k3_int_2(k2_int_5(X1,X2),k1_int_2(X3))=k3_int_2(k4_xcmplx_0(k2_int_5(X1,X2)),X3)|k1_int_2(k4_xcmplx_0(k2_int_5(X1,X2)))=k4_xcmplx_0(k2_int_5(X1,X2))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_127458, plain, (v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)))|v7_ordinal1(k6_xcmplx_0(np__2,np__2)))).#trainneg
cnf(i_0_127475, plain, (v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)))|v7_ordinal1(k6_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_127473, plain, (v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)))|k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)))!=k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_127490, plain, (v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)))|k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)))!=k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)))).#trainneg
cnf(i_0_127507, plain, (v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)))|k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)))!=k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)))).#trainneg
cnf(i_0_3860, plain, (k3_int_2(X1,k1_int_2(X2))=k3_int_2(k4_xcmplx_0(X1),X2)|k1_int_2(k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_127492, plain, (v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)))|v7_ordinal1(k6_xcmplx_0(np__2,np__1)))).#trainneg
cnf(i_0_127509, plain, (v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)))|v7_ordinal1(k6_xcmplx_0(np__1,np__1)))).#trainneg
cnf(i_0_3871, plain, (k3_int_2(k4_xcmplx_0(k7_xcmplx_0(X1,X2)),X3)=k3_int_2(k7_xcmplx_0(X1,X2),X3)|k1_int_2(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_128716, plain, (k3_int_2(k4_xcmplx_0(np__2),np__0)=k3_int_2(np__2,np__0)|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_3872, plain, (k3_int_2(k4_xcmplx_0(k5_xcmplx_0(X1)),X2)=k3_int_2(k5_xcmplx_0(X1),X2)|k1_int_2(k4_xcmplx_0(k5_xcmplx_0(X1)))=k4_xcmplx_0(k5_xcmplx_0(X1))|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_127479, plain, (r1_xxreal_0(k6_xcmplx_0(np__1,np__2),np__1)|r1_xxreal_0(np__2,k6_xcmplx_0(np__1,np__2))|v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_127483, plain, (r1_xxreal_0(k6_xcmplx_0(np__1,np__2),np__0)|r1_xxreal_0(np__1,k6_xcmplx_0(np__1,np__2))|v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_127496, plain, (r1_xxreal_0(k6_xcmplx_0(np__2,np__1),np__1)|r1_xxreal_0(np__2,k6_xcmplx_0(np__2,np__1))|v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1))))).#trainneg
cnf(i_0_3873, plain, (k3_int_2(k4_xcmplx_0(k6_xcmplx_0(X1,X2)),X3)=k3_int_2(k6_xcmplx_0(X1,X2),X3)|k1_int_2(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))=k4_xcmplx_0(k6_xcmplx_0(X1,X2))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_127500, plain, (r1_xxreal_0(k6_xcmplx_0(np__2,np__1),np__0)|r1_xxreal_0(np__1,k6_xcmplx_0(np__2,np__1))|v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1))))).#trainneg
cnf(i_0_127465, plain, (r1_xxreal_0(k6_xcmplx_0(np__2,np__2),np__0)|r1_xxreal_0(np__1,k6_xcmplx_0(np__2,np__2))|m1_subset_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),k5_numbers))).#trainneg
cnf(i_0_127482, plain, (r1_xxreal_0(k6_xcmplx_0(np__1,np__2),np__0)|r1_xxreal_0(np__1,k6_xcmplx_0(np__1,np__2))|m1_subset_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)),k5_numbers))).#trainneg
cnf(i_0_127499, plain, (r1_xxreal_0(k6_xcmplx_0(np__2,np__1),np__0)|r1_xxreal_0(np__1,k6_xcmplx_0(np__2,np__1))|m1_subset_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)),k5_numbers))).#trainneg
cnf(i_0_127516, plain, (r1_xxreal_0(k6_xcmplx_0(np__1,np__1),np__0)|r1_xxreal_0(np__1,k6_xcmplx_0(np__1,np__1))|m1_subset_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)),k5_numbers))).#trainneg
cnf(i_0_3874, plain, (k3_int_2(k4_xcmplx_0(k1_newton(X1,X2)),X3)=k3_int_2(k1_newton(X1,X2),X3)|k1_int_2(k4_xcmplx_0(k1_newton(X1,X2)))=k4_xcmplx_0(k1_newton(X1,X2))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_126760, plain, (r1_int_1(X1,k7_xcmplx_0(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k7_xcmplx_0(X2,X3)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X2,X3)))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_126767, plain, (r1_int_1(X1,k1_newton(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_126768, plain, (r1_int_1(X1,k1_newton(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_3875, plain, (k3_int_2(k4_xcmplx_0(k1_newton(X1,X2)),X3)=k3_int_2(k1_newton(X1,X2),X3)|k1_int_2(k4_xcmplx_0(k1_newton(X1,X2)))=k4_xcmplx_0(k1_newton(X1,X2))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_130587, plain, (v1_xcmplx_0(k6_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_131461, plain, (v1_xcmplx_0(k6_xcmplx_0(np__2,np__1)))).#trainneg
cnf(i_0_128860, plain, (v1_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_128875, plain, (v1_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1))))).#trainneg
cnf(i_0_133255, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~r1_int_1(X1,np__1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3876, plain, (k3_int_2(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)=k3_int_2(k4_xcmplx_0(X1),X2)|k1_int_2(k4_xcmplx_0(k4_xcmplx_0(X1)))=k4_xcmplx_0(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_129748, plain, (v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_130590, plain, (v1_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_131464, plain, (v1_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1))))).#trainneg
cnf(i_0_129908, plain, (v1_xreal_0(k6_xcmplx_0(np__1,np__2))|~v1_xreal_0(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_3880, plain, (k3_int_2(k4_xcmplx_0(k2_int_5(X1,X2)),X3)=k3_int_2(k2_int_5(X1,X2),X3)|k1_int_2(k4_xcmplx_0(k2_int_5(X1,X2)))=k4_xcmplx_0(k2_int_5(X1,X2))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_132405, plain, (v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1))))).#trainneg
cnf(i_0_130297, plain, (v1_xreal_0(k6_xcmplx_0(np__1,np__2))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_130782, plain, (v1_xreal_0(k6_xcmplx_0(np__2,np__1))|~v1_xreal_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1))))).#trainneg
cnf(i_0_3881, plain, (k3_int_2(k4_xcmplx_0(k1_int_2(X1)),X2)=k3_int_2(k1_int_2(X1),X2)|k1_int_2(k4_xcmplx_0(k1_int_2(X1)))=k4_xcmplx_0(k1_int_2(X1))|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_131171, plain, (v1_xreal_0(k6_xcmplx_0(np__2,np__1))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1))))).#trainneg
cnf(i_0_130647, plain, (v1_xreal_0(k6_xcmplx_0(np__1,np__2))|~m1_subset_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)),k1_numbers))).#trainneg
cnf(i_0_131521, plain, (v1_xreal_0(k6_xcmplx_0(np__2,np__1))|~m1_subset_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)),k1_numbers))).#trainneg
cnf(i_0_129722, plain, (v1_xreal_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)))|~v7_ordinal1(k6_xcmplx_0(np__2,np__2)))).#trainneg
cnf(i_0_132379, plain, (v1_xreal_0(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)))|~v7_ordinal1(k6_xcmplx_0(np__1,np__1)))).#trainneg
cnf(i_0_125859, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(X2))).#trainneg
cnf(i_0_3882, plain, (k3_int_2(k4_xcmplx_0(X1),X2)=k3_int_2(X1,X2)|k1_int_2(k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_129073, plain, (v1_xreal_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_3975, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),np__2)|~r1_int_1(k5_xcmplx_0(X1),np__2)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_126271, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~r1_int_1(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3979, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__2)|~r1_int_1(k4_xcmplx_0(X1),np__2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_129789, plain, (r1_int_1(k6_xcmplx_0(np__2,np__2),np__1)|~r1_int_1(k6_xcmplx_0(np__2,np__2),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_131730, plain, (v1_xreal_0(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1))))).#trainneg
cnf(i_0_3982, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))),np__2)|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),np__2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_129790, plain, (r1_int_1(k6_xcmplx_0(np__2,np__2),np__2)|~r1_int_1(k6_xcmplx_0(np__2,np__2),k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_130643, plain, (r1_int_1(np__1,k6_xcmplx_0(np__1,np__2))|~r1_int_1(k4_xcmplx_0(np__1),k6_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_130644, plain, (r1_int_1(np__2,k6_xcmplx_0(np__1,np__2))|~r1_int_1(k4_xcmplx_0(np__2),k6_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_130663, plain, (r1_int_1(k6_xcmplx_0(np__1,np__2),np__1)|~r1_int_1(k6_xcmplx_0(np__1,np__2),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_130664, plain, (r1_int_1(k6_xcmplx_0(np__1,np__2),np__2)|~r1_int_1(k6_xcmplx_0(np__1,np__2),k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_124827, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~r1_int_1(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_125624, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~r1_int_1(X1,X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_3984, plain, (r1_int_1(k4_xcmplx_0(k1_int_2(X1)),np__2)|~r1_int_1(k1_int_2(X1),np__2)|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_131517, plain, (r1_int_1(np__1,k6_xcmplx_0(np__2,np__1))|~r1_int_1(k4_xcmplx_0(np__1),k6_xcmplx_0(np__2,np__1)))).#trainneg
cnf(i_0_131518, plain, (r1_int_1(np__2,k6_xcmplx_0(np__2,np__1))|~r1_int_1(k4_xcmplx_0(np__2),k6_xcmplx_0(np__2,np__1)))).#trainneg
cnf(i_0_131537, plain, (r1_int_1(k6_xcmplx_0(np__2,np__1),np__1)|~r1_int_1(k6_xcmplx_0(np__2,np__1),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_131538, plain, (r1_int_1(k6_xcmplx_0(np__2,np__1),np__2)|~r1_int_1(k6_xcmplx_0(np__2,np__1),k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_132446, plain, (r1_int_1(k6_xcmplx_0(np__1,np__1),np__1)|~r1_int_1(k6_xcmplx_0(np__1,np__1),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_132447, plain, (r1_int_1(k6_xcmplx_0(np__1,np__1),np__2)|~r1_int_1(k6_xcmplx_0(np__1,np__1),k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_121882, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_132882, plain, (k3_int_2(np__0,k4_xcmplx_0(np__2))=k3_int_2(np__2,np__0)|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_126490, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(X2))).#trainneg
cnf(i_0_4004, plain, (m1_subset_1(k5_xcmplx_0(X1),k5_numbers)|m1_subset_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k5_numbers)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_133078, plain, (k3_int_2(k4_xcmplx_0(X1),X2)=k3_int_2(X1,X2)|k1_int_2(k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_103021, plain, (r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__2),X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4008, plain, (m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k5_numbers)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_130083, plain, (v1_xreal_0(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)))|~v1_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)))|~m1_subset_1(k6_xcmplx_0(np__1,np__2),k1_numbers))).#trainneg
cnf(i_0_123374, plain, (r1_int_1(X1,k6_xcmplx_0(X2,X3))|~r1_int_1(k4_xcmplx_0(X1),k6_xcmplx_0(X2,X3))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(X3)|~v1_int_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_130957, plain, (v1_xreal_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)))|~v1_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)))|~m1_subset_1(k6_xcmplx_0(np__2,np__1),k1_numbers))).#trainneg
cnf(i_0_103027, plain, (r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__1),X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4047, plain, (k3_int_2(k1_int_2(X1),X2)=k3_int_2(X1,k4_xcmplx_0(X2))|k1_int_2(X2)=X2|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_int_1(X2)|~v1_int_1(k1_int_2(X1)))).#trainneg
cnf(i_0_127346, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k5_numbers)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),X2))).#trainneg
cnf(i_0_4048, plain, (k3_int_2(k1_int_2(X1),X2)=k3_int_2(k4_xcmplx_0(X2),X1)|k1_int_2(X2)=X2|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(k1_int_2(X1)))).#trainneg
cnf(i_0_4053, plain, (k3_int_2(X1,k4_xcmplx_0(np__2))=k3_int_2(X1,np__2)|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2)|~v1_int_1(X1))).#trainneg
cnf(i_0_91497, plain, (r1_int_1(np__2,X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_91502, plain, (r1_int_1(np__1,X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_129048, plain, (r1_int_1(np__0,k6_xcmplx_0(np__2,np__2))|~r1_int_1(np__0,k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_129051, plain, (r1_int_1(k6_xcmplx_0(np__2,np__2),np__0)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),np__0)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_129123, plain, (r1_int_1(k6_xcmplx_0(np__2,np__2),np__1)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),np__1)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_129149, plain, (r1_int_1(k6_xcmplx_0(np__2,np__2),np__2)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),np__2)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_108949, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_136494, plain, (m1_subset_1(X1,k5_numbers)|m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_133932, plain, (v1_xreal_0(k6_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_4055, plain, (k3_int_2(X1,k4_xcmplx_0(k7_xcmplx_0(X2,X3)))=k3_int_2(X1,k7_xcmplx_0(X2,X3))|k1_int_2(k4_xcmplx_0(k7_xcmplx_0(X2,X3)))=k4_xcmplx_0(k7_xcmplx_0(X2,X3))|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X2,X3)))|~v1_int_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_134018, plain, (v1_xreal_0(k6_xcmplx_0(np__2,np__1)))).#trainneg
cnf(i_0_136447, plain, (v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(X1)))|m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4056, plain, (k3_int_2(X1,k4_xcmplx_0(k5_xcmplx_0(X2)))=k3_int_2(X1,k5_xcmplx_0(X2))|k1_int_2(k4_xcmplx_0(k5_xcmplx_0(X2)))=k4_xcmplx_0(k5_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_133906, plain, (v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_133912, plain, (v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1))))).#trainneg
cnf(i_0_137287, plain, (k3_int_2(X1,k1_int_2(X2))=k3_int_2(k4_xcmplx_0(X2),X1)|k1_int_2(X2)=X2|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_4057, plain, (k3_int_2(X1,k4_xcmplx_0(k6_xcmplx_0(X2,X3)))=k3_int_2(X1,k6_xcmplx_0(X2,X3))|k1_int_2(k4_xcmplx_0(k6_xcmplx_0(X2,X3)))=k4_xcmplx_0(k6_xcmplx_0(X2,X3))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(X2,X3)))|~v1_int_1(X1)|~v1_int_1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_4058, plain, (k3_int_2(X1,k4_xcmplx_0(k1_newton(X2,X3)))=k3_int_2(X1,k1_newton(X2,X3))|k1_int_2(k4_xcmplx_0(k1_newton(X2,X3)))=k4_xcmplx_0(k1_newton(X2,X3))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(X1)|~v7_ordinal1(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_133886, plain, (k3_int_2(k4_xcmplx_0(X1),X2)=k3_int_2(X1,X2)|k1_int_2(X1)=X1|~v7_ordinal1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_4059, plain, (k3_int_2(X1,k4_xcmplx_0(k1_newton(X2,X3)))=k3_int_2(X1,k1_newton(X2,X3))|k1_int_2(k4_xcmplx_0(k1_newton(X2,X3)))=k4_xcmplx_0(k1_newton(X2,X3))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(X1)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_135012, plain, (r1_int_1(np__1,k6_xcmplx_0(np__1,np__2))|~r1_int_1(np__1,k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_4060, plain, (k3_int_2(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))=k3_int_2(X1,k4_xcmplx_0(X2))|k1_int_2(k4_xcmplx_0(k4_xcmplx_0(X2)))=k4_xcmplx_0(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_135047, plain, (r1_int_1(np__2,k6_xcmplx_0(np__1,np__2))|~r1_int_1(np__2,k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_4064, plain, (k3_int_2(X1,k4_xcmplx_0(k2_int_5(X2,X3)))=k3_int_2(X1,k2_int_5(X2,X3))|k1_int_2(k4_xcmplx_0(k2_int_5(X2,X3)))=k4_xcmplx_0(k2_int_5(X2,X3))|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_1(X1)|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_135708, plain, (r1_int_1(np__1,k6_xcmplx_0(np__2,np__1))|~r1_int_1(np__1,k4_xcmplx_0(k6_xcmplx_0(np__2,np__1))))).#trainneg
cnf(i_0_135743, plain, (r1_int_1(np__2,k6_xcmplx_0(np__2,np__1))|~r1_int_1(np__2,k4_xcmplx_0(k6_xcmplx_0(np__2,np__1))))).#trainneg
cnf(i_0_4065, plain, (k3_int_2(X1,k4_xcmplx_0(k1_int_2(X2)))=k3_int_2(X1,k1_int_2(X2))|k1_int_2(k4_xcmplx_0(k1_int_2(X2)))=k4_xcmplx_0(k1_int_2(X2))|~v1_int_1(k4_xcmplx_0(k1_int_2(X2)))|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_136848, plain, (k3_int_2(X1,k1_int_2(X2))=k3_int_2(X1,k4_xcmplx_0(X2))|k1_int_2(X2)=X2|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_137308, plain, (k3_int_2(k4_xcmplx_0(X1),X2)=k3_int_2(k4_xcmplx_0(X2),X1)|k1_int_2(X2)=X2|k1_int_2(X1)=X1|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_137314, plain, (k3_int_2(np__0,np__2)=k3_int_2(np__2,np__0)|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_4146, plain, (r1_int_1(k5_xcmplx_0(X1),np__2)|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),np__2)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_134411, plain, (k3_int_2(X2,k4_xcmplx_0(X1))=k3_int_2(X1,X2)|k1_int_2(k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_136274, plain, (k3_int_2(X2,k4_xcmplx_0(X1))=k3_int_2(X1,X2)|k1_int_2(k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4150, plain, (r1_int_1(k4_xcmplx_0(X1),np__2)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4151, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__2)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))),np__2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v1_int_1(X1))).#trainneg
cnf(i_0_89269, plain, (r1_int_1(X1,np__1)|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4152, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__2)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))),np__2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4153, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),np__2)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))),np__2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_124991, plain, (r1_int_1(k2_int_5(X1,X2),k4_xcmplx_0(X3))|v7_ordinal1(k2_int_5(X1,X2))|~r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),X3)|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_136516, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_int_1(X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_4155, plain, (r1_int_1(k1_int_2(X1),np__2)|~r1_int_1(k4_xcmplx_0(k1_int_2(X1)),np__2)|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_119895, plain, (v7_ordinal1(k6_xcmplx_0(X1,np__0))|k4_xcmplx_0(k6_xcmplx_0(X1,np__0))!=k6_xcmplx_0(X1,np__0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_129245, plain, (r1_int_1(np__1,k6_xcmplx_0(np__2,np__2))|~r1_int_1(np__1,k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_129256, plain, (r1_int_1(np__2,k6_xcmplx_0(np__2,np__2))|~r1_int_1(np__2,k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_81832, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~r1_int_1(k1_newton(X1,X2),X3)|~v7_ordinal1(X3)|~v1_int_1(k1_newton(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_136849, plain, (k3_int_2(k1_int_2(X2),X1)=k3_int_2(X1,k4_xcmplx_0(X2))|k1_int_2(X2)=X2|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_109218, plain, (r1_int_1(np__1,k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__1),X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4160, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_129922, plain, (r1_int_1(np__0,k6_xcmplx_0(np__1,np__2))|~r1_int_1(np__0,k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_129925, plain, (r1_int_1(k6_xcmplx_0(np__1,np__2),np__0)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)),np__0)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_129997, plain, (r1_int_1(k6_xcmplx_0(np__1,np__2),np__1)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)),np__1)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_11394, plain, (v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))!=k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_4231, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),np__1)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_130023, plain, (r1_int_1(k6_xcmplx_0(np__1,np__2),np__2)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)),np__2)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_91417, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~r1_int_1(np__1,k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_140807, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v1_int_1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_138016, plain, (v7_ordinal1(k4_xcmplx_0(X1))|m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_138693, plain, (k3_int_2(X1,k4_xcmplx_0(X2))=k3_int_2(X1,X2)|k1_int_2(X2)=X2|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_4233, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),np__1)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(np__1))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4241, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),np__1)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~r1_int_1(k1_newton(X1,X2),np__1)|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_139260, plain, (k3_int_2(k4_xcmplx_0(X2),X1)=k3_int_2(X1,X2)|k1_int_2(X2)=X2|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_138463, plain, (k3_int_2(X2,k4_xcmplx_0(X1))=k3_int_2(X1,X2)|k1_int_2(X1)=X1|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_138920, plain, (k3_int_2(X1,k4_xcmplx_0(X2))=k3_int_2(k4_xcmplx_0(X2),X1)|k1_int_2(X2)=X2|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_139202, plain, (k3_int_2(k4_xcmplx_0(np__1),np__2)=k3_int_2(np__2,k4_xcmplx_0(np__1))|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_126578, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X2))).#trainneg
cnf(i_0_127349, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k5_numbers)|~v7_ordinal1(X2)|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_101109, plain, (r1_int_1(k4_xcmplx_0(X1),np__2)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_139514, plain, (r1_int_1(k6_xcmplx_0(np__2,np__1),np__2)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)),np__2))).#trainneg
cnf(i_0_139515, plain, (r1_int_1(k6_xcmplx_0(np__1,np__1),np__2)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)),np__2))).#trainneg
cnf(i_0_138437, plain, (k3_int_2(X1,k4_xcmplx_0(X2))=k3_int_2(X1,X2)|k1_int_2(k4_xcmplx_0(X2))=k4_xcmplx_0(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_139490, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),np__2)|~r1_int_1(k6_xcmplx_0(np__2,np__2),np__2))).#trainneg
cnf(i_0_139491, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)),np__2)|~r1_int_1(k6_xcmplx_0(np__1,np__2),np__2))).#trainneg
cnf(i_0_139492, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)),np__2)|~r1_int_1(k6_xcmplx_0(np__2,np__1),np__2))).#trainneg
cnf(i_0_139493, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)),np__2)|~r1_int_1(k6_xcmplx_0(np__1,np__1),np__2))).#trainneg
cnf(i_0_108181, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_86437, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_82861, plain, (v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))))|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))))|~v1_int_1(X1))).#trainneg
cnf(i_0_66817, plain, (v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4364, plain, (r1_int_1(X1,k5_xcmplx_0(X2))|~r1_int_1(X1,k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4368, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4369, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_xcmplx_0(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_144723, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),X2))).#trainneg
cnf(i_0_4370, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_142915, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4371, plain, (r1_int_1(X1,k4_xcmplx_0(k5_xcmplx_0(X2)))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X2))))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X2))))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_141923, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v1_int_1(X1)|~r1_int_1(X1,X2))).#trainneg
cnf(i_0_4373, plain, (r1_int_1(X1,k1_int_2(X2))|~r1_int_1(X1,k4_xcmplx_0(k1_int_2(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k1_int_2(X2)))|~v1_xcmplx_0(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_4393, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__0)|r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(np__1))|~r1_int_1(X1,np__1)|~v1_int_1(X1))).#trainneg
cnf(i_0_142034, plain, (k3_int_2(k4_xcmplx_0(X2),X1)=k3_int_2(X1,X2)|k1_int_2(X2)=X2|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_141629, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(X1))).#trainneg
cnf(i_0_143612, plain, (r1_int_1(k6_xcmplx_0(np__2,np__2),k4_xcmplx_0(np__2))|~r1_int_1(k6_xcmplx_0(np__2,np__2),np__2))).#trainneg
cnf(i_0_4418, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),np__0)|r1_xxreal_0(X1,k4_xcmplx_0(np__2))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(np__2))|~r1_int_1(X1,np__2)|~v1_int_1(X1))).#trainneg
cnf(i_0_142614, plain, (k3_int_2(np__2,k4_xcmplx_0(np__1))=k3_int_2(np__1,np__2)|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_4441, plain, (v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))!=k4_xcmplx_0(k6_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_143650, plain, (r1_int_1(k6_xcmplx_0(np__1,np__2),k4_xcmplx_0(np__2))|~r1_int_1(k6_xcmplx_0(np__1,np__2),np__2))).#trainneg
cnf(i_0_4442, plain, (v7_ordinal1(k4_xcmplx_0(k1_newton(X1,X2)))|k4_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2)))!=k4_xcmplx_0(k1_newton(X1,X2))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_4448, plain, (v7_ordinal1(k4_xcmplx_0(k2_int_5(X1,X2)))|k4_xcmplx_0(k4_xcmplx_0(k2_int_5(X1,X2)))!=k4_xcmplx_0(k2_int_5(X1,X2))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_4454, plain, (v1_xreal_0(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~v1_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4455, plain, (v1_xreal_0(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1)))|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_146547, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_4456, plain, (v1_xreal_0(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v1_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~m1_subset_1(k6_xcmplx_0(X1,X2),k1_numbers)|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_4457, plain, (v1_xreal_0(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2)))|~m1_subset_1(k1_newton(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4458, plain, (v1_xreal_0(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2)))|~m1_subset_1(k1_newton(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_145019, plain, (r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(X1))|~r1_int_1(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4459, plain, (v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_145020, plain, (r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))|~r1_int_1(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4463, plain, (v1_xreal_0(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_xcmplx_0(k4_xcmplx_0(k2_int_5(X1,X2)))|~m1_subset_1(k2_int_5(X1,X2),k1_numbers)|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_143688, plain, (r1_int_1(k6_xcmplx_0(np__2,np__1),k4_xcmplx_0(np__2))|~r1_int_1(k6_xcmplx_0(np__2,np__1),np__2))).#trainneg
cnf(i_0_4464, plain, (v1_xreal_0(k4_xcmplx_0(k1_int_2(X1)))|~v1_xcmplx_0(k4_xcmplx_0(k1_int_2(X1)))|~m1_subset_1(k1_int_2(X1),k1_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_142460, plain, (k3_int_2(np__1,k4_xcmplx_0(X1))=np__1|k1_int_2(X1)=X1|~v1_int_1(X1))).#trainneg
cnf(i_0_143726, plain, (r1_int_1(k6_xcmplx_0(np__1,np__1),k4_xcmplx_0(np__2))|~r1_int_1(k6_xcmplx_0(np__1,np__1),np__2))).#trainneg
cnf(i_0_138013, plain, (v1_xreal_0(k4_xcmplx_0(X1))|m1_subset_1(X1,k5_numbers)|~v3_membered(k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_130796, plain, (r1_int_1(np__0,k6_xcmplx_0(np__2,np__1))|~r1_int_1(np__0,k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1))))).#trainneg
cnf(i_0_145731, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),np__1))).#trainneg
cnf(i_0_138104, plain, (v1_xreal_0(k4_xcmplx_0(X1))|v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v3_membered(k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_130799, plain, (r1_int_1(k6_xcmplx_0(np__2,np__1),np__0)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)),np__0)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1))))).#trainneg
cnf(i_0_130871, plain, (r1_int_1(k6_xcmplx_0(np__2,np__1),np__1)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)),np__1)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1))))).#trainneg
cnf(i_0_147390, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),np__0)|r1_xxreal_0(X1,k4_xcmplx_0(np__2))|~r1_int_1(X1,np__2)|~v7_ordinal1(X1)|k4_xcmplx_0(np__2)!=np__2)).#trainneg
cnf(i_0_4505, plain, (v7_ordinal1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|m1_subset_1(k7_xcmplx_0(X1,X2),k5_numbers)|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_131705, plain, (r1_int_1(np__0,k6_xcmplx_0(np__1,np__1))|~r1_int_1(np__0,k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1))))).#trainneg
cnf(i_0_146875, plain, (k3_int_2(k4_xcmplx_0(X2),X1)=k3_int_2(k4_xcmplx_0(X1),X2)|k1_int_2(X2)=X2|k1_int_2(X1)=X1|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_146659, plain, (r1_int_1(X1,k2_int_5(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|~v7_ordinal1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_1(X1)|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_4506, plain, (v7_ordinal1(k4_xcmplx_0(k5_xcmplx_0(X1)))|m1_subset_1(k5_xcmplx_0(X1),k5_numbers)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_131708, plain, (r1_int_1(k6_xcmplx_0(np__1,np__1),np__0)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)),np__0)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1))))).#trainneg
cnf(i_0_131780, plain, (r1_int_1(k6_xcmplx_0(np__1,np__1),np__1)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)),np__1)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1))))).#trainneg
cnf(i_0_131902, plain, (r1_int_1(np__1,k6_xcmplx_0(np__1,np__1))|~r1_int_1(np__1,k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1))))).#trainneg
cnf(i_0_145819, plain, (r1_int_1(X1,np__2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(np__2))|~v1_xcmplx_0(X1)|~r1_int_1(X1,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_4507, plain, (v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|m1_subset_1(k6_xcmplx_0(X1,X2),k5_numbers)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_131913, plain, (r1_int_1(np__2,k6_xcmplx_0(np__1,np__1))|~r1_int_1(np__2,k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1))))).#trainneg
cnf(i_0_145305, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_146874, plain, (k3_int_2(k4_xcmplx_0(X2),X1)=k3_int_2(X2,k4_xcmplx_0(X1))|k1_int_2(X2)=X2|k1_int_2(X1)=X1|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_142893, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4508, plain, (v7_ordinal1(k4_xcmplx_0(k1_newton(X1,X2)))|m1_subset_1(k1_newton(X1,X2),k5_numbers)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4509, plain, (v7_ordinal1(k4_xcmplx_0(k1_newton(X1,X2)))|m1_subset_1(k1_newton(X1,X2),k5_numbers)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_4511, plain, (v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k5_numbers)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v1_int_1(X1))).#trainneg
cnf(i_0_4512, plain, (v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k5_numbers)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_148107, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4513, plain, (v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|m1_subset_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k5_numbers)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4514, plain, (v7_ordinal1(k4_xcmplx_0(k2_int_5(X1,X2)))|m1_subset_1(k2_int_5(X1,X2),k5_numbers)|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_150165, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),np__1))).#trainneg
cnf(i_0_150111, plain, (r1_int_1(X1,np__0)|~r1_int_1(k4_xcmplx_0(X1),np__0)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_150113, plain, (r1_int_1(X1,np__2)|~r1_int_1(k4_xcmplx_0(X1),np__2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4518, plain, (r1_int_1(k4_xcmplx_0(np__1),k5_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1)|~r1_int_1(np__1,k4_xcmplx_0(k5_xcmplx_0(X1))))).#trainneg
cnf(i_0_4519, plain, (r1_int_1(k4_xcmplx_0(np__1),k1_newton(X1,X2))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(X1)|~r1_int_1(np__1,k4_xcmplx_0(k1_newton(X1,X2))))).#trainneg
cnf(i_0_4521, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~v1_int_1(X1)|~r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_4526, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~v7_ordinal1(k4_xcmplx_0(np__1))|~v7_ordinal1(X1)|~r1_int_1(np__1,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_4552, plain, (r1_int_1(k4_xcmplx_0(np__2),k5_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1)|~r1_int_1(np__2,k4_xcmplx_0(k5_xcmplx_0(X1))))).#trainneg
cnf(i_0_4553, plain, (r1_int_1(k4_xcmplx_0(np__2),k1_newton(X1,X2))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(X1)|~r1_int_1(np__2,k4_xcmplx_0(k1_newton(X1,X2))))).#trainneg
cnf(i_0_150006, plain, (k3_int_2(X1,k4_xcmplx_0(X2))=k3_int_2(X2,k4_xcmplx_0(X1))|k1_int_2(X2)=X2|k1_int_2(X1)=X1|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4555, plain, (r1_int_1(k4_xcmplx_0(np__2),X1)|~v1_int_1(X1)|~r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_4560, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(k4_xcmplx_0(np__2),X1)|~v7_ordinal1(k4_xcmplx_0(np__2))|~v7_ordinal1(X1)|~r1_int_1(np__2,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_146855, plain, (k3_int_2(k4_xcmplx_0(np__1),X1)=np__1|k1_int_2(X1)=X1|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_146867, plain, (k3_int_2(k4_xcmplx_0(X2),X1)=k3_int_2(X2,k4_xcmplx_0(X1))|k1_int_2(X2)=X2|k1_int_2(X1)=X1|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_65434, plain, (k6_xcmplx_0(X1,np__0)=k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_118370, plain, (k1_int_2(X1)=X1|k18_complex1(k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_109475, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4589, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),np__2)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~r1_int_1(k5_xcmplx_0(X1),k4_xcmplx_0(np__2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_131689, plain, (k3_int_2(X1,k1_int_2(X2))=k3_int_2(k4_xcmplx_0(X1),X2)|k1_int_2(k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_129147, plain, (r1_int_1(k6_xcmplx_0(np__2,np__2),np__2)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),k4_xcmplx_0(np__2))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_4591, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),np__2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(np__2))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_129247, plain, (r1_int_1(k6_xcmplx_0(np__2,np__2),np__1)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_131804, plain, (r1_int_1(k6_xcmplx_0(np__1,np__1),np__2)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)),k4_xcmplx_0(np__2))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1))))).#trainneg
cnf(i_0_131904, plain, (r1_int_1(k6_xcmplx_0(np__1,np__1),np__1)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)),k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1))))).#trainneg
cnf(i_0_84395, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X2))|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(X1),np__1)|~v1_int_1(k4_xcmplx_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_98871, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~v7_ordinal1(k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_4599, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),np__2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~r1_int_1(k1_newton(X1,X2),np__2)|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6244, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__2)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(np__2))|~v1_int_1(X1))).#trainneg
cnf(i_0_8235, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__1)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v1_int_1(X1))).#trainneg
cnf(i_0_128857, plain, (r1_int_1(X1,k2_int_5(X2,X3))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_2(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_1(X2))).#trainneg
cnf(i_0_23768, plain, (r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(k4_xcmplx_0(X1)))|~r1_int_1(k4_xcmplx_0(np__2),X1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_4609, plain, (r1_int_1(np__2,k6_xcmplx_0(k2_int_5(X1,k4_xcmplx_0(np__2)),k1_newton(X1,k3_nat_d(k7_nat_d(k4_xcmplx_0(np__2),np__1),np__2))))|r1_xxreal_0(k4_xcmplx_0(np__2),np__2)|~v1_int_1(k6_xcmplx_0(k2_int_5(X1,k4_xcmplx_0(np__2)),k1_newton(X1,k3_nat_d(k7_nat_d(k4_xcmplx_0(np__2),np__1),np__2))))|k3_int_2(X1,k4_xcmplx_0(np__2))!=np__1|~v1_int_2(k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(np__2))|~v1_int_1(k1_newton(X1,k3_nat_d(k7_nat_d(k4_xcmplx_0(np__2),np__1),np__2)))|~v1_int_1(X1))).#trainneg
cnf(i_0_23769, plain, (r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(k4_xcmplx_0(X1)))|~r1_int_1(k4_xcmplx_0(np__1),X1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_106024, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X2))|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_143298, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__2)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_146866, plain, (k3_int_2(k4_xcmplx_0(X2),X1)=k3_int_2(k4_xcmplx_0(X1),X2)|k1_int_2(X2)=X2|k1_int_2(X1)=X1|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_153003, plain, (r1_xxreal_0(X1,np__2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_4638, plain, (r1_int_1(np__1,k6_xcmplx_0(k2_int_5(X1,k4_xcmplx_0(np__1)),k1_newton(X1,k3_nat_d(k7_nat_d(k4_xcmplx_0(np__1),np__1),np__2))))|r1_xxreal_0(k4_xcmplx_0(np__1),np__2)|~v1_int_1(k6_xcmplx_0(k2_int_5(X1,k4_xcmplx_0(np__1)),k1_newton(X1,k3_nat_d(k7_nat_d(k4_xcmplx_0(np__1),np__1),np__2))))|k3_int_2(X1,k4_xcmplx_0(np__1))!=np__1|~v1_int_2(k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(np__1))|~v1_int_1(k1_newton(X1,k3_nat_d(k7_nat_d(k4_xcmplx_0(np__1),np__1),np__2)))|~v1_int_1(X1))).#trainneg
cnf(i_0_150435, plain, (v7_ordinal1(k4_xcmplx_0(X1))|m1_subset_1(X1,k5_numbers)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_151219, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~m1_subset_1(X1,k5_numbers)|~r1_int_1(np__2,k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))))).#trainneg
cnf(i_0_4670, plain, (r1_int_1(X1,k4_xcmplx_0(k5_xcmplx_0(X2)))|~r1_int_1(X1,k5_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_150718, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(X1,np__1))).#trainneg
cnf(i_0_4673, plain, (r1_int_1(X1,k4_xcmplx_0(k1_newton(X2,X3)))|~r1_int_1(X1,k1_newton(X2,X3))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_4674, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_152265, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_int_1(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4675, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_xcmplx_0(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_4676, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2))))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_150554, plain, (v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_153649, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_xxreal_0(X2,np__1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)),np__1)|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_4677, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X2))))|~r1_int_1(X1,k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X2))))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_153652, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~v7_ordinal1(k4_xcmplx_0(np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4679, plain, (r1_int_1(X1,k4_xcmplx_0(k1_int_2(X2)))|~r1_int_1(X1,k1_int_2(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k1_int_2(X2)))|~v1_xcmplx_0(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_126179, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(X1))).#trainneg
cnf(i_0_113459, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_int_1(X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_154824, plain, (r1_int_1(X1,X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_153894, negated_conjecture, (~r1_int_1(k4_xcmplx_0(esk2_0),k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_153922, plain, (v7_ordinal1(X1)|v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_155759, plain, (r1_xxreal_0(X1,np__1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~v7_ordinal1(k4_xcmplx_0(np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_154759, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(X1,k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_153928, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(np__2,k1_newton(k4_xcmplx_0(np__1),X1))|~v1_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_4713, plain, (r1_int_1(k4_xcmplx_0(np__1),k1_newton(X1,X2))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(X1)|~r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2)))))).#trainneg
cnf(i_0_153938, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(np__2,k1_newton(k4_xcmplx_0(np__1),X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4715, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~v1_int_1(X1)|~r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_4718, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_4721, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~v7_ordinal1(k4_xcmplx_0(np__1))|~v7_ordinal1(X1)|~r1_int_1(np__1,k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_4756, plain, (r1_int_1(np__0,np__1)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(np__0,k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4759, plain, (r1_int_1(np__1,np__0)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),np__0)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4766, plain, (k6_xcmplx_0(k4_xcmplx_0(k1_newton(X2,X3)),k1_newton(k4_xcmplx_0(np__1),X1))=k6_xcmplx_0(np__1,k1_newton(X2,X3))|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v7_ordinal1(X3)|~v1_xcmplx_0(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4767, plain, (r1_int_1(np__1,k1_newton(X2,X3))|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_154894, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2))).#trainneg
cnf(i_0_154915, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4768, plain, (r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),k4_xcmplx_0(k1_newton(X2,X3)))|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(np__1,k1_newton(X2,X3))|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_99130, plain, (r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),X3)|~v1_int_1(k4_xcmplx_0(X3))|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_xcmplx_0(X3)|~r1_int_1(k2_int_5(X1,X2),X3)|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_4771, plain, (r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),k4_xcmplx_0(k5_xcmplx_0(X2)))|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(np__1,k5_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_148160, plain, (r1_int_1(k4_xcmplx_0(np__2),X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(np__2))|~v1_xcmplx_0(X1)|~r1_int_1(np__2,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_4786, plain, (k3_int_2(k1_newton(k4_xcmplx_0(np__1),X1),X2)=k3_int_2(np__1,X2)|k1_int_2(k1_newton(k4_xcmplx_0(np__1),X1))=k1_newton(k4_xcmplx_0(np__1),X1)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v1_int_1(X2)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4790, plain, (r1_int_1(np__1,np__1)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),np__1)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_157753, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4797, plain, (r1_int_1(np__1,np__2)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),np__2)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4825, plain, (k3_int_2(X1,k1_newton(k4_xcmplx_0(np__1),X2))=k3_int_2(X1,np__1)|k1_int_2(k1_newton(k4_xcmplx_0(np__1),X2))=k1_newton(k4_xcmplx_0(np__1),X2)|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~v1_int_1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4827, plain, (k6_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X3),k4_xcmplx_0(k1_newton(X1,X2)))=k6_xcmplx_0(k1_newton(X1,X2),np__1)|k1_newton(k4_xcmplx_0(np__1),X3)=np__1|~v7_ordinal1(X2)|~v1_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X3))|~v1_xcmplx_0(X1)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_158230, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~r1_int_1(X1,X2)|~v7_ordinal1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_4829, plain, (r1_int_1(k1_newton(X1,X2),k1_newton(k4_xcmplx_0(np__1),X3))|k1_newton(k4_xcmplx_0(np__1),X3)=np__1|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),np__1)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X3))|~v1_xcmplx_0(X1)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_157341, negated_conjecture, (~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(esk2_0)),np__1)|~v7_ordinal1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_4833, plain, (k6_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,np__1)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v7_ordinal1(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4836, plain, (r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),k4_xcmplx_0(X2))|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(np__1,X2)|~v7_ordinal1(X1)|~v1_int_1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4837, plain, (r1_int_1(np__1,X2)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4849, plain, (r1_int_1(k4_xcmplx_0(np__2),k1_newton(X1,X2))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(X1)|~r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2)))))).#trainneg
cnf(i_0_4851, plain, (r1_int_1(k4_xcmplx_0(np__2),X1)|~v1_int_1(X1)|~r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_158159, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,X2)|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),k4_xcmplx_0(X2))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4854, plain, (r1_int_1(k4_xcmplx_0(np__2),X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_159050, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(np__2,k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~v7_ordinal1(k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~v1_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_4857, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(k4_xcmplx_0(np__2),X1)|~v7_ordinal1(k4_xcmplx_0(np__2))|~v7_ordinal1(X1)|~r1_int_1(np__2,k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_4878, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2))),X3)=k6_xcmplx_0(k1_newton(X1,X2),k4_xcmplx_0(k4_xcmplx_0(X3)))|~v7_ordinal1(X2)|~v1_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_4879, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)=k6_xcmplx_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v7_ordinal1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_83885, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1)|~r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_4880, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)=k6_xcmplx_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_158185, plain, (r1_int_1(np__0,k4_xcmplx_0(k4_xcmplx_0(X1)))|~r1_int_1(np__0,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_156961, plain, (r1_int_1(k4_xcmplx_0(np__2),X1)|~v7_ordinal1(X1)|~r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_149828, plain, (r1_int_1(np__2,k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(np__2),X1)|~v7_ordinal1(k4_xcmplx_0(np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4881, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)=k6_xcmplx_0(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_146952, plain, (k3_int_2(X2,X1)=k3_int_2(X1,X2)|k1_int_2(X1)=X1|~v7_ordinal1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_4883, plain, (v1_int_1(k6_xcmplx_0(k4_xcmplx_0(X2),X1))|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_4884, plain, (k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(X2),X1)))=k6_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_4889, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2))),X3)=k6_xcmplx_0(k1_newton(X1,X2),k4_xcmplx_0(k4_xcmplx_0(X3)))|~v7_ordinal1(X2)|~v1_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_4890, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)=k6_xcmplx_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4891, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)=k6_xcmplx_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4892, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)=k6_xcmplx_0(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_4894, plain, (v1_int_1(k6_xcmplx_0(k4_xcmplx_0(X2),X1))|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_129050, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),np__0)|~r1_int_1(k6_xcmplx_0(np__2,np__2),np__0)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_129114, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),np__1)|~r1_int_1(k6_xcmplx_0(np__2,np__2),np__1)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_4895, plain, (k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(X2),X1)))=k6_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_160795, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v7_ordinal1(k4_xcmplx_0(np__2))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(np__2,k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_160662, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,X2)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_129923, plain, (r1_int_1(np__0,k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)))|~r1_int_1(np__0,k6_xcmplx_0(np__1,np__2))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_160660, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,X2)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2))).#trainneg
cnf(i_0_129924, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)),np__0)|~r1_int_1(k6_xcmplx_0(np__1,np__2),np__0)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_129988, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)),np__1)|~r1_int_1(k6_xcmplx_0(np__1,np__2),np__1)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_159982, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|~v7_ordinal1(X1)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_162078, plain, (k3_int_2(X1,k4_xcmplx_0(np__1))=np__1|k1_int_2(X1)=X1|~v7_ordinal1(X1))).#trainneg
cnf(i_0_161280, plain, (k6_xcmplx_0(np__2,k4_xcmplx_0(k4_xcmplx_0(X1)))=k6_xcmplx_0(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_130117, plain, (r1_int_1(np__1,k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)))|~r1_int_1(np__1,k6_xcmplx_0(np__1,np__2))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_130135, plain, (r1_int_1(np__2,k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)))|~r1_int_1(np__2,k6_xcmplx_0(np__1,np__2))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_4976, plain, (v1_xreal_0(k1_newton(X1,X2))|~v7_ordinal1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_161281, plain, (k6_xcmplx_0(np__1,k4_xcmplx_0(k4_xcmplx_0(X1)))=k6_xcmplx_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4977, plain, (v1_xreal_0(k1_newton(X1,X2))|~v7_ordinal1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_160175, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))=k6_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_130797, plain, (r1_int_1(np__0,k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)))|~r1_int_1(np__0,k6_xcmplx_0(np__2,np__1))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1))))).#trainneg
cnf(i_0_130798, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)),np__0)|~r1_int_1(k6_xcmplx_0(np__2,np__1),np__0)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1))))).#trainneg
cnf(i_0_162102, plain, (k3_int_2(X2,k4_xcmplx_0(X1))=k3_int_2(X1,k4_xcmplx_0(X2))|k1_int_2(X2)=X2|k1_int_2(X1)=X1|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_120545, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4982, plain, (v1_xreal_0(k2_int_5(X1,X2))|~v7_ordinal1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_130862, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)),np__1)|~r1_int_1(k6_xcmplx_0(np__2,np__1),np__1)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1))))).#trainneg
cnf(i_0_103345, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~r1_int_1(X1,k4_xcmplx_0(np__1))|~v7_ordinal1(X1)|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_161256, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(k4_xcmplx_0(X1)))=k6_xcmplx_0(np__0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_130991, plain, (r1_int_1(np__1,k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)))|~r1_int_1(np__1,k6_xcmplx_0(np__2,np__1))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1))))).#trainneg
cnf(i_0_131009, plain, (r1_int_1(np__2,k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)))|~r1_int_1(np__2,k6_xcmplx_0(np__2,np__1))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1))))).#trainneg
cnf(i_0_88623, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__0)|r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_160329, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)),X2)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_5027, plain, (v1_xreal_0(k1_int_2(X1))|~v1_int_1(X1)|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k1_int_2(X1))))|~v1_xcmplx_0(k4_xcmplx_0(k1_int_2(X1))))).#trainneg
cnf(i_0_133778, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X2)))|~v1_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X2))|~r1_int_1(k4_xcmplx_0(X1),np__1)|~v7_ordinal1(k4_xcmplx_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_5028, plain, (v1_xreal_0(k1_int_2(X1))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k1_int_2(X1))))|~v1_xcmplx_0(k4_xcmplx_0(k1_int_2(X1))))).#trainneg
cnf(i_0_131707, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)),np__0)|~r1_int_1(k6_xcmplx_0(np__1,np__1),np__0)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1))))).#trainneg
cnf(i_0_138671, plain, (k3_int_2(X1,k4_xcmplx_0(X2))=k3_int_2(X1,X2)|k1_int_2(k4_xcmplx_0(X2))=k4_xcmplx_0(X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_105310, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_131771, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)),np__1)|~r1_int_1(k6_xcmplx_0(np__1,np__1),np__1)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1))))).#trainneg
cnf(i_0_160197, plain, (r1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)),X2)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v1_int_1(X2)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~r1_int_1(np__1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_138253, plain, (k3_int_2(k1_int_2(X2),X1)=k3_int_2(k4_xcmplx_0(X2),X1)|k1_int_2(X2)=X2|~v1_int_1(k1_int_2(X2))|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_164054, plain, (r1_int_1(k6_xcmplx_0(np__2,np__2),k4_xcmplx_0(np__1))|~r1_int_1(k6_xcmplx_0(np__2,np__2),np__1))).#trainneg
cnf(i_0_162206, plain, (v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v1_int_1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_5097, plain, (k3_int_2(k4_xcmplx_0(X1),k7_xcmplx_0(X2,X3))=k3_int_2(X1,k4_xcmplx_0(k7_xcmplx_0(X2,X3)))|k1_int_2(k4_xcmplx_0(k7_xcmplx_0(X2,X3)))=k4_xcmplx_0(k7_xcmplx_0(X2,X3))|k1_int_2(X1)=X1|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X2,X3)))|~v1_int_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_165942, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),np__1)=k6_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_5099, plain, (k3_int_2(k4_xcmplx_0(X1),k6_xcmplx_0(X2,X3))=k3_int_2(X1,k4_xcmplx_0(k6_xcmplx_0(X2,X3)))|k1_int_2(k4_xcmplx_0(k6_xcmplx_0(X2,X3)))=k4_xcmplx_0(k6_xcmplx_0(X2,X3))|k1_int_2(X1)=X1|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(X2,X3)))|~v1_int_1(X1)|~v1_int_1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_162180, plain, (v1_int_1(k6_xcmplx_0(np__0,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_165502, plain, (r1_int_1(k6_xcmplx_0(np__1,np__2),k4_xcmplx_0(np__1))|~r1_int_1(k6_xcmplx_0(np__1,np__2),np__1))).#trainneg
cnf(i_0_167117, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,k4_xcmplx_0(X2))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2))).#trainneg
cnf(i_0_5100, plain, (k3_int_2(k4_xcmplx_0(X1),k1_newton(X2,X3))=k3_int_2(X1,k4_xcmplx_0(k1_newton(X2,X3)))|k1_int_2(k4_xcmplx_0(k1_newton(X2,X3)))=k4_xcmplx_0(k1_newton(X2,X3))|k1_int_2(X1)=X1|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(X1)|~v7_ordinal1(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_5101, plain, (k3_int_2(k4_xcmplx_0(X1),k1_newton(X2,X3))=k3_int_2(X1,k4_xcmplx_0(k1_newton(X2,X3)))|k1_int_2(k4_xcmplx_0(k1_newton(X2,X3)))=k4_xcmplx_0(k1_newton(X2,X3))|k1_int_2(X1)=X1|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_int_1(X1)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_166615, plain, (r1_int_1(k6_xcmplx_0(np__2,np__1),k4_xcmplx_0(np__1))|~r1_int_1(k6_xcmplx_0(np__2,np__1),np__1))).#trainneg
cnf(i_0_165121, plain, (k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))))=k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_5106, plain, (k3_int_2(k4_xcmplx_0(X1),k2_int_5(X2,X3))=k3_int_2(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|k1_int_2(k4_xcmplx_0(k2_int_5(X2,X3)))=k4_xcmplx_0(k2_int_5(X2,X3))|k1_int_2(X1)=X1|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_1(X1)|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_167756, plain, (r1_int_1(k6_xcmplx_0(np__1,np__1),k4_xcmplx_0(np__1))|~r1_int_1(k6_xcmplx_0(np__1,np__1),np__1))).#trainneg
cnf(i_0_5108, plain, (k3_int_2(k4_xcmplx_0(X1),X2)=k3_int_2(X1,k4_xcmplx_0(X2))|k1_int_2(k4_xcmplx_0(X2))=k4_xcmplx_0(X2)|k1_int_2(X1)=X1|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_162230, plain, (v1_int_1(k6_xcmplx_0(np__0,X1))|~v1_int_1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_165892, plain, (v1_int_1(k6_xcmplx_0(X1,np__1))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5113, plain, (k3_int_2(k4_xcmplx_0(k7_xcmplx_0(X1,X2)),X3)=k3_int_2(k7_xcmplx_0(X1,X2),k4_xcmplx_0(X3))|k1_int_2(X3)=X3|k1_int_2(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5115, plain, (k3_int_2(k4_xcmplx_0(k6_xcmplx_0(X1,X2)),X3)=k3_int_2(k6_xcmplx_0(X1,X2),k4_xcmplx_0(X3))|k1_int_2(X3)=X3|k1_int_2(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))=k4_xcmplx_0(k6_xcmplx_0(X1,X2))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_5116, plain, (k3_int_2(k4_xcmplx_0(k1_newton(X1,X2)),X3)=k3_int_2(k1_newton(X1,X2),k4_xcmplx_0(X3))|k1_int_2(X3)=X3|k1_int_2(k4_xcmplx_0(k1_newton(X1,X2)))=k4_xcmplx_0(k1_newton(X1,X2))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5117, plain, (k3_int_2(k4_xcmplx_0(k1_newton(X1,X2)),X3)=k3_int_2(k1_newton(X1,X2),k4_xcmplx_0(X3))|k1_int_2(X3)=X3|k1_int_2(k4_xcmplx_0(k1_newton(X1,X2)))=k4_xcmplx_0(k1_newton(X1,X2))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_163714, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(k4_xcmplx_0(X1)))=k6_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5122, plain, (k3_int_2(k4_xcmplx_0(k2_int_5(X1,X2)),X3)=k3_int_2(k2_int_5(X1,X2),k4_xcmplx_0(X3))|k1_int_2(X3)=X3|k1_int_2(k4_xcmplx_0(k2_int_5(X1,X2)))=k4_xcmplx_0(k2_int_5(X1,X2))|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_5124, plain, (k3_int_2(k4_xcmplx_0(X1),X2)=k3_int_2(X1,k4_xcmplx_0(X2))|k1_int_2(X2)=X2|k1_int_2(k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_121875, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_162079, plain, (k3_int_2(X1,k4_xcmplx_0(np__1))=k3_int_2(X1,np__1)|k1_int_2(X1)=X1|~v7_ordinal1(X1))).#trainneg
cnf(i_0_143610, plain, (r1_xxreal_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),np__2)|~v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)))|~r1_int_1(k6_xcmplx_0(np__2,np__2),np__2))).#trainneg
cnf(i_0_143648, plain, (r1_xxreal_0(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)),np__2)|~v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)))|~r1_int_1(k6_xcmplx_0(np__1,np__2),np__2))).#trainneg
cnf(i_0_166740, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X2))|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_162080, plain, (k3_int_2(X1,k4_xcmplx_0(np__1))=k3_int_2(np__1,X1)|k1_int_2(X1)=X1|~v7_ordinal1(X1))).#trainneg
cnf(i_0_143686, plain, (r1_xxreal_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)),np__2)|~v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)))|~r1_int_1(k6_xcmplx_0(np__2,np__1),np__2))).#trainneg
cnf(i_0_143724, plain, (r1_xxreal_0(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)),np__2)|~v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)))|~r1_int_1(k6_xcmplx_0(np__1,np__1),np__2))).#trainneg
cnf(i_0_166738, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_88884, plain, (r1_int_1(k2_int_5(X1,X2),k4_xcmplx_0(X3))|~r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),X3)|~v7_ordinal1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v7_ordinal1(X3)|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_120322, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(X1))).#trainneg
cnf(i_0_88868, plain, (r1_int_1(k4_xcmplx_0(X1),k2_int_5(X2,X3))|~r1_int_1(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_164052, plain, (r1_xxreal_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),np__1)|~v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)))|~r1_int_1(k6_xcmplx_0(np__2,np__2),np__1))).#trainneg
cnf(i_0_165500, plain, (r1_xxreal_0(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)),np__1)|~v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)))|~r1_int_1(k6_xcmplx_0(np__1,np__2),np__1))).#trainneg
cnf(i_0_108936, plain, (r1_int_1(k4_xcmplx_0(X1),np__2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(np__2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_167637, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5627, plain, (r1_xxreal_0(k1_int_2(X1),X2)|r1_xxreal_0(X2,k1_int_2(X1))|~v1_xreal_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_168214, plain, (v1_int_1(k6_xcmplx_0(k4_xcmplx_0(np__1),np__2)))).#trainneg
cnf(i_0_5668, plain, (k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(X1,np__0)))=k6_xcmplx_0(X1,np__0)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5672, plain, (v7_ordinal1(k6_xcmplx_0(X1,np__0))|k4_xcmplx_0(k6_xcmplx_0(X1,np__0))!=k6_xcmplx_0(X1,np__0)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_168451, plain, (v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_171615, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(np__1))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5673, plain, (v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(X1,np__0)))|k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(X1,np__0)))!=k4_xcmplx_0(k6_xcmplx_0(X1,np__0))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_170917, plain, (r1_int_1(np__1,X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(np__1),X1)|~v7_ordinal1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_5686, plain, (k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(np__0,X1)))=k6_xcmplx_0(np__0,X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_166613, plain, (r1_xxreal_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)),np__1)|~v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__1)))|~r1_int_1(k6_xcmplx_0(np__2,np__1),np__1))).#trainneg
cnf(i_0_170820, plain, (k3_int_2(k4_xcmplx_0(np__1),np__2)=k3_int_2(np__1,np__2)|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_5690, plain, (v7_ordinal1(k6_xcmplx_0(np__0,X1))|k4_xcmplx_0(k6_xcmplx_0(np__0,X1))!=k6_xcmplx_0(np__0,X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_167754, plain, (r1_xxreal_0(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)),np__1)|~v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)))|~r1_int_1(k6_xcmplx_0(np__1,np__1),np__1))).#trainneg
cnf(i_0_170885, plain, (k3_int_2(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_int_2(X1,X2)|k1_int_2(X1)=X1|k1_int_2(X2)=X2|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5691, plain, (v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__0,X1)))|k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(np__0,X1)))!=k4_xcmplx_0(k6_xcmplx_0(np__0,X1))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_129145, plain, (r1_int_1(k4_xcmplx_0(np__2),k6_xcmplx_0(np__2,np__2))|~r1_int_1(np__2,k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_129248, plain, (r1_int_1(k4_xcmplx_0(np__1),k6_xcmplx_0(np__2,np__2))|~r1_int_1(np__1,k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_129295, plain, (r1_int_1(k6_xcmplx_0(np__2,np__2),k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),np__2)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_129296, plain, (r1_int_1(k6_xcmplx_0(np__2,np__2),k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),np__1)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_170897, plain, (r1_int_1(np__2,X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(np__2),X1)|~v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_5861, plain, (k3_int_2(X2,k1_int_2(X1))=k3_int_2(X1,k4_xcmplx_0(X2))|k1_int_2(X2)=X2|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v7_ordinal1(k1_int_2(X1))|~v1_int_1(X2))).#trainneg
cnf(i_0_168447, plain, (v1_int_1(k6_xcmplx_0(X1,np__0))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_131802, plain, (r1_int_1(k4_xcmplx_0(np__2),k6_xcmplx_0(np__1,np__1))|~r1_int_1(np__2,k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1))))).#trainneg
cnf(i_0_5869, plain, (k3_int_2(k1_newton(k4_xcmplx_0(np__1),X2),X1)=k3_int_2(X1,np__1)|k1_int_2(k1_newton(k4_xcmplx_0(np__1),X2))=k1_newton(k4_xcmplx_0(np__1),X2)|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v7_ordinal1(X1)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_168408, plain, (k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(np__0,X1)))=k6_xcmplx_0(np__0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_131905, plain, (r1_int_1(k4_xcmplx_0(np__1),k6_xcmplx_0(np__1,np__1))|~r1_int_1(np__1,k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1))))).#trainneg
cnf(i_0_131952, plain, (r1_int_1(k6_xcmplx_0(np__1,np__1),k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)),np__2)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1))))).#trainneg
cnf(i_0_131953, plain, (r1_int_1(k6_xcmplx_0(np__1,np__1),k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1)),np__1)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__1))))).#trainneg
cnf(i_0_5870, plain, (k3_int_2(k4_xcmplx_0(np__2),X1)=k3_int_2(X1,np__2)|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_129122, plain, (r1_int_1(k6_xcmplx_0(np__2,np__2),k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_129279, plain, (r1_int_1(k4_xcmplx_0(np__2),k6_xcmplx_0(np__2,np__2))|~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_138453, plain, (k3_int_2(k4_xcmplx_0(X1),X2)=k3_int_2(X1,k4_xcmplx_0(X2))|k1_int_2(X2)=X2|k1_int_2(X1)=X1|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_173912, plain, (v7_ordinal1(X1)|k4_xcmplx_0(X1)!=X1|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5872, plain, (k3_int_2(k4_xcmplx_0(k7_xcmplx_0(X2,X3)),X1)=k3_int_2(X1,k7_xcmplx_0(X2,X3))|k1_int_2(k4_xcmplx_0(k7_xcmplx_0(X2,X3)))=k4_xcmplx_0(k7_xcmplx_0(X2,X3))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X2,X3)))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_176345, plain, (v1_int_1(X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5873, plain, (k3_int_2(k4_xcmplx_0(k5_xcmplx_0(X2)),X1)=k3_int_2(X1,k5_xcmplx_0(X2))|k1_int_2(k4_xcmplx_0(k5_xcmplx_0(X2)))=k4_xcmplx_0(k5_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_177974, plain, (k3_int_2(X1,k4_xcmplx_0(X2))=k3_int_2(X1,X2)|k1_int_2(X1)=X1|k1_int_2(X2)=X2|~v7_ordinal1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_5874, plain, (k3_int_2(k4_xcmplx_0(k6_xcmplx_0(X2,X3)),X1)=k3_int_2(X1,k6_xcmplx_0(X2,X3))|k1_int_2(k4_xcmplx_0(k6_xcmplx_0(X2,X3)))=k4_xcmplx_0(k6_xcmplx_0(X2,X3))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(X2,X3)))|~v1_int_1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_5875, plain, (k3_int_2(k4_xcmplx_0(k1_newton(X2,X3)),X1)=k3_int_2(X1,k1_newton(X2,X3))|k1_int_2(k4_xcmplx_0(k1_newton(X2,X3)))=k4_xcmplx_0(k1_newton(X2,X3))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_177856, plain, (k3_int_2(k4_xcmplx_0(X1),np__1)=k3_int_2(X1,np__1)|k1_int_2(X1)=X1|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5876, plain, (k3_int_2(k4_xcmplx_0(k1_newton(X2,X3)),X1)=k3_int_2(X1,k1_newton(X2,X3))|k1_int_2(k4_xcmplx_0(k1_newton(X2,X3)))=k4_xcmplx_0(k1_newton(X2,X3))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_176437, plain, (k3_int_2(k4_xcmplx_0(np__1),X1)=k3_int_2(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5877, plain, (k3_int_2(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)=k3_int_2(X1,k4_xcmplx_0(X2))|k1_int_2(k4_xcmplx_0(k4_xcmplx_0(X2)))=k4_xcmplx_0(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5881, plain, (k3_int_2(k4_xcmplx_0(k2_int_5(X2,X3)),X1)=k3_int_2(X1,k2_int_5(X2,X3))|k1_int_2(k4_xcmplx_0(k2_int_5(X2,X3)))=k4_xcmplx_0(k2_int_5(X2,X3))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_5882, plain, (k3_int_2(k4_xcmplx_0(k1_int_2(X2)),X1)=k3_int_2(X1,k1_int_2(X2))|k1_int_2(k4_xcmplx_0(k1_int_2(X2)))=k4_xcmplx_0(k1_int_2(X2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k1_int_2(X2)))|~v1_int_1(X2))).#trainneg
cnf(i_0_5883, plain, (k3_int_2(k4_xcmplx_0(X2),X1)=k3_int_2(X1,X2)|k1_int_2(k4_xcmplx_0(X2))=k4_xcmplx_0(X2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_176287, plain, (k3_int_2(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))=k3_int_2(X1,np__1)|k1_int_2(X1)=X1|~v1_int_1(X1))).#trainneg
cnf(i_0_177893, plain, (k3_int_2(X1,k4_xcmplx_0(X2))=k3_int_2(X2,k4_xcmplx_0(X1))|k1_int_2(X1)=X1|k1_int_2(X2)=X2|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5895, plain, (k3_int_2(np__1,k4_xcmplx_0(X1))=k3_int_2(np__1,X1)|k1_int_2(X1)=X1|~v1_int_1(X1))).#trainneg
cnf(i_0_172790, plain, (k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(np__1),np__2)))=k6_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_175623, plain, (k3_int_2(k4_xcmplx_0(X2),k4_xcmplx_0(X1))=k3_int_2(X1,X2)|k1_int_2(X2)=X2|k1_int_2(X1)=X1|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_124234, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(k4_xcmplx_0(X2),np__1)|~v7_ordinal1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(X2,k1_newton(k4_xcmplx_0(np__1),X1))|~v7_ordinal1(X2)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_87317, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_172794, plain, (v7_ordinal1(k6_xcmplx_0(k4_xcmplx_0(np__1),np__2))|k4_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(np__1),np__2))!=k6_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_177608, plain, (k3_int_2(np__0,k4_xcmplx_0(np__2))=k3_int_2(np__0,np__2)|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_69936, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1))).#trainneg
cnf(i_0_154904, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_180916, plain, (v1_int_1(k4_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(np__1),np__2))))).#trainneg
cnf(i_0_179971, plain, (v1_xreal_0(k6_xcmplx_0(k4_xcmplx_0(np__1),np__2))|~v1_xreal_0(k4_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(np__1),np__2))))).#trainneg
cnf(i_0_179973, plain, (v1_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(np__1),np__2))|~v1_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(np__1),np__2))))).#trainneg
cnf(i_0_180278, plain, (v1_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(np__1),np__2))|~m1_subset_1(k4_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(np__1),np__2)),k1_numbers))).#trainneg
cnf(i_0_180841, plain, (v1_xreal_0(k6_xcmplx_0(k4_xcmplx_0(np__1),np__2))|~m1_subset_1(k4_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(np__1),np__2)),k1_numbers))).#trainneg
cnf(i_0_179107, plain, (k3_int_2(X2,k4_xcmplx_0(X1))=k3_int_2(X1,X2)|k1_int_2(X1)=X1|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_181929, plain, (r1_int_1(k2_int_5(X1,X2),X3)|~r1_int_1(k2_int_5(X1,X2),k4_xcmplx_0(k4_xcmplx_0(X3)))|~v7_ordinal1(X3)|~v1_int_1(k2_int_5(X1,X2))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_178040, plain, (v7_ordinal1(k4_xcmplx_0(X1))|k4_xcmplx_0(k4_xcmplx_0(X1))!=k4_xcmplx_0(X1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_178017, plain, (v7_ordinal1(k5_xcmplx_0(X1))|k4_xcmplx_0(k5_xcmplx_0(X1))!=k5_xcmplx_0(X1)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6047, plain, (r1_int_1(X1,np__2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_6048, plain, (r1_int_1(X1,np__1)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_181837, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1))).#trainneg
cnf(i_0_181872, plain, (r1_int_1(np__0,X1)|~r1_int_1(np__0,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_179360, plain, (k3_int_2(X2,k4_xcmplx_0(X1))=k3_int_2(X2,X1)|k1_int_2(k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_178353, plain, (k3_int_2(k4_xcmplx_0(X1),X2)=k3_int_2(X2,X1)|k1_int_2(X2)=X2|k1_int_2(X1)=X1|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_126489, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(X2))).#trainneg
cnf(i_0_178960, plain, (k3_int_2(X1,np__1)=k3_int_2(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_153900, plain, (~r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(k1_nat_1(X1,np__1)),k4_xcmplx_0(np__2))|~v7_ordinal1(k1_nat_1(X1,np__1)))).#trainneg
cnf(i_0_160336, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_73595, plain, (r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),k2_int_5(X2,X3))|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v7_ordinal1(k2_int_5(X2,X3))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~r1_int_1(np__1,k2_int_5(X2,X3))|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_1(X2)|~v1_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_159209, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1)|~r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(k4_xcmplx_0(np__2))|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_6086, plain, (k3_int_2(k1_newton(k4_xcmplx_0(np__1),X2),X1)=k3_int_2(X1,np__1)|k1_int_2(k1_newton(k4_xcmplx_0(np__1),X2))=k1_newton(k4_xcmplx_0(np__1),X2)|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v1_int_1(X1)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_178886, plain, (k3_int_2(k4_xcmplx_0(X1),np__1)=k3_int_2(k4_xcmplx_0(np__1),X1)|k1_int_2(X1)=X1|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_6087, plain, (k3_int_2(k4_xcmplx_0(np__2),X1)=k3_int_2(X1,np__2)|k1_int_2(k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2)|~v1_int_1(X1))).#trainneg
cnf(i_0_177790, plain, (r1_int_1(k4_xcmplx_0(np__2),k6_xcmplx_0(np__2,np__2))|~r1_int_1(np__2,k6_xcmplx_0(np__2,np__2)))).#trainneg
cnf(i_0_6089, plain, (k3_int_2(k4_xcmplx_0(k7_xcmplx_0(X2,X3)),X1)=k3_int_2(X1,k7_xcmplx_0(X2,X3))|k1_int_2(k4_xcmplx_0(k7_xcmplx_0(X2,X3)))=k4_xcmplx_0(k7_xcmplx_0(X2,X3))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X2,X3)))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_178899, plain, (k3_int_2(X1,k4_xcmplx_0(np__1))=k3_int_2(X1,np__1)|~v7_ordinal1(k4_xcmplx_0(np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6090, plain, (k3_int_2(k4_xcmplx_0(k5_xcmplx_0(X2)),X1)=k3_int_2(X1,k5_xcmplx_0(X2))|k1_int_2(k4_xcmplx_0(k5_xcmplx_0(X2)))=k4_xcmplx_0(k5_xcmplx_0(X2))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_6091, plain, (k3_int_2(k4_xcmplx_0(k6_xcmplx_0(X2,X3)),X1)=k3_int_2(X1,k6_xcmplx_0(X2,X3))|k1_int_2(k4_xcmplx_0(k6_xcmplx_0(X2,X3)))=k4_xcmplx_0(k6_xcmplx_0(X2,X3))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(X2,X3)))|~v1_int_1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_181874, plain, (r1_int_1(k4_xcmplx_0(np__2),X1)|~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_181875, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_182260, plain, (v1_xreal_0(k6_xcmplx_0(k4_xcmplx_0(np__1),np__2)))).#trainneg
cnf(i_0_182278, plain, (v1_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(np__1),np__2)))).#trainneg
cnf(i_0_6092, plain, (k3_int_2(k4_xcmplx_0(k1_newton(X2,X3)),X1)=k3_int_2(X1,k1_newton(X2,X3))|k1_int_2(k4_xcmplx_0(k1_newton(X2,X3)))=k4_xcmplx_0(k1_newton(X2,X3))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_184399, plain, (r1_int_1(np__2,k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)))|~r1_int_1(np__2,k6_xcmplx_0(np__2,np__2)))).#trainneg
cnf(i_0_6093, plain, (k3_int_2(k4_xcmplx_0(k1_newton(X2,X3)),X1)=k3_int_2(X1,k1_newton(X2,X3))|k1_int_2(k4_xcmplx_0(k1_newton(X2,X3)))=k4_xcmplx_0(k1_newton(X2,X3))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_6094, plain, (k3_int_2(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)=k3_int_2(X1,k4_xcmplx_0(X2))|k1_int_2(k4_xcmplx_0(k4_xcmplx_0(X2)))=k4_xcmplx_0(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6098, plain, (k3_int_2(k4_xcmplx_0(k2_int_5(X2,X3)),X1)=k3_int_2(X1,k2_int_5(X2,X3))|k1_int_2(k4_xcmplx_0(k2_int_5(X2,X3)))=k4_xcmplx_0(k2_int_5(X2,X3))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_183472, plain, (k3_int_2(X1,k4_xcmplx_0(X2))=k3_int_2(X1,X2)|k1_int_2(X2)=X2|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_182406, plain, (k3_int_2(X1,k4_xcmplx_0(X2))=k3_int_2(X1,X2)|k1_int_2(X2)=X2|k1_int_2(X1)=X1|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_184090, plain, (r1_xxreal_0(np__2,X1)|r1_xxreal_0(X1,np__0)|~r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1)|k4_xcmplx_0(np__2)!=np__2)).#trainneg
cnf(i_0_6099, plain, (k3_int_2(k4_xcmplx_0(k1_int_2(X2)),X1)=k3_int_2(X1,k1_int_2(X2))|k1_int_2(k4_xcmplx_0(k1_int_2(X2)))=k4_xcmplx_0(k1_int_2(X2))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(k1_int_2(X2)))|~v1_int_1(X2))).#trainneg
cnf(i_0_6100, plain, (k3_int_2(k4_xcmplx_0(X2),X1)=k3_int_2(X1,X2)|k1_int_2(k4_xcmplx_0(X2))=k4_xcmplx_0(X2)|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6101, plain, (v1_int_1(k6_xcmplx_0(X2,X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_182390, plain, (k3_int_2(k4_xcmplx_0(X1),X2)=k3_int_2(X1,X2)|k1_int_2(X2)=X2|k1_int_2(X1)=X1|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6102, plain, (k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(X2,X1)))=k6_xcmplx_0(X2,X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_182386, plain, (k3_int_2(k4_xcmplx_0(X2),X1)=k3_int_2(X2,X1)|k1_int_2(X2)=X2|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_124233, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(k4_xcmplx_0(X2),np__1)|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(X2,k1_newton(k4_xcmplx_0(np__1),X1))|~v7_ordinal1(X2)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_88516, plain, (r1_int_1(X1,k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X2)))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X2)))|~r1_int_1(k4_xcmplx_0(X1),np__1)|~v1_int_1(k4_xcmplx_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_138459, plain, (k3_int_2(X1,k1_int_2(X2))=k3_int_2(k4_xcmplx_0(X1),X2)|k1_int_2(X1)=X1|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_int_1(k1_int_2(X2)))).#trainneg
cnf(i_0_136233, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_178678, plain, (k3_int_2(k4_xcmplx_0(X1),X2)=k3_int_2(X2,X1)|k1_int_2(k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6109, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)=k6_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_124236, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X2)),np__1)|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(X2,k1_newton(k4_xcmplx_0(np__1),X1))|~v7_ordinal1(X2)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_6115, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k7_xcmplx_0(X2,X3))=k6_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(X2,X3)),X1)|~v7_ordinal1(X1)|~v1_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(X2,X3)))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_185296, plain, (v1_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(np__1),np__2))))).#trainneg
cnf(i_0_6117, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k6_xcmplx_0(X2,X3))=k6_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(X2,X3)),X1)|~v7_ordinal1(X1)|~v1_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(X2,X3)))|~v1_int_1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_185914, plain, (k3_int_2(k4_xcmplx_0(X2),X1)=k3_int_2(X1,X2)|k1_int_2(X2)=X2|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6126, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),X2)=k6_xcmplx_0(k4_xcmplx_0(X2),X1)|~v7_ordinal1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6131, plain, (k6_xcmplx_0(k7_xcmplx_0(X1,X2),k4_xcmplx_0(X3))=k6_xcmplx_0(X3,k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~v7_ordinal1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_28324, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__2))|~r1_int_1(X1,np__2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6133, plain, (k6_xcmplx_0(k6_xcmplx_0(X1,X2),k4_xcmplx_0(X3))=k6_xcmplx_0(X3,k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v1_xcmplx_0(X3)|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_188750, plain, (k3_int_2(X2,k4_xcmplx_0(X1))=k3_int_2(X2,X1)|k1_int_2(k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6148, plain, (r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)),X2)|~v1_int_1(k5_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_186082, plain, (k3_int_2(k4_xcmplx_0(X1),X2)=k3_int_2(X2,X1)|k1_int_2(X2)=X2|k1_int_2(X1)=X1|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6150, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k4_xcmplx_0(X3))|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~v1_int_1(k1_newton(X1,X2))|~v1_int_1(X3)|~v1_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6151, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),k4_xcmplx_0(X3))|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~v1_int_1(k1_newton(X1,X2))|~v1_int_1(X3)|~v1_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_180848, plain, (v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(np__1),np__2)))|k4_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(np__1),np__2))!=k6_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_6152, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_124244, plain, (r1_int_1(k4_xcmplx_0(X1),k7_xcmplx_0(X2,X3))|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X2,X3)))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_int_1(X1,k4_xcmplx_0(k7_xcmplx_0(X2,X3)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6156, plain, (r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),k4_xcmplx_0(X3))|~r1_int_1(k4_xcmplx_0(k2_int_5(X1,X2)),X3)|~v1_int_1(k2_int_5(X1,X2))|~v1_int_1(X3)|~v1_xcmplx_0(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_6157, plain, (r1_int_1(k4_xcmplx_0(k1_int_2(X1)),k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(k1_int_2(X1)),X2)|~v1_int_1(k1_int_2(X1))|~v1_int_1(X2)|~v1_xcmplx_0(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_129280, plain, (r1_int_1(k4_xcmplx_0(np__1),k6_xcmplx_0(np__2,np__2))|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_129344, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),np__1)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_129345, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),np__2)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),k4_xcmplx_0(np__2))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_129366, plain, (r1_int_1(k6_xcmplx_0(np__2,np__2),k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),k4_xcmplx_0(np__2))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2))))).#trainneg
cnf(i_0_87261, plain, (r1_int_1(X1,np__1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),np__1))).#trainneg
cnf(i_0_185011, plain, (r1_int_1(k4_xcmplx_0(np__2),X1)|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_192591, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_191042, plain, (k3_int_2(X1,k4_xcmplx_0(X2))=k3_int_2(X1,X2)|k1_int_2(X2)=X2|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xcmplx_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_192606, plain, (r1_int_1(k4_xcmplx_0(np__1),k6_xcmplx_0(np__2,np__2))|~r1_int_1(np__1,k6_xcmplx_0(np__2,np__2)))).#trainneg
cnf(i_0_6175, plain, (r1_int_1(X1,np__2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_6176, plain, (r1_int_1(X1,np__1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1))))).#trainneg
cnf(i_0_185192, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_6209, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)),k1_newton(X2,X3))=k6_xcmplx_0(X1,k1_newton(X2,X3))|~v7_ordinal1(X3)|~v1_xcmplx_0(k4_xcmplx_0(k1_newton(X2,X3)))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v7_ordinal1(X1)|~v7_ordinal1(k1_newton(X2,X3)))).#trainneg
cnf(i_0_6210, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)=k6_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6216, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k7_xcmplx_0(X2,X3))=k6_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(X2,X3)),X1)|~v7_ordinal1(k4_xcmplx_0(k7_xcmplx_0(X2,X3)))|~v7_ordinal1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_103265, plain, (r1_int_1(X1,k4_xcmplx_0(k2_int_5(X2,X3)))|~v7_ordinal1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(X1),k2_int_5(X2,X3))|~v1_int_2(X3)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_6218, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k6_xcmplx_0(X2,X3))=k6_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(X2,X3)),X1)|~v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(X2,X3)))|~v7_ordinal1(X1)|~v1_int_1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_6219, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k1_newton(X2,X3))=k6_xcmplx_0(k4_xcmplx_0(k1_newton(X2,X3)),X1)|~v7_ordinal1(k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_192680, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),np__1)|~r1_int_1(k6_xcmplx_0(np__2,np__2),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_6220, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k1_newton(X2,X3))=k6_xcmplx_0(k4_xcmplx_0(k1_newton(X2,X3)),X1)|~v7_ordinal1(k4_xcmplx_0(k1_newton(X2,X3)))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_160205, plain, (r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~r1_int_1(np__1,k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_6225, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k2_int_5(X2,X3))=k6_xcmplx_0(k4_xcmplx_0(k2_int_5(X2,X3)),X1)|~v7_ordinal1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v7_ordinal1(X1)|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_192751, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),np__2)|~r1_int_1(k6_xcmplx_0(np__2,np__2),k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_136227, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_168028, plain, (k3_int_2(k4_xcmplx_0(X1),k1_int_2(X2))=k3_int_2(k1_int_2(X1),X2)|k1_int_2(X1)=X1|~v1_int_1(X2)|~v1_int_1(k1_int_2(X1))|~v1_int_1(k1_int_2(X2))|~v1_int_1(X1))).#trainneg
cnf(i_0_184052, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X2)))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X2)))|~v1_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)),np__1)|~v7_ordinal1(X1)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~v1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_160330, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)),X2)|~v7_ordinal1(X2)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_109215, plain, (r1_int_1(np__0,k4_xcmplx_0(X1))|~r1_int_1(np__0,X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6232, plain, (k6_xcmplx_0(k7_xcmplx_0(X1,X2),k4_xcmplx_0(X3))=k6_xcmplx_0(X3,k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~v7_ordinal1(X3)|~v7_ordinal1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6234, plain, (k6_xcmplx_0(k6_xcmplx_0(X1,X2),k4_xcmplx_0(X3))=k6_xcmplx_0(X3,k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v7_ordinal1(X3)|~v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_6235, plain, (k6_xcmplx_0(k1_newton(X1,X2),k4_xcmplx_0(X3))=k6_xcmplx_0(X3,k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X3)|~v7_ordinal1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6236, plain, (k6_xcmplx_0(k1_newton(X1,X2),k4_xcmplx_0(X3))=k6_xcmplx_0(X3,k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X3)|~v7_ordinal1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_193516, plain, (r1_int_1(k4_xcmplx_0(X1),np__1)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v1_int_1(X1))).#trainneg
cnf(i_0_6241, plain, (k6_xcmplx_0(k2_int_5(X1,X2),k4_xcmplx_0(X3))=k6_xcmplx_0(X3,k4_xcmplx_0(k2_int_5(X1,X2)))|~v7_ordinal1(X3)|~v7_ordinal1(k4_xcmplx_0(k2_int_5(X1,X2)))|~v1_int_2(X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_193847, plain, (r1_int_1(X1,np__2)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),np__2)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_193515, plain, (r1_int_1(k4_xcmplx_0(X1),np__2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__2))|~v1_int_1(X1))).#trainneg
cnf(i_0_193751, plain, (r1_int_1(np__1,k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)))|~r1_int_1(np__1,k6_xcmplx_0(np__2,np__2)))).#trainneg
cnf(i_0_193544, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6254, plain, (r1_int_1(X1,k4_xcmplx_0(np__2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_64668, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))))|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_193463, plain, (~r1_xxreal_0(np__2,X1)|~v7_ordinal1(k1_nat_1(X1,np__1))|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(k1_nat_1(X1,np__1)),np__2)|~v1_int_1(k1_nat_1(X1,np__1)))).#trainneg
cnf(i_0_193336, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_72036, plain, (r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),k4_xcmplx_0(X2))|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(np__1,X2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_121360, plain, (r1_int_1(k4_xcmplx_0(X1),k2_int_5(X2,X3))|~r1_int_1(X1,k2_int_5(X2,X3))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_6322, plain, (v1_xreal_0(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1)|~m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))),k1_numbers))).#trainneg
cnf(i_0_70151, plain, (r1_int_1(np__1,k4_xcmplx_0(X2))|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),k4_xcmplx_0(X2))|~v1_int_1(X2)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_69688, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_108404, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__0)|r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(X2,X1)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_71207, plain, (r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),X3)|~r1_int_1(k1_newton(X1,X2),X3)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6648, plain, (k3_int_2(np__2,X1)=k3_int_2(X1,np__2)|v7_ordinal1(k4_xcmplx_0(np__2))|~v1_int_1(X1))).#trainneg
cnf(i_0_199267, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1)|~v1_int_1(X2)|~r1_int_1(X1,k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_6822, plain, (r1_int_1(k1_newton(X1,X2),k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X3)))|k1_newton(k4_xcmplx_0(np__1),X3)=np__1|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X3))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(k1_newton(X1,X2)),np__1)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_6826, plain, (r1_int_1(X1,k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X2)))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(X1),np__1)|~v1_int_1(k4_xcmplx_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_199821, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~v1_int_1(X1))).#trainneg
cnf(i_0_198274, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(X1),np__2)|~v7_ordinal1(k4_xcmplx_0(np__2))|~v1_int_1(X1))).#trainneg
cnf(i_0_198134, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6837, plain, (r1_int_1(X1,k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X2)))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X2)))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(X1),np__1)|~v1_int_1(k4_xcmplx_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_200110, plain, (r1_int_1(np__1,X2)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_6838, plain, (r1_xxreal_0(k1_newton(k4_xcmplx_0(np__1),X1),np__0)|r1_xxreal_0(X2,k1_newton(k4_xcmplx_0(np__1),X1))|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v7_ordinal1(X2)|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1))|~r1_int_1(X2,np__1)|~v1_int_1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_198132, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,np__2)|~v7_ordinal1(k4_xcmplx_0(np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_198275, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(X1),np__1)|~v7_ordinal1(k4_xcmplx_0(np__1))|~v1_int_1(X1))).#trainneg
cnf(i_0_198135, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6843, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(k4_xcmplx_0(np__2),np__1)|~m1_subset_1(X1,k5_numbers)|~r1_int_1(np__2,k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_198492, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v1_int_1(X1)|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_197593, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6844, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(k4_xcmplx_0(np__2),np__1)|~m1_subset_1(X1,k5_numbers)|~r1_int_1(np__2,k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_199679, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)),X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))))).#trainneg
cnf(i_0_6845, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(k4_xcmplx_0(np__1),np__1)|~m1_subset_1(X1,k5_numbers)|~r1_int_1(np__1,k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_6846, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(k4_xcmplx_0(np__1),np__1)|~m1_subset_1(X1,k5_numbers)|~r1_int_1(np__1,k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_6849, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)),np__1)|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k5_xcmplx_0(X2),k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_200124, plain, (k3_int_2(np__2,k4_xcmplx_0(np__1))=np__1|v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_6853, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(k4_xcmplx_0(X2),np__1)|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(X2,k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(X2))).#trainneg
cnf(i_0_6857, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(k4_xcmplx_0(X2),np__1)|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(X2,k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_200111, plain, (k1_newton(k4_xcmplx_0(np__1),k3_nat_d(k7_nat_d(X1,np__1),np__2))=np__1|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),X1),k4_xcmplx_0(np__1),X1)|r1_xxreal_0(X1,np__2)|k3_int_2(k4_xcmplx_0(np__1),X1)!=np__1|~v1_int_2(X1)|~v7_ordinal1(X1)|k4_xcmplx_0(k3_nat_d(k7_nat_d(X1,np__1),np__2))!=k3_nat_d(k7_nat_d(X1,np__1),np__2)|~v1_int_1(k3_nat_d(k7_nat_d(X1,np__1),np__2)))).#trainneg
cnf(i_0_6859, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(k4_xcmplx_0(X2),np__1)|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(X2,k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~v1_int_1(X2))).#trainneg
cnf(i_0_7188, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__0)|r1_xxreal_0(np__2,k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(np__2,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_195332, plain, (r1_xxreal_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),np__1)|~v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)))|~r1_int_1(k6_xcmplx_0(np__2,np__2),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_7210, plain, (r1_int_1(np__2,k5_xcmplx_0(X1))|~r1_int_1(np__2,k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_195854, plain, (r1_xxreal_0(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)),np__2)|~v7_ordinal1(k4_xcmplx_0(k6_xcmplx_0(np__2,np__2)))|~r1_int_1(k6_xcmplx_0(np__2,np__2),k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_129996, plain, (r1_int_1(k6_xcmplx_0(np__1,np__2),k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)),k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_130141, plain, (r1_int_1(np__1,k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)))|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_200755, plain, (r1_int_1(X1,k1_newton(k4_xcmplx_0(np__1),X2))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_201587, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_7212, plain, (r1_int_1(np__2,k1_newton(X1,X2))|~r1_int_1(np__2,k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_130143, plain, (r1_int_1(np__2,k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)))|~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_130153, plain, (r1_int_1(k4_xcmplx_0(np__2),k6_xcmplx_0(np__1,np__2))|~r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_130154, plain, (r1_int_1(k4_xcmplx_0(np__1),k6_xcmplx_0(np__1,np__2))|~r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_130218, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)),np__1)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)),k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_130219, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)),np__2)|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)),k4_xcmplx_0(np__2))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_202815, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X2)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|~v1_int_1(X2))).#trainneg
cnf(i_0_7213, plain, (r1_int_1(np__2,k1_newton(X1,X2))|~r1_int_1(np__2,k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_130240, plain, (r1_int_1(k6_xcmplx_0(np__1,np__2),k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)),k4_xcmplx_0(np__2))|~v1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_202362, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_203860, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7217, plain, (r1_int_1(np__2,k4_xcmplx_0(k5_xcmplx_0(X1)))|~r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7219, plain, (r1_int_1(np__2,k1_int_2(X1))|~r1_int_1(np__2,k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_203578, plain, (r1_int_1(k6_xcmplx_0(np__1,np__2),k4_xcmplx_0(np__1))|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)),np__1))).#trainneg
cnf(i_0_203433, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k4_xcmplx_0(np__2),k1_newton(k4_xcmplx_0(np__1),X1))|~v1_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_203083, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k4_xcmplx_0(np__2),k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_199349, plain, (r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v7_ordinal1(X2)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~r1_int_1(np__1,X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_7393, plain, (r1_int_1(np__1,k4_xcmplx_0(k5_xcmplx_0(X1)))|~r1_int_1(np__1,k5_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_203651, plain, (r1_int_1(np__1,k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)))|~r1_int_1(k4_xcmplx_0(np__1),k6_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_199347, plain, (r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1))|~v7_ordinal1(X2)|~r1_int_1(np__1,X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_7402, plain, (r1_int_1(np__1,k4_xcmplx_0(k1_int_2(X1)))|~r1_int_1(np__1,k1_int_2(X1))|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_199337, plain, (r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),k2_int_5(X2,X3))|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v1_int_2(X3)|~v7_ordinal1(X3)|~v1_int_1(k4_xcmplx_0(k2_int_5(X2,X3)))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(X2)|~r1_int_1(np__1,k2_int_5(X2,X3))|~v1_xcmplx_0(k2_int_5(X2,X3))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_7419, plain, (k3_int_2(k4_xcmplx_0(np__2),k1_int_2(X1))=k3_int_2(k4_xcmplx_0(np__2),X1)|k18_complex1(np__2)=np__2|~v1_int_1(X1))).#trainneg
cnf(i_0_204115, plain, (r1_int_1(k4_xcmplx_0(np__2),k6_xcmplx_0(np__1,np__2))|~r1_int_1(np__2,k6_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_7426, plain, (k3_int_2(k4_xcmplx_0(np__2),X1)=k3_int_2(X1,k4_xcmplx_0(np__2))|k18_complex1(np__2)=np__2|~v1_int_1(X1))).#trainneg
cnf(i_0_204029, plain, (r1_int_1(np__2,k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)))|~r1_int_1(k4_xcmplx_0(np__2),k6_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_199680, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)),X2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_198921, plain, (~r1_int_1(k4_xcmplx_0(k2_xcmplx_0(X1,np__1)),np__2)|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7429, plain, (k3_int_2(k1_int_2(X1),k4_xcmplx_0(np__2))=k3_int_2(X1,k4_xcmplx_0(np__2))|k18_complex1(np__2)=np__2|~v1_int_1(X1))).#trainneg
cnf(i_0_202367, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_198456, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_204188, plain, (r1_int_1(k4_xcmplx_0(np__1),k6_xcmplx_0(np__1,np__2))|~r1_int_1(np__1,k6_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_199344, plain, (r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_xcmplx_0(X2)|~r1_int_1(np__1,X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_199677, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)),X2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~v1_int_1(X2))).#trainneg
cnf(i_0_7444, plain, (r1_int_1(X1,k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X2)))|k1_newton(k4_xcmplx_0(np__1),X2)=np__1|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X2))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~v1_int_1(k4_xcmplx_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_7456, plain, (r1_xxreal_0(k1_newton(k4_xcmplx_0(np__1),X1),np__0)|r1_xxreal_0(X2,k1_newton(k4_xcmplx_0(np__1),X1))|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v7_ordinal1(X2)|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1))|~r1_int_1(k4_xcmplx_0(X2),np__1)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_206003, plain, (k3_int_2(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2))=np__1|k18_complex1(np__2)=np__2)).#trainneg
cnf(i_0_7665, plain, (v1_int_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(X1,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_207693, plain, (r1_int_1(X1,np__1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~v1_int_1(X1))).#trainneg
cnf(i_0_7667, plain, (v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k1_newton(X1,X2))))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7671, plain, (v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1)))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7699, plain, (r1_int_1(np__1,np__1)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(np__1,k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_204561, plain, (r1_int_1(k6_xcmplx_0(np__1,np__2),k4_xcmplx_0(np__2))|~r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)),np__2))).#trainneg
cnf(i_0_7703, plain, (r1_int_1(np__1,k5_xcmplx_0(X1))|~r1_int_1(np__1,k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7705, plain, (r1_int_1(np__1,k1_newton(X1,X2))|~r1_int_1(np__1,k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7706, plain, (r1_int_1(np__1,k1_newton(X1,X2))|~r1_int_1(np__1,k4_xcmplx_0(k1_newton(X1,X2)))|~v1_int_1(k4_xcmplx_0(k1_newton(X1,X2)))|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_197592, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),np__1)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7710, plain, (r1_int_1(np__1,k4_xcmplx_0(k5_xcmplx_0(X1)))|~r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|~v1_int_1(k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7712, plain, (r1_int_1(np__1,k1_int_2(X1))|~r1_int_1(np__1,k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_206339, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|~v1_int_1(X2))).#trainneg
cnf(i_0_7716, plain, (k3_int_2(k4_xcmplx_0(np__2),k1_int_2(X1))=k3_int_2(k4_xcmplx_0(np__2),X1)|k1_int_2(np__2)=np__2|~v1_int_1(X1))).#trainneg
cnf(i_0_205555, plain, (r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_xcmplx_0(X2)|~r1_int_1(np__1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_7723, plain, (k3_int_2(k4_xcmplx_0(np__2),X1)=k3_int_2(X1,k4_xcmplx_0(np__2))|k1_int_2(np__2)=np__2|~v1_int_1(X1))).#trainneg
cnf(i_0_199244, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1)|~v1_int_1(X2)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_7726, plain, (k3_int_2(k1_int_2(X1),k4_xcmplx_0(np__2))=k3_int_2(X1,k4_xcmplx_0(np__2))|k1_int_2(np__2)=np__2|~v1_int_1(X1))).#trainneg
cnf(i_0_199352, plain, (r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~r1_int_1(np__1,X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_198922, plain, (~r1_int_1(k4_xcmplx_0(k2_xcmplx_0(np__1,X1)),np__2)|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(k2_xcmplx_0(np__1,X1))|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_7728, plain, (k3_int_2(k4_xcmplx_0(X1),k4_xcmplx_0(np__2))=k3_int_2(X1,k4_xcmplx_0(np__2))|k1_int_2(X1)=X1|k1_int_2(np__2)=np__2|~v1_int_1(X1))).#trainneg
cnf(i_0_7761, plain, (r1_int_1(np__2,k4_xcmplx_0(k5_xcmplx_0(X1)))|~r1_int_1(np__2,k5_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_204114, plain, (r1_int_1(k4_xcmplx_0(np__2),k6_xcmplx_0(np__1,np__2))|~r1_int_1(np__2,k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_198303, plain, (r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_202364, plain, (r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7770, plain, (r1_int_1(np__2,k4_xcmplx_0(k1_int_2(X1)))|~r1_int_1(np__2,k1_int_2(X1))|~v1_int_1(k4_xcmplx_0(k1_int_2(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_7879, plain, (k3_int_2(np__2,k1_int_2(X1))=k3_int_2(np__2,X1)|m1_subset_1(k4_xcmplx_0(np__2),k5_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_7886, plain, (k3_int_2(np__2,X1)=k3_int_2(X1,np__2)|m1_subset_1(k4_xcmplx_0(np__2),k5_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_208729, plain, (k3_int_2(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2))=np__1|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_7889, plain, (k3_int_2(k1_int_2(X1),np__2)=k3_int_2(X1,np__2)|m1_subset_1(k4_xcmplx_0(np__2),k5_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_7891, plain, (k3_int_2(k4_xcmplx_0(X1),np__2)=k3_int_2(X1,np__2)|k1_int_2(X1)=X1|m1_subset_1(k4_xcmplx_0(np__2),k5_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_208842, plain, (r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(k4_xcmplx_0(X1)))|~r1_int_1(np__2,k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8013, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__1)|r1_xxreal_0(np__2,k4_xcmplx_0(X1))|k4_xcmplx_0(k4_xcmplx_0(X1))!=k4_xcmplx_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_208623, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~r1_int_1(np__1,k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_208843, plain, (r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(k4_xcmplx_0(X1)))|~r1_int_1(np__1,k4_xcmplx_0(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8014, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__0)|r1_xxreal_0(np__1,k4_xcmplx_0(X1))|k4_xcmplx_0(k4_xcmplx_0(X1))!=k4_xcmplx_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_204187, plain, (r1_int_1(k4_xcmplx_0(np__1),k6_xcmplx_0(np__1,np__2))|~r1_int_1(np__1,k4_xcmplx_0(k6_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_208957, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(X2)|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_204263, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)),np__1)|~r1_int_1(k6_xcmplx_0(np__1,np__2),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_204336, plain, (r1_int_1(k4_xcmplx_0(k6_xcmplx_0(np__1,np__2)),np__2)|~r1_int_1(k6_xcmplx_0(np__1,np__2),k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_209398, plain, (k3_int_2(k4_xcmplx_0(np__2),k4_xcmplx_0(np__2))=k3_int_2(np__2,k4_xcmplx_0(np__2))|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_209328, plain, (~r1_int_1(k4_xcmplx_0(k2_xcmplx_0(X1,np__1)),np__2)|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_208064, plain, (r1_int_1(np__1,k1_newton(k4_xcmplx_0(np__1),X1))|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~r1_int_1(k4_xcmplx_0(np__1),np__1)|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_70274, plain, (v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))))|k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))))!=k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1))))|~v1_int_1(X1))).#trainneg
cnf(i_0_87314, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_121364, plain, (r1_int_1(np__1,k4_xcmplx_0(X2))|k1_newton(k4_xcmplx_0(np__1),X1)=np__1|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),k4_xcmplx_0(X2))|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_208508, plain, (k3_int_2(k4_xcmplx_0(np__2),X1)=k3_int_2(np__2,X1)|k1_int_2(np__2)=np__2|~v1_int_1(X1))).#trainneg
cnf(i_0_208509, plain, (k3_int_2(k4_xcmplx_0(np__2),X1)=k3_int_2(X1,np__2)|k1_int_2(np__2)=np__2|~v1_int_1(X1))).#trainneg
cnf(i_0_211785, plain, (r1_int_1(X1,X2)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_210463, negated_conjecture, (r1_int_1(esk2_0,k6_xcmplx_0(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1))|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),k4_xcmplx_0(np__1),esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1)).#trainneg
cnf(i_0_212663, plain, (k3_int_2(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2))=k3_int_2(k4_xcmplx_0(np__1),np__2)|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_210461, negated_conjecture, (r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1,esk2_0)|r2_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0),k4_xcmplx_0(np__1),esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1)).#trainneg
cnf(i_0_211659, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_211669, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_212519, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)),X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))))).#trainneg
cnf(i_0_212222, plain, (r1_int_1(np__2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(np__2,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_212225, plain, (r1_int_1(k4_xcmplx_0(np__2),k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(np__2,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_212231, plain, (r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(np__1,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_212234, plain, (r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X1))|~r1_int_1(np__1,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_211657, plain, (r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_212520, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)),X2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_212471, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k1_newton(k4_xcmplx_0(np__1),X1),X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_211917, plain, (~r1_xxreal_0(np__2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__1))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k2_xcmplx_0(X1,np__1),k4_xcmplx_0(np__2))|~v1_int_1(k4_xcmplx_0(k2_xcmplx_0(X1,np__1)))|~v1_xcmplx_0(k2_xcmplx_0(X1,np__1)))).#trainneg
cnf(i_0_210493, plain, (k3_int_2(np__2,k4_xcmplx_0(np__1))=np__1|m1_subset_1(k4_xcmplx_0(np__2),k5_numbers))).#trainneg
cnf(i_0_211868, plain, (k3_int_2(k4_xcmplx_0(np__2),np__2)=k3_int_2(np__2,k4_xcmplx_0(np__2))|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_213252, negated_conjecture, (r1_int_1(esk2_0,k6_xcmplx_0(k2_int_5(k4_xcmplx_0(np__1),esk2_0),k4_xcmplx_0(np__1)))|r1_int_1(esk2_0,k6_xcmplx_0(k2_int_5(k4_xcmplx_0(np__1),esk2_0),np__1))|~v1_int_1(esk2_0)|~v1_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0))|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1)).#trainneg
cnf(i_0_215747, plain, (k3_int_2(np__2,k4_xcmplx_0(np__2))=k3_int_2(np__2,np__2)|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_8246, plain, (r1_int_1(X1,k4_xcmplx_0(np__1))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~r1_int_1(k4_xcmplx_0(X1),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_213272, negated_conjecture, (r2_int_1(np__1,k4_xcmplx_0(np__1),esk2_0)|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1|k4_xcmplx_0(np__1)!=np__0)).#trainneg
cnf(i_0_8379, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(np__1,k1_newton(k4_xcmplx_0(np__1),X1)))).#trainneg
cnf(i_0_8380, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(np__1,k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))))).#trainneg
cnf(i_0_213264, negated_conjecture, (r1_int_1(esk2_0,k6_xcmplx_0(k2_int_5(k4_xcmplx_0(np__1),esk2_0),k4_xcmplx_0(np__1)))|~v1_int_1(esk2_0)|~v1_int_1(k2_int_5(k4_xcmplx_0(np__1),esk2_0))|k2_int_5(k4_xcmplx_0(np__1),esk2_0)!=np__1)).#trainneg
cnf(i_0_8383, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)),k4_xcmplx_0(np__1))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k5_xcmplx_0(X2),k1_newton(k4_xcmplx_0(np__1),X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8384, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)),k4_xcmplx_0(np__1))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k4_xcmplx_0(k5_xcmplx_0(X2)))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k5_xcmplx_0(X2),k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_212914, plain, (r1_int_1(X1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_8385, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(k4_xcmplx_0(k1_newton(X2,X3)),k4_xcmplx_0(np__1))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k4_xcmplx_0(k1_newton(X2,X3)))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k1_newton(X2,X3),k1_newton(k4_xcmplx_0(np__1),X1))|~v7_ordinal1(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_212543, plain, (k3_int_2(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2))=k3_int_2(np__2,k4_xcmplx_0(np__1))|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_212659, plain, (k3_int_2(k4_xcmplx_0(np__2),k4_xcmplx_0(np__2))=k3_int_2(k4_xcmplx_0(np__2),np__2)|k1_int_2(np__2)=np__2)).#trainneg
cnf(i_0_8387, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(k4_xcmplx_0(X2),k4_xcmplx_0(np__1))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(X2,k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(X2))).#trainneg
cnf(i_0_214806, plain, (r1_int_1(k4_xcmplx_0(np__1),X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1)|~r1_int_1(np__1,X1))).#trainneg
cnf(i_0_8388, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(k4_xcmplx_0(X2),k4_xcmplx_0(np__1))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(X2,k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~v1_int_1(X2))).#trainneg
cnf(i_0_8392, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(k4_xcmplx_0(X2),k4_xcmplx_0(np__1))|~v1_int_1(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(X2,k1_newton(k4_xcmplx_0(np__1),X1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_212517, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k1_newton(k4_xcmplx_0(np__1),X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)),X2)|~v7_ordinal1(X1)|~v1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~v1_int_1(X2))).#trainneg
cnf(i_0_213208, plain, (r1_int_1(k4_xcmplx_0(X1),X2)|~v7_ordinal1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~r1_int_1(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_int_1(X1))).#trainneg
cnf(i_0_213253, negated_conjecture, (r2_int_1(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1),esk2_0)|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__1|k2_int_5(k4_xcmplx_0(np__1),esk2_0)=np__0|k3_int_2(k4_xcmplx_0(np__1),esk2_0)!=np__1)).#trainneg
cnf(i_0_212113, plain, (k1_newton(k4_xcmplx_0(np__1),X1)=np__1|r1_int_1(np__1,k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(X1,k5_numbers)|~r1_int_1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)),X2)|~v7_ordinal1(k4_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1)))|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_xcmplx_0(k1_newton(k4_xcmplx_0(np__1),X1))|~v1_xcmplx_0(X2))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 80
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 111
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 109
# Processed clauses                    : 26382
# ...of these trivial                  : 187
# ...subsumed                          : 23384
# ...remaining for further processing  : 2811
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 436
# Backward-rewritten                   : 123
# Generated clauses                    : 218100
# ...of the previous two non-trivial   : 79988
# Contextual simplify-reflections      : 697
# Paramodulations                      : 218087
# Factorizations                       : 10
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 2249
#    Positive orientable unit clauses  : 117
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 23
#    Non-unit-clauses                  : 2109
# Current number of unprocessed clauses: 51576
# ...number of literals in the above   : 378349
# Current number of archived formulas  : 0
# Current number of archived clauses   : 3373
# Proof object given clauses           : 202
# Proof search given clauses           : 2811
# Clause-clause subsumption calls (NU) : 1647650
# Rec. Clause-clause subsumption calls : 348217
# Non-unit clause-clause subsumptions  : 22874
# Unit Clause-clause subsumption calls : 6174
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 122
# BW rewrite match successes           : 33
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5031164

# -------------------------------------------------
# User time                : 68.301 s
# System time              : 1.160 s
# Total time               : 69.461 s
# Maximum resident set size: 260428 pages
68.30user 1.19system 2:30.76elapsed 46%CPU (0avgtext+0avgdata 260784maxresident)k
24inputs+1152outputs (0major+257311minor)pagefaults 0swaps
