% Refutation found. Thanks to Tanya!
% SZS status Theorem for t36_ordinal5
% SZS output start Proof for t36_ordinal5
fof(f2863755,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f710,f720,f725,f730,f740,f745,f750,f755,f765,f770,f775,f785,f795,f807,f812,f817,f826,f827,f828,f829,f834,f835,f836,f842,f843,f844,f850,f851,f852,f858,f859,f860,f866,f867,f868,f874,f875,f876,f882,f883,f884,f890,f891,f892,f897,f898,f899,f901,f902,f903,f904,f905,f906,f907,f909,f910,f911,f913,f914,f915,f917,f918,f919,f921,f922,f923,f925,f926,f927,f929,f930,f931,f933,f934,f935,f983,f984,f985,f987,f988,f989,f991,f992,f993,f995,f996,f997,f999,f1000,f1001,f1003,f1004,f1005,f1007,f1008,f1009,f1010,f1011,f1012,f1013,f1015,f1016,f1017,f1019,f1020,f1021,f1022,f1026,f1027,f1028,f1030,f1031,f1032,f1034,f1035,f1036,f1038,f1039,f1040,f1042,f1043,f1044,f1046,f1047,f1048,f1050,f1051,f1052,f1054,f1055,f1056,f1058,f1059,f1060,f1062,f1063,f1064,f1073,f1074,f1075,f1246,f1251,f1256,f1261,f1266,f1271,f1289,f1295,f1301,f1307,f1359,f1377,f1391,f1397,f1495,f1498,f1541,f1575,f1647,f1725,f1895,f2223,f2239,f2395,f2424,f2432,f2836,f2918,f5353,f5712,f5713,f6616,f8488,f9366,f15317,f15450,f23132,f28505,f31201,f50373,f64213,f256572,f268833,f280140,f280269,f283194,f346790,f567577,f567593,f584568,f616838,f622614,f622757,f625196,f628860,f630924,f656277,f1016851,f1029048,f1029081,f1057514,f1237352,f1336124,f1336290,f1338674,f1386539,f1386597,f1970633,f1977665,f1982807,f1987569,f1987571,f1992184,f1993103,f1994774,f1995028,f1995062,f1995104,f1999881,f1999884,f2000076,f2000569,f2006066,f2006133,f2009572,f2030268,f2068141,f2149625,f2345839,f2346351,f2863754])).
fof(f2863754,plain,(
  ~spl14_28 | ~spl14_29 | ~spl14_30 | spl14_557 | spl14_31432),
  inference(avatar_contradiction_clause,[],[f2863753])).
fof(f2863753,plain,(
  $false | (~spl14_28 | ~spl14_29 | ~spl14_30 | spl14_557 | spl14_31432)),
  inference(subsumption_resolution,[],[f2863752,f865])).
fof(f865,plain,(
  v1_relat_1(sK7) | ~spl14_30),
  inference(avatar_component_clause,[],[f863])).
fof(f863,plain,(
  spl14_30 <=> v1_relat_1(sK7)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_30])])).
fof(f2863752,plain,(
  ~v1_relat_1(sK7) | (~spl14_28 | ~spl14_29 | spl14_557 | spl14_31432)),
  inference(subsumption_resolution,[],[f2863751,f857])).
fof(f857,plain,(
  v1_funct_1(sK7) | ~spl14_29),
  inference(avatar_component_clause,[],[f855])).
fof(f855,plain,(
  spl14_29 <=> v1_funct_1(sK7)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_29])])).
fof(f2863751,plain,(
  ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | (~spl14_28 | spl14_557 | spl14_31432)),
  inference(subsumption_resolution,[],[f2863750,f849])).
fof(f849,plain,(
  v5_ordinal1(sK7) | ~spl14_28),
  inference(avatar_component_clause,[],[f847])).
fof(f847,plain,(
  spl14_28 <=> v5_ordinal1(sK7)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_28])])).
fof(f2863750,plain,(
  ~v5_ordinal1(sK7) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | (spl14_557 | spl14_31432)),
  inference(subsumption_resolution,[],[f2863748,f8280])).
fof(f8280,plain,(
  ~v2_ordinal5(sK7) | spl14_557),
  inference(avatar_component_clause,[],[f8279])).
fof(f8279,plain,(
  spl14_557 <=> v2_ordinal5(sK7)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_557])])).
fof(f2863748,plain,(
  v2_ordinal5(sK7) | ~v5_ordinal1(sK7) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | spl14_31432),
  inference(resolution,[],[f745214,f659])).
fof(f659,plain,(
  ( ! [X0] : (r2_hidden(sK10(X0),sK11(X0)) | v2_ordinal5(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0] : (((v2_ordinal5(X0) | ((~r1_tarski(k1_funct_1(X0,sK10(X0)),k1_funct_1(X0,sK11(X0))) & r2_hidden(sK11(X0),k9_xtuple_0(X0)) & r2_hidden(sK10(X0),sK11(X0)) & v3_ordinal1(sK11(X0))) & v3_ordinal1(sK10(X0)))) & (! [X3] : (! [X4] : (r1_tarski(k1_funct_1(X0,X3),k1_funct_1(X0,X4)) | ~r2_hidden(X4,k9_xtuple_0(X0)) | ~r2_hidden(X3,X4) | ~v3_ordinal1(X4)) | ~v3_ordinal1(X3)) | ~v2_ordinal5(X0))) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11])],[f209,f211,f210])).
fof(f210,plain,(
  ! [X0] : (? [X1] : (? [X2] : (~r1_tarski(k1_funct_1(X0,X1),k1_funct_1(X0,X2)) & r2_hidden(X2,k9_xtuple_0(X0)) & r2_hidden(X1,X2) & v3_ordinal1(X2)) & v3_ordinal1(X1)) => (? [X2] : (~r1_tarski(k1_funct_1(X0,sK10(X0)),k1_funct_1(X0,X2)) & r2_hidden(X2,k9_xtuple_0(X0)) & r2_hidden(sK10(X0),X2) & v3_ordinal1(X2)) & v3_ordinal1(sK10(X0))))),
  introduced(choice_axiom,[])).
fof(f211,plain,(
  ! [X0] : (? [X2] : (~r1_tarski(k1_funct_1(X0,sK10(X0)),k1_funct_1(X0,X2)) & r2_hidden(X2,k9_xtuple_0(X0)) & r2_hidden(sK10(X0),X2) & v3_ordinal1(X2)) => (~r1_tarski(k1_funct_1(X0,sK10(X0)),k1_funct_1(X0,sK11(X0))) & r2_hidden(sK11(X0),k9_xtuple_0(X0)) & r2_hidden(sK10(X0),sK11(X0)) & v3_ordinal1(sK11(X0))))),
  introduced(choice_axiom,[])).
fof(f209,plain,(
  ! [X0] : (((v2_ordinal5(X0) | ? [X1] : (? [X2] : (~r1_tarski(k1_funct_1(X0,X1),k1_funct_1(X0,X2)) & r2_hidden(X2,k9_xtuple_0(X0)) & r2_hidden(X1,X2) & v3_ordinal1(X2)) & v3_ordinal1(X1))) & (! [X3] : (! [X4] : (r1_tarski(k1_funct_1(X0,X3),k1_funct_1(X0,X4)) | ~r2_hidden(X4,k9_xtuple_0(X0)) | ~r2_hidden(X3,X4) | ~v3_ordinal1(X4)) | ~v3_ordinal1(X3)) | ~v2_ordinal5(X0))) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))),
  inference(rectify,[],[f208])).
fof(f208,plain,(
  ! [X0] : (((v2_ordinal5(X0) | ? [X1] : (? [X2] : (~r1_tarski(k1_funct_1(X0,X1),k1_funct_1(X0,X2)) & r2_hidden(X2,k9_xtuple_0(X0)) & r2_hidden(X1,X2) & v3_ordinal1(X2)) & v3_ordinal1(X1))) & (! [X1] : (! [X2] : (r1_tarski(k1_funct_1(X0,X1),k1_funct_1(X0,X2)) | ~r2_hidden(X2,k9_xtuple_0(X0)) | ~r2_hidden(X1,X2) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v2_ordinal5(X0))) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0] : ((v2_ordinal5(X0) <=> ! [X1] : (! [X2] : (r1_tarski(k1_funct_1(X0,X1),k1_funct_1(X0,X2)) | ~r2_hidden(X2,k9_xtuple_0(X0)) | ~r2_hidden(X1,X2) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1))) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X0] : ((v2_ordinal5(X0) <=> ! [X1] : (! [X2] : ((r1_tarski(k1_funct_1(X0,X1),k1_funct_1(X0,X2)) | (~r2_hidden(X2,k9_xtuple_0(X0)) | ~r2_hidden(X1,X2))) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1))) | (~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)))),
  inference(ennf_transformation,[],[f17])).
fof(f17,axiom,(
  ! [X0] : ((v5_ordinal1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => (v2_ordinal5(X0) <=> ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => ((r2_hidden(X2,k9_xtuple_0(X0)) & r2_hidden(X1,X2)) => r1_tarski(k1_funct_1(X0,X1),k1_funct_1(X0,X2)))))))),
  file('minsubpreds65/t36_ordinal5',d2_ordinal5)).
fof(f745214,plain,(
  ~r2_hidden(sK10(sK7),sK11(sK7)) | spl14_31432),
  inference(avatar_component_clause,[],[f745212])).
fof(f745212,plain,(
  spl14_31432 <=> r2_hidden(sK10(sK7),sK11(sK7))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_31432])])).
fof(f2346351,plain,(
  ~spl14_193 | spl14_14 | spl14_137),
  inference(avatar_split_clause,[],[f2346350,f2113,f772,f2892])).
fof(f2892,plain,(
  spl14_193 <=> v7_ordinal1(sK5(sK13))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_193])])).
fof(f772,plain,(
  spl14_14 <=> v1_xboole_0(k4_ordinal1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_14])])).
fof(f2113,plain,(
  spl14_137 <=> r2_hidden(k1_ordinal1(sK5(sK13)),k4_ordinal1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_137])])).
fof(f2346350,plain,(
  ~v7_ordinal1(sK5(sK13)) | (spl14_14 | spl14_137)),
  inference(subsumption_resolution,[],[f2346339,f1563])).
fof(f1563,plain,(
  ( ! [X0] : (r2_hidden(k4_card_1(X0),k4_ordinal1) | ~v7_ordinal1(X0)) ) | spl14_14),
  inference(subsumption_resolution,[],[f1555,f774])).
fof(f774,plain,(
  ~v1_xboole_0(k4_ordinal1) | spl14_14),
  inference(avatar_component_clause,[],[f772])).
fof(f1555,plain,(
  ( ! [X0] : (v1_xboole_0(k4_ordinal1) | r2_hidden(k4_card_1(X0),k4_ordinal1) | ~v7_ordinal1(X0)) )),
  inference(resolution,[],[f674,f246])).
fof(f246,plain,(
  ( ! [X0] : (m1_subset_1(k4_card_1(X0),k4_ordinal1) | ~v7_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0] : (m1_subset_1(k4_card_1(X0),k4_ordinal1) | ~v7_ordinal1(X0))),
  inference(ennf_transformation,[],[f67])).
fof(f67,axiom,(
  ! [X0] : (v7_ordinal1(X0) => m1_subset_1(k4_card_1(X0),k4_ordinal1))),
  file('minsubpreds65/t36_ordinal5',dt_k4_card_1)).
fof(f674,plain,(
  ( ! [X0,X1] : (~m1_subset_1(X0,X1) | v1_xboole_0(X1) | r2_hidden(X0,X1)) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1] : (r2_hidden(X0,X1) | v1_xboole_0(X1) | ~m1_subset_1(X0,X1))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1] : ((r2_hidden(X0,X1) | v1_xboole_0(X1)) | ~m1_subset_1(X0,X1))),
  inference(ennf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1)))),
  inference(flattening,[],[f56])).
fof(f56,axiom,(
  ! [X0] : ! [X1] : (m1_subset_1(X0,X1) => (r2_hidden(X0,X1) | v1_xboole_0(X1)))),
  file('minsubpreds65/t36_ordinal5',t2_subset)).
fof(f2346339,plain,(
  ~r2_hidden(k4_card_1(sK5(sK13)),k4_ordinal1) | ~v7_ordinal1(sK5(sK13)) | spl14_137),
  inference(superposition,[],[f2115,f247])).
fof(f247,plain,(
  ( ! [X0] : (k4_card_1(X0) = k1_ordinal1(X0) | ~v7_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0] : (k4_card_1(X0) = k1_ordinal1(X0) | ~v7_ordinal1(X0))),
  inference(ennf_transformation,[],[f35])).
fof(f35,axiom,(
  ! [X0] : (v7_ordinal1(X0) => k4_card_1(X0) = k1_ordinal1(X0))),
  file('minsubpreds65/t36_ordinal5',redefinition_k4_card_1)).
fof(f2115,plain,(
  ~r2_hidden(k1_ordinal1(sK5(sK13)),k4_ordinal1) | spl14_137),
  inference(avatar_component_clause,[],[f2113])).
fof(f2345839,plain,(
  spl14_75 | ~spl14_22),
  inference(avatar_split_clause,[],[f2345789,f809,f1454])).
fof(f1454,plain,(
  spl14_75 <=> m1_subset_1(sK8,sK9)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_75])])).
fof(f809,plain,(
  spl14_22 <=> r2_hidden(sK8,sK9)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_22])])).
fof(f2345789,plain,(
  m1_subset_1(sK8,sK9) | ~spl14_22),
  inference(resolution,[],[f811,f675])).
fof(f675,plain,(
  ( ! [X0,X1] : (~r2_hidden(X0,X1) | m1_subset_1(X0,X1)) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1] : (m1_subset_1(X0,X1) | ~r2_hidden(X0,X1))),
  inference(ennf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1))),
  inference(flattening,[],[f51])).
fof(f51,axiom,(
  ! [X0] : ! [X1] : (r2_hidden(X0,X1) => m1_subset_1(X0,X1))),
  file('minsubpreds65/t36_ordinal5',t1_subset)).
fof(f811,plain,(
  r2_hidden(sK8,sK9) | ~spl14_22),
  inference(avatar_component_clause,[],[f809])).
fof(f2149625,plain,(
  ~spl14_31432 | ~spl14_13 | ~spl14_25 | ~spl14_28 | ~spl14_29 | ~spl14_30 | ~spl14_86 | spl14_557 | ~spl14_558 | ~spl14_1540 | ~spl14_1546),
  inference(avatar_split_clause,[],[f2149624,f22785,f22685,f8283,f8279,f1572,f863,f855,f847,f824,f767,f745212])).
fof(f767,plain,(
  spl14_13 <=> v3_ordinal1(k4_ordinal1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_13])])).
fof(f824,plain,(
  spl14_25 <=> ! [X5] : (k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(sK7,X5) | ~v3_ordinal1(X5) | ~r2_hidden(X5,sK6))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_25])])).
fof(f1572,plain,(
  spl14_86 <=> r2_hidden(np__1,k4_ordinal1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_86])])).
fof(f8283,plain,(
  spl14_558 <=> r2_hidden(sK11(sK7),sK6)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_558])])).
fof(f22685,plain,(
  spl14_1540 <=> r2_hidden(sK10(sK7),sK6)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_1540])])).
fof(f22785,plain,(
  spl14_1546 <=> v3_ordinal1(sK11(sK7))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_1546])])).
fof(f2149624,plain,(
  ~r2_hidden(sK10(sK7),sK11(sK7)) | (~spl14_13 | ~spl14_25 | ~spl14_28 | ~spl14_29 | ~spl14_30 | ~spl14_86 | spl14_557 | ~spl14_558 | ~spl14_1540 | ~spl14_1546)),
  inference(subsumption_resolution,[],[f2149623,f865])).
fof(f2149623,plain,(
  ~r2_hidden(sK10(sK7),sK11(sK7)) | ~v1_relat_1(sK7) | (~spl14_13 | ~spl14_25 | ~spl14_28 | ~spl14_29 | ~spl14_86 | spl14_557 | ~spl14_558 | ~spl14_1540 | ~spl14_1546)),
  inference(subsumption_resolution,[],[f2149622,f857])).
fof(f2149622,plain,(
  ~r2_hidden(sK10(sK7),sK11(sK7)) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | (~spl14_13 | ~spl14_25 | ~spl14_28 | ~spl14_86 | spl14_557 | ~spl14_558 | ~spl14_1540 | ~spl14_1546)),
  inference(subsumption_resolution,[],[f2149621,f849])).
fof(f2149621,plain,(
  ~r2_hidden(sK10(sK7),sK11(sK7)) | ~v5_ordinal1(sK7) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | (~spl14_13 | ~spl14_25 | ~spl14_86 | spl14_557 | ~spl14_558 | ~spl14_1540 | ~spl14_1546)),
  inference(subsumption_resolution,[],[f2149620,f8280])).
fof(f2149620,plain,(
  ~r2_hidden(sK10(sK7),sK11(sK7)) | v2_ordinal5(sK7) | ~v5_ordinal1(sK7) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | (~spl14_13 | ~spl14_25 | ~spl14_86 | ~spl14_558 | ~spl14_1540 | ~spl14_1546)),
  inference(subsumption_resolution,[],[f2149619,f8285])).
fof(f8285,plain,(
  r2_hidden(sK11(sK7),sK6) | ~spl14_558),
  inference(avatar_component_clause,[],[f8283])).
fof(f2149619,plain,(
  ~r2_hidden(sK10(sK7),sK11(sK7)) | ~r2_hidden(sK11(sK7),sK6) | v2_ordinal5(sK7) | ~v5_ordinal1(sK7) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | (~spl14_13 | ~spl14_25 | ~spl14_86 | ~spl14_1540 | ~spl14_1546)),
  inference(subsumption_resolution,[],[f2149618,f22687])).
fof(f22687,plain,(
  r2_hidden(sK10(sK7),sK6) | ~spl14_1540),
  inference(avatar_component_clause,[],[f22685])).
fof(f2149618,plain,(
  ~r2_hidden(sK10(sK7),sK11(sK7)) | ~r2_hidden(sK10(sK7),sK6) | ~r2_hidden(sK11(sK7),sK6) | v2_ordinal5(sK7) | ~v5_ordinal1(sK7) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | (~spl14_13 | ~spl14_25 | ~spl14_86 | ~spl14_1546)),
  inference(subsumption_resolution,[],[f2149566,f22786])).
fof(f22786,plain,(
  v3_ordinal1(sK11(sK7)) | ~spl14_1546),
  inference(avatar_component_clause,[],[f22785])).
fof(f2149566,plain,(
  ~v3_ordinal1(sK11(sK7)) | ~r2_hidden(sK10(sK7),sK11(sK7)) | ~r2_hidden(sK10(sK7),sK6) | ~r2_hidden(sK11(sK7),sK6) | v2_ordinal5(sK7) | ~v5_ordinal1(sK7) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | (~spl14_13 | ~spl14_25 | ~spl14_86)),
  inference(resolution,[],[f2020670,f661])).
fof(f661,plain,(
  ( ! [X0] : (~r1_tarski(k1_funct_1(X0,sK10(X0)),k1_funct_1(X0,sK11(X0))) | v2_ordinal5(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f212])).
fof(f2020670,plain,(
  ( ! [X0,X1] : (r1_tarski(k1_funct_1(sK7,X1),k1_funct_1(sK7,X0)) | ~v3_ordinal1(X0) | ~r2_hidden(X1,X0) | ~r2_hidden(X1,sK6) | ~r2_hidden(X0,sK6)) ) | (~spl14_13 | ~spl14_25 | ~spl14_86)),
  inference(duplicate_literal_removal,[],[f2020629])).
fof(f2020629,plain,(
  ( ! [X0,X1] : (r1_tarski(k1_funct_1(sK7,X1),k1_funct_1(sK7,X0)) | ~v3_ordinal1(X0) | ~r2_hidden(X1,X0) | ~r2_hidden(X1,sK6) | ~v3_ordinal1(X0) | ~r2_hidden(X0,sK6)) ) | (~spl14_13 | ~spl14_25 | ~spl14_86)),
  inference(superposition,[],[f1999323,f825])).
fof(f825,plain,(
  ( ! [X5] : (k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(sK7,X5) | ~v3_ordinal1(X5) | ~r2_hidden(X5,sK6)) ) | ~spl14_25),
  inference(avatar_component_clause,[],[f824])).
fof(f1999323,plain,(
  ( ! [X722,X721] : (r1_tarski(k1_funct_1(sK7,X721),k12_ordinal2(k4_ordinal1,X722)) | ~v3_ordinal1(X722) | ~r2_hidden(X721,X722) | ~r2_hidden(X721,sK6)) ) | (~spl14_13 | ~spl14_25 | ~spl14_86)),
  inference(subsumption_resolution,[],[f1999322,f6680])).
fof(f6680,plain,(
  ( ! [X2,X3] : (~r2_hidden(X2,X3) | ~v3_ordinal1(X3) | v3_ordinal1(X2)) )),
  inference(subsumption_resolution,[],[f6679,f257])).
fof(f257,plain,(
  ( ! [X0] : (v3_ordinal1(k2_ordinal5(X0)) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0] : ((v4_ordinal5(k2_ordinal5(X0)) & v3_ordinal1(k2_ordinal5(X0))) | ~v3_ordinal1(X0))),
  inference(ennf_transformation,[],[f66])).
fof(f66,axiom,(
  ! [X0] : (v3_ordinal1(X0) => (v4_ordinal5(k2_ordinal5(X0)) & v3_ordinal1(k2_ordinal5(X0))))),
  file('minsubpreds65/t36_ordinal5',dt_k2_ordinal5)).
fof(f6679,plain,(
  ( ! [X2,X3] : (~r2_hidden(X2,X3) | ~v3_ordinal1(X3) | v3_ordinal1(X2) | ~v3_ordinal1(k2_ordinal5(X3))) )),
  inference(resolution,[],[f3246,f276])).
fof(f276,plain,(
  ( ! [X0,X1] : (~m1_subset_1(X1,X0) | v3_ordinal1(X1) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0] : (! [X1] : (v3_ordinal1(X1) | ~m1_subset_1(X1,X0)) | ~v3_ordinal1(X0))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0] : (v3_ordinal1(X0) => ! [X1] : (m1_subset_1(X1,X0) => v3_ordinal1(X1)))),
  file('minsubpreds65/t36_ordinal5',cc2_ordinal2)).
fof(f3246,plain,(
  ( ! [X8,X7] : (m1_subset_1(X8,k2_ordinal5(X7)) | ~r2_hidden(X8,X7) | ~v3_ordinal1(X7)) )),
  inference(resolution,[],[f3242,f675])).
fof(f3242,plain,(
  ( ! [X0,X1] : (r2_hidden(X1,k2_ordinal5(X0)) | ~v3_ordinal1(X0) | ~r2_hidden(X1,X0)) )),
  inference(subsumption_resolution,[],[f1690,f1351])).
fof(f1351,plain,(
  ( ! [X0] : (v1_ordinal1(k2_ordinal5(X0)) | ~v3_ordinal1(X0)) )),
  inference(resolution,[],[f256,f257])).
fof(f256,plain,(
  ( ! [X0] : (~v3_ordinal1(X0) | v1_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0] : (v1_ordinal1(X0) | ~v3_ordinal1(X0))),
  inference(ennf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : (v3_ordinal1(X0) => v1_ordinal1(X0))),
  inference(pure_predicate_removal,[],[f5])).
fof(f5,axiom,(
  ! [X0] : (v3_ordinal1(X0) => (v2_ordinal1(X0) & v1_ordinal1(X0)))),
  file('minsubpreds65/t36_ordinal5',cc1_ordinal1)).
fof(f1690,plain,(
  ( ! [X0,X1] : (~v3_ordinal1(X0) | r2_hidden(X1,k2_ordinal5(X0)) | ~r2_hidden(X1,X0) | ~v1_ordinal1(k2_ordinal5(X0))) )),
  inference(resolution,[],[f1689,f686])).
fof(f686,plain,(
  ( ! [X2,X0,X1] : (~r2_hidden(X1,X2) | r2_hidden(X0,X2) | ~r2_hidden(X0,X1) | ~v1_ordinal1(X2)) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2] : (r2_hidden(X0,X2) | ~r2_hidden(X1,X2) | ~r2_hidden(X0,X1) | ~v1_ordinal1(X2))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2] : ((r2_hidden(X0,X2) | (~r2_hidden(X1,X2) | ~r2_hidden(X0,X1))) | ~v1_ordinal1(X2))),
  inference(ennf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2] : (v1_ordinal1(X2) => ((r2_hidden(X1,X2) & r2_hidden(X0,X1)) => r2_hidden(X0,X2)))),
  inference(flattening,[],[f46])).
fof(f46,axiom,(
  ! [X0] : ! [X1] : ! [X2] : (v1_ordinal1(X2) => ((r2_hidden(X1,X2) & r2_hidden(X0,X1)) => r2_hidden(X0,X2)))),
  file('minsubpreds65/t36_ordinal5',t10_ordinal1)).
fof(f1689,plain,(
  ( ! [X0] : (r2_hidden(X0,k2_ordinal5(X0)) | ~v3_ordinal1(X0)) )),
  inference(subsumption_resolution,[],[f1688,f257])).
fof(f1688,plain,(
  ( ! [X0] : (r2_hidden(X0,k2_ordinal5(X0)) | ~v3_ordinal1(k2_ordinal5(X0)) | ~v3_ordinal1(X0)) )),
  inference(subsumption_resolution,[],[f704,f258])).
fof(f258,plain,(
  ( ! [X0] : (v4_ordinal5(k2_ordinal5(X0)) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f111])).
fof(f704,plain,(
  ( ! [X0] : (r2_hidden(X0,k2_ordinal5(X0)) | ~v4_ordinal5(k2_ordinal5(X0)) | ~v3_ordinal1(k2_ordinal5(X0)) | ~v3_ordinal1(X0)) )),
  inference(equality_resolution,[],[f277])).
fof(f277,plain,(
  ( ! [X0,X1] : (r2_hidden(X0,X1) | k2_ordinal5(X0) != X1 | ~v4_ordinal5(X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0] : (! [X1] : (((k2_ordinal5(X0) = X1 | (~r1_ordinal1(X1,sK3(X0,X1)) & r2_hidden(X0,sK3(X0,X1)) & v4_ordinal5(sK3(X0,X1)) & v3_ordinal1(sK3(X0,X1))) | ~r2_hidden(X0,X1)) & ((! [X3] : (r1_ordinal1(X1,X3) | ~r2_hidden(X0,X3) | ~v4_ordinal5(X3) | ~v3_ordinal1(X3)) & r2_hidden(X0,X1)) | k2_ordinal5(X0) != X1)) | ~v4_ordinal5(X1) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f193,f194])).
fof(f194,plain,(
  ! [X1,X0] : (? [X2] : (~r1_ordinal1(X1,X2) & r2_hidden(X0,X2) & v4_ordinal5(X2) & v3_ordinal1(X2)) => (~r1_ordinal1(X1,sK3(X0,X1)) & r2_hidden(X0,sK3(X0,X1)) & v4_ordinal5(sK3(X0,X1)) & v3_ordinal1(sK3(X0,X1))))),
  introduced(choice_axiom,[])).
fof(f193,plain,(
  ! [X0] : (! [X1] : (((k2_ordinal5(X0) = X1 | ? [X2] : (~r1_ordinal1(X1,X2) & r2_hidden(X0,X2) & v4_ordinal5(X2) & v3_ordinal1(X2)) | ~r2_hidden(X0,X1)) & ((! [X3] : (r1_ordinal1(X1,X3) | ~r2_hidden(X0,X3) | ~v4_ordinal5(X3) | ~v3_ordinal1(X3)) & r2_hidden(X0,X1)) | k2_ordinal5(X0) != X1)) | ~v4_ordinal5(X1) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0))),
  inference(rectify,[],[f192])).
fof(f192,plain,(
  ! [X0] : (! [X1] : (((k2_ordinal5(X0) = X1 | ? [X2] : (~r1_ordinal1(X1,X2) & r2_hidden(X0,X2) & v4_ordinal5(X2) & v3_ordinal1(X2)) | ~r2_hidden(X0,X1)) & ((! [X2] : (r1_ordinal1(X1,X2) | ~r2_hidden(X0,X2) | ~v4_ordinal5(X2) | ~v3_ordinal1(X2)) & r2_hidden(X0,X1)) | k2_ordinal5(X0) != X1)) | ~v4_ordinal5(X1) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0))),
  inference(flattening,[],[f191])).
fof(f191,plain,(
  ! [X0] : (! [X1] : (((k2_ordinal5(X0) = X1 | (? [X2] : (~r1_ordinal1(X1,X2) & r2_hidden(X0,X2) & v4_ordinal5(X2) & v3_ordinal1(X2)) | ~r2_hidden(X0,X1))) & ((! [X2] : (r1_ordinal1(X1,X2) | ~r2_hidden(X0,X2) | ~v4_ordinal5(X2) | ~v3_ordinal1(X2)) & r2_hidden(X0,X1)) | k2_ordinal5(X0) != X1)) | ~v4_ordinal5(X1) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0))),
  inference(nnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0] : (! [X1] : ((k2_ordinal5(X0) = X1 <=> (! [X2] : (r1_ordinal1(X1,X2) | ~r2_hidden(X0,X2) | ~v4_ordinal5(X2) | ~v3_ordinal1(X2)) & r2_hidden(X0,X1))) | ~v4_ordinal5(X1) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0] : (! [X1] : ((k2_ordinal5(X0) = X1 <=> (! [X2] : ((r1_ordinal1(X1,X2) | ~r2_hidden(X0,X2)) | (~v4_ordinal5(X2) | ~v3_ordinal1(X2))) & r2_hidden(X0,X1))) | (~v4_ordinal5(X1) | ~v3_ordinal1(X1))) | ~v3_ordinal1(X0))),
  inference(ennf_transformation,[],[f20])).
fof(f20,axiom,(
  ! [X0] : (v3_ordinal1(X0) => ! [X1] : ((v4_ordinal5(X1) & v3_ordinal1(X1)) => (k2_ordinal5(X0) = X1 <=> (! [X2] : ((v4_ordinal5(X2) & v3_ordinal1(X2)) => (r2_hidden(X0,X2) => r1_ordinal1(X1,X2))) & r2_hidden(X0,X1)))))),
  file('minsubpreds65/t36_ordinal5',d6_ordinal5)).
fof(f1999322,plain,(
  ( ! [X722,X721] : (r1_tarski(k1_funct_1(sK7,X721),k12_ordinal2(k4_ordinal1,X722)) | ~v3_ordinal1(X722) | ~v3_ordinal1(X721) | ~r2_hidden(X721,X722) | ~r2_hidden(X721,sK6)) ) | (~spl14_13 | ~spl14_25 | ~spl14_86)),
  inference(subsumption_resolution,[],[f1999321,f769])).
fof(f769,plain,(
  v3_ordinal1(k4_ordinal1) | ~spl14_13),
  inference(avatar_component_clause,[],[f767])).
fof(f1999321,plain,(
  ( ! [X722,X721] : (r1_tarski(k1_funct_1(sK7,X721),k12_ordinal2(k4_ordinal1,X722)) | ~v3_ordinal1(X722) | ~v3_ordinal1(X721) | ~v3_ordinal1(k4_ordinal1) | ~r2_hidden(X721,X722) | ~r2_hidden(X721,sK6)) ) | (~spl14_25 | ~spl14_86)),
  inference(subsumption_resolution,[],[f1997552,f1574])).
fof(f1574,plain,(
  r2_hidden(np__1,k4_ordinal1) | ~spl14_86),
  inference(avatar_component_clause,[],[f1572])).
fof(f1997552,plain,(
  ( ! [X722,X721] : (r1_tarski(k1_funct_1(sK7,X721),k12_ordinal2(k4_ordinal1,X722)) | ~r2_hidden(np__1,k4_ordinal1) | ~v3_ordinal1(X722) | ~v3_ordinal1(X721) | ~v3_ordinal1(k4_ordinal1) | ~r2_hidden(X721,X722) | ~r2_hidden(X721,sK6)) ) | ~spl14_25),
  inference(duplicate_literal_removal,[],[f1997231])).
fof(f1997231,plain,(
  ( ! [X722,X721] : (r1_tarski(k1_funct_1(sK7,X721),k12_ordinal2(k4_ordinal1,X722)) | ~r2_hidden(np__1,k4_ordinal1) | ~v3_ordinal1(X722) | ~v3_ordinal1(X721) | ~v3_ordinal1(k4_ordinal1) | ~r2_hidden(X721,X722) | ~v3_ordinal1(X721) | ~r2_hidden(X721,sK6)) ) | ~spl14_25),
  inference(superposition,[],[f4366,f825])).
fof(f4366,plain,(
  ( ! [X6,X4,X5] : (r1_tarski(k12_ordinal2(X6,X4),k12_ordinal2(X6,X5)) | ~r2_hidden(np__1,X6) | ~v3_ordinal1(X5) | ~v3_ordinal1(X4) | ~v3_ordinal1(X6) | ~r2_hidden(X4,X5)) )),
  inference(subsumption_resolution,[],[f1837,f1611])).
fof(f1611,plain,(
  ( ! [X0,X1] : (v1_ordinal1(k12_ordinal2(X1,X0)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) )),
  inference(resolution,[],[f681,f256])).
fof(f681,plain,(
  ( ! [X0,X1] : (v3_ordinal1(k12_ordinal2(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1] : (v3_ordinal1(k12_ordinal2(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0,X1] : (v3_ordinal1(k12_ordinal2(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0)))),
  inference(ennf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => v3_ordinal1(k12_ordinal2(X0,X1)))),
  inference(flattening,[],[f21])).
fof(f21,axiom,(
  ! [X0] : ! [X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => v3_ordinal1(k12_ordinal2(X0,X1)))),
  file('minsubpreds65/t36_ordinal5',dt_k12_ordinal2)).
fof(f1837,plain,(
  ( ! [X6,X4,X5] : (~r2_hidden(X4,X5) | ~r2_hidden(np__1,X6) | ~v3_ordinal1(X5) | ~v3_ordinal1(X4) | ~v3_ordinal1(X6) | r1_tarski(k12_ordinal2(X6,X4),k12_ordinal2(X6,X5)) | ~v1_ordinal1(k12_ordinal2(X6,X5))) )),
  inference(resolution,[],[f274,f669])).
fof(f669,plain,(
  ( ! [X2,X0] : (~r2_hidden(X2,X0) | r1_tarski(X2,X0) | ~v1_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0] : ((v1_ordinal1(X0) | (~r1_tarski(sK12(X0),X0) & r2_hidden(sK12(X0),X0))) & (! [X2] : (r1_tarski(X2,X0) | ~r2_hidden(X2,X0)) | ~v1_ordinal1(X0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f214,f215])).
fof(f215,plain,(
  ! [X0] : (? [X1] : (~r1_tarski(X1,X0) & r2_hidden(X1,X0)) => (~r1_tarski(sK12(X0),X0) & r2_hidden(sK12(X0),X0)))),
  introduced(choice_axiom,[])).
fof(f214,plain,(
  ! [X0] : ((v1_ordinal1(X0) | ? [X1] : (~r1_tarski(X1,X0) & r2_hidden(X1,X0))) & (! [X2] : (r1_tarski(X2,X0) | ~r2_hidden(X2,X0)) | ~v1_ordinal1(X0)))),
  inference(rectify,[],[f213])).
fof(f213,plain,(
  ! [X0] : ((v1_ordinal1(X0) | ? [X1] : (~r1_tarski(X1,X0) & r2_hidden(X1,X0))) & (! [X1] : (r1_tarski(X1,X0) | ~r2_hidden(X1,X0)) | ~v1_ordinal1(X0)))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0] : (v1_ordinal1(X0) <=> ! [X1] : (r1_tarski(X1,X0) | ~r2_hidden(X1,X0)))),
  inference(ennf_transformation,[],[f16])).
fof(f16,axiom,(
  ! [X0] : (v1_ordinal1(X0) <=> ! [X1] : (r2_hidden(X1,X0) => r1_tarski(X1,X0)))),
  file('minsubpreds65/t36_ordinal5',d2_ordinal1)).
fof(f274,plain,(
  ( ! [X2,X0,X1] : (r2_hidden(k12_ordinal2(X0,X1),k12_ordinal2(X0,X2)) | ~r2_hidden(X1,X2) | ~r2_hidden(np__1,X0) | ~v3_ordinal1(X2) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0] : (! [X1] : (! [X2] : (r2_hidden(k12_ordinal2(X0,X1),k12_ordinal2(X0,X2)) | ~r2_hidden(X1,X2) | ~r2_hidden(np__1,X0) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0] : (! [X1] : (! [X2] : ((r2_hidden(k12_ordinal2(X0,X1),k12_ordinal2(X0,X2)) | (~r2_hidden(X1,X2) | ~r2_hidden(np__1,X0))) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0))),
  inference(ennf_transformation,[],[f54])).
fof(f54,axiom,(
  ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => ((r2_hidden(X1,X2) & r2_hidden(np__1,X0)) => r2_hidden(k12_ordinal2(X0,X1),k12_ordinal2(X0,X2))))))),
  file('minsubpreds65/t36_ordinal5',t24_ordinal4)).
fof(f2068141,plain,(
  ~spl14_4802 | ~spl14_13 | ~spl14_54 | ~spl14_153 | ~spl14_429 | spl14_432),
  inference(avatar_split_clause,[],[f2068140,f6666,f6642,f2322,f1304,f767,f99130])).
fof(f99130,plain,(
  spl14_4802 <=> r1_ordinal1(sK10(sK13),sK11(sK13))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_4802])])).
fof(f1304,plain,(
  spl14_54 <=> r2_hidden(k6_numbers,k4_ordinal1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_54])])).
fof(f2322,plain,(
  spl14_153 <=> v3_ordinal1(sK11(sK13))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_153])])).
fof(f6642,plain,(
  spl14_429 <=> v3_ordinal1(sK10(sK13))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_429])])).
fof(f6666,plain,(
  spl14_432 <=> r1_tarski(k1_ordinal5(k4_ordinal1,sK10(sK13)),k1_ordinal5(k4_ordinal1,sK11(sK13)))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_432])])).
fof(f2068140,plain,(
  ~r1_ordinal1(sK10(sK13),sK11(sK13)) | (~spl14_13 | ~spl14_54 | ~spl14_153 | ~spl14_429 | spl14_432)),
  inference(subsumption_resolution,[],[f2068139,f1305])).
fof(f1305,plain,(
  r2_hidden(k6_numbers,k4_ordinal1) | ~spl14_54),
  inference(avatar_component_clause,[],[f1304])).
fof(f2068139,plain,(
  ~r1_ordinal1(sK10(sK13),sK11(sK13)) | ~r2_hidden(k6_numbers,k4_ordinal1) | (~spl14_13 | ~spl14_153 | ~spl14_429 | spl14_432)),
  inference(subsumption_resolution,[],[f2068138,f6643])).
fof(f6643,plain,(
  v3_ordinal1(sK10(sK13)) | ~spl14_429),
  inference(avatar_component_clause,[],[f6642])).
fof(f2068138,plain,(
  ~r1_ordinal1(sK10(sK13),sK11(sK13)) | ~v3_ordinal1(sK10(sK13)) | ~r2_hidden(k6_numbers,k4_ordinal1) | (~spl14_13 | ~spl14_153 | spl14_432)),
  inference(subsumption_resolution,[],[f2068137,f2324])).
fof(f2324,plain,(
  v3_ordinal1(sK11(sK13)) | ~spl14_153),
  inference(avatar_component_clause,[],[f2322])).
fof(f2068137,plain,(
  ~r1_ordinal1(sK10(sK13),sK11(sK13)) | ~v3_ordinal1(sK11(sK13)) | ~v3_ordinal1(sK10(sK13)) | ~r2_hidden(k6_numbers,k4_ordinal1) | (~spl14_13 | spl14_432)),
  inference(subsumption_resolution,[],[f2068130,f769])).
fof(f2068130,plain,(
  ~r1_ordinal1(sK10(sK13),sK11(sK13)) | ~v3_ordinal1(k4_ordinal1) | ~v3_ordinal1(sK11(sK13)) | ~v3_ordinal1(sK10(sK13)) | ~r2_hidden(k6_numbers,k4_ordinal1) | spl14_432),
  inference(resolution,[],[f6668,f1891])).
fof(f1891,plain,(
  ( ! [X2,X0,X1] : (r1_tarski(k1_ordinal5(X0,X1),k1_ordinal5(X0,X2)) | ~r1_ordinal1(X1,X2) | ~v3_ordinal1(X0) | ~v3_ordinal1(X2) | ~v3_ordinal1(X1) | ~r2_hidden(k6_numbers,X0)) )),
  inference(subsumption_resolution,[],[f1890,f680])).
fof(f680,plain,(
  ( ! [X0,X1] : (v3_ordinal1(k1_ordinal5(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1] : (v3_ordinal1(k1_ordinal5(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0))),
  inference(flattening,[],[f168])).
fof(f168,plain,(
  ! [X0,X1] : (v3_ordinal1(k1_ordinal5(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0)))),
  inference(ennf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => v3_ordinal1(k1_ordinal5(X0,X1)))),
  inference(flattening,[],[f22])).
fof(f22,axiom,(
  ! [X0] : ! [X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => v3_ordinal1(k1_ordinal5(X0,X1)))),
  file('minsubpreds65/t36_ordinal5',dt_k1_ordinal5)).
fof(f1890,plain,(
  ( ! [X2,X0,X1] : (~r2_hidden(k6_numbers,X0) | ~r1_ordinal1(X1,X2) | ~v3_ordinal1(X0) | ~v3_ordinal1(X2) | ~v3_ordinal1(X1) | r1_tarski(k1_ordinal5(X0,X1),k1_ordinal5(X0,X2)) | ~v3_ordinal1(k1_ordinal5(X0,X1))) )),
  inference(subsumption_resolution,[],[f1885,f680])).
fof(f1885,plain,(
  ( ! [X2,X0,X1] : (~r2_hidden(k6_numbers,X0) | ~r1_ordinal1(X1,X2) | ~v3_ordinal1(X0) | ~v3_ordinal1(X2) | ~v3_ordinal1(X1) | r1_tarski(k1_ordinal5(X0,X1),k1_ordinal5(X0,X2)) | ~v3_ordinal1(k1_ordinal5(X0,X2)) | ~v3_ordinal1(k1_ordinal5(X0,X1))) )),
  inference(resolution,[],[f275,f682])).
fof(f682,plain,(
  ( ! [X0,X1] : (~r1_ordinal1(X0,X1) | r1_tarski(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1] : (((r1_ordinal1(X0,X1) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~r1_ordinal1(X0,X1))) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0))),
  inference(nnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0))),
  inference(flattening,[],[f172])).
fof(f172,plain,(
  ! [X0,X1] : ((r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)) | (~v3_ordinal1(X1) | ~v3_ordinal1(X0)))),
  inference(ennf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => (r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)))),
  inference(flattening,[],[f38])).
fof(f38,axiom,(
  ! [X0] : ! [X1] : ((v3_ordinal1(X1) & v3_ordinal1(X0)) => (r1_ordinal1(X0,X1) <=> r1_tarski(X0,X1)))),
  file('minsubpreds65/t36_ordinal5',redefinition_r1_ordinal1)).
fof(f275,plain,(
  ( ! [X2,X0,X1] : (r1_ordinal1(k1_ordinal5(X2,X0),k1_ordinal5(X2,X1)) | ~r2_hidden(k6_numbers,X2) | ~r1_ordinal1(X0,X1) | ~v3_ordinal1(X2) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0] : (! [X1] : (! [X2] : (r1_ordinal1(k1_ordinal5(X2,X0),k1_ordinal5(X2,X1)) | ~r2_hidden(k6_numbers,X2) | ~r1_ordinal1(X0,X1) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0] : (! [X1] : (! [X2] : ((r1_ordinal1(k1_ordinal5(X2,X0),k1_ordinal5(X2,X1)) | (~r2_hidden(k6_numbers,X2) | ~r1_ordinal1(X0,X1))) | ~v3_ordinal1(X2)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0))),
  inference(ennf_transformation,[],[f53])).
fof(f53,axiom,(
  ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => ((r2_hidden(k6_numbers,X2) & r1_ordinal1(X0,X1)) => r1_ordinal1(k1_ordinal5(X2,X0),k1_ordinal5(X2,X1))))))),
  file('minsubpreds65/t36_ordinal5',t21_ordinal5)).
fof(f6668,plain,(
  ~r1_tarski(k1_ordinal5(k4_ordinal1,sK10(sK13)),k1_ordinal5(k4_ordinal1,sK11(sK13))) | spl14_432),
  inference(avatar_component_clause,[],[f6666])).
fof(f2030268,plain,(
  spl14_2736 | ~spl14_3 | ~spl14_8 | ~spl14_13 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_132 | spl14_139 | ~spl14_145),
  inference(avatar_split_clause,[],[f2030267,f2237,f2144,f2048,f1268,f1263,f1258,f1253,f1248,f767,f742,f717,f49957])).
fof(f49957,plain,(
  spl14_2736 <=> k1_ordinal5(k4_ordinal1,k4_ordinal1) = k8_ordinal2(sK13)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_2736])])).
fof(f717,plain,(
  spl14_3 <=> v4_ordinal1(k4_ordinal1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_3])])).
fof(f742,plain,(
  spl14_8 <=> k6_numbers = k1_xboole_0),
  introduced(avatar_definition,[new_symbols(naming,[spl14_8])])).
fof(f1248,plain,(
  spl14_44 <=> k4_ordinal1 = k9_xtuple_0(sK13)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_44])])).
fof(f1253,plain,(
  spl14_45 <=> v1_ordinal2(sK13)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_45])])).
fof(f1258,plain,(
  spl14_46 <=> v1_funct_1(sK13)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_46])])).
fof(f1263,plain,(
  spl14_47 <=> v1_relat_1(sK13)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_47])])).
fof(f1268,plain,(
  spl14_48 <=> v5_ordinal1(sK13)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_48])])).
fof(f2048,plain,(
  spl14_132 <=> ! [X0] : (~v7_ordinal1(X0) | k1_funct_1(sK13,X0) = k1_ordinal5(k4_ordinal1,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_132])])).
fof(f2144,plain,(
  spl14_139 <=> k6_numbers = k4_ordinal1),
  introduced(avatar_definition,[new_symbols(naming,[spl14_139])])).
fof(f2237,plain,(
  spl14_145 <=> ! [X0] : (r2_hidden(sK2(k4_ordinal1,X0,sK13),k4_ordinal1) | ~v3_ordinal1(X0) | k8_ordinal2(sK13) = k1_ordinal5(X0,k4_ordinal1))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_145])])).
fof(f2030267,plain,(
  k1_ordinal5(k4_ordinal1,k4_ordinal1) = k8_ordinal2(sK13) | (~spl14_3 | ~spl14_8 | ~spl14_13 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_132 | spl14_139 | ~spl14_145)),
  inference(subsumption_resolution,[],[f2030263,f769])).
fof(f2030263,plain,(
  k1_ordinal5(k4_ordinal1,k4_ordinal1) = k8_ordinal2(sK13) | ~v3_ordinal1(k4_ordinal1) | (~spl14_3 | ~spl14_8 | ~spl14_13 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_132 | spl14_139 | ~spl14_145)),
  inference(equality_resolution,[],[f1981470])).
fof(f1981470,plain,(
  ( ! [X718] : (k1_ordinal5(k4_ordinal1,sK2(k4_ordinal1,X718,sK13)) != k1_ordinal5(X718,sK2(k4_ordinal1,X718,sK13)) | k8_ordinal2(sK13) = k1_ordinal5(X718,k4_ordinal1) | ~v3_ordinal1(X718)) ) | (~spl14_3 | ~spl14_8 | ~spl14_13 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_132 | spl14_139 | ~spl14_145)),
  inference(subsumption_resolution,[],[f1980625,f3635])).
fof(f3635,plain,(
  ( ! [X1] : (k8_ordinal2(sK13) = k1_ordinal5(X1,k4_ordinal1) | ~v3_ordinal1(X1) | v7_ordinal1(sK2(k4_ordinal1,X1,sK13))) ) | ~spl14_145),
  inference(resolution,[],[f2238,f673])).
fof(f673,plain,(
  ( ! [X0] : (~r2_hidden(X0,k4_ordinal1) | v7_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0] : ((v7_ordinal1(X0) | ~r2_hidden(X0,k4_ordinal1)) & (r2_hidden(X0,k4_ordinal1) | ~v7_ordinal1(X0)))),
  inference(nnf_transformation,[],[f15])).
fof(f15,axiom,(
  ! [X0] : (v7_ordinal1(X0) <=> r2_hidden(X0,k4_ordinal1))),
  file('minsubpreds65/t36_ordinal5',d12_ordinal1)).
fof(f2238,plain,(
  ( ! [X0] : (r2_hidden(sK2(k4_ordinal1,X0,sK13),k4_ordinal1) | ~v3_ordinal1(X0) | k8_ordinal2(sK13) = k1_ordinal5(X0,k4_ordinal1)) ) | ~spl14_145),
  inference(avatar_component_clause,[],[f2237])).
fof(f1980625,plain,(
  ( ! [X718] : (k1_ordinal5(k4_ordinal1,sK2(k4_ordinal1,X718,sK13)) != k1_ordinal5(X718,sK2(k4_ordinal1,X718,sK13)) | k8_ordinal2(sK13) = k1_ordinal5(X718,k4_ordinal1) | ~v3_ordinal1(X718) | ~v7_ordinal1(sK2(k4_ordinal1,X718,sK13))) ) | (~spl14_3 | ~spl14_8 | ~spl14_13 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_132 | spl14_139)),
  inference(superposition,[],[f2272,f2049])).
fof(f2049,plain,(
  ( ! [X0] : (k1_funct_1(sK13,X0) = k1_ordinal5(k4_ordinal1,X0) | ~v7_ordinal1(X0)) ) | ~spl14_132),
  inference(avatar_component_clause,[],[f2048])).
fof(f2272,plain,(
  ( ! [X0] : (k1_funct_1(sK13,sK2(k4_ordinal1,X0,sK13)) != k1_ordinal5(X0,sK2(k4_ordinal1,X0,sK13)) | k8_ordinal2(sK13) = k1_ordinal5(X0,k4_ordinal1) | ~v3_ordinal1(X0)) ) | (~spl14_3 | ~spl14_8 | ~spl14_13 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | spl14_139)),
  inference(subsumption_resolution,[],[f2271,f769])).
fof(f2271,plain,(
  ( ! [X0] : (k1_funct_1(sK13,sK2(k4_ordinal1,X0,sK13)) != k1_ordinal5(X0,sK2(k4_ordinal1,X0,sK13)) | k8_ordinal2(sK13) = k1_ordinal5(X0,k4_ordinal1) | ~v3_ordinal1(X0) | ~v3_ordinal1(k4_ordinal1)) ) | (~spl14_3 | ~spl14_8 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | spl14_139)),
  inference(subsumption_resolution,[],[f2270,f719])).
fof(f719,plain,(
  v4_ordinal1(k4_ordinal1) | ~spl14_3),
  inference(avatar_component_clause,[],[f717])).
fof(f2270,plain,(
  ( ! [X0] : (k1_funct_1(sK13,sK2(k4_ordinal1,X0,sK13)) != k1_ordinal5(X0,sK2(k4_ordinal1,X0,sK13)) | k8_ordinal2(sK13) = k1_ordinal5(X0,k4_ordinal1) | ~v4_ordinal1(k4_ordinal1) | ~v3_ordinal1(X0) | ~v3_ordinal1(k4_ordinal1)) ) | (~spl14_8 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | spl14_139)),
  inference(subsumption_resolution,[],[f2269,f1265])).
fof(f1265,plain,(
  v1_relat_1(sK13) | ~spl14_47),
  inference(avatar_component_clause,[],[f1263])).
fof(f2269,plain,(
  ( ! [X0] : (k1_funct_1(sK13,sK2(k4_ordinal1,X0,sK13)) != k1_ordinal5(X0,sK2(k4_ordinal1,X0,sK13)) | k8_ordinal2(sK13) = k1_ordinal5(X0,k4_ordinal1) | ~v1_relat_1(sK13) | ~v4_ordinal1(k4_ordinal1) | ~v3_ordinal1(X0) | ~v3_ordinal1(k4_ordinal1)) ) | (~spl14_8 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_48 | spl14_139)),
  inference(subsumption_resolution,[],[f2268,f1260])).
fof(f1260,plain,(
  v1_funct_1(sK13) | ~spl14_46),
  inference(avatar_component_clause,[],[f1258])).
fof(f2268,plain,(
  ( ! [X0] : (k1_funct_1(sK13,sK2(k4_ordinal1,X0,sK13)) != k1_ordinal5(X0,sK2(k4_ordinal1,X0,sK13)) | k8_ordinal2(sK13) = k1_ordinal5(X0,k4_ordinal1) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | ~v4_ordinal1(k4_ordinal1) | ~v3_ordinal1(X0) | ~v3_ordinal1(k4_ordinal1)) ) | (~spl14_8 | ~spl14_44 | ~spl14_45 | ~spl14_48 | spl14_139)),
  inference(subsumption_resolution,[],[f2267,f1270])).
fof(f1270,plain,(
  v5_ordinal1(sK13) | ~spl14_48),
  inference(avatar_component_clause,[],[f1268])).
fof(f2267,plain,(
  ( ! [X0] : (k1_funct_1(sK13,sK2(k4_ordinal1,X0,sK13)) != k1_ordinal5(X0,sK2(k4_ordinal1,X0,sK13)) | k8_ordinal2(sK13) = k1_ordinal5(X0,k4_ordinal1) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | ~v4_ordinal1(k4_ordinal1) | ~v3_ordinal1(X0) | ~v3_ordinal1(k4_ordinal1)) ) | (~spl14_8 | ~spl14_44 | ~spl14_45 | spl14_139)),
  inference(subsumption_resolution,[],[f2266,f1255])).
fof(f1255,plain,(
  v1_ordinal2(sK13) | ~spl14_45),
  inference(avatar_component_clause,[],[f1253])).
fof(f2266,plain,(
  ( ! [X0] : (k1_funct_1(sK13,sK2(k4_ordinal1,X0,sK13)) != k1_ordinal5(X0,sK2(k4_ordinal1,X0,sK13)) | k8_ordinal2(sK13) = k1_ordinal5(X0,k4_ordinal1) | ~v1_ordinal2(sK13) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | ~v4_ordinal1(k4_ordinal1) | ~v3_ordinal1(X0) | ~v3_ordinal1(k4_ordinal1)) ) | (~spl14_8 | ~spl14_44 | spl14_139)),
  inference(subsumption_resolution,[],[f2265,f2145])).
fof(f2145,plain,(
  k6_numbers != k4_ordinal1 | spl14_139),
  inference(avatar_component_clause,[],[f2144])).
fof(f2265,plain,(
  ( ! [X0] : (k1_funct_1(sK13,sK2(k4_ordinal1,X0,sK13)) != k1_ordinal5(X0,sK2(k4_ordinal1,X0,sK13)) | k8_ordinal2(sK13) = k1_ordinal5(X0,k4_ordinal1) | k6_numbers = k4_ordinal1 | ~v1_ordinal2(sK13) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | ~v4_ordinal1(k4_ordinal1) | ~v3_ordinal1(X0) | ~v3_ordinal1(k4_ordinal1)) ) | (~spl14_8 | ~spl14_44)),
  inference(superposition,[],[f2254,f1250])).
fof(f1250,plain,(
  k4_ordinal1 = k9_xtuple_0(sK13) | ~spl14_44),
  inference(avatar_component_clause,[],[f1248])).
fof(f2254,plain,(
  ( ! [X2,X1] : (k1_funct_1(X2,sK2(k9_xtuple_0(X2),X1,X2)) != k1_ordinal5(X1,sK2(k9_xtuple_0(X2),X1,X2)) | k8_ordinal2(X2) = k1_ordinal5(X1,k9_xtuple_0(X2)) | k6_numbers = k9_xtuple_0(X2) | ~v1_ordinal2(X2) | ~v5_ordinal1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v4_ordinal1(k9_xtuple_0(X2)) | ~v3_ordinal1(X1) | ~v3_ordinal1(k9_xtuple_0(X2))) ) | ~spl14_8),
  inference(forward_demodulation,[],[f700,f744])).
fof(f744,plain,(
  k6_numbers = k1_xboole_0 | ~spl14_8),
  inference(avatar_component_clause,[],[f742])).
fof(f700,plain,(
  ( ! [X2,X1] : (k8_ordinal2(X2) = k1_ordinal5(X1,k9_xtuple_0(X2)) | k1_funct_1(X2,sK2(k9_xtuple_0(X2),X1,X2)) != k1_ordinal5(X1,sK2(k9_xtuple_0(X2),X1,X2)) | ~v1_ordinal2(X2) | ~v5_ordinal1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | k1_xboole_0 = k9_xtuple_0(X2) | ~v4_ordinal1(k9_xtuple_0(X2)) | ~v3_ordinal1(X1) | ~v3_ordinal1(k9_xtuple_0(X2))) )),
  inference(equality_resolution,[],[f271])).
fof(f271,plain,(
  ( ! [X2,X0,X1] : (k1_ordinal5(X1,X0) = k8_ordinal2(X2) | k1_funct_1(X2,sK2(X0,X1,X2)) != k1_ordinal5(X1,sK2(X0,X1,X2)) | k9_xtuple_0(X2) != X0 | ~v1_ordinal2(X2) | ~v5_ordinal1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | k1_xboole_0 = X0 | ~v4_ordinal1(X0) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0] : (! [X1] : (! [X2] : (k1_ordinal5(X1,X0) = k8_ordinal2(X2) | (k1_funct_1(X2,sK2(X0,X1,X2)) != k1_ordinal5(X1,sK2(X0,X1,X2)) & r2_hidden(sK2(X0,X1,X2),X0) & v3_ordinal1(sK2(X0,X1,X2))) | k9_xtuple_0(X2) != X0 | ~v1_ordinal2(X2) | ~v5_ordinal1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) | k1_xboole_0 = X0 | ~v4_ordinal1(X0) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f119,f188])).
fof(f188,plain,(
  ! [X2,X1,X0] : (? [X3] : (k1_funct_1(X2,X3) != k1_ordinal5(X1,X3) & r2_hidden(X3,X0) & v3_ordinal1(X3)) => (k1_funct_1(X2,sK2(X0,X1,X2)) != k1_ordinal5(X1,sK2(X0,X1,X2)) & r2_hidden(sK2(X0,X1,X2),X0) & v3_ordinal1(sK2(X0,X1,X2))))),
  introduced(choice_axiom,[])).
fof(f119,plain,(
  ! [X0] : (! [X1] : (! [X2] : (k1_ordinal5(X1,X0) = k8_ordinal2(X2) | ? [X3] : (k1_funct_1(X2,X3) != k1_ordinal5(X1,X3) & r2_hidden(X3,X0) & v3_ordinal1(X3)) | k9_xtuple_0(X2) != X0 | ~v1_ordinal2(X2) | ~v5_ordinal1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2)) | k1_xboole_0 = X0 | ~v4_ordinal1(X0) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0] : (! [X1] : (((! [X2] : ((k1_ordinal5(X1,X0) = k8_ordinal2(X2) | (? [X3] : ((k1_funct_1(X2,X3) != k1_ordinal5(X1,X3) & r2_hidden(X3,X0)) & v3_ordinal1(X3)) | k9_xtuple_0(X2) != X0)) | (~v1_ordinal2(X2) | ~v5_ordinal1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2))) | k1_xboole_0 = X0) | ~v4_ordinal1(X0)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0))),
  inference(ennf_transformation,[],[f49])).
fof(f49,axiom,(
  ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (v4_ordinal1(X0) => (! [X2] : ((v1_ordinal2(X2) & v5_ordinal1(X2) & v1_funct_1(X2) & v1_relat_1(X2)) => ((! [X3] : (v3_ordinal1(X3) => (r2_hidden(X3,X0) => k1_funct_1(X2,X3) = k1_ordinal5(X1,X3))) & k9_xtuple_0(X2) = X0) => k1_ordinal5(X1,X0) = k8_ordinal2(X2))) | k1_xboole_0 = X0))))),
  file('minsubpreds65/t36_ordinal5',t15_ordinal5)).
fof(f2009572,plain,(
  ~spl14_432 | ~spl14_132 | ~spl14_200 | spl14_202),
  inference(avatar_split_clause,[],[f2009571,f3076,f3062,f2048,f6666])).
fof(f3062,plain,(
  spl14_200 <=> v7_ordinal1(sK10(sK13))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_200])])).
fof(f3076,plain,(
  spl14_202 <=> r1_tarski(k1_funct_1(sK13,sK10(sK13)),k1_ordinal5(k4_ordinal1,sK11(sK13)))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_202])])).
fof(f2009571,plain,(
  ~r1_tarski(k1_ordinal5(k4_ordinal1,sK10(sK13)),k1_ordinal5(k4_ordinal1,sK11(sK13))) | (~spl14_132 | ~spl14_200 | spl14_202)),
  inference(subsumption_resolution,[],[f2009566,f3063])).
fof(f3063,plain,(
  v7_ordinal1(sK10(sK13)) | ~spl14_200),
  inference(avatar_component_clause,[],[f3062])).
fof(f2009566,plain,(
  ~r1_tarski(k1_ordinal5(k4_ordinal1,sK10(sK13)),k1_ordinal5(k4_ordinal1,sK11(sK13))) | ~v7_ordinal1(sK10(sK13)) | (~spl14_132 | spl14_202)),
  inference(superposition,[],[f3078,f2049])).
fof(f3078,plain,(
  ~r1_tarski(k1_funct_1(sK13,sK10(sK13)),k1_ordinal5(k4_ordinal1,sK11(sK13))) | spl14_202),
  inference(avatar_component_clause,[],[f3076])).
fof(f2006133,plain,(
  spl14_688 | spl14_688 | ~spl14_27 | ~spl14_28 | ~spl14_29 | ~spl14_30 | spl14_631),
  inference(avatar_split_clause,[],[f2006132,f9069,f863,f855,f847,f839,f10052,f10052])).
fof(f10052,plain,(
  spl14_688 <=> ! [X1] : (~r1_ordinal2(X1,sK7) | ~v3_ordinal1(X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_688])])).
fof(f839,plain,(
  spl14_27 <=> v1_ordinal2(sK7)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_27])])).
fof(f9069,plain,(
  spl14_631 <=> r1_ordinal2(k8_ordinal2(sK7),sK7)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_631])])).
fof(f2006132,plain,(
  ( ! [X0,X1] : (~r1_ordinal2(X0,sK7) | ~v3_ordinal1(X0) | ~r1_ordinal2(X1,sK7) | ~v3_ordinal1(X1)) ) | (~spl14_27 | ~spl14_28 | ~spl14_29 | ~spl14_30 | spl14_631)),
  inference(subsumption_resolution,[],[f2006131,f849])).
fof(f2006131,plain,(
  ( ! [X0,X1] : (~r1_ordinal2(X0,sK7) | ~v3_ordinal1(X0) | ~r1_ordinal2(X1,sK7) | ~v3_ordinal1(X1) | ~v5_ordinal1(sK7)) ) | (~spl14_27 | ~spl14_29 | ~spl14_30 | spl14_631)),
  inference(subsumption_resolution,[],[f2006130,f865])).
fof(f2006130,plain,(
  ( ! [X0,X1] : (~r1_ordinal2(X0,sK7) | ~v3_ordinal1(X0) | ~r1_ordinal2(X1,sK7) | ~v3_ordinal1(X1) | ~v1_relat_1(sK7) | ~v5_ordinal1(sK7)) ) | (~spl14_27 | ~spl14_29 | spl14_631)),
  inference(subsumption_resolution,[],[f2006129,f857])).
fof(f2006129,plain,(
  ( ! [X0,X1] : (~r1_ordinal2(X0,sK7) | ~v3_ordinal1(X0) | ~r1_ordinal2(X1,sK7) | ~v3_ordinal1(X1) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | ~v5_ordinal1(sK7)) ) | (~spl14_27 | spl14_631)),
  inference(subsumption_resolution,[],[f2006123,f841])).
fof(f841,plain,(
  v1_ordinal2(sK7) | ~spl14_27),
  inference(avatar_component_clause,[],[f839])).
fof(f2006123,plain,(
  ( ! [X0,X1] : (~r1_ordinal2(X0,sK7) | ~v3_ordinal1(X0) | ~r1_ordinal2(X1,sK7) | ~v3_ordinal1(X1) | ~v1_ordinal2(sK7) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | ~v5_ordinal1(sK7)) ) | spl14_631),
  inference(duplicate_literal_removal,[],[f2006122])).
fof(f2006122,plain,(
  ( ! [X0,X1] : (~r1_ordinal2(X0,sK7) | ~r1_ordinal2(X0,sK7) | ~v3_ordinal1(X0) | ~r1_ordinal2(X1,sK7) | ~v3_ordinal1(X1) | ~v1_ordinal2(sK7) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | ~v5_ordinal1(sK7)) ) | spl14_631),
  inference(superposition,[],[f9071,f288])).
fof(f288,plain,(
  ( ! [X2,X0,X1] : (k8_ordinal2(X0) = X1 | ~r1_ordinal2(X1,X0) | ~v3_ordinal1(X1) | ~r1_ordinal2(X2,X0) | ~v3_ordinal1(X2) | ~v1_ordinal2(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v5_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0] : (! [X1] : (((k8_ordinal2(X0) = X1 | ~r1_ordinal2(X1,X0)) & (r1_ordinal2(X1,X0) | k8_ordinal2(X0) != X1)) | ~v3_ordinal1(X1)) | ! [X2] : (~r1_ordinal2(X2,X0) | ~v3_ordinal1(X2)) | ~v1_ordinal2(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v5_ordinal1(X0))),
  inference(rectify,[],[f196])).
fof(f196,plain,(
  ! [X0] : (! [X2] : (((k8_ordinal2(X0) = X2 | ~r1_ordinal2(X2,X0)) & (r1_ordinal2(X2,X0) | k8_ordinal2(X0) != X2)) | ~v3_ordinal1(X2)) | ! [X1] : (~r1_ordinal2(X1,X0) | ~v3_ordinal1(X1)) | ~v1_ordinal2(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v5_ordinal1(X0))),
  inference(nnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0] : (! [X2] : ((k8_ordinal2(X0) = X2 <=> r1_ordinal2(X2,X0)) | ~v3_ordinal1(X2)) | ! [X1] : (~r1_ordinal2(X1,X0) | ~v3_ordinal1(X1)) | ~v1_ordinal2(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v5_ordinal1(X0))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0] : ((! [X2] : ((k8_ordinal2(X0) = X2 <=> r1_ordinal2(X2,X0)) | ~v3_ordinal1(X2)) | ! [X1] : (~r1_ordinal2(X1,X0) | ~v3_ordinal1(X1))) | (~v1_ordinal2(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v5_ordinal1(X0)))),
  inference(ennf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : ((v1_ordinal2(X0) & v1_funct_1(X0) & v1_relat_1(X0) & v5_ordinal1(X0)) => (? [X1] : (r1_ordinal2(X1,X0) & v3_ordinal1(X1)) => ! [X2] : (v3_ordinal1(X2) => (k8_ordinal2(X0) = X2 <=> r1_ordinal2(X2,X0)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X0] : ((v1_ordinal2(X0) & v1_funct_1(X0) & v1_relat_1(X0) & v5_ordinal1(X0)) => (? [X1] : (r1_ordinal2(X1,X0) & v3_ordinal1(X1)) => ! [X1] : (v3_ordinal1(X1) => (k8_ordinal2(X0) = X1 <=> r1_ordinal2(X1,X0)))))),
  file('minsubpreds65/t36_ordinal5',d10_ordinal2)).
fof(f9071,plain,(
  ~r1_ordinal2(k8_ordinal2(sK7),sK7) | spl14_631),
  inference(avatar_component_clause,[],[f9069])).
fof(f2006066,plain,(
  spl14_195 | ~spl14_62 | ~spl14_192),
  inference(avatar_split_clause,[],[f2006058,f2887,f1356,f2904])).
fof(f2904,plain,(
  spl14_195 <=> m1_subset_1(sK5(sK13),k4_ordinal1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_195])])).
fof(f1356,plain,(
  spl14_62 <=> v1_ordinal1(k4_ordinal1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_62])])).
fof(f2887,plain,(
  spl14_192 <=> v7_ordinal1(k1_ordinal1(sK5(sK13)))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_192])])).
fof(f2006058,plain,(
  m1_subset_1(sK5(sK13),k4_ordinal1) | (~spl14_62 | ~spl14_192)),
  inference(resolution,[],[f2889,f6497])).
fof(f6497,plain,(
  ( ! [X8] : (~v7_ordinal1(k1_ordinal1(X8)) | m1_subset_1(X8,k4_ordinal1)) ) | ~spl14_62),
  inference(resolution,[],[f3643,f240])).
fof(f240,plain,(
  ( ! [X0] : (r2_hidden(X0,k1_ordinal1(X0))) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,axiom,(
  ! [X0] : r2_hidden(X0,k1_ordinal1(X0))),
  file('minsubpreds65/t36_ordinal5',t6_ordinal1)).
fof(f3643,plain,(
  ( ! [X0,X1] : (~r2_hidden(X0,X1) | m1_subset_1(X0,k4_ordinal1) | ~v7_ordinal1(X1)) ) | ~spl14_62),
  inference(resolution,[],[f1673,f1535])).
fof(f1535,plain,(
  ( ! [X1] : (r1_tarski(X1,k4_ordinal1) | ~v7_ordinal1(X1)) ) | ~spl14_62),
  inference(subsumption_resolution,[],[f1530,f1358])).
fof(f1358,plain,(
  v1_ordinal1(k4_ordinal1) | ~spl14_62),
  inference(avatar_component_clause,[],[f1356])).
fof(f1530,plain,(
  ( ! [X1] : (r1_tarski(X1,k4_ordinal1) | ~v1_ordinal1(k4_ordinal1) | ~v7_ordinal1(X1)) )),
  inference(resolution,[],[f669,f672])).
fof(f672,plain,(
  ( ! [X0] : (r2_hidden(X0,k4_ordinal1) | ~v7_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f217])).
fof(f1673,plain,(
  ( ! [X2,X0,X1] : (~r1_tarski(X2,X1) | ~r2_hidden(X0,X2) | m1_subset_1(X0,X1)) )),
  inference(resolution,[],[f687,f685])).
fof(f685,plain,(
  ( ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1] : ((m1_subset_1(X0,k1_zfmisc_1(X1)) | ~r1_tarski(X0,X1)) & (r1_tarski(X0,X1) | ~m1_subset_1(X0,k1_zfmisc_1(X1))))),
  inference(nnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1))),
  inference(flattening,[],[f58])).
fof(f58,axiom,(
  ! [X0] : ! [X1] : (m1_subset_1(X0,k1_zfmisc_1(X1)) <=> r1_tarski(X0,X1))),
  file('minsubpreds65/t36_ordinal5',t3_subset)).
fof(f687,plain,(
  ( ! [X2,X0,X1] : (~m1_subset_1(X1,k1_zfmisc_1(X2)) | m1_subset_1(X0,X2) | ~r2_hidden(X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2] : (m1_subset_1(X0,X2) | ~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2] : (m1_subset_1(X0,X2) | (~m1_subset_1(X1,k1_zfmisc_1(X2)) | ~r2_hidden(X0,X1)))),
  inference(ennf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2))),
  inference(flattening,[],[f60])).
fof(f60,axiom,(
  ! [X0] : ! [X1] : ! [X2] : ((m1_subset_1(X1,k1_zfmisc_1(X2)) & r2_hidden(X0,X1)) => m1_subset_1(X0,X2))),
  file('minsubpreds65/t36_ordinal5',t4_subset)).
fof(f2889,plain,(
  v7_ordinal1(k1_ordinal1(sK5(sK13))) | ~spl14_192),
  inference(avatar_component_clause,[],[f2887])).
fof(f2000569,plain,(
  ~spl14_150 | spl14_172 | ~spl14_134),
  inference(avatar_split_clause,[],[f2000559,f2064,f2618,f2301])).
fof(f2301,plain,(
  spl14_150 <=> v3_ordinal1(k3_card_3(sK13))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_150])])).
fof(f2618,plain,(
  spl14_172 <=> v4_ordinal5(k3_card_3(sK13))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_172])])).
fof(f2064,plain,(
  spl14_134 <=> k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_134])])).
fof(f2000559,plain,(
  v4_ordinal5(k3_card_3(sK13)) | ~v3_ordinal1(k3_card_3(sK13)) | ~spl14_134),
  inference(trivial_inequality_removal,[],[f2000102])).
fof(f2000102,plain,(
  k3_card_3(sK13) != k3_card_3(sK13) | v4_ordinal5(k3_card_3(sK13)) | ~v3_ordinal1(k3_card_3(sK13)) | ~spl14_134),
  inference(superposition,[],[f260,f2066])).
fof(f2066,plain,(
  k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | ~spl14_134),
  inference(avatar_component_clause,[],[f2064])).
fof(f260,plain,(
  ( ! [X0] : (k12_ordinal2(k4_ordinal1,X0) != X0 | v4_ordinal5(X0) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0] : (((v4_ordinal5(X0) | k12_ordinal2(k4_ordinal1,X0) != X0) & (k12_ordinal2(k4_ordinal1,X0) = X0 | ~v4_ordinal5(X0))) | ~v3_ordinal1(X0))),
  inference(nnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0] : ((v4_ordinal5(X0) <=> k12_ordinal2(k4_ordinal1,X0) = X0) | ~v3_ordinal1(X0))),
  inference(ennf_transformation,[],[f19])).
fof(f19,axiom,(
  ! [X0] : (v3_ordinal1(X0) => (v4_ordinal5(X0) <=> k12_ordinal2(k4_ordinal1,X0) = X0))),
  file('minsubpreds65/t36_ordinal5',d5_ordinal5)).
fof(f2000076,plain,(
  spl14_192 | ~spl14_137),
  inference(avatar_split_clause,[],[f2000011,f2113,f2887])).
fof(f2000011,plain,(
  v7_ordinal1(k1_ordinal1(sK5(sK13))) | ~spl14_137),
  inference(resolution,[],[f2114,f673])).
fof(f2114,plain,(
  r2_hidden(k1_ordinal1(sK5(sK13)),k4_ordinal1) | ~spl14_137),
  inference(avatar_component_clause,[],[f2113])).
fof(f1999884,plain,(
  spl14_195 | ~spl14_128),
  inference(avatar_split_clause,[],[f1999830,f1998,f2904])).
fof(f1998,plain,(
  spl14_128 <=> r2_hidden(sK5(sK13),k4_ordinal1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_128])])).
fof(f1999830,plain,(
  m1_subset_1(sK5(sK13),k4_ordinal1) | ~spl14_128),
  inference(resolution,[],[f2000,f675])).
fof(f2000,plain,(
  r2_hidden(sK5(sK13),k4_ordinal1) | ~spl14_128),
  inference(avatar_component_clause,[],[f1998])).
fof(f1999881,plain,(
  spl14_193 | ~spl14_128),
  inference(avatar_split_clause,[],[f1999828,f1998,f2892])).
fof(f1999828,plain,(
  v7_ordinal1(sK5(sK13)) | ~spl14_128),
  inference(resolution,[],[f2000,f673])).
fof(f1995104,plain,(
  spl14_562 | spl14_577 | ~spl14_8 | ~spl14_26 | ~spl14_27 | ~spl14_28 | ~spl14_29 | ~spl14_30 | ~spl14_32 | ~spl14_33),
  inference(avatar_split_clause,[],[f1995103,f887,f879,f863,f855,f847,f839,f831,f742,f8456,f8316])).
fof(f8316,plain,(
  spl14_562 <=> k6_numbers = sK6),
  introduced(avatar_definition,[new_symbols(naming,[spl14_562])])).
fof(f8456,plain,(
  spl14_577 <=> ! [X13] : (k1_funct_1(sK7,sK1(sK6,X13,sK7)) != k12_ordinal2(X13,sK1(sK6,X13,sK7)) | ~v3_ordinal1(X13) | k8_ordinal2(sK7) = k12_ordinal2(X13,sK6))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_577])])).
fof(f831,plain,(
  spl14_26 <=> sK6 = k9_xtuple_0(sK7)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_26])])).
fof(f879,plain,(
  spl14_32 <=> v4_ordinal1(sK6)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_32])])).
fof(f887,plain,(
  spl14_33 <=> v3_ordinal1(sK6)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_33])])).
fof(f1995103,plain,(
  ( ! [X13] : (k1_funct_1(sK7,sK1(sK6,X13,sK7)) != k12_ordinal2(X13,sK1(sK6,X13,sK7)) | k8_ordinal2(sK7) = k12_ordinal2(X13,sK6) | k6_numbers = sK6 | ~v3_ordinal1(X13)) ) | (~spl14_8 | ~spl14_26 | ~spl14_27 | ~spl14_28 | ~spl14_29 | ~spl14_30 | ~spl14_32 | ~spl14_33)),
  inference(subsumption_resolution,[],[f1995102,f889])).
fof(f889,plain,(
  v3_ordinal1(sK6) | ~spl14_33),
  inference(avatar_component_clause,[],[f887])).
fof(f1995102,plain,(
  ( ! [X13] : (k1_funct_1(sK7,sK1(sK6,X13,sK7)) != k12_ordinal2(X13,sK1(sK6,X13,sK7)) | k8_ordinal2(sK7) = k12_ordinal2(X13,sK6) | k6_numbers = sK6 | ~v3_ordinal1(X13) | ~v3_ordinal1(sK6)) ) | (~spl14_8 | ~spl14_26 | ~spl14_27 | ~spl14_28 | ~spl14_29 | ~spl14_30 | ~spl14_32)),
  inference(subsumption_resolution,[],[f1995101,f881])).
fof(f881,plain,(
  v4_ordinal1(sK6) | ~spl14_32),
  inference(avatar_component_clause,[],[f879])).
fof(f1995101,plain,(
  ( ! [X13] : (k1_funct_1(sK7,sK1(sK6,X13,sK7)) != k12_ordinal2(X13,sK1(sK6,X13,sK7)) | k8_ordinal2(sK7) = k12_ordinal2(X13,sK6) | k6_numbers = sK6 | ~v4_ordinal1(sK6) | ~v3_ordinal1(X13) | ~v3_ordinal1(sK6)) ) | (~spl14_8 | ~spl14_26 | ~spl14_27 | ~spl14_28 | ~spl14_29 | ~spl14_30)),
  inference(subsumption_resolution,[],[f1995100,f849])).
fof(f1995100,plain,(
  ( ! [X13] : (k1_funct_1(sK7,sK1(sK6,X13,sK7)) != k12_ordinal2(X13,sK1(sK6,X13,sK7)) | k8_ordinal2(sK7) = k12_ordinal2(X13,sK6) | k6_numbers = sK6 | ~v5_ordinal1(sK7) | ~v4_ordinal1(sK6) | ~v3_ordinal1(X13) | ~v3_ordinal1(sK6)) ) | (~spl14_8 | ~spl14_26 | ~spl14_27 | ~spl14_29 | ~spl14_30)),
  inference(subsumption_resolution,[],[f1995099,f865])).
fof(f1995099,plain,(
  ( ! [X13] : (k1_funct_1(sK7,sK1(sK6,X13,sK7)) != k12_ordinal2(X13,sK1(sK6,X13,sK7)) | k8_ordinal2(sK7) = k12_ordinal2(X13,sK6) | k6_numbers = sK6 | ~v1_relat_1(sK7) | ~v5_ordinal1(sK7) | ~v4_ordinal1(sK6) | ~v3_ordinal1(X13) | ~v3_ordinal1(sK6)) ) | (~spl14_8 | ~spl14_26 | ~spl14_27 | ~spl14_29)),
  inference(subsumption_resolution,[],[f1995098,f857])).
fof(f1995098,plain,(
  ( ! [X13] : (k1_funct_1(sK7,sK1(sK6,X13,sK7)) != k12_ordinal2(X13,sK1(sK6,X13,sK7)) | k8_ordinal2(sK7) = k12_ordinal2(X13,sK6) | k6_numbers = sK6 | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | ~v5_ordinal1(sK7) | ~v4_ordinal1(sK6) | ~v3_ordinal1(X13) | ~v3_ordinal1(sK6)) ) | (~spl14_8 | ~spl14_26 | ~spl14_27)),
  inference(subsumption_resolution,[],[f1994797,f841])).
fof(f1994797,plain,(
  ( ! [X13] : (k1_funct_1(sK7,sK1(sK6,X13,sK7)) != k12_ordinal2(X13,sK1(sK6,X13,sK7)) | k8_ordinal2(sK7) = k12_ordinal2(X13,sK6) | k6_numbers = sK6 | ~v1_ordinal2(sK7) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | ~v5_ordinal1(sK7) | ~v4_ordinal1(sK6) | ~v3_ordinal1(X13) | ~v3_ordinal1(sK6)) ) | (~spl14_8 | ~spl14_26)),
  inference(superposition,[],[f2253,f833])).
fof(f833,plain,(
  sK6 = k9_xtuple_0(sK7) | ~spl14_26),
  inference(avatar_component_clause,[],[f831])).
fof(f2253,plain,(
  ( ! [X2,X1] : (k1_funct_1(X2,sK1(k9_xtuple_0(X2),X1,X2)) != k12_ordinal2(X1,sK1(k9_xtuple_0(X2),X1,X2)) | k8_ordinal2(X2) = k12_ordinal2(X1,k9_xtuple_0(X2)) | k6_numbers = k9_xtuple_0(X2) | ~v1_ordinal2(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v5_ordinal1(X2) | ~v4_ordinal1(k9_xtuple_0(X2)) | ~v3_ordinal1(X1) | ~v3_ordinal1(k9_xtuple_0(X2))) ) | ~spl14_8),
  inference(forward_demodulation,[],[f697,f744])).
fof(f697,plain,(
  ( ! [X2,X1] : (k8_ordinal2(X2) = k12_ordinal2(X1,k9_xtuple_0(X2)) | k1_funct_1(X2,sK1(k9_xtuple_0(X2),X1,X2)) != k12_ordinal2(X1,sK1(k9_xtuple_0(X2),X1,X2)) | ~v1_ordinal2(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v5_ordinal1(X2) | k1_xboole_0 = k9_xtuple_0(X2) | ~v4_ordinal1(k9_xtuple_0(X2)) | ~v3_ordinal1(X1) | ~v3_ordinal1(k9_xtuple_0(X2))) )),
  inference(equality_resolution,[],[f268])).
fof(f268,plain,(
  ( ! [X2,X0,X1] : (k8_ordinal2(X2) = k12_ordinal2(X1,X0) | k1_funct_1(X2,sK1(X0,X1,X2)) != k12_ordinal2(X1,sK1(X0,X1,X2)) | k9_xtuple_0(X2) != X0 | ~v1_ordinal2(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v5_ordinal1(X2) | k1_xboole_0 = X0 | ~v4_ordinal1(X0) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0] : (! [X1] : (! [X2] : (k8_ordinal2(X2) = k12_ordinal2(X1,X0) | (k1_funct_1(X2,sK1(X0,X1,X2)) != k12_ordinal2(X1,sK1(X0,X1,X2)) & r2_hidden(sK1(X0,X1,X2),X0) & v3_ordinal1(sK1(X0,X1,X2))) | k9_xtuple_0(X2) != X0 | ~v1_ordinal2(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v5_ordinal1(X2)) | k1_xboole_0 = X0 | ~v4_ordinal1(X0) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1])],[f117,f186])).
fof(f186,plain,(
  ! [X2,X1,X0] : (? [X3] : (k1_funct_1(X2,X3) != k12_ordinal2(X1,X3) & r2_hidden(X3,X0) & v3_ordinal1(X3)) => (k1_funct_1(X2,sK1(X0,X1,X2)) != k12_ordinal2(X1,sK1(X0,X1,X2)) & r2_hidden(sK1(X0,X1,X2),X0) & v3_ordinal1(sK1(X0,X1,X2))))),
  introduced(choice_axiom,[])).
fof(f117,plain,(
  ! [X0] : (! [X1] : (! [X2] : (k8_ordinal2(X2) = k12_ordinal2(X1,X0) | ? [X3] : (k1_funct_1(X2,X3) != k12_ordinal2(X1,X3) & r2_hidden(X3,X0) & v3_ordinal1(X3)) | k9_xtuple_0(X2) != X0 | ~v1_ordinal2(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v5_ordinal1(X2)) | k1_xboole_0 = X0 | ~v4_ordinal1(X0) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : (! [X1] : (((! [X2] : ((k8_ordinal2(X2) = k12_ordinal2(X1,X0) | (? [X3] : ((k1_funct_1(X2,X3) != k12_ordinal2(X1,X3) & r2_hidden(X3,X0)) & v3_ordinal1(X3)) | k9_xtuple_0(X2) != X0)) | (~v1_ordinal2(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v5_ordinal1(X2))) | k1_xboole_0 = X0) | ~v4_ordinal1(X0)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0))),
  inference(ennf_transformation,[],[f59])).
fof(f59,axiom,(
  ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (v4_ordinal1(X0) => (! [X2] : ((v1_ordinal2(X2) & v1_funct_1(X2) & v1_relat_1(X2) & v5_ordinal1(X2)) => ((! [X3] : (v3_ordinal1(X3) => (r2_hidden(X3,X0) => k1_funct_1(X2,X3) = k12_ordinal2(X1,X3))) & k9_xtuple_0(X2) = X0) => k8_ordinal2(X2) = k12_ordinal2(X1,X0))) | k1_xboole_0 = X0))))),
  file('minsubpreds65/t36_ordinal5',t45_ordinal2)).
fof(f1995062,plain,(
  spl14_562 | spl14_569 | ~spl14_8 | ~spl14_26 | ~spl14_27 | ~spl14_28 | ~spl14_29 | ~spl14_30 | ~spl14_32 | ~spl14_33),
  inference(avatar_split_clause,[],[f1995061,f887,f879,f863,f855,f847,f839,f831,f742,f8376,f8316])).
fof(f8376,plain,(
  spl14_569 <=> ! [X3] : (r2_hidden(sK1(sK6,X3,sK7),sK6) | ~v3_ordinal1(X3) | k8_ordinal2(sK7) = k12_ordinal2(X3,sK6))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_569])])).
fof(f1995061,plain,(
  ( ! [X3] : (r2_hidden(sK1(sK6,X3,sK7),sK6) | k8_ordinal2(sK7) = k12_ordinal2(X3,sK6) | k6_numbers = sK6 | ~v3_ordinal1(X3)) ) | (~spl14_8 | ~spl14_26 | ~spl14_27 | ~spl14_28 | ~spl14_29 | ~spl14_30 | ~spl14_32 | ~spl14_33)),
  inference(subsumption_resolution,[],[f1995060,f889])).
fof(f1995060,plain,(
  ( ! [X3] : (r2_hidden(sK1(sK6,X3,sK7),sK6) | k8_ordinal2(sK7) = k12_ordinal2(X3,sK6) | k6_numbers = sK6 | ~v3_ordinal1(X3) | ~v3_ordinal1(sK6)) ) | (~spl14_8 | ~spl14_26 | ~spl14_27 | ~spl14_28 | ~spl14_29 | ~spl14_30 | ~spl14_32)),
  inference(subsumption_resolution,[],[f1995059,f881])).
fof(f1995059,plain,(
  ( ! [X3] : (r2_hidden(sK1(sK6,X3,sK7),sK6) | k8_ordinal2(sK7) = k12_ordinal2(X3,sK6) | k6_numbers = sK6 | ~v4_ordinal1(sK6) | ~v3_ordinal1(X3) | ~v3_ordinal1(sK6)) ) | (~spl14_8 | ~spl14_26 | ~spl14_27 | ~spl14_28 | ~spl14_29 | ~spl14_30)),
  inference(subsumption_resolution,[],[f1995058,f849])).
fof(f1995058,plain,(
  ( ! [X3] : (r2_hidden(sK1(sK6,X3,sK7),sK6) | k8_ordinal2(sK7) = k12_ordinal2(X3,sK6) | k6_numbers = sK6 | ~v5_ordinal1(sK7) | ~v4_ordinal1(sK6) | ~v3_ordinal1(X3) | ~v3_ordinal1(sK6)) ) | (~spl14_8 | ~spl14_26 | ~spl14_27 | ~spl14_29 | ~spl14_30)),
  inference(subsumption_resolution,[],[f1995057,f865])).
fof(f1995057,plain,(
  ( ! [X3] : (r2_hidden(sK1(sK6,X3,sK7),sK6) | k8_ordinal2(sK7) = k12_ordinal2(X3,sK6) | k6_numbers = sK6 | ~v1_relat_1(sK7) | ~v5_ordinal1(sK7) | ~v4_ordinal1(sK6) | ~v3_ordinal1(X3) | ~v3_ordinal1(sK6)) ) | (~spl14_8 | ~spl14_26 | ~spl14_27 | ~spl14_29)),
  inference(subsumption_resolution,[],[f1995056,f857])).
fof(f1995056,plain,(
  ( ! [X3] : (r2_hidden(sK1(sK6,X3,sK7),sK6) | k8_ordinal2(sK7) = k12_ordinal2(X3,sK6) | k6_numbers = sK6 | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | ~v5_ordinal1(sK7) | ~v4_ordinal1(sK6) | ~v3_ordinal1(X3) | ~v3_ordinal1(sK6)) ) | (~spl14_8 | ~spl14_26 | ~spl14_27)),
  inference(subsumption_resolution,[],[f1994789,f841])).
fof(f1994789,plain,(
  ( ! [X3] : (r2_hidden(sK1(sK6,X3,sK7),sK6) | k8_ordinal2(sK7) = k12_ordinal2(X3,sK6) | k6_numbers = sK6 | ~v1_ordinal2(sK7) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | ~v5_ordinal1(sK7) | ~v4_ordinal1(sK6) | ~v3_ordinal1(X3) | ~v3_ordinal1(sK6)) ) | (~spl14_8 | ~spl14_26)),
  inference(superposition,[],[f2182,f833])).
fof(f2182,plain,(
  ( ! [X2,X1] : (r2_hidden(sK1(k9_xtuple_0(X2),X1,X2),k9_xtuple_0(X2)) | k8_ordinal2(X2) = k12_ordinal2(X1,k9_xtuple_0(X2)) | k6_numbers = k9_xtuple_0(X2) | ~v1_ordinal2(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v5_ordinal1(X2) | ~v4_ordinal1(k9_xtuple_0(X2)) | ~v3_ordinal1(X1) | ~v3_ordinal1(k9_xtuple_0(X2))) ) | ~spl14_8),
  inference(forward_demodulation,[],[f698,f744])).
fof(f698,plain,(
  ( ! [X2,X1] : (k8_ordinal2(X2) = k12_ordinal2(X1,k9_xtuple_0(X2)) | r2_hidden(sK1(k9_xtuple_0(X2),X1,X2),k9_xtuple_0(X2)) | ~v1_ordinal2(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v5_ordinal1(X2) | k1_xboole_0 = k9_xtuple_0(X2) | ~v4_ordinal1(k9_xtuple_0(X2)) | ~v3_ordinal1(X1) | ~v3_ordinal1(k9_xtuple_0(X2))) )),
  inference(equality_resolution,[],[f267])).
fof(f267,plain,(
  ( ! [X2,X0,X1] : (k8_ordinal2(X2) = k12_ordinal2(X1,X0) | r2_hidden(sK1(X0,X1,X2),X0) | k9_xtuple_0(X2) != X0 | ~v1_ordinal2(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v5_ordinal1(X2) | k1_xboole_0 = X0 | ~v4_ordinal1(X0) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f187])).
fof(f1995028,plain,(
  spl14_558 | ~spl14_26 | ~spl14_28 | ~spl14_29 | ~spl14_30 | spl14_557),
  inference(avatar_split_clause,[],[f1995027,f8279,f863,f855,f847,f831,f8283])).
fof(f1995027,plain,(
  r2_hidden(sK11(sK7),sK6) | (~spl14_26 | ~spl14_28 | ~spl14_29 | ~spl14_30 | spl14_557)),
  inference(subsumption_resolution,[],[f1995026,f865])).
fof(f1995026,plain,(
  r2_hidden(sK11(sK7),sK6) | ~v1_relat_1(sK7) | (~spl14_26 | ~spl14_28 | ~spl14_29 | spl14_557)),
  inference(subsumption_resolution,[],[f1995025,f857])).
fof(f1995025,plain,(
  r2_hidden(sK11(sK7),sK6) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | (~spl14_26 | ~spl14_28 | spl14_557)),
  inference(subsumption_resolution,[],[f1995024,f849])).
fof(f1995024,plain,(
  r2_hidden(sK11(sK7),sK6) | ~v5_ordinal1(sK7) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | (~spl14_26 | spl14_557)),
  inference(subsumption_resolution,[],[f1994780,f8280])).
fof(f1994780,plain,(
  r2_hidden(sK11(sK7),sK6) | v2_ordinal5(sK7) | ~v5_ordinal1(sK7) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | ~spl14_26),
  inference(superposition,[],[f660,f833])).
fof(f660,plain,(
  ( ! [X0] : (r2_hidden(sK11(X0),k9_xtuple_0(X0)) | v2_ordinal5(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f212])).
fof(f1994774,plain,(
  spl14_551 | ~spl14_33),
  inference(avatar_split_clause,[],[f1994771,f887,f8212])).
fof(f8212,plain,(
  spl14_551 <=> v1_ordinal1(sK6)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_551])])).
fof(f1994771,plain,(
  v1_ordinal1(sK6) | ~spl14_33),
  inference(resolution,[],[f889,f256])).
fof(f1993103,plain,(
  ~spl14_631 | spl14_20 | ~spl14_580),
  inference(avatar_split_clause,[],[f1993102,f8480,f801,f9069])).
fof(f801,plain,(
  spl14_20 <=> r1_ordinal2(k12_ordinal2(k4_ordinal1,sK6),sK7)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_20])])).
fof(f8480,plain,(
  spl14_580 <=> k12_ordinal2(k4_ordinal1,sK6) = k8_ordinal2(sK7)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_580])])).
fof(f1993102,plain,(
  ~r1_ordinal2(k8_ordinal2(sK7),sK7) | (spl14_20 | ~spl14_580)),
  inference(forward_demodulation,[],[f803,f8482])).
fof(f8482,plain,(
  k12_ordinal2(k4_ordinal1,sK6) = k8_ordinal2(sK7) | ~spl14_580),
  inference(avatar_component_clause,[],[f8480])).
fof(f803,plain,(
  ~r1_ordinal2(k12_ordinal2(k4_ordinal1,sK6),sK7) | spl14_20),
  inference(avatar_component_clause,[],[f801])).
fof(f1992184,plain,(
  spl14_180 | ~spl14_150 | spl14_14 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_98 | ~spl14_3675),
  inference(avatar_split_clause,[],[f1992183,f78931,f1718,f1268,f1263,f1258,f1253,f1248,f772,f2301,f2700])).
fof(f2700,plain,(
  spl14_180 <=> r2_hidden(np__0,k3_card_3(sK13))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_180])])).
fof(f1718,plain,(
  spl14_98 <=> v2_ordinal5(sK13)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_98])])).
fof(f78931,plain,(
  spl14_3675 <=> ! [X0] : (r2_hidden(np__0,X0) | ~v3_ordinal1(X0) | ~r1_ordinal2(X0,sK13))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_3675])])).
fof(f1992183,plain,(
  ~v3_ordinal1(k3_card_3(sK13)) | r2_hidden(np__0,k3_card_3(sK13)) | (spl14_14 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_98 | ~spl14_3675)),
  inference(subsumption_resolution,[],[f1992182,f774])).
fof(f1992182,plain,(
  v1_xboole_0(k4_ordinal1) | ~v3_ordinal1(k3_card_3(sK13)) | r2_hidden(np__0,k3_card_3(sK13)) | (~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_98 | ~spl14_3675)),
  inference(forward_demodulation,[],[f1992181,f1250])).
fof(f1992181,plain,(
  ~v3_ordinal1(k3_card_3(sK13)) | r2_hidden(np__0,k3_card_3(sK13)) | v1_xboole_0(k9_xtuple_0(sK13)) | (~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_98 | ~spl14_3675)),
  inference(subsumption_resolution,[],[f1971981,f1720])).
fof(f1720,plain,(
  v2_ordinal5(sK13) | ~spl14_98),
  inference(avatar_component_clause,[],[f1718])).
fof(f1971981,plain,(
  ~v3_ordinal1(k3_card_3(sK13)) | r2_hidden(np__0,k3_card_3(sK13)) | v1_xboole_0(k9_xtuple_0(sK13)) | ~v2_ordinal5(sK13) | (~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_3675)),
  inference(subsumption_resolution,[],[f1971980,f1265])).
fof(f1971980,plain,(
  ~v3_ordinal1(k3_card_3(sK13)) | r2_hidden(np__0,k3_card_3(sK13)) | v1_xboole_0(k9_xtuple_0(sK13)) | ~v2_ordinal5(sK13) | ~v1_relat_1(sK13) | (~spl14_45 | ~spl14_46 | ~spl14_48 | ~spl14_3675)),
  inference(subsumption_resolution,[],[f1971979,f1260])).
fof(f1971979,plain,(
  ~v3_ordinal1(k3_card_3(sK13)) | r2_hidden(np__0,k3_card_3(sK13)) | v1_xboole_0(k9_xtuple_0(sK13)) | ~v2_ordinal5(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | (~spl14_45 | ~spl14_48 | ~spl14_3675)),
  inference(subsumption_resolution,[],[f1971978,f1270])).
fof(f1971978,plain,(
  ~v3_ordinal1(k3_card_3(sK13)) | r2_hidden(np__0,k3_card_3(sK13)) | v1_xboole_0(k9_xtuple_0(sK13)) | ~v2_ordinal5(sK13) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | (~spl14_45 | ~spl14_3675)),
  inference(subsumption_resolution,[],[f1657894,f1255])).
fof(f1657894,plain,(
  ~v3_ordinal1(k3_card_3(sK13)) | r2_hidden(np__0,k3_card_3(sK13)) | v1_xboole_0(k9_xtuple_0(sK13)) | ~v2_ordinal5(sK13) | ~v1_ordinal2(sK13) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | ~spl14_3675),
  inference(resolution,[],[f78932,f291])).
fof(f291,plain,(
  ( ! [X0] : (r1_ordinal2(k3_card_3(X0),X0) | v1_xboole_0(k9_xtuple_0(X0)) | ~v2_ordinal5(X0) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0] : (r1_ordinal2(k3_card_3(X0),X0) | v1_xboole_0(k9_xtuple_0(X0)) | ~v2_ordinal5(X0) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0] : ((r1_ordinal2(k3_card_3(X0),X0) | v1_xboole_0(k9_xtuple_0(X0))) | (~v2_ordinal5(X0) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)))),
  inference(ennf_transformation,[],[f64])).
fof(f64,axiom,(
  ! [X0] : ((v2_ordinal5(X0) & v1_ordinal2(X0) & v5_ordinal1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => (~v1_xboole_0(k9_xtuple_0(X0)) => r1_ordinal2(k3_card_3(X0),X0)))),
  file('minsubpreds65/t36_ordinal5',t6_ordinal5)).
fof(f78932,plain,(
  ( ! [X0] : (~r1_ordinal2(X0,sK13) | ~v3_ordinal1(X0) | r2_hidden(np__0,X0)) ) | ~spl14_3675),
  inference(avatar_component_clause,[],[f78931])).
fof(f1987571,plain,(
  spl14_128 | spl14_142 | ~spl14_34 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_65 | ~spl14_132 | ~spl14_133),
  inference(avatar_split_clause,[],[f1987570,f2051,f2048,f1374,f1268,f1263,f1258,f1253,f1248,f895,f2194,f1998])).
fof(f2194,plain,(
  spl14_142 <=> ! [X0] : (r1_ordinal1(k3_card_3(sK13),X0) | k12_ordinal2(k4_ordinal1,X0) != X0 | ~r1_ordinal1(np__1,X0) | ~v3_ordinal1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_142])])).
fof(f895,plain,(
  spl14_34 <=> ! [X1,X0] : (r1_ordinal1(k3_card_3(X0),X1) | ~v1_relat_1(X0) | ~v1_funct_1(X0) | ~v5_ordinal1(X0) | ~v1_ordinal2(X0) | k4_ordinal1 != k9_xtuple_0(X0) | np__1 != k1_funct_1(X0,k6_numbers) | r2_hidden(sK5(X0),k4_ordinal1) | ~v3_ordinal1(X1) | ~r1_ordinal1(np__1,X1) | k12_ordinal2(k4_ordinal1,X1) != X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_34])])).
fof(f1374,plain,(
  spl14_65 <=> v7_ordinal1(k6_numbers)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_65])])).
fof(f2051,plain,(
  spl14_133 <=> np__1 = k1_ordinal5(k4_ordinal1,k6_numbers)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_133])])).
fof(f1987570,plain,(
  ( ! [X779] : (r1_ordinal1(k3_card_3(sK13),X779) | r2_hidden(sK5(sK13),k4_ordinal1) | ~v3_ordinal1(X779) | ~r1_ordinal1(np__1,X779) | k12_ordinal2(k4_ordinal1,X779) != X779) ) | (~spl14_34 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_65 | ~spl14_132 | ~spl14_133)),
  inference(subsumption_resolution,[],[f1981620,f1376])).
fof(f1376,plain,(
  v7_ordinal1(k6_numbers) | ~spl14_65),
  inference(avatar_component_clause,[],[f1374])).
fof(f1981620,plain,(
  ( ! [X779] : (r1_ordinal1(k3_card_3(sK13),X779) | r2_hidden(sK5(sK13),k4_ordinal1) | ~v3_ordinal1(X779) | ~r1_ordinal1(np__1,X779) | k12_ordinal2(k4_ordinal1,X779) != X779 | ~v7_ordinal1(k6_numbers)) ) | (~spl14_34 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_132 | ~spl14_133)),
  inference(subsumption_resolution,[],[f1981619,f1250])).
fof(f1981619,plain,(
  ( ! [X779] : (k4_ordinal1 != k9_xtuple_0(sK13) | r1_ordinal1(k3_card_3(sK13),X779) | r2_hidden(sK5(sK13),k4_ordinal1) | ~v3_ordinal1(X779) | ~r1_ordinal1(np__1,X779) | k12_ordinal2(k4_ordinal1,X779) != X779 | ~v7_ordinal1(k6_numbers)) ) | (~spl14_34 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_132 | ~spl14_133)),
  inference(subsumption_resolution,[],[f1981618,f1255])).
fof(f1981618,plain,(
  ( ! [X779] : (~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | r1_ordinal1(k3_card_3(sK13),X779) | r2_hidden(sK5(sK13),k4_ordinal1) | ~v3_ordinal1(X779) | ~r1_ordinal1(np__1,X779) | k12_ordinal2(k4_ordinal1,X779) != X779 | ~v7_ordinal1(k6_numbers)) ) | (~spl14_34 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_132 | ~spl14_133)),
  inference(subsumption_resolution,[],[f1981617,f1270])).
fof(f1981617,plain,(
  ( ! [X779] : (~v5_ordinal1(sK13) | ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | r1_ordinal1(k3_card_3(sK13),X779) | r2_hidden(sK5(sK13),k4_ordinal1) | ~v3_ordinal1(X779) | ~r1_ordinal1(np__1,X779) | k12_ordinal2(k4_ordinal1,X779) != X779 | ~v7_ordinal1(k6_numbers)) ) | (~spl14_34 | ~spl14_46 | ~spl14_47 | ~spl14_132 | ~spl14_133)),
  inference(subsumption_resolution,[],[f1981616,f1260])).
fof(f1981616,plain,(
  ( ! [X779] : (~v1_funct_1(sK13) | ~v5_ordinal1(sK13) | ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | r1_ordinal1(k3_card_3(sK13),X779) | r2_hidden(sK5(sK13),k4_ordinal1) | ~v3_ordinal1(X779) | ~r1_ordinal1(np__1,X779) | k12_ordinal2(k4_ordinal1,X779) != X779 | ~v7_ordinal1(k6_numbers)) ) | (~spl14_34 | ~spl14_47 | ~spl14_132 | ~spl14_133)),
  inference(subsumption_resolution,[],[f1981615,f1265])).
fof(f1981615,plain,(
  ( ! [X779] : (~v1_relat_1(sK13) | ~v1_funct_1(sK13) | ~v5_ordinal1(sK13) | ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | r1_ordinal1(k3_card_3(sK13),X779) | r2_hidden(sK5(sK13),k4_ordinal1) | ~v3_ordinal1(X779) | ~r1_ordinal1(np__1,X779) | k12_ordinal2(k4_ordinal1,X779) != X779 | ~v7_ordinal1(k6_numbers)) ) | (~spl14_34 | ~spl14_132 | ~spl14_133)),
  inference(subsumption_resolution,[],[f1980691,f2052])).
fof(f2052,plain,(
  np__1 = k1_ordinal5(k4_ordinal1,k6_numbers) | ~spl14_133),
  inference(avatar_component_clause,[],[f2051])).
fof(f1980691,plain,(
  ( ! [X779] : (np__1 != k1_ordinal5(k4_ordinal1,k6_numbers) | ~v1_relat_1(sK13) | ~v1_funct_1(sK13) | ~v5_ordinal1(sK13) | ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | r1_ordinal1(k3_card_3(sK13),X779) | r2_hidden(sK5(sK13),k4_ordinal1) | ~v3_ordinal1(X779) | ~r1_ordinal1(np__1,X779) | k12_ordinal2(k4_ordinal1,X779) != X779 | ~v7_ordinal1(k6_numbers)) ) | (~spl14_34 | ~spl14_132)),
  inference(superposition,[],[f896,f2049])).
fof(f896,plain,(
  ( ! [X0,X1] : (np__1 != k1_funct_1(X0,k6_numbers) | ~v1_relat_1(X0) | ~v1_funct_1(X0) | ~v5_ordinal1(X0) | ~v1_ordinal2(X0) | k4_ordinal1 != k9_xtuple_0(X0) | r1_ordinal1(k3_card_3(X0),X1) | r2_hidden(sK5(X0),k4_ordinal1) | ~v3_ordinal1(X1) | ~r1_ordinal1(np__1,X1) | k12_ordinal2(k4_ordinal1,X1) != X1) ) | ~spl14_34),
  inference(avatar_component_clause,[],[f895])).
fof(f1987569,plain,(
  spl14_128 | spl14_134 | ~spl14_37 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_65 | ~spl14_132 | ~spl14_133),
  inference(avatar_split_clause,[],[f1987568,f2051,f2048,f1374,f1268,f1263,f1258,f1253,f1248,f1024,f2064,f1998])).
fof(f1024,plain,(
  spl14_37 <=> ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | ~v1_relat_1(X0) | ~v1_funct_1(X0) | ~v5_ordinal1(X0) | ~v1_ordinal2(X0) | k4_ordinal1 != k9_xtuple_0(X0) | np__1 != k1_funct_1(X0,k6_numbers) | r2_hidden(sK5(X0),k4_ordinal1))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_37])])).
fof(f1987568,plain,(
  k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | r2_hidden(sK5(sK13),k4_ordinal1) | (~spl14_37 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_65 | ~spl14_132 | ~spl14_133)),
  inference(subsumption_resolution,[],[f1981629,f1376])).
fof(f1981629,plain,(
  k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | r2_hidden(sK5(sK13),k4_ordinal1) | ~v7_ordinal1(k6_numbers) | (~spl14_37 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_132 | ~spl14_133)),
  inference(subsumption_resolution,[],[f1981628,f1250])).
fof(f1981628,plain,(
  k4_ordinal1 != k9_xtuple_0(sK13) | k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | r2_hidden(sK5(sK13),k4_ordinal1) | ~v7_ordinal1(k6_numbers) | (~spl14_37 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_132 | ~spl14_133)),
  inference(subsumption_resolution,[],[f1981627,f1255])).
fof(f1981627,plain,(
  ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | r2_hidden(sK5(sK13),k4_ordinal1) | ~v7_ordinal1(k6_numbers) | (~spl14_37 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_132 | ~spl14_133)),
  inference(subsumption_resolution,[],[f1981626,f1270])).
fof(f1981626,plain,(
  ~v5_ordinal1(sK13) | ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | r2_hidden(sK5(sK13),k4_ordinal1) | ~v7_ordinal1(k6_numbers) | (~spl14_37 | ~spl14_46 | ~spl14_47 | ~spl14_132 | ~spl14_133)),
  inference(subsumption_resolution,[],[f1981625,f1260])).
fof(f1981625,plain,(
  ~v1_funct_1(sK13) | ~v5_ordinal1(sK13) | ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | r2_hidden(sK5(sK13),k4_ordinal1) | ~v7_ordinal1(k6_numbers) | (~spl14_37 | ~spl14_47 | ~spl14_132 | ~spl14_133)),
  inference(subsumption_resolution,[],[f1981624,f1265])).
fof(f1981624,plain,(
  ~v1_relat_1(sK13) | ~v1_funct_1(sK13) | ~v5_ordinal1(sK13) | ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | r2_hidden(sK5(sK13),k4_ordinal1) | ~v7_ordinal1(k6_numbers) | (~spl14_37 | ~spl14_132 | ~spl14_133)),
  inference(subsumption_resolution,[],[f1980693,f2052])).
fof(f1980693,plain,(
  np__1 != k1_ordinal5(k4_ordinal1,k6_numbers) | ~v1_relat_1(sK13) | ~v1_funct_1(sK13) | ~v5_ordinal1(sK13) | ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | r2_hidden(sK5(sK13),k4_ordinal1) | ~v7_ordinal1(k6_numbers) | (~spl14_37 | ~spl14_132)),
  inference(superposition,[],[f1025,f2049])).
fof(f1025,plain,(
  ( ! [X0] : (np__1 != k1_funct_1(X0,k6_numbers) | ~v1_relat_1(X0) | ~v1_funct_1(X0) | ~v5_ordinal1(X0) | ~v1_ordinal2(X0) | k4_ordinal1 != k9_xtuple_0(X0) | k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1)) ) | ~spl14_37),
  inference(avatar_component_clause,[],[f1024])).
fof(f1982807,plain,(
  ~spl14_5 | ~spl14_8 | spl14_256 | ~spl14_4176),
  inference(avatar_contradiction_clause,[],[f1982806])).
fof(f1982806,plain,(
  $false | (~spl14_5 | ~spl14_8 | spl14_256 | ~spl14_4176)),
  inference(subsumption_resolution,[],[f1982771,f729])).
fof(f729,plain,(
  v1_xboole_0(np__0) | ~spl14_5),
  inference(avatar_component_clause,[],[f727])).
fof(f727,plain,(
  spl14_5 <=> v1_xboole_0(np__0)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_5])])).
fof(f1982771,plain,(
  ~v1_xboole_0(np__0) | (~spl14_8 | spl14_256 | ~spl14_4176)),
  inference(resolution,[],[f87729,f349925])).
fof(f349925,plain,(
  ( ! [X0] : (~r1_ordinal2(k2_ordinal5(X0),sK13) | ~v1_xboole_0(X0)) ) | (~spl14_8 | spl14_256)),
  inference(superposition,[],[f3752,f1350])).
fof(f1350,plain,(
  ( ! [X0] : (k6_numbers = X0 | ~v1_xboole_0(X0)) ) | ~spl14_8),
  inference(forward_demodulation,[],[f253,f744])).
fof(f253,plain,(
  ( ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0)) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0] : (k1_xboole_0 = X0 | ~v1_xboole_0(X0))),
  inference(ennf_transformation,[],[f62])).
fof(f62,axiom,(
  ! [X0] : (v1_xboole_0(X0) => k1_xboole_0 = X0)),
  file('minsubpreds65/t36_ordinal5',t6_boole)).
fof(f3752,plain,(
  ~r1_ordinal2(k2_ordinal5(k6_numbers),sK13) | spl14_256),
  inference(avatar_component_clause,[],[f3750])).
fof(f3750,plain,(
  spl14_256 <=> r1_ordinal2(k2_ordinal5(k6_numbers),sK13)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_256])])).
fof(f87729,plain,(
  r1_ordinal2(k2_ordinal5(np__0),sK13) | ~spl14_4176),
  inference(avatar_component_clause,[],[f87728])).
fof(f87728,plain,(
  spl14_4176 <=> r1_ordinal2(k2_ordinal5(np__0),sK13)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_4176])])).
fof(f1977665,plain,(
  spl14_4176 | ~spl14_4174 | ~spl14_19821),
  inference(avatar_split_clause,[],[f1977664,f440104,f87713,f87728])).
fof(f87713,plain,(
  spl14_4174 <=> r1_ordinal2(k3_card_3(sK13),sK13)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_4174])])).
fof(f440104,plain,(
  spl14_19821 <=> k2_ordinal5(np__0) = k3_card_3(sK13)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_19821])])).
fof(f1977664,plain,(
  r1_ordinal2(k2_ordinal5(np__0),sK13) | (~spl14_4174 | ~spl14_19821)),
  inference(forward_demodulation,[],[f87714,f440106])).
fof(f440106,plain,(
  k2_ordinal5(np__0) = k3_card_3(sK13) | ~spl14_19821),
  inference(avatar_component_clause,[],[f440104])).
fof(f87714,plain,(
  r1_ordinal2(k3_card_3(sK13),sK13) | ~spl14_4174),
  inference(avatar_component_clause,[],[f87713])).
fof(f1970633,plain,(
  spl14_132 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_68 | ~spl14_133 | ~spl14_15782),
  inference(avatar_split_clause,[],[f1970632,f313626,f2051,f1394,f1268,f1263,f1258,f1253,f2048])).
fof(f1394,plain,(
  spl14_68 <=> np__1 = k1_funct_1(sK13,k6_numbers)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_68])])).
fof(f313626,plain,(
  spl14_15782 <=> k1_funct_1(sK13,k4_card_1(sK4(sK13))) = k1_ordinal5(k4_ordinal1,k4_card_1(sK4(sK13)))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_15782])])).
fof(f1970632,plain,(
  ( ! [X137] : (~v7_ordinal1(X137) | k1_funct_1(sK13,X137) = k1_ordinal5(k4_ordinal1,X137)) ) | (~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_68 | ~spl14_133 | ~spl14_15782)),
  inference(subsumption_resolution,[],[f1970631,f1265])).
fof(f1970631,plain,(
  ( ! [X137] : (~v7_ordinal1(X137) | k1_funct_1(sK13,X137) = k1_ordinal5(k4_ordinal1,X137) | ~v1_relat_1(sK13)) ) | (~spl14_45 | ~spl14_46 | ~spl14_48 | ~spl14_68 | ~spl14_133 | ~spl14_15782)),
  inference(subsumption_resolution,[],[f1970630,f1260])).
fof(f1970630,plain,(
  ( ! [X137] : (~v7_ordinal1(X137) | k1_funct_1(sK13,X137) = k1_ordinal5(k4_ordinal1,X137) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13)) ) | (~spl14_45 | ~spl14_48 | ~spl14_68 | ~spl14_133 | ~spl14_15782)),
  inference(subsumption_resolution,[],[f1970629,f1270])).
fof(f1970629,plain,(
  ( ! [X137] : (~v7_ordinal1(X137) | k1_funct_1(sK13,X137) = k1_ordinal5(k4_ordinal1,X137) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13)) ) | (~spl14_45 | ~spl14_68 | ~spl14_133 | ~spl14_15782)),
  inference(subsumption_resolution,[],[f1970628,f1255])).
fof(f1970628,plain,(
  ( ! [X137] : (~v7_ordinal1(X137) | k1_funct_1(sK13,X137) = k1_ordinal5(k4_ordinal1,X137) | ~v1_ordinal2(sK13) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13)) ) | (~spl14_68 | ~spl14_133 | ~spl14_15782)),
  inference(subsumption_resolution,[],[f1969849,f1396])).
fof(f1396,plain,(
  np__1 = k1_funct_1(sK13,k6_numbers) | ~spl14_68),
  inference(avatar_component_clause,[],[f1394])).
fof(f1969849,plain,(
  ( ! [X137] : (np__1 != k1_funct_1(sK13,k6_numbers) | ~v7_ordinal1(X137) | k1_funct_1(sK13,X137) = k1_ordinal5(k4_ordinal1,X137) | ~v1_ordinal2(sK13) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13)) ) | (~spl14_133 | ~spl14_15782)),
  inference(trivial_inequality_removal,[],[f1969818])).
fof(f1969818,plain,(
  ( ! [X137] : (k1_ordinal5(k4_ordinal1,k4_card_1(sK4(sK13))) != k1_ordinal5(k4_ordinal1,k4_card_1(sK4(sK13))) | np__1 != k1_funct_1(sK13,k6_numbers) | ~v7_ordinal1(X137) | k1_funct_1(sK13,X137) = k1_ordinal5(k4_ordinal1,X137) | ~v1_ordinal2(sK13) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13)) ) | (~spl14_133 | ~spl14_15782)),
  inference(superposition,[],[f2434,f313628])).
fof(f313628,plain,(
  k1_funct_1(sK13,k4_card_1(sK4(sK13))) = k1_ordinal5(k4_ordinal1,k4_card_1(sK4(sK13))) | ~spl14_15782),
  inference(avatar_component_clause,[],[f313626])).
fof(f2434,plain,(
  ( ! [X0,X1] : (k1_funct_1(X0,k4_card_1(sK4(X0))) != k1_ordinal5(k4_ordinal1,k4_card_1(sK4(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | ~v7_ordinal1(X1) | k1_funct_1(X0,X1) = k1_ordinal5(k4_ordinal1,X1) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) ) | ~spl14_133),
  inference(subsumption_resolution,[],[f2433,f2398])).
fof(f2398,plain,(
  ( ! [X0,X1] : (np__1 != k1_funct_1(X0,k6_numbers) | ~v7_ordinal1(X1) | v7_ordinal1(sK4(X0)) | k1_funct_1(X0,X1) = k1_ordinal5(k4_ordinal1,X1) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) ) | ~spl14_133),
  inference(backward_demodulation,[],[f293,f2052])).
fof(f293,plain,(
  ( ! [X0,X1] : (k1_funct_1(X0,k6_numbers) != k1_ordinal5(k4_ordinal1,k6_numbers) | ~v7_ordinal1(X1) | v7_ordinal1(sK4(X0)) | k1_funct_1(X0,X1) = k1_ordinal5(k4_ordinal1,X1) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0] : (! [X1] : (k1_funct_1(X0,X1) = k1_ordinal5(k4_ordinal1,X1) | ~v7_ordinal1(X1)) | (k1_funct_1(X0,k1_nat_1(sK4(X0),np__1)) != k1_ordinal5(k4_ordinal1,k1_nat_1(sK4(X0),np__1)) & k1_funct_1(X0,sK4(X0)) = k1_ordinal5(k4_ordinal1,sK4(X0)) & v7_ordinal1(sK4(X0))) | k1_funct_1(X0,k6_numbers) != k1_ordinal5(k4_ordinal1,k6_numbers) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f198,f199])).
fof(f199,plain,(
  ! [X0] : (? [X2] : (k1_funct_1(X0,k1_nat_1(X2,np__1)) != k1_ordinal5(k4_ordinal1,k1_nat_1(X2,np__1)) & k1_funct_1(X0,X2) = k1_ordinal5(k4_ordinal1,X2) & v7_ordinal1(X2)) => (k1_funct_1(X0,k1_nat_1(sK4(X0),np__1)) != k1_ordinal5(k4_ordinal1,k1_nat_1(sK4(X0),np__1)) & k1_funct_1(X0,sK4(X0)) = k1_ordinal5(k4_ordinal1,sK4(X0)) & v7_ordinal1(sK4(X0))))),
  introduced(choice_axiom,[])).
fof(f198,plain,(
  ! [X0] : (! [X1] : (k1_funct_1(X0,X1) = k1_ordinal5(k4_ordinal1,X1) | ~v7_ordinal1(X1)) | ? [X2] : (k1_funct_1(X0,k1_nat_1(X2,np__1)) != k1_ordinal5(k4_ordinal1,k1_nat_1(X2,np__1)) & k1_funct_1(X0,X2) = k1_ordinal5(k4_ordinal1,X2) & v7_ordinal1(X2)) | k1_funct_1(X0,k6_numbers) != k1_ordinal5(k4_ordinal1,k6_numbers) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))),
  inference(rectify,[],[f145])).
fof(f145,plain,(
  ! [X0] : (! [X2] : (k1_funct_1(X0,X2) = k1_ordinal5(k4_ordinal1,X2) | ~v7_ordinal1(X2)) | ? [X1] : (k1_funct_1(X0,k1_nat_1(X1,np__1)) != k1_ordinal5(k4_ordinal1,k1_nat_1(X1,np__1)) & k1_funct_1(X0,X1) = k1_ordinal5(k4_ordinal1,X1) & v7_ordinal1(X1)) | k1_funct_1(X0,k6_numbers) != k1_ordinal5(k4_ordinal1,k6_numbers) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0] : ((! [X2] : (k1_funct_1(X0,X2) = k1_ordinal5(k4_ordinal1,X2) | ~v7_ordinal1(X2)) | (? [X1] : ((k1_funct_1(X0,k1_nat_1(X1,np__1)) != k1_ordinal5(k4_ordinal1,k1_nat_1(X1,np__1)) & k1_funct_1(X0,X1) = k1_ordinal5(k4_ordinal1,X1)) & v7_ordinal1(X1)) | k1_funct_1(X0,k6_numbers) != k1_ordinal5(k4_ordinal1,k6_numbers))) | (~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)))),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : ((v1_ordinal2(X0) & v5_ordinal1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => ((! [X1] : (v7_ordinal1(X1) => (k1_funct_1(X0,X1) = k1_ordinal5(k4_ordinal1,X1) => k1_funct_1(X0,k1_nat_1(X1,np__1)) = k1_ordinal5(k4_ordinal1,k1_nat_1(X1,np__1)))) & k1_funct_1(X0,k6_numbers) = k1_ordinal5(k4_ordinal1,k6_numbers)) => ! [X2] : (v7_ordinal1(X2) => k1_funct_1(X0,X2) = k1_ordinal5(k4_ordinal1,X2))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X0] : ((v1_ordinal2(X0) & v5_ordinal1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => ((! [X1] : (v7_ordinal1(X1) => (k1_funct_1(X0,X1) = k1_ordinal5(k4_ordinal1,X1) => k1_funct_1(X0,k1_nat_1(X1,np__1)) = k1_ordinal5(k4_ordinal1,k1_nat_1(X1,np__1)))) & k1_funct_1(X0,k6_numbers) = k1_ordinal5(k4_ordinal1,k6_numbers)) => ! [X1] : (v7_ordinal1(X1) => k1_funct_1(X0,X1) = k1_ordinal5(k4_ordinal1,X1))))),
  file('minsubpreds65/t36_ordinal5',s2_nat_1__e15_53__ordinal5)).
fof(f2433,plain,(
  ( ! [X0,X1] : (k1_funct_1(X0,k4_card_1(sK4(X0))) != k1_ordinal5(k4_ordinal1,k4_card_1(sK4(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | ~v7_ordinal1(X1) | k1_funct_1(X0,X1) = k1_ordinal5(k4_ordinal1,X1) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | ~v7_ordinal1(sK4(X0))) ) | ~spl14_133),
  inference(superposition,[],[f2396,f248])).
fof(f248,plain,(
  ( ! [X0] : (k4_card_1(X0) = k1_nat_1(X0,np__1) | ~v7_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0] : (k4_card_1(X0) = k1_nat_1(X0,np__1) | ~v7_ordinal1(X0))),
  inference(ennf_transformation,[],[f57])).
fof(f57,axiom,(
  ! [X0] : (v7_ordinal1(X0) => k4_card_1(X0) = k1_nat_1(X0,np__1))),
  file('minsubpreds65/t36_ordinal5',t38_nat_1)).
fof(f2396,plain,(
  ( ! [X0,X1] : (k1_funct_1(X0,k1_nat_1(sK4(X0),np__1)) != k1_ordinal5(k4_ordinal1,k1_nat_1(sK4(X0),np__1)) | np__1 != k1_funct_1(X0,k6_numbers) | ~v7_ordinal1(X1) | k1_funct_1(X0,X1) = k1_ordinal5(k4_ordinal1,X1) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) ) | ~spl14_133),
  inference(backward_demodulation,[],[f295,f2052])).
fof(f295,plain,(
  ( ! [X0,X1] : (k1_funct_1(X0,k1_nat_1(sK4(X0),np__1)) != k1_ordinal5(k4_ordinal1,k1_nat_1(sK4(X0),np__1)) | ~v7_ordinal1(X1) | k1_funct_1(X0,X1) = k1_ordinal5(k4_ordinal1,X1) | k1_funct_1(X0,k6_numbers) != k1_ordinal5(k4_ordinal1,k6_numbers) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f200])).
fof(f1386597,plain,(
  spl14_19821 | ~spl14_79 | ~spl14_150 | ~spl14_172 | ~spl14_180 | spl14_19823),
  inference(avatar_split_clause,[],[f1386596,f440112,f2700,f2618,f2301,f1490,f440104])).
fof(f1490,plain,(
  spl14_79 <=> v3_ordinal1(np__0)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_79])])).
fof(f440112,plain,(
  spl14_19823 <=> v3_ordinal1(sK3(np__0,k3_card_3(sK13)))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_19823])])).
fof(f1386596,plain,(
  k2_ordinal5(np__0) = k3_card_3(sK13) | (~spl14_79 | ~spl14_150 | ~spl14_172 | ~spl14_180 | spl14_19823)),
  inference(subsumption_resolution,[],[f1386595,f1492])).
fof(f1492,plain,(
  v3_ordinal1(np__0) | ~spl14_79),
  inference(avatar_component_clause,[],[f1490])).
fof(f1386595,plain,(
  k2_ordinal5(np__0) = k3_card_3(sK13) | ~v3_ordinal1(np__0) | (~spl14_150 | ~spl14_172 | ~spl14_180 | spl14_19823)),
  inference(subsumption_resolution,[],[f1386594,f2302])).
fof(f2302,plain,(
  v3_ordinal1(k3_card_3(sK13)) | ~spl14_150),
  inference(avatar_component_clause,[],[f2301])).
fof(f1386594,plain,(
  k2_ordinal5(np__0) = k3_card_3(sK13) | ~v3_ordinal1(k3_card_3(sK13)) | ~v3_ordinal1(np__0) | (~spl14_172 | ~spl14_180 | spl14_19823)),
  inference(subsumption_resolution,[],[f1386593,f2620])).
fof(f2620,plain,(
  v4_ordinal5(k3_card_3(sK13)) | ~spl14_172),
  inference(avatar_component_clause,[],[f2618])).
fof(f1386593,plain,(
  k2_ordinal5(np__0) = k3_card_3(sK13) | ~v4_ordinal5(k3_card_3(sK13)) | ~v3_ordinal1(k3_card_3(sK13)) | ~v3_ordinal1(np__0) | (~spl14_180 | spl14_19823)),
  inference(subsumption_resolution,[],[f1386591,f2702])).
fof(f2702,plain,(
  r2_hidden(np__0,k3_card_3(sK13)) | ~spl14_180),
  inference(avatar_component_clause,[],[f2700])).
fof(f1386591,plain,(
  k2_ordinal5(np__0) = k3_card_3(sK13) | ~r2_hidden(np__0,k3_card_3(sK13)) | ~v4_ordinal5(k3_card_3(sK13)) | ~v3_ordinal1(k3_card_3(sK13)) | ~v3_ordinal1(np__0) | spl14_19823),
  inference(resolution,[],[f440114,f279])).
fof(f279,plain,(
  ( ! [X0,X1] : (v3_ordinal1(sK3(X0,X1)) | k2_ordinal5(X0) = X1 | ~r2_hidden(X0,X1) | ~v4_ordinal5(X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f195])).
fof(f440114,plain,(
  ~v3_ordinal1(sK3(np__0,k3_card_3(sK13))) | spl14_19823),
  inference(avatar_component_clause,[],[f440112])).
fof(f1386539,plain,(
  spl14_19821 | ~spl14_79 | ~spl14_150 | ~spl14_172 | ~spl14_180 | spl14_19822),
  inference(avatar_split_clause,[],[f1386538,f440108,f2700,f2618,f2301,f1490,f440104])).
fof(f440108,plain,(
  spl14_19822 <=> v4_ordinal5(sK3(np__0,k3_card_3(sK13)))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_19822])])).
fof(f1386538,plain,(
  k2_ordinal5(np__0) = k3_card_3(sK13) | (~spl14_79 | ~spl14_150 | ~spl14_172 | ~spl14_180 | spl14_19822)),
  inference(subsumption_resolution,[],[f1386537,f1492])).
fof(f1386537,plain,(
  k2_ordinal5(np__0) = k3_card_3(sK13) | ~v3_ordinal1(np__0) | (~spl14_150 | ~spl14_172 | ~spl14_180 | spl14_19822)),
  inference(subsumption_resolution,[],[f1386536,f2302])).
fof(f1386536,plain,(
  k2_ordinal5(np__0) = k3_card_3(sK13) | ~v3_ordinal1(k3_card_3(sK13)) | ~v3_ordinal1(np__0) | (~spl14_172 | ~spl14_180 | spl14_19822)),
  inference(subsumption_resolution,[],[f1386535,f2620])).
fof(f1386535,plain,(
  k2_ordinal5(np__0) = k3_card_3(sK13) | ~v4_ordinal5(k3_card_3(sK13)) | ~v3_ordinal1(k3_card_3(sK13)) | ~v3_ordinal1(np__0) | (~spl14_180 | spl14_19822)),
  inference(subsumption_resolution,[],[f1386533,f2702])).
fof(f1386533,plain,(
  k2_ordinal5(np__0) = k3_card_3(sK13) | ~r2_hidden(np__0,k3_card_3(sK13)) | ~v4_ordinal5(k3_card_3(sK13)) | ~v3_ordinal1(k3_card_3(sK13)) | ~v3_ordinal1(np__0) | spl14_19822),
  inference(resolution,[],[f440110,f280])).
fof(f280,plain,(
  ( ! [X0,X1] : (v4_ordinal5(sK3(X0,X1)) | k2_ordinal5(X0) = X1 | ~r2_hidden(X0,X1) | ~v4_ordinal5(X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f195])).
fof(f440110,plain,(
  ~v4_ordinal5(sK3(np__0,k3_card_3(sK13))) | spl14_19822),
  inference(avatar_component_clause,[],[f440108])).
fof(f1338674,plain,(
  spl14_273 | ~spl14_131),
  inference(avatar_split_clause,[],[f1338672,f2044,f4161])).
fof(f4161,plain,(
  spl14_273 <=> v3_ordinal1(sK4(sK13))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_273])])).
fof(f2044,plain,(
  spl14_131 <=> v7_ordinal1(sK4(sK13))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_131])])).
fof(f1338672,plain,(
  v3_ordinal1(sK4(sK13)) | ~spl14_131),
  inference(resolution,[],[f2046,f249])).
fof(f249,plain,(
  ( ! [X0] : (~v7_ordinal1(X0) | v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0] : ((v7_ordinal1(X0) & v3_ordinal1(X0)) | ~v7_ordinal1(X0))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0] : (v7_ordinal1(X0) => (v7_ordinal1(X0) & v3_ordinal1(X0)))),
  file('minsubpreds65/t36_ordinal5',cc1_nat_1)).
fof(f2046,plain,(
  v7_ordinal1(sK4(sK13)) | ~spl14_131),
  inference(avatar_component_clause,[],[f2044])).
fof(f1336290,plain,(
  ~spl14_131 | spl14_15782 | ~spl14_1255),
  inference(avatar_split_clause,[],[f694731,f18956,f313626,f2044])).
fof(f18956,plain,(
  spl14_1255 <=> k1_funct_1(sK13,k1_ordinal1(sK4(sK13))) = k1_ordinal5(k4_ordinal1,k1_ordinal1(sK4(sK13)))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_1255])])).
fof(f694731,plain,(
  k1_funct_1(sK13,k4_card_1(sK4(sK13))) = k1_ordinal5(k4_ordinal1,k4_card_1(sK4(sK13))) | ~v7_ordinal1(sK4(sK13)) | ~spl14_1255),
  inference(superposition,[],[f18958,f247])).
fof(f18958,plain,(
  k1_funct_1(sK13,k1_ordinal1(sK4(sK13))) = k1_ordinal5(k4_ordinal1,k1_ordinal1(sK4(sK13))) | ~spl14_1255),
  inference(avatar_component_clause,[],[f18956])).
fof(f1336124,plain,(
  ~spl14_46 | ~spl14_47 | ~spl14_48 | spl14_98 | spl14_4802),
  inference(avatar_contradiction_clause,[],[f1336123])).
fof(f1336123,plain,(
  $false | (~spl14_46 | ~spl14_47 | ~spl14_48 | spl14_98 | spl14_4802)),
  inference(subsumption_resolution,[],[f1336122,f1270])).
fof(f1336122,plain,(
  ~v5_ordinal1(sK13) | (~spl14_46 | ~spl14_47 | spl14_98 | spl14_4802)),
  inference(subsumption_resolution,[],[f1336121,f1719])).
fof(f1719,plain,(
  ~v2_ordinal5(sK13) | spl14_98),
  inference(avatar_component_clause,[],[f1718])).
fof(f1336121,plain,(
  v2_ordinal5(sK13) | ~v5_ordinal1(sK13) | (~spl14_46 | ~spl14_47 | spl14_4802)),
  inference(subsumption_resolution,[],[f1336120,f1265])).
fof(f1336120,plain,(
  ~v1_relat_1(sK13) | v2_ordinal5(sK13) | ~v5_ordinal1(sK13) | (~spl14_46 | spl14_4802)),
  inference(subsumption_resolution,[],[f1336118,f1260])).
fof(f1336118,plain,(
  ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | v2_ordinal5(sK13) | ~v5_ordinal1(sK13) | spl14_4802),
  inference(resolution,[],[f99132,f3773])).
fof(f3773,plain,(
  ( ! [X5] : (r1_ordinal1(sK10(X5),sK11(X5)) | ~v1_funct_1(X5) | ~v1_relat_1(X5) | v2_ordinal5(X5) | ~v5_ordinal1(X5)) )),
  inference(subsumption_resolution,[],[f3772,f657])).
fof(f657,plain,(
  ( ! [X0] : (v3_ordinal1(sK10(X0)) | v2_ordinal5(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f212])).
fof(f3772,plain,(
  ( ! [X5] : (~v5_ordinal1(X5) | ~v1_funct_1(X5) | ~v1_relat_1(X5) | v2_ordinal5(X5) | r1_ordinal1(sK10(X5),sK11(X5)) | ~v3_ordinal1(sK10(X5))) )),
  inference(subsumption_resolution,[],[f3771,f658])).
fof(f658,plain,(
  ( ! [X0] : (v3_ordinal1(sK11(X0)) | v2_ordinal5(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f212])).
fof(f3771,plain,(
  ( ! [X5] : (~v5_ordinal1(X5) | ~v1_funct_1(X5) | ~v1_relat_1(X5) | v2_ordinal5(X5) | r1_ordinal1(sK10(X5),sK11(X5)) | ~v3_ordinal1(sK11(X5)) | ~v3_ordinal1(sK10(X5))) )),
  inference(subsumption_resolution,[],[f3770,f256])).
fof(f3770,plain,(
  ( ! [X5] : (~v5_ordinal1(X5) | ~v1_funct_1(X5) | ~v1_relat_1(X5) | v2_ordinal5(X5) | ~v1_ordinal1(sK11(X5)) | r1_ordinal1(sK10(X5),sK11(X5)) | ~v3_ordinal1(sK11(X5)) | ~v3_ordinal1(sK10(X5))) )),
  inference(resolution,[],[f1708,f683])).
fof(f683,plain,(
  ( ! [X0,X1] : (~r1_tarski(X0,X1) | r1_ordinal1(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f218])).
fof(f1708,plain,(
  ( ! [X2] : (r1_tarski(sK10(X2),sK11(X2)) | ~v5_ordinal1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | v2_ordinal5(X2) | ~v1_ordinal1(sK11(X2))) )),
  inference(resolution,[],[f659,f669])).
fof(f99132,plain,(
  ~r1_ordinal1(sK10(sK13),sK11(sK13)) | spl14_4802),
  inference(avatar_component_clause,[],[f99130])).
fof(f1237352,plain,(
  spl14_98 | ~spl14_202 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_132 | ~spl14_147),
  inference(avatar_split_clause,[],[f1237351,f2282,f2048,f1268,f1263,f1258,f3076,f1718])).
fof(f2282,plain,(
  spl14_147 <=> v7_ordinal1(sK11(sK13))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_147])])).
fof(f1237351,plain,(
  ~r1_tarski(k1_funct_1(sK13,sK10(sK13)),k1_ordinal5(k4_ordinal1,sK11(sK13))) | v2_ordinal5(sK13) | (~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_132 | ~spl14_147)),
  inference(subsumption_resolution,[],[f1237350,f2284])).
fof(f2284,plain,(
  v7_ordinal1(sK11(sK13)) | ~spl14_147),
  inference(avatar_component_clause,[],[f2282])).
fof(f1237350,plain,(
  ~r1_tarski(k1_funct_1(sK13,sK10(sK13)),k1_ordinal5(k4_ordinal1,sK11(sK13))) | v2_ordinal5(sK13) | ~v7_ordinal1(sK11(sK13)) | (~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_132)),
  inference(subsumption_resolution,[],[f1237349,f1265])).
fof(f1237349,plain,(
  ~r1_tarski(k1_funct_1(sK13,sK10(sK13)),k1_ordinal5(k4_ordinal1,sK11(sK13))) | v2_ordinal5(sK13) | ~v1_relat_1(sK13) | ~v7_ordinal1(sK11(sK13)) | (~spl14_46 | ~spl14_48 | ~spl14_132)),
  inference(subsumption_resolution,[],[f1237348,f1260])).
fof(f1237348,plain,(
  ~r1_tarski(k1_funct_1(sK13,sK10(sK13)),k1_ordinal5(k4_ordinal1,sK11(sK13))) | v2_ordinal5(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | ~v7_ordinal1(sK11(sK13)) | (~spl14_48 | ~spl14_132)),
  inference(subsumption_resolution,[],[f615931,f1270])).
fof(f615931,plain,(
  ~r1_tarski(k1_funct_1(sK13,sK10(sK13)),k1_ordinal5(k4_ordinal1,sK11(sK13))) | v2_ordinal5(sK13) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | ~v7_ordinal1(sK11(sK13)) | ~spl14_132),
  inference(superposition,[],[f661,f2049])).
fof(f1057514,plain,(
  spl14_429 | ~spl14_13 | ~spl14_409),
  inference(avatar_split_clause,[],[f1057513,f6260,f767,f6642])).
fof(f6260,plain,(
  spl14_409 <=> r2_hidden(sK10(sK13),k4_ordinal1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_409])])).
fof(f1057513,plain,(
  v3_ordinal1(sK10(sK13)) | (~spl14_13 | ~spl14_409)),
  inference(subsumption_resolution,[],[f1057480,f769])).
fof(f1057480,plain,(
  ~v3_ordinal1(k4_ordinal1) | v3_ordinal1(sK10(sK13)) | ~spl14_409),
  inference(resolution,[],[f6262,f6680])).
fof(f6262,plain,(
  r2_hidden(sK10(sK13),k4_ordinal1) | ~spl14_409),
  inference(avatar_component_clause,[],[f6260])).
fof(f1029081,plain,(
  ~spl14_137 | spl14_142 | ~spl14_21 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_68 | ~spl14_120),
  inference(avatar_split_clause,[],[f1029080,f1937,f1394,f1268,f1263,f1258,f1253,f1248,f805,f2194,f2113])).
fof(f805,plain,(
  spl14_21 <=> ! [X1,X0] : (r1_ordinal1(k3_card_3(X0),X1) | ~v1_relat_1(X0) | ~v1_funct_1(X0) | ~v5_ordinal1(X0) | ~v1_ordinal2(X0) | k4_ordinal1 != k9_xtuple_0(X0) | np__1 != k1_funct_1(X0,k6_numbers) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | ~v3_ordinal1(X1) | ~r1_ordinal1(np__1,X1) | k12_ordinal2(k4_ordinal1,X1) != X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_21])])).
fof(f1937,plain,(
  spl14_120 <=> v3_ordinal1(sK5(sK13))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_120])])).
fof(f1029080,plain,(
  ( ! [X2] : (r1_ordinal1(k3_card_3(sK13),X2) | ~v3_ordinal1(X2) | ~r1_ordinal1(np__1,X2) | k12_ordinal2(k4_ordinal1,X2) != X2 | ~r2_hidden(k1_ordinal1(sK5(sK13)),k4_ordinal1)) ) | (~spl14_21 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_68 | ~spl14_120)),
  inference(subsumption_resolution,[],[f1029079,f1939])).
fof(f1939,plain,(
  v3_ordinal1(sK5(sK13)) | ~spl14_120),
  inference(avatar_component_clause,[],[f1937])).
fof(f1029079,plain,(
  ( ! [X2] : (r1_ordinal1(k3_card_3(sK13),X2) | ~v3_ordinal1(X2) | ~r1_ordinal1(np__1,X2) | k12_ordinal2(k4_ordinal1,X2) != X2 | ~r2_hidden(k1_ordinal1(sK5(sK13)),k4_ordinal1) | ~v3_ordinal1(sK5(sK13))) ) | (~spl14_21 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_68)),
  inference(subsumption_resolution,[],[f1029078,f1396])).
fof(f1029078,plain,(
  ( ! [X2] : (np__1 != k1_funct_1(sK13,k6_numbers) | r1_ordinal1(k3_card_3(sK13),X2) | ~v3_ordinal1(X2) | ~r1_ordinal1(np__1,X2) | k12_ordinal2(k4_ordinal1,X2) != X2 | ~r2_hidden(k1_ordinal1(sK5(sK13)),k4_ordinal1) | ~v3_ordinal1(sK5(sK13))) ) | (~spl14_21 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48)),
  inference(subsumption_resolution,[],[f1029077,f1250])).
fof(f1029077,plain,(
  ( ! [X2] : (k4_ordinal1 != k9_xtuple_0(sK13) | np__1 != k1_funct_1(sK13,k6_numbers) | r1_ordinal1(k3_card_3(sK13),X2) | ~v3_ordinal1(X2) | ~r1_ordinal1(np__1,X2) | k12_ordinal2(k4_ordinal1,X2) != X2 | ~r2_hidden(k1_ordinal1(sK5(sK13)),k4_ordinal1) | ~v3_ordinal1(sK5(sK13))) ) | (~spl14_21 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48)),
  inference(subsumption_resolution,[],[f1029076,f1255])).
fof(f1029076,plain,(
  ( ! [X2] : (~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | np__1 != k1_funct_1(sK13,k6_numbers) | r1_ordinal1(k3_card_3(sK13),X2) | ~v3_ordinal1(X2) | ~r1_ordinal1(np__1,X2) | k12_ordinal2(k4_ordinal1,X2) != X2 | ~r2_hidden(k1_ordinal1(sK5(sK13)),k4_ordinal1) | ~v3_ordinal1(sK5(sK13))) ) | (~spl14_21 | ~spl14_46 | ~spl14_47 | ~spl14_48)),
  inference(subsumption_resolution,[],[f1029075,f1270])).
fof(f1029075,plain,(
  ( ! [X2] : (~v5_ordinal1(sK13) | ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | np__1 != k1_funct_1(sK13,k6_numbers) | r1_ordinal1(k3_card_3(sK13),X2) | ~v3_ordinal1(X2) | ~r1_ordinal1(np__1,X2) | k12_ordinal2(k4_ordinal1,X2) != X2 | ~r2_hidden(k1_ordinal1(sK5(sK13)),k4_ordinal1) | ~v3_ordinal1(sK5(sK13))) ) | (~spl14_21 | ~spl14_46 | ~spl14_47)),
  inference(subsumption_resolution,[],[f1029074,f1260])).
fof(f1029074,plain,(
  ( ! [X2] : (~v1_funct_1(sK13) | ~v5_ordinal1(sK13) | ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | np__1 != k1_funct_1(sK13,k6_numbers) | r1_ordinal1(k3_card_3(sK13),X2) | ~v3_ordinal1(X2) | ~r1_ordinal1(np__1,X2) | k12_ordinal2(k4_ordinal1,X2) != X2 | ~r2_hidden(k1_ordinal1(sK5(sK13)),k4_ordinal1) | ~v3_ordinal1(sK5(sK13))) ) | (~spl14_21 | ~spl14_47)),
  inference(subsumption_resolution,[],[f1029058,f1265])).
fof(f1029058,plain,(
  ( ! [X2] : (~v1_relat_1(sK13) | ~v1_funct_1(sK13) | ~v5_ordinal1(sK13) | ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | np__1 != k1_funct_1(sK13,k6_numbers) | r1_ordinal1(k3_card_3(sK13),X2) | ~v3_ordinal1(X2) | ~r1_ordinal1(np__1,X2) | k12_ordinal2(k4_ordinal1,X2) != X2 | ~r2_hidden(k1_ordinal1(sK5(sK13)),k4_ordinal1) | ~v3_ordinal1(sK5(sK13))) ) | ~spl14_21),
  inference(trivial_inequality_removal,[],[f1029056])).
fof(f1029056,plain,(
  ( ! [X2] : (k1_funct_1(sK13,k1_ordinal1(sK5(sK13))) != k1_funct_1(sK13,k1_ordinal1(sK5(sK13))) | ~v1_relat_1(sK13) | ~v1_funct_1(sK13) | ~v5_ordinal1(sK13) | ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | np__1 != k1_funct_1(sK13,k6_numbers) | r1_ordinal1(k3_card_3(sK13),X2) | ~v3_ordinal1(X2) | ~r1_ordinal1(np__1,X2) | k12_ordinal2(k4_ordinal1,X2) != X2 | ~r2_hidden(k1_ordinal1(sK5(sK13)),k4_ordinal1) | ~v3_ordinal1(sK5(sK13))) ) | ~spl14_21),
  inference(superposition,[],[f806,f695])).
fof(f695,plain,(
  ( ! [X2] : (k1_funct_1(sK13,k1_ordinal1(X2)) = k12_ordinal2(k4_ordinal1,k1_funct_1(sK13,X2)) | ~r2_hidden(k1_ordinal1(X2),k4_ordinal1) | ~v3_ordinal1(X2)) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X1] : (k1_funct_1(sK13,X1) = k8_ordinal2(k5_relat_1(sK13,X1)) | k1_xboole_0 = X1 | ~v4_ordinal1(X1) | ~r2_hidden(X1,k4_ordinal1) | ~v3_ordinal1(X1)) & ! [X2] : (k1_funct_1(sK13,k1_ordinal1(X2)) = k12_ordinal2(k4_ordinal1,k1_funct_1(sK13,X2)) | ~r2_hidden(k1_ordinal1(X2),k4_ordinal1) | ~v3_ordinal1(X2)) & (np__1 = k1_funct_1(sK13,k1_xboole_0) | ~r2_hidden(k1_xboole_0,k4_ordinal1)) & k4_ordinal1 = k9_xtuple_0(sK13) & v1_ordinal2(sK13) & v1_funct_1(sK13) & v1_relat_1(sK13) & v5_ordinal1(sK13)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f180,f220])).
fof(f220,plain,(
  ? [X0] : (! [X1] : (k1_funct_1(X0,X1) = k8_ordinal2(k5_relat_1(X0,X1)) | k1_xboole_0 = X1 | ~v4_ordinal1(X1) | ~r2_hidden(X1,k4_ordinal1) | ~v3_ordinal1(X1)) & ! [X2] : (k1_funct_1(X0,k1_ordinal1(X2)) = k12_ordinal2(k4_ordinal1,k1_funct_1(X0,X2)) | ~r2_hidden(k1_ordinal1(X2),k4_ordinal1) | ~v3_ordinal1(X2)) & (np__1 = k1_funct_1(X0,k1_xboole_0) | ~r2_hidden(k1_xboole_0,k4_ordinal1)) & k4_ordinal1 = k9_xtuple_0(X0) & v1_ordinal2(X0) & v1_funct_1(X0) & v1_relat_1(X0) & v5_ordinal1(X0)) => (! [X1] : (k1_funct_1(sK13,X1) = k8_ordinal2(k5_relat_1(sK13,X1)) | k1_xboole_0 = X1 | ~v4_ordinal1(X1) | ~r2_hidden(X1,k4_ordinal1) | ~v3_ordinal1(X1)) & ! [X2] : (k1_funct_1(sK13,k1_ordinal1(X2)) = k12_ordinal2(k4_ordinal1,k1_funct_1(sK13,X2)) | ~r2_hidden(k1_ordinal1(X2),k4_ordinal1) | ~v3_ordinal1(X2)) & (np__1 = k1_funct_1(sK13,k1_xboole_0) | ~r2_hidden(k1_xboole_0,k4_ordinal1)) & k4_ordinal1 = k9_xtuple_0(sK13) & v1_ordinal2(sK13) & v1_funct_1(sK13) & v1_relat_1(sK13) & v5_ordinal1(sK13))),
  introduced(choice_axiom,[])).
fof(f180,plain,(
  ? [X0] : (! [X1] : (k1_funct_1(X0,X1) = k8_ordinal2(k5_relat_1(X0,X1)) | k1_xboole_0 = X1 | ~v4_ordinal1(X1) | ~r2_hidden(X1,k4_ordinal1) | ~v3_ordinal1(X1)) & ! [X2] : (k1_funct_1(X0,k1_ordinal1(X2)) = k12_ordinal2(k4_ordinal1,k1_funct_1(X0,X2)) | ~r2_hidden(k1_ordinal1(X2),k4_ordinal1) | ~v3_ordinal1(X2)) & (np__1 = k1_funct_1(X0,k1_xboole_0) | ~r2_hidden(k1_xboole_0,k4_ordinal1)) & k4_ordinal1 = k9_xtuple_0(X0) & v1_ordinal2(X0) & v1_funct_1(X0) & v1_relat_1(X0) & v5_ordinal1(X0))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ? [X0] : (! [X1] : (((k1_funct_1(X0,X1) = k8_ordinal2(k5_relat_1(X0,X1)) | k1_xboole_0 = X1) | (~v4_ordinal1(X1) | ~r2_hidden(X1,k4_ordinal1))) | ~v3_ordinal1(X1)) & ! [X2] : ((k1_funct_1(X0,k1_ordinal1(X2)) = k12_ordinal2(k4_ordinal1,k1_funct_1(X0,X2)) | ~r2_hidden(k1_ordinal1(X2),k4_ordinal1)) | ~v3_ordinal1(X2)) & (np__1 = k1_funct_1(X0,k1_xboole_0) | ~r2_hidden(k1_xboole_0,k4_ordinal1)) & k4_ordinal1 = k9_xtuple_0(X0) & v1_ordinal2(X0) & v1_funct_1(X0) & v1_relat_1(X0) & v5_ordinal1(X0))),
  inference(ennf_transformation,[],[f91])).
fof(f91,plain,(
  ? [X0] : (! [X1] : (v3_ordinal1(X1) => ((v4_ordinal1(X1) & r2_hidden(X1,k4_ordinal1)) => (k1_funct_1(X0,X1) = k8_ordinal2(k5_relat_1(X0,X1)) | k1_xboole_0 = X1))) & ! [X2] : (v3_ordinal1(X2) => (r2_hidden(k1_ordinal1(X2),k4_ordinal1) => k1_funct_1(X0,k1_ordinal1(X2)) = k12_ordinal2(k4_ordinal1,k1_funct_1(X0,X2)))) & (r2_hidden(k1_xboole_0,k4_ordinal1) => np__1 = k1_funct_1(X0,k1_xboole_0)) & k4_ordinal1 = k9_xtuple_0(X0) & v1_ordinal2(X0) & v1_funct_1(X0) & v1_relat_1(X0) & v5_ordinal1(X0))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ? [X0] : (! [X1] : (v3_ordinal1(X1) => ((v4_ordinal1(X1) & r2_hidden(X1,k4_ordinal1)) => (k1_funct_1(X0,X1) = k8_ordinal2(k5_relat_1(X0,X1)) | k1_xboole_0 = X1))) & ! [X1] : (v3_ordinal1(X1) => (r2_hidden(k1_ordinal1(X1),k4_ordinal1) => k1_funct_1(X0,k1_ordinal1(X1)) = k12_ordinal2(k4_ordinal1,k1_funct_1(X0,X1)))) & (r2_hidden(k1_xboole_0,k4_ordinal1) => np__1 = k1_funct_1(X0,k1_xboole_0)) & k4_ordinal1 = k9_xtuple_0(X0) & v1_ordinal2(X0) & v1_funct_1(X0) & v1_relat_1(X0) & v5_ordinal1(X0))),
  file('minsubpreds65/t36_ordinal5',s11_ordinal2__e3_53__ordinal5)).
fof(f806,plain,(
  ( ! [X0,X1] : (k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | ~v1_relat_1(X0) | ~v1_funct_1(X0) | ~v5_ordinal1(X0) | ~v1_ordinal2(X0) | k4_ordinal1 != k9_xtuple_0(X0) | np__1 != k1_funct_1(X0,k6_numbers) | r1_ordinal1(k3_card_3(X0),X1) | ~v3_ordinal1(X1) | ~r1_ordinal1(np__1,X1) | k12_ordinal2(k4_ordinal1,X1) != X1) ) | ~spl14_21),
  inference(avatar_component_clause,[],[f805])).
fof(f1029048,plain,(
  ~spl14_137 | spl14_134 | ~spl14_36 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_68 | ~spl14_120),
  inference(avatar_split_clause,[],[f1029047,f1937,f1394,f1268,f1263,f1258,f1253,f1248,f981,f2064,f2113])).
fof(f981,plain,(
  spl14_36 <=> ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | ~v1_relat_1(X0) | ~v1_funct_1(X0) | ~v5_ordinal1(X0) | ~v1_ordinal2(X0) | k4_ordinal1 != k9_xtuple_0(X0) | np__1 != k1_funct_1(X0,k6_numbers) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_36])])).
fof(f1029047,plain,(
  k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | ~r2_hidden(k1_ordinal1(sK5(sK13)),k4_ordinal1) | (~spl14_36 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_68 | ~spl14_120)),
  inference(subsumption_resolution,[],[f1029046,f1939])).
fof(f1029046,plain,(
  k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | ~r2_hidden(k1_ordinal1(sK5(sK13)),k4_ordinal1) | ~v3_ordinal1(sK5(sK13)) | (~spl14_36 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_68)),
  inference(subsumption_resolution,[],[f1029045,f1396])).
fof(f1029045,plain,(
  np__1 != k1_funct_1(sK13,k6_numbers) | k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | ~r2_hidden(k1_ordinal1(sK5(sK13)),k4_ordinal1) | ~v3_ordinal1(sK5(sK13)) | (~spl14_36 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48)),
  inference(subsumption_resolution,[],[f1029044,f1250])).
fof(f1029044,plain,(
  k4_ordinal1 != k9_xtuple_0(sK13) | np__1 != k1_funct_1(sK13,k6_numbers) | k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | ~r2_hidden(k1_ordinal1(sK5(sK13)),k4_ordinal1) | ~v3_ordinal1(sK5(sK13)) | (~spl14_36 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48)),
  inference(subsumption_resolution,[],[f1029043,f1255])).
fof(f1029043,plain,(
  ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | np__1 != k1_funct_1(sK13,k6_numbers) | k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | ~r2_hidden(k1_ordinal1(sK5(sK13)),k4_ordinal1) | ~v3_ordinal1(sK5(sK13)) | (~spl14_36 | ~spl14_46 | ~spl14_47 | ~spl14_48)),
  inference(subsumption_resolution,[],[f1029042,f1270])).
fof(f1029042,plain,(
  ~v5_ordinal1(sK13) | ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | np__1 != k1_funct_1(sK13,k6_numbers) | k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | ~r2_hidden(k1_ordinal1(sK5(sK13)),k4_ordinal1) | ~v3_ordinal1(sK5(sK13)) | (~spl14_36 | ~spl14_46 | ~spl14_47)),
  inference(subsumption_resolution,[],[f1029041,f1260])).
fof(f1029041,plain,(
  ~v1_funct_1(sK13) | ~v5_ordinal1(sK13) | ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | np__1 != k1_funct_1(sK13,k6_numbers) | k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | ~r2_hidden(k1_ordinal1(sK5(sK13)),k4_ordinal1) | ~v3_ordinal1(sK5(sK13)) | (~spl14_36 | ~spl14_47)),
  inference(subsumption_resolution,[],[f1029025,f1265])).
fof(f1029025,plain,(
  ~v1_relat_1(sK13) | ~v1_funct_1(sK13) | ~v5_ordinal1(sK13) | ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | np__1 != k1_funct_1(sK13,k6_numbers) | k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | ~r2_hidden(k1_ordinal1(sK5(sK13)),k4_ordinal1) | ~v3_ordinal1(sK5(sK13)) | ~spl14_36),
  inference(trivial_inequality_removal,[],[f1029023])).
fof(f1029023,plain,(
  k1_funct_1(sK13,k1_ordinal1(sK5(sK13))) != k1_funct_1(sK13,k1_ordinal1(sK5(sK13))) | ~v1_relat_1(sK13) | ~v1_funct_1(sK13) | ~v5_ordinal1(sK13) | ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | np__1 != k1_funct_1(sK13,k6_numbers) | k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | ~r2_hidden(k1_ordinal1(sK5(sK13)),k4_ordinal1) | ~v3_ordinal1(sK5(sK13)) | ~spl14_36),
  inference(superposition,[],[f982,f695])).
fof(f982,plain,(
  ( ! [X0] : (k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | ~v1_relat_1(X0) | ~v1_funct_1(X0) | ~v5_ordinal1(X0) | ~v1_ordinal2(X0) | k4_ordinal1 != k9_xtuple_0(X0) | np__1 != k1_funct_1(X0,k6_numbers) | k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0))) ) | ~spl14_36),
  inference(avatar_component_clause,[],[f981])).
fof(f1016851,plain,(
  spl14_231 | spl14_3675 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_882),
  inference(avatar_split_clause,[],[f1016850,f13055,f1268,f1263,f1258,f1253,f78931,f3475])).
fof(f3475,plain,(
  spl14_231 <=> ! [X1] : (~r1_ordinal2(X1,sK13) | ~v3_ordinal1(X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_231])])).
fof(f13055,plain,(
  spl14_882 <=> r2_hidden(np__0,k8_ordinal2(sK13))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_882])])).
fof(f1016850,plain,(
  ( ! [X0,X1] : (r2_hidden(np__0,X0) | ~r1_ordinal2(X0,sK13) | ~v3_ordinal1(X0) | ~r1_ordinal2(X1,sK13) | ~v3_ordinal1(X1)) ) | (~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_882)),
  inference(subsumption_resolution,[],[f1016849,f1270])).
fof(f1016849,plain,(
  ( ! [X0,X1] : (r2_hidden(np__0,X0) | ~r1_ordinal2(X0,sK13) | ~v3_ordinal1(X0) | ~r1_ordinal2(X1,sK13) | ~v3_ordinal1(X1) | ~v5_ordinal1(sK13)) ) | (~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_882)),
  inference(subsumption_resolution,[],[f1016848,f1265])).
fof(f1016848,plain,(
  ( ! [X0,X1] : (r2_hidden(np__0,X0) | ~r1_ordinal2(X0,sK13) | ~v3_ordinal1(X0) | ~r1_ordinal2(X1,sK13) | ~v3_ordinal1(X1) | ~v1_relat_1(sK13) | ~v5_ordinal1(sK13)) ) | (~spl14_45 | ~spl14_46 | ~spl14_882)),
  inference(subsumption_resolution,[],[f1016847,f1260])).
fof(f1016847,plain,(
  ( ! [X0,X1] : (r2_hidden(np__0,X0) | ~r1_ordinal2(X0,sK13) | ~v3_ordinal1(X0) | ~r1_ordinal2(X1,sK13) | ~v3_ordinal1(X1) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | ~v5_ordinal1(sK13)) ) | (~spl14_45 | ~spl14_882)),
  inference(subsumption_resolution,[],[f659870,f1255])).
fof(f659870,plain,(
  ( ! [X0,X1] : (r2_hidden(np__0,X0) | ~r1_ordinal2(X0,sK13) | ~v3_ordinal1(X0) | ~r1_ordinal2(X1,sK13) | ~v3_ordinal1(X1) | ~v1_ordinal2(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | ~v5_ordinal1(sK13)) ) | ~spl14_882),
  inference(superposition,[],[f13057,f288])).
fof(f13057,plain,(
  r2_hidden(np__0,k8_ordinal2(sK13)) | ~spl14_882),
  inference(avatar_component_clause,[],[f13055])).
fof(f656277,plain,(
  spl14_882 | ~spl14_12 | ~spl14_13 | ~spl14_54 | ~spl14_79 | ~spl14_95 | ~spl14_117 | ~spl14_2736),
  inference(avatar_split_clause,[],[f656276,f49957,f1893,f1644,f1490,f1304,f767,f762,f13055])).
fof(f762,plain,(
  spl14_12 <=> k1_ordinal1(np__0) = np__1),
  introduced(avatar_definition,[new_symbols(naming,[spl14_12])])).
fof(f1644,plain,(
  spl14_95 <=> r1_ordinal1(k6_numbers,k4_ordinal1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_95])])).
fof(f1893,plain,(
  spl14_117 <=> ! [X1,X0] : (r1_ordinal1(np__1,k1_ordinal5(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) | ~r1_ordinal1(k6_numbers,X1) | ~r2_hidden(k6_numbers,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_117])])).
fof(f656276,plain,(
  r2_hidden(np__0,k8_ordinal2(sK13)) | (~spl14_12 | ~spl14_13 | ~spl14_54 | ~spl14_79 | ~spl14_95 | ~spl14_117 | ~spl14_2736)),
  inference(subsumption_resolution,[],[f656275,f1305])).
fof(f656275,plain,(
  r2_hidden(np__0,k8_ordinal2(sK13)) | ~r2_hidden(k6_numbers,k4_ordinal1) | (~spl14_12 | ~spl14_13 | ~spl14_79 | ~spl14_95 | ~spl14_117 | ~spl14_2736)),
  inference(subsumption_resolution,[],[f656274,f1646])).
fof(f1646,plain,(
  r1_ordinal1(k6_numbers,k4_ordinal1) | ~spl14_95),
  inference(avatar_component_clause,[],[f1644])).
fof(f656274,plain,(
  r2_hidden(np__0,k8_ordinal2(sK13)) | ~r1_ordinal1(k6_numbers,k4_ordinal1) | ~r2_hidden(k6_numbers,k4_ordinal1) | (~spl14_12 | ~spl14_13 | ~spl14_79 | ~spl14_117 | ~spl14_2736)),
  inference(subsumption_resolution,[],[f656121,f769])).
fof(f656121,plain,(
  r2_hidden(np__0,k8_ordinal2(sK13)) | ~v3_ordinal1(k4_ordinal1) | ~r1_ordinal1(k6_numbers,k4_ordinal1) | ~r2_hidden(k6_numbers,k4_ordinal1) | (~spl14_12 | ~spl14_79 | ~spl14_117 | ~spl14_2736)),
  inference(duplicate_literal_removal,[],[f656029])).
fof(f656029,plain,(
  r2_hidden(np__0,k8_ordinal2(sK13)) | ~v3_ordinal1(k4_ordinal1) | ~r1_ordinal1(k6_numbers,k4_ordinal1) | ~r2_hidden(k6_numbers,k4_ordinal1) | ~v3_ordinal1(k4_ordinal1) | (~spl14_12 | ~spl14_79 | ~spl14_117 | ~spl14_2736)),
  inference(superposition,[],[f3694,f49959])).
fof(f49959,plain,(
  k1_ordinal5(k4_ordinal1,k4_ordinal1) = k8_ordinal2(sK13) | ~spl14_2736),
  inference(avatar_component_clause,[],[f49957])).
fof(f3694,plain,(
  ( ! [X0,X1] : (r2_hidden(np__0,k1_ordinal5(X1,X0)) | ~v3_ordinal1(X1) | ~r1_ordinal1(k6_numbers,X0) | ~r2_hidden(k6_numbers,X1) | ~v3_ordinal1(X0)) ) | (~spl14_12 | ~spl14_79 | ~spl14_117)),
  inference(subsumption_resolution,[],[f3681,f680])).
fof(f3681,plain,(
  ( ! [X0,X1] : (~v3_ordinal1(X0) | ~v3_ordinal1(X1) | ~r1_ordinal1(k6_numbers,X0) | ~r2_hidden(k6_numbers,X1) | r2_hidden(np__0,k1_ordinal5(X1,X0)) | ~v3_ordinal1(k1_ordinal5(X1,X0))) ) | (~spl14_12 | ~spl14_79 | ~spl14_117)),
  inference(resolution,[],[f1894,f1666])).
fof(f1666,plain,(
  ( ! [X0] : (~r1_ordinal1(np__1,X0) | r2_hidden(np__0,X0) | ~v3_ordinal1(X0)) ) | (~spl14_12 | ~spl14_79)),
  inference(subsumption_resolution,[],[f1663,f1492])).
fof(f1663,plain,(
  ( ! [X0] : (~r1_ordinal1(np__1,X0) | r2_hidden(np__0,X0) | ~v3_ordinal1(X0) | ~v3_ordinal1(np__0)) ) | ~spl14_12),
  inference(superposition,[],[f273,f764])).
fof(f764,plain,(
  k1_ordinal1(np__0) = np__1 | ~spl14_12),
  inference(avatar_component_clause,[],[f762])).
fof(f273,plain,(
  ( ! [X0,X1] : (~r1_ordinal1(k1_ordinal1(X0),X1) | r2_hidden(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0] : (! [X1] : (((r2_hidden(X0,X1) | ~r1_ordinal1(k1_ordinal1(X0),X1)) & (r1_ordinal1(k1_ordinal1(X0),X1) | ~r2_hidden(X0,X1))) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0))),
  inference(nnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0] : (! [X1] : ((r2_hidden(X0,X1) <=> r1_ordinal1(k1_ordinal1(X0),X1)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0))),
  inference(ennf_transformation,[],[f52])).
fof(f52,axiom,(
  ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => (r2_hidden(X0,X1) <=> r1_ordinal1(k1_ordinal1(X0),X1))))),
  file('minsubpreds65/t36_ordinal5',t21_ordinal1)).
fof(f1894,plain,(
  ( ! [X0,X1] : (r1_ordinal1(np__1,k1_ordinal5(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0) | ~r1_ordinal1(k6_numbers,X1) | ~r2_hidden(k6_numbers,X0)) ) | ~spl14_117),
  inference(avatar_component_clause,[],[f1893])).
fof(f630924,plain,(
  spl14_153 | ~spl14_13 | ~spl14_149),
  inference(avatar_split_clause,[],[f630923,f2294,f767,f2322])).
fof(f2294,plain,(
  spl14_149 <=> m1_subset_1(sK11(sK13),k4_ordinal1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_149])])).
fof(f630923,plain,(
  v3_ordinal1(sK11(sK13)) | (~spl14_13 | ~spl14_149)),
  inference(subsumption_resolution,[],[f630922,f769])).
fof(f630922,plain,(
  v3_ordinal1(sK11(sK13)) | ~v3_ordinal1(k4_ordinal1) | ~spl14_149),
  inference(resolution,[],[f2296,f276])).
fof(f2296,plain,(
  m1_subset_1(sK11(sK13),k4_ordinal1) | ~spl14_149),
  inference(avatar_component_clause,[],[f2294])).
fof(f628860,plain,(
  spl14_580 | ~spl14_13 | ~spl14_569 | spl14_625),
  inference(avatar_split_clause,[],[f628859,f8979,f8376,f767,f8480])).
fof(f8979,plain,(
  spl14_625 <=> r2_hidden(sK1(sK6,k4_ordinal1,sK7),sK6)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_625])])).
fof(f628859,plain,(
  k12_ordinal2(k4_ordinal1,sK6) = k8_ordinal2(sK7) | (~spl14_13 | ~spl14_569 | spl14_625)),
  inference(subsumption_resolution,[],[f628818,f769])).
fof(f628818,plain,(
  ~v3_ordinal1(k4_ordinal1) | k12_ordinal2(k4_ordinal1,sK6) = k8_ordinal2(sK7) | (~spl14_569 | spl14_625)),
  inference(resolution,[],[f8377,f8981])).
fof(f8981,plain,(
  ~r2_hidden(sK1(sK6,k4_ordinal1,sK7),sK6) | spl14_625),
  inference(avatar_component_clause,[],[f8979])).
fof(f8377,plain,(
  ( ! [X3] : (r2_hidden(sK1(sK6,X3,sK7),sK6) | ~v3_ordinal1(X3) | k8_ordinal2(sK7) = k12_ordinal2(X3,sK6)) ) | ~spl14_569),
  inference(avatar_component_clause,[],[f8376])).
fof(f625196,plain,(
  spl14_147 | ~spl14_99),
  inference(avatar_split_clause,[],[f625153,f1722,f2282])).
fof(f1722,plain,(
  spl14_99 <=> r2_hidden(sK11(sK13),k4_ordinal1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_99])])).
fof(f625153,plain,(
  v7_ordinal1(sK11(sK13)) | ~spl14_99),
  inference(resolution,[],[f1724,f673])).
fof(f1724,plain,(
  r2_hidden(sK11(sK13),k4_ordinal1) | ~spl14_99),
  inference(avatar_component_clause,[],[f1722])).
fof(f622757,plain,(
  spl14_557 | spl14_1540 | ~spl14_26 | ~spl14_28 | ~spl14_29 | ~spl14_30 | ~spl14_551),
  inference(avatar_split_clause,[],[f622756,f8212,f863,f855,f847,f831,f22685,f8279])).
fof(f622756,plain,(
  r2_hidden(sK10(sK7),sK6) | v2_ordinal5(sK7) | (~spl14_26 | ~spl14_28 | ~spl14_29 | ~spl14_30 | ~spl14_551)),
  inference(subsumption_resolution,[],[f622755,f8214])).
fof(f8214,plain,(
  v1_ordinal1(sK6) | ~spl14_551),
  inference(avatar_component_clause,[],[f8212])).
fof(f622755,plain,(
  r2_hidden(sK10(sK7),sK6) | v2_ordinal5(sK7) | ~v1_ordinal1(sK6) | (~spl14_26 | ~spl14_28 | ~spl14_29 | ~spl14_30)),
  inference(subsumption_resolution,[],[f622754,f849])).
fof(f622754,plain,(
  r2_hidden(sK10(sK7),sK6) | ~v5_ordinal1(sK7) | v2_ordinal5(sK7) | ~v1_ordinal1(sK6) | (~spl14_26 | ~spl14_29 | ~spl14_30)),
  inference(subsumption_resolution,[],[f622753,f865])).
fof(f622753,plain,(
  r2_hidden(sK10(sK7),sK6) | ~v1_relat_1(sK7) | ~v5_ordinal1(sK7) | v2_ordinal5(sK7) | ~v1_ordinal1(sK6) | (~spl14_26 | ~spl14_29)),
  inference(subsumption_resolution,[],[f622477,f857])).
fof(f622477,plain,(
  r2_hidden(sK10(sK7),sK6) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | ~v5_ordinal1(sK7) | v2_ordinal5(sK7) | ~v1_ordinal1(sK6) | ~spl14_26),
  inference(superposition,[],[f3884,f833])).
fof(f3884,plain,(
  ( ! [X7] : (r2_hidden(sK10(X7),k9_xtuple_0(X7)) | ~v1_funct_1(X7) | ~v1_relat_1(X7) | ~v5_ordinal1(X7) | v2_ordinal5(X7) | ~v1_ordinal1(k9_xtuple_0(X7))) )),
  inference(duplicate_literal_removal,[],[f3882])).
fof(f3882,plain,(
  ( ! [X7] : (~v5_ordinal1(X7) | ~v1_funct_1(X7) | ~v1_relat_1(X7) | r2_hidden(sK10(X7),k9_xtuple_0(X7)) | v2_ordinal5(X7) | ~v1_ordinal1(k9_xtuple_0(X7)) | v2_ordinal5(X7) | ~v5_ordinal1(X7) | ~v1_funct_1(X7) | ~v1_relat_1(X7)) )),
  inference(resolution,[],[f1710,f659])).
fof(f1710,plain,(
  ( ! [X0,X1] : (~r2_hidden(X1,sK11(X0)) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0) | r2_hidden(X1,k9_xtuple_0(X0)) | v2_ordinal5(X0) | ~v1_ordinal1(k9_xtuple_0(X0))) )),
  inference(resolution,[],[f660,f686])).
fof(f622614,plain,(
  spl14_562 | spl14_563 | ~spl14_8 | ~spl14_26 | ~spl14_27 | ~spl14_28 | ~spl14_29 | ~spl14_30 | ~spl14_32 | ~spl14_33),
  inference(avatar_split_clause,[],[f622613,f887,f879,f863,f855,f847,f839,f831,f742,f8320,f8316])).
fof(f8320,plain,(
  spl14_563 <=> ! [X0] : (k8_ordinal2(sK7) = k12_ordinal2(X0,sK6) | ~v3_ordinal1(X0) | v3_ordinal1(sK1(sK6,X0,sK7)))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_563])])).
fof(f622613,plain,(
  ( ! [X0] : (k8_ordinal2(sK7) = k12_ordinal2(X0,sK6) | k6_numbers = sK6 | v3_ordinal1(sK1(sK6,X0,sK7)) | ~v3_ordinal1(X0)) ) | (~spl14_8 | ~spl14_26 | ~spl14_27 | ~spl14_28 | ~spl14_29 | ~spl14_30 | ~spl14_32 | ~spl14_33)),
  inference(subsumption_resolution,[],[f622612,f889])).
fof(f622612,plain,(
  ( ! [X0] : (k8_ordinal2(sK7) = k12_ordinal2(X0,sK6) | k6_numbers = sK6 | v3_ordinal1(sK1(sK6,X0,sK7)) | ~v3_ordinal1(X0) | ~v3_ordinal1(sK6)) ) | (~spl14_8 | ~spl14_26 | ~spl14_27 | ~spl14_28 | ~spl14_29 | ~spl14_30 | ~spl14_32)),
  inference(subsumption_resolution,[],[f622611,f881])).
fof(f622611,plain,(
  ( ! [X0] : (k8_ordinal2(sK7) = k12_ordinal2(X0,sK6) | k6_numbers = sK6 | v3_ordinal1(sK1(sK6,X0,sK7)) | ~v4_ordinal1(sK6) | ~v3_ordinal1(X0) | ~v3_ordinal1(sK6)) ) | (~spl14_8 | ~spl14_26 | ~spl14_27 | ~spl14_28 | ~spl14_29 | ~spl14_30)),
  inference(subsumption_resolution,[],[f622610,f849])).
fof(f622610,plain,(
  ( ! [X0] : (k8_ordinal2(sK7) = k12_ordinal2(X0,sK6) | k6_numbers = sK6 | v3_ordinal1(sK1(sK6,X0,sK7)) | ~v5_ordinal1(sK7) | ~v4_ordinal1(sK6) | ~v3_ordinal1(X0) | ~v3_ordinal1(sK6)) ) | (~spl14_8 | ~spl14_26 | ~spl14_27 | ~spl14_29 | ~spl14_30)),
  inference(subsumption_resolution,[],[f622609,f865])).
fof(f622609,plain,(
  ( ! [X0] : (k8_ordinal2(sK7) = k12_ordinal2(X0,sK6) | k6_numbers = sK6 | v3_ordinal1(sK1(sK6,X0,sK7)) | ~v1_relat_1(sK7) | ~v5_ordinal1(sK7) | ~v4_ordinal1(sK6) | ~v3_ordinal1(X0) | ~v3_ordinal1(sK6)) ) | (~spl14_8 | ~spl14_26 | ~spl14_27 | ~spl14_29)),
  inference(subsumption_resolution,[],[f622608,f857])).
fof(f622608,plain,(
  ( ! [X0] : (k8_ordinal2(sK7) = k12_ordinal2(X0,sK6) | k6_numbers = sK6 | v3_ordinal1(sK1(sK6,X0,sK7)) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | ~v5_ordinal1(sK7) | ~v4_ordinal1(sK6) | ~v3_ordinal1(X0) | ~v3_ordinal1(sK6)) ) | (~spl14_8 | ~spl14_26 | ~spl14_27)),
  inference(subsumption_resolution,[],[f622450,f841])).
fof(f622450,plain,(
  ( ! [X0] : (k8_ordinal2(sK7) = k12_ordinal2(X0,sK6) | k6_numbers = sK6 | v3_ordinal1(sK1(sK6,X0,sK7)) | ~v1_ordinal2(sK7) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | ~v5_ordinal1(sK7) | ~v4_ordinal1(sK6) | ~v3_ordinal1(X0) | ~v3_ordinal1(sK6)) ) | (~spl14_8 | ~spl14_26)),
  inference(superposition,[],[f2101,f833])).
fof(f2101,plain,(
  ( ! [X2,X1] : (k8_ordinal2(X2) = k12_ordinal2(X1,k9_xtuple_0(X2)) | k6_numbers = k9_xtuple_0(X2) | v3_ordinal1(sK1(k9_xtuple_0(X2),X1,X2)) | ~v1_ordinal2(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v5_ordinal1(X2) | ~v4_ordinal1(k9_xtuple_0(X2)) | ~v3_ordinal1(X1) | ~v3_ordinal1(k9_xtuple_0(X2))) ) | ~spl14_8),
  inference(forward_demodulation,[],[f699,f744])).
fof(f699,plain,(
  ( ! [X2,X1] : (k8_ordinal2(X2) = k12_ordinal2(X1,k9_xtuple_0(X2)) | v3_ordinal1(sK1(k9_xtuple_0(X2),X1,X2)) | ~v1_ordinal2(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v5_ordinal1(X2) | k1_xboole_0 = k9_xtuple_0(X2) | ~v4_ordinal1(k9_xtuple_0(X2)) | ~v3_ordinal1(X1) | ~v3_ordinal1(k9_xtuple_0(X2))) )),
  inference(equality_resolution,[],[f266])).
fof(f266,plain,(
  ( ! [X2,X0,X1] : (k8_ordinal2(X2) = k12_ordinal2(X1,X0) | v3_ordinal1(sK1(X0,X1,X2)) | k9_xtuple_0(X2) != X0 | ~v1_ordinal2(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v5_ordinal1(X2) | k1_xboole_0 = X0 | ~v4_ordinal1(X0) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f187])).
fof(f616838,plain,(
  sK6 != k9_xtuple_0(sK7) | ~v1_xboole_0(k9_xtuple_0(sK7)) | v1_xboole_0(sK6)),
  introduced(theory_tautology_sat_conflict,[])).
fof(f584568,plain,(
  spl14_567 | ~spl14_631 | ~spl14_13 | spl14_20 | ~spl14_563),
  inference(avatar_split_clause,[],[f584567,f8320,f801,f767,f9069,f8348])).
fof(f8348,plain,(
  spl14_567 <=> v3_ordinal1(sK1(sK6,k4_ordinal1,sK7))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_567])])).
fof(f584567,plain,(
  ~r1_ordinal2(k8_ordinal2(sK7),sK7) | v3_ordinal1(sK1(sK6,k4_ordinal1,sK7)) | (~spl14_13 | spl14_20 | ~spl14_563)),
  inference(subsumption_resolution,[],[f584153,f769])).
fof(f584153,plain,(
  ~r1_ordinal2(k8_ordinal2(sK7),sK7) | ~v3_ordinal1(k4_ordinal1) | v3_ordinal1(sK1(sK6,k4_ordinal1,sK7)) | (spl14_20 | ~spl14_563)),
  inference(superposition,[],[f803,f8321])).
fof(f8321,plain,(
  ( ! [X0] : (k8_ordinal2(sK7) = k12_ordinal2(X0,sK6) | ~v3_ordinal1(X0) | v3_ordinal1(sK1(sK6,X0,sK7))) ) | ~spl14_563),
  inference(avatar_component_clause,[],[f8320])).
fof(f567593,plain,(
  ~spl14_180 | spl14_19821 | ~spl14_19822 | ~spl14_19823 | ~spl14_12 | ~spl14_79 | ~spl14_142 | ~spl14_150 | ~spl14_172),
  inference(avatar_split_clause,[],[f567592,f2618,f2301,f2194,f1490,f762,f440112,f440108,f440104,f2700])).
fof(f567592,plain,(
  ~v3_ordinal1(sK3(np__0,k3_card_3(sK13))) | ~v4_ordinal5(sK3(np__0,k3_card_3(sK13))) | k2_ordinal5(np__0) = k3_card_3(sK13) | ~r2_hidden(np__0,k3_card_3(sK13)) | (~spl14_12 | ~spl14_79 | ~spl14_142 | ~spl14_150 | ~spl14_172)),
  inference(subsumption_resolution,[],[f567591,f2302])).
fof(f567591,plain,(
  ~v3_ordinal1(sK3(np__0,k3_card_3(sK13))) | ~v4_ordinal5(sK3(np__0,k3_card_3(sK13))) | k2_ordinal5(np__0) = k3_card_3(sK13) | ~r2_hidden(np__0,k3_card_3(sK13)) | ~v3_ordinal1(k3_card_3(sK13)) | (~spl14_12 | ~spl14_79 | ~spl14_142 | ~spl14_150 | ~spl14_172)),
  inference(subsumption_resolution,[],[f567590,f2620])).
fof(f567590,plain,(
  ~v3_ordinal1(sK3(np__0,k3_card_3(sK13))) | ~v4_ordinal5(sK3(np__0,k3_card_3(sK13))) | k2_ordinal5(np__0) = k3_card_3(sK13) | ~r2_hidden(np__0,k3_card_3(sK13)) | ~v4_ordinal5(k3_card_3(sK13)) | ~v3_ordinal1(k3_card_3(sK13)) | (~spl14_12 | ~spl14_79 | ~spl14_142 | ~spl14_150 | ~spl14_172)),
  inference(subsumption_resolution,[],[f440098,f1492])).
fof(f440098,plain,(
  ~v3_ordinal1(sK3(np__0,k3_card_3(sK13))) | ~v4_ordinal5(sK3(np__0,k3_card_3(sK13))) | k2_ordinal5(np__0) = k3_card_3(sK13) | ~r2_hidden(np__0,k3_card_3(sK13)) | ~v3_ordinal1(np__0) | ~v4_ordinal5(k3_card_3(sK13)) | ~v3_ordinal1(k3_card_3(sK13)) | (~spl14_12 | ~spl14_79 | ~spl14_142 | ~spl14_150 | ~spl14_172)),
  inference(duplicate_literal_removal,[],[f440091])).
fof(f440091,plain,(
  ~v3_ordinal1(sK3(np__0,k3_card_3(sK13))) | ~v4_ordinal5(sK3(np__0,k3_card_3(sK13))) | k2_ordinal5(np__0) = k3_card_3(sK13) | ~r2_hidden(np__0,k3_card_3(sK13)) | ~v3_ordinal1(np__0) | ~v3_ordinal1(sK3(np__0,k3_card_3(sK13))) | k2_ordinal5(np__0) = k3_card_3(sK13) | ~r2_hidden(np__0,k3_card_3(sK13)) | ~v4_ordinal5(k3_card_3(sK13)) | ~v3_ordinal1(k3_card_3(sK13)) | (~spl14_12 | ~spl14_79 | ~spl14_142 | ~spl14_150 | ~spl14_172)),
  inference(resolution,[],[f7246,f2697])).
fof(f2697,plain,(
  ( ! [X0] : (r1_ordinal1(np__1,sK3(np__0,X0)) | ~v3_ordinal1(sK3(np__0,X0)) | k2_ordinal5(np__0) = X0 | ~r2_hidden(np__0,X0) | ~v4_ordinal5(X0) | ~v3_ordinal1(X0)) ) | (~spl14_12 | ~spl14_79)),
  inference(subsumption_resolution,[],[f2654,f1492])).
fof(f2654,plain,(
  ( ! [X0] : (r1_ordinal1(np__1,sK3(np__0,X0)) | ~v3_ordinal1(sK3(np__0,X0)) | k2_ordinal5(np__0) = X0 | ~r2_hidden(np__0,X0) | ~v4_ordinal5(X0) | ~v3_ordinal1(X0) | ~v3_ordinal1(np__0)) ) | (~spl14_12 | ~spl14_79)),
  inference(resolution,[],[f1660,f281])).
fof(f281,plain,(
  ( ! [X0,X1] : (r2_hidden(X0,sK3(X0,X1)) | k2_ordinal5(X0) = X1 | ~r2_hidden(X0,X1) | ~v4_ordinal5(X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f195])).
fof(f1660,plain,(
  ( ! [X0] : (~r2_hidden(np__0,X0) | r1_ordinal1(np__1,X0) | ~v3_ordinal1(X0)) ) | (~spl14_12 | ~spl14_79)),
  inference(subsumption_resolution,[],[f1657,f1492])).
fof(f1657,plain,(
  ( ! [X0] : (r1_ordinal1(np__1,X0) | ~r2_hidden(np__0,X0) | ~v3_ordinal1(X0) | ~v3_ordinal1(np__0)) ) | ~spl14_12),
  inference(superposition,[],[f272,f764])).
fof(f272,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k1_ordinal1(X0),X1) | ~r2_hidden(X0,X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f190])).
fof(f7246,plain,(
  ( ! [X1] : (~r1_ordinal1(np__1,sK3(X1,k3_card_3(sK13))) | ~v3_ordinal1(sK3(X1,k3_card_3(sK13))) | ~v4_ordinal5(sK3(X1,k3_card_3(sK13))) | k3_card_3(sK13) = k2_ordinal5(X1) | ~r2_hidden(X1,k3_card_3(sK13)) | ~v3_ordinal1(X1)) ) | (~spl14_142 | ~spl14_150 | ~spl14_172)),
  inference(subsumption_resolution,[],[f7245,f2302])).
fof(f7245,plain,(
  ( ! [X1] : (~r1_ordinal1(np__1,sK3(X1,k3_card_3(sK13))) | ~v3_ordinal1(sK3(X1,k3_card_3(sK13))) | ~v4_ordinal5(sK3(X1,k3_card_3(sK13))) | k3_card_3(sK13) = k2_ordinal5(X1) | ~r2_hidden(X1,k3_card_3(sK13)) | ~v3_ordinal1(k3_card_3(sK13)) | ~v3_ordinal1(X1)) ) | (~spl14_142 | ~spl14_172)),
  inference(subsumption_resolution,[],[f7242,f2620])).
fof(f7242,plain,(
  ( ! [X1] : (~r1_ordinal1(np__1,sK3(X1,k3_card_3(sK13))) | ~v3_ordinal1(sK3(X1,k3_card_3(sK13))) | ~v4_ordinal5(sK3(X1,k3_card_3(sK13))) | k3_card_3(sK13) = k2_ordinal5(X1) | ~r2_hidden(X1,k3_card_3(sK13)) | ~v4_ordinal5(k3_card_3(sK13)) | ~v3_ordinal1(k3_card_3(sK13)) | ~v3_ordinal1(X1)) ) | ~spl14_142),
  inference(resolution,[],[f3606,f282])).
fof(f282,plain,(
  ( ! [X0,X1] : (~r1_ordinal1(X1,sK3(X0,X1)) | k2_ordinal5(X0) = X1 | ~r2_hidden(X0,X1) | ~v4_ordinal5(X1) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f195])).
fof(f3606,plain,(
  ( ! [X0] : (r1_ordinal1(k3_card_3(sK13),X0) | ~r1_ordinal1(np__1,X0) | ~v3_ordinal1(X0) | ~v4_ordinal5(X0)) ) | ~spl14_142),
  inference(trivial_inequality_removal,[],[f3605])).
fof(f3605,plain,(
  ( ! [X0] : (X0 != X0 | r1_ordinal1(k3_card_3(sK13),X0) | ~r1_ordinal1(np__1,X0) | ~v3_ordinal1(X0) | ~v4_ordinal5(X0)) ) | ~spl14_142),
  inference(duplicate_literal_removal,[],[f3596])).
fof(f3596,plain,(
  ( ! [X0] : (X0 != X0 | r1_ordinal1(k3_card_3(sK13),X0) | ~r1_ordinal1(np__1,X0) | ~v3_ordinal1(X0) | ~v4_ordinal5(X0) | ~v3_ordinal1(X0)) ) | ~spl14_142),
  inference(superposition,[],[f2195,f259])).
fof(f259,plain,(
  ( ! [X0] : (k12_ordinal2(k4_ordinal1,X0) = X0 | ~v4_ordinal5(X0) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f181])).
fof(f2195,plain,(
  ( ! [X0] : (k12_ordinal2(k4_ordinal1,X0) != X0 | r1_ordinal1(k3_card_3(sK13),X0) | ~r1_ordinal1(np__1,X0) | ~v3_ordinal1(X0)) ) | ~spl14_142),
  inference(avatar_component_clause,[],[f2194])).
fof(f567577,plain,(
  spl14_120 | spl14_134 | ~spl14_38 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_68),
  inference(avatar_split_clause,[],[f567576,f1394,f1268,f1263,f1258,f1253,f1248,f1067,f2064,f1937])).
fof(f1067,plain,(
  spl14_38 <=> ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | ~v1_relat_1(X0) | ~v1_funct_1(X0) | ~v5_ordinal1(X0) | ~v1_ordinal2(X0) | k4_ordinal1 != k9_xtuple_0(X0) | np__1 != k1_funct_1(X0,k6_numbers) | v3_ordinal1(sK5(X0)))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_38])])).
fof(f567576,plain,(
  k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | v3_ordinal1(sK5(sK13)) | (~spl14_38 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_68)),
  inference(subsumption_resolution,[],[f567575,f1250])).
fof(f567575,plain,(
  k4_ordinal1 != k9_xtuple_0(sK13) | k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | v3_ordinal1(sK5(sK13)) | (~spl14_38 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_68)),
  inference(subsumption_resolution,[],[f567574,f1255])).
fof(f567574,plain,(
  ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | v3_ordinal1(sK5(sK13)) | (~spl14_38 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_68)),
  inference(subsumption_resolution,[],[f567573,f1270])).
fof(f567573,plain,(
  ~v5_ordinal1(sK13) | ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | v3_ordinal1(sK5(sK13)) | (~spl14_38 | ~spl14_46 | ~spl14_47 | ~spl14_68)),
  inference(subsumption_resolution,[],[f567572,f1260])).
fof(f567572,plain,(
  ~v1_funct_1(sK13) | ~v5_ordinal1(sK13) | ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | v3_ordinal1(sK5(sK13)) | (~spl14_38 | ~spl14_47 | ~spl14_68)),
  inference(subsumption_resolution,[],[f16659,f1265])).
fof(f16659,plain,(
  ~v1_relat_1(sK13) | ~v1_funct_1(sK13) | ~v5_ordinal1(sK13) | ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | v3_ordinal1(sK5(sK13)) | (~spl14_38 | ~spl14_68)),
  inference(trivial_inequality_removal,[],[f16652])).
fof(f16652,plain,(
  np__1 != np__1 | ~v1_relat_1(sK13) | ~v1_funct_1(sK13) | ~v5_ordinal1(sK13) | ~v1_ordinal2(sK13) | k4_ordinal1 != k9_xtuple_0(sK13) | k3_card_3(sK13) = k12_ordinal2(k4_ordinal1,k3_card_3(sK13)) | v3_ordinal1(sK5(sK13)) | (~spl14_38 | ~spl14_68)),
  inference(superposition,[],[f1068,f1396])).
fof(f1068,plain,(
  ( ! [X0] : (np__1 != k1_funct_1(X0,k6_numbers) | ~v1_relat_1(X0) | ~v1_funct_1(X0) | ~v5_ordinal1(X0) | ~v1_ordinal2(X0) | k4_ordinal1 != k9_xtuple_0(X0) | k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | v3_ordinal1(sK5(X0))) ) | ~spl14_38),
  inference(avatar_component_clause,[],[f1067])).
fof(f346790,plain,(
  ~spl14_30 | ~spl14_29 | ~spl14_28 | ~spl14_27 | ~spl14_557 | spl14_14155 | ~spl14_688),
  inference(avatar_split_clause,[],[f346789,f10052,f276432,f8279,f839,f847,f855,f863])).
fof(f276432,plain,(
  spl14_14155 <=> v1_xboole_0(k9_xtuple_0(sK7))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_14155])])).
fof(f346789,plain,(
  v1_xboole_0(k9_xtuple_0(sK7)) | ~v2_ordinal5(sK7) | ~v1_ordinal2(sK7) | ~v5_ordinal1(sK7) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | ~spl14_688),
  inference(subsumption_resolution,[],[f293596,f292])).
fof(f292,plain,(
  ( ! [X0] : (v3_ordinal1(k3_card_3(X0)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0] : (v3_ordinal1(k3_card_3(X0)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))),
  inference(flattening,[],[f142])).
fof(f142,plain,(
  ! [X0] : (v3_ordinal1(k3_card_3(X0)) | (~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)))),
  inference(ennf_transformation,[],[f28])).
fof(f28,axiom,(
  ! [X0] : ((v1_ordinal2(X0) & v5_ordinal1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => v3_ordinal1(k3_card_3(X0)))),
  file('minsubpreds65/t36_ordinal5',fc1_card_5)).
fof(f293596,plain,(
  ~v3_ordinal1(k3_card_3(sK7)) | v1_xboole_0(k9_xtuple_0(sK7)) | ~v2_ordinal5(sK7) | ~v1_ordinal2(sK7) | ~v5_ordinal1(sK7) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | ~spl14_688),
  inference(resolution,[],[f10053,f291])).
fof(f10053,plain,(
  ( ! [X1] : (~r1_ordinal2(X1,sK7) | ~v3_ordinal1(X1)) ) | ~spl14_688),
  inference(avatar_component_clause,[],[f10052])).
fof(f283194,plain,(
  spl14_1255 | ~spl14_13 | ~spl14_273 | ~spl14_277),
  inference(avatar_split_clause,[],[f283193,f4187,f4161,f767,f18956])).
fof(f4187,plain,(
  spl14_277 <=> k1_funct_1(sK13,k1_ordinal1(sK4(sK13))) = k12_ordinal2(k4_ordinal1,k1_ordinal5(k4_ordinal1,sK4(sK13)))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_277])])).
fof(f283193,plain,(
  k1_funct_1(sK13,k1_ordinal1(sK4(sK13))) = k1_ordinal5(k4_ordinal1,k1_ordinal1(sK4(sK13))) | (~spl14_13 | ~spl14_273 | ~spl14_277)),
  inference(subsumption_resolution,[],[f283192,f769])).
fof(f283192,plain,(
  k1_funct_1(sK13,k1_ordinal1(sK4(sK13))) = k1_ordinal5(k4_ordinal1,k1_ordinal1(sK4(sK13))) | ~v3_ordinal1(k4_ordinal1) | (~spl14_273 | ~spl14_277)),
  inference(subsumption_resolution,[],[f282842,f4163])).
fof(f4163,plain,(
  v3_ordinal1(sK4(sK13)) | ~spl14_273),
  inference(avatar_component_clause,[],[f4161])).
fof(f282842,plain,(
  k1_funct_1(sK13,k1_ordinal1(sK4(sK13))) = k1_ordinal5(k4_ordinal1,k1_ordinal1(sK4(sK13))) | ~v3_ordinal1(sK4(sK13)) | ~v3_ordinal1(k4_ordinal1) | ~spl14_277),
  inference(superposition,[],[f265,f4189])).
fof(f4189,plain,(
  k1_funct_1(sK13,k1_ordinal1(sK4(sK13))) = k12_ordinal2(k4_ordinal1,k1_ordinal5(k4_ordinal1,sK4(sK13))) | ~spl14_277),
  inference(avatar_component_clause,[],[f4187])).
fof(f265,plain,(
  ( ! [X0,X1] : (k1_ordinal5(X0,k1_ordinal1(X1)) = k12_ordinal2(X0,k1_ordinal5(X0,X1)) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0] : (! [X1] : (k1_ordinal5(X0,k1_ordinal1(X1)) = k12_ordinal2(X0,k1_ordinal5(X0,X1)) | ~v3_ordinal1(X1)) | ~v3_ordinal1(X0))),
  inference(ennf_transformation,[],[f48])).
fof(f48,axiom,(
  ! [X0] : (v3_ordinal1(X0) => ! [X1] : (v3_ordinal1(X1) => k1_ordinal5(X0,k1_ordinal1(X1)) = k12_ordinal2(X0,k1_ordinal5(X0,X1))))),
  file('minsubpreds65/t36_ordinal5',t14_ordinal5)).
fof(f280269,plain,(
  ~spl14_30 | ~spl14_29 | ~spl14_28 | spl14_557 | spl14_1546),
  inference(avatar_split_clause,[],[f280268,f22785,f8279,f847,f855,f863])).
fof(f280268,plain,(
  ~v5_ordinal1(sK7) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | (spl14_557 | spl14_1546)),
  inference(subsumption_resolution,[],[f280267,f8280])).
fof(f280267,plain,(
  v2_ordinal5(sK7) | ~v5_ordinal1(sK7) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | spl14_1546),
  inference(resolution,[],[f22787,f658])).
fof(f22787,plain,(
  ~v3_ordinal1(sK11(sK7)) | spl14_1546),
  inference(avatar_component_clause,[],[f22785])).
fof(f280140,plain,(
  ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | spl14_229 | ~spl14_256 | ~spl14_257),
  inference(avatar_contradiction_clause,[],[f280129])).
fof(f280129,plain,(
  $false | (~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | spl14_229 | ~spl14_256 | ~spl14_257)),
  inference(unit_resulting_resolution,[],[f1270,f1265,f1260,f1255,f3755,f3751,f3755,f3751,f3448,f288])).
fof(f3448,plain,(
  k2_ordinal5(k6_numbers) != k8_ordinal2(sK13) | spl14_229),
  inference(avatar_component_clause,[],[f3446])).
fof(f3446,plain,(
  spl14_229 <=> k2_ordinal5(k6_numbers) = k8_ordinal2(sK13)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_229])])).
fof(f3751,plain,(
  r1_ordinal2(k2_ordinal5(k6_numbers),sK13) | ~spl14_256),
  inference(avatar_component_clause,[],[f3750])).
fof(f3755,plain,(
  v3_ordinal1(k2_ordinal5(k6_numbers)) | ~spl14_257),
  inference(avatar_component_clause,[],[f3754])).
fof(f3754,plain,(
  spl14_257 <=> v3_ordinal1(k2_ordinal5(k6_numbers))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_257])])).
fof(f268833,plain,(
  ~spl14_273 | ~spl14_276 | spl14_277 | ~spl14_135),
  inference(avatar_split_clause,[],[f51637,f2084,f4187,f4183,f4161])).
fof(f4183,plain,(
  spl14_276 <=> r2_hidden(k1_ordinal1(sK4(sK13)),k4_ordinal1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_276])])).
fof(f2084,plain,(
  spl14_135 <=> k1_funct_1(sK13,sK4(sK13)) = k1_ordinal5(k4_ordinal1,sK4(sK13))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_135])])).
fof(f51637,plain,(
  k1_funct_1(sK13,k1_ordinal1(sK4(sK13))) = k12_ordinal2(k4_ordinal1,k1_ordinal5(k4_ordinal1,sK4(sK13))) | ~r2_hidden(k1_ordinal1(sK4(sK13)),k4_ordinal1) | ~v3_ordinal1(sK4(sK13)) | ~spl14_135),
  inference(superposition,[],[f695,f2086])).
fof(f2086,plain,(
  k1_funct_1(sK13,sK4(sK13)) = k1_ordinal5(k4_ordinal1,sK4(sK13)) | ~spl14_135),
  inference(avatar_component_clause,[],[f2084])).
fof(f256572,plain,(
  ~spl14_98 | spl14_14 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | spl14_4174),
  inference(avatar_split_clause,[],[f256571,f87713,f1268,f1263,f1258,f1253,f1248,f772,f1718])).
fof(f256571,plain,(
  ~v2_ordinal5(sK13) | (spl14_14 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | spl14_4174)),
  inference(subsumption_resolution,[],[f256570,f774])).
fof(f256570,plain,(
  v1_xboole_0(k4_ordinal1) | ~v2_ordinal5(sK13) | (~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | spl14_4174)),
  inference(forward_demodulation,[],[f256569,f1250])).
fof(f256569,plain,(
  v1_xboole_0(k9_xtuple_0(sK13)) | ~v2_ordinal5(sK13) | (~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | spl14_4174)),
  inference(subsumption_resolution,[],[f256568,f1265])).
fof(f256568,plain,(
  v1_xboole_0(k9_xtuple_0(sK13)) | ~v2_ordinal5(sK13) | ~v1_relat_1(sK13) | (~spl14_45 | ~spl14_46 | ~spl14_48 | spl14_4174)),
  inference(subsumption_resolution,[],[f256567,f1260])).
fof(f256567,plain,(
  v1_xboole_0(k9_xtuple_0(sK13)) | ~v2_ordinal5(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | (~spl14_45 | ~spl14_48 | spl14_4174)),
  inference(subsumption_resolution,[],[f256566,f1270])).
fof(f256566,plain,(
  v1_xboole_0(k9_xtuple_0(sK13)) | ~v2_ordinal5(sK13) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | (~spl14_45 | spl14_4174)),
  inference(subsumption_resolution,[],[f87869,f1255])).
fof(f87869,plain,(
  v1_xboole_0(k9_xtuple_0(sK13)) | ~v2_ordinal5(sK13) | ~v1_ordinal2(sK13) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | spl14_4174),
  inference(resolution,[],[f87715,f291])).
fof(f87715,plain,(
  ~r1_ordinal2(k3_card_3(sK13),sK13) | spl14_4174),
  inference(avatar_component_clause,[],[f87713])).
fof(f64213,plain,(
  spl14_98 | spl14_409 | ~spl14_44 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_62),
  inference(avatar_split_clause,[],[f64212,f1356,f1268,f1263,f1258,f1248,f6260,f1718])).
fof(f64212,plain,(
  r2_hidden(sK10(sK13),k4_ordinal1) | v2_ordinal5(sK13) | (~spl14_44 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_62)),
  inference(subsumption_resolution,[],[f64211,f1358])).
fof(f64211,plain,(
  r2_hidden(sK10(sK13),k4_ordinal1) | v2_ordinal5(sK13) | ~v1_ordinal1(k4_ordinal1) | (~spl14_44 | ~spl14_46 | ~spl14_47 | ~spl14_48)),
  inference(subsumption_resolution,[],[f64210,f1270])).
fof(f64210,plain,(
  r2_hidden(sK10(sK13),k4_ordinal1) | ~v5_ordinal1(sK13) | v2_ordinal5(sK13) | ~v1_ordinal1(k4_ordinal1) | (~spl14_44 | ~spl14_46 | ~spl14_47)),
  inference(subsumption_resolution,[],[f64209,f1265])).
fof(f64209,plain,(
  r2_hidden(sK10(sK13),k4_ordinal1) | ~v1_relat_1(sK13) | ~v5_ordinal1(sK13) | v2_ordinal5(sK13) | ~v1_ordinal1(k4_ordinal1) | (~spl14_44 | ~spl14_46)),
  inference(subsumption_resolution,[],[f13411,f1260])).
fof(f13411,plain,(
  r2_hidden(sK10(sK13),k4_ordinal1) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | ~v5_ordinal1(sK13) | v2_ordinal5(sK13) | ~v1_ordinal1(k4_ordinal1) | ~spl14_44),
  inference(superposition,[],[f3884,f1250])).
fof(f50373,plain,(
  ~spl14_24 | ~spl14_23 | ~spl14_22 | ~spl14_13 | spl14_19 | ~spl14_86),
  inference(avatar_split_clause,[],[f50372,f1572,f797,f767,f809,f814,f819])).
fof(f819,plain,(
  spl14_24 <=> v3_ordinal1(sK8)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_24])])).
fof(f814,plain,(
  spl14_23 <=> v3_ordinal1(sK9)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_23])])).
fof(f797,plain,(
  spl14_19 <=> r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_19])])).
fof(f50372,plain,(
  ~r2_hidden(sK8,sK9) | ~v3_ordinal1(sK9) | ~v3_ordinal1(sK8) | (~spl14_13 | spl14_19 | ~spl14_86)),
  inference(subsumption_resolution,[],[f50371,f769])).
fof(f50371,plain,(
  ~r2_hidden(sK8,sK9) | ~v3_ordinal1(sK9) | ~v3_ordinal1(sK8) | ~v3_ordinal1(k4_ordinal1) | (spl14_19 | ~spl14_86)),
  inference(subsumption_resolution,[],[f42676,f1574])).
fof(f42676,plain,(
  ~r2_hidden(sK8,sK9) | ~r2_hidden(np__1,k4_ordinal1) | ~v3_ordinal1(sK9) | ~v3_ordinal1(sK8) | ~v3_ordinal1(k4_ordinal1) | spl14_19),
  inference(resolution,[],[f799,f274])).
fof(f799,plain,(
  ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | spl14_19),
  inference(avatar_component_clause,[],[f797])).
fof(f31201,plain,(
  spl14_1 | ~spl14_3 | ~spl14_8 | ~spl14_13 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_132 | spl14_139 | ~spl14_145 | ~spl14_229),
  inference(avatar_contradiction_clause,[],[f31200])).
fof(f31200,plain,(
  $false | (spl14_1 | ~spl14_3 | ~spl14_8 | ~spl14_13 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_132 | spl14_139 | ~spl14_145 | ~spl14_229)),
  inference(subsumption_resolution,[],[f31199,f769])).
fof(f31199,plain,(
  ~v3_ordinal1(k4_ordinal1) | (spl14_1 | ~spl14_3 | ~spl14_8 | ~spl14_13 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_132 | spl14_139 | ~spl14_145 | ~spl14_229)),
  inference(subsumption_resolution,[],[f31197,f709])).
fof(f709,plain,(
  k2_ordinal5(k6_numbers) != k1_ordinal5(k4_ordinal1,k4_ordinal1) | spl14_1),
  inference(avatar_component_clause,[],[f707])).
fof(f707,plain,(
  spl14_1 <=> k2_ordinal5(k6_numbers) = k1_ordinal5(k4_ordinal1,k4_ordinal1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_1])])).
fof(f31197,plain,(
  k2_ordinal5(k6_numbers) = k1_ordinal5(k4_ordinal1,k4_ordinal1) | ~v3_ordinal1(k4_ordinal1) | (~spl14_3 | ~spl14_8 | ~spl14_13 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_132 | spl14_139 | ~spl14_145 | ~spl14_229)),
  inference(equality_resolution,[],[f31187])).
fof(f31187,plain,(
  ( ! [X0] : (k1_ordinal5(X0,sK2(k4_ordinal1,X0,sK13)) != k1_ordinal5(k4_ordinal1,sK2(k4_ordinal1,X0,sK13)) | k2_ordinal5(k6_numbers) = k1_ordinal5(X0,k4_ordinal1) | ~v3_ordinal1(X0)) ) | (~spl14_3 | ~spl14_8 | ~spl14_13 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_132 | spl14_139 | ~spl14_145 | ~spl14_229)),
  inference(subsumption_resolution,[],[f29179,f28745])).
fof(f28745,plain,(
  ( ! [X1] : (k2_ordinal5(k6_numbers) = k1_ordinal5(X1,k4_ordinal1) | ~v3_ordinal1(X1) | v7_ordinal1(sK2(k4_ordinal1,X1,sK13))) ) | (~spl14_145 | ~spl14_229)),
  inference(backward_demodulation,[],[f3635,f3447])).
fof(f3447,plain,(
  k2_ordinal5(k6_numbers) = k8_ordinal2(sK13) | ~spl14_229),
  inference(avatar_component_clause,[],[f3446])).
fof(f29179,plain,(
  ( ! [X0] : (k1_ordinal5(X0,sK2(k4_ordinal1,X0,sK13)) != k1_ordinal5(k4_ordinal1,sK2(k4_ordinal1,X0,sK13)) | k2_ordinal5(k6_numbers) = k1_ordinal5(X0,k4_ordinal1) | ~v3_ordinal1(X0) | ~v7_ordinal1(sK2(k4_ordinal1,X0,sK13))) ) | (~spl14_3 | ~spl14_8 | ~spl14_13 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_132 | spl14_139 | ~spl14_229)),
  inference(superposition,[],[f28737,f2049])).
fof(f28737,plain,(
  ( ! [X0] : (k1_funct_1(sK13,sK2(k4_ordinal1,X0,sK13)) != k1_ordinal5(X0,sK2(k4_ordinal1,X0,sK13)) | k2_ordinal5(k6_numbers) = k1_ordinal5(X0,k4_ordinal1) | ~v3_ordinal1(X0)) ) | (~spl14_3 | ~spl14_8 | ~spl14_13 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | spl14_139 | ~spl14_229)),
  inference(backward_demodulation,[],[f2272,f3447])).
fof(f28505,plain,(
  spl14_120 | ~spl14_13 | ~spl14_195),
  inference(avatar_split_clause,[],[f28504,f2904,f767,f1937])).
fof(f28504,plain,(
  v3_ordinal1(sK5(sK13)) | (~spl14_13 | ~spl14_195)),
  inference(subsumption_resolution,[],[f24291,f769])).
fof(f24291,plain,(
  v3_ordinal1(sK5(sK13)) | ~v3_ordinal1(k4_ordinal1) | ~spl14_195),
  inference(resolution,[],[f2906,f276])).
fof(f2906,plain,(
  m1_subset_1(sK5(sK13),k4_ordinal1) | ~spl14_195),
  inference(avatar_component_clause,[],[f2904])).
fof(f23132,plain,(
  ~spl14_69 | spl14_257),
  inference(avatar_contradiction_clause,[],[f23131])).
fof(f23131,plain,(
  $false | (~spl14_69 | spl14_257)),
  inference(subsumption_resolution,[],[f23126,f1404])).
fof(f1404,plain,(
  v3_ordinal1(k6_numbers) | ~spl14_69),
  inference(avatar_component_clause,[],[f1402])).
fof(f1402,plain,(
  spl14_69 <=> v3_ordinal1(k6_numbers)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_69])])).
fof(f23126,plain,(
  ~v3_ordinal1(k6_numbers) | spl14_257),
  inference(resolution,[],[f3756,f257])).
fof(f3756,plain,(
  ~v3_ordinal1(k2_ordinal5(k6_numbers)) | spl14_257),
  inference(avatar_component_clause,[],[f3754])).
fof(f15450,plain,(
  ~spl14_625 | ~spl14_567 | ~spl14_33 | ~spl14_32 | ~spl14_30 | ~spl14_29 | spl14_562 | spl14_580 | ~spl14_8 | ~spl14_13 | ~spl14_25 | ~spl14_26 | ~spl14_27 | ~spl14_28),
  inference(avatar_split_clause,[],[f15449,f847,f839,f831,f824,f767,f742,f8480,f8316,f855,f863,f879,f887,f8348,f8979])).
fof(f15449,plain,(
  k12_ordinal2(k4_ordinal1,sK6) = k8_ordinal2(sK7) | k6_numbers = sK6 | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | ~v4_ordinal1(sK6) | ~v3_ordinal1(sK6) | ~v3_ordinal1(sK1(sK6,k4_ordinal1,sK7)) | ~r2_hidden(sK1(sK6,k4_ordinal1,sK7),sK6) | (~spl14_8 | ~spl14_13 | ~spl14_25 | ~spl14_26 | ~spl14_27 | ~spl14_28)),
  inference(subsumption_resolution,[],[f9863,f849])).
fof(f9863,plain,(
  k12_ordinal2(k4_ordinal1,sK6) = k8_ordinal2(sK7) | k6_numbers = sK6 | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | ~v5_ordinal1(sK7) | ~v4_ordinal1(sK6) | ~v3_ordinal1(sK6) | ~v3_ordinal1(sK1(sK6,k4_ordinal1,sK7)) | ~r2_hidden(sK1(sK6,k4_ordinal1,sK7),sK6) | (~spl14_8 | ~spl14_13 | ~spl14_25 | ~spl14_26 | ~spl14_27)),
  inference(subsumption_resolution,[],[f9853,f841])).
fof(f9853,plain,(
  k12_ordinal2(k4_ordinal1,sK6) = k8_ordinal2(sK7) | k6_numbers = sK6 | ~v1_ordinal2(sK7) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | ~v5_ordinal1(sK7) | ~v4_ordinal1(sK6) | ~v3_ordinal1(sK6) | ~v3_ordinal1(sK1(sK6,k4_ordinal1,sK7)) | ~r2_hidden(sK1(sK6,k4_ordinal1,sK7),sK6) | (~spl14_8 | ~spl14_13 | ~spl14_25 | ~spl14_26)),
  inference(trivial_inequality_removal,[],[f9846])).
fof(f9846,plain,(
  k1_funct_1(sK7,sK1(sK6,k4_ordinal1,sK7)) != k1_funct_1(sK7,sK1(sK6,k4_ordinal1,sK7)) | k12_ordinal2(k4_ordinal1,sK6) = k8_ordinal2(sK7) | k6_numbers = sK6 | ~v1_ordinal2(sK7) | ~v1_funct_1(sK7) | ~v1_relat_1(sK7) | ~v5_ordinal1(sK7) | ~v4_ordinal1(sK6) | ~v3_ordinal1(sK6) | ~v3_ordinal1(sK1(sK6,k4_ordinal1,sK7)) | ~r2_hidden(sK1(sK6,k4_ordinal1,sK7),sK6) | (~spl14_8 | ~spl14_13 | ~spl14_25 | ~spl14_26)),
  inference(superposition,[],[f8714,f833])).
fof(f8714,plain,(
  ( ! [X47] : (k1_funct_1(sK7,sK1(k9_xtuple_0(X47),k4_ordinal1,X47)) != k1_funct_1(X47,sK1(k9_xtuple_0(X47),k4_ordinal1,X47)) | k8_ordinal2(X47) = k12_ordinal2(k4_ordinal1,k9_xtuple_0(X47)) | k6_numbers = k9_xtuple_0(X47) | ~v1_ordinal2(X47) | ~v1_funct_1(X47) | ~v1_relat_1(X47) | ~v5_ordinal1(X47) | ~v4_ordinal1(k9_xtuple_0(X47)) | ~v3_ordinal1(k9_xtuple_0(X47)) | ~v3_ordinal1(sK1(k9_xtuple_0(X47),k4_ordinal1,X47)) | ~r2_hidden(sK1(k9_xtuple_0(X47),k4_ordinal1,X47),sK6)) ) | (~spl14_8 | ~spl14_13 | ~spl14_25)),
  inference(subsumption_resolution,[],[f8548,f769])).
fof(f8548,plain,(
  ( ! [X47] : (k1_funct_1(sK7,sK1(k9_xtuple_0(X47),k4_ordinal1,X47)) != k1_funct_1(X47,sK1(k9_xtuple_0(X47),k4_ordinal1,X47)) | k8_ordinal2(X47) = k12_ordinal2(k4_ordinal1,k9_xtuple_0(X47)) | k6_numbers = k9_xtuple_0(X47) | ~v1_ordinal2(X47) | ~v1_funct_1(X47) | ~v1_relat_1(X47) | ~v5_ordinal1(X47) | ~v4_ordinal1(k9_xtuple_0(X47)) | ~v3_ordinal1(k4_ordinal1) | ~v3_ordinal1(k9_xtuple_0(X47)) | ~v3_ordinal1(sK1(k9_xtuple_0(X47),k4_ordinal1,X47)) | ~r2_hidden(sK1(k9_xtuple_0(X47),k4_ordinal1,X47),sK6)) ) | (~spl14_8 | ~spl14_25)),
  inference(superposition,[],[f2253,f825])).
fof(f15317,plain,(
  ~spl14_23 | spl14_24 | ~spl14_75),
  inference(avatar_split_clause,[],[f8726,f1454,f819,f814])).
fof(f8726,plain,(
  v3_ordinal1(sK8) | ~v3_ordinal1(sK9) | ~spl14_75),
  inference(resolution,[],[f1456,f276])).
fof(f1456,plain,(
  m1_subset_1(sK8,sK9) | ~spl14_75),
  inference(avatar_component_clause,[],[f1454])).
fof(f9366,plain,(
  ~spl14_625 | ~spl14_567 | spl14_580 | ~spl14_13 | ~spl14_25 | ~spl14_577),
  inference(avatar_split_clause,[],[f9365,f8456,f824,f767,f8480,f8348,f8979])).
fof(f9365,plain,(
  k12_ordinal2(k4_ordinal1,sK6) = k8_ordinal2(sK7) | ~v3_ordinal1(sK1(sK6,k4_ordinal1,sK7)) | ~r2_hidden(sK1(sK6,k4_ordinal1,sK7),sK6) | (~spl14_13 | ~spl14_25 | ~spl14_577)),
  inference(subsumption_resolution,[],[f9364,f769])).
fof(f9364,plain,(
  ~v3_ordinal1(k4_ordinal1) | k12_ordinal2(k4_ordinal1,sK6) = k8_ordinal2(sK7) | ~v3_ordinal1(sK1(sK6,k4_ordinal1,sK7)) | ~r2_hidden(sK1(sK6,k4_ordinal1,sK7),sK6) | (~spl14_25 | ~spl14_577)),
  inference(trivial_inequality_removal,[],[f9362])).
fof(f9362,plain,(
  k1_funct_1(sK7,sK1(sK6,k4_ordinal1,sK7)) != k1_funct_1(sK7,sK1(sK6,k4_ordinal1,sK7)) | ~v3_ordinal1(k4_ordinal1) | k12_ordinal2(k4_ordinal1,sK6) = k8_ordinal2(sK7) | ~v3_ordinal1(sK1(sK6,k4_ordinal1,sK7)) | ~r2_hidden(sK1(sK6,k4_ordinal1,sK7),sK6) | (~spl14_25 | ~spl14_577)),
  inference(superposition,[],[f8457,f825])).
fof(f8457,plain,(
  ( ! [X13] : (k1_funct_1(sK7,sK1(sK6,X13,sK7)) != k12_ordinal2(X13,sK1(sK6,X13,sK7)) | ~v3_ordinal1(X13) | k8_ordinal2(sK7) = k12_ordinal2(X13,sK6)) ) | ~spl14_577),
  inference(avatar_component_clause,[],[f8456])).
fof(f8488,plain,(
  k6_numbers != sK6 | k6_numbers != k1_xboole_0 | k1_xboole_0 != o_0_0_xboole_0 | v1_xboole_0(sK6) | ~v1_xboole_0(o_0_0_xboole_0)),
  introduced(theory_tautology_sat_conflict,[])).
fof(f6616,plain,(
  spl14_200 | ~spl14_409),
  inference(avatar_split_clause,[],[f6604,f6260,f3062])).
fof(f6604,plain,(
  v7_ordinal1(sK10(sK13)) | ~spl14_409),
  inference(resolution,[],[f6262,f673])).
fof(f5713,plain,(
  spl14_98 | spl14_149 | ~spl14_44 | ~spl14_46 | ~spl14_47 | ~spl14_48),
  inference(avatar_split_clause,[],[f3347,f1268,f1263,f1258,f1248,f2294,f1718])).
fof(f3347,plain,(
  m1_subset_1(sK11(sK13),k4_ordinal1) | v2_ordinal5(sK13) | (~spl14_44 | ~spl14_46 | ~spl14_47 | ~spl14_48)),
  inference(subsumption_resolution,[],[f3346,f1265])).
fof(f3346,plain,(
  m1_subset_1(sK11(sK13),k4_ordinal1) | ~v1_relat_1(sK13) | v2_ordinal5(sK13) | (~spl14_44 | ~spl14_46 | ~spl14_48)),
  inference(subsumption_resolution,[],[f3345,f1260])).
fof(f3345,plain,(
  m1_subset_1(sK11(sK13),k4_ordinal1) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | v2_ordinal5(sK13) | (~spl14_44 | ~spl14_48)),
  inference(subsumption_resolution,[],[f3344,f1270])).
fof(f3344,plain,(
  m1_subset_1(sK11(sK13),k4_ordinal1) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | v2_ordinal5(sK13) | ~spl14_44),
  inference(superposition,[],[f1712,f1250])).
fof(f1712,plain,(
  ( ! [X3] : (m1_subset_1(sK11(X3),k9_xtuple_0(X3)) | ~v5_ordinal1(X3) | ~v1_funct_1(X3) | ~v1_relat_1(X3) | v2_ordinal5(X3)) )),
  inference(resolution,[],[f660,f675])).
fof(f5712,plain,(
  ~spl14_98 | spl14_14 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_150 | ~spl14_231),
  inference(avatar_split_clause,[],[f5711,f3475,f2301,f1268,f1263,f1258,f1253,f1248,f772,f1718])).
fof(f5711,plain,(
  ~v2_ordinal5(sK13) | (spl14_14 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_150 | ~spl14_231)),
  inference(subsumption_resolution,[],[f5710,f774])).
fof(f5710,plain,(
  v1_xboole_0(k4_ordinal1) | ~v2_ordinal5(sK13) | (~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_150 | ~spl14_231)),
  inference(forward_demodulation,[],[f5709,f1250])).
fof(f5709,plain,(
  v1_xboole_0(k9_xtuple_0(sK13)) | ~v2_ordinal5(sK13) | (~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_150 | ~spl14_231)),
  inference(subsumption_resolution,[],[f5708,f1265])).
fof(f5708,plain,(
  v1_xboole_0(k9_xtuple_0(sK13)) | ~v2_ordinal5(sK13) | ~v1_relat_1(sK13) | (~spl14_45 | ~spl14_46 | ~spl14_48 | ~spl14_150 | ~spl14_231)),
  inference(subsumption_resolution,[],[f5707,f1260])).
fof(f5707,plain,(
  v1_xboole_0(k9_xtuple_0(sK13)) | ~v2_ordinal5(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | (~spl14_45 | ~spl14_48 | ~spl14_150 | ~spl14_231)),
  inference(subsumption_resolution,[],[f5706,f1270])).
fof(f5706,plain,(
  v1_xboole_0(k9_xtuple_0(sK13)) | ~v2_ordinal5(sK13) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | (~spl14_45 | ~spl14_150 | ~spl14_231)),
  inference(subsumption_resolution,[],[f5705,f1255])).
fof(f5705,plain,(
  v1_xboole_0(k9_xtuple_0(sK13)) | ~v2_ordinal5(sK13) | ~v1_ordinal2(sK13) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | (~spl14_150 | ~spl14_231)),
  inference(subsumption_resolution,[],[f4025,f2302])).
fof(f4025,plain,(
  ~v3_ordinal1(k3_card_3(sK13)) | v1_xboole_0(k9_xtuple_0(sK13)) | ~v2_ordinal5(sK13) | ~v1_ordinal2(sK13) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | ~spl14_231),
  inference(resolution,[],[f3476,f291])).
fof(f3476,plain,(
  ( ! [X1] : (~r1_ordinal2(X1,sK13) | ~v3_ordinal1(X1)) ) | ~spl14_231),
  inference(avatar_component_clause,[],[f3475])).
fof(f5353,plain,(
  ~spl14_131 | spl14_14 | spl14_276),
  inference(avatar_split_clause,[],[f5352,f4183,f772,f2044])).
fof(f5352,plain,(
  ~v7_ordinal1(sK4(sK13)) | (spl14_14 | spl14_276)),
  inference(subsumption_resolution,[],[f5324,f1563])).
fof(f5324,plain,(
  ~r2_hidden(k4_card_1(sK4(sK13)),k4_ordinal1) | ~v7_ordinal1(sK4(sK13)) | spl14_276),
  inference(superposition,[],[f4185,f247])).
fof(f4185,plain,(
  ~r2_hidden(k1_ordinal1(sK4(sK13)),k4_ordinal1) | spl14_276),
  inference(avatar_component_clause,[],[f4183])).
fof(f2918,plain,(
  ~spl14_3 | ~spl14_13 | ~spl14_120 | ~spl14_128 | spl14_137),
  inference(avatar_contradiction_clause,[],[f2917])).
fof(f2917,plain,(
  $false | (~spl14_3 | ~spl14_13 | ~spl14_120 | ~spl14_128 | spl14_137)),
  inference(subsumption_resolution,[],[f2916,f769])).
fof(f2916,plain,(
  ~v3_ordinal1(k4_ordinal1) | (~spl14_3 | ~spl14_120 | ~spl14_128 | spl14_137)),
  inference(subsumption_resolution,[],[f2915,f719])).
fof(f2915,plain,(
  ~v4_ordinal1(k4_ordinal1) | ~v3_ordinal1(k4_ordinal1) | (~spl14_120 | ~spl14_128 | spl14_137)),
  inference(subsumption_resolution,[],[f2914,f1939])).
fof(f2914,plain,(
  ~v3_ordinal1(sK5(sK13)) | ~v4_ordinal1(k4_ordinal1) | ~v3_ordinal1(k4_ordinal1) | (~spl14_128 | spl14_137)),
  inference(subsumption_resolution,[],[f2909,f2000])).
fof(f2909,plain,(
  ~r2_hidden(sK5(sK13),k4_ordinal1) | ~v3_ordinal1(sK5(sK13)) | ~v4_ordinal1(k4_ordinal1) | ~v3_ordinal1(k4_ordinal1) | spl14_137),
  inference(resolution,[],[f2115,f261])).
fof(f261,plain,(
  ( ! [X2,X0] : (r2_hidden(k1_ordinal1(X2),X0) | ~r2_hidden(X2,X0) | ~v3_ordinal1(X2) | ~v4_ordinal1(X0) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0] : (((v4_ordinal1(X0) | (~r2_hidden(k1_ordinal1(sK0(X0)),X0) & r2_hidden(sK0(X0),X0) & v3_ordinal1(sK0(X0)))) & (! [X2] : (r2_hidden(k1_ordinal1(X2),X0) | ~r2_hidden(X2,X0) | ~v3_ordinal1(X2)) | ~v4_ordinal1(X0))) | ~v3_ordinal1(X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f183,f184])).
fof(f184,plain,(
  ! [X0] : (? [X1] : (~r2_hidden(k1_ordinal1(X1),X0) & r2_hidden(X1,X0) & v3_ordinal1(X1)) => (~r2_hidden(k1_ordinal1(sK0(X0)),X0) & r2_hidden(sK0(X0),X0) & v3_ordinal1(sK0(X0))))),
  introduced(choice_axiom,[])).
fof(f183,plain,(
  ! [X0] : (((v4_ordinal1(X0) | ? [X1] : (~r2_hidden(k1_ordinal1(X1),X0) & r2_hidden(X1,X0) & v3_ordinal1(X1))) & (! [X2] : (r2_hidden(k1_ordinal1(X2),X0) | ~r2_hidden(X2,X0) | ~v3_ordinal1(X2)) | ~v4_ordinal1(X0))) | ~v3_ordinal1(X0))),
  inference(rectify,[],[f182])).
fof(f182,plain,(
  ! [X0] : (((v4_ordinal1(X0) | ? [X1] : (~r2_hidden(k1_ordinal1(X1),X0) & r2_hidden(X1,X0) & v3_ordinal1(X1))) & (! [X1] : (r2_hidden(k1_ordinal1(X1),X0) | ~r2_hidden(X1,X0) | ~v3_ordinal1(X1)) | ~v4_ordinal1(X0))) | ~v3_ordinal1(X0))),
  inference(nnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0] : ((v4_ordinal1(X0) <=> ! [X1] : (r2_hidden(k1_ordinal1(X1),X0) | ~r2_hidden(X1,X0) | ~v3_ordinal1(X1))) | ~v3_ordinal1(X0))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0] : ((v4_ordinal1(X0) <=> ! [X1] : ((r2_hidden(k1_ordinal1(X1),X0) | ~r2_hidden(X1,X0)) | ~v3_ordinal1(X1))) | ~v3_ordinal1(X0))),
  inference(ennf_transformation,[],[f55])).
fof(f55,axiom,(
  ! [X0] : (v3_ordinal1(X0) => (v4_ordinal1(X0) <=> ! [X1] : (v3_ordinal1(X1) => (r2_hidden(X1,X0) => r2_hidden(k1_ordinal1(X1),X0)))))),
  file('minsubpreds65/t36_ordinal5',t28_ordinal1)).
fof(f2836,plain,(
  ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | spl14_150),
  inference(avatar_contradiction_clause,[],[f2835])).
fof(f2835,plain,(
  $false | (~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | spl14_150)),
  inference(subsumption_resolution,[],[f2834,f1265])).
fof(f2834,plain,(
  ~v1_relat_1(sK13) | (~spl14_45 | ~spl14_46 | ~spl14_48 | spl14_150)),
  inference(subsumption_resolution,[],[f2833,f1260])).
fof(f2833,plain,(
  ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | (~spl14_45 | ~spl14_48 | spl14_150)),
  inference(subsumption_resolution,[],[f2832,f1270])).
fof(f2832,plain,(
  ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | (~spl14_45 | spl14_150)),
  inference(subsumption_resolution,[],[f2830,f1255])).
fof(f2830,plain,(
  ~v1_ordinal2(sK13) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | spl14_150),
  inference(resolution,[],[f2303,f292])).
fof(f2303,plain,(
  ~v3_ordinal1(k3_card_3(sK13)) | spl14_150),
  inference(avatar_component_clause,[],[f2301])).
fof(f2432,plain,(
  spl14_135 | spl14_132 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_68 | ~spl14_133),
  inference(avatar_split_clause,[],[f2431,f2051,f1394,f1268,f1263,f1258,f1253,f2048,f2084])).
fof(f2431,plain,(
  ( ! [X0] : (~v7_ordinal1(X0) | k1_funct_1(sK13,sK4(sK13)) = k1_ordinal5(k4_ordinal1,sK4(sK13)) | k1_funct_1(sK13,X0) = k1_ordinal5(k4_ordinal1,X0)) ) | (~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_68 | ~spl14_133)),
  inference(subsumption_resolution,[],[f2430,f1265])).
fof(f2430,plain,(
  ( ! [X0] : (~v7_ordinal1(X0) | k1_funct_1(sK13,sK4(sK13)) = k1_ordinal5(k4_ordinal1,sK4(sK13)) | k1_funct_1(sK13,X0) = k1_ordinal5(k4_ordinal1,X0) | ~v1_relat_1(sK13)) ) | (~spl14_45 | ~spl14_46 | ~spl14_48 | ~spl14_68 | ~spl14_133)),
  inference(subsumption_resolution,[],[f2429,f1260])).
fof(f2429,plain,(
  ( ! [X0] : (~v7_ordinal1(X0) | k1_funct_1(sK13,sK4(sK13)) = k1_ordinal5(k4_ordinal1,sK4(sK13)) | k1_funct_1(sK13,X0) = k1_ordinal5(k4_ordinal1,X0) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13)) ) | (~spl14_45 | ~spl14_48 | ~spl14_68 | ~spl14_133)),
  inference(subsumption_resolution,[],[f2428,f1270])).
fof(f2428,plain,(
  ( ! [X0] : (~v7_ordinal1(X0) | k1_funct_1(sK13,sK4(sK13)) = k1_ordinal5(k4_ordinal1,sK4(sK13)) | k1_funct_1(sK13,X0) = k1_ordinal5(k4_ordinal1,X0) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13)) ) | (~spl14_45 | ~spl14_68 | ~spl14_133)),
  inference(subsumption_resolution,[],[f2427,f1255])).
fof(f2427,plain,(
  ( ! [X0] : (~v7_ordinal1(X0) | k1_funct_1(sK13,sK4(sK13)) = k1_ordinal5(k4_ordinal1,sK4(sK13)) | k1_funct_1(sK13,X0) = k1_ordinal5(k4_ordinal1,X0) | ~v1_ordinal2(sK13) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13)) ) | (~spl14_68 | ~spl14_133)),
  inference(trivial_inequality_removal,[],[f2426])).
fof(f2426,plain,(
  ( ! [X0] : (np__1 != np__1 | ~v7_ordinal1(X0) | k1_funct_1(sK13,sK4(sK13)) = k1_ordinal5(k4_ordinal1,sK4(sK13)) | k1_funct_1(sK13,X0) = k1_ordinal5(k4_ordinal1,X0) | ~v1_ordinal2(sK13) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13)) ) | (~spl14_68 | ~spl14_133)),
  inference(superposition,[],[f2397,f1396])).
fof(f2397,plain,(
  ( ! [X0,X1] : (np__1 != k1_funct_1(X0,k6_numbers) | ~v7_ordinal1(X1) | k1_funct_1(X0,sK4(X0)) = k1_ordinal5(k4_ordinal1,sK4(X0)) | k1_funct_1(X0,X1) = k1_ordinal5(k4_ordinal1,X1) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) ) | ~spl14_133),
  inference(backward_demodulation,[],[f294,f2052])).
fof(f294,plain,(
  ( ! [X0,X1] : (k1_funct_1(X0,k6_numbers) != k1_ordinal5(k4_ordinal1,k6_numbers) | ~v7_ordinal1(X1) | k1_funct_1(X0,sK4(X0)) = k1_ordinal5(k4_ordinal1,sK4(X0)) | k1_funct_1(X0,X1) = k1_ordinal5(k4_ordinal1,X1) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f200])).
fof(f2424,plain,(
  spl14_131 | spl14_132 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_68 | ~spl14_133),
  inference(avatar_split_clause,[],[f2423,f2051,f1394,f1268,f1263,f1258,f1253,f2048,f2044])).
fof(f2423,plain,(
  ( ! [X0] : (~v7_ordinal1(X0) | v7_ordinal1(sK4(sK13)) | k1_funct_1(sK13,X0) = k1_ordinal5(k4_ordinal1,X0)) ) | (~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48 | ~spl14_68 | ~spl14_133)),
  inference(subsumption_resolution,[],[f2422,f1265])).
fof(f2422,plain,(
  ( ! [X0] : (~v7_ordinal1(X0) | v7_ordinal1(sK4(sK13)) | k1_funct_1(sK13,X0) = k1_ordinal5(k4_ordinal1,X0) | ~v1_relat_1(sK13)) ) | (~spl14_45 | ~spl14_46 | ~spl14_48 | ~spl14_68 | ~spl14_133)),
  inference(subsumption_resolution,[],[f2421,f1260])).
fof(f2421,plain,(
  ( ! [X0] : (~v7_ordinal1(X0) | v7_ordinal1(sK4(sK13)) | k1_funct_1(sK13,X0) = k1_ordinal5(k4_ordinal1,X0) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13)) ) | (~spl14_45 | ~spl14_48 | ~spl14_68 | ~spl14_133)),
  inference(subsumption_resolution,[],[f2420,f1270])).
fof(f2420,plain,(
  ( ! [X0] : (~v7_ordinal1(X0) | v7_ordinal1(sK4(sK13)) | k1_funct_1(sK13,X0) = k1_ordinal5(k4_ordinal1,X0) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13)) ) | (~spl14_45 | ~spl14_68 | ~spl14_133)),
  inference(subsumption_resolution,[],[f2419,f1255])).
fof(f2419,plain,(
  ( ! [X0] : (~v7_ordinal1(X0) | v7_ordinal1(sK4(sK13)) | k1_funct_1(sK13,X0) = k1_ordinal5(k4_ordinal1,X0) | ~v1_ordinal2(sK13) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13)) ) | (~spl14_68 | ~spl14_133)),
  inference(trivial_inequality_removal,[],[f2418])).
fof(f2418,plain,(
  ( ! [X0] : (np__1 != np__1 | ~v7_ordinal1(X0) | v7_ordinal1(sK4(sK13)) | k1_funct_1(sK13,X0) = k1_ordinal5(k4_ordinal1,X0) | ~v1_ordinal2(sK13) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13)) ) | (~spl14_68 | ~spl14_133)),
  inference(superposition,[],[f2398,f1396])).
fof(f2395,plain,(
  ~spl14_13 | spl14_133),
  inference(avatar_contradiction_clause,[],[f2394])).
fof(f2394,plain,(
  $false | (~spl14_13 | spl14_133)),
  inference(subsumption_resolution,[],[f2392,f769])).
fof(f2392,plain,(
  ~v3_ordinal1(k4_ordinal1) | spl14_133),
  inference(trivial_inequality_removal,[],[f2391])).
fof(f2391,plain,(
  np__1 != np__1 | ~v3_ordinal1(k4_ordinal1) | spl14_133),
  inference(superposition,[],[f2053,f255])).
fof(f255,plain,(
  ( ! [X0] : (np__1 = k1_ordinal5(X0,k6_numbers) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0] : (np__1 = k1_ordinal5(X0,k6_numbers) | ~v3_ordinal1(X0))),
  inference(ennf_transformation,[],[f47])).
fof(f47,axiom,(
  ! [X0] : (v3_ordinal1(X0) => np__1 = k1_ordinal5(X0,k6_numbers))),
  file('minsubpreds65/t36_ordinal5',t13_ordinal5)).
fof(f2053,plain,(
  np__1 != k1_ordinal5(k4_ordinal1,k6_numbers) | spl14_133),
  inference(avatar_component_clause,[],[f2051])).
fof(f2239,plain,(
  spl14_139 | spl14_145 | ~spl14_3 | ~spl14_8 | ~spl14_13 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48),
  inference(avatar_split_clause,[],[f2235,f1268,f1263,f1258,f1253,f1248,f767,f742,f717,f2237,f2144])).
fof(f2235,plain,(
  ( ! [X0] : (r2_hidden(sK2(k4_ordinal1,X0,sK13),k4_ordinal1) | k8_ordinal2(sK13) = k1_ordinal5(X0,k4_ordinal1) | k6_numbers = k4_ordinal1 | ~v3_ordinal1(X0)) ) | (~spl14_3 | ~spl14_8 | ~spl14_13 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48)),
  inference(subsumption_resolution,[],[f2234,f769])).
fof(f2234,plain,(
  ( ! [X0] : (r2_hidden(sK2(k4_ordinal1,X0,sK13),k4_ordinal1) | k8_ordinal2(sK13) = k1_ordinal5(X0,k4_ordinal1) | k6_numbers = k4_ordinal1 | ~v3_ordinal1(X0) | ~v3_ordinal1(k4_ordinal1)) ) | (~spl14_3 | ~spl14_8 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48)),
  inference(subsumption_resolution,[],[f2233,f719])).
fof(f2233,plain,(
  ( ! [X0] : (r2_hidden(sK2(k4_ordinal1,X0,sK13),k4_ordinal1) | k8_ordinal2(sK13) = k1_ordinal5(X0,k4_ordinal1) | k6_numbers = k4_ordinal1 | ~v4_ordinal1(k4_ordinal1) | ~v3_ordinal1(X0) | ~v3_ordinal1(k4_ordinal1)) ) | (~spl14_8 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_47 | ~spl14_48)),
  inference(subsumption_resolution,[],[f2232,f1265])).
fof(f2232,plain,(
  ( ! [X0] : (r2_hidden(sK2(k4_ordinal1,X0,sK13),k4_ordinal1) | k8_ordinal2(sK13) = k1_ordinal5(X0,k4_ordinal1) | k6_numbers = k4_ordinal1 | ~v1_relat_1(sK13) | ~v4_ordinal1(k4_ordinal1) | ~v3_ordinal1(X0) | ~v3_ordinal1(k4_ordinal1)) ) | (~spl14_8 | ~spl14_44 | ~spl14_45 | ~spl14_46 | ~spl14_48)),
  inference(subsumption_resolution,[],[f2231,f1260])).
fof(f2231,plain,(
  ( ! [X0] : (r2_hidden(sK2(k4_ordinal1,X0,sK13),k4_ordinal1) | k8_ordinal2(sK13) = k1_ordinal5(X0,k4_ordinal1) | k6_numbers = k4_ordinal1 | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | ~v4_ordinal1(k4_ordinal1) | ~v3_ordinal1(X0) | ~v3_ordinal1(k4_ordinal1)) ) | (~spl14_8 | ~spl14_44 | ~spl14_45 | ~spl14_48)),
  inference(subsumption_resolution,[],[f2230,f1270])).
fof(f2230,plain,(
  ( ! [X0] : (r2_hidden(sK2(k4_ordinal1,X0,sK13),k4_ordinal1) | k8_ordinal2(sK13) = k1_ordinal5(X0,k4_ordinal1) | k6_numbers = k4_ordinal1 | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | ~v4_ordinal1(k4_ordinal1) | ~v3_ordinal1(X0) | ~v3_ordinal1(k4_ordinal1)) ) | (~spl14_8 | ~spl14_44 | ~spl14_45)),
  inference(subsumption_resolution,[],[f2227,f1255])).
fof(f2227,plain,(
  ( ! [X0] : (r2_hidden(sK2(k4_ordinal1,X0,sK13),k4_ordinal1) | k8_ordinal2(sK13) = k1_ordinal5(X0,k4_ordinal1) | k6_numbers = k4_ordinal1 | ~v1_ordinal2(sK13) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | ~v4_ordinal1(k4_ordinal1) | ~v3_ordinal1(X0) | ~v3_ordinal1(k4_ordinal1)) ) | (~spl14_8 | ~spl14_44)),
  inference(superposition,[],[f2183,f1250])).
fof(f2183,plain,(
  ( ! [X2,X1] : (r2_hidden(sK2(k9_xtuple_0(X2),X1,X2),k9_xtuple_0(X2)) | k8_ordinal2(X2) = k1_ordinal5(X1,k9_xtuple_0(X2)) | k6_numbers = k9_xtuple_0(X2) | ~v1_ordinal2(X2) | ~v5_ordinal1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | ~v4_ordinal1(k9_xtuple_0(X2)) | ~v3_ordinal1(X1) | ~v3_ordinal1(k9_xtuple_0(X2))) ) | ~spl14_8),
  inference(forward_demodulation,[],[f701,f744])).
fof(f701,plain,(
  ( ! [X2,X1] : (k8_ordinal2(X2) = k1_ordinal5(X1,k9_xtuple_0(X2)) | r2_hidden(sK2(k9_xtuple_0(X2),X1,X2),k9_xtuple_0(X2)) | ~v1_ordinal2(X2) | ~v5_ordinal1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | k1_xboole_0 = k9_xtuple_0(X2) | ~v4_ordinal1(k9_xtuple_0(X2)) | ~v3_ordinal1(X1) | ~v3_ordinal1(k9_xtuple_0(X2))) )),
  inference(equality_resolution,[],[f270])).
fof(f270,plain,(
  ( ! [X2,X0,X1] : (k1_ordinal5(X1,X0) = k8_ordinal2(X2) | r2_hidden(sK2(X0,X1,X2),X0) | k9_xtuple_0(X2) != X0 | ~v1_ordinal2(X2) | ~v5_ordinal1(X2) | ~v1_funct_1(X2) | ~v1_relat_1(X2) | k1_xboole_0 = X0 | ~v4_ordinal1(X0) | ~v3_ordinal1(X1) | ~v3_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f189])).
fof(f2223,plain,(
  k1_xboole_0 != o_0_0_xboole_0 | k6_numbers != k1_xboole_0 | k6_numbers != k4_ordinal1 | v1_xboole_0(k4_ordinal1) | ~v1_xboole_0(o_0_0_xboole_0)),
  introduced(theory_tautology_sat_conflict,[])).
fof(f1895,plain,(
  ~spl14_69 | spl14_117),
  inference(avatar_split_clause,[],[f1889,f1893,f1402])).
fof(f1889,plain,(
  ( ! [X0,X1] : (r1_ordinal1(np__1,k1_ordinal5(X0,X1)) | ~r2_hidden(k6_numbers,X0) | ~r1_ordinal1(k6_numbers,X1) | ~v3_ordinal1(X0) | ~v3_ordinal1(X1) | ~v3_ordinal1(k6_numbers)) )),
  inference(duplicate_literal_removal,[],[f1886])).
fof(f1886,plain,(
  ( ! [X0,X1] : (r1_ordinal1(np__1,k1_ordinal5(X0,X1)) | ~r2_hidden(k6_numbers,X0) | ~r1_ordinal1(k6_numbers,X1) | ~v3_ordinal1(X0) | ~v3_ordinal1(X1) | ~v3_ordinal1(k6_numbers) | ~v3_ordinal1(X0)) )),
  inference(superposition,[],[f275,f255])).
fof(f1725,plain,(
  spl14_98 | spl14_99 | ~spl14_44 | ~spl14_46 | ~spl14_47 | ~spl14_48),
  inference(avatar_split_clause,[],[f1716,f1268,f1263,f1258,f1248,f1722,f1718])).
fof(f1716,plain,(
  r2_hidden(sK11(sK13),k4_ordinal1) | v2_ordinal5(sK13) | (~spl14_44 | ~spl14_46 | ~spl14_47 | ~spl14_48)),
  inference(subsumption_resolution,[],[f1715,f1265])).
fof(f1715,plain,(
  r2_hidden(sK11(sK13),k4_ordinal1) | v2_ordinal5(sK13) | ~v1_relat_1(sK13) | (~spl14_44 | ~spl14_46 | ~spl14_48)),
  inference(subsumption_resolution,[],[f1714,f1260])).
fof(f1714,plain,(
  r2_hidden(sK11(sK13),k4_ordinal1) | v2_ordinal5(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | (~spl14_44 | ~spl14_48)),
  inference(subsumption_resolution,[],[f1713,f1270])).
fof(f1713,plain,(
  r2_hidden(sK11(sK13),k4_ordinal1) | v2_ordinal5(sK13) | ~v5_ordinal1(sK13) | ~v1_funct_1(sK13) | ~v1_relat_1(sK13) | ~spl14_44),
  inference(superposition,[],[f660,f1250])).
fof(f1647,plain,(
  ~spl14_69 | spl14_95 | ~spl14_13 | ~spl14_82),
  inference(avatar_split_clause,[],[f1642,f1538,f767,f1644,f1402])).
fof(f1538,plain,(
  spl14_82 <=> r1_tarski(k6_numbers,k4_ordinal1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_82])])).
fof(f1642,plain,(
  r1_ordinal1(k6_numbers,k4_ordinal1) | ~v3_ordinal1(k6_numbers) | (~spl14_13 | ~spl14_82)),
  inference(subsumption_resolution,[],[f1641,f769])).
fof(f1641,plain,(
  r1_ordinal1(k6_numbers,k4_ordinal1) | ~v3_ordinal1(k4_ordinal1) | ~v3_ordinal1(k6_numbers) | ~spl14_82),
  inference(resolution,[],[f683,f1540])).
fof(f1540,plain,(
  r1_tarski(k6_numbers,k4_ordinal1) | ~spl14_82),
  inference(avatar_component_clause,[],[f1538])).
fof(f1575,plain,(
  spl14_86 | spl14_14 | ~spl14_52),
  inference(avatar_split_clause,[],[f1570,f1292,f772,f1572])).
fof(f1292,plain,(
  spl14_52 <=> m1_subset_1(np__1,k4_ordinal1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_52])])).
fof(f1570,plain,(
  r2_hidden(np__1,k4_ordinal1) | (spl14_14 | ~spl14_52)),
  inference(subsumption_resolution,[],[f1557,f774])).
fof(f1557,plain,(
  v1_xboole_0(k4_ordinal1) | r2_hidden(np__1,k4_ordinal1) | ~spl14_52),
  inference(resolution,[],[f674,f1294])).
fof(f1294,plain,(
  m1_subset_1(np__1,k4_ordinal1) | ~spl14_52),
  inference(avatar_component_clause,[],[f1292])).
fof(f1541,plain,(
  spl14_82 | ~spl14_54 | ~spl14_62),
  inference(avatar_split_clause,[],[f1536,f1356,f1304,f1538])).
fof(f1536,plain,(
  r1_tarski(k6_numbers,k4_ordinal1) | (~spl14_54 | ~spl14_62)),
  inference(subsumption_resolution,[],[f1531,f1358])).
fof(f1531,plain,(
  r1_tarski(k6_numbers,k4_ordinal1) | ~v1_ordinal1(k4_ordinal1) | ~spl14_54),
  inference(resolution,[],[f669,f1305])).
fof(f1498,plain,(
  spl14_79 | ~spl14_13 | ~spl14_53),
  inference(avatar_split_clause,[],[f1497,f1298,f767,f1490])).
fof(f1298,plain,(
  spl14_53 <=> m1_subset_1(np__0,k4_ordinal1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_53])])).
fof(f1497,plain,(
  v3_ordinal1(np__0) | (~spl14_13 | ~spl14_53)),
  inference(subsumption_resolution,[],[f1481,f769])).
fof(f1481,plain,(
  v3_ordinal1(np__0) | ~v3_ordinal1(k4_ordinal1) | ~spl14_53),
  inference(resolution,[],[f276,f1300])).
fof(f1300,plain,(
  m1_subset_1(np__0,k4_ordinal1) | ~spl14_53),
  inference(avatar_component_clause,[],[f1298])).
fof(f1495,plain,(
  spl14_69 | ~spl14_13 | ~spl14_51),
  inference(avatar_split_clause,[],[f1494,f1286,f767,f1402])).
fof(f1286,plain,(
  spl14_51 <=> m1_subset_1(k6_numbers,k4_ordinal1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_51])])).
fof(f1494,plain,(
  v3_ordinal1(k6_numbers) | (~spl14_13 | ~spl14_51)),
  inference(subsumption_resolution,[],[f1479,f769])).
fof(f1479,plain,(
  v3_ordinal1(k6_numbers) | ~v3_ordinal1(k4_ordinal1) | ~spl14_51),
  inference(resolution,[],[f276,f1288])).
fof(f1288,plain,(
  m1_subset_1(k6_numbers,k4_ordinal1) | ~spl14_51),
  inference(avatar_component_clause,[],[f1286])).
fof(f1397,plain,(
  spl14_68 | ~spl14_8 | ~spl14_43),
  inference(avatar_split_clause,[],[f1392,f1243,f742,f1394])).
fof(f1243,plain,(
  spl14_43 <=> np__1 = k1_funct_1(sK13,k1_xboole_0)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_43])])).
fof(f1392,plain,(
  np__1 = k1_funct_1(sK13,k6_numbers) | (~spl14_8 | ~spl14_43)),
  inference(forward_demodulation,[],[f1245,f744])).
fof(f1245,plain,(
  np__1 = k1_funct_1(sK13,k1_xboole_0) | ~spl14_43),
  inference(avatar_component_clause,[],[f1243])).
fof(f1391,plain,(
  ~spl14_65 | spl14_54),
  inference(avatar_split_clause,[],[f1390,f1304,f1374])).
fof(f1390,plain,(
  ~v7_ordinal1(k6_numbers) | spl14_54),
  inference(resolution,[],[f672,f1306])).
fof(f1306,plain,(
  ~r2_hidden(k6_numbers,k4_ordinal1) | spl14_54),
  inference(avatar_component_clause,[],[f1304])).
fof(f1377,plain,(
  spl14_65 | ~spl14_51),
  inference(avatar_split_clause,[],[f1370,f1286,f1374])).
fof(f1370,plain,(
  v7_ordinal1(k6_numbers) | ~spl14_51),
  inference(resolution,[],[f284,f1288])).
fof(f284,plain,(
  ( ! [X0] : (~m1_subset_1(X0,k4_ordinal1) | v7_ordinal1(X0)) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0] : (v7_ordinal1(X0) | ~m1_subset_1(X0,k4_ordinal1))),
  inference(ennf_transformation,[],[f13])).
fof(f13,axiom,(
  ! [X0] : (m1_subset_1(X0,k4_ordinal1) => v7_ordinal1(X0))),
  file('minsubpreds65/t36_ordinal5',cc8_ordinal1)).
fof(f1359,plain,(
  spl14_62 | ~spl14_13),
  inference(avatar_split_clause,[],[f1352,f767,f1356])).
fof(f1352,plain,(
  v1_ordinal1(k4_ordinal1) | ~spl14_13),
  inference(resolution,[],[f256,f769])).
fof(f1307,plain,(
  ~spl14_54 | ~spl14_8 | spl14_42),
  inference(avatar_split_clause,[],[f1302,f1239,f742,f1304])).
fof(f1239,plain,(
  spl14_42 <=> r2_hidden(k1_xboole_0,k4_ordinal1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_42])])).
fof(f1302,plain,(
  ~r2_hidden(k6_numbers,k4_ordinal1) | (~spl14_8 | spl14_42)),
  inference(forward_demodulation,[],[f1241,f744])).
fof(f1241,plain,(
  ~r2_hidden(k1_xboole_0,k4_ordinal1) | spl14_42),
  inference(avatar_component_clause,[],[f1239])).
fof(f1301,plain,(
  spl14_53 | ~spl14_9 | ~spl14_18),
  inference(avatar_split_clause,[],[f1296,f792,f747,f1298])).
fof(f747,plain,(
  spl14_9 <=> k4_ordinal1 = k5_numbers),
  introduced(avatar_definition,[new_symbols(naming,[spl14_9])])).
fof(f792,plain,(
  spl14_18 <=> m1_subset_1(np__0,k5_numbers)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_18])])).
fof(f1296,plain,(
  m1_subset_1(np__0,k4_ordinal1) | (~spl14_9 | ~spl14_18)),
  inference(forward_demodulation,[],[f794,f749])).
fof(f749,plain,(
  k4_ordinal1 = k5_numbers | ~spl14_9),
  inference(avatar_component_clause,[],[f747])).
fof(f794,plain,(
  m1_subset_1(np__0,k5_numbers) | ~spl14_18),
  inference(avatar_component_clause,[],[f792])).
fof(f1295,plain,(
  spl14_52 | ~spl14_9 | ~spl14_16),
  inference(avatar_split_clause,[],[f1290,f782,f747,f1292])).
fof(f782,plain,(
  spl14_16 <=> m1_subset_1(np__1,k5_numbers)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_16])])).
fof(f1290,plain,(
  m1_subset_1(np__1,k4_ordinal1) | (~spl14_9 | ~spl14_16)),
  inference(forward_demodulation,[],[f784,f749])).
fof(f784,plain,(
  m1_subset_1(np__1,k5_numbers) | ~spl14_16),
  inference(avatar_component_clause,[],[f782])).
fof(f1289,plain,(
  spl14_51 | ~spl14_8 | ~spl14_10),
  inference(avatar_split_clause,[],[f1284,f752,f742,f1286])).
fof(f752,plain,(
  spl14_10 <=> m1_subset_1(k1_xboole_0,k4_ordinal1)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_10])])).
fof(f1284,plain,(
  m1_subset_1(k6_numbers,k4_ordinal1) | (~spl14_8 | ~spl14_10)),
  inference(forward_demodulation,[],[f754,f744])).
fof(f754,plain,(
  m1_subset_1(k1_xboole_0,k4_ordinal1) | ~spl14_10),
  inference(avatar_component_clause,[],[f752])).
fof(f1271,plain,(
  spl14_48),
  inference(avatar_split_clause,[],[f689,f1268])).
fof(f689,plain,(
  v5_ordinal1(sK13)),
  inference(cnf_transformation,[],[f221])).
fof(f1266,plain,(
  spl14_47),
  inference(avatar_split_clause,[],[f690,f1263])).
fof(f690,plain,(
  v1_relat_1(sK13)),
  inference(cnf_transformation,[],[f221])).
fof(f1261,plain,(
  spl14_46),
  inference(avatar_split_clause,[],[f691,f1258])).
fof(f691,plain,(
  v1_funct_1(sK13)),
  inference(cnf_transformation,[],[f221])).
fof(f1256,plain,(
  spl14_45),
  inference(avatar_split_clause,[],[f692,f1253])).
fof(f692,plain,(
  v1_ordinal2(sK13)),
  inference(cnf_transformation,[],[f221])).
fof(f1251,plain,(
  spl14_44),
  inference(avatar_split_clause,[],[f693,f1248])).
fof(f693,plain,(
  k4_ordinal1 = k9_xtuple_0(sK13)),
  inference(cnf_transformation,[],[f221])).
fof(f1246,plain,(
  ~spl14_42 | spl14_43),
  inference(avatar_split_clause,[],[f694,f1243,f1239])).
fof(f694,plain,(
  np__1 = k1_funct_1(sK13,k1_xboole_0) | ~r2_hidden(k1_xboole_0,k4_ordinal1)),
  inference(cnf_transformation,[],[f221])).
fof(f1075,plain,(
  spl14_23 | spl14_25 | spl14_38),
  inference(avatar_split_clause,[],[f449,f1067,f824,f814])).
fof(f449,plain,(
  ( ! [X0,X5] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | v3_ordinal1(sK5(X0)) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(sK7,X5) | ~r2_hidden(X5,sK6) | ~v3_ordinal1(X5) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0] : ((! [X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1)) & k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) & r1_ordinal1(np__1,k3_card_3(X0))) | (k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) & r2_hidden(sK5(X0),k4_ordinal1) & v3_ordinal1(sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ((~r1_ordinal2(k12_ordinal2(k4_ordinal1,sK6),sK7) & ! [X5] : (k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(sK7,X5) | ~r2_hidden(X5,sK6) | ~v3_ordinal1(X5)) & sK6 = k9_xtuple_0(sK7) & v1_ordinal2(sK7) & v5_ordinal1(sK7) & v1_funct_1(sK7) & v1_relat_1(sK7)) & ~v1_xboole_0(sK6) & v4_ordinal1(sK6) & v3_ordinal1(sK6)) | ((~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) & r2_hidden(sK8,sK9) & v3_ordinal1(sK9)) & v3_ordinal1(sK8)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6,sK7,sK8,sK9])],[f201,f206,f205,f204,f203,f202])).
fof(f202,plain,(
  ! [X0] : (? [X2] : (k1_funct_1(X0,k1_ordinal1(X2)) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,X2)) & r2_hidden(X2,k4_ordinal1) & v3_ordinal1(X2)) => (k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) & r2_hidden(sK5(X0),k4_ordinal1) & v3_ordinal1(sK5(X0))))),
  introduced(choice_axiom,[])).
fof(f203,plain,(
  ? [X3] : (? [X4] : (~r1_ordinal2(k12_ordinal2(k4_ordinal1,X3),X4) & ! [X5] : (k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(X4,X5) | ~r2_hidden(X5,X3) | ~v3_ordinal1(X5)) & k9_xtuple_0(X4) = X3 & v1_ordinal2(X4) & v5_ordinal1(X4) & v1_funct_1(X4) & v1_relat_1(X4)) & ~v1_xboole_0(X3) & v4_ordinal1(X3) & v3_ordinal1(X3)) => (? [X4] : (~r1_ordinal2(k12_ordinal2(k4_ordinal1,sK6),X4) & ! [X5] : (k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(X4,X5) | ~r2_hidden(X5,sK6) | ~v3_ordinal1(X5)) & k9_xtuple_0(X4) = sK6 & v1_ordinal2(X4) & v5_ordinal1(X4) & v1_funct_1(X4) & v1_relat_1(X4)) & ~v1_xboole_0(sK6) & v4_ordinal1(sK6) & v3_ordinal1(sK6))),
  introduced(choice_axiom,[])).
fof(f204,plain,(
  ? [X4] : (~r1_ordinal2(k12_ordinal2(k4_ordinal1,sK6),X4) & ! [X5] : (k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(X4,X5) | ~r2_hidden(X5,sK6) | ~v3_ordinal1(X5)) & k9_xtuple_0(X4) = sK6 & v1_ordinal2(X4) & v5_ordinal1(X4) & v1_funct_1(X4) & v1_relat_1(X4)) => (~r1_ordinal2(k12_ordinal2(k4_ordinal1,sK6),sK7) & ! [X5] : (k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(sK7,X5) | ~r2_hidden(X5,sK6) | ~v3_ordinal1(X5)) & sK6 = k9_xtuple_0(sK7) & v1_ordinal2(sK7) & v5_ordinal1(sK7) & v1_funct_1(sK7) & v1_relat_1(sK7))),
  introduced(choice_axiom,[])).
fof(f205,plain,(
  ? [X6] : (? [X7] : (~r2_hidden(k12_ordinal2(k4_ordinal1,X6),k12_ordinal2(k4_ordinal1,X7)) & r2_hidden(X6,X7) & v3_ordinal1(X7)) & v3_ordinal1(X6)) => (? [X7] : (~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,X7)) & r2_hidden(sK8,X7) & v3_ordinal1(X7)) & v3_ordinal1(sK8))),
  introduced(choice_axiom,[])).
fof(f206,plain,(
  ? [X7] : (~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,X7)) & r2_hidden(sK8,X7) & v3_ordinal1(X7)) => (~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) & r2_hidden(sK8,sK9) & v3_ordinal1(sK9))),
  introduced(choice_axiom,[])).
fof(f201,plain,(
  ! [X0] : ((! [X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1)) & k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) & r1_ordinal1(np__1,k3_card_3(X0))) | ? [X2] : (k1_funct_1(X0,k1_ordinal1(X2)) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,X2)) & r2_hidden(X2,k4_ordinal1) & v3_ordinal1(X2)) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ? [X3] : (? [X4] : (~r1_ordinal2(k12_ordinal2(k4_ordinal1,X3),X4) & ! [X5] : (k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(X4,X5) | ~r2_hidden(X5,X3) | ~v3_ordinal1(X5)) & k9_xtuple_0(X4) = X3 & v1_ordinal2(X4) & v5_ordinal1(X4) & v1_funct_1(X4) & v1_relat_1(X4)) & ~v1_xboole_0(X3) & v4_ordinal1(X3) & v3_ordinal1(X3)) | ? [X6] : (? [X7] : (~r2_hidden(k12_ordinal2(k4_ordinal1,X6),k12_ordinal2(k4_ordinal1,X7)) & r2_hidden(X6,X7) & v3_ordinal1(X7)) & v3_ordinal1(X6)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))),
  inference(rectify,[],[f147])).
fof(f147,plain,(
  ! [X0] : ((! [X7] : (r1_ordinal1(k3_card_3(X0),X7) | k12_ordinal2(k4_ordinal1,X7) != X7 | ~r1_ordinal1(np__1,X7) | ~v3_ordinal1(X7)) & k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) & r1_ordinal1(np__1,k3_card_3(X0))) | ? [X1] : (k1_funct_1(X0,k1_ordinal1(X1)) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,X1)) & r2_hidden(X1,k4_ordinal1) & v3_ordinal1(X1)) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ? [X2] : (? [X3] : (~r1_ordinal2(k12_ordinal2(k4_ordinal1,X2),X3) & ! [X4] : (k1_funct_1(X3,X4) = k12_ordinal2(k4_ordinal1,X4) | ~r2_hidden(X4,X2) | ~v3_ordinal1(X4)) & k9_xtuple_0(X3) = X2 & v1_ordinal2(X3) & v5_ordinal1(X3) & v1_funct_1(X3) & v1_relat_1(X3)) & ~v1_xboole_0(X2) & v4_ordinal1(X2) & v3_ordinal1(X2)) | ? [X5] : (? [X6] : (~r2_hidden(k12_ordinal2(k4_ordinal1,X5),k12_ordinal2(k4_ordinal1,X6)) & r2_hidden(X5,X6) & v3_ordinal1(X6)) & v3_ordinal1(X5)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0))),
  inference(flattening,[],[f146])).
fof(f146,plain,(
  ! [X0] : (((! [X7] : ((r1_ordinal1(k3_card_3(X0),X7) | (k12_ordinal2(k4_ordinal1,X7) != X7 | ~r1_ordinal1(np__1,X7))) | ~v3_ordinal1(X7)) & k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) & r1_ordinal1(np__1,k3_card_3(X0))) | (? [X1] : ((k1_funct_1(X0,k1_ordinal1(X1)) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,X1)) & r2_hidden(X1,k4_ordinal1)) & v3_ordinal1(X1)) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ? [X2] : (((? [X3] : ((~r1_ordinal2(k12_ordinal2(k4_ordinal1,X2),X3) & (! [X4] : ((k1_funct_1(X3,X4) = k12_ordinal2(k4_ordinal1,X4) | ~r2_hidden(X4,X2)) | ~v3_ordinal1(X4)) & k9_xtuple_0(X3) = X2)) & (v1_ordinal2(X3) & v5_ordinal1(X3) & v1_funct_1(X3) & v1_relat_1(X3))) & ~v1_xboole_0(X2)) & v4_ordinal1(X2)) & v3_ordinal1(X2)) | ? [X5] : (? [X6] : ((~r2_hidden(k12_ordinal2(k4_ordinal1,X5),k12_ordinal2(k4_ordinal1,X6)) & r2_hidden(X5,X6)) & v3_ordinal1(X6)) & v3_ordinal1(X5)))) | (~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)))),
  inference(ennf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : ((v1_ordinal2(X0) & v5_ordinal1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => ((! [X1] : (v3_ordinal1(X1) => (r2_hidden(X1,k4_ordinal1) => k1_funct_1(X0,k1_ordinal1(X1)) = k12_ordinal2(k4_ordinal1,k1_funct_1(X0,X1)))) & np__1 = k1_funct_1(X0,k6_numbers) & k4_ordinal1 = k9_xtuple_0(X0) & ! [X2] : (v3_ordinal1(X2) => (v4_ordinal1(X2) => (! [X3] : ((v1_ordinal2(X3) & v5_ordinal1(X3) & v1_funct_1(X3) & v1_relat_1(X3)) => ((! [X4] : (v3_ordinal1(X4) => (r2_hidden(X4,X2) => k1_funct_1(X3,X4) = k12_ordinal2(k4_ordinal1,X4))) & k9_xtuple_0(X3) = X2) => r1_ordinal2(k12_ordinal2(k4_ordinal1,X2),X3))) | v1_xboole_0(X2)))) & ! [X5] : (v3_ordinal1(X5) => ! [X6] : (v3_ordinal1(X6) => (r2_hidden(X5,X6) => r2_hidden(k12_ordinal2(k4_ordinal1,X5),k12_ordinal2(k4_ordinal1,X6)))))) => (! [X7] : (v3_ordinal1(X7) => ((k12_ordinal2(k4_ordinal1,X7) = X7 & r1_ordinal1(np__1,X7)) => r1_ordinal1(k3_card_3(X0),X7))) & k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) & r1_ordinal1(np__1,k3_card_3(X0)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X0] : ((v1_ordinal2(X0) & v5_ordinal1(X0) & v1_funct_1(X0) & v1_relat_1(X0)) => ((! [X1] : (v3_ordinal1(X1) => (r2_hidden(X1,k4_ordinal1) => k1_funct_1(X0,k1_ordinal1(X1)) = k12_ordinal2(k4_ordinal1,k1_funct_1(X0,X1)))) & np__1 = k1_funct_1(X0,k6_numbers) & k4_ordinal1 = k9_xtuple_0(X0) & ! [X1] : (v3_ordinal1(X1) => (v4_ordinal1(X1) => (! [X2] : ((v1_ordinal2(X2) & v5_ordinal1(X2) & v1_funct_1(X2) & v1_relat_1(X2)) => ((! [X3] : (v3_ordinal1(X3) => (r2_hidden(X3,X1) => k1_funct_1(X2,X3) = k12_ordinal2(k4_ordinal1,X3))) & k9_xtuple_0(X2) = X1) => r1_ordinal2(k12_ordinal2(k4_ordinal1,X1),X2))) | v1_xboole_0(X1)))) & ! [X1] : (v3_ordinal1(X1) => ! [X2] : (v3_ordinal1(X2) => (r2_hidden(X1,X2) => r2_hidden(k12_ordinal2(k4_ordinal1,X1),k12_ordinal2(k4_ordinal1,X2)))))) => (! [X1] : (v3_ordinal1(X1) => ((k12_ordinal2(k4_ordinal1,X1) = X1 & r1_ordinal1(np__1,X1)) => r1_ordinal1(k3_card_3(X0),X1))) & k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) & r1_ordinal1(np__1,k3_card_3(X0)))))),
  file('minsubpreds65/t36_ordinal5',s1_ordinal5__e10_53__ordinal5)).
fof(f1074,plain,(
  spl14_22 | spl14_25 | spl14_38),
  inference(avatar_split_clause,[],[f450,f1067,f824,f809])).
fof(f450,plain,(
  ( ! [X0,X5] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | v3_ordinal1(sK5(X0)) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(sK7,X5) | ~r2_hidden(X5,sK6) | ~v3_ordinal1(X5) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1073,plain,(
  ~spl14_19 | spl14_25 | spl14_38),
  inference(avatar_split_clause,[],[f451,f1067,f824,f797])).
fof(f451,plain,(
  ( ! [X0,X5] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | v3_ordinal1(sK5(X0)) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(sK7,X5) | ~r2_hidden(X5,sK6) | ~v3_ordinal1(X5) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1064,plain,(
  spl14_23 | spl14_33 | spl14_37),
  inference(avatar_split_clause,[],[f457,f1024,f887,f814])).
fof(f457,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v3_ordinal1(sK6) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1063,plain,(
  spl14_22 | spl14_33 | spl14_37),
  inference(avatar_split_clause,[],[f458,f1024,f887,f809])).
fof(f458,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v3_ordinal1(sK6) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1062,plain,(
  ~spl14_19 | spl14_33 | spl14_37),
  inference(avatar_split_clause,[],[f459,f1024,f887,f797])).
fof(f459,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v3_ordinal1(sK6) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1060,plain,(
  spl14_23 | spl14_32 | spl14_37),
  inference(avatar_split_clause,[],[f461,f1024,f879,f814])).
fof(f461,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v4_ordinal1(sK6) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1059,plain,(
  spl14_22 | spl14_32 | spl14_37),
  inference(avatar_split_clause,[],[f462,f1024,f879,f809])).
fof(f462,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v4_ordinal1(sK6) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1058,plain,(
  ~spl14_19 | spl14_32 | spl14_37),
  inference(avatar_split_clause,[],[f463,f1024,f879,f797])).
fof(f463,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v4_ordinal1(sK6) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1056,plain,(
  spl14_23 | ~spl14_31 | spl14_37),
  inference(avatar_split_clause,[],[f465,f1024,f871,f814])).
fof(f871,plain,(
  spl14_31 <=> v1_xboole_0(sK6)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_31])])).
fof(f465,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~v1_xboole_0(sK6) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1055,plain,(
  spl14_22 | ~spl14_31 | spl14_37),
  inference(avatar_split_clause,[],[f466,f1024,f871,f809])).
fof(f466,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~v1_xboole_0(sK6) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1054,plain,(
  ~spl14_19 | ~spl14_31 | spl14_37),
  inference(avatar_split_clause,[],[f467,f1024,f871,f797])).
fof(f467,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~v1_xboole_0(sK6) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1052,plain,(
  spl14_23 | spl14_30 | spl14_37),
  inference(avatar_split_clause,[],[f469,f1024,f863,f814])).
fof(f469,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_relat_1(sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1051,plain,(
  spl14_22 | spl14_30 | spl14_37),
  inference(avatar_split_clause,[],[f470,f1024,f863,f809])).
fof(f470,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_relat_1(sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1050,plain,(
  ~spl14_19 | spl14_30 | spl14_37),
  inference(avatar_split_clause,[],[f471,f1024,f863,f797])).
fof(f471,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_relat_1(sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1048,plain,(
  spl14_23 | spl14_29 | spl14_37),
  inference(avatar_split_clause,[],[f473,f1024,f855,f814])).
fof(f473,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_funct_1(sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1047,plain,(
  spl14_22 | spl14_29 | spl14_37),
  inference(avatar_split_clause,[],[f474,f1024,f855,f809])).
fof(f474,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_funct_1(sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1046,plain,(
  ~spl14_19 | spl14_29 | spl14_37),
  inference(avatar_split_clause,[],[f475,f1024,f855,f797])).
fof(f475,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_funct_1(sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1044,plain,(
  spl14_23 | spl14_28 | spl14_37),
  inference(avatar_split_clause,[],[f477,f1024,f847,f814])).
fof(f477,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v5_ordinal1(sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1043,plain,(
  spl14_22 | spl14_28 | spl14_37),
  inference(avatar_split_clause,[],[f478,f1024,f847,f809])).
fof(f478,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v5_ordinal1(sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1042,plain,(
  ~spl14_19 | spl14_28 | spl14_37),
  inference(avatar_split_clause,[],[f479,f1024,f847,f797])).
fof(f479,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v5_ordinal1(sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1040,plain,(
  spl14_23 | spl14_27 | spl14_37),
  inference(avatar_split_clause,[],[f481,f1024,f839,f814])).
fof(f481,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_ordinal2(sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1039,plain,(
  spl14_22 | spl14_27 | spl14_37),
  inference(avatar_split_clause,[],[f482,f1024,f839,f809])).
fof(f482,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_ordinal2(sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1038,plain,(
  ~spl14_19 | spl14_27 | spl14_37),
  inference(avatar_split_clause,[],[f483,f1024,f839,f797])).
fof(f483,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_ordinal2(sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1036,plain,(
  spl14_23 | spl14_26 | spl14_37),
  inference(avatar_split_clause,[],[f485,f1024,f831,f814])).
fof(f485,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | sK6 = k9_xtuple_0(sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1035,plain,(
  spl14_22 | spl14_26 | spl14_37),
  inference(avatar_split_clause,[],[f486,f1024,f831,f809])).
fof(f486,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | sK6 = k9_xtuple_0(sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1034,plain,(
  ~spl14_19 | spl14_26 | spl14_37),
  inference(avatar_split_clause,[],[f487,f1024,f831,f797])).
fof(f487,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | sK6 = k9_xtuple_0(sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1032,plain,(
  spl14_23 | spl14_25 | spl14_37),
  inference(avatar_split_clause,[],[f489,f1024,f824,f814])).
fof(f489,plain,(
  ( ! [X0,X5] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(sK7,X5) | ~r2_hidden(X5,sK6) | ~v3_ordinal1(X5) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1031,plain,(
  spl14_22 | spl14_25 | spl14_37),
  inference(avatar_split_clause,[],[f490,f1024,f824,f809])).
fof(f490,plain,(
  ( ! [X0,X5] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(sK7,X5) | ~r2_hidden(X5,sK6) | ~v3_ordinal1(X5) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1030,plain,(
  ~spl14_19 | spl14_25 | spl14_37),
  inference(avatar_split_clause,[],[f491,f1024,f824,f797])).
fof(f491,plain,(
  ( ! [X0,X5] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(sK7,X5) | ~r2_hidden(X5,sK6) | ~v3_ordinal1(X5) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1028,plain,(
  spl14_23 | ~spl14_20 | spl14_37),
  inference(avatar_split_clause,[],[f493,f1024,f801,f814])).
fof(f493,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~r1_ordinal2(k12_ordinal2(k4_ordinal1,sK6),sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1027,plain,(
  spl14_22 | ~spl14_20 | spl14_37),
  inference(avatar_split_clause,[],[f494,f1024,f801,f809])).
fof(f494,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~r1_ordinal2(k12_ordinal2(k4_ordinal1,sK6),sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1026,plain,(
  ~spl14_19 | ~spl14_20 | spl14_37),
  inference(avatar_split_clause,[],[f495,f1024,f801,f797])).
fof(f495,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~r1_ordinal2(k12_ordinal2(k4_ordinal1,sK6),sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1022,plain,(
  spl14_24 | spl14_33 | spl14_36),
  inference(avatar_split_clause,[],[f496,f981,f887,f819])).
fof(f496,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v3_ordinal1(sK6) | v3_ordinal1(sK8) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1021,plain,(
  spl14_23 | spl14_33 | spl14_36),
  inference(avatar_split_clause,[],[f497,f981,f887,f814])).
fof(f497,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v3_ordinal1(sK6) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1020,plain,(
  spl14_22 | spl14_33 | spl14_36),
  inference(avatar_split_clause,[],[f498,f981,f887,f809])).
fof(f498,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v3_ordinal1(sK6) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1019,plain,(
  ~spl14_19 | spl14_33 | spl14_36),
  inference(avatar_split_clause,[],[f499,f981,f887,f797])).
fof(f499,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v3_ordinal1(sK6) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1017,plain,(
  spl14_23 | spl14_32 | spl14_36),
  inference(avatar_split_clause,[],[f501,f981,f879,f814])).
fof(f501,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v4_ordinal1(sK6) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1016,plain,(
  spl14_22 | spl14_32 | spl14_36),
  inference(avatar_split_clause,[],[f502,f981,f879,f809])).
fof(f502,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v4_ordinal1(sK6) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1015,plain,(
  ~spl14_19 | spl14_32 | spl14_36),
  inference(avatar_split_clause,[],[f503,f981,f879,f797])).
fof(f503,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v4_ordinal1(sK6) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1013,plain,(
  spl14_23 | ~spl14_31 | spl14_36),
  inference(avatar_split_clause,[],[f505,f981,f871,f814])).
fof(f505,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~v1_xboole_0(sK6) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1012,plain,(
  spl14_22 | ~spl14_31 | spl14_36),
  inference(avatar_split_clause,[],[f506,f981,f871,f809])).
fof(f506,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~v1_xboole_0(sK6) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1011,plain,(
  ~spl14_19 | ~spl14_31 | spl14_36),
  inference(avatar_split_clause,[],[f507,f981,f871,f797])).
fof(f507,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~v1_xboole_0(sK6) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1010,plain,(
  spl14_24 | spl14_30 | spl14_36),
  inference(avatar_split_clause,[],[f508,f981,f863,f819])).
fof(f508,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_relat_1(sK7) | v3_ordinal1(sK8) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1009,plain,(
  spl14_23 | spl14_30 | spl14_36),
  inference(avatar_split_clause,[],[f509,f981,f863,f814])).
fof(f509,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_relat_1(sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1008,plain,(
  spl14_22 | spl14_30 | spl14_36),
  inference(avatar_split_clause,[],[f510,f981,f863,f809])).
fof(f510,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_relat_1(sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1007,plain,(
  ~spl14_19 | spl14_30 | spl14_36),
  inference(avatar_split_clause,[],[f511,f981,f863,f797])).
fof(f511,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_relat_1(sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1005,plain,(
  spl14_23 | spl14_29 | spl14_36),
  inference(avatar_split_clause,[],[f513,f981,f855,f814])).
fof(f513,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_funct_1(sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1004,plain,(
  spl14_22 | spl14_29 | spl14_36),
  inference(avatar_split_clause,[],[f514,f981,f855,f809])).
fof(f514,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_funct_1(sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1003,plain,(
  ~spl14_19 | spl14_29 | spl14_36),
  inference(avatar_split_clause,[],[f515,f981,f855,f797])).
fof(f515,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_funct_1(sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1001,plain,(
  spl14_23 | spl14_28 | spl14_36),
  inference(avatar_split_clause,[],[f517,f981,f847,f814])).
fof(f517,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v5_ordinal1(sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f1000,plain,(
  spl14_22 | spl14_28 | spl14_36),
  inference(avatar_split_clause,[],[f518,f981,f847,f809])).
fof(f518,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v5_ordinal1(sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f999,plain,(
  ~spl14_19 | spl14_28 | spl14_36),
  inference(avatar_split_clause,[],[f519,f981,f847,f797])).
fof(f519,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v5_ordinal1(sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f997,plain,(
  spl14_23 | spl14_27 | spl14_36),
  inference(avatar_split_clause,[],[f521,f981,f839,f814])).
fof(f521,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_ordinal2(sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f996,plain,(
  spl14_22 | spl14_27 | spl14_36),
  inference(avatar_split_clause,[],[f522,f981,f839,f809])).
fof(f522,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_ordinal2(sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f995,plain,(
  ~spl14_19 | spl14_27 | spl14_36),
  inference(avatar_split_clause,[],[f523,f981,f839,f797])).
fof(f523,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_ordinal2(sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f993,plain,(
  spl14_23 | spl14_26 | spl14_36),
  inference(avatar_split_clause,[],[f525,f981,f831,f814])).
fof(f525,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | sK6 = k9_xtuple_0(sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f992,plain,(
  spl14_22 | spl14_26 | spl14_36),
  inference(avatar_split_clause,[],[f526,f981,f831,f809])).
fof(f526,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | sK6 = k9_xtuple_0(sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f991,plain,(
  ~spl14_19 | spl14_26 | spl14_36),
  inference(avatar_split_clause,[],[f527,f981,f831,f797])).
fof(f527,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | sK6 = k9_xtuple_0(sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f989,plain,(
  spl14_23 | spl14_25 | spl14_36),
  inference(avatar_split_clause,[],[f529,f981,f824,f814])).
fof(f529,plain,(
  ( ! [X0,X5] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(sK7,X5) | ~r2_hidden(X5,sK6) | ~v3_ordinal1(X5) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f988,plain,(
  spl14_22 | spl14_25 | spl14_36),
  inference(avatar_split_clause,[],[f530,f981,f824,f809])).
fof(f530,plain,(
  ( ! [X0,X5] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(sK7,X5) | ~r2_hidden(X5,sK6) | ~v3_ordinal1(X5) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f987,plain,(
  ~spl14_19 | spl14_25 | spl14_36),
  inference(avatar_split_clause,[],[f531,f981,f824,f797])).
fof(f531,plain,(
  ( ! [X0,X5] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(sK7,X5) | ~r2_hidden(X5,sK6) | ~v3_ordinal1(X5) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f985,plain,(
  spl14_23 | ~spl14_20 | spl14_36),
  inference(avatar_split_clause,[],[f533,f981,f801,f814])).
fof(f533,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~r1_ordinal2(k12_ordinal2(k4_ordinal1,sK6),sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f984,plain,(
  spl14_22 | ~spl14_20 | spl14_36),
  inference(avatar_split_clause,[],[f534,f981,f801,f809])).
fof(f534,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~r1_ordinal2(k12_ordinal2(k4_ordinal1,sK6),sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f983,plain,(
  ~spl14_19 | ~spl14_20 | spl14_36),
  inference(avatar_split_clause,[],[f535,f981,f801,f797])).
fof(f535,plain,(
  ( ! [X0] : (k3_card_3(X0) = k12_ordinal2(k4_ordinal1,k3_card_3(X0)) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~r1_ordinal2(k12_ordinal2(k4_ordinal1,sK6),sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f935,plain,(
  spl14_23 | spl14_33 | spl14_34),
  inference(avatar_split_clause,[],[f577,f895,f887,f814])).
fof(f577,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v3_ordinal1(sK6) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f934,plain,(
  spl14_22 | spl14_33 | spl14_34),
  inference(avatar_split_clause,[],[f578,f895,f887,f809])).
fof(f578,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v3_ordinal1(sK6) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f933,plain,(
  ~spl14_19 | spl14_33 | spl14_34),
  inference(avatar_split_clause,[],[f579,f895,f887,f797])).
fof(f579,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v3_ordinal1(sK6) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f931,plain,(
  spl14_23 | spl14_32 | spl14_34),
  inference(avatar_split_clause,[],[f581,f895,f879,f814])).
fof(f581,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v4_ordinal1(sK6) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f930,plain,(
  spl14_22 | spl14_32 | spl14_34),
  inference(avatar_split_clause,[],[f582,f895,f879,f809])).
fof(f582,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v4_ordinal1(sK6) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f929,plain,(
  ~spl14_19 | spl14_32 | spl14_34),
  inference(avatar_split_clause,[],[f583,f895,f879,f797])).
fof(f583,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v4_ordinal1(sK6) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f927,plain,(
  spl14_23 | ~spl14_31 | spl14_34),
  inference(avatar_split_clause,[],[f585,f895,f871,f814])).
fof(f585,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~v1_xboole_0(sK6) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f926,plain,(
  spl14_22 | ~spl14_31 | spl14_34),
  inference(avatar_split_clause,[],[f586,f895,f871,f809])).
fof(f586,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~v1_xboole_0(sK6) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f925,plain,(
  ~spl14_19 | ~spl14_31 | spl14_34),
  inference(avatar_split_clause,[],[f587,f895,f871,f797])).
fof(f587,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~v1_xboole_0(sK6) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f923,plain,(
  spl14_23 | spl14_30 | spl14_34),
  inference(avatar_split_clause,[],[f589,f895,f863,f814])).
fof(f589,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_relat_1(sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f922,plain,(
  spl14_22 | spl14_30 | spl14_34),
  inference(avatar_split_clause,[],[f590,f895,f863,f809])).
fof(f590,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_relat_1(sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f921,plain,(
  ~spl14_19 | spl14_30 | spl14_34),
  inference(avatar_split_clause,[],[f591,f895,f863,f797])).
fof(f591,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_relat_1(sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f919,plain,(
  spl14_23 | spl14_29 | spl14_34),
  inference(avatar_split_clause,[],[f593,f895,f855,f814])).
fof(f593,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_funct_1(sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f918,plain,(
  spl14_22 | spl14_29 | spl14_34),
  inference(avatar_split_clause,[],[f594,f895,f855,f809])).
fof(f594,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_funct_1(sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f917,plain,(
  ~spl14_19 | spl14_29 | spl14_34),
  inference(avatar_split_clause,[],[f595,f895,f855,f797])).
fof(f595,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_funct_1(sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f915,plain,(
  spl14_23 | spl14_28 | spl14_34),
  inference(avatar_split_clause,[],[f597,f895,f847,f814])).
fof(f597,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v5_ordinal1(sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f914,plain,(
  spl14_22 | spl14_28 | spl14_34),
  inference(avatar_split_clause,[],[f598,f895,f847,f809])).
fof(f598,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v5_ordinal1(sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f913,plain,(
  ~spl14_19 | spl14_28 | spl14_34),
  inference(avatar_split_clause,[],[f599,f895,f847,f797])).
fof(f599,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v5_ordinal1(sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f911,plain,(
  spl14_23 | spl14_27 | spl14_34),
  inference(avatar_split_clause,[],[f601,f895,f839,f814])).
fof(f601,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_ordinal2(sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f910,plain,(
  spl14_22 | spl14_27 | spl14_34),
  inference(avatar_split_clause,[],[f602,f895,f839,f809])).
fof(f602,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_ordinal2(sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f909,plain,(
  ~spl14_19 | spl14_27 | spl14_34),
  inference(avatar_split_clause,[],[f603,f895,f839,f797])).
fof(f603,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_ordinal2(sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f907,plain,(
  spl14_23 | spl14_26 | spl14_34),
  inference(avatar_split_clause,[],[f605,f895,f831,f814])).
fof(f605,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | sK6 = k9_xtuple_0(sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f906,plain,(
  spl14_22 | spl14_26 | spl14_34),
  inference(avatar_split_clause,[],[f606,f895,f831,f809])).
fof(f606,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | sK6 = k9_xtuple_0(sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f905,plain,(
  ~spl14_19 | spl14_26 | spl14_34),
  inference(avatar_split_clause,[],[f607,f895,f831,f797])).
fof(f607,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | sK6 = k9_xtuple_0(sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f904,plain,(
  spl14_24 | spl14_25 | spl14_34),
  inference(avatar_split_clause,[],[f608,f895,f824,f819])).
fof(f608,plain,(
  ( ! [X0,X5,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(sK7,X5) | ~r2_hidden(X5,sK6) | ~v3_ordinal1(X5) | v3_ordinal1(sK8) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f903,plain,(
  spl14_23 | spl14_25 | spl14_34),
  inference(avatar_split_clause,[],[f609,f895,f824,f814])).
fof(f609,plain,(
  ( ! [X0,X5,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(sK7,X5) | ~r2_hidden(X5,sK6) | ~v3_ordinal1(X5) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f902,plain,(
  spl14_22 | spl14_25 | spl14_34),
  inference(avatar_split_clause,[],[f610,f895,f824,f809])).
fof(f610,plain,(
  ( ! [X0,X5,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(sK7,X5) | ~r2_hidden(X5,sK6) | ~v3_ordinal1(X5) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f901,plain,(
  ~spl14_19 | spl14_25 | spl14_34),
  inference(avatar_split_clause,[],[f611,f895,f824,f797])).
fof(f611,plain,(
  ( ! [X0,X5,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(sK7,X5) | ~r2_hidden(X5,sK6) | ~v3_ordinal1(X5) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f899,plain,(
  spl14_23 | ~spl14_20 | spl14_34),
  inference(avatar_split_clause,[],[f613,f895,f801,f814])).
fof(f613,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~r1_ordinal2(k12_ordinal2(k4_ordinal1,sK6),sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f898,plain,(
  spl14_22 | ~spl14_20 | spl14_34),
  inference(avatar_split_clause,[],[f614,f895,f801,f809])).
fof(f614,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~r1_ordinal2(k12_ordinal2(k4_ordinal1,sK6),sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f897,plain,(
  ~spl14_19 | ~spl14_20 | spl14_34),
  inference(avatar_split_clause,[],[f615,f895,f801,f797])).
fof(f615,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | r2_hidden(sK5(X0),k4_ordinal1) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~r1_ordinal2(k12_ordinal2(k4_ordinal1,sK6),sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f892,plain,(
  spl14_23 | spl14_33 | spl14_21),
  inference(avatar_split_clause,[],[f617,f805,f887,f814])).
fof(f617,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v3_ordinal1(sK6) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f891,plain,(
  spl14_22 | spl14_33 | spl14_21),
  inference(avatar_split_clause,[],[f618,f805,f887,f809])).
fof(f618,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v3_ordinal1(sK6) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f890,plain,(
  ~spl14_19 | spl14_33 | spl14_21),
  inference(avatar_split_clause,[],[f619,f805,f887,f797])).
fof(f619,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v3_ordinal1(sK6) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f884,plain,(
  spl14_23 | spl14_32 | spl14_21),
  inference(avatar_split_clause,[],[f621,f805,f879,f814])).
fof(f621,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v4_ordinal1(sK6) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f883,plain,(
  spl14_22 | spl14_32 | spl14_21),
  inference(avatar_split_clause,[],[f622,f805,f879,f809])).
fof(f622,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v4_ordinal1(sK6) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f882,plain,(
  ~spl14_19 | spl14_32 | spl14_21),
  inference(avatar_split_clause,[],[f623,f805,f879,f797])).
fof(f623,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v4_ordinal1(sK6) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f876,plain,(
  spl14_23 | ~spl14_31 | spl14_21),
  inference(avatar_split_clause,[],[f625,f805,f871,f814])).
fof(f625,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~v1_xboole_0(sK6) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f875,plain,(
  spl14_22 | ~spl14_31 | spl14_21),
  inference(avatar_split_clause,[],[f626,f805,f871,f809])).
fof(f626,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~v1_xboole_0(sK6) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f874,plain,(
  ~spl14_19 | ~spl14_31 | spl14_21),
  inference(avatar_split_clause,[],[f627,f805,f871,f797])).
fof(f627,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~v1_xboole_0(sK6) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f868,plain,(
  spl14_23 | spl14_30 | spl14_21),
  inference(avatar_split_clause,[],[f629,f805,f863,f814])).
fof(f629,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_relat_1(sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f867,plain,(
  spl14_22 | spl14_30 | spl14_21),
  inference(avatar_split_clause,[],[f630,f805,f863,f809])).
fof(f630,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_relat_1(sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f866,plain,(
  ~spl14_19 | spl14_30 | spl14_21),
  inference(avatar_split_clause,[],[f631,f805,f863,f797])).
fof(f631,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_relat_1(sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f860,plain,(
  spl14_23 | spl14_29 | spl14_21),
  inference(avatar_split_clause,[],[f633,f805,f855,f814])).
fof(f633,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_funct_1(sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f859,plain,(
  spl14_22 | spl14_29 | spl14_21),
  inference(avatar_split_clause,[],[f634,f805,f855,f809])).
fof(f634,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_funct_1(sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f858,plain,(
  ~spl14_19 | spl14_29 | spl14_21),
  inference(avatar_split_clause,[],[f635,f805,f855,f797])).
fof(f635,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_funct_1(sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f852,plain,(
  spl14_23 | spl14_28 | spl14_21),
  inference(avatar_split_clause,[],[f637,f805,f847,f814])).
fof(f637,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v5_ordinal1(sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f851,plain,(
  spl14_22 | spl14_28 | spl14_21),
  inference(avatar_split_clause,[],[f638,f805,f847,f809])).
fof(f638,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v5_ordinal1(sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f850,plain,(
  ~spl14_19 | spl14_28 | spl14_21),
  inference(avatar_split_clause,[],[f639,f805,f847,f797])).
fof(f639,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v5_ordinal1(sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f844,plain,(
  spl14_23 | spl14_27 | spl14_21),
  inference(avatar_split_clause,[],[f641,f805,f839,f814])).
fof(f641,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_ordinal2(sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f843,plain,(
  spl14_22 | spl14_27 | spl14_21),
  inference(avatar_split_clause,[],[f642,f805,f839,f809])).
fof(f642,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_ordinal2(sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f842,plain,(
  ~spl14_19 | spl14_27 | spl14_21),
  inference(avatar_split_clause,[],[f643,f805,f839,f797])).
fof(f643,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | v1_ordinal2(sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f836,plain,(
  spl14_23 | spl14_26 | spl14_21),
  inference(avatar_split_clause,[],[f645,f805,f831,f814])).
fof(f645,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | sK6 = k9_xtuple_0(sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f835,plain,(
  spl14_22 | spl14_26 | spl14_21),
  inference(avatar_split_clause,[],[f646,f805,f831,f809])).
fof(f646,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | sK6 = k9_xtuple_0(sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f834,plain,(
  ~spl14_19 | spl14_26 | spl14_21),
  inference(avatar_split_clause,[],[f647,f805,f831,f797])).
fof(f647,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | sK6 = k9_xtuple_0(sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f829,plain,(
  spl14_24 | spl14_25 | spl14_21),
  inference(avatar_split_clause,[],[f648,f805,f824,f819])).
fof(f648,plain,(
  ( ! [X0,X5,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(sK7,X5) | ~r2_hidden(X5,sK6) | ~v3_ordinal1(X5) | v3_ordinal1(sK8) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f828,plain,(
  spl14_23 | spl14_25 | spl14_21),
  inference(avatar_split_clause,[],[f649,f805,f824,f814])).
fof(f649,plain,(
  ( ! [X0,X5,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(sK7,X5) | ~r2_hidden(X5,sK6) | ~v3_ordinal1(X5) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f827,plain,(
  spl14_22 | spl14_25 | spl14_21),
  inference(avatar_split_clause,[],[f650,f805,f824,f809])).
fof(f650,plain,(
  ( ! [X0,X5,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(sK7,X5) | ~r2_hidden(X5,sK6) | ~v3_ordinal1(X5) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f826,plain,(
  ~spl14_19 | spl14_25 | spl14_21),
  inference(avatar_split_clause,[],[f651,f805,f824,f797])).
fof(f651,plain,(
  ( ! [X0,X5,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | k12_ordinal2(k4_ordinal1,X5) = k1_funct_1(sK7,X5) | ~r2_hidden(X5,sK6) | ~v3_ordinal1(X5) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f817,plain,(
  spl14_23 | ~spl14_20 | spl14_21),
  inference(avatar_split_clause,[],[f653,f805,f801,f814])).
fof(f653,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~r1_ordinal2(k12_ordinal2(k4_ordinal1,sK6),sK7) | v3_ordinal1(sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f812,plain,(
  spl14_22 | ~spl14_20 | spl14_21),
  inference(avatar_split_clause,[],[f654,f805,f801,f809])).
fof(f654,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~r1_ordinal2(k12_ordinal2(k4_ordinal1,sK6),sK7) | r2_hidden(sK8,sK9) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f807,plain,(
  ~spl14_19 | ~spl14_20 | spl14_21),
  inference(avatar_split_clause,[],[f655,f805,f801,f797])).
fof(f655,plain,(
  ( ! [X0,X1] : (r1_ordinal1(k3_card_3(X0),X1) | k12_ordinal2(k4_ordinal1,X1) != X1 | ~r1_ordinal1(np__1,X1) | ~v3_ordinal1(X1) | k1_funct_1(X0,k1_ordinal1(sK5(X0))) != k12_ordinal2(k4_ordinal1,k1_funct_1(X0,sK5(X0))) | np__1 != k1_funct_1(X0,k6_numbers) | k4_ordinal1 != k9_xtuple_0(X0) | ~r1_ordinal2(k12_ordinal2(k4_ordinal1,sK6),sK7) | ~r2_hidden(k12_ordinal2(k4_ordinal1,sK8),k12_ordinal2(k4_ordinal1,sK9)) | ~v1_ordinal2(X0) | ~v5_ordinal1(X0) | ~v1_funct_1(X0) | ~v1_relat_1(X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f795,plain,(
  spl14_18),
  inference(avatar_split_clause,[],[f238,f792])).
fof(f238,plain,(
  m1_subset_1(np__0,k5_numbers)),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  m1_subset_1(np__0,k1_numbers) & m1_subset_1(np__0,k5_numbers)),
  inference(pure_predicate_removal,[],[f44])).
fof(f44,axiom,(
  m1_subset_1(np__0,k1_numbers) & m1_subset_1(np__0,k5_numbers) & m2_subset_1(np__0,k1_numbers,k5_numbers)),
  file('minsubpreds65/t36_ordinal5',spc0_numerals)).
fof(f785,plain,(
  spl14_16),
  inference(avatar_split_clause,[],[f236,f782])).
fof(f236,plain,(
  m1_subset_1(np__1,k5_numbers)),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  m1_subset_1(np__1,k1_numbers) & m1_subset_1(np__1,k5_numbers)),
  inference(pure_predicate_removal,[],[f93])).
fof(f93,plain,(
  m1_subset_1(np__1,k1_numbers) & m1_subset_1(np__1,k5_numbers) & m2_subset_1(np__1,k1_numbers,k5_numbers)),
  inference(pure_predicate_removal,[],[f45])).
fof(f45,axiom,(
  m1_subset_1(np__1,k1_numbers) & m1_subset_1(np__1,k5_numbers) & m2_subset_1(np__1,k1_numbers,k5_numbers) & v2_xxreal_0(np__1)),
  file('minsubpreds65/t36_ordinal5',spc1_numerals)).
fof(f775,plain,(
  ~spl14_14),
  inference(avatar_split_clause,[],[f234,f772])).
fof(f234,plain,(
  ~v1_xboole_0(k4_ordinal1)),
  inference(cnf_transformation,[],[f33])).
fof(f33,axiom,(
  v3_ordinal1(k4_ordinal1) & ~v1_xboole_0(k4_ordinal1)),
  file('minsubpreds65/t36_ordinal5',fc6_ordinal1)).
fof(f770,plain,(
  spl14_13),
  inference(avatar_split_clause,[],[f235,f767])).
fof(f235,plain,(
  v3_ordinal1(k4_ordinal1)),
  inference(cnf_transformation,[],[f33])).
fof(f765,plain,(
  spl14_12),
  inference(avatar_split_clause,[],[f233,f762])).
fof(f233,plain,(
  k1_ordinal1(np__0) = np__1),
  inference(cnf_transformation,[],[f39])).
fof(f39,axiom,(
  k1_ordinal1(np__0) = np__1),
  file('minsubpreds65/t36_ordinal5',rqSucc__k1_ordinal1__r0_r1)).
fof(f755,plain,(
  spl14_10),
  inference(avatar_split_clause,[],[f231,f752])).
fof(f231,plain,(
  m1_subset_1(k1_xboole_0,k4_ordinal1)),
  inference(cnf_transformation,[],[f50])).
fof(f50,axiom,(
  m1_subset_1(k1_xboole_0,k4_ordinal1)),
  file('minsubpreds65/t36_ordinal5',t1_numerals)).
fof(f750,plain,(
  spl14_9),
  inference(avatar_split_clause,[],[f230,f747])).
fof(f230,plain,(
  k4_ordinal1 = k5_numbers),
  inference(cnf_transformation,[],[f36])).
fof(f36,axiom,(
  k4_ordinal1 = k5_numbers),
  file('minsubpreds65/t36_ordinal5',redefinition_k5_numbers)).
fof(f745,plain,(
  spl14_8),
  inference(avatar_split_clause,[],[f229,f742])).
fof(f229,plain,(
  k6_numbers = k1_xboole_0),
  inference(cnf_transformation,[],[f37])).
fof(f37,axiom,(
  k6_numbers = k1_xboole_0),
  file('minsubpreds65/t36_ordinal5',redefinition_k6_numbers)).
fof(f740,plain,(
  spl14_7),
  inference(avatar_split_clause,[],[f228,f737])).
fof(f737,plain,(
  spl14_7 <=> k1_xboole_0 = o_0_0_xboole_0),
  introduced(avatar_definition,[new_symbols(naming,[spl14_7])])).
fof(f228,plain,(
  k1_xboole_0 = o_0_0_xboole_0),
  inference(cnf_transformation,[],[f18])).
fof(f18,axiom,(
  k1_xboole_0 = o_0_0_xboole_0),
  file('minsubpreds65/t36_ordinal5',d2_xboole_0)).
fof(f730,plain,(
  spl14_5),
  inference(avatar_split_clause,[],[f226,f727])).
fof(f226,plain,(
  v1_xboole_0(np__0)),
  inference(cnf_transformation,[],[f43])).
fof(f43,axiom,(
  v1_xboole_0(np__0)),
  file('minsubpreds65/t36_ordinal5',spc0_boole)).
fof(f725,plain,(
  spl14_4),
  inference(avatar_split_clause,[],[f225,f722])).
fof(f722,plain,(
  spl14_4 <=> v1_xboole_0(o_0_0_xboole_0)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_4])])).
fof(f225,plain,(
  v1_xboole_0(o_0_0_xboole_0)),
  inference(cnf_transformation,[],[f24])).
fof(f24,axiom,(
  v1_xboole_0(o_0_0_xboole_0)),
  file('minsubpreds65/t36_ordinal5',dt_o_0_0_xboole_0)).
fof(f720,plain,(
  spl14_3),
  inference(avatar_split_clause,[],[f224,f717])).
fof(f224,plain,(
  v4_ordinal1(k4_ordinal1)),
  inference(cnf_transformation,[],[f29])).
fof(f29,axiom,(
  v4_ordinal1(k4_ordinal1)),
  file('minsubpreds65/t36_ordinal5',fc1_ordinal5)).
fof(f710,plain,(
  ~spl14_1),
  inference(avatar_split_clause,[],[f222,f707])).
fof(f222,plain,(
  k2_ordinal5(k6_numbers) != k1_ordinal5(k4_ordinal1,k4_ordinal1)),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  k2_ordinal5(k6_numbers) != k1_ordinal5(k4_ordinal1,k4_ordinal1)),
  inference(flattening,[],[f2])).
fof(f2,negated_conjecture,(
  ~k2_ordinal5(k6_numbers) = k1_ordinal5(k4_ordinal1,k4_ordinal1)),
  inference(negated_conjecture,[],[f1])).
fof(f1,conjecture,(
  k2_ordinal5(k6_numbers) = k1_ordinal5(k4_ordinal1,k4_ordinal1)),
  file('minsubpreds65/t36_ordinal5',t36_ordinal5)).
% SZS output end Proof for t36_ordinal5
% ------------------------------
% Version: Vampire 4.5.1 (commit 110f4142 on 2020-10-16 16:55:15 +0200)
% Termination reason: Refutation

% >>> Input
% Input formulas: 73
% >>> Preprocessing
% Introduced skolems: 14
% Pure predicates: 4
% 
% >>> Saturation
% Initial clauses: 475
% Generated clauses: 2392519
% Active clauses: 30495
% Passive clauses: 543073
% Final active clauses: 9237
% Final passive clauses: 236587
% Main loop iterations started: 38065
% Innermost selections: 25377
% 
% >>> Simplifying Inferences
% Duplicate literals: 146659
% Trivial inequalities: 1335
% Fw subsumption resolutions: 762690
% Fw demodulations: 82745
% Bw demodulations: 1804
% 
% >>> Deletion Inferences
% Simple tautologies: 6345
% Equational tautologies: 8
% Forward subsumptions: 555782
% Backward subsumptions: 1500
% Fw demodulations to eq. taut.: 6
% 
% >>> Generating Inferences
% Binary resolution: 679036
% Unit resulting resolution: 2946
% Factoring: 7
% Forward superposition: 381940
% Backward superposition: 150704
% Self superposition: 71
% Equality resolution: 48
% 
% >>> AVATAR
% Split clauses: 238434
% Split components: 559728
% Sat splitting refutations: 4917
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 308927
% SAT solver unit clauses: 25
% SAT solver binary clauses: 7629
% 
% Memory used [KB]: 2697397
% Time elapsed: 503.976 s
% ------------------------------
% ------------------------------
% Time measurement results:
% parsing: 0.003 s
% property evaluation: 0.001 s
% preprocessing: 0.001 s
% SAT solver time: 181.936 s
% minimizing solver time: 11.624 s
% sat proof minimization: 0.382 s
% term sharing: 11.207 s
% splitting model update: 1.626 s ( own 0.144 s ) 
% congruence closure: 167.665 s ( own 156.041 s ) 
% unit clause index maintenance: 0.104 s
% non unit clause index maintenance: 0.201 s ( own 0.193 s ) 
% forward subsumption index maintenance: 0.031 s ( own 0.029 s ) 
% binary resolution index maintenance: 0.124 s ( own 0.123 s ) 
% backward subsumption index maintenance: 0.326 s ( own 0.303 s ) 
% backward superposition index maintenance: 0.421 s
% forward superposition index maintenance: 0.004 s
% backward demodulation index maintenance: 0.600 s ( own 0.592 s ) 
% forward demodulation index maintenance: 0.005 s
% splitting component index maintenance: 0.133 s ( own 0.125 s ) 
% splitting component index usage: 1.815 s ( own 1.572 s ) 
% forward subsumption: 37.502 s
% forward subsumption resolution: 22.588 s
% backward subsumption: 0.451 s
% forward demodulation: 1.746 s ( own 1.715 s ) 
% backward demodulation: 0.002 s
% resolution: 3.886 s ( own 2.252 s ) 
% unit resulting resolution: 29.770 s ( own 22.063 s ) 
% superposition: 3.630 s ( own 2.090 s ) 
% literal selection: 0.010 s
% passive container maintenance: 24.190 s ( own 5.815 s ) 
% neural model evaluation: 18.493 s
% other: 503.976 s ( own 26.185 s ) 

