2020-06-01 09:59:13.582755: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_8_bb1/premsel_enigma_01_2020_T30_loop02_epoch_389
2020-06-01 09:59:13.647493: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-06-01 09:59:13.762556: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-06-01 09:59:14.175809: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-06-01 09:59:15.663432: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 2080711 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_8_bb1/premsel_enigma_01_2020_T30_loop02_epoch_389' loaded (query=256; context=768; weigths=real; len_mult=0.000000).
# Preprocessing time       : 3.030 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((X2!=k1_xboole_0=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X3)))&(X2=k1_xboole_0=>(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0)))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', d1_funct_2)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', redefinition_k6_numbers)).
fof(dt_k6_weierstr, axiom, ![X1, X2]:(((((((~(v2_struct_0(X1))&v6_metric_1(X1))&v7_metric_1(X1))&v8_metric_1(X1))&v9_metric_1(X1))&l1_metric_1(X1))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))=>((v1_funct_1(k6_weierstr(X1,X2))&v1_funct_2(k6_weierstr(X1,X2),u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr)))&m1_subset_1(k6_weierstr(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr)))))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', dt_k6_weierstr)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>k1_relset_1(X1,X2)=k9_xtuple_0(X2)), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', redefinition_k1_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', cc1_relset_1)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', cc2_relset_1)).
fof(t37_hausdorf, conjecture, ![X1]:((((((~(v2_struct_0(X1))&v6_metric_1(X1))&v7_metric_1(X1))&v8_metric_1(X1))&v9_metric_1(X1))&l1_metric_1(X1))=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))=>![X3]:((~(v1_xboole_0(X3))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))=>(((v2_compts_1(X2,k3_pcomps_1(X1))&v2_compts_1(X3,k3_pcomps_1(X1)))&k1_hausdorf(X1,X2,X3)=k6_numbers)=>X2=X3)))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', t37_hausdorf)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', d10_xboole_0)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', t3_subset)).
fof(t7_boole, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', t7_boole)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', d3_tarski)).
fof(d6_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2, X3]:(X3=k7_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:((r2_hidden(X5,k9_xtuple_0(X1))&r2_hidden(X5,X2))&X4=k1_funct_1(X1,X5))))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', d6_funct_1)).
fof(fc1_waybel_2, axiom, ![X1, X2, X3, X4]:(((((((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&v1_funct_1(X3))&v1_funct_2(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&~(v1_xboole_0(X4)))&m1_subset_1(X4,k1_zfmisc_1(X1)))=>~(v1_xboole_0(k7_relat_1(X3,X4)))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', fc1_waybel_2)).
fof(cc1_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', cc1_subset_1)).
fof(reflexivity_r1_tarski, axiom, ![X1, X2]:r1_tarski(X1,X1), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', reflexivity_r1_tarski)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', fc2_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', dt_l1_pre_topc)).
fof(t12_topmetr, axiom, ![X1]:(l1_metric_1(X1)=>(u1_struct_0(X1)=u1_struct_0(k3_pcomps_1(X1))&u1_pre_topc(k3_pcomps_1(X1))=k2_pcomps_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', t12_topmetr)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', t6_boole)).
fof(d8_weierstr, axiom, ![X1]:((((((~(v2_struct_0(X1))&v6_metric_1(X1))&v7_metric_1(X1))&v8_metric_1(X1))&v9_metric_1(X1))&l1_metric_1(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))=>k8_weierstr(X1,X2,X3)=k2_weierstr(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3))))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', d8_weierstr)).
fof(redefinition_k7_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k7_relset_1(X1,X2,X3,X4)=k7_relat_1(X3,X4)), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', redefinition_k7_relset_1)).
fof(fc2_topmetr, axiom, ((~(v2_struct_0(k3_topmetr))&v1_pre_topc(k3_topmetr))&v2_pre_topc(k3_topmetr)), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', fc2_topmetr)).
fof(dt_k3_pcomps_1, axiom, ![X1]:(l1_metric_1(X1)=>l1_pre_topc(k3_pcomps_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', dt_k3_pcomps_1)).
fof(dt_k7_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>m1_subset_1(k7_relset_1(X1,X2,X3,X4),k1_zfmisc_1(X2))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', dt_k7_relset_1)).
fof(rc2_subset_1, axiom, ![X1]:?[X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))&v1_xboole_0(X2)), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', rc2_subset_1)).
fof(fc4_topmetr, axiom, ![X1]:((v3_topmetr(X1)&l1_struct_0(X1))=>v3_membered(u1_struct_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', fc4_topmetr)).
fof(t11_hausdorf, axiom, ![X1]:((~(v1_xboole_0(X1))&m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))=>?[X2]:(((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(k1_numbers)))&X1=X2)&k2_weierstr(X1)=k4_seq_4(X2))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', t11_hausdorf)).
fof(dt_k3_topmetr, axiom, (v2_pre_topc(k3_topmetr)&l1_pre_topc(k3_topmetr)), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', dt_k3_topmetr)).
fof(fc2_pcomps_1, axiom, ![X1]:((~(v2_struct_0(X1))&l1_metric_1(X1))=>~(v2_struct_0(k3_pcomps_1(X1)))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', fc2_pcomps_1)).
fof(d1_seq_4, axiom, ![X1]:(v3_membered(X1)=>(v4_xxreal_2(X1)=>(v1_xboole_0(X1)|![X2]:(v1_xreal_0(X2)=>(X2=k2_seq_4(X1)<=>(![X3]:(v1_xreal_0(X3)=>(r2_hidden(X3,X1)=>r1_xxreal_0(X3,X2)))&![X3]:(v1_xreal_0(X3)=>~((~(r1_xxreal_0(X3,k6_numbers))&![X4]:(v1_xreal_0(X4)=>~((r2_hidden(X4,X1)&~(r1_xxreal_0(X4,k10_binop_2(X2,X3))))))))))))))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', d1_seq_4)).
fof(fc5_topmetr, axiom, (v2_pre_topc(k3_topmetr)&v3_topmetr(k3_topmetr)), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', fc5_topmetr)).
fof(d2_xboole_0, axiom, k1_xboole_0=o_0_0_xboole_0, file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', d2_xboole_0)).
fof(cc21_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v3_membered(X2))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', cc21_membered)).
fof(redefinition_k4_seq_4, axiom, ![X1]:(m1_subset_1(X1,k1_zfmisc_1(k1_numbers))=>k4_seq_4(X1)=k2_seq_4(X1)), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', redefinition_k4_seq_4)).
fof(dt_o_0_0_xboole_0, axiom, v1_xboole_0(o_0_0_xboole_0), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', dt_o_0_0_xboole_0)).
fof(t25_hausdorf, axiom, ![X1]:((((((~(v2_struct_0(X1))&v6_metric_1(X1))&v7_metric_1(X1))&v8_metric_1(X1))&v9_metric_1(X1))&l1_metric_1(X1))=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))=>![X3]:((~(v1_xboole_0(X3))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(k1_numbers))=>(((X4=k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3)&v2_compts_1(X2,k3_pcomps_1(X1)))&v2_compts_1(X3,k3_pcomps_1(X1)))=>v4_xxreal_2(X4)))))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', t25_hausdorf)).
fof(rc3_tbsp_1, axiom, ![X1]:(((((~(v2_struct_0(X1))&v6_metric_1(X1))&v8_metric_1(X1))&v9_metric_1(X1))&l1_metric_1(X1))=>?[X2]:((m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))&~(v1_xboole_0(X2)))&v1_finset_1(X2))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', rc3_tbsp_1)).
fof(cc14_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xreal_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', cc14_membered)).
fof(dt_k8_weierstr, axiom, ![X1, X2, X3]:((((((((~(v2_struct_0(X1))&v6_metric_1(X1))&v7_metric_1(X1))&v8_metric_1(X1))&v9_metric_1(X1))&l1_metric_1(X1))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))=>m1_subset_1(k8_weierstr(X1,X2,X3),k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', dt_k8_weierstr)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', fc3_membered)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', t2_subset)).
fof(t1_hausdorf, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>((r1_xxreal_0(k6_numbers,X1)&k4_xxreal_0(X1,X2)=k6_numbers)=>X1=k6_numbers))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', t1_hausdorf)).
fof(commutativity_k4_xxreal_0, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>k4_xxreal_0(X1,X2)=k4_xxreal_0(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', commutativity_k4_xxreal_0)).
fof(redefinition_k2_square_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&m1_subset_1(X2,k1_numbers))=>k2_square_1(X1,X2)=k4_xxreal_0(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', redefinition_k2_square_1)).
fof(d1_hausdorf, axiom, ![X1]:((((((~(v2_struct_0(X1))&v6_metric_1(X1))&v7_metric_1(X1))&v8_metric_1(X1))&v9_metric_1(X1))&l1_metric_1(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))=>k1_hausdorf(X1,X2,X3)=k2_square_1(k8_weierstr(X1,X2,X3),k8_weierstr(X1,X3,X2))))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', d1_hausdorf)).
fof(cc4_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xxreal_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', cc4_xreal_0)).
fof(t28_hausdorf, axiom, ![X1]:((((((~(v2_struct_0(X1))&v6_metric_1(X1))&v7_metric_1(X1))&v8_metric_1(X1))&v9_metric_1(X1))&l1_metric_1(X1))=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))=>![X3]:((~(v1_xboole_0(X3))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))=>((v2_compts_1(X2,k3_pcomps_1(X1))&v2_compts_1(X3,k3_pcomps_1(X1)))=>r1_xxreal_0(k6_numbers,k8_weierstr(X1,X2,X3)))))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', t28_hausdorf)).
fof(rc4_xreal_0, axiom, ?[X1]:(((v1_xboole_0(X1)&v1_xcmplx_0(X1))&v1_xxreal_0(X1))&v1_xreal_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', rc4_xreal_0)).
fof(t1_xxreal_0, axiom, ![X1]:(v1_xxreal_0(X1)=>![X2]:(v1_xxreal_0(X2)=>((r1_xxreal_0(X1,X2)&r1_xxreal_0(X2,X1))=>X1=X2))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', t1_xxreal_0)).
fof(t9_jordan1k, axiom, ![X1]:((((((~(v2_struct_0(X1))&v6_metric_1(X1))&v7_metric_1(X1))&v8_metric_1(X1))&v9_metric_1(X1))&l1_metric_1(X1))=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))=>![X3]:(m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))=>r1_xxreal_0(k6_numbers,k3_funct_2(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3))))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', t9_jordan1k)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', t4_subset)).
fof(redefinition_k3_funct_2, axiom, ![X1, X2, X3, X4]:(((((~(v1_xboole_0(X1))&v1_funct_1(X3))&v1_funct_2(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&m1_subset_1(X4,X1))=>k3_funct_2(X1,X2,X3,X4)=k1_funct_1(X3,X4)), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', redefinition_k3_funct_2)).
fof(commutativity_k1_hausdorf, axiom, ![X1, X2, X3]:((((((((~(v2_struct_0(X1))&v6_metric_1(X1))&v7_metric_1(X1))&v8_metric_1(X1))&v9_metric_1(X1))&l1_metric_1(X1))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))=>k1_hausdorf(X1,X2,X3)=k1_hausdorf(X1,X3,X2)), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', commutativity_k1_hausdorf)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', fc1_subset_1)).
fof(cc1_xreal_0, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>v1_xreal_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', cc1_xreal_0)).
fof(t9_hausdorf, axiom, ![X1]:((((((~(v2_struct_0(X1))&v6_metric_1(X1))&v7_metric_1(X1))&v8_metric_1(X1))&v9_metric_1(X1))&l1_metric_1(X1))=>![X2]:(((~(v1_xboole_0(X2))&v4_pre_topc(X2,k3_pcomps_1(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,X2)<=>k1_funct_1(k6_weierstr(X1,X2),X3)=k6_numbers)))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', t9_hausdorf)).
fof(t7_compts_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v8_pre_topc(X1)&v2_compts_1(X2,X1))=>v4_pre_topc(X2,X1)))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', t7_compts_1)).
fof(fc1_hausdorf, axiom, ![X1]:((((((~(v2_struct_0(X1))&v6_metric_1(X1))&v7_metric_1(X1))&v8_metric_1(X1))&v9_metric_1(X1))&l1_metric_1(X1))=>v8_pre_topc(k3_pcomps_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', fc1_hausdorf)).
fof(fc1_pcomps_1, axiom, ![X1]:(l1_metric_1(X1)=>(v1_pre_topc(k3_pcomps_1(X1))&v2_pre_topc(k3_pcomps_1(X1)))), file('/scratch/mptp/allpred/preds__0.005/t37_hausdorf', fc1_pcomps_1)).
fof(c_0_59, plain, ![X22, X23, X24]:(((~v1_funct_2(X24,X22,X23)|X22=k1_relset_1(X22,X24)|X23=k1_xboole_0|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(X22,X23))))&(X22!=k1_relset_1(X22,X24)|v1_funct_2(X24,X22,X23)|X23=k1_xboole_0|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(X22,X23)))))&((~v1_funct_2(X24,X22,X23)|X24=k1_xboole_0|X23!=k1_xboole_0|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(X22,X23))))&(X24!=k1_xboole_0|v1_funct_2(X24,X22,X23)|X23!=k1_xboole_0|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(X22,X23)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_2])])])).
cnf(c_0_60, plain, (X2=k1_relset_1(X2,X1)|X3=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_61, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
fof(c_0_62, plain, ![X50, X51]:(((v1_funct_1(k6_weierstr(X50,X51))|(v2_struct_0(X50)|~v6_metric_1(X50)|~v7_metric_1(X50)|~v8_metric_1(X50)|~v9_metric_1(X50)|~l1_metric_1(X50)|~m1_subset_1(X51,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X50))))))&(v1_funct_2(k6_weierstr(X50,X51),u1_struct_0(k3_pcomps_1(X50)),u1_struct_0(k3_topmetr))|(v2_struct_0(X50)|~v6_metric_1(X50)|~v7_metric_1(X50)|~v8_metric_1(X50)|~v9_metric_1(X50)|~l1_metric_1(X50)|~m1_subset_1(X51,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X50)))))))&(m1_subset_1(k6_weierstr(X50,X51),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X50)),u1_struct_0(k3_topmetr))))|(v2_struct_0(X50)|~v6_metric_1(X50)|~v7_metric_1(X50)|~v8_metric_1(X50)|~v9_metric_1(X50)|~l1_metric_1(X50)|~m1_subset_1(X51,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X50))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_weierstr])])])])).
cnf(c_0_63, plain, (X3=k1_xboole_0|k1_relset_1(X2,X1)=X2|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_60])).
cnf(c_0_64, plain, (k1_xboole_0=k6_numbers), inference(evalgc,[status(thm)],[c_0_61])).
cnf(c_0_65, plain, (v1_funct_2(k6_weierstr(X1,X2),u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr))|v2_struct_0(X1)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(split_conjunct,[status(thm)],[c_0_62])).
fof(c_0_66, plain, ![X67, X68]:(~v1_relat_1(X68)|~v4_relat_1(X68,X67)|k1_relset_1(X67,X68)=k9_xtuple_0(X68)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_67, plain, (k1_relset_1(X1,X2)=X1|X3=k6_numbers|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(rw,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_68, plain, (v2_struct_0(X1)|v1_funct_2(k6_weierstr(X1,X2),u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(evalgc,[status(thm)],[c_0_65])).
cnf(c_0_69, plain, (m1_subset_1(k6_weierstr(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr))))|v2_struct_0(X1)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(split_conjunct,[status(thm)],[c_0_62])).
fof(c_0_70, plain, ![X122, X123, X124]:(~m1_subset_1(X124,k1_zfmisc_1(k2_zfmisc_1(X122,X123)))|v1_relat_1(X124)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_71, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_72, plain, (k1_relset_1(u1_struct_0(k3_pcomps_1(X1)),k6_weierstr(X1,X2))=u1_struct_0(k3_pcomps_1(X1))|u1_struct_0(k3_topmetr)=k6_numbers|v2_struct_0(X1)|~m1_subset_1(k6_weierstr(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68])])).
cnf(c_0_73, plain, (v2_struct_0(X1)|m1_subset_1(k6_weierstr(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr))))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(evalgc,[status(thm)],[c_0_69])).
cnf(c_0_74, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_70])).
fof(c_0_75, plain, ![X166, X167, X168]:((v4_relat_1(X168,X166)|~m1_subset_1(X168,k1_zfmisc_1(k2_zfmisc_1(X166,X167))))&(v5_relat_1(X168,X167)|~m1_subset_1(X168,k1_zfmisc_1(k2_zfmisc_1(X166,X167))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
cnf(c_0_76, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_71])).
cnf(c_0_77, plain, (k1_relset_1(u1_struct_0(k3_pcomps_1(X1)),k6_weierstr(X1,X2))=u1_struct_0(k3_pcomps_1(X1))|u1_struct_0(k3_topmetr)=k6_numbers|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(csr,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_78, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_74])).
cnf(c_0_79, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_75])).
fof(c_0_80, negated_conjecture, ~(![X1]:((((((~(v2_struct_0(X1))&v6_metric_1(X1))&v7_metric_1(X1))&v8_metric_1(X1))&v9_metric_1(X1))&l1_metric_1(X1))=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))=>![X3]:((~(v1_xboole_0(X3))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))=>(((v2_compts_1(X2,k3_pcomps_1(X1))&v2_compts_1(X3,k3_pcomps_1(X1)))&k1_hausdorf(X1,X2,X3)=k6_numbers)=>X2=X3))))), inference(assume_negation,[status(cth)],[t37_hausdorf])).
fof(c_0_81, plain, ![X9, X10]:(((r1_tarski(X9,X10)|X9!=X10)&(r1_tarski(X10,X9)|X9!=X10))&(~r1_tarski(X9,X10)|~r1_tarski(X10,X9)|X9=X10)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_82, plain, ![X100, X101]:((~m1_subset_1(X100,k1_zfmisc_1(X101))|r1_tarski(X100,X101))&(~r1_tarski(X100,X101)|m1_subset_1(X100,k1_zfmisc_1(X101)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_83, plain, ![X126, X127]:(~r2_hidden(X126,X127)|~v1_xboole_0(X127)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
fof(c_0_84, plain, ![X26, X27, X28, X29, X30]:((~r1_tarski(X26,X27)|(~r2_hidden(X28,X26)|r2_hidden(X28,X27)))&((r2_hidden(esk5_2(X29,X30),X29)|r1_tarski(X29,X30))&(~r2_hidden(esk5_2(X29,X30),X30)|r1_tarski(X29,X30)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_85, plain, ![X38, X39, X40, X41, X43, X44, X45, X46, X48]:(((((r2_hidden(esk6_4(X38,X39,X40,X41),k9_xtuple_0(X38))|~r2_hidden(X41,X40)|X40!=k7_relat_1(X38,X39)|(~v1_relat_1(X38)|~v1_funct_1(X38)))&(r2_hidden(esk6_4(X38,X39,X40,X41),X39)|~r2_hidden(X41,X40)|X40!=k7_relat_1(X38,X39)|(~v1_relat_1(X38)|~v1_funct_1(X38))))&(X41=k1_funct_1(X38,esk6_4(X38,X39,X40,X41))|~r2_hidden(X41,X40)|X40!=k7_relat_1(X38,X39)|(~v1_relat_1(X38)|~v1_funct_1(X38))))&(~r2_hidden(X44,k9_xtuple_0(X38))|~r2_hidden(X44,X39)|X43!=k1_funct_1(X38,X44)|r2_hidden(X43,X40)|X40!=k7_relat_1(X38,X39)|(~v1_relat_1(X38)|~v1_funct_1(X38))))&((~r2_hidden(esk7_3(X38,X45,X46),X46)|(~r2_hidden(X48,k9_xtuple_0(X38))|~r2_hidden(X48,X45)|esk7_3(X38,X45,X46)!=k1_funct_1(X38,X48))|X46=k7_relat_1(X38,X45)|(~v1_relat_1(X38)|~v1_funct_1(X38)))&(((r2_hidden(esk8_3(X38,X45,X46),k9_xtuple_0(X38))|r2_hidden(esk7_3(X38,X45,X46),X46)|X46=k7_relat_1(X38,X45)|(~v1_relat_1(X38)|~v1_funct_1(X38)))&(r2_hidden(esk8_3(X38,X45,X46),X45)|r2_hidden(esk7_3(X38,X45,X46),X46)|X46=k7_relat_1(X38,X45)|(~v1_relat_1(X38)|~v1_funct_1(X38))))&(esk7_3(X38,X45,X46)=k1_funct_1(X38,esk8_3(X38,X45,X46))|r2_hidden(esk7_3(X38,X45,X46),X46)|X46=k7_relat_1(X38,X45)|(~v1_relat_1(X38)|~v1_funct_1(X38)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_funct_1])])])])])])).
cnf(c_0_86, plain, (k9_xtuple_0(k6_weierstr(X1,X2))=u1_struct_0(k3_pcomps_1(X1))|u1_struct_0(k3_topmetr)=k6_numbers|v2_struct_0(X1)|~v4_relat_1(k6_weierstr(X1,X2),u1_struct_0(k3_pcomps_1(X1)))|~v1_relat_1(k6_weierstr(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77])])).
cnf(c_0_87, plain, (v1_relat_1(k6_weierstr(X1,X2))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_73])])).
cnf(c_0_88, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_79])).
fof(c_0_89, plain, ![X128, X129, X130, X131]:(v1_xboole_0(X128)|v1_xboole_0(X129)|~v1_funct_1(X130)|~v1_funct_2(X130,X128,X129)|~m1_subset_1(X130,k1_zfmisc_1(k2_zfmisc_1(X128,X129)))|v1_xboole_0(X131)|~m1_subset_1(X131,k1_zfmisc_1(X128))|~v1_xboole_0(k7_relat_1(X130,X131))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_waybel_2])])])).
fof(c_0_90, plain, ![X112, X113]:(~v1_xboole_0(X112)|(~m1_subset_1(X113,k1_zfmisc_1(X112))|v1_xboole_0(X113))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])).
fof(c_0_91, plain, ![X132]:r1_tarski(X132,X132), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r1_tarski])])).
fof(c_0_92, plain, ![X172]:(v2_struct_0(X172)|~l1_struct_0(X172)|~v1_xboole_0(u1_struct_0(X172))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_93, plain, ![X133]:(~l1_pre_topc(X133)|l1_struct_0(X133)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_94, plain, ![X25]:((u1_struct_0(X25)=u1_struct_0(k3_pcomps_1(X25))|~l1_metric_1(X25))&(u1_pre_topc(k3_pcomps_1(X25))=k2_pcomps_1(X25)|~l1_metric_1(X25))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_topmetr])])])).
fof(c_0_95, negated_conjecture, ((((((~v2_struct_0(esk1_0)&v6_metric_1(esk1_0))&v7_metric_1(esk1_0))&v8_metric_1(esk1_0))&v9_metric_1(esk1_0))&l1_metric_1(esk1_0))&((~v1_xboole_0(esk2_0)&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(esk1_0)))))&((~v1_xboole_0(esk3_0)&m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(esk1_0)))))&(((v2_compts_1(esk2_0,k3_pcomps_1(esk1_0))&v2_compts_1(esk3_0,k3_pcomps_1(esk1_0)))&k1_hausdorf(esk1_0,esk2_0,esk3_0)=k6_numbers)&esk2_0!=esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_80])])])])).
cnf(c_0_96, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_97, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_98, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_99, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_84])).
fof(c_0_100, plain, ![X81]:(~v1_xboole_0(X81)|X81=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
cnf(c_0_101, plain, (r2_hidden(X4,X5)|~r2_hidden(X1,k9_xtuple_0(X2))|~r2_hidden(X1,X3)|X4!=k1_funct_1(X2,X1)|X5!=k7_relat_1(X2,X3)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_102, plain, (k9_xtuple_0(k6_weierstr(X1,X2))=u1_struct_0(k3_pcomps_1(X1))|u1_struct_0(k3_topmetr)=k6_numbers|v2_struct_0(X1)|~v4_relat_1(k6_weierstr(X1,X2),u1_struct_0(k3_pcomps_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(csr,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_103, plain, (v4_relat_1(k6_weierstr(X1,X2),u1_struct_0(k3_pcomps_1(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_73])])).
fof(c_0_104, plain, ![X52, X53, X54]:(v2_struct_0(X52)|~v6_metric_1(X52)|~v7_metric_1(X52)|~v8_metric_1(X52)|~v9_metric_1(X52)|~l1_metric_1(X52)|(~m1_subset_1(X53,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X52))))|(~m1_subset_1(X54,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X52))))|k8_weierstr(X52,X53,X54)=k2_weierstr(k7_relset_1(u1_struct_0(k3_pcomps_1(X52)),u1_struct_0(k3_topmetr),k6_weierstr(X52,X53),X54))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_weierstr])])])])).
fof(c_0_105, plain, ![X55, X56, X57, X58]:(~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(X55,X56)))|k7_relset_1(X55,X56,X57,X58)=k7_relat_1(X57,X58)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_relset_1])])).
cnf(c_0_106, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X4)|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~v1_xboole_0(k7_relat_1(X3,X4))), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_107, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_108, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_109, plain, (r1_tarski(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_110, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_111, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_93])).
fof(c_0_112, plain, ((~v2_struct_0(k3_topmetr)&v1_pre_topc(k3_topmetr))&v2_pre_topc(k3_topmetr)), inference(fof_simplification,[status(thm)],[fc2_topmetr])).
fof(c_0_113, plain, ![X114]:(~l1_metric_1(X114)|l1_pre_topc(k3_pcomps_1(X114))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_pcomps_1])])).
cnf(c_0_114, plain, (u1_struct_0(X1)=u1_struct_0(k3_pcomps_1(X1))|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_115, negated_conjecture, (l1_metric_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_116, plain, (X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_96])).
cnf(c_0_117, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_97])).
cnf(c_0_118, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_98])).
cnf(c_0_119, plain, (r1_tarski(X1,X2)|r2_hidden(esk5_2(X1,X2),X1)), inference(evalgc,[status(thm)],[c_0_99])).
fof(c_0_120, plain, ![X92, X93, X94, X95]:(~m1_subset_1(X94,k1_zfmisc_1(k2_zfmisc_1(X92,X93)))|m1_subset_1(k7_relset_1(X92,X93,X94,X95),k1_zfmisc_1(X93))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_relset_1])])).
cnf(c_0_121, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_100])).
fof(c_0_122, plain, ![X120]:(m1_subset_1(esk13_1(X120),k1_zfmisc_1(X120))&v1_xboole_0(esk13_1(X120))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc2_subset_1])])).
cnf(c_0_123, plain, (r2_hidden(X4,X5)|X5!=k7_relat_1(X2,X3)|X4!=k1_funct_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~r2_hidden(X1,X3)|~r2_hidden(X1,k9_xtuple_0(X2))), inference(evalgc,[status(thm)],[c_0_101])).
cnf(c_0_124, plain, (k9_xtuple_0(k6_weierstr(X1,X2))=u1_struct_0(k3_pcomps_1(X1))|u1_struct_0(k3_topmetr)=k6_numbers|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103])])).
cnf(c_0_125, plain, (v1_funct_1(k6_weierstr(X1,X2))|v2_struct_0(X1)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(split_conjunct,[status(thm)],[c_0_62])).
fof(c_0_126, plain, ![X218]:(~v3_topmetr(X218)|~l1_struct_0(X218)|v3_membered(u1_struct_0(X218))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc4_topmetr])])).
fof(c_0_127, plain, ![X15]:((((~v1_xboole_0(esk4_1(X15))|(v1_xboole_0(X15)|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(k3_topmetr)))))&(m1_subset_1(esk4_1(X15),k1_zfmisc_1(k1_numbers))|(v1_xboole_0(X15)|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(k3_topmetr))))))&(X15=esk4_1(X15)|(v1_xboole_0(X15)|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(k3_topmetr))))))&(k2_weierstr(X15)=k4_seq_4(esk4_1(X15))|(v1_xboole_0(X15)|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(k3_topmetr)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t11_hausdorf])])])])])).
cnf(c_0_128, plain, (v2_struct_0(X1)|k8_weierstr(X1,X2,X3)=k2_weierstr(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_129, plain, (k7_relset_1(X2,X3,X1,X4)=k7_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_130, plain, (v1_xboole_0(X4)|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~v1_xboole_0(k7_relat_1(X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_106])).
cnf(c_0_131, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_107])).
cnf(c_0_132, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_108])).
cnf(c_0_133, plain, (r1_tarski(X1,X1)), inference(evalgc,[status(thm)],[c_0_109])).
cnf(c_0_134, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_110])).
cnf(c_0_135, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[c_0_111])).
cnf(c_0_136, plain, (l1_pre_topc(k3_topmetr)), inference(split_conjunct,[status(thm)],[dt_k3_topmetr])).
cnf(c_0_137, plain, (~v2_struct_0(k3_topmetr)), inference(split_conjunct,[status(thm)],[c_0_112])).
fof(c_0_138, plain, ![X148]:(v2_struct_0(X148)|~l1_metric_1(X148)|~v2_struct_0(k3_pcomps_1(X148))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_pcomps_1])])])).
cnf(c_0_139, plain, (l1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_140, plain, (X1=k1_xboole_0|~v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_141, plain, (u1_struct_0(k3_pcomps_1(X1))=u1_struct_0(X1)|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[c_0_114])).
cnf(c_0_142, negated_conjecture, (l1_metric_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_115])).
cnf(c_0_143, negated_conjecture, (v9_metric_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_144, negated_conjecture, (v8_metric_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_145, negated_conjecture, (v7_metric_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_146, negated_conjecture, (v6_metric_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_147, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_148, plain, (X1=X2|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117])])).
cnf(c_0_149, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119])])).
cnf(c_0_150, plain, (m1_subset_1(k7_relset_1(X2,X3,X1,X4),k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_151, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_121])).
cnf(c_0_152, plain, (v1_xboole_0(esk13_1(X1))), inference(split_conjunct,[status(thm)],[c_0_122])).
fof(c_0_153, plain, ![X71, X72, X73, X74, X78]:(((~v1_xreal_0(X73)|(~r2_hidden(X73,X71)|r1_xxreal_0(X73,X72))|X72!=k2_seq_4(X71)|~v1_xreal_0(X72)|v1_xboole_0(X71)|~v4_xxreal_2(X71)|~v3_membered(X71))&((v1_xreal_0(esk9_3(X71,X72,X74))|r1_xxreal_0(X74,k6_numbers)|~v1_xreal_0(X74)|X72!=k2_seq_4(X71)|~v1_xreal_0(X72)|v1_xboole_0(X71)|~v4_xxreal_2(X71)|~v3_membered(X71))&((r2_hidden(esk9_3(X71,X72,X74),X71)|r1_xxreal_0(X74,k6_numbers)|~v1_xreal_0(X74)|X72!=k2_seq_4(X71)|~v1_xreal_0(X72)|v1_xboole_0(X71)|~v4_xxreal_2(X71)|~v3_membered(X71))&(~r1_xxreal_0(esk9_3(X71,X72,X74),k10_binop_2(X72,X74))|r1_xxreal_0(X74,k6_numbers)|~v1_xreal_0(X74)|X72!=k2_seq_4(X71)|~v1_xreal_0(X72)|v1_xboole_0(X71)|~v4_xxreal_2(X71)|~v3_membered(X71)))))&(((v1_xreal_0(esk11_2(X71,X72))|v1_xreal_0(esk10_2(X71,X72))|X72=k2_seq_4(X71)|~v1_xreal_0(X72)|v1_xboole_0(X71)|~v4_xxreal_2(X71)|~v3_membered(X71))&((~r1_xxreal_0(esk11_2(X71,X72),k6_numbers)|v1_xreal_0(esk10_2(X71,X72))|X72=k2_seq_4(X71)|~v1_xreal_0(X72)|v1_xboole_0(X71)|~v4_xxreal_2(X71)|~v3_membered(X71))&(~v1_xreal_0(X78)|(~r2_hidden(X78,X71)|r1_xxreal_0(X78,k10_binop_2(X72,esk11_2(X71,X72))))|v1_xreal_0(esk10_2(X71,X72))|X72=k2_seq_4(X71)|~v1_xreal_0(X72)|v1_xboole_0(X71)|~v4_xxreal_2(X71)|~v3_membered(X71))))&(((v1_xreal_0(esk11_2(X71,X72))|r2_hidden(esk10_2(X71,X72),X71)|X72=k2_seq_4(X71)|~v1_xreal_0(X72)|v1_xboole_0(X71)|~v4_xxreal_2(X71)|~v3_membered(X71))&((~r1_xxreal_0(esk11_2(X71,X72),k6_numbers)|r2_hidden(esk10_2(X71,X72),X71)|X72=k2_seq_4(X71)|~v1_xreal_0(X72)|v1_xboole_0(X71)|~v4_xxreal_2(X71)|~v3_membered(X71))&(~v1_xreal_0(X78)|(~r2_hidden(X78,X71)|r1_xxreal_0(X78,k10_binop_2(X72,esk11_2(X71,X72))))|r2_hidden(esk10_2(X71,X72),X71)|X72=k2_seq_4(X71)|~v1_xreal_0(X72)|v1_xboole_0(X71)|~v4_xxreal_2(X71)|~v3_membered(X71))))&((v1_xreal_0(esk11_2(X71,X72))|~r1_xxreal_0(esk10_2(X71,X72),X72)|X72=k2_seq_4(X71)|~v1_xreal_0(X72)|v1_xboole_0(X71)|~v4_xxreal_2(X71)|~v3_membered(X71))&((~r1_xxreal_0(esk11_2(X71,X72),k6_numbers)|~r1_xxreal_0(esk10_2(X71,X72),X72)|X72=k2_seq_4(X71)|~v1_xreal_0(X72)|v1_xboole_0(X71)|~v4_xxreal_2(X71)|~v3_membered(X71))&(~v1_xreal_0(X78)|(~r2_hidden(X78,X71)|r1_xxreal_0(X78,k10_binop_2(X72,esk11_2(X71,X72))))|~r1_xxreal_0(esk10_2(X71,X72),X72)|X72=k2_seq_4(X71)|~v1_xreal_0(X72)|v1_xboole_0(X71)|~v4_xxreal_2(X71)|~v3_membered(X71))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_seq_4])])])])])])).
cnf(c_0_154, plain, (r2_hidden(X1,X2)|u1_struct_0(k3_topmetr)=k6_numbers|v2_struct_0(X3)|X2!=k7_relat_1(k6_weierstr(X3,X4),X5)|X1!=k1_funct_1(k6_weierstr(X3,X4),X6)|~v1_funct_1(k6_weierstr(X3,X4))|~v1_relat_1(k6_weierstr(X3,X4))|~r2_hidden(X6,u1_struct_0(k3_pcomps_1(X3)))|~r2_hidden(X6,X5)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_124])])).
cnf(c_0_155, plain, (v2_struct_0(X1)|v1_funct_1(k6_weierstr(X1,X2))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(evalgc,[status(thm)],[c_0_125])).
cnf(c_0_156, plain, (v3_membered(u1_struct_0(X1))|~v3_topmetr(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_126])).
cnf(c_0_157, plain, (v3_topmetr(k3_topmetr)), inference(split_conjunct,[status(thm)],[fc5_topmetr])).
cnf(c_0_158, plain, (k2_weierstr(X1)=k4_seq_4(esk4_1(X1))|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr)))), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_159, plain, (X1=esk4_1(X1)|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr)))), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_160, plain, (k2_weierstr(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3))=k8_weierstr(X1,X2,X3)|v2_struct_0(X1)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(evalgc,[status(thm)],[c_0_128])).
cnf(c_0_161, plain, (k7_relset_1(X2,X3,X1,X4)=k7_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_129])).
cnf(c_0_162, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,X4,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X1)))|~m1_subset_1(X2,k1_zfmisc_1(X4))|~v1_xboole_0(k7_relat_1(X3,X2))), inference(csr,[status(thm)],[c_0_130, c_0_131])).
cnf(c_0_163, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_133])])).
cnf(c_0_164, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135])])).
cnf(c_0_165, plain, (l1_pre_topc(k3_topmetr)), inference(evalgc,[status(thm)],[c_0_136])).
cnf(c_0_166, plain, (~v2_struct_0(k3_topmetr)), inference(evalgc,[status(thm)],[c_0_137])).
cnf(c_0_167, plain, (v2_struct_0(X1)|~l1_metric_1(X1)|~v2_struct_0(k3_pcomps_1(X1))), inference(split_conjunct,[status(thm)],[c_0_138])).
cnf(c_0_168, plain, (l1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[c_0_139])).
cnf(c_0_169, plain, (X1=k1_xboole_0|X3!=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_140])).
cnf(c_0_170, negated_conjecture, (u1_struct_0(k3_pcomps_1(esk1_0))=u1_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_142])])).
cnf(c_0_171, negated_conjecture, (v9_metric_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_143])).
cnf(c_0_172, negated_conjecture, (v8_metric_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_144])).
cnf(c_0_173, negated_conjecture, (v7_metric_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_145])).
cnf(c_0_174, negated_conjecture, (v6_metric_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_146])).
cnf(c_0_175, negated_conjecture, (~v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_147])).
cnf(c_0_176, plain, (X1=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_149])])).
cnf(c_0_177, plain, (m1_subset_1(k7_relset_1(X2,X3,X1,X4),k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_150])).
cnf(c_0_178, plain, (m1_subset_1(esk13_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_122])).
cnf(c_0_179, plain, (X1=k6_numbers|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_151, c_0_64])).
cnf(c_0_180, plain, (v1_xboole_0(esk13_1(X1))), inference(evalgc,[status(thm)],[c_0_152])).
cnf(c_0_181, plain, (k1_xboole_0=o_0_0_xboole_0), inference(split_conjunct,[status(thm)],[d2_xboole_0])).
cnf(c_0_182, plain, (r1_xxreal_0(X1,X3)|v1_xboole_0(X2)|~v1_xreal_0(X1)|~r2_hidden(X1,X2)|X3!=k2_seq_4(X2)|~v1_xreal_0(X3)|~v4_xxreal_2(X2)|~v3_membered(X2)), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_183, plain, (u1_struct_0(k3_topmetr)=k6_numbers|r2_hidden(X1,X2)|v2_struct_0(X3)|X2!=k7_relat_1(k6_weierstr(X3,X4),X5)|X1!=k1_funct_1(k6_weierstr(X3,X4),X6)|~r2_hidden(X6,u1_struct_0(k3_pcomps_1(X3)))|~r2_hidden(X6,X5)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_154, c_0_87]), c_0_155])).
fof(c_0_184, plain, ![X178, X179]:(~v3_membered(X178)|(~m1_subset_1(X179,k1_zfmisc_1(X178))|v3_membered(X179))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc21_membered])])])).
cnf(c_0_185, plain, (v3_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~v3_topmetr(X1)), inference(evalgc,[status(thm)],[c_0_156])).
cnf(c_0_186, plain, (v3_topmetr(k3_topmetr)), inference(evalgc,[status(thm)],[c_0_157])).
fof(c_0_187, plain, ![X153]:(~m1_subset_1(X153,k1_zfmisc_1(k1_numbers))|k4_seq_4(X153)=k2_seq_4(X153)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_seq_4])])).
cnf(c_0_188, plain, (k4_seq_4(esk4_1(X1))=k2_weierstr(X1)|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr)))), inference(evalgc,[status(thm)],[c_0_158])).
cnf(c_0_189, plain, (esk4_1(X1)=X1|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr)))), inference(evalgc,[status(thm)],[c_0_159])).
cnf(c_0_190, plain, (k2_weierstr(k7_relat_1(k6_weierstr(X1,X2),X3))=k8_weierstr(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(k6_weierstr(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_161])])).
cnf(c_0_191, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v1_xboole_0(k7_relat_1(X3,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_163])])).
cnf(c_0_192, plain, (~v1_xboole_0(u1_struct_0(k3_topmetr))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_165]), c_0_166])])).
cnf(c_0_193, plain, (v2_struct_0(X1)|~l1_metric_1(X1)|~v2_struct_0(k3_pcomps_1(X1))), inference(evalgc,[status(thm)],[c_0_167])).
cnf(c_0_194, plain, (v2_struct_0(k3_pcomps_1(X1))|~v1_xboole_0(u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_168])])).
cnf(c_0_195, plain, (X1=k6_numbers|X2!=k6_numbers|~v1_funct_2(X1,X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_169, c_0_64]), c_0_64])).
cnf(c_0_196, negated_conjecture, (v1_funct_2(k6_weierstr(esk1_0,X1),u1_struct_0(esk1_0),u1_struct_0(k3_topmetr))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_170]), c_0_142]), c_0_171]), c_0_172]), c_0_173]), c_0_174])]), c_0_175])])).
cnf(c_0_197, plain, (k7_relset_1(X2,X1,X3,X4)=X1|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_177])])).
cnf(c_0_198, plain, (m1_subset_1(esk13_1(X1),k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_178])).
cnf(c_0_199, plain, (esk13_1(X1)=k6_numbers), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_180])])).
cnf(c_0_200, plain, (v1_xboole_0(o_0_0_xboole_0)), inference(split_conjunct,[status(thm)],[dt_o_0_0_xboole_0])).
cnf(c_0_201, plain, (o_0_0_xboole_0=k1_xboole_0), inference(evalgc,[status(thm)],[c_0_181])).
cnf(c_0_202, plain, (v1_xboole_0(X2)|r1_xxreal_0(X1,X3)|X3!=k2_seq_4(X2)|~v4_xxreal_2(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v3_membered(X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_182])).
cnf(c_0_203, plain, (u1_struct_0(k3_topmetr)=k6_numbers|r2_hidden(X1,k7_relat_1(k6_weierstr(X2,X3),X4))|v2_struct_0(X2)|X1!=k1_funct_1(k6_weierstr(X2,X3),X5)|~r2_hidden(X5,u1_struct_0(k3_pcomps_1(X2)))|~r2_hidden(X5,X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_183])])).
cnf(c_0_204, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_184])).
cnf(c_0_205, plain, (m1_subset_1(k7_relat_1(X3,X4),k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_161])])).
cnf(c_0_206, plain, (v3_membered(u1_struct_0(k3_topmetr))|~l1_struct_0(k3_topmetr)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_186])])).
cnf(c_0_207, plain, (k4_seq_4(X1)=k2_seq_4(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_187])).
cnf(c_0_208, plain, (k2_weierstr(X1)=k4_seq_4(X1)|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_189])])).
cnf(c_0_209, plain, (k2_weierstr(k7_relat_1(k6_weierstr(X1,X2),X3))=k8_weierstr(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(csr,[status(thm)],[c_0_190, c_0_73])).
cnf(c_0_210, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(k1_numbers))|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr)))), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_211, plain, (v1_xboole_0(u1_struct_0(k3_pcomps_1(X1)))|v2_struct_0(X1)|~v1_funct_1(k6_weierstr(X1,X2))|~m1_subset_1(k6_weierstr(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr))))|~v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_68]), c_0_192])])).
cnf(c_0_212, plain, (v2_struct_0(X1)|~l1_metric_1(X1)|~v1_xboole_0(u1_struct_0(k3_pcomps_1(X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_194])])).
cnf(c_0_213, negated_conjecture, (k6_weierstr(esk1_0,X1)=k6_numbers|u1_struct_0(k3_topmetr)!=k6_numbers|~m1_subset_1(k6_weierstr(esk1_0,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_196])])).
cnf(c_0_214, negated_conjecture, (m1_subset_1(k6_weierstr(esk1_0,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_170]), c_0_142]), c_0_171]), c_0_172]), c_0_173]), c_0_174])]), c_0_175])])).
cnf(c_0_215, plain, (X2=k7_relat_1(X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_197])])).
cnf(c_0_216, plain, (m1_subset_1(k6_numbers,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_198, c_0_199])])).
cnf(c_0_217, plain, (v1_xboole_0(o_0_0_xboole_0)), inference(evalgc,[status(thm)],[c_0_200])).
cnf(c_0_218, plain, (o_0_0_xboole_0=k6_numbers), inference(rw,[status(thm)],[c_0_201, c_0_64])).
fof(c_0_219, plain, ![X11, X12, X13, X14]:(v2_struct_0(X11)|~v6_metric_1(X11)|~v7_metric_1(X11)|~v8_metric_1(X11)|~v9_metric_1(X11)|~l1_metric_1(X11)|(v1_xboole_0(X12)|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X11))))|(v1_xboole_0(X13)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X11))))|(~m1_subset_1(X14,k1_zfmisc_1(k1_numbers))|(X14!=k7_relset_1(u1_struct_0(k3_pcomps_1(X11)),u1_struct_0(k3_topmetr),k6_weierstr(X11,X12),X13)|~v2_compts_1(X12,k3_pcomps_1(X11))|~v2_compts_1(X13,k3_pcomps_1(X11))|v4_xxreal_2(X14)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t25_hausdorf])])])])).
cnf(c_0_220, plain, (r1_xxreal_0(X1,X2)|X2!=k2_seq_4(X3)|~v3_membered(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r2_hidden(X1,X3)|~v4_xxreal_2(X3)), inference(csr,[status(thm)],[c_0_202, c_0_118])).
cnf(c_0_221, plain, (u1_struct_0(k3_topmetr)=k6_numbers|r2_hidden(k1_funct_1(k6_weierstr(X1,X2),X3),k7_relat_1(k6_weierstr(X1,X2),X4))|v2_struct_0(X1)|~r2_hidden(X3,u1_struct_0(k3_pcomps_1(X1)))|~r2_hidden(X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_203])])).
cnf(c_0_222, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_204])).
cnf(c_0_223, plain, (m1_subset_1(k7_relat_1(k6_weierstr(X1,X2),X3),k1_zfmisc_1(u1_struct_0(k3_topmetr)))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_73])])).
cnf(c_0_224, plain, (v3_membered(u1_struct_0(k3_topmetr))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_135]), c_0_165])])])).
cnf(c_0_225, plain, (k4_seq_4(X1)=k2_seq_4(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_207])).
cnf(c_0_226, plain, (k4_seq_4(k7_relat_1(k6_weierstr(X1,X2),X3))=k8_weierstr(X1,X2,X3)|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X3))|v2_struct_0(X1)|~m1_subset_1(k7_relat_1(k6_weierstr(X1,X2),X3),k1_zfmisc_1(u1_struct_0(k3_topmetr)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_209])])).
cnf(c_0_227, plain, (v1_xboole_0(X1)|m1_subset_1(esk4_1(X1),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr)))), inference(evalgc,[status(thm)],[c_0_210])).
cnf(c_0_228, plain, (v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_211, c_0_73]), c_0_155]), c_0_212])).
cnf(c_0_229, negated_conjecture, (k6_weierstr(esk1_0,X1)=k6_numbers|u1_struct_0(k3_topmetr)!=k6_numbers|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[c_0_213, c_0_214])).
cnf(c_0_230, plain, (X1=k7_relat_1(k6_numbers,X2)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_216])])).
cnf(c_0_231, plain, (v1_xboole_0(k6_numbers)), inference(rw,[status(thm)],[c_0_217, c_0_218])).
fof(c_0_232, plain, ![X268]:(((m1_subset_1(esk28_1(X268),k1_zfmisc_1(u1_struct_0(X268)))|(v2_struct_0(X268)|~v6_metric_1(X268)|~v8_metric_1(X268)|~v9_metric_1(X268)|~l1_metric_1(X268)))&(~v1_xboole_0(esk28_1(X268))|(v2_struct_0(X268)|~v6_metric_1(X268)|~v8_metric_1(X268)|~v9_metric_1(X268)|~l1_metric_1(X268))))&(v1_finset_1(esk28_1(X268))|(v2_struct_0(X268)|~v6_metric_1(X268)|~v8_metric_1(X268)|~v9_metric_1(X268)|~l1_metric_1(X268)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[rc3_tbsp_1])])])])])).
cnf(c_0_233, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|v4_xxreal_2(X4)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|X4!=k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3)|~v2_compts_1(X2,k3_pcomps_1(X1))|~v2_compts_1(X3,k3_pcomps_1(X1))), inference(split_conjunct,[status(thm)],[c_0_219])).
cnf(c_0_234, plain, (r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),X4)|u1_struct_0(k3_topmetr)=k6_numbers|v2_struct_0(X1)|X4!=k2_seq_4(k7_relat_1(k6_weierstr(X1,X2),X5))|~v3_membered(k7_relat_1(k6_weierstr(X1,X2),X5))|~v1_xreal_0(X4)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~v4_xxreal_2(k7_relat_1(k6_weierstr(X1,X2),X5))|~r2_hidden(X3,u1_struct_0(k3_pcomps_1(X1)))|~r2_hidden(X3,X5)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_221])])).
cnf(c_0_235, plain, (v3_membered(k7_relat_1(k6_weierstr(X1,X2),X3))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_223]), c_0_224])])])).
cnf(c_0_236, plain, (k2_seq_4(k7_relat_1(k6_weierstr(X1,X2),X3))=k8_weierstr(X1,X2,X3)|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X3))|v2_struct_0(X1)|~m1_subset_1(k7_relat_1(k6_weierstr(X1,X2),X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k7_relat_1(k6_weierstr(X1,X2),X3),k1_zfmisc_1(u1_struct_0(k3_topmetr)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_226])])).
cnf(c_0_237, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_189])])).
cnf(c_0_238, negated_conjecture, (~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_xboole_0(k7_relat_1(k6_numbers,u1_struct_0(esk1_0)))|u1_struct_0(k3_topmetr)!=k6_numbers), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_229]), c_0_170]), c_0_170]), c_0_142]), c_0_171]), c_0_172]), c_0_173]), c_0_174])]), c_0_175])])).
cnf(c_0_239, plain, (k7_relat_1(k6_numbers,X1)=k6_numbers), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_231])])).
cnf(c_0_240, plain, (m1_subset_1(esk28_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~v6_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_232])).
cnf(c_0_241, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_242, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|v4_xxreal_2(X4)|X4!=k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(evalgc,[status(thm)],[c_0_233])).
cnf(c_0_243, plain, (u1_struct_0(k3_topmetr)=k6_numbers|r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),X4)|v2_struct_0(X1)|X4!=k2_seq_4(k7_relat_1(k6_weierstr(X1,X2),X5))|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~v1_xreal_0(X4)|~r2_hidden(X3,u1_struct_0(k3_pcomps_1(X1)))|~r2_hidden(X3,X5)|~v4_xxreal_2(k7_relat_1(k6_weierstr(X1,X2),X5))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(csr,[status(thm)],[c_0_234, c_0_235])).
cnf(c_0_244, plain, (k2_seq_4(k7_relat_1(k6_weierstr(X1,X2),X3))=k8_weierstr(X1,X2,X3)|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X3))|v2_struct_0(X1)|~m1_subset_1(k7_relat_1(k6_weierstr(X1,X2),X3),k1_zfmisc_1(u1_struct_0(k3_topmetr)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(csr,[status(thm)],[c_0_236, c_0_237])).
cnf(c_0_245, negated_conjecture, (u1_struct_0(k3_topmetr)!=k6_numbers|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_238, c_0_239]), c_0_231])])])).
cnf(c_0_246, plain, (v2_struct_0(X1)|m1_subset_1(esk28_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~v6_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[c_0_240])).
cnf(c_0_247, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(evalgc,[status(thm)],[c_0_241])).
cnf(c_0_248, plain, (v4_xxreal_2(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|v2_struct_0(X4)|X1!=k7_relat_1(k6_weierstr(X4,X2),X3)|~v2_compts_1(X3,k3_pcomps_1(X4))|~v2_compts_1(X2,k3_pcomps_1(X4))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~l1_metric_1(X4)|~v9_metric_1(X4)|~v8_metric_1(X4)|~v7_metric_1(X4)|~v6_metric_1(X4)|~m1_subset_1(k6_weierstr(X4,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X4)),u1_struct_0(k3_topmetr))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_161])])).
cnf(c_0_249, plain, (u1_struct_0(k3_topmetr)=k6_numbers|r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),X4)|v2_struct_0(X1)|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X5))|X4!=k8_weierstr(X1,X2,X5)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~v1_xreal_0(X4)|~r2_hidden(X3,u1_struct_0(k3_pcomps_1(X1)))|~r2_hidden(X3,X5)|~v4_xxreal_2(k7_relat_1(k6_weierstr(X1,X2),X5))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(k7_relat_1(k6_weierstr(X1,X2),X5),k1_zfmisc_1(u1_struct_0(k3_topmetr)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_244])])).
cnf(c_0_250, negated_conjecture, (u1_struct_0(k3_topmetr)!=k6_numbers), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_246]), c_0_142]), c_0_171]), c_0_172]), c_0_174])]), c_0_175])])).
cnf(c_0_251, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_117])])).
cnf(c_0_252, plain, (v4_xxreal_2(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|v2_struct_0(X4)|X1!=k7_relat_1(k6_weierstr(X4,X3),X2)|~v2_compts_1(X2,k3_pcomps_1(X4))|~v2_compts_1(X3,k3_pcomps_1(X4))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~l1_metric_1(X4)|~v9_metric_1(X4)|~v8_metric_1(X4)|~v7_metric_1(X4)|~v6_metric_1(X4)), inference(csr,[status(thm)],[c_0_248, c_0_73])).
cnf(c_0_253, plain, (r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),X4)|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X5))|v2_struct_0(X1)|X4!=k8_weierstr(X1,X2,X5)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~v1_xreal_0(X4)|~r2_hidden(X3,X5)|~v4_xxreal_2(k7_relat_1(k6_weierstr(X1,X2),X5))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[c_0_249, c_0_250]), c_0_223]), c_0_251])).
cnf(c_0_254, plain, (v4_xxreal_2(k7_relat_1(k6_weierstr(X1,X2),X3))|v1_xboole_0(X3)|v1_xboole_0(X2)|v2_struct_0(X1)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(k7_relat_1(k6_weierstr(X1,X2),X3),k1_zfmisc_1(k1_numbers))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_252])])).
fof(c_0_255, plain, ![X142, X143]:(~v3_membered(X142)|(~m1_subset_1(X143,X142)|v1_xreal_0(X143))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc14_membered])])])).
fof(c_0_256, plain, ![X89, X90, X91]:(v2_struct_0(X89)|~v6_metric_1(X89)|~v7_metric_1(X89)|~v8_metric_1(X89)|~v9_metric_1(X89)|~l1_metric_1(X89)|~m1_subset_1(X90,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X89))))|~m1_subset_1(X91,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X89))))|m1_subset_1(k8_weierstr(X89,X90,X91),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k8_weierstr])])])).
cnf(c_0_257, plain, (r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),X4)|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X5))|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X5)|X4!=k8_weierstr(X1,X2,X5)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~v1_xreal_0(X4)|~r2_hidden(X3,X5)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_compts_1(X5,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(k7_relat_1(k6_weierstr(X1,X2),X5),k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_254])])).
cnf(c_0_258, plain, (m1_subset_1(k7_relat_1(k6_weierstr(X1,X2),X3),k1_zfmisc_1(k1_numbers))|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X3))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_223])])).
cnf(c_0_259, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_255])).
cnf(c_0_260, plain, (v2_struct_0(X1)|m1_subset_1(k8_weierstr(X1,X2,X3),k1_numbers)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(split_conjunct,[status(thm)],[c_0_256])).
cnf(c_0_261, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
cnf(c_0_262, plain, (r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),X4)|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X5))|v1_xboole_0(X2)|v2_struct_0(X1)|X4!=k8_weierstr(X1,X2,X5)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~v1_xreal_0(X4)|~r2_hidden(X3,X5)|~v2_compts_1(X5,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_257, c_0_258]), c_0_118])).
cnf(c_0_263, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_259])).
cnf(c_0_264, plain, (v2_struct_0(X1)|m1_subset_1(k8_weierstr(X1,X2,X3),k1_numbers)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(evalgc,[status(thm)],[c_0_260])).
cnf(c_0_265, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_261])).
fof(c_0_266, plain, ![X185, X186]:(~m1_subset_1(X185,X186)|(v1_xboole_0(X186)|r2_hidden(X185,X186))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_267, plain, ![X65, X66]:(~v1_xreal_0(X65)|(~v1_xreal_0(X66)|(~r1_xxreal_0(k6_numbers,X65)|k4_xxreal_0(X65,X66)!=k6_numbers|X65=k6_numbers))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_hausdorf])])])).
fof(c_0_268, plain, ![X82, X83]:(~v1_xxreal_0(X82)|~v1_xxreal_0(X83)|k4_xxreal_0(X82,X83)=k4_xxreal_0(X83,X82)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k4_xxreal_0])])).
fof(c_0_269, plain, ![X69, X70]:(~m1_subset_1(X69,k1_numbers)|~m1_subset_1(X70,k1_numbers)|k2_square_1(X69,X70)=k4_xxreal_0(X69,X70)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_square_1])])).
fof(c_0_270, plain, ![X59, X60, X61]:(v2_struct_0(X59)|~v6_metric_1(X59)|~v7_metric_1(X59)|~v8_metric_1(X59)|~v9_metric_1(X59)|~l1_metric_1(X59)|(~m1_subset_1(X60,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X59))))|(~m1_subset_1(X61,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X59))))|k1_hausdorf(X59,X60,X61)=k2_square_1(k8_weierstr(X59,X60,X61),k8_weierstr(X59,X61,X60))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_hausdorf])])])])).
cnf(c_0_271, plain, (r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),k8_weierstr(X1,X2,X4))|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X4))|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~v1_xreal_0(k8_weierstr(X1,X2,X4))|~r2_hidden(X3,X4)|~v2_compts_1(X4,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_262])])).
cnf(c_0_272, plain, (v1_xreal_0(k8_weierstr(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_264]), c_0_265])])])).
cnf(c_0_273, plain, (X1=k1_funct_1(X2,esk6_4(X2,X3,X4,X1))|~r2_hidden(X1,X4)|X4!=k7_relat_1(X2,X3)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_274, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_266])).
cnf(c_0_275, plain, (X1=k6_numbers|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k6_numbers,X1)|k4_xxreal_0(X1,X2)!=k6_numbers), inference(split_conjunct,[status(thm)],[c_0_267])).
cnf(c_0_276, plain, (k4_xxreal_0(X1,X2)=k4_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_268])).
fof(c_0_277, plain, ![X159]:(~v1_xreal_0(X159)|v1_xxreal_0(X159)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_xreal_0])])).
cnf(c_0_278, plain, (k2_square_1(X1,X2)=k4_xxreal_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_269])).
cnf(c_0_279, plain, (v2_struct_0(X1)|k1_hausdorf(X1,X2,X3)=k2_square_1(k8_weierstr(X1,X2,X3),k8_weierstr(X1,X3,X2))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(split_conjunct,[status(thm)],[c_0_270])).
cnf(c_0_280, plain, (r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),k8_weierstr(X1,X2,X4))|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X4))|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~r2_hidden(X3,X4)|~v2_compts_1(X4,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(csr,[status(thm)],[c_0_271, c_0_272])).
cnf(c_0_281, plain, (k1_funct_1(X2,esk6_4(X2,X3,X4,X1))=X1|X4!=k7_relat_1(X2,X3)|~v1_relat_1(X2)|~v1_funct_1(X2)|~r2_hidden(X1,X4)), inference(evalgc,[status(thm)],[c_0_273])).
cnf(c_0_282, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_274])).
cnf(c_0_283, plain, (X1=k6_numbers|k4_xxreal_0(X1,X2)!=k6_numbers|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k6_numbers,X1)), inference(evalgc,[status(thm)],[c_0_275])).
cnf(c_0_284, plain, (k4_xxreal_0(X1,X2)=k4_xxreal_0(X2,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_276])).
cnf(c_0_285, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_277])).
cnf(c_0_286, plain, (k2_square_1(X1,X2)=k4_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_278])).
cnf(c_0_287, plain, (k2_square_1(k8_weierstr(X1,X2,X3),k8_weierstr(X1,X3,X2))=k1_hausdorf(X1,X2,X3)|v2_struct_0(X1)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(evalgc,[status(thm)],[c_0_279])).
cnf(c_0_288, plain, (r1_xxreal_0(X5,k8_weierstr(X1,X2,X6))|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X6))|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_xreal_0(X5)|~r2_hidden(esk6_4(k6_weierstr(X1,X2),X3,X4,X5),X6)|~v2_compts_1(X6,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|X4!=k7_relat_1(k6_weierstr(X1,X2),X3)|~v1_funct_1(k6_weierstr(X1,X2))|~v1_relat_1(k6_weierstr(X1,X2))|~r2_hidden(X5,X4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_281])])).
cnf(c_0_289, plain, (r2_hidden(esk6_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|X3!=k7_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_290, plain, (r2_hidden(X1,X2)|v1_xboole_0(k9_xtuple_0(X3))|X2!=k7_relat_1(X3,X4)|X1!=k1_funct_1(X3,X5)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(X5,X4)|~m1_subset_1(X5,k9_xtuple_0(X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_282])])).
cnf(c_0_291, plain, (X1=k6_numbers|k4_xxreal_0(X2,X1)!=k6_numbers|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k6_numbers,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_284])])).
cnf(c_0_292, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_285])).
cnf(c_0_293, plain, (k4_xxreal_0(k8_weierstr(X1,X2,X3),k8_weierstr(X1,X3,X2))=k1_hausdorf(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(k8_weierstr(X1,X3,X2),k1_numbers)|~m1_subset_1(k8_weierstr(X1,X2,X3),k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_287])])).
cnf(c_0_294, plain, (r1_xxreal_0(X1,k8_weierstr(X2,X3,X4))|v1_xboole_0(k7_relat_1(k6_weierstr(X2,X3),X4))|v1_xboole_0(X3)|v2_struct_0(X2)|X5!=k7_relat_1(k6_weierstr(X2,X3),X6)|~v1_xreal_0(X1)|~r2_hidden(esk6_4(k6_weierstr(X2,X3),X6,X5,X1),X4)|~r2_hidden(X1,X5)|~v2_compts_1(X4,k3_pcomps_1(X2))|~v2_compts_1(X3,k3_pcomps_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_288, c_0_87]), c_0_155])).
cnf(c_0_295, plain, (r2_hidden(esk6_4(X1,X2,X3,X4),X2)|X3!=k7_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X4,X3)), inference(evalgc,[status(thm)],[c_0_289])).
cnf(c_0_296, plain, (r2_hidden(X1,X2)|v1_xboole_0(u1_struct_0(k3_pcomps_1(X3)))|u1_struct_0(k3_topmetr)=k6_numbers|v2_struct_0(X3)|X2!=k7_relat_1(k6_weierstr(X3,X4),X5)|X1!=k1_funct_1(k6_weierstr(X3,X4),X6)|~v1_funct_1(k6_weierstr(X3,X4))|~v1_relat_1(k6_weierstr(X3,X4))|~r2_hidden(X6,X5)|~m1_subset_1(X6,u1_struct_0(k3_pcomps_1(X3)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_124])])).
cnf(c_0_297, plain, (X1=k6_numbers|k4_xxreal_0(X2,X1)!=k6_numbers|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k6_numbers,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_291, c_0_292]), c_0_292])).
cnf(c_0_298, plain, (k4_xxreal_0(k8_weierstr(X1,X2,X3),k8_weierstr(X1,X3,X2))=k1_hausdorf(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_293, c_0_264]), c_0_264])).
fof(c_0_299, plain, ![X32, X33, X34]:(v2_struct_0(X32)|~v6_metric_1(X32)|~v7_metric_1(X32)|~v8_metric_1(X32)|~v9_metric_1(X32)|~l1_metric_1(X32)|(v1_xboole_0(X33)|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X32))))|(v1_xboole_0(X34)|~m1_subset_1(X34,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X32))))|(~v2_compts_1(X33,k3_pcomps_1(X32))|~v2_compts_1(X34,k3_pcomps_1(X32))|r1_xxreal_0(k6_numbers,k8_weierstr(X32,X33,X34)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t28_hausdorf])])])])).
fof(c_0_300, plain, (((v1_xboole_0(esk12_0)&v1_xcmplx_0(esk12_0))&v1_xxreal_0(esk12_0))&v1_xreal_0(esk12_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc4_xreal_0])])).
cnf(c_0_301, plain, (r1_xxreal_0(X1,k8_weierstr(X2,X3,X4))|v1_xboole_0(k7_relat_1(k6_weierstr(X2,X3),X4))|v1_xboole_0(X3)|v2_struct_0(X2)|X5!=k7_relat_1(k6_weierstr(X2,X3),X4)|~v1_xreal_0(X1)|~r2_hidden(X1,X5)|~v2_compts_1(X4,k3_pcomps_1(X2))|~v2_compts_1(X3,k3_pcomps_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v1_funct_1(k6_weierstr(X2,X3))|~v1_relat_1(k6_weierstr(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_294, c_0_295])])).
cnf(c_0_302, plain, (u1_struct_0(k3_topmetr)=k6_numbers|r2_hidden(X1,X2)|v2_struct_0(X3)|X2!=k7_relat_1(k6_weierstr(X3,X4),X5)|X1!=k1_funct_1(k6_weierstr(X3,X4),X6)|~r2_hidden(X6,X5)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X6,u1_struct_0(k3_pcomps_1(X3)))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_296, c_0_87]), c_0_155]), c_0_212])).
cnf(c_0_303, plain, (k8_weierstr(X1,X2,X3)=k6_numbers|v2_struct_0(X1)|k1_hausdorf(X1,X3,X2)!=k6_numbers|~v1_xreal_0(k8_weierstr(X1,X3,X2))|~v1_xreal_0(k8_weierstr(X1,X2,X3))|~r1_xxreal_0(k6_numbers,k8_weierstr(X1,X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_297, c_0_298])])).
cnf(c_0_304, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|r1_xxreal_0(k6_numbers,k8_weierstr(X1,X2,X3))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~v2_compts_1(X2,k3_pcomps_1(X1))|~v2_compts_1(X3,k3_pcomps_1(X1))), inference(split_conjunct,[status(thm)],[c_0_299])).
fof(c_0_305, plain, ![X20, X21]:(~v1_xxreal_0(X20)|(~v1_xxreal_0(X21)|(~r1_xxreal_0(X20,X21)|~r1_xxreal_0(X21,X20)|X20=X21))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xxreal_0])])])).
fof(c_0_306, plain, ![X35, X36, X37]:(v2_struct_0(X35)|~v6_metric_1(X35)|~v7_metric_1(X35)|~v8_metric_1(X35)|~v9_metric_1(X35)|~l1_metric_1(X35)|(v1_xboole_0(X36)|~m1_subset_1(X36,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X35))))|(~m1_subset_1(X37,u1_struct_0(k3_pcomps_1(X35)))|r1_xxreal_0(k6_numbers,k3_funct_2(u1_struct_0(k3_pcomps_1(X35)),u1_struct_0(k3_topmetr),k6_weierstr(X35,X36),X37))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_jordan1k])])])])).
cnf(c_0_307, plain, (v1_xboole_0(esk12_0)), inference(split_conjunct,[status(thm)],[c_0_300])).
cnf(c_0_308, plain, (r1_xxreal_0(X1,k8_weierstr(X2,X3,X4))|v1_xboole_0(k7_relat_1(k6_weierstr(X2,X3),X4))|v1_xboole_0(X3)|v2_struct_0(X2)|X5!=k7_relat_1(k6_weierstr(X2,X3),X4)|~v1_xreal_0(X1)|~r2_hidden(X1,X5)|~v2_compts_1(X4,k3_pcomps_1(X2))|~v2_compts_1(X3,k3_pcomps_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_301, c_0_87]), c_0_155])).
cnf(c_0_309, plain, (u1_struct_0(k3_topmetr)=k6_numbers|r2_hidden(X1,k7_relat_1(k6_weierstr(X2,X3),X4))|v2_struct_0(X2)|X1!=k1_funct_1(k6_weierstr(X2,X3),X5)|~r2_hidden(X5,X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X5,u1_struct_0(k3_pcomps_1(X2)))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_302])])).
cnf(c_0_310, plain, (k8_weierstr(X1,X2,X3)=k6_numbers|v2_struct_0(X1)|k1_hausdorf(X1,X3,X2)!=k6_numbers|~r1_xxreal_0(k6_numbers,k8_weierstr(X1,X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_303, c_0_272]), c_0_272])).
cnf(c_0_311, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|r1_xxreal_0(k6_numbers,k8_weierstr(X1,X2,X3))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(evalgc,[status(thm)],[c_0_304])).
cnf(c_0_312, plain, (k8_weierstr(X1,X2,X3)=k6_numbers|v2_struct_0(X1)|k1_hausdorf(X1,X2,X3)!=k6_numbers|~v1_xreal_0(k8_weierstr(X1,X3,X2))|~v1_xreal_0(k8_weierstr(X1,X2,X3))|~r1_xxreal_0(k6_numbers,k8_weierstr(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_298])])).
cnf(c_0_313, plain, (X1=X2|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_305])).
cnf(c_0_314, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|r1_xxreal_0(k6_numbers,k3_funct_2(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_306])).
cnf(c_0_315, plain, (v1_xxreal_0(esk12_0)), inference(split_conjunct,[status(thm)],[c_0_300])).
cnf(c_0_316, plain, (v1_xboole_0(esk12_0)), inference(evalgc,[status(thm)],[c_0_307])).
cnf(c_0_317, plain, (r1_xxreal_0(X1,k8_weierstr(X2,X3,X4))|v1_xboole_0(k7_relat_1(k6_weierstr(X2,X3),X4))|v1_xboole_0(X3)|v2_struct_0(X2)|~v1_xreal_0(X1)|~r2_hidden(X1,k7_relat_1(k6_weierstr(X2,X3),X4))|~v2_compts_1(X4,k3_pcomps_1(X2))|~v2_compts_1(X3,k3_pcomps_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_308])])).
cnf(c_0_318, plain, (u1_struct_0(k3_topmetr)=k6_numbers|r2_hidden(k1_funct_1(k6_weierstr(X1,X2),X3),k7_relat_1(k6_weierstr(X1,X2),X4))|v2_struct_0(X1)|~r2_hidden(X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_309])])).
fof(c_0_319, plain, ![X103, X104, X105]:(~r2_hidden(X103,X104)|~m1_subset_1(X104,k1_zfmisc_1(X105))|m1_subset_1(X103,X105)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_320, plain, (k8_weierstr(X1,X2,X3)=k6_numbers|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|k1_hausdorf(X1,X3,X2)!=k6_numbers|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_311])])).
cnf(c_0_321, plain, (k8_weierstr(X1,X2,X3)=k6_numbers|v2_struct_0(X1)|k1_hausdorf(X1,X2,X3)!=k6_numbers|~r1_xxreal_0(k6_numbers,k8_weierstr(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_312, c_0_272]), c_0_272])).
cnf(c_0_322, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_323, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_324, plain, (X1=X2|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_313])).
cnf(c_0_325, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|r1_xxreal_0(k6_numbers,k3_funct_2(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(evalgc,[status(thm)],[c_0_314])).
cnf(c_0_326, plain, (v1_xxreal_0(esk12_0)), inference(evalgc,[status(thm)],[c_0_315])).
cnf(c_0_327, plain, (esk12_0=k6_numbers), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_316])])).
fof(c_0_328, plain, ![X96, X97, X98, X99]:(v1_xboole_0(X96)|~v1_funct_1(X98)|~v1_funct_2(X98,X96,X97)|~m1_subset_1(X98,k1_zfmisc_1(k2_zfmisc_1(X96,X97)))|~m1_subset_1(X99,X96)|k3_funct_2(X96,X97,X98,X99)=k1_funct_1(X98,X99)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_funct_2])])])).
cnf(c_0_329, plain, (r1_xxreal_0(X1,k8_weierstr(X2,X3,X4))|v1_xboole_0(X3)|v2_struct_0(X2)|~v1_xreal_0(X1)|~r2_hidden(X1,k7_relat_1(k6_weierstr(X2,X3),X4))|~v2_compts_1(X4,k3_pcomps_1(X2))|~v2_compts_1(X3,k3_pcomps_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)), inference(csr,[status(thm)],[c_0_317, c_0_118])).
cnf(c_0_330, plain, (r2_hidden(k1_funct_1(k6_weierstr(X1,X2),X3),k7_relat_1(k6_weierstr(X1,X2),X4))|v2_struct_0(X1)|~r2_hidden(X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(sr,[status(thm)],[c_0_318, c_0_250])).
cnf(c_0_331, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_319])).
fof(c_0_332, plain, ![X62, X63, X64]:(v2_struct_0(X62)|~v6_metric_1(X62)|~v7_metric_1(X62)|~v8_metric_1(X62)|~v9_metric_1(X62)|~l1_metric_1(X62)|~m1_subset_1(X63,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X62))))|~m1_subset_1(X64,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X62))))|k1_hausdorf(X62,X63,X64)=k1_hausdorf(X62,X64,X63)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[commutativity_k1_hausdorf])])])).
cnf(c_0_333, plain, (k2_square_1(k6_numbers,k8_weierstr(X1,X3,X2))=k1_hausdorf(X1,X2,X3)|v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X3,X2)!=k6_numbers|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_287, c_0_320])])).
cnf(c_0_334, plain, (k8_weierstr(X1,X2,X3)=k6_numbers|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|k1_hausdorf(X1,X2,X3)!=k6_numbers|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_321, c_0_311])])).
cnf(c_0_335, negated_conjecture, (k1_hausdorf(esk1_0,esk2_0,esk3_0)=k6_numbers), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_336, negated_conjecture, (v2_compts_1(esk2_0,k3_pcomps_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_337, negated_conjecture, (v2_compts_1(esk3_0,k3_pcomps_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_338, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(esk1_0))))), inference(evalgc,[status(thm)],[c_0_322])).
cnf(c_0_339, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(esk1_0))))), inference(evalgc,[status(thm)],[c_0_323])).
cnf(c_0_340, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_341, negated_conjecture, (~v1_xboole_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_342, plain, (k3_funct_2(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3)=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|~r1_xxreal_0(k3_funct_2(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3),k6_numbers)|~v1_xxreal_0(k6_numbers)|~v1_xxreal_0(k3_funct_2(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_324, c_0_325])])).
cnf(c_0_343, plain, (v1_xxreal_0(k6_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_326, c_0_327])])).
cnf(c_0_344, plain, (v1_xboole_0(X1)|k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|~v1_funct_1(X2)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~m1_subset_1(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_328])).
cnf(c_0_345, plain, (r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),k8_weierstr(X1,X2,X4))|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~v2_compts_1(X4,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_330])])).
cnf(c_0_346, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_331])).
cnf(c_0_347, plain, (v2_struct_0(X1)|k1_hausdorf(X1,X2,X3)=k1_hausdorf(X1,X3,X2)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(split_conjunct,[status(thm)],[c_0_332])).
cnf(c_0_348, plain, (k1_hausdorf(X1,X3,X2)=k2_square_1(k6_numbers,k6_numbers)|v1_xboole_0(X3)|v1_xboole_0(X2)|v2_struct_0(X1)|k1_hausdorf(X1,X2,X3)!=k6_numbers|~v2_compts_1(X2,k3_pcomps_1(X1))|~v2_compts_1(X3,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_333, c_0_334])])).
cnf(c_0_349, negated_conjecture, (k1_hausdorf(esk1_0,esk2_0,esk3_0)=k6_numbers), inference(evalgc,[status(thm)],[c_0_335])).
cnf(c_0_350, negated_conjecture, (v2_compts_1(esk2_0,k3_pcomps_1(esk1_0))), inference(evalgc,[status(thm)],[c_0_336])).
cnf(c_0_351, negated_conjecture, (v2_compts_1(esk3_0,k3_pcomps_1(esk1_0))), inference(evalgc,[status(thm)],[c_0_337])).
cnf(c_0_352, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_338, c_0_170])])).
cnf(c_0_353, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_339, c_0_170])])).
cnf(c_0_354, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(evalgc,[status(thm)],[c_0_340])).
cnf(c_0_355, negated_conjecture, (~v1_xboole_0(esk3_0)), inference(evalgc,[status(thm)],[c_0_341])).
cnf(c_0_356, plain, (k3_funct_2(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3)=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|~r1_xxreal_0(k3_funct_2(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3),k6_numbers)|~v1_xxreal_0(k3_funct_2(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_342, c_0_343])])).
cnf(c_0_357, plain, (k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(evalgc,[status(thm)],[c_0_344])).
cnf(c_0_358, plain, (r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),k8_weierstr(X1,X2,X4))|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~r2_hidden(X3,X4)|~v2_compts_1(X4,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(csr,[status(thm)],[c_0_345, c_0_346])).
cnf(c_0_359, plain, (k1_hausdorf(X1,X2,X3)=k1_hausdorf(X1,X3,X2)|v2_struct_0(X1)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(evalgc,[status(thm)],[c_0_347])).
cnf(c_0_360, negated_conjecture, (k1_hausdorf(esk1_0,esk3_0,esk2_0)=k2_square_1(k6_numbers,k6_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_348, c_0_349]), c_0_350]), c_0_351]), c_0_170]), c_0_352]), c_0_170]), c_0_353]), c_0_142]), c_0_171]), c_0_172]), c_0_173]), c_0_174])]), c_0_354]), c_0_355]), c_0_175])])).
cnf(c_0_361, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k7_relset_1(X3,X2,X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_177])])).
cnf(c_0_362, plain, (k1_funct_1(k6_weierstr(X1,X2),X3)=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(k3_pcomps_1(X1)))|~r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),k6_numbers)|~v1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v1_funct_1(k6_weierstr(X1,X2))|~v1_funct_2(k6_weierstr(X1,X2),u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr))|~m1_subset_1(k6_weierstr(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_357])])).
cnf(c_0_363, plain, (r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),k6_numbers)|v1_xboole_0(X2)|v2_struct_0(X1)|v1_xboole_0(X4)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~r2_hidden(X3,X4)|~v2_compts_1(X4,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X2,X4)!=k6_numbers), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_358, c_0_334])])).
cnf(c_0_364, negated_conjecture, (k2_square_1(k6_numbers,k6_numbers)=k6_numbers), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_359, c_0_360]), c_0_349]), c_0_170]), c_0_352]), c_0_170]), c_0_353]), c_0_142]), c_0_171]), c_0_172]), c_0_173]), c_0_174])]), c_0_175])])).
cnf(c_0_365, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k7_relat_1(X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_161])])).
cnf(c_0_366, plain, (k1_funct_1(k6_weierstr(X1,X2),X3)=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|~r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),k6_numbers)|~v1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_362, c_0_73]), c_0_68]), c_0_155]), c_0_131])).
cnf(c_0_367, plain, (r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),k6_numbers)|v1_xboole_0(X2)|v2_struct_0(X1)|k1_hausdorf(X1,X2,X4)!=k6_numbers|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~r2_hidden(X3,X4)|~v2_compts_1(X4,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(csr,[status(thm)],[c_0_363, c_0_118])).
cnf(c_0_368, negated_conjecture, (k1_hausdorf(esk1_0,esk3_0,esk2_0)=k6_numbers), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_360, c_0_364])])).
cnf(c_0_369, plain, (r2_hidden(X1,u1_struct_0(k3_topmetr))|v2_struct_0(X2)|~r2_hidden(X1,k7_relat_1(k6_weierstr(X2,X3),X4))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_365, c_0_73])])).
fof(c_0_370, plain, ![X219]:~v1_xboole_0(k1_zfmisc_1(X219)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
cnf(c_0_371, plain, (k1_funct_1(k6_weierstr(X1,X2),X3)=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|~r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),k6_numbers)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_292])])).
cnf(c_0_372, negated_conjecture, (r1_xxreal_0(k1_funct_1(k6_weierstr(esk1_0,esk3_0),X1),k6_numbers)|~v1_xreal_0(k1_funct_1(k6_weierstr(esk1_0,esk3_0),X1))|~r2_hidden(X1,esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_368]), c_0_350]), c_0_351]), c_0_170]), c_0_353]), c_0_170]), c_0_352]), c_0_142]), c_0_171]), c_0_172]), c_0_173]), c_0_174])]), c_0_355]), c_0_175])])).
fof(c_0_373, plain, ![X125]:(~m1_subset_1(X125,k1_numbers)|v1_xreal_0(X125)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_xreal_0])])).
cnf(c_0_374, plain, (m1_subset_1(X1,k1_numbers)|v1_xboole_0(X2)|~r2_hidden(X1,esk4_1(X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_topmetr)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_346, c_0_227])])).
cnf(c_0_375, plain, (r2_hidden(k1_funct_1(k6_weierstr(X1,X2),X3),u1_struct_0(k3_topmetr))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_369, c_0_330])])).
cnf(c_0_376, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_370])).
cnf(c_0_377, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,esk3_0),X1)=k6_numbers|~v1_xreal_0(k1_funct_1(k6_weierstr(esk1_0,esk3_0),X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_371, c_0_372]), c_0_170]), c_0_353]), c_0_170]), c_0_142]), c_0_171]), c_0_172]), c_0_173]), c_0_174])]), c_0_355]), c_0_175])])).
cnf(c_0_378, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_346, c_0_352])])).
cnf(c_0_379, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_373])).
cnf(c_0_380, plain, (m1_subset_1(X1,k1_numbers)|v1_xboole_0(X2)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_topmetr)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_189])])).
cnf(c_0_381, plain, (r2_hidden(k1_funct_1(k6_weierstr(X1,X2),X3),u1_struct_0(k3_topmetr))|v2_struct_0(X1)|v1_xboole_0(X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(X3,X4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_375, c_0_282])])).
cnf(c_0_382, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_376])).
cnf(c_0_383, negated_conjecture, (r1_xxreal_0(k1_funct_1(k6_weierstr(esk1_0,esk2_0),X1),k6_numbers)|~v1_xreal_0(k1_funct_1(k6_weierstr(esk1_0,esk2_0),X1))|~r2_hidden(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_349]), c_0_351]), c_0_350]), c_0_170]), c_0_352]), c_0_170]), c_0_353]), c_0_142]), c_0_171]), c_0_172]), c_0_173]), c_0_174])]), c_0_354]), c_0_175])])).
cnf(c_0_384, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,esk3_0),X1)=k6_numbers|~v1_xreal_0(k1_funct_1(k6_weierstr(esk1_0,esk3_0),X1))|~r2_hidden(X1,esk2_0)), inference(csr,[status(thm)],[c_0_377, c_0_378])).
cnf(c_0_385, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_379])).
cnf(c_0_386, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_topmetr)))), inference(csr,[status(thm)],[c_0_380, c_0_118])).
cnf(c_0_387, plain, (r2_hidden(k1_funct_1(k6_weierstr(X1,X2),X3),u1_struct_0(k3_topmetr))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_381, c_0_163]), c_0_382])])).
cnf(c_0_388, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,esk2_0),X1)=k6_numbers|~v1_xreal_0(k1_funct_1(k6_weierstr(esk1_0,esk2_0),X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_371, c_0_383]), c_0_170]), c_0_352]), c_0_170]), c_0_142]), c_0_171]), c_0_172]), c_0_173]), c_0_174])]), c_0_354]), c_0_175])])).
cnf(c_0_389, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_346, c_0_353])])).
fof(c_0_390, plain, ![X17, X18, X19]:((~r2_hidden(X19,X18)|k1_funct_1(k6_weierstr(X17,X18),X19)=k6_numbers|~m1_subset_1(X19,u1_struct_0(X17))|(v1_xboole_0(X18)|~v4_pre_topc(X18,k3_pcomps_1(X17))|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X17)))))|(v2_struct_0(X17)|~v6_metric_1(X17)|~v7_metric_1(X17)|~v8_metric_1(X17)|~v9_metric_1(X17)|~l1_metric_1(X17)))&(k1_funct_1(k6_weierstr(X17,X18),X19)!=k6_numbers|r2_hidden(X19,X18)|~m1_subset_1(X19,u1_struct_0(X17))|(v1_xboole_0(X18)|~v4_pre_topc(X18,k3_pcomps_1(X17))|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X17)))))|(v2_struct_0(X17)|~v6_metric_1(X17)|~v7_metric_1(X17)|~v8_metric_1(X17)|~v9_metric_1(X17)|~l1_metric_1(X17)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_hausdorf])])])])])).
cnf(c_0_391, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,esk3_0),X1)=k6_numbers|~r2_hidden(X1,esk2_0)|~m1_subset_1(k1_funct_1(k6_weierstr(esk1_0,esk3_0),X1),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_384, c_0_385])])).
cnf(c_0_392, plain, (m1_subset_1(k1_funct_1(k6_weierstr(X1,X2),X3),k1_numbers)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_386, c_0_387]), c_0_163])])])).
fof(c_0_393, plain, ![X195, X196]:(~v2_pre_topc(X195)|~l1_pre_topc(X195)|(~m1_subset_1(X196,k1_zfmisc_1(u1_struct_0(X195)))|(~v8_pre_topc(X195)|~v2_compts_1(X196,X195)|v4_pre_topc(X196,X195)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_compts_1])])])).
fof(c_0_394, plain, ![X211]:(v2_struct_0(X211)|~v6_metric_1(X211)|~v7_metric_1(X211)|~v8_metric_1(X211)|~v9_metric_1(X211)|~l1_metric_1(X211)|v8_pre_topc(k3_pcomps_1(X211))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_hausdorf])])])).
cnf(c_0_395, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,esk2_0),X1)=k6_numbers|~v1_xreal_0(k1_funct_1(k6_weierstr(esk1_0,esk2_0),X1))|~r2_hidden(X1,esk3_0)), inference(csr,[status(thm)],[c_0_388, c_0_389])).
cnf(c_0_396, plain, (r2_hidden(X3,X2)|v1_xboole_0(X2)|v2_struct_0(X1)|k1_funct_1(k6_weierstr(X1,X2),X3)!=k6_numbers|~m1_subset_1(X3,u1_struct_0(X1))|~v4_pre_topc(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_390])).
cnf(c_0_397, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,esk3_0),X1)=k6_numbers|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_392]), c_0_170]), c_0_353]), c_0_170]), c_0_142]), c_0_171]), c_0_172]), c_0_173]), c_0_174])]), c_0_175])])).
cnf(c_0_398, plain, (v4_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v8_pre_topc(X1)|~v2_compts_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_393])).
cnf(c_0_399, plain, (v2_struct_0(X1)|v8_pre_topc(k3_pcomps_1(X1))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_394])).
fof(c_0_400, plain, ![X108]:((v1_pre_topc(k3_pcomps_1(X108))|~l1_metric_1(X108))&(v2_pre_topc(k3_pcomps_1(X108))|~l1_metric_1(X108))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_pcomps_1])])])).
cnf(c_0_401, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,esk2_0),X1)=k6_numbers|~r2_hidden(X1,esk3_0)|~m1_subset_1(k1_funct_1(k6_weierstr(esk1_0,esk2_0),X1),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_385])])).
cnf(c_0_402, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|r2_hidden(X3,X2)|k1_funct_1(k6_weierstr(X1,X2),X3)!=k6_numbers|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v4_pre_topc(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(evalgc,[status(thm)],[c_0_396])).
cnf(c_0_403, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,esk3_0),X1)=k6_numbers|~r2_hidden(X1,esk2_0)), inference(csr,[status(thm)],[c_0_397, c_0_378])).
cnf(c_0_404, plain, (v4_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~v2_compts_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[c_0_398])).
cnf(c_0_405, plain, (v2_struct_0(X1)|v8_pre_topc(k3_pcomps_1(X1))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[c_0_399])).
cnf(c_0_406, plain, (v2_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_400])).
cnf(c_0_407, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,esk2_0),X1)=k6_numbers|~r2_hidden(X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_401, c_0_392]), c_0_170]), c_0_352]), c_0_170]), c_0_142]), c_0_171]), c_0_172]), c_0_173]), c_0_174])]), c_0_175])])).
cnf(c_0_408, negated_conjecture, (r2_hidden(X1,esk3_0)|~v4_pre_topc(esk3_0,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_403]), c_0_170]), c_0_353]), c_0_142]), c_0_171]), c_0_172]), c_0_173]), c_0_174])]), c_0_355]), c_0_175])])).
cnf(c_0_409, plain, (v4_pre_topc(X1,k3_pcomps_1(X2))|v2_struct_0(X2)|~v2_pre_topc(k3_pcomps_1(X2))|~l1_pre_topc(k3_pcomps_1(X2))|~v2_compts_1(X1,k3_pcomps_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_404, c_0_405])])).
cnf(c_0_410, plain, (v2_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[c_0_406])).
cnf(c_0_411, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,esk2_0),X1)=k6_numbers|~r2_hidden(X1,esk3_0)), inference(csr,[status(thm)],[c_0_407, c_0_389])).
cnf(c_0_412, plain, (r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_413, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk2_0)|~v4_pre_topc(esk3_0,k3_pcomps_1(esk1_0))), inference(csr,[status(thm)],[c_0_408, c_0_378])).
cnf(c_0_414, plain, (v4_pre_topc(X1,k3_pcomps_1(X2))|v2_struct_0(X2)|~v2_compts_1(X1,k3_pcomps_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_409, c_0_168]), c_0_410])).
cnf(c_0_415, negated_conjecture, (r2_hidden(X1,esk2_0)|~v4_pre_topc(esk2_0,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_411]), c_0_170]), c_0_352]), c_0_142]), c_0_171]), c_0_172]), c_0_173]), c_0_174])]), c_0_354]), c_0_175])])).
cnf(c_0_416, plain, (r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(evalgc,[status(thm)],[c_0_412])).
cnf(c_0_417, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_413, c_0_414]), c_0_351]), c_0_170]), c_0_353]), c_0_142]), c_0_171]), c_0_172]), c_0_173]), c_0_174])]), c_0_175])])).
cnf(c_0_418, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk3_0)|~v4_pre_topc(esk2_0,k3_pcomps_1(esk1_0))), inference(csr,[status(thm)],[c_0_415, c_0_389])).
cnf(c_0_419, negated_conjecture, (r1_tarski(X1,esk3_0)|~r2_hidden(esk5_2(X1,esk3_0),esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_416, c_0_417])])).
cnf(c_0_420, negated_conjecture, (esk2_0!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_421, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_418, c_0_414]), c_0_350]), c_0_170]), c_0_352]), c_0_142]), c_0_171]), c_0_172]), c_0_173]), c_0_174])]), c_0_175])])).
cnf(c_0_422, negated_conjecture, (r1_tarski(esk2_0,esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_419, c_0_119])])).
cnf(c_0_423, negated_conjecture, (esk3_0!=esk2_0), inference(evalgc,[status(thm)],[c_0_420])).
cnf(c_0_424, negated_conjecture, (r2_hidden(esk5_2(esk3_0,X1),esk2_0)|r1_tarski(esk3_0,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_421, c_0_119])])).
cnf(c_0_425, negated_conjecture, (~r1_tarski(esk3_0,esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_422]), c_0_423])])).
cnf(c_0_426, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_416, c_0_424]), c_0_425]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 427
# Proof object clause steps            : 314
# Proof object formula steps           : 113
# Proof object conjectures             : 69
# Proof object clause conjectures      : 66
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 84
# Proof object initial formulas used   : 59
# Proof object generating inferences   : 100
# Proof object simplifying inferences  : 247
# Training examples: 189 positive, 1124 negative
# Training: Positive examples begin
cnf(c_0_64, plain, (k1_xboole_0=k6_numbers)).# trainpos
cnf(c_0_201, plain, (o_0_0_xboole_0=k1_xboole_0)).# trainpos
cnf(c_0_423, negated_conjecture, (esk3_0!=esk2_0)).# trainpos
cnf(c_0_174, negated_conjecture, (v6_metric_1(esk1_0))).# trainpos
cnf(c_0_173, negated_conjecture, (v7_metric_1(esk1_0))).# trainpos
cnf(c_0_172, negated_conjecture, (v8_metric_1(esk1_0))).# trainpos
cnf(c_0_171, negated_conjecture, (v9_metric_1(esk1_0))).# trainpos
cnf(c_0_142, negated_conjecture, (l1_metric_1(esk1_0))).# trainpos
cnf(c_0_175, negated_conjecture, (~v2_struct_0(esk1_0))).# trainpos
cnf(c_0_354, negated_conjecture, (~v1_xboole_0(esk2_0))).# trainpos
cnf(c_0_217, plain, (v1_xboole_0(o_0_0_xboole_0))).# trainpos
cnf(c_0_283, plain, (X1=k6_numbers|k4_xxreal_0(X1,X2)!=k6_numbers|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k6_numbers,X1))).# trainpos
cnf(c_0_316, plain, (v1_xboole_0(esk12_0))).# trainpos
cnf(c_0_166, plain, (~v2_struct_0(k3_topmetr))).# trainpos
cnf(c_0_355, negated_conjecture, (~v1_xboole_0(esk3_0))).# trainpos
cnf(c_0_326, plain, (v1_xxreal_0(esk12_0))).# trainpos
cnf(c_0_265, plain, (v3_membered(k1_numbers))).# trainpos
cnf(c_0_165, plain, (l1_pre_topc(k3_topmetr))).# trainpos
cnf(c_0_350, negated_conjecture, (v2_compts_1(esk2_0,k3_pcomps_1(esk1_0)))).# trainpos
cnf(c_0_351, negated_conjecture, (v2_compts_1(esk3_0,k3_pcomps_1(esk1_0)))).# trainpos
cnf(c_0_292, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_186, plain, (v3_topmetr(k3_topmetr))).# trainpos
cnf(c_0_180, plain, (v1_xboole_0(esk13_1(X1)))).# trainpos
cnf(c_0_133, plain, (r1_tarski(X1,X1))).# trainpos
cnf(c_0_382, plain, (~v1_xboole_0(k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_151, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_349, negated_conjecture, (k1_hausdorf(esk1_0,esk2_0,esk3_0)=k6_numbers)).# trainpos
cnf(c_0_338, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(esk1_0)))))).# trainpos
cnf(c_0_339, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(esk1_0)))))).# trainpos
cnf(c_0_118, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_135, plain, (l1_struct_0(X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_198, plain, (m1_subset_1(esk13_1(X1),k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_141, plain, (u1_struct_0(k3_pcomps_1(X1))=u1_struct_0(X1)|~l1_metric_1(X1))).# trainpos
cnf(c_0_385, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_168, plain, (l1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1))).# trainpos
cnf(c_0_410, plain, (v2_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1))).# trainpos
cnf(c_0_225, plain, (k4_seq_4(X1)=k2_seq_4(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_132, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_117, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_193, plain, (v2_struct_0(X1)|~l1_metric_1(X1)|~v2_struct_0(k3_pcomps_1(X1)))).# trainpos
cnf(c_0_134, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).# trainpos
cnf(c_0_185, plain, (v3_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~v3_topmetr(X1))).# trainpos
cnf(c_0_282, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_405, plain, (v2_struct_0(X1)|v8_pre_topc(k3_pcomps_1(X1))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1))).# trainpos
cnf(c_0_263, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_116, plain, (X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_119, plain, (r1_tarski(X1,X2)|r2_hidden(esk5_2(X1,X2),X1))).# trainpos
cnf(c_0_416, plain, (r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2))).# trainpos
cnf(c_0_78, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_131, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_222, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_346, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_227, plain, (v1_xboole_0(X1)|m1_subset_1(esk4_1(X1),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).# trainpos
cnf(c_0_284, plain, (k4_xxreal_0(X1,X2)=k4_xxreal_0(X2,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_88, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_76, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).# trainpos
cnf(c_0_189, plain, (esk4_1(X1)=X1|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).# trainpos
cnf(c_0_188, plain, (k4_seq_4(esk4_1(X1))=k2_weierstr(X1)|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).# trainpos
cnf(c_0_286, plain, (k2_square_1(X1,X2)=k4_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_247, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1))).# trainpos
cnf(c_0_246, plain, (v2_struct_0(X1)|m1_subset_1(esk28_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~v6_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1))).# trainpos
cnf(c_0_169, plain, (X1=k1_xboole_0|X3!=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_324, plain, (X1=X2|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_404, plain, (v4_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~v2_compts_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_202, plain, (v1_xboole_0(X2)|r1_xxreal_0(X1,X3)|X3!=k2_seq_4(X2)|~v4_xxreal_2(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v3_membered(X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_161, plain, (k7_relset_1(X2,X3,X1,X4)=k7_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_177, plain, (m1_subset_1(k7_relset_1(X2,X3,X1,X4),k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_63, plain, (X3=k1_xboole_0|k1_relset_1(X2,X1)=X2|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_295, plain, (r2_hidden(esk6_4(X1,X2,X3,X4),X2)|X3!=k7_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X4,X3))).# trainpos
cnf(c_0_281, plain, (k1_funct_1(X2,esk6_4(X2,X3,X4,X1))=X1|X4!=k7_relat_1(X2,X3)|~v1_relat_1(X2)|~v1_funct_1(X2)|~r2_hidden(X1,X4))).# trainpos
cnf(c_0_359, plain, (k1_hausdorf(X1,X2,X3)=k1_hausdorf(X1,X3,X2)|v2_struct_0(X1)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))).# trainpos
cnf(c_0_155, plain, (v2_struct_0(X1)|v1_funct_1(k6_weierstr(X1,X2))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))).# trainpos
cnf(c_0_402, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|r2_hidden(X3,X2)|k1_funct_1(k6_weierstr(X1,X2),X3)!=k6_numbers|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v4_pre_topc(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))).# trainpos
cnf(c_0_123, plain, (r2_hidden(X4,X5)|X5!=k7_relat_1(X2,X3)|X4!=k1_funct_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~r2_hidden(X1,X3)|~r2_hidden(X1,k9_xtuple_0(X2)))).# trainpos
cnf(c_0_357, plain, (k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).# trainpos
cnf(c_0_130, plain, (v1_xboole_0(X4)|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~v1_xboole_0(k7_relat_1(X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_68, plain, (v2_struct_0(X1)|v1_funct_2(k6_weierstr(X1,X2),u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))).# trainpos
cnf(c_0_73, plain, (v2_struct_0(X1)|m1_subset_1(k6_weierstr(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr))))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))).# trainpos
cnf(c_0_264, plain, (v2_struct_0(X1)|m1_subset_1(k8_weierstr(X1,X2,X3),k1_numbers)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))).# trainpos
cnf(c_0_311, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|r1_xxreal_0(k6_numbers,k8_weierstr(X1,X2,X3))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))).# trainpos
cnf(c_0_287, plain, (k2_square_1(k8_weierstr(X1,X2,X3),k8_weierstr(X1,X3,X2))=k1_hausdorf(X1,X2,X3)|v2_struct_0(X1)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))).# trainpos
cnf(c_0_325, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|r1_xxreal_0(k6_numbers,k3_funct_2(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))).# trainpos
cnf(c_0_160, plain, (k2_weierstr(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3))=k8_weierstr(X1,X2,X3)|v2_struct_0(X1)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))).# trainpos
cnf(c_0_242, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|v4_xxreal_2(X4)|X4!=k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))).# trainpos
cnf(c_0_170, negated_conjecture, (u1_struct_0(k3_pcomps_1(esk1_0))=u1_struct_0(esk1_0))).# trainpos
cnf(c_0_208, plain, (k2_weierstr(X1)=k4_seq_4(X1)|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).# trainpos
cnf(c_0_251, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_149, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1))).# trainpos
cnf(c_0_327, plain, (esk12_0=k6_numbers)).# trainpos
cnf(c_0_237, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).# trainpos
cnf(c_0_163, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_353, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_199, plain, (esk13_1(X1)=k6_numbers)).# trainpos
cnf(c_0_343, plain, (v1_xxreal_0(k6_numbers))).# trainpos
cnf(c_0_352, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_409, plain, (v4_pre_topc(X1,k3_pcomps_1(X2))|v2_struct_0(X2)|~v2_pre_topc(k3_pcomps_1(X2))|~l1_pre_topc(k3_pcomps_1(X2))|~v2_compts_1(X1,k3_pcomps_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).# trainpos
cnf(c_0_248, plain, (v4_xxreal_2(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|v2_struct_0(X4)|X1!=k7_relat_1(k6_weierstr(X4,X2),X3)|~v2_compts_1(X3,k3_pcomps_1(X4))|~v2_compts_1(X2,k3_pcomps_1(X4))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~l1_metric_1(X4)|~v9_metric_1(X4)|~v8_metric_1(X4)|~v7_metric_1(X4)|~v6_metric_1(X4)|~m1_subset_1(k6_weierstr(X4,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X4)),u1_struct_0(k3_topmetr)))))).# trainpos
cnf(c_0_272, plain, (v1_xreal_0(k8_weierstr(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_190, plain, (k2_weierstr(k7_relat_1(k6_weierstr(X1,X2),X3))=k8_weierstr(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(k6_weierstr(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr)))))).# trainpos
cnf(c_0_196, negated_conjecture, (v1_funct_2(k6_weierstr(esk1_0,X1),u1_struct_0(esk1_0),u1_struct_0(k3_topmetr))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_164, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1))).# trainpos
cnf(c_0_214, negated_conjecture, (m1_subset_1(k6_weierstr(esk1_0,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_87, plain, (v1_relat_1(k6_weierstr(X1,X2))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_293, plain, (k4_xxreal_0(k8_weierstr(X1,X2,X3),k8_weierstr(X1,X3,X2))=k1_hausdorf(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(k8_weierstr(X1,X3,X2),k1_numbers)|~m1_subset_1(k8_weierstr(X1,X2,X3),k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_291, plain, (X1=k6_numbers|k4_xxreal_0(X2,X1)!=k6_numbers|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k6_numbers,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2))).# trainpos
cnf(c_0_206, plain, (v3_membered(u1_struct_0(k3_topmetr))|~l1_struct_0(k3_topmetr))).# trainpos
cnf(c_0_290, plain, (r2_hidden(X1,X2)|v1_xboole_0(k9_xtuple_0(X3))|X2!=k7_relat_1(X3,X4)|X1!=k1_funct_1(X3,X5)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(X5,X4)|~m1_subset_1(X5,k9_xtuple_0(X3)))).# trainpos
cnf(c_0_342, plain, (k3_funct_2(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3)=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|~r1_xxreal_0(k3_funct_2(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3),k6_numbers)|~v1_xxreal_0(k6_numbers)|~v1_xxreal_0(k3_funct_2(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_192, plain, (~v1_xboole_0(u1_struct_0(k3_topmetr)))).# trainpos
cnf(c_0_378, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0))).# trainpos
cnf(c_0_216, plain, (m1_subset_1(k6_numbers,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_254, plain, (v4_xxreal_2(k7_relat_1(k6_weierstr(X1,X2),X3))|v1_xboole_0(X3)|v1_xboole_0(X2)|v2_struct_0(X1)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(k7_relat_1(k6_weierstr(X1,X2),X3),k1_zfmisc_1(k1_numbers))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_226, plain, (k4_seq_4(k7_relat_1(k6_weierstr(X1,X2),X3))=k8_weierstr(X1,X2,X3)|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X3))|v2_struct_0(X1)|~m1_subset_1(k7_relat_1(k6_weierstr(X1,X2),X3),k1_zfmisc_1(u1_struct_0(k3_topmetr)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_224, plain, (v3_membered(u1_struct_0(k3_topmetr)))).# trainpos
cnf(c_0_194, plain, (v2_struct_0(k3_pcomps_1(X1))|~v1_xboole_0(u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1))).# trainpos
cnf(c_0_362, plain, (k1_funct_1(k6_weierstr(X1,X2),X3)=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(k3_pcomps_1(X1)))|~r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),k6_numbers)|~v1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v1_funct_1(k6_weierstr(X1,X2))|~v1_funct_2(k6_weierstr(X1,X2),u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr))|~m1_subset_1(k6_weierstr(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr)))))).# trainpos
cnf(c_0_213, negated_conjecture, (k6_weierstr(esk1_0,X1)=k6_numbers|u1_struct_0(k3_topmetr)!=k6_numbers|~m1_subset_1(k6_weierstr(esk1_0,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_72, plain, (k1_relset_1(u1_struct_0(k3_pcomps_1(X1)),k6_weierstr(X1,X2))=u1_struct_0(k3_pcomps_1(X1))|u1_struct_0(k3_topmetr)=k6_numbers|v2_struct_0(X1)|~m1_subset_1(k6_weierstr(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_191, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v1_xboole_0(k7_relat_1(X3,X2)))).# trainpos
cnf(c_0_389, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk3_0))).# trainpos
cnf(c_0_212, plain, (v2_struct_0(X1)|~l1_metric_1(X1)|~v1_xboole_0(u1_struct_0(k3_pcomps_1(X1))))).# trainpos
cnf(c_0_236, plain, (k2_seq_4(k7_relat_1(k6_weierstr(X1,X2),X3))=k8_weierstr(X1,X2,X3)|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X3))|v2_struct_0(X1)|~m1_subset_1(k7_relat_1(k6_weierstr(X1,X2),X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k7_relat_1(k6_weierstr(X1,X2),X3),k1_zfmisc_1(u1_struct_0(k3_topmetr)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_371, plain, (k1_funct_1(k6_weierstr(X1,X2),X3)=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|~r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),k6_numbers)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3)))).# trainpos
cnf(c_0_86, plain, (k9_xtuple_0(k6_weierstr(X1,X2))=u1_struct_0(k3_pcomps_1(X1))|u1_struct_0(k3_topmetr)=k6_numbers|v2_struct_0(X1)|~v4_relat_1(k6_weierstr(X1,X2),u1_struct_0(k3_pcomps_1(X1)))|~v1_relat_1(k6_weierstr(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_148, plain, (X1=X2|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_211, plain, (v1_xboole_0(u1_struct_0(k3_pcomps_1(X1)))|v2_struct_0(X1)|~v1_funct_1(k6_weierstr(X1,X2))|~m1_subset_1(k6_weierstr(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr))))|~v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_238, negated_conjecture, (~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_xboole_0(k7_relat_1(k6_numbers,u1_struct_0(esk1_0)))|u1_struct_0(k3_topmetr)!=k6_numbers)).# trainpos
cnf(c_0_176, plain, (X1=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))|~v1_xboole_0(X1))).# trainpos
cnf(c_0_303, plain, (k8_weierstr(X1,X2,X3)=k6_numbers|v2_struct_0(X1)|k1_hausdorf(X1,X3,X2)!=k6_numbers|~v1_xreal_0(k8_weierstr(X1,X3,X2))|~v1_xreal_0(k8_weierstr(X1,X2,X3))|~r1_xxreal_0(k6_numbers,k8_weierstr(X1,X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_103, plain, (v4_relat_1(k6_weierstr(X1,X2),u1_struct_0(k3_pcomps_1(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_312, plain, (k8_weierstr(X1,X2,X3)=k6_numbers|v2_struct_0(X1)|k1_hausdorf(X1,X2,X3)!=k6_numbers|~v1_xreal_0(k8_weierstr(X1,X3,X2))|~v1_xreal_0(k8_weierstr(X1,X2,X3))|~r1_xxreal_0(k6_numbers,k8_weierstr(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_334, plain, (k8_weierstr(X1,X2,X3)=k6_numbers|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|k1_hausdorf(X1,X2,X3)!=k6_numbers|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1)))).# trainpos
cnf(c_0_320, plain, (k8_weierstr(X1,X2,X3)=k6_numbers|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|k1_hausdorf(X1,X3,X2)!=k6_numbers|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1)))).# trainpos
cnf(c_0_197, plain, (k7_relset_1(X2,X1,X3,X4)=X1|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1))))).# trainpos
cnf(c_0_124, plain, (k9_xtuple_0(k6_weierstr(X1,X2))=u1_struct_0(k3_pcomps_1(X1))|u1_struct_0(k3_topmetr)=k6_numbers|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_374, plain, (m1_subset_1(X1,k1_numbers)|v1_xboole_0(X2)|~r2_hidden(X1,esk4_1(X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).# trainpos
cnf(c_0_361, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k7_relset_1(X3,X2,X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).# trainpos
cnf(c_0_205, plain, (m1_subset_1(k7_relat_1(X3,X4),k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_223, plain, (m1_subset_1(k7_relat_1(k6_weierstr(X1,X2),X3),k1_zfmisc_1(u1_struct_0(k3_topmetr)))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_380, plain, (m1_subset_1(X1,k1_numbers)|v1_xboole_0(X2)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).# trainpos
cnf(c_0_215, plain, (X2=k7_relat_1(X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X2))).# trainpos
cnf(c_0_154, plain, (r2_hidden(X1,X2)|u1_struct_0(k3_topmetr)=k6_numbers|v2_struct_0(X3)|X2!=k7_relat_1(k6_weierstr(X3,X4),X5)|X1!=k1_funct_1(k6_weierstr(X3,X4),X6)|~v1_funct_1(k6_weierstr(X3,X4))|~v1_relat_1(k6_weierstr(X3,X4))|~r2_hidden(X6,u1_struct_0(k3_pcomps_1(X3)))|~r2_hidden(X6,X5)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).# trainpos
cnf(c_0_333, plain, (k2_square_1(k6_numbers,k8_weierstr(X1,X3,X2))=k1_hausdorf(X1,X2,X3)|v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X3,X2)!=k6_numbers|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1)))).# trainpos
cnf(c_0_258, plain, (m1_subset_1(k7_relat_1(k6_weierstr(X1,X2),X3),k1_zfmisc_1(k1_numbers))|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X3))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_203, plain, (u1_struct_0(k3_topmetr)=k6_numbers|r2_hidden(X1,k7_relat_1(k6_weierstr(X2,X3),X4))|v2_struct_0(X2)|X1!=k1_funct_1(k6_weierstr(X2,X3),X5)|~r2_hidden(X5,u1_struct_0(k3_pcomps_1(X2)))|~r2_hidden(X5,X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).# trainpos
cnf(c_0_235, plain, (v3_membered(k7_relat_1(k6_weierstr(X1,X2),X3))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_230, plain, (X1=k7_relat_1(k6_numbers,X2)|~v1_xboole_0(X1))).# trainpos
cnf(c_0_239, plain, (k7_relat_1(k6_numbers,X1)=k6_numbers)).# trainpos
cnf(c_0_221, plain, (u1_struct_0(k3_topmetr)=k6_numbers|r2_hidden(k1_funct_1(k6_weierstr(X1,X2),X3),k7_relat_1(k6_weierstr(X1,X2),X4))|v2_struct_0(X1)|~r2_hidden(X3,u1_struct_0(k3_pcomps_1(X1)))|~r2_hidden(X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_348, plain, (k1_hausdorf(X1,X3,X2)=k2_square_1(k6_numbers,k6_numbers)|v1_xboole_0(X3)|v1_xboole_0(X2)|v2_struct_0(X1)|k1_hausdorf(X1,X2,X3)!=k6_numbers|~v2_compts_1(X2,k3_pcomps_1(X1))|~v2_compts_1(X3,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_296, plain, (r2_hidden(X1,X2)|v1_xboole_0(u1_struct_0(k3_pcomps_1(X3)))|u1_struct_0(k3_topmetr)=k6_numbers|v2_struct_0(X3)|X2!=k7_relat_1(k6_weierstr(X3,X4),X5)|X1!=k1_funct_1(k6_weierstr(X3,X4),X6)|~v1_funct_1(k6_weierstr(X3,X4))|~v1_relat_1(k6_weierstr(X3,X4))|~r2_hidden(X6,X5)|~m1_subset_1(X6,u1_struct_0(k3_pcomps_1(X3)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).# trainpos
cnf(c_0_365, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k7_relat_1(X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).# trainpos
cnf(c_0_245, negated_conjecture, (u1_struct_0(k3_topmetr)!=k6_numbers|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_309, plain, (u1_struct_0(k3_topmetr)=k6_numbers|r2_hidden(X1,k7_relat_1(k6_weierstr(X2,X3),X4))|v2_struct_0(X2)|X1!=k1_funct_1(k6_weierstr(X2,X3),X5)|~r2_hidden(X5,X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X5,u1_struct_0(k3_pcomps_1(X2)))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).# trainpos
cnf(c_0_360, negated_conjecture, (k1_hausdorf(esk1_0,esk3_0,esk2_0)=k2_square_1(k6_numbers,k6_numbers))).# trainpos
cnf(c_0_234, plain, (r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),X4)|u1_struct_0(k3_topmetr)=k6_numbers|v2_struct_0(X1)|X4!=k2_seq_4(k7_relat_1(k6_weierstr(X1,X2),X5))|~v3_membered(k7_relat_1(k6_weierstr(X1,X2),X5))|~v1_xreal_0(X4)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~v4_xxreal_2(k7_relat_1(k6_weierstr(X1,X2),X5))|~r2_hidden(X3,u1_struct_0(k3_pcomps_1(X1)))|~r2_hidden(X3,X5)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_369, plain, (r2_hidden(X1,u1_struct_0(k3_topmetr))|v2_struct_0(X2)|~r2_hidden(X1,k7_relat_1(k6_weierstr(X2,X3),X4))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).# trainpos
cnf(c_0_250, negated_conjecture, (u1_struct_0(k3_topmetr)!=k6_numbers)).# trainpos
cnf(c_0_364, negated_conjecture, (k2_square_1(k6_numbers,k6_numbers)=k6_numbers)).# trainpos
cnf(c_0_318, plain, (u1_struct_0(k3_topmetr)=k6_numbers|r2_hidden(k1_funct_1(k6_weierstr(X1,X2),X3),k7_relat_1(k6_weierstr(X1,X2),X4))|v2_struct_0(X1)|~r2_hidden(X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_249, plain, (u1_struct_0(k3_topmetr)=k6_numbers|r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),X4)|v2_struct_0(X1)|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X5))|X4!=k8_weierstr(X1,X2,X5)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~v1_xreal_0(X4)|~r2_hidden(X3,u1_struct_0(k3_pcomps_1(X1)))|~r2_hidden(X3,X5)|~v4_xxreal_2(k7_relat_1(k6_weierstr(X1,X2),X5))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(k7_relat_1(k6_weierstr(X1,X2),X5),k1_zfmisc_1(u1_struct_0(k3_topmetr)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))).# trainpos
cnf(c_0_368, negated_conjecture, (k1_hausdorf(esk1_0,esk3_0,esk2_0)=k6_numbers)).# trainpos
cnf(c_0_257, plain, (r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),X4)|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X5))|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X5)|X4!=k8_weierstr(X1,X2,X5)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~v1_xreal_0(X4)|~r2_hidden(X3,X5)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_compts_1(X5,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(k7_relat_1(k6_weierstr(X1,X2),X5),k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_375, plain, (r2_hidden(k1_funct_1(k6_weierstr(X1,X2),X3),u1_struct_0(k3_topmetr))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1))))).# trainpos
cnf(c_0_271, plain, (r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),k8_weierstr(X1,X2,X4))|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X4))|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~v1_xreal_0(k8_weierstr(X1,X2,X4))|~r2_hidden(X3,X4)|~v2_compts_1(X4,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_381, plain, (r2_hidden(k1_funct_1(k6_weierstr(X1,X2),X3),u1_struct_0(k3_topmetr))|v2_struct_0(X1)|v1_xboole_0(X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(X3,X4))).# trainpos
cnf(c_0_387, plain, (r2_hidden(k1_funct_1(k6_weierstr(X1,X2),X3),u1_struct_0(k3_topmetr))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_288, plain, (r1_xxreal_0(X5,k8_weierstr(X1,X2,X6))|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X6))|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_xreal_0(X5)|~r2_hidden(esk6_4(k6_weierstr(X1,X2),X3,X4,X5),X6)|~v2_compts_1(X6,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|X4!=k7_relat_1(k6_weierstr(X1,X2),X3)|~v1_funct_1(k6_weierstr(X1,X2))|~v1_relat_1(k6_weierstr(X1,X2))|~r2_hidden(X5,X4))).# trainpos
cnf(c_0_392, plain, (m1_subset_1(k1_funct_1(k6_weierstr(X1,X2),X3),k1_numbers)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_301, plain, (r1_xxreal_0(X1,k8_weierstr(X2,X3,X4))|v1_xboole_0(k7_relat_1(k6_weierstr(X2,X3),X4))|v1_xboole_0(X3)|v2_struct_0(X2)|X5!=k7_relat_1(k6_weierstr(X2,X3),X4)|~v1_xreal_0(X1)|~r2_hidden(X1,X5)|~v2_compts_1(X4,k3_pcomps_1(X2))|~v2_compts_1(X3,k3_pcomps_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v1_funct_1(k6_weierstr(X2,X3))|~v1_relat_1(k6_weierstr(X2,X3)))).# trainpos
cnf(c_0_317, plain, (r1_xxreal_0(X1,k8_weierstr(X2,X3,X4))|v1_xboole_0(k7_relat_1(k6_weierstr(X2,X3),X4))|v1_xboole_0(X3)|v2_struct_0(X2)|~v1_xreal_0(X1)|~r2_hidden(X1,k7_relat_1(k6_weierstr(X2,X3),X4))|~v2_compts_1(X4,k3_pcomps_1(X2))|~v2_compts_1(X3,k3_pcomps_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).# trainpos
cnf(c_0_345, plain, (r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),k8_weierstr(X1,X2,X4))|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~v2_compts_1(X4,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1))))).# trainpos
cnf(c_0_363, plain, (r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),k6_numbers)|v1_xboole_0(X2)|v2_struct_0(X1)|v1_xboole_0(X4)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~r2_hidden(X3,X4)|~v2_compts_1(X4,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X2,X4)!=k6_numbers)).# trainpos
cnf(c_0_383, negated_conjecture, (r1_xxreal_0(k1_funct_1(k6_weierstr(esk1_0,esk2_0),X1),k6_numbers)|~v1_xreal_0(k1_funct_1(k6_weierstr(esk1_0,esk2_0),X1))|~r2_hidden(X1,esk3_0))).# trainpos
cnf(c_0_372, negated_conjecture, (r1_xxreal_0(k1_funct_1(k6_weierstr(esk1_0,esk3_0),X1),k6_numbers)|~v1_xreal_0(k1_funct_1(k6_weierstr(esk1_0,esk3_0),X1))|~r2_hidden(X1,esk2_0))).# trainpos
cnf(c_0_388, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,esk2_0),X1)=k6_numbers|~v1_xreal_0(k1_funct_1(k6_weierstr(esk1_0,esk2_0),X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk3_0))).# trainpos
cnf(c_0_377, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,esk3_0),X1)=k6_numbers|~v1_xreal_0(k1_funct_1(k6_weierstr(esk1_0,esk3_0),X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0))).# trainpos
cnf(c_0_391, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,esk3_0),X1)=k6_numbers|~r2_hidden(X1,esk2_0)|~m1_subset_1(k1_funct_1(k6_weierstr(esk1_0,esk3_0),X1),k1_numbers))).# trainpos
cnf(c_0_401, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,esk2_0),X1)=k6_numbers|~r2_hidden(X1,esk3_0)|~m1_subset_1(k1_funct_1(k6_weierstr(esk1_0,esk2_0),X1),k1_numbers))).# trainpos
cnf(c_0_407, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,esk2_0),X1)=k6_numbers|~r2_hidden(X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_397, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,esk3_0),X1)=k6_numbers|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_408, negated_conjecture, (r2_hidden(X1,esk3_0)|~v4_pre_topc(esk3_0,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0))).# trainpos
cnf(c_0_417, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk2_0))).# trainpos
cnf(c_0_419, negated_conjecture, (r1_tarski(X1,esk3_0)|~r2_hidden(esk5_2(X1,esk3_0),esk2_0))).# trainpos
cnf(c_0_422, negated_conjecture, (r1_tarski(esk2_0,esk3_0))).# trainpos
cnf(c_0_415, negated_conjecture, (r2_hidden(X1,esk2_0)|~v4_pre_topc(esk2_0,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk3_0))).# trainpos
cnf(c_0_425, negated_conjecture, (~r1_tarski(esk3_0,esk2_0))).# trainpos
cnf(c_0_421, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk3_0))).# trainpos
cnf(c_0_424, negated_conjecture, (r2_hidden(esk5_2(esk3_0,X1),esk2_0)|r1_tarski(esk3_0,X1))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_122, plain, (k5_numbers=k4_ordinal1)).#trainneg
cnf(i_0_176, plain, (v2_struct_0(esk18_0))).#trainneg
cnf(i_0_190, plain, (v2_struct_0(esk20_0))).#trainneg
cnf(i_0_195, plain, (l1_metric_1(esk22_0))).#trainneg
cnf(i_0_224, plain, (~v2_struct_0(esk27_0))).#trainneg
cnf(i_0_255, plain, (~v2_struct_0(esk32_0))).#trainneg
cnf(i_0_141, plain, (l1_struct_0(X1)|~l1_metric_1(X1))).#trainneg
cnf(i_0_212, plain, (v1_relat_1(esk25_0))).#trainneg
cnf(i_0_211, plain, (v1_funct_1(esk25_0))).#trainneg
cnf(i_0_88, plain, (v1_xreal_0(esk12_0))).#trainneg
cnf(i_0_200, plain, (v1_xreal_0(esk23_0))).#trainneg
cnf(i_0_177, plain, (l1_pre_topc(esk18_0))).#trainneg
cnf(i_0_191, plain, (l1_pre_topc(esk20_0))).#trainneg
cnf(i_0_225, plain, (l1_pre_topc(esk27_0))).#trainneg
cnf(i_0_249, plain, (l1_pre_topc(esk30_0))).#trainneg
cnf(i_0_256, plain, (l1_pre_topc(esk32_0))).#trainneg
cnf(i_0_115, plain, (v1_pre_topc(k3_topmetr))).#trainneg
cnf(i_0_175, plain, (v1_pre_topc(esk18_0))).#trainneg
cnf(i_0_189, plain, (v1_pre_topc(esk20_0))).#trainneg
cnf(i_0_232, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_254, plain, (v1_pre_topc(esk32_0))).#trainneg
cnf(i_0_94, plain, (v2_pre_topc(k3_topmetr))).#trainneg
cnf(i_0_257, plain, (v6_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_174, plain, (v2_pre_topc(esk18_0))).#trainneg
cnf(i_0_223, plain, (v2_pre_topc(esk27_0))).#trainneg
cnf(i_0_253, plain, (v2_pre_topc(esk32_0))).#trainneg
cnf(i_0_90, plain, (v1_xcmplx_0(esk12_0))).#trainneg
cnf(i_0_217, plain, (l1_struct_0(esk26_0))).#trainneg
cnf(i_0_137, plain, (v6_membered(k4_ordinal1))).#trainneg
cnf(i_0_210, plain, (v4_valued_0(esk25_0))).#trainneg
cnf(i_0_222, plain, (v7_pre_topc(esk27_0))).#trainneg
cnf(i_0_124, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_186, plain, (~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_xboole_0(u1_pre_topc(X1)))).#trainneg
cnf(i_0_159, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_161, plain, (v3_membered(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_187, plain, (~v1_subset_1(esk19_1(X1),X1))).#trainneg
cnf(i_0_197, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_145, plain, (v2_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_218, plain, (v4_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_126, plain, (m1_subset_1(esk14_1(X1),X1))).#trainneg
cnf(i_0_151, plain, (v1_xboole_0(esk15_2(X1,X2)))).#trainneg
cnf(i_0_150, plain, (v1_relat_1(esk15_2(X1,X2)))).#trainneg
cnf(i_0_226, plain, (v5_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_219, plain, (v1_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_16, plain, (r1_tarski(X1,X2)|X1!=X2)).#trainneg
cnf(i_0_17, plain, (r1_tarski(X1,X2)|X1!=X2)).#trainneg
cnf(i_0_188, plain, (m1_subset_1(esk19_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_149, plain, (v4_relat_1(esk15_2(X1,X2),X1))).#trainneg
cnf(i_0_148, plain, (v5_relat_1(esk15_2(X1,X2),X2))).#trainneg
cnf(i_0_207, plain, (v1_xboole_0(X1)|v1_zfmisc_1(esk24_1(X1)))).#trainneg
cnf(i_0_163, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk16_1(X1)))).#trainneg
cnf(i_0_108, plain, (X1=X2|g1_pre_topc(X3,X1)!=g1_pre_topc(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X3))))).#trainneg
cnf(i_0_119, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_220, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainneg
cnf(i_0_140, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_208, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk24_1(X1)))).#trainneg
cnf(i_0_87, plain, (v1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_180, plain, (v1_xreal_0(k2_seq_4(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_30, plain, (k2_pcomps_1(X1)=u1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_147, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_85, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_201, plain, (v3_valued_0(k2_zfmisc_1(X2,X1))|~v3_membered(X1))).#trainneg
cnf(i_0_227, plain, (v2_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_214, plain, (v7_pre_topc(X1)|~l1_pre_topc(X1)|~v8_pre_topc(X1))).#trainneg
cnf(i_0_164, plain, (v1_xboole_0(X1)|m1_subset_1(esk16_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_209, plain, (v1_xboole_0(X1)|m1_subset_1(esk24_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_152, plain, (m1_subset_1(esk15_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_230, plain, (v2_struct_0(X1)|~v6_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~v1_xboole_0(esk28_1(X1)))).#trainneg
cnf(i_0_153, plain, (m1_subset_1(k4_seq_4(X1),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_109, plain, (X1=X2|g1_pre_topc(X1,X3)!=g1_pre_topc(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_138, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_120, plain, (m1_subset_1(k2_weierstr(X1),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_136, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_146, plain, (k4_xxreal_0(X1,X1)=X1|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_160, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_259, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_184, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_250, plain, (v4_pre_topc(esk31_1(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_204, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_22, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk4_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_193, plain, (v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_xboole_0(esk21_1(X1)))).#trainneg
cnf(i_0_192, plain, (v2_struct_0(X1)|v4_pre_topc(esk21_1(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_196, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_143, plain, (v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))).#trainneg
cnf(i_0_199, plain, (v3_valued_0(X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_258, plain, (v2_struct_0(X1)|v6_tbsp_1(X2,X1)|~l1_metric_1(X1)|~v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_132, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_83, plain, (g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=X1|~l1_pre_topc(X1)|~v1_pre_topc(X1))).#trainneg
cnf(i_0_185, plain, (v1_xreal_0(k4_xxreal_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_156, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_183, plain, (v3_membered(k7_relat_1(X1,X2))|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_98, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_121, plain, (m1_subset_1(k2_pcomps_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_metric_1(X1))).#trainneg
cnf(i_0_181, plain, (l1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_182, plain, (v1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_130, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_110, plain, (k6_xcmplx_0(X1,X2)=k10_binop_2(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_203, plain, (v1_xboole_0(X1)|~v2_struct_0(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_229, plain, (v2_struct_0(X1)|v1_finset_1(esk28_1(X1))|~v6_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1))).#trainneg
cnf(i_0_128, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_155, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_129, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_162, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_198, plain, (v4_pre_topc(X2,X1)|~v1_xboole_0(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_206, plain, (v2_struct_0(X1)|v6_tbsp_1(X2,X1)|~v6_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_205, plain, (v1_xxreal_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v2_valued_0(X1))).#trainneg
cnf(i_0_166, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_111, plain, (m1_subset_1(k10_binop_2(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_113, plain, (k2_square_1(X1,X1)=X1|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_251, plain, (m1_subset_1(esk31_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_142, plain, (v2_struct_0(X1)|v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_158, plain, (v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_157, plain, (v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_213, plain, (v2_valued_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_252, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_69, plain, (k2_square_1(X1,X2)=k2_square_1(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_117, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_92, plain, (m1_subset_1(k2_square_1(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_135, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(k7_relat_1(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k9_xtuple_0(X1))))).#trainneg
cnf(i_0_194, plain, (v2_struct_0(X1)|m1_subset_1(esk21_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_172, plain, (v2_struct_0(X2)|v1_relat_1(esk17_2(X1,X2))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_133, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_125, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_169, plain, (v2_struct_0(X2)|v1_funct_1(esk17_2(X1,X2))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_26, plain, (v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_248, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m2_subset_1(esk29_2(X1,X2),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_28, plain, (X3=k1_xboole_0|v1_funct_2(X2,X1,X3)|k1_relset_1(X1,X2)!=X1|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_171, plain, (v2_struct_0(X2)|v4_relat_1(esk17_2(X1,X2),u1_struct_0(X1))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_170, plain, (v2_struct_0(X2)|v5_relat_1(esk17_2(X1,X2),u1_struct_0(X2))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_167, plain, (v2_struct_0(X2)|v5_pre_topc(esk17_2(X1,X2),X1,X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_45, plain, (r2_hidden(esk6_4(X1,X2,X3,X4),k9_xtuple_0(X1))|X3!=k7_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_64, plain, (X2=k2_seq_4(X1)|v1_xboole_0(X1)|v1_xreal_0(esk10_2(X1,X2))|v1_xreal_0(esk11_2(X1,X2))|~v4_xxreal_2(X1)|~v1_xreal_0(X2)|~v3_membered(X1))).#trainneg
cnf(i_0_61, plain, (X2=k2_seq_4(X1)|v1_xboole_0(X1)|v1_xreal_0(esk11_2(X1,X2))|r2_hidden(esk10_2(X1,X2),X1)|~v4_xxreal_2(X1)|~v1_xreal_0(X2)|~v3_membered(X1))).#trainneg
cnf(i_0_63, plain, (X2=k2_seq_4(X1)|v1_xboole_0(X1)|v1_xreal_0(esk10_2(X1,X2))|~v4_xxreal_2(X1)|~v1_xreal_0(X2)|~v3_membered(X1)|~r1_xxreal_0(esk11_2(X1,X2),k6_numbers))).#trainneg
cnf(i_0_65, plain, (v1_xboole_0(X1)|r1_xxreal_0(X3,k6_numbers)|X2!=k2_seq_4(X1)|~v4_xxreal_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v3_membered(X1)|~r1_xxreal_0(esk9_3(X1,X2,X3),k10_binop_2(X2,X3)))).#trainneg
cnf(i_0_168, plain, (v2_struct_0(X2)|v1_funct_2(esk17_2(X1,X2),u1_struct_0(X1),u1_struct_0(X2))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_58, plain, (X2=k2_seq_4(X1)|v1_xboole_0(X1)|v1_xreal_0(esk11_2(X1,X2))|~v4_xxreal_2(X1)|~v1_xreal_0(X2)|~v3_membered(X1)|~r1_xxreal_0(esk10_2(X1,X2),X2))).#trainneg
cnf(i_0_60, plain, (X2=k2_seq_4(X1)|v1_xboole_0(X1)|r2_hidden(esk10_2(X1,X2),X1)|~v4_xxreal_2(X1)|~v1_xreal_0(X2)|~v3_membered(X1)|~r1_xxreal_0(esk11_2(X1,X2),k6_numbers))).#trainneg
cnf(i_0_39, plain, (X3=k7_relat_1(X1,X2)|r2_hidden(esk7_3(X1,X2,X3),X3)|r2_hidden(esk8_3(X1,X2,X3),X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_40, plain, (X3=k7_relat_1(X1,X2)|r2_hidden(esk7_3(X1,X2,X3),X3)|r2_hidden(esk8_3(X1,X2,X3),k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_173, plain, (v2_struct_0(X2)|m1_subset_1(esk17_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_57, plain, (X2=k2_seq_4(X1)|v1_xboole_0(X1)|~v4_xxreal_2(X1)|~v1_xreal_0(X2)|~v3_membered(X1)|~r1_xxreal_0(esk10_2(X1,X2),X2)|~r1_xxreal_0(esk11_2(X1,X2),k6_numbers))).#trainneg
cnf(i_0_38, plain, (X3=k7_relat_1(X1,X2)|k1_funct_1(X1,esk8_3(X1,X2,X3))=esk7_3(X1,X2,X3)|r2_hidden(esk7_3(X1,X2,X3),X3)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_67, plain, (v1_xboole_0(X1)|r1_xxreal_0(X3,k6_numbers)|v1_xreal_0(esk9_3(X1,X2,X3))|X2!=k2_seq_4(X1)|~v4_xxreal_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v3_membered(X1))).#trainneg
cnf(i_0_66, plain, (v1_xboole_0(X1)|r1_xxreal_0(X3,k6_numbers)|r2_hidden(esk9_3(X1,X2,X3),X1)|X2!=k2_seq_4(X1)|~v4_xxreal_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v3_membered(X1))).#trainneg
cnf(i_0_99, plain, (v1_xboole_0(X1)|m1_subset_1(k3_funct_2(X1,X3,X2,X4),X3)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_41, plain, (X3=k7_relat_1(X1,X2)|esk7_3(X1,X2,X3)!=k1_funct_1(X1,X4)|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X4,X2)|~r2_hidden(X4,k9_xtuple_0(X1))|~r2_hidden(esk7_3(X1,X2,X3),X3))).#trainneg
cnf(i_0_62, plain, (X3=k2_seq_4(X2)|v1_xboole_0(X2)|v1_xreal_0(esk10_2(X2,X3))|r1_xxreal_0(X1,k10_binop_2(X3,esk11_2(X2,X3)))|~v4_xxreal_2(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v3_membered(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_72, plain, (v2_struct_0(X1)|m1_subset_1(k1_hausdorf(X1,X2,X3),k1_numbers)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))).#trainneg
cnf(i_0_24, plain, (k1_funct_1(k6_weierstr(X3,X2),X1)=k6_numbers|v2_struct_0(X3)|v1_xboole_0(X2)|~v6_metric_1(X3)|~v7_metric_1(X3)|~v8_metric_1(X3)|~v9_metric_1(X3)|~l1_metric_1(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X1,u1_struct_0(X3))|~v4_pre_topc(X2,k3_pcomps_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3)))))).#trainneg
cnf(i_0_59, plain, (X3=k2_seq_4(X2)|v1_xboole_0(X2)|r2_hidden(esk10_2(X2,X3),X2)|r1_xxreal_0(X1,k10_binop_2(X3,esk11_2(X2,X3)))|~v4_xxreal_2(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v3_membered(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_73, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|v5_pre_topc(k6_weierstr(X1,X2),k3_pcomps_1(X1),k3_topmetr)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))).#trainneg
cnf(i_0_56, plain, (X3=k2_seq_4(X2)|v1_xboole_0(X2)|r1_xxreal_0(X1,k10_binop_2(X3,esk11_2(X2,X3)))|~v4_xxreal_2(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v3_membered(X2)|~r2_hidden(X1,X2)|~r1_xxreal_0(esk10_2(X2,X3),X3))).#trainneg
cnf(i_0_347, plain, (r1_tarski(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(esk5_2(X1,X2),X2))).#trainneg
cnf(i_0_291, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_288, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_metric_1(X1))).#trainneg
cnf(i_0_372, plain, (v3_membered(k4_ordinal1))).#trainneg
cnf(i_0_88, plain, (v1_xreal_0(k6_numbers))).#trainneg
cnf(i_0_789, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_90, plain, (v1_xcmplx_0(k6_numbers))).#trainneg
cnf(i_0_263, plain, (v4_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_262, plain, (esk15_2(X1,X2)=k6_numbers)).#trainneg
cnf(i_0_266, plain, (u1_struct_0(k3_pcomps_1(esk22_0))=u1_struct_0(esk22_0))).#trainneg
cnf(i_0_767, negated_conjecture, (v4_xxreal_2(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|X1!=k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X2),X3)|~v2_compts_1(X3,k3_pcomps_1(esk1_0))|~v2_compts_1(X2,k3_pcomps_1(esk1_0))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_733, plain, (k8_weierstr(X1,X2,X3)=k6_numbers|v1_xboole_0(X2)|v1_xboole_0(X3)|v2_struct_0(X1)|~r1_xxreal_0(k8_weierstr(X1,X2,X3),k6_numbers)|~v1_xxreal_0(k6_numbers)|~v1_xxreal_0(k8_weierstr(X1,X2,X3))|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_764, negated_conjecture, (k2_weierstr(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2))=k8_weierstr(esk1_0,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_717, plain, (k6_numbers=X5|v2_struct_0(X1)|X4!=k7_relat_1(k6_weierstr(X1,X2),X3)|~v1_funct_1(k6_weierstr(X1,X2))|~v1_relat_1(k6_weierstr(X1,X2))|~r2_hidden(X5,X4)|~r2_hidden(esk6_4(k6_weierstr(X1,X2),X3,X4,X5),X2)|~v4_pre_topc(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(esk6_4(k6_weierstr(X1,X2),X3,X4,X5),u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_383, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk19_1(X2)))).#trainneg
cnf(i_0_386, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk14_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_596, plain, (~v1_xboole_0(X2)|X3!=k7_relat_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_285, plain, (m1_subset_1(k2_seq_4(X1),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_381, plain, (v2_membered(k4_ordinal1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_290, plain, (v2_struct_0(esk26_0)|~v1_xboole_0(u1_struct_0(esk26_0)))).#trainneg
cnf(i_0_308, plain, (v1_xxreal_0(k4_ordinal1)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_809, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_289, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_287, plain, (v1_xboole_0(u1_struct_0(esk26_0))|~v2_struct_0(esk26_0))).#trainneg
cnf(i_0_818, plain, (m1_subset_1(u1_struct_0(k3_topmetr),k1_zfmisc_1(k1_numbers))|v1_xboole_0(u1_struct_0(k3_topmetr)))).#trainneg
cnf(i_0_763, negated_conjecture, (r1_xxreal_0(k6_numbers,k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2))|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_742, plain, (r1_xxreal_0(k6_numbers,k1_funct_1(k6_weierstr(X1,X2),X3))|v1_xboole_0(X2)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(k3_pcomps_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v1_funct_1(k6_weierstr(X1,X2))|~v1_funct_2(k6_weierstr(X1,X2),u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr))|~m1_subset_1(k6_weierstr(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr)))))).#trainneg
cnf(i_0_571, plain, (r1_xxreal_0(X1,X2)|v1_xboole_0(X3)|X2!=k2_seq_4(X3)|~v3_membered(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v4_xxreal_2(X3)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_745, plain, (v4_xxreal_2(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3))|v1_xboole_0(X2)|v1_xboole_0(X3)|v2_struct_0(X1)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3),k1_zfmisc_1(k1_numbers))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_570, plain, (r1_xxreal_0(esk5_2(X1,X2),X3)|r1_tarski(X1,X2)|X3!=k2_seq_4(X1)|~v3_membered(X1)|~v1_xreal_0(X3)|~v1_xreal_0(esk5_2(X1,X2))|~v4_xxreal_2(X1))).#trainneg
cnf(i_0_694, negated_conjecture, (m1_subset_1(k6_numbers,k1_numbers))).#trainneg
cnf(i_0_438, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,k2_pcomps_1(X2))|~l1_metric_1(X2))).#trainneg
cnf(i_0_356, plain, (v1_xboole_0(esk19_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_341, plain, (v7_ordinal1(k4_ordinal1)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_409, plain, (v3_valued_0(k4_ordinal1)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_319, plain, (v1_xreal_0(k4_ordinal1)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_330, plain, (v1_xcmplx_0(k4_ordinal1)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_297, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_806, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_299, plain, (k4_xxreal_0(X1,X1)=X1|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_300, plain, (v1_xxreal_0(esk24_1(X1))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_605, plain, (r2_hidden(esk6_4(k6_weierstr(X1,X2),X3,X4,X5),X2)|v1_xboole_0(X2)|v2_struct_0(X1)|X5!=k6_numbers|~v4_pre_topc(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(esk6_4(k6_weierstr(X1,X2),X3,X4,X5),u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|X4!=k7_relat_1(k6_weierstr(X1,X2),X3)|~v1_funct_1(k6_weierstr(X1,X2))|~v1_relat_1(k6_weierstr(X1,X2))|~r2_hidden(X5,X4))).#trainneg
cnf(i_0_737, plain, (k2_square_1(k8_weierstr(X1,X3,X2),k8_weierstr(X1,X2,X3))=k1_hausdorf(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(k8_weierstr(X1,X3,X2),k1_numbers)|~m1_subset_1(k8_weierstr(X1,X2,X3),k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_768, plain, (k4_seq_4(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3))=k8_weierstr(X1,X2,X3)|v2_struct_0(X1)|v1_xboole_0(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_644, plain, (m1_subset_1(esk9_3(X1,X2,X3),X1)|r1_xxreal_0(X3,k6_numbers)|v1_xboole_0(X1)|X2!=k2_seq_4(X1)|~v3_membered(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v4_xxreal_2(X1))).#trainneg
cnf(i_0_788, plain, (v2_struct_0(esk22_0)|~v1_xboole_0(u1_struct_0(esk22_0)))).#trainneg
cnf(i_0_150, plain, (v1_relat_1(k6_numbers))).#trainneg
cnf(i_0_1000, negated_conjecture, (v7_ordinal1(k6_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_301, plain, (v1_xxreal_0(esk19_1(X1))|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_148, plain, (v5_relat_1(k6_numbers,X2))).#trainneg
cnf(i_0_916, plain, (~v1_xboole_0(u1_struct_0(esk32_0)))).#trainneg
cnf(i_0_918, plain, (~v1_xboole_0(u1_struct_0(esk27_0)))).#trainneg
cnf(i_0_149, plain, (v4_relat_1(k6_numbers,X1))).#trainneg
cnf(i_0_1028, plain, (v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_302, plain, (v1_xxreal_0(esk16_1(X1))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_304, plain, (v1_xxreal_0(esk14_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_873, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0))).#trainneg
cnf(i_0_1051, plain, (k2_seq_4(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3))=k8_weierstr(X1,X2,X3)|v1_xboole_0(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3))|v2_struct_0(X1)|~m1_subset_1(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k7_relset_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3),k1_zfmisc_1(u1_struct_0(k3_topmetr)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_950, negated_conjecture, (v1_relat_1(k6_weierstr(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_876, negated_conjecture, (v4_xxreal_2(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v2_compts_1(X2,k3_pcomps_1(esk1_0))|~v2_compts_1(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_917, plain, (v2_struct_0(esk30_0)|~v1_xboole_0(u1_struct_0(esk30_0)))).#trainneg
cnf(i_0_910, plain, (v1_xboole_0(u1_struct_0(esk20_0)))).#trainneg
cnf(i_0_911, plain, (v1_xboole_0(u1_struct_0(esk18_0)))).#trainneg
cnf(i_0_309, plain, (v1_xxreal_0(k4_seq_4(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_905, plain, (v1_xxreal_0(k4_ordinal1)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_832, negated_conjecture, (v3_membered(esk3_0)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_833, negated_conjecture, (v2_membered(esk3_0)|~v2_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_842, plain, (v3_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_839, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_310, plain, (v1_xxreal_0(k2_weierstr(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_311, plain, (v1_xreal_0(esk24_1(X1))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1014, plain, (v3_membered(u1_struct_0(k3_topmetr))|~l1_metric_1(k3_topmetr))).#trainneg
cnf(i_0_836, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_992, plain, (k4_xxreal_0(k8_weierstr(X1,X3,X2),k8_weierstr(X1,X2,X3))=k1_hausdorf(X1,X2,X3)|v2_struct_0(X1)|~v1_xxreal_0(k8_weierstr(X1,X3,X2))|~v1_xxreal_0(k8_weierstr(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_958, plain, (r1_xxreal_0(k6_numbers,X5)|v1_xboole_0(X2)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(esk6_4(k6_weierstr(X1,X2),X3,X4,X5),u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|X4!=k7_relat_1(k6_weierstr(X1,X2),X3)|~v1_funct_1(k6_weierstr(X1,X2))|~v1_relat_1(k6_weierstr(X1,X2))|~r2_hidden(X5,X4))).#trainneg
cnf(i_0_877, negated_conjecture, (v4_xxreal_2(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|X1!=k7_relat_1(k6_weierstr(esk1_0,X3),X2)|~v2_compts_1(X2,k3_pcomps_1(esk1_0))|~v2_compts_1(X3,k3_pcomps_1(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(k6_weierstr(esk1_0,X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr)))))).#trainneg
cnf(i_0_878, plain, (k8_weierstr(X1,X2,X3)=k6_numbers|v1_xboole_0(X3)|v1_xboole_0(X2)|v2_struct_0(X1)|~r1_xxreal_0(k8_weierstr(X1,X2,X3),k6_numbers)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v1_xreal_0(k8_weierstr(X1,X2,X3)))).#trainneg
cnf(i_0_887, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),k6_weierstr(esk1_0,X1))=u1_struct_0(esk1_0)|u1_struct_0(k3_topmetr)=k6_numbers|~m1_subset_1(k6_weierstr(esk1_0,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1045, plain, (k4_xxreal_0(k8_weierstr(X1,X2,X3),k8_weierstr(X1,X3,X2))=k1_hausdorf(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(k8_weierstr(X1,X3,X2),k1_numbers)|~m1_subset_1(k8_weierstr(X1,X2,X3),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_428, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,u1_pre_topc(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1009, plain, (esk19_1(X1)=k6_numbers|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1017, plain, (v1_xcmplx_0(k4_ordinal1)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_312, plain, (v1_xreal_0(esk19_1(X1))|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_390, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k4_ordinal1))).#trainneg
cnf(i_0_869, negated_conjecture, (v3_membered(esk2_0)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_870, negated_conjecture, (v2_membered(esk2_0)|~v2_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_841, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_807, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_313, plain, (v1_xreal_0(esk16_1(X1))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_315, plain, (v1_xreal_0(esk14_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_959, plain, (r1_xxreal_0(k6_numbers,k6_numbers)|v1_xboole_0(X2)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~r2_hidden(X3,X2)|~v4_pre_topc(X2,k3_pcomps_1(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_1059, plain, (k3_funct_2(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3)=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|~r1_xxreal_0(k3_funct_2(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3),k6_numbers)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v1_xreal_0(k3_funct_2(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr),k6_weierstr(X1,X2),X3)))).#trainneg
cnf(i_0_937, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_funct_1(k6_weierstr(esk1_0,X1),X2))|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~v1_funct_1(k6_weierstr(esk1_0,X1))|~v1_funct_2(k6_weierstr(esk1_0,X1),u1_struct_0(esk1_0),u1_struct_0(k3_topmetr))|~m1_subset_1(k6_weierstr(esk1_0,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr)))))).#trainneg
cnf(i_0_953, negated_conjecture, (v4_relat_1(k6_weierstr(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_901, plain, (v1_xreal_0(k2_seq_4(X1))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1002, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,u1_pre_topc(k3_pcomps_1(X2)))|~l1_metric_1(X2))).#trainneg
cnf(i_0_359, plain, (v1_xboole_0(esk14_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_804, plain, (v1_xboole_0(u1_struct_0(esk22_0))|~v2_struct_0(esk22_0))).#trainneg
cnf(i_0_320, plain, (v1_xreal_0(k4_seq_4(X1))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_781, plain, (r2_hidden(X1,k1_numbers)|~r2_hidden(X1,k4_ordinal1))).#trainneg
cnf(i_0_1010, plain, (v1_xreal_0(k4_ordinal1)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_827, negated_conjecture, (v1_xxreal_0(esk3_0)|~v2_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_808, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_906, plain, (v1_xreal_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_321, plain, (v1_xreal_0(k2_weierstr(X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_322, plain, (v1_xcmplx_0(esk24_1(X1))|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_882, negated_conjecture, (k2_weierstr(k7_relat_1(k6_weierstr(esk1_0,X1),X2))=k8_weierstr(esk1_0,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k6_weierstr(esk1_0,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr)))))).#trainneg
cnf(i_0_936, negated_conjecture, (k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2)=k6_numbers|v1_xboole_0(X1)|~r1_xxreal_0(k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2),k6_numbers)|~v1_xxreal_0(k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_991, plain, (X1=esk5_2(X2,X3)|r1_tarski(X2,X3)|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|X1!=k2_seq_4(X2)|~v3_membered(X2)|~v1_xreal_0(esk5_2(X2,X3))|~v1_xreal_0(X1)|~v4_xxreal_2(X2))).#trainneg
cnf(i_0_914, plain, (v1_xboole_0(u1_struct_0(k3_pcomps_1(X1)))|~v2_struct_0(k3_pcomps_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_880, negated_conjecture, (k4_seq_4(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2))=k8_weierstr(esk1_0,X1,X2)|v1_xboole_0(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2))|~m1_subset_1(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2),k1_zfmisc_1(u1_struct_0(k3_topmetr)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_791, plain, (X1=X2|~r1_tarski(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_813, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_908, plain, (v1_xboole_0(u1_struct_0(esk30_0))|~v2_struct_0(esk30_0))).#trainneg
cnf(i_0_413, plain, (X1=u1_struct_0(X2)|g1_pre_topc(X1,X3)!=X2|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v1_pre_topc(X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_323, plain, (v1_xcmplx_0(esk19_1(X1))|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_828, negated_conjecture, (v1_xreal_0(esk3_0)|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_829, negated_conjecture, (v1_xcmplx_0(esk3_0)|~v1_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_468, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_324, plain, (v1_xcmplx_0(esk16_1(X1))|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_326, plain, (v1_xcmplx_0(esk14_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_645, plain, (r1_xxreal_0(esk9_3(X1,X2,X3),X4)|r1_xxreal_0(X3,k6_numbers)|v1_xboole_0(X1)|X4!=k2_seq_4(X1)|~v3_membered(X1)|~v1_xreal_0(X4)|~v1_xreal_0(esk9_3(X1,X2,X3))|~v4_xxreal_2(X1)|X2!=k2_seq_4(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_818, plain, (m1_subset_1(u1_struct_0(k3_topmetr),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1281, negated_conjecture, (r1_xxreal_0(k6_numbers,X4)|v1_xboole_0(X1)|~v1_funct_1(k6_weierstr(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk6_4(k6_weierstr(esk1_0,X1),X2,X3,X4),u1_struct_0(esk1_0))|X3!=k7_relat_1(k6_weierstr(esk1_0,X1),X2)|~v1_relat_1(k6_weierstr(esk1_0,X1))|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_1110, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k6_numbers))).#trainneg
cnf(i_0_1142, plain, (u1_struct_0(esk20_0)=k6_numbers)).#trainneg
cnf(i_0_1115, plain, (v3_valued_0(k6_numbers))).#trainneg
cnf(i_0_411, plain, (X1=u1_pre_topc(X2)|g1_pre_topc(X3,X1)!=X2|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X3)))|~v1_pre_topc(X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_331, plain, (v1_xcmplx_0(k4_seq_4(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1166, plain, (r1_xxreal_0(k6_numbers,k6_numbers))).#trainneg
cnf(i_0_1156, negated_conjecture, (v2_membered(esk3_0)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1220, negated_conjecture, (v2_membered(esk2_0)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1158, plain, (v3_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_332, plain, (v1_xcmplx_0(k2_weierstr(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_333, plain, (v7_ordinal1(esk24_1(X1))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1087, negated_conjecture, (m1_subset_1(esk5_2(esk2_0,X1),u1_struct_0(esk1_0))|r1_tarski(esk2_0,X1))).#trainneg
cnf(i_0_1136, negated_conjecture, (k2_seq_4(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2))=k8_weierstr(esk1_0,X1,X2)|v1_xboole_0(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2))|~m1_subset_1(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2),k1_zfmisc_1(u1_struct_0(k3_topmetr)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1365, negated_conjecture, (m1_subset_1(esk6_4(X1,esk3_0,X2,X3),u1_struct_0(esk1_0))|X2!=k7_relat_1(X1,esk3_0)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_1128, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~r2_hidden(esk5_2(X1,u1_struct_0(esk1_0)),esk2_0))).#trainneg
cnf(i_0_1086, negated_conjecture, (m1_subset_1(esk6_4(X1,esk2_0,X2,X3),u1_struct_0(esk1_0))|X2!=k7_relat_1(X1,esk2_0)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_1144, plain, (u1_struct_0(esk18_0)=k6_numbers)).#trainneg
cnf(i_0_602, plain, (~v1_xboole_0(k9_xtuple_0(X1))|X3!=k7_relat_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_334, plain, (v7_ordinal1(esk19_1(X1))|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1163, plain, (r1_xxreal_0(k4_ordinal1,k4_ordinal1)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_830, negated_conjecture, (v7_ordinal1(esk3_0)|~v6_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_335, plain, (v7_ordinal1(esk16_1(X1))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_337, plain, (v7_ordinal1(esk14_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_1191, negated_conjecture, (r1_xxreal_0(k6_numbers,X1)|v1_xboole_0(X2)|X3!=k7_relat_1(k6_weierstr(esk1_0,X2),X4)|~r2_hidden(X1,X3)|~m1_subset_1(esk6_4(k6_weierstr(esk1_0,X2),X4,X3,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1178, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~r2_hidden(esk5_2(X1,u1_struct_0(esk1_0)),esk3_0))).#trainneg
cnf(i_0_1192, negated_conjecture, (v4_xxreal_2(k7_relat_1(k6_weierstr(esk1_0,X1),X2))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v2_compts_1(X2,k3_pcomps_1(esk1_0))|~v2_compts_1(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k7_relat_1(k6_weierstr(esk1_0,X1),X2),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1193, negated_conjecture, (k9_xtuple_0(k6_weierstr(esk1_0,X1))=u1_struct_0(esk1_0)|u1_struct_0(k3_topmetr)=k6_numbers|~v4_relat_1(k6_weierstr(esk1_0,X1),u1_struct_0(esk1_0))|~v1_relat_1(k6_weierstr(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1277, negated_conjecture, (k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2)=k6_numbers|v1_xboole_0(X1)|~v1_xreal_0(k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2))|~r1_xxreal_0(k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2),k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1234, plain, (k4_xxreal_0(X1,X1)=X1|~v1_xxreal_0(X1)|~v3_membered(X2))).#trainneg
cnf(i_0_1292, plain, (esk14_1(k1_zfmisc_1(X1))=k6_numbers|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1305, negated_conjecture, (v1_xxreal_0(esk3_0)|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1339, negated_conjecture, (~v1_funct_1(k6_weierstr(esk1_0,X1))|~m1_subset_1(k6_weierstr(esk1_0,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr))))|~v1_xboole_0(k7_relat_1(k6_weierstr(esk1_0,X1),u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_342, plain, (v7_ordinal1(k4_seq_4(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1225, plain, (r1_xxreal_0(k6_numbers,k4_ordinal1)|r1_xxreal_0(k4_ordinal1,k6_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_864, negated_conjecture, (v1_xxreal_0(esk2_0)|~v2_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_865, negated_conjecture, (v1_xreal_0(esk2_0)|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1306, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_343, plain, (v7_ordinal1(k2_weierstr(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_1366, negated_conjecture, (m1_subset_1(esk5_2(esk3_0,X1),u1_struct_0(esk1_0))|r1_tarski(esk3_0,X1))).#trainneg
cnf(i_0_1308, plain, (k4_xxreal_0(X1,X1)=X1|~v1_xxreal_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_1279, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,X1),X2)=k6_numbers|v1_xboole_0(X1)|~r1_xxreal_0(k1_funct_1(k6_weierstr(esk1_0,X1),X2),k6_numbers)|~v1_xxreal_0(k1_funct_1(k6_weierstr(esk1_0,X1),X2))|~v1_funct_1(k6_weierstr(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_693, plain, (k2_square_1(X1,X1)=X1|v2_struct_0(X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_1238, plain, (k1_funct_1(k6_weierstr(X1,X2),X3)=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|~r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),k6_numbers)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_valued_0(k6_weierstr(X1,X2))|~v1_funct_1(k6_weierstr(X1,X2))|~v1_relat_1(k6_weierstr(X1,X2)))).#trainneg
cnf(i_0_382, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk24_1(X2)))).#trainneg
cnf(i_0_1343, plain, (k2_zfmisc_1(X1,X2)=k6_numbers|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1351, negated_conjecture, (v1_xcmplx_0(esk3_0)|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_346, plain, (m1_subset_1(esk5_2(X1,X2),X1)|r1_tarski(X1,X2))).#trainneg
cnf(i_0_1101, plain, (v2_valued_0(k6_numbers)|~v2_membered(X1))).#trainneg
cnf(i_0_866, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v1_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_867, negated_conjecture, (v7_ordinal1(esk2_0)|~v6_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1105, plain, (v3_membered(k6_numbers)|~v3_membered(X1))).#trainneg
cnf(i_0_1106, plain, (v2_membered(k6_numbers)|~v2_membered(X1))).#trainneg
cnf(i_0_349, plain, (v1_relat_1(esk24_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_350, plain, (v1_relat_1(esk19_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1235, plain, (X5=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|~r1_xxreal_0(X5,k6_numbers)|~v1_xxreal_0(X5)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(esk6_4(k6_weierstr(X1,X2),X3,X4,X5),u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|X4!=k7_relat_1(k6_weierstr(X1,X2),X3)|~v1_funct_1(k6_weierstr(X1,X2))|~v1_relat_1(k6_weierstr(X1,X2))|~r2_hidden(X5,X4))).#trainneg
cnf(i_0_1327, negated_conjecture, (v1_xreal_0(k8_weierstr(esk1_0,X1,X2))|v1_xboole_0(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2))|~m1_subset_1(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2),k1_zfmisc_1(u1_struct_0(k3_topmetr)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_881, negated_conjecture, (m1_subset_1(k8_weierstr(esk1_0,X1,X2),k1_numbers)|~m1_subset_1(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2),k1_zfmisc_1(u1_struct_0(k3_topmetr)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1329, negated_conjecture, (k4_seq_4(k7_relat_1(k6_weierstr(esk1_0,X1),X2))=k8_weierstr(esk1_0,X1,X2)|v1_xboole_0(k7_relat_1(k6_weierstr(esk1_0,X1),X2))|~m1_subset_1(k7_relat_1(k6_weierstr(esk1_0,X1),X2),k1_zfmisc_1(u1_struct_0(k3_topmetr)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k6_weierstr(esk1_0,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr)))))).#trainneg
cnf(i_0_1462, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,X3),esk2_0)|~v1_relat_1(k6_weierstr(esk1_0,X3))|~r2_hidden(esk6_4(k6_weierstr(esk1_0,X3),esk2_0,X2,X1),X3)|~r2_hidden(X1,X2)|~v4_pre_topc(X3,k3_pcomps_1(esk1_0))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_1(k6_weierstr(esk1_0,X3)))).#trainneg
cnf(i_0_1452, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,X3),esk3_0)|~v1_relat_1(k6_weierstr(esk1_0,X3))|~r2_hidden(esk6_4(k6_weierstr(esk1_0,X3),esk3_0,X2,X1),X3)|~r2_hidden(X1,X2)|~v4_pre_topc(X3,k3_pcomps_1(esk1_0))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_1(k6_weierstr(esk1_0,X3)))).#trainneg
cnf(i_0_1488, negated_conjecture, (r1_xxreal_0(k6_numbers,X1)|v1_xboole_0(X2)|X3!=k7_relat_1(k6_weierstr(esk1_0,X2),esk3_0)|~r2_hidden(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_1(k6_weierstr(esk1_0,X2))|~v1_relat_1(k6_weierstr(esk1_0,X2)))).#trainneg
cnf(i_0_1489, negated_conjecture, (r1_xxreal_0(k6_numbers,X1)|v1_xboole_0(X2)|X3!=k7_relat_1(k6_weierstr(esk1_0,X2),esk2_0)|~r2_hidden(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_1(k6_weierstr(esk1_0,X2))|~v1_relat_1(k6_weierstr(esk1_0,X2)))).#trainneg
cnf(i_0_1598, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1636, plain, (v3_membered(k6_numbers))).#trainneg
cnf(i_0_351, plain, (v1_relat_1(esk16_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_353, plain, (v1_relat_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_364, plain, (v3_membered(esk24_1(X1))|v1_xboole_0(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_1382, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,u1_struct_0(k3_topmetr)))).#trainneg
cnf(i_0_1491, negated_conjecture, (r1_tarski(esk3_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1438, negated_conjecture, (r1_tarski(esk2_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1451, negated_conjecture, (r2_hidden(esk6_4(k6_weierstr(esk1_0,X1),esk3_0,X2,X3),X1)|v1_xboole_0(X1)|X2!=k7_relat_1(k6_weierstr(esk1_0,X1),esk3_0)|X3!=k6_numbers|~r2_hidden(X3,X2)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_1(k6_weierstr(esk1_0,X1))|~v1_relat_1(k6_weierstr(esk1_0,X1)))).#trainneg
cnf(i_0_1658, negated_conjecture, (k2_seq_4(k7_relat_1(k6_weierstr(esk1_0,X1),X2))=k8_weierstr(esk1_0,X1,X2)|v1_xboole_0(k7_relat_1(k6_weierstr(esk1_0,X1),X2))|~m1_subset_1(k7_relat_1(k6_weierstr(esk1_0,X1),X2),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k7_relat_1(k6_weierstr(esk1_0,X1),X2),k1_zfmisc_1(u1_struct_0(k3_topmetr)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1440, plain, (X5=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_xreal_0(X5)|~r1_xxreal_0(X5,k6_numbers)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(esk6_4(k6_weierstr(X1,X2),X3,X4,X5),u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|X4!=k7_relat_1(k6_weierstr(X1,X2),X3)|~v1_funct_1(k6_weierstr(X1,X2))|~v1_relat_1(k6_weierstr(X1,X2))|~r2_hidden(X5,X4))).#trainneg
cnf(i_0_1564, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,X1),X2)=k6_numbers|v1_xboole_0(X1)|~v1_funct_1(k6_weierstr(esk1_0,X1))|~r1_xxreal_0(k1_funct_1(k6_weierstr(esk1_0,X1),X2),k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~v2_valued_0(k6_weierstr(esk1_0,X1))|~v1_relat_1(k6_weierstr(esk1_0,X1)))).#trainneg
cnf(i_0_1581, plain, (X5=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|~v2_valued_0(k6_weierstr(X1,X2))|~r1_xxreal_0(X5,k6_numbers)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(esk6_4(k6_weierstr(X1,X2),X3,X4,X5),u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|X4!=k7_relat_1(k6_weierstr(X1,X2),X3)|~v1_funct_1(k6_weierstr(X1,X2))|~v1_relat_1(k6_weierstr(X1,X2))|~r2_hidden(X5,X4))).#trainneg
cnf(i_0_1461, negated_conjecture, (r2_hidden(esk6_4(k6_weierstr(esk1_0,X1),esk2_0,X2,X3),X1)|v1_xboole_0(X1)|X2!=k7_relat_1(k6_weierstr(esk1_0,X1),esk2_0)|X3!=k6_numbers|~r2_hidden(X3,X2)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_1(k6_weierstr(esk1_0,X1))|~v1_relat_1(k6_weierstr(esk1_0,X1)))).#trainneg
cnf(i_0_365, plain, (v3_membered(esk19_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_1633, plain, (v2_valued_0(k6_numbers)|~v3_membered(X1))).#trainneg
cnf(i_0_1642, negated_conjecture, (v2_membered(k6_numbers)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_366, plain, (v3_membered(esk16_1(X1))|v1_xboole_0(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_368, plain, (v3_membered(esk14_1(k1_zfmisc_1(X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_1626, plain, (v1_xreal_0(esk5_2(X1,X2))|r1_tarski(X1,X2)|~v3_membered(X1))).#trainneg
cnf(i_0_1561, negated_conjecture, (X4=k6_numbers|v1_xboole_0(X1)|~v1_funct_1(k6_weierstr(esk1_0,X1))|~r1_xxreal_0(X4,k6_numbers)|~v1_xxreal_0(X4)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk6_4(k6_weierstr(esk1_0,X1),X2,X3,X4),u1_struct_0(esk1_0))|X3!=k7_relat_1(k6_weierstr(esk1_0,X1),X2)|~v1_relat_1(k6_weierstr(esk1_0,X1))|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_1650, negated_conjecture, (X1=k6_numbers|v1_xboole_0(X2)|X3!=k7_relat_1(k6_weierstr(esk1_0,X2),X4)|~r1_xxreal_0(X1,k6_numbers)|~v1_xxreal_0(X1)|~r2_hidden(X1,X3)|~m1_subset_1(esk6_4(k6_weierstr(esk1_0,X2),X4,X3,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1539, plain, (X1=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1503, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,X1),X2)=k6_numbers|v1_xboole_0(X1)|~v1_xreal_0(k1_funct_1(k6_weierstr(esk1_0,X1),X2))|~r1_xxreal_0(k1_funct_1(k6_weierstr(esk1_0,X1),X2),k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~v1_funct_1(k6_weierstr(esk1_0,X1))|~v1_funct_2(k6_weierstr(esk1_0,X1),u1_struct_0(esk1_0),u1_struct_0(k3_topmetr))|~m1_subset_1(k6_weierstr(esk1_0,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr)))))).#trainneg
cnf(i_0_1383, plain, (v1_xboole_0(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,k1_numbers,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,X1)))|~v1_xboole_0(k7_relat_1(X2,u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_384, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk16_1(X2)))).#trainneg
cnf(i_0_811, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_373, plain, (v2_membered(esk24_1(X1))|v1_xboole_0(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_1645, plain, (v2_membered(k6_numbers)|~v3_membered(X1))).#trainneg
cnf(i_0_1167, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_374, plain, (v2_membered(esk19_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_375, plain, (v2_membered(esk16_1(X1))|v1_xboole_0(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_1284, plain, (k4_xxreal_0(X1,X1)=X1|~v1_xxreal_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1654, negated_conjecture, (v1_xreal_0(k8_weierstr(esk1_0,X1,X2))|v1_xboole_0(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k6_weierstr(esk1_0,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr)))))).#trainneg
cnf(i_0_1332, plain, (X1=k6_numbers|v1_funct_2(k6_weierstr(X2,X3),u1_struct_0(k3_pcomps_1(X2)),X1)|u1_struct_0(k3_topmetr)=k6_numbers|v2_struct_0(X2)|~m1_subset_1(k6_weierstr(X2,X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_838, negated_conjecture, (v1_xboole_0(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,u1_struct_0(esk1_0),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X1)))|~v1_xboole_0(k7_relat_1(X2,esk3_0)))).#trainneg
cnf(i_0_904, plain, (k2_square_1(X1,X1)=X1|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1360, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_569, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_1764, plain, (v2_valued_0(k6_numbers))).#trainneg
cnf(i_0_1847, plain, (v2_membered(k6_numbers))).#trainneg
cnf(i_0_1669, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,esk2_0),esk2_0)|~v1_funct_1(k6_weierstr(esk1_0,esk2_0))|~r2_hidden(X1,X2)|~v4_pre_topc(esk2_0,k3_pcomps_1(esk1_0))|~v1_relat_1(k6_weierstr(esk1_0,esk2_0)))).#trainneg
cnf(i_0_377, plain, (v2_membered(esk14_1(k1_zfmisc_1(X1)))|~v2_membered(X1))).#trainneg
cnf(i_0_1907, plain, (r1_xxreal_0(k6_numbers,esk23_0)|r1_xxreal_0(esk23_0,k6_numbers))).#trainneg
cnf(i_0_1530, negated_conjecture, (v1_xxreal_0(esk2_0)|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1634, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_391, plain, (v1_xxreal_0(esk4_1(X1))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_392, plain, (v1_xreal_0(esk4_1(X1))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_1674, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,esk3_0),esk3_0)|~v1_funct_1(k6_weierstr(esk1_0,esk3_0))|~r2_hidden(X1,X2)|~v4_pre_topc(esk3_0,k3_pcomps_1(esk1_0))|~v1_relat_1(k6_weierstr(esk1_0,esk3_0)))).#trainneg
cnf(i_0_1734, negated_conjecture, (X1=k6_numbers|v1_xboole_0(X2)|X3!=k7_relat_1(k6_weierstr(esk1_0,X2),X4)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k6_numbers)|~r2_hidden(X1,X3)|~m1_subset_1(esk6_4(k6_weierstr(esk1_0,X2),X4,X3,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1890, negated_conjecture, (~v1_funct_1(k6_weierstr(esk1_0,X1))|~m1_subset_1(k6_weierstr(esk1_0,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr))))|~v1_xboole_0(k7_relat_1(k6_weierstr(esk1_0,X1),esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1667, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,u1_struct_0(esk1_0)),esk2_0)|~v1_funct_1(k6_weierstr(esk1_0,u1_struct_0(esk1_0)))|~r2_hidden(X1,X2)|~v4_pre_topc(u1_struct_0(esk1_0),k3_pcomps_1(esk1_0))|~r2_hidden(esk6_4(k6_weierstr(esk1_0,u1_struct_0(esk1_0)),esk2_0,X2,X1),esk2_0))).#trainneg
cnf(i_0_1666, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,u1_struct_0(esk1_0)),esk2_0)|~v1_funct_1(k6_weierstr(esk1_0,u1_struct_0(esk1_0)))|~r2_hidden(X1,X2)|~v4_pre_topc(u1_struct_0(esk1_0),k3_pcomps_1(esk1_0))|~r2_hidden(esk6_4(k6_weierstr(esk1_0,u1_struct_0(esk1_0)),esk2_0,X2,X1),esk3_0))).#trainneg
cnf(i_0_1672, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,u1_struct_0(esk1_0)),esk3_0)|~v1_funct_1(k6_weierstr(esk1_0,u1_struct_0(esk1_0)))|~r2_hidden(X1,X2)|~v4_pre_topc(u1_struct_0(esk1_0),k3_pcomps_1(esk1_0))|~r2_hidden(esk6_4(k6_weierstr(esk1_0,u1_struct_0(esk1_0)),esk3_0,X2,X1),esk2_0))).#trainneg
cnf(i_0_1806, plain, (X1=k6_numbers|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_1465, plain, (m1_subset_1(esk31_1(esk18_0),k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_393, plain, (v1_xcmplx_0(esk4_1(X1))|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_1798, plain, (k4_ordinal1=k1_numbers|~m1_subset_1(k1_numbers,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_1016, plain, (v1_xreal_0(k4_ordinal1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1155, negated_conjecture, (v3_membered(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1168, plain, (r1_xxreal_0(X1,X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_394, plain, (v7_ordinal1(esk4_1(X1))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_1671, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,u1_struct_0(esk1_0)),esk3_0)|~v1_funct_1(k6_weierstr(esk1_0,u1_struct_0(esk1_0)))|~r2_hidden(X1,X2)|~v4_pre_topc(u1_struct_0(esk1_0),k3_pcomps_1(esk1_0))|~r2_hidden(esk6_4(k6_weierstr(esk1_0,u1_struct_0(esk1_0)),esk3_0,X2,X1),esk3_0))).#trainneg
cnf(i_0_1803, plain, (k7_relset_1(X2,X1,X3,X4)=X1|~m1_subset_1(X1,k1_zfmisc_1(k7_relset_1(X2,X1,X3,X4)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1))))).#trainneg
cnf(i_0_1673, negated_conjecture, (k6_numbers=X1|v1_xboole_0(X3)|X2!=k7_relat_1(k6_weierstr(esk1_0,X3),esk3_0)|~v1_funct_1(k6_weierstr(esk1_0,X3))|~r2_hidden(X1,X2)|~v4_pre_topc(X3,k3_pcomps_1(esk1_0))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk6_4(k6_weierstr(esk1_0,X3),esk3_0,X2,X1),X3))).#trainneg
cnf(i_0_1668, negated_conjecture, (k6_numbers=X1|v1_xboole_0(X3)|X2!=k7_relat_1(k6_weierstr(esk1_0,X3),esk2_0)|~v1_funct_1(k6_weierstr(esk1_0,X3))|~r2_hidden(X1,X2)|~v4_pre_topc(X3,k3_pcomps_1(esk1_0))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk6_4(k6_weierstr(esk1_0,X3),esk2_0,X2,X1),X3))).#trainneg
cnf(i_0_1783, negated_conjecture, (X1=k6_numbers|v1_xboole_0(X2)|X3!=k7_relat_1(k6_weierstr(esk1_0,X2),esk3_0)|~r1_xxreal_0(X1,k6_numbers)|~v1_xxreal_0(X1)|~r2_hidden(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_1(k6_weierstr(esk1_0,X2))|~v1_relat_1(k6_weierstr(esk1_0,X2)))).#trainneg
cnf(i_0_1784, negated_conjecture, (X1=k6_numbers|v1_xboole_0(X2)|X3!=k7_relat_1(k6_weierstr(esk1_0,X2),esk2_0)|~r1_xxreal_0(X1,k6_numbers)|~v1_xxreal_0(X1)|~r2_hidden(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_1(k6_weierstr(esk1_0,X2))|~v1_relat_1(k6_weierstr(esk1_0,X2)))).#trainneg
cnf(i_0_1880, negated_conjecture, (k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2)=k6_numbers|v1_xreal_0(k8_weierstr(esk1_0,X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1703, negated_conjecture, (u1_struct_0(esk1_0)=esk3_0|~r1_tarski(u1_struct_0(esk1_0),esk3_0))).#trainneg
cnf(i_0_470, plain, (v1_xboole_0(esk19_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_396, plain, (v3_membered(esk4_1(X1))|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_1218, plain, (m1_subset_1(X1,k1_numbers)|v1_xboole_0(k4_ordinal1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_1219, negated_conjecture, (v3_membered(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_834, negated_conjecture, (v3_valued_0(esk3_0)|~v3_valued_0(u1_struct_0(esk1_0))|~v1_relat_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1596, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1602, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_397, plain, (v2_membered(esk4_1(X1))|v1_xboole_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_1739, negated_conjecture, (X1=k6_numbers|v1_xboole_0(X2)|X3!=k7_relat_1(k6_weierstr(esk1_0,X2),X4)|~v2_valued_0(k6_weierstr(esk1_0,X2))|~r1_xxreal_0(X1,k6_numbers)|~r2_hidden(X1,X3)|~m1_subset_1(esk6_4(k6_weierstr(esk1_0,X2),X4,X3,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1311, negated_conjecture, (v1_xreal_0(k8_weierstr(esk1_0,X1,X2))|~m1_subset_1(k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr),k6_weierstr(esk1_0,X1),X2),k1_zfmisc_1(u1_struct_0(k3_topmetr)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1001, negated_conjecture, (k2_square_1(X1,X1)=X1|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_598, plain, (r1_xxreal_0(esk6_4(X1,X2,X3,X4),X5)|X5!=k2_seq_4(X2)|~v3_membered(X2)|~v1_xreal_0(X5)|~v1_xreal_0(esk6_4(X1,X2,X3,X4))|~v4_xxreal_2(X2)|X3!=k7_relat_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_1195, negated_conjecture, (X1=k6_numbers|v1_funct_2(k6_weierstr(esk1_0,X2),u1_struct_0(esk1_0),X1)|u1_struct_0(k3_topmetr)=k6_numbers|~m1_subset_1(k6_weierstr(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_875, negated_conjecture, (v1_xboole_0(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,u1_struct_0(esk1_0),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X1)))|~v1_xboole_0(k7_relat_1(X2,esk2_0)))).#trainneg
cnf(i_0_1702, negated_conjecture, (u1_struct_0(esk1_0)=esk3_0|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_812, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1378, plain, (v2_membered(u1_struct_0(k3_topmetr))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1707, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_400, plain, (v3_valued_0(esk24_1(X1))|v1_xboole_0(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_871, negated_conjecture, (v3_valued_0(esk2_0)|~v3_valued_0(u1_struct_0(esk1_0))|~v1_relat_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1350, negated_conjecture, (v1_xreal_0(esk3_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1599, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1094, plain, (v7_ordinal1(k6_numbers)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_401, plain, (v3_valued_0(esk19_1(X1))|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_402, plain, (v3_valued_0(esk16_1(X1))|v1_xboole_0(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1624, plain, (r1_xxreal_0(esk5_2(X1,X2),X3)|v1_xboole_0(X1)|r1_tarski(X1,X2)|X3!=k2_seq_4(X1)|~v3_membered(X1)|~v1_xreal_0(X3)|~v4_xxreal_2(X1))).#trainneg
cnf(i_0_690, plain, (v1_xreal_0(k1_hausdorf(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_1657, negated_conjecture, (m1_subset_1(k8_weierstr(esk1_0,X1,X2),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k6_weierstr(esk1_0,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr)))))).#trainneg
cnf(i_0_1744, negated_conjecture, (m1_subset_1(esk6_4(k6_weierstr(esk1_0,esk3_0),esk2_0,X1,X2),u1_struct_0(esk1_0))|X1!=k7_relat_1(k6_weierstr(esk1_0,esk3_0),esk2_0)|X2!=k6_numbers|~v1_funct_1(k6_weierstr(esk1_0,esk3_0))|~r2_hidden(X2,X1)|~v4_pre_topc(esk3_0,k3_pcomps_1(esk1_0)))).#trainneg
cnf(i_0_1583, negated_conjecture, (~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_xboole_0(k7_relat_1(k6_weierstr(esk1_0,X1),u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1504, plain, (k4_xxreal_0(X1,X1)=X1|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1316, negated_conjecture, (v1_xreal_0(k8_weierstr(esk1_0,X1,X2))|~m1_subset_1(k7_relat_1(k6_weierstr(esk1_0,X1),X2),k1_zfmisc_1(u1_struct_0(k3_topmetr)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_698, plain, (k6_weierstr(X1,X2)=k6_numbers|v2_struct_0(X1)|u1_struct_0(k3_topmetr)!=k6_numbers|~m1_subset_1(k6_weierstr(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_topmetr))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_2016, plain, (esk19_1(k6_numbers)=k6_numbers)).#trainneg
cnf(i_0_2034, plain, (esk31_1(esk18_0)=k6_numbers)).#trainneg
cnf(i_0_2018, plain, (esk14_1(k1_zfmisc_1(k6_numbers))=k6_numbers)).#trainneg
cnf(i_0_2010, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,u1_struct_0(esk1_0)),esk2_0)|~v1_funct_1(k6_weierstr(esk1_0,u1_struct_0(esk1_0)))|~r2_hidden(X1,X2)|~v4_pre_topc(u1_struct_0(esk1_0),k3_pcomps_1(esk1_0))|~v1_relat_1(k6_weierstr(esk1_0,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_404, plain, (v3_valued_0(esk14_1(k1_zfmisc_1(X1)))|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2044, plain, (r1_xxreal_0(k6_numbers,k4_ordinal1)|r1_xxreal_0(k4_ordinal1,k6_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1531, negated_conjecture, (v1_xreal_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1372, plain, (v1_xxreal_0(u1_struct_0(k3_topmetr))|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2098, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2104, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_406, plain, (v3_valued_0(esk4_1(X1))|v1_xboole_0(X1)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_410, plain, (u1_pre_topc(X1)=X2|X1!=g1_pre_topc(X3,X2)|~m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2156, negated_conjecture, (m1_subset_1(k7_relat_1(k6_weierstr(esk1_0,X1),X2),k1_zfmisc_1(u1_struct_0(k3_topmetr)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2199, plain, (k8_weierstr(X1,X2,X2)=k1_hausdorf(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v1_xxreal_0(k8_weierstr(X1,X2,X2)))).#trainneg
cnf(i_0_2132, negated_conjecture, (k8_weierstr(X1,X2,X2)=k1_hausdorf(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(k8_weierstr(X1,X2,X2),k1_numbers))).#trainneg
cnf(i_0_2050, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,u1_struct_0(esk1_0)),esk3_0)|~v1_funct_1(k6_weierstr(esk1_0,u1_struct_0(esk1_0)))|~r2_hidden(X1,X2)|~v4_pre_topc(u1_struct_0(esk1_0),k3_pcomps_1(esk1_0))|~v1_relat_1(k6_weierstr(esk1_0,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2130, negated_conjecture, (v1_xreal_0(k8_weierstr(esk1_0,X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2087, plain, (v3_membered(X1)|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_1375, plain, (v7_ordinal1(u1_struct_0(k3_topmetr))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_412, plain, (u1_struct_0(X1)=X2|X1!=g1_pre_topc(X2,X3)|~m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1303, plain, (r2_hidden(X1,k1_numbers)|v1_xboole_0(k4_ordinal1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_1373, plain, (v1_xreal_0(u1_struct_0(k3_topmetr))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1374, plain, (v1_xcmplx_0(u1_struct_0(k3_topmetr))|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1361, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2060, plain, (r2_hidden(esk6_4(k6_weierstr(X1,X2),X3,X4,X5),u1_struct_0(k3_pcomps_1(X1)))|u1_struct_0(k3_topmetr)=k6_numbers|v2_struct_0(X1)|X4!=k7_relat_1(k6_weierstr(X1,X2),X3)|~v1_funct_1(k6_weierstr(X1,X2))|~v1_relat_1(k6_weierstr(X1,X2))|~r2_hidden(X5,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_2159, plain, (m1_subset_1(k7_relat_1(k2_zfmisc_1(X1,X2),X3),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_2115, negated_conjecture, (X1=k6_numbers|v1_xboole_0(X2)|X3!=k7_relat_1(k6_weierstr(esk1_0,X2),esk3_0)|~v2_valued_0(k6_weierstr(esk1_0,X2))|~r1_xxreal_0(X1,k6_numbers)|~r2_hidden(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_1(k6_weierstr(esk1_0,X2))|~v1_relat_1(k6_weierstr(esk1_0,X2)))).#trainneg
cnf(i_0_1991, negated_conjecture, (X1=k6_numbers|v1_xboole_0(X2)|X3!=k7_relat_1(k6_weierstr(esk1_0,X2),esk3_0)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k6_numbers)|~r2_hidden(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_1(k6_weierstr(esk1_0,X2))|~v1_relat_1(k6_weierstr(esk1_0,X2)))).#trainneg
cnf(i_0_1948, negated_conjecture, (k6_numbers=X1|~v1_funct_1(k6_weierstr(esk1_0,esk2_0))|~v1_relat_1(k6_weierstr(esk1_0,esk2_0))|~r2_hidden(X1,k7_relat_1(k6_weierstr(esk1_0,esk2_0),esk2_0))|~v4_pre_topc(esk2_0,k3_pcomps_1(esk1_0)))).#trainneg
cnf(i_0_2162, negated_conjecture, (~v1_funct_1(k6_weierstr(esk1_0,X1))|~m1_subset_1(k6_weierstr(esk1_0,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_topmetr))))|~v1_xboole_0(k7_relat_1(k6_weierstr(esk1_0,X1),esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2061, plain, (v1_xboole_0(k6_weierstr(X1,X2))|v1_xboole_0(X3)|u1_struct_0(k3_topmetr)=k6_numbers|v2_struct_0(X1)|~v1_relat_1(k6_weierstr(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_1984, plain, (k2_square_1(k6_numbers,k8_weierstr(X1,X3,X2))=k1_hausdorf(X1,X3,X2)|v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X3,X2)!=k6_numbers|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1)))).#trainneg
cnf(i_0_1986, plain, (k2_square_1(k8_weierstr(X1,X2,X3),k6_numbers)=k1_hausdorf(X1,X2,X3)|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X2,X3)!=k6_numbers|~v2_compts_1(X2,k3_pcomps_1(X1))|~v2_compts_1(X3,k3_pcomps_1(X1)))).#trainneg
cnf(i_0_2083, plain, (esk19_1(k2_zfmisc_1(X1,X2))=k6_numbers|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1379, plain, (v3_valued_0(u1_struct_0(k3_topmetr))|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2007, negated_conjecture, (~v1_funct_1(k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_xboole_0(k7_relat_1(k6_numbers,esk3_0))|u1_struct_0(k3_topmetr)!=k6_numbers)).#trainneg
cnf(i_0_419, plain, (v1_xxreal_0(u1_pre_topc(X1))|~v2_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1109, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k6_numbers))).#trainneg
cnf(i_0_1517, negated_conjecture, (r1_xxreal_0(esk3_0,esk3_0)|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1952, negated_conjecture, (r1_xxreal_0(esk2_0,esk2_0)|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1600, plain, (v3_valued_0(X1)|~v3_membered(X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1601, plain, (v2_valued_0(X1)|~v2_membered(X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_420, plain, (v1_xreal_0(u1_pre_topc(X1))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_421, plain, (v1_xcmplx_0(u1_pre_topc(X1))|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1983, plain, (k4_xxreal_0(k6_numbers,k8_weierstr(X1,X3,X2))=k1_hausdorf(X1,X3,X2)|v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X3,X2)!=k6_numbers|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1)))).#trainneg
cnf(i_0_1989, plain, (k4_xxreal_0(k8_weierstr(X1,X2,X3),k6_numbers)=k1_hausdorf(X1,X2,X3)|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X2,X3)!=k6_numbers|~v2_compts_1(X2,k3_pcomps_1(X1))|~v2_compts_1(X3,k3_pcomps_1(X1)))).#trainneg
cnf(i_0_1992, negated_conjecture, (X1=k6_numbers|v1_xboole_0(X2)|X3!=k7_relat_1(k6_weierstr(esk1_0,X2),esk2_0)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k6_numbers)|~r2_hidden(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_1(k6_weierstr(esk1_0,X2))|~v1_relat_1(k6_weierstr(esk1_0,X2)))).#trainneg
cnf(i_0_2374, negated_conjecture, (k8_weierstr(X1,X2,X2)=k1_hausdorf(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_2493, plain, (m1_subset_1(esk6_4(X1,X2,X3,X4),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_topmetr)))|X3!=k7_relat_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_2461, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relat_1(k2_zfmisc_1(X3,X2),X4)))).#trainneg
cnf(i_0_2304, plain, (v4_pre_topc(k6_numbers,esk18_0))).#trainneg
cnf(i_0_2429, negated_conjecture, (k6_numbers=X1|u1_struct_0(k3_topmetr)=k6_numbers|X2!=k7_relat_1(k6_weierstr(esk1_0,u1_struct_0(esk1_0)),esk3_0)|~v1_funct_1(k6_weierstr(esk1_0,u1_struct_0(esk1_0)))|~r2_hidden(X1,X2)|~v4_pre_topc(u1_struct_0(esk1_0),k3_pcomps_1(esk1_0)))).#trainneg
cnf(i_0_422, plain, (v7_ordinal1(u1_pre_topc(X1))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2294, plain, (~v1_subset_1(k6_numbers,k6_numbers))).#trainneg
cnf(i_0_2322, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_424, plain, (v3_membered(u1_pre_topc(X1))|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2494, plain, (m1_subset_1(esk5_2(X1,X2),k1_numbers)|r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_2363, negated_conjecture, (m1_subset_1(k7_relat_1(k6_weierstr(esk1_0,X1),X2),k1_zfmisc_1(k1_numbers))|v1_xboole_0(k7_relat_1(k6_weierstr(esk1_0,X1),X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2528, plain, (r2_hidden(esk5_2(k6_numbers,X1),X2)|r1_tarski(k6_numbers,X1))).#trainneg
cnf(i_0_2477, negated_conjecture, (esk5_2(k7_relat_1(k6_weierstr(esk1_0,esk2_0),esk2_0),X1)=k6_numbers|r1_tarski(k7_relat_1(k6_weierstr(esk1_0,esk2_0),esk2_0),X1)|~v1_funct_1(k6_weierstr(esk1_0,esk2_0))|~v1_relat_1(k6_weierstr(esk1_0,esk2_0))|~v4_pre_topc(esk2_0,k3_pcomps_1(esk1_0)))).#trainneg
cnf(i_0_2167, plain, (k2_zfmisc_1(X1,X2)=k6_numbers|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2321, plain, (v1_xxreal_0(u1_struct_0(k3_topmetr))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2430, negated_conjecture, (k6_numbers=X1|u1_struct_0(k3_topmetr)=k6_numbers|X2!=k7_relat_1(k6_weierstr(esk1_0,u1_struct_0(esk1_0)),esk2_0)|~v1_funct_1(k6_weierstr(esk1_0,u1_struct_0(esk1_0)))|~r2_hidden(X1,X2)|~v4_pre_topc(u1_struct_0(esk1_0),k3_pcomps_1(esk1_0)))).#trainneg
cnf(i_0_425, plain, (v2_membered(u1_pre_topc(X1))|~v2_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2381, plain, (v1_xreal_0(k4_ordinal1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_2385, negated_conjecture, (v3_membered(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_426, plain, (v3_valued_0(u1_pre_topc(X1))|~v3_valued_0(k1_zfmisc_1(u1_struct_0(X1)))|~v1_relat_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_429, plain, (v1_xxreal_0(k2_pcomps_1(X1))|~v2_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_metric_1(X1))).#trainneg
cnf(i_0_2489, plain, (m1_subset_1(esk6_4(X1,X2,X3,X4),k1_numbers)|~m1_subset_1(k9_xtuple_0(X1),k1_zfmisc_1(u1_struct_0(k3_topmetr)))|X3!=k7_relat_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_2116, negated_conjecture, (X1=k6_numbers|v1_xboole_0(X2)|X3!=k7_relat_1(k6_weierstr(esk1_0,X2),esk2_0)|~v2_valued_0(k6_weierstr(esk1_0,X2))|~r1_xxreal_0(X1,k6_numbers)|~r2_hidden(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_1(k6_weierstr(esk1_0,X2))|~v1_relat_1(k6_weierstr(esk1_0,X2)))).#trainneg
cnf(i_0_2353, negated_conjecture, (v3_membered(k7_relat_1(k6_weierstr(esk1_0,X1),X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1985, plain, (k4_xxreal_0(k6_numbers,k8_weierstr(X1,X3,X2))=k1_hausdorf(X1,X2,X3)|v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X3,X2)!=k6_numbers|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1)))).#trainneg
cnf(i_0_2125, plain, (r2_hidden(esk5_2(k7_relset_1(X1,X2,X3,X4),X5),X2)|r1_tarski(k7_relset_1(X1,X2,X3,X4),X5)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2124, plain, (r2_hidden(esk6_4(X1,k7_relset_1(X2,X3,X4,X5),X6,X7),X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|X6!=k7_relat_1(X1,k7_relset_1(X2,X3,X4,X5))|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X7,X6))).#trainneg
cnf(i_0_1972, plain, (k4_xxreal_0(k6_numbers,k8_weierstr(X1,X3,X2))=k1_hausdorf(X1,X3,X2)|v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X2,X3)!=k6_numbers|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1)))).#trainneg
cnf(i_0_1312, plain, (v1_xreal_0(k4_seq_4(X1))|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_430, plain, (v1_xreal_0(k2_pcomps_1(X1))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_metric_1(X1))).#trainneg
cnf(i_0_1923, plain, (r1_xxreal_0(k6_numbers,X1)|r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2386, negated_conjecture, (v3_membered(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_2382, negated_conjecture, (v1_xreal_0(esk3_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_2719, plain, (r1_tarski(k6_numbers,X1))).#trainneg
cnf(i_0_2559, plain, (v2_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_431, plain, (v1_xcmplx_0(k2_pcomps_1(X1))|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_metric_1(X1))).#trainneg
cnf(i_0_432, plain, (v7_ordinal1(k2_pcomps_1(X1))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_metric_1(X1))).#trainneg
cnf(i_0_2816, plain, (r1_tarski(k7_relset_1(X1,X2,X3,X4),X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2849, plain, (r2_hidden(esk6_4(X1,k7_relat_1(X4,X5),X6,X7),X3)|X6!=k7_relat_1(X1,k7_relat_1(X4,X5))|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X7,X6)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2402, plain, (v1_xcmplx_0(u1_struct_0(k3_topmetr))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2629, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,k1_numbers),esk3_0)|~v1_funct_1(k6_weierstr(esk1_0,k1_numbers))|~r2_hidden(X1,X2)|~v4_pre_topc(k1_numbers,k3_pcomps_1(esk1_0))|~m1_subset_1(k1_numbers,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_relat_1(k6_weierstr(esk1_0,k1_numbers))|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_2404, plain, (r1_xxreal_0(X1,esk23_0)|r1_xxreal_0(esk23_0,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2383, negated_conjecture, (v1_xreal_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_2387, plain, (v1_xreal_0(u1_struct_0(k3_topmetr))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2752, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_434, plain, (v3_membered(k2_pcomps_1(X1))|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_metric_1(X1))).#trainneg
cnf(i_0_435, plain, (v2_membered(k2_pcomps_1(X1))|~v2_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_metric_1(X1))).#trainneg
cnf(i_0_2603, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_hausdorf(X1,X2,X2))|v1_xboole_0(X2)|v2_struct_0(X1)|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_2626, plain, (v1_xreal_0(esk6_4(X1,X2,X3,X4))|X3!=k7_relat_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X4,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_2464, negated_conjecture, (X1=k6_numbers|v1_xboole_0(X2)|~v2_valued_0(k6_weierstr(esk1_0,X2))|~v1_funct_1(k6_weierstr(esk1_0,X2))|~r1_xxreal_0(X1,k6_numbers)|~r2_hidden(X1,k7_relat_1(k6_weierstr(esk1_0,X2),esk3_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1973, plain, (k2_square_1(k6_numbers,k8_weierstr(X1,X3,X2))=k1_hausdorf(X1,X3,X2)|v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X2,X3)!=k6_numbers|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1)))).#trainneg
cnf(i_0_1396, plain, (g1_pre_topc(k6_numbers,u1_pre_topc(esk20_0))=esk20_0)).#trainneg
cnf(i_0_2630, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,k1_numbers),esk2_0)|~v1_funct_1(k6_weierstr(esk1_0,k1_numbers))|~r2_hidden(X1,X2)|~v4_pre_topc(k1_numbers,k3_pcomps_1(esk1_0))|~m1_subset_1(k1_numbers,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_relat_1(k6_weierstr(esk1_0,k1_numbers))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_436, plain, (v3_valued_0(k2_pcomps_1(X1))|~v3_valued_0(k1_zfmisc_1(u1_struct_0(X1)))|~v1_relat_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_metric_1(X1))).#trainneg
cnf(i_0_2667, negated_conjecture, (v3_valued_0(esk3_0)|~v1_relat_1(u1_struct_0(esk1_0))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2668, negated_conjecture, (v3_valued_0(esk2_0)|~v1_relat_1(u1_struct_0(esk1_0))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_439, plain, (m1_subset_1(u1_pre_topc(k3_pcomps_1(X1)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_metric_1(X1))).#trainneg
cnf(i_0_442, plain, (v4_relat_1(esk24_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1966, negated_conjecture, (k6_numbers=X1|~v1_funct_1(k6_weierstr(esk1_0,esk3_0))|~v1_relat_1(k6_weierstr(esk1_0,esk3_0))|~r2_hidden(X1,k7_relat_1(k6_weierstr(esk1_0,esk3_0),esk3_0))|~v4_pre_topc(esk3_0,k3_pcomps_1(esk1_0)))).#trainneg
cnf(i_0_1977, plain, (k2_square_1(k8_weierstr(X1,X2,X3),k6_numbers)=k1_hausdorf(X1,X3,X2)|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X3,X2)!=k6_numbers|~v2_compts_1(X2,k3_pcomps_1(X1))|~v2_compts_1(X3,k3_pcomps_1(X1)))).#trainneg
cnf(i_0_1969, plain, (k2_square_1(k6_numbers,k8_weierstr(X1,X3,X2))=k1_hausdorf(X1,X2,X3)|v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X2,X3)!=k6_numbers|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1)))).#trainneg
cnf(i_0_2021, plain, (k3_funct_2(X1,k1_zfmisc_1(k6_numbers),X2,X3)=k6_numbers|v1_xboole_0(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,X1,k1_zfmisc_1(k6_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k6_numbers))))|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_769, plain, (m1_subset_1(k4_seq_4(X1),k1_numbers)|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_1464, plain, (g1_pre_topc(k6_numbers,u1_pre_topc(esk18_0))=esk18_0)).#trainneg
cnf(i_0_2479, negated_conjecture, (~v1_funct_1(k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_xboole_0(k7_relat_1(k6_numbers,esk2_0))|u1_struct_0(k3_topmetr)!=k6_numbers)).#trainneg
cnf(i_0_443, plain, (v4_relat_1(esk19_1(k2_zfmisc_1(X1,X2)),X1))).#trainneg
cnf(i_0_2399, plain, (v1_xreal_0(u1_struct_0(k3_topmetr))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2401, plain, (v1_xreal_0(u1_struct_0(k3_topmetr))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_444, plain, (v4_relat_1(esk16_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_446, plain, (v4_relat_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_1975, plain, (k2_square_1(k8_weierstr(X1,X2,X3),k6_numbers)=k1_hausdorf(X1,X2,X3)|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X3,X2)!=k6_numbers|~v2_compts_1(X2,k3_pcomps_1(X1))|~v2_compts_1(X3,k3_pcomps_1(X1)))).#trainneg
cnf(i_0_1978, plain, (k4_xxreal_0(k8_weierstr(X1,X2,X3),k6_numbers)=k1_hausdorf(X1,X2,X3)|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X3,X2)!=k6_numbers|~v2_compts_1(X2,k3_pcomps_1(X1))|~v2_compts_1(X3,k3_pcomps_1(X1)))).#trainneg
cnf(i_0_2173, plain, (X1=esk5_2(X2,X3)|r1_tarski(X2,X3)|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|X1!=k2_seq_4(X2)|~v3_membered(X2)|~v1_xreal_0(X1)|~v4_xxreal_2(X2))).#trainneg
cnf(i_0_1976, plain, (k4_xxreal_0(k8_weierstr(X1,X2,X3),k6_numbers)=k1_hausdorf(X1,X3,X2)|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X3,X2)!=k6_numbers|~v2_compts_1(X2,k3_pcomps_1(X1))|~v2_compts_1(X3,k3_pcomps_1(X1)))).#trainneg
cnf(i_0_2607, negated_conjecture, (k1_hausdorf(X1,X2,X2)=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|~r1_xxreal_0(k1_hausdorf(X1,X2,X2),k6_numbers)|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_1987, plain, (k4_xxreal_0(k8_weierstr(X1,X2,X3),k6_numbers)=k1_hausdorf(X1,X3,X2)|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X2,X3)!=k6_numbers|~v2_compts_1(X2,k3_pcomps_1(X1))|~v2_compts_1(X3,k3_pcomps_1(X1)))).#trainneg
cnf(i_0_1988, plain, (k2_square_1(k8_weierstr(X1,X2,X3),k6_numbers)=k1_hausdorf(X1,X3,X2)|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X2,X3)!=k6_numbers|~v2_compts_1(X2,k3_pcomps_1(X1))|~v2_compts_1(X3,k3_pcomps_1(X1)))).#trainneg
cnf(i_0_3045, negated_conjecture, (esk5_2(k7_relat_1(k6_weierstr(esk1_0,esk3_0),esk3_0),X1)=k6_numbers|r1_tarski(k7_relat_1(k6_weierstr(esk1_0,esk3_0),esk3_0),X1)|~v1_funct_1(k6_weierstr(esk1_0,esk3_0))|~v1_relat_1(k6_weierstr(esk1_0,esk3_0))|~v4_pre_topc(esk3_0,k3_pcomps_1(esk1_0)))).#trainneg
cnf(i_0_2941, plain, (r1_xxreal_0(esk23_0,esk23_0))).#trainneg
cnf(i_0_448, plain, (v5_relat_1(esk24_1(k2_zfmisc_1(X1,X2)),X2)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2530, negated_conjecture, (r1_xxreal_0(esk3_0,esk3_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2894, plain, (v2_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2971, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_449, plain, (v5_relat_1(esk19_1(k2_zfmisc_1(X1,X2)),X2))).#trainneg
cnf(i_0_450, plain, (v5_relat_1(esk16_1(k2_zfmisc_1(X1,X2)),X2)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2906, plain, (r1_tarski(k7_relat_1(X3,X4),X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2887, plain, (X1=k6_numbers|~r1_tarski(X1,k6_numbers))).#trainneg
cnf(i_0_2088, plain, (m1_subset_1(esk14_1(k4_ordinal1),k1_numbers)|v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_452, plain, (v5_relat_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_2949, plain, (r1_xxreal_0(esk23_0,k4_ordinal1)|r1_xxreal_0(k4_ordinal1,esk23_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2533, negated_conjecture, (r1_xxreal_0(esk2_0,esk2_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2532, negated_conjecture, (r1_xxreal_0(esk3_0,esk3_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_454, plain, (v1_xreal_0(k10_binop_2(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3009, plain, (u1_struct_0(X1)=k6_numbers|X1!=esk20_0|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3007, plain, (X1=k6_numbers|g1_pre_topc(X1,X2)!=esk20_0|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_458, plain, (k2_seq_4(esk4_1(X1))=k2_weierstr(X1)|v1_xboole_0(X1)|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_2964, plain, (r1_xxreal_0(esk23_0,X1)|r1_xxreal_0(X1,esk23_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2535, negated_conjecture, (r1_xxreal_0(esk2_0,esk2_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_1120, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1231, plain, (r1_xxreal_0(k6_numbers,X1)|r1_xxreal_0(X1,k6_numbers)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3066, plain, (k1_funct_1(X2,X3)=k6_numbers|v1_xboole_0(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,X1,k1_zfmisc_1(k6_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k6_numbers))))|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_3088, plain, (u1_struct_0(X1)=k6_numbers|X1!=esk18_0|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_814, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_2868, plain, (r1_xxreal_0(esk14_1(k1_numbers),k6_numbers)|r1_xxreal_0(k6_numbers,esk14_1(k1_numbers)))).#trainneg
cnf(i_0_3086, plain, (X1=k6_numbers|g1_pre_topc(X1,X2)!=esk18_0|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_462, plain, (v1_xboole_0(esk19_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2394, plain, (r1_tarski(X1,k1_numbers)|v1_xboole_0(k4_ordinal1)|~m1_subset_1(esk5_2(X1,k1_numbers),k4_ordinal1))).#trainneg
cnf(i_0_1708, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|~r1_tarski(u1_struct_0(esk1_0),esk2_0))).#trainneg
cnf(i_0_1807, plain, (u1_struct_0(k3_topmetr)=k1_numbers|~m1_subset_1(k1_numbers,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_1521, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|r1_xxreal_0(esk3_0,X1)|~v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1956, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|r1_xxreal_0(esk2_0,X1)|~v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_465, plain, (v1_xboole_0(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3123, plain, (k4_xxreal_0(k6_numbers,k8_weierstr(X1,X2,X3))=k1_hausdorf(X1,X3,X2)|v1_xboole_0(X3)|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_xxreal_0(k8_weierstr(X1,X2,X3))|k1_hausdorf(X1,X3,X2)!=k6_numbers|~v2_compts_1(X2,k3_pcomps_1(X1))|~v2_compts_1(X3,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_2817, plain, (r2_hidden(esk5_2(k7_relat_1(X3,X4),X5),X2)|r1_tarski(k7_relat_1(X3,X4),X5)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2645, plain, (m1_subset_1(X1,k1_numbers)|v1_xboole_0(k7_relat_1(k6_weierstr(X2,X3),X4))|v2_struct_0(X2)|~r2_hidden(X1,k7_relat_1(k6_weierstr(X2,X3),X4))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_531, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,X3))|~v4_relat_1(X3,X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_817, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_1397, plain, (m1_subset_1(esk31_1(esk20_0),k1_zfmisc_1(k6_numbers))|~v2_pre_topc(esk20_0))).#trainneg
cnf(i_0_3005, plain, (X1=u1_pre_topc(esk20_0)|g1_pre_topc(X2,X1)!=esk20_0|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_1154, plain, (r1_xxreal_0(X1,k4_ordinal1)|r1_xxreal_0(k4_ordinal1,X1)|~v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_473, plain, (v1_xboole_0(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3266, plain, (k2_weierstr(X1)=k2_seq_4(X1)|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_3224, plain, (r1_tarski(k7_relat_1(k6_weierstr(X1,X2),X3),u1_struct_0(k3_topmetr))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3432, plain, (v3_valued_0(k2_zfmisc_1(X1,k1_numbers)))).#trainneg
cnf(i_0_3254, plain, (u1_struct_0(esk32_0)=k6_numbers|esk20_0!=esk32_0)).#trainneg
cnf(i_0_475, plain, (v3_valued_0(esk24_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_3328, plain, (r1_tarski(k4_ordinal1,k1_numbers)|v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_1362, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_476, plain, (v3_valued_0(esk19_1(k2_zfmisc_1(X1,X2)))|~v3_membered(X2))).#trainneg
cnf(i_0_477, plain, (v3_valued_0(esk16_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_3389, plain, (r2_hidden(esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4),u1_struct_0(k3_topmetr))|v2_struct_0(X1)|r1_tarski(k7_relat_1(k6_weierstr(X1,X2),X3),X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3223, negated_conjecture, (r1_tarski(k7_relat_1(k6_weierstr(esk1_0,X1),X2),u1_struct_0(k3_topmetr))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3174, negated_conjecture, (r1_tarski(k7_relat_1(k6_weierstr(esk1_0,esk3_0),esk3_0),X1)|~r2_hidden(k6_numbers,X1)|~v1_funct_1(k6_weierstr(esk1_0,esk3_0))|~v1_relat_1(k6_weierstr(esk1_0,esk3_0))|~v4_pre_topc(esk3_0,k3_pcomps_1(esk1_0)))).#trainneg
cnf(i_0_3417, plain, (m1_subset_1(esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4),k1_numbers)|v2_struct_0(X1)|r1_tarski(k7_relat_1(k6_weierstr(X1,X2),X3),X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3422, plain, (m1_subset_1(X1,X2)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~r2_hidden(X1,k9_xtuple_0(X3)))).#trainneg
cnf(i_0_3445, plain, (esk31_1(esk20_0)=k6_numbers|~v2_pre_topc(esk20_0))).#trainneg
cnf(i_0_3260, plain, (u1_struct_0(k3_topmetr)=k6_numbers|esk20_0!=k3_topmetr)).#trainneg
cnf(i_0_3084, plain, (X1=u1_pre_topc(esk18_0)|g1_pre_topc(X2,X1)!=esk18_0|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_479, plain, (v3_valued_0(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v3_membered(X2))).#trainneg
cnf(i_0_3440, plain, (v1_xboole_0(esk31_1(esk20_0))|~v2_pre_topc(esk20_0))).#trainneg
cnf(i_0_3441, plain, (v3_membered(esk31_1(esk20_0))|~v2_pre_topc(esk20_0))).#trainneg
cnf(i_0_483, plain, (r1_xxreal_0(X1,k1_funct_1(X2,X3))|r1_xxreal_0(k1_funct_1(X2,X3),X1)|~v1_xxreal_0(X1)|~v2_valued_0(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_484, plain, (v1_xxreal_0(k10_binop_2(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3405, negated_conjecture, (r2_hidden(esk5_2(k7_relat_1(k6_weierstr(esk1_0,X1),X2),X3),u1_struct_0(k3_topmetr))|r1_tarski(k7_relat_1(k6_weierstr(esk1_0,X1),X2),X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2191, negated_conjecture, (r2_hidden(esk6_4(k6_weierstr(esk1_0,esk3_0),esk2_0,X1,X2),esk3_0)|X1!=k7_relat_1(k6_weierstr(esk1_0,esk3_0),esk2_0)|X2!=k6_numbers|~r2_hidden(X2,X1)|~v4_pre_topc(esk3_0,k3_pcomps_1(esk1_0))|~v1_funct_1(k6_weierstr(esk1_0,esk3_0)))).#trainneg
cnf(i_0_2620, negated_conjecture, (k2_square_1(k8_weierstr(X1,X2,X2),k1_hausdorf(X1,X2,X2))=k1_hausdorf(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_2051, plain, (k7_relat_1(X3,X4)=X2|~m1_subset_1(X2,k1_zfmisc_1(k7_relat_1(X3,X4)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2621, negated_conjecture, (k4_xxreal_0(k8_weierstr(X1,X2,X2),k1_hausdorf(X1,X2,X2))=k1_hausdorf(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_2135, negated_conjecture, (k4_xxreal_0(X1,X1)=X1|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1974, plain, (k4_xxreal_0(k6_numbers,k8_weierstr(X1,X3,X2))=k1_hausdorf(X1,X2,X3)|v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X2,X3)!=k6_numbers|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1)))).#trainneg
cnf(i_0_2450, plain, (k7_relat_1(k2_zfmisc_1(X1,k6_numbers),X2)=k6_numbers)).#trainneg
cnf(i_0_3308, plain, (u1_struct_0(esk32_0)=k6_numbers|esk32_0!=esk18_0)).#trainneg
cnf(i_0_3442, plain, (v2_membered(esk31_1(esk20_0))|~v2_pre_topc(esk20_0))).#trainneg
cnf(i_0_3443, plain, (v3_valued_0(esk31_1(esk20_0))|~v2_pre_topc(esk20_0))).#trainneg
cnf(i_0_3301, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v3_membered(k1_zfmisc_1(X1))|~v3_membered(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1518, negated_conjecture, (r1_xxreal_0(k6_numbers,esk3_0)|r1_xxreal_0(esk3_0,k6_numbers)|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_486, plain, (v1_xcmplx_0(k10_binop_2(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_487, plain, (v7_ordinal1(k10_binop_2(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2014, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,u1_struct_0(esk1_0)),esk3_0)|~v1_funct_1(k6_weierstr(esk1_0,u1_struct_0(esk1_0)))|~r2_hidden(X1,X2)|~v4_pre_topc(u1_struct_0(esk1_0),k3_pcomps_1(esk1_0))|~m1_subset_1(esk6_4(k6_weierstr(esk1_0,u1_struct_0(esk1_0)),esk3_0,X2,X1),esk2_0))).#trainneg
cnf(i_0_2453, plain, (v3_membered(k7_relat_1(k2_zfmisc_1(X1,X2),X3))|~v3_membered(X2))).#trainneg
cnf(i_0_2139, plain, (X1=esk6_4(X2,X3,X4,X5)|~r1_xxreal_0(X1,esk6_4(X2,X3,X4,X5))|~v1_xxreal_0(esk6_4(X2,X3,X4,X5))|~v1_xxreal_0(X1)|X4!=k7_relat_1(X2,X3)|X1!=k2_seq_4(X3)|~v3_membered(X3)|~v1_xreal_0(esk6_4(X2,X3,X4,X5))|~v1_xreal_0(X1)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X5,X4)|~v4_xxreal_2(X3))).#trainneg
cnf(i_0_2492, negated_conjecture, (m1_subset_1(esk6_4(k6_weierstr(esk1_0,X1),esk2_0,X2,X3),k1_numbers)|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr)))|X2!=k7_relat_1(k6_weierstr(esk1_0,X1),esk2_0)|X3!=k6_numbers|~v1_funct_1(k6_weierstr(esk1_0,X1))|~r2_hidden(X3,X2)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3202, plain, (v3_valued_0(esk19_1(k2_zfmisc_1(X1,k1_numbers))))).#trainneg
cnf(i_0_3314, plain, (u1_struct_0(k3_topmetr)=k6_numbers|esk18_0!=k3_topmetr)).#trainneg
cnf(i_0_494, plain, (v1_xxreal_0(esk31_1(X1))|~v2_membered(k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3242, plain, (v1_xreal_0(esk14_1(k4_ordinal1))|v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_3436, plain, (v1_xxreal_0(esk31_1(esk20_0))|~v2_membered(k1_zfmisc_1(k6_numbers))|~v2_pre_topc(esk20_0))).#trainneg
cnf(i_0_3437, plain, (v1_xreal_0(esk31_1(esk20_0))|~v3_membered(k1_zfmisc_1(k6_numbers))|~v2_pre_topc(esk20_0))).#trainneg
cnf(i_0_1953, negated_conjecture, (r1_xxreal_0(k6_numbers,esk2_0)|r1_xxreal_0(esk2_0,k6_numbers)|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3337, negated_conjecture, (r1_xxreal_0(esk3_0,esk2_0)|r1_xxreal_0(esk2_0,esk3_0)|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_495, plain, (v1_xreal_0(esk31_1(X1))|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_496, plain, (v1_xcmplx_0(esk31_1(X1))|~v1_membered(k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2904, plain, (k7_relset_1(X2,X1,X3,X4)=X1|~r1_tarski(X1,k7_relset_1(X2,X1,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1))))).#trainneg
cnf(i_0_2111, plain, (m1_subset_1(esk6_4(X1,esk4_1(X2),X3,X4),k1_numbers)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_topmetr)))|X3!=k7_relat_1(X1,esk4_1(X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_3481, plain, (k4_seq_4(X1)=k2_seq_4(X1)|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_3594, plain, (m1_subset_1(esk6_4(X1,X2,X3,X4),X5)|~v4_relat_1(X1,X5)|~v1_relat_1(X1)|X3!=k7_relat_1(X1,X2)|~v1_funct_1(X1)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_3484, plain, (v1_xreal_0(k2_seq_4(X1))|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_3518, plain, (esk20_0!=esk32_0)).#trainneg
cnf(i_0_2599, plain, (k1_hausdorf(X1,X2,X3)=k6_numbers|v1_xboole_0(X3)|v1_xboole_0(X2)|v2_struct_0(X1)|k1_hausdorf(X1,X3,X2)!=k6_numbers|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_497, plain, (v7_ordinal1(esk31_1(X1))|~v6_membered(k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3598, plain, (v4_pre_topc(k6_numbers,esk20_0)|~v2_pre_topc(esk20_0))).#trainneg
cnf(i_0_3711, plain, (esk32_0!=esk18_0)).#trainneg
cnf(i_0_498, plain, (v1_xboole_0(esk31_1(X1))|~v1_xboole_0(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_499, plain, (v3_membered(esk31_1(X1))|~v3_membered(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3780, plain, (k7_relat_1(X3,X4)=X2|~r1_tarski(X2,k7_relat_1(X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3488, negated_conjecture, (k4_xxreal_0(k6_numbers,k6_numbers)=k6_numbers)).#trainneg
cnf(i_0_3553, plain, (m1_subset_1(k1_funct_1(k6_weierstr(X1,X2),X3),k1_numbers)|v2_struct_0(X1)|~m1_subset_1(k7_relat_1(k6_weierstr(X1,X2),X4),k1_zfmisc_1(u1_struct_0(k3_topmetr)))|~r2_hidden(X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3564, plain, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(X1,X2),X4))|v2_struct_0(X1)|~r2_hidden(X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~r2_hidden(X3,X2)|~v4_pre_topc(X2,k3_pcomps_1(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_3482, plain, (m1_subset_1(k2_seq_4(X1),k1_numbers)|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_1402, plain, (m1_subset_1(u1_pre_topc(esk20_0),k1_zfmisc_1(k1_zfmisc_1(k6_numbers))))).#trainneg
cnf(i_0_500, plain, (v2_membered(esk31_1(X1))|~v2_membered(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_501, plain, (v3_valued_0(esk31_1(X1))|~v3_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_503, plain, (m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,esk31_1(X2))|~v2_pre_topc(X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_3663, negated_conjecture, (k4_xxreal_0(k1_hausdorf(X1,X2,X2),k8_weierstr(X1,X2,X2))=k1_hausdorf(X1,X2,X2)|v2_struct_0(X1)|~v1_xxreal_0(k8_weierstr(X1,X2,X2))|~v1_xxreal_0(k1_hausdorf(X1,X2,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3639, negated_conjecture, (k2_square_1(k1_hausdorf(X1,X2,X2),k8_weierstr(X1,X2,X2))=k1_hausdorf(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(k1_hausdorf(X1,X2,X2),k1_numbers)|~m1_subset_1(k8_weierstr(X1,X2,X2),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3558, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,esk2_0),X1)=k6_numbers|~v1_funct_1(k6_weierstr(esk1_0,esk2_0))|~v1_relat_1(k6_weierstr(esk1_0,esk2_0))|~v4_pre_topc(esk2_0,k3_pcomps_1(esk1_0))|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3551, plain, (v2_struct_0(X1)|~v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X4))|~r2_hidden(X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3449, plain, (m1_subset_1(X1,k6_numbers)|~r2_hidden(X1,esk31_1(esk20_0))|~v2_pre_topc(esk20_0))).#trainneg
cnf(i_0_1470, plain, (m1_subset_1(u1_pre_topc(esk18_0),k1_zfmisc_1(k1_zfmisc_1(k6_numbers))))).#trainneg
cnf(i_0_507, plain, (v2_valued_0(esk24_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v2_membered(X2))).#trainneg
cnf(i_0_3717, negated_conjecture, (r1_xxreal_0(esk3_0,k6_numbers)|r1_xxreal_0(k6_numbers,esk3_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3524, plain, (k4_ordinal1=k1_numbers|~r1_tarski(k1_numbers,k4_ordinal1))).#trainneg
cnf(i_0_508, plain, (v2_valued_0(esk19_1(k2_zfmisc_1(X1,X2)))|~v2_membered(X2))).#trainneg
cnf(i_0_509, plain, (v2_valued_0(esk16_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v2_membered(X2))).#trainneg
cnf(i_0_3662, negated_conjecture, (k8_weierstr(X1,X2,X2)=k6_numbers|v2_struct_0(X1)|k1_hausdorf(X1,X2,X2)!=k6_numbers|~v1_xreal_0(k1_hausdorf(X1,X2,X2))|~v1_xreal_0(k8_weierstr(X1,X2,X2))|~r1_xxreal_0(k6_numbers,k8_weierstr(X1,X2,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3584, negated_conjecture, (r1_xxreal_0(k1_funct_1(k6_weierstr(esk1_0,X1),X2),X3)|v1_xboole_0(k7_relat_1(k6_weierstr(esk1_0,X1),X4))|v1_xboole_0(X4)|v1_xboole_0(X1)|X3!=k8_weierstr(esk1_0,X1,X4)|~v1_xreal_0(k1_funct_1(k6_weierstr(esk1_0,X1),X2))|~v1_xreal_0(X3)|~r2_hidden(X2,X4)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_compts_1(X4,k3_pcomps_1(esk1_0))|~v2_compts_1(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(k7_relat_1(k6_weierstr(esk1_0,X1),X4),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3596, plain, (m1_subset_1(esk5_2(k9_xtuple_0(X1),X2),X3)|r1_tarski(k9_xtuple_0(X1),X2)|~v4_relat_1(X1,X3)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3859, plain, (k6_numbers=X1|v2_struct_0(X2)|X3!=k7_relat_1(k6_weierstr(X2,X4),X5)|~v1_relat_1(k6_weierstr(X2,X4))|~r2_hidden(esk6_4(k6_weierstr(X2,X4),X5,X3,X1),X4)|~r2_hidden(X1,X3)|~v4_pre_topc(X4,k3_pcomps_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v4_relat_1(k6_weierstr(X2,X4),u1_struct_0(X2))|~v1_funct_1(k6_weierstr(X2,X4)))).#trainneg
cnf(i_0_511, plain, (v2_valued_0(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v2_membered(X2))).#trainneg
cnf(i_0_1520, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|r1_xxreal_0(esk3_0,X1)|~v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_513, plain, (k2_square_1(X2,X1)=k4_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4001, negated_conjecture, (k6_numbers=X3|X2!=k7_relat_1(k6_weierstr(esk1_0,esk2_0),X1)|~v1_funct_1(k6_weierstr(esk1_0,esk2_0))|~v1_relat_1(k6_weierstr(esk1_0,esk2_0))|~r2_hidden(X3,X2)|~r2_hidden(esk6_4(k6_weierstr(esk1_0,esk2_0),X1,X2,X3),esk2_0)|~v4_pre_topc(esk2_0,k3_pcomps_1(esk1_0)))).#trainneg
cnf(i_0_3937, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),X2))|~r2_hidden(esk6_4(X3,esk3_0,X4,X5),X2)|~r2_hidden(esk6_4(X3,esk3_0,X4,X5),X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|X4!=k7_relat_1(X3,esk3_0)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(X5,X4))).#trainneg
cnf(i_0_3907, negated_conjecture, (r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),X4)|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X2))|v1_xboole_0(X2)|v2_struct_0(X1)|X4!=k1_hausdorf(X1,X2,X2)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~v1_xreal_0(X4)|~r2_hidden(X3,X2)|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3823, plain, (r1_tarski(k7_relat_1(esk6_4(X1,X2,X3,X4),X5),X6)|X3!=k7_relat_1(X1,X2)|~v4_relat_1(X1,k1_zfmisc_1(k2_zfmisc_1(X7,X6)))|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_3938, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),X2))|~r2_hidden(esk6_4(X3,esk2_0,X4,X5),X2)|~r2_hidden(esk6_4(X3,esk2_0,X4,X5),X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|X4!=k7_relat_1(X3,esk2_0)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(X5,X4))).#trainneg
cnf(i_0_3969, plain, (m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~r2_hidden(X1,u1_pre_topc(esk20_0)))).#trainneg
cnf(i_0_3965, plain, (u1_pre_topc(esk20_0)=k1_zfmisc_1(k6_numbers)|~m1_subset_1(k1_zfmisc_1(k6_numbers),k1_zfmisc_1(u1_pre_topc(esk20_0))))).#trainneg
cnf(i_0_515, plain, (v1_xxreal_0(k1_relset_1(X1,X2))|~v2_membered(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_3961, plain, (v3_membered(u1_pre_topc(esk20_0))|~v3_membered(k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_3962, plain, (v2_membered(u1_pre_topc(esk20_0))|~v2_membered(k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_1955, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|r1_xxreal_0(esk2_0,X1)|~v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_516, plain, (v1_xreal_0(k1_relset_1(X1,X2))|~v3_membered(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_517, plain, (v1_xcmplx_0(k1_relset_1(X1,X2))|~v1_membered(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_3941, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),X2))|r1_tarski(esk2_0,X3)|~r2_hidden(esk5_2(esk2_0,X3),X2)|~r2_hidden(esk5_2(esk2_0,X3),X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_4122, plain, (r1_tarski(k9_xtuple_0(X1),X2)|v1_xboole_0(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4074, negated_conjecture, (k8_weierstr(X1,X2,X2)=k6_numbers|v2_struct_0(X1)|v1_xboole_0(X2)|k1_hausdorf(X1,X2,X2)!=k6_numbers|~v1_xreal_0(k1_hausdorf(X1,X2,X2))|~v1_xreal_0(k8_weierstr(X1,X2,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_compts_1(X2,k3_pcomps_1(X1)))).#trainneg
cnf(i_0_3940, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),X2))|r1_tarski(esk3_0,X3)|~r2_hidden(esk5_2(esk3_0,X3),X2)|~r2_hidden(esk5_2(esk3_0,X3),X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3858, plain, (r2_hidden(esk6_4(k6_weierstr(X1,X2),X3,X4,X5),X2)|v1_xboole_0(X2)|v2_struct_0(X1)|X4!=k7_relat_1(k6_weierstr(X1,X2),X3)|X5!=k6_numbers|~r2_hidden(X5,X4)|~v4_pre_topc(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v4_relat_1(k6_weierstr(X1,X2),u1_struct_0(X1))|~v1_funct_1(k6_weierstr(X1,X2))|~v1_relat_1(k6_weierstr(X1,X2)))).#trainneg
cnf(i_0_3844, plain, (r1_xxreal_0(esk6_4(X1,X2,X3,X4),X5)|v1_xboole_0(X6)|X5!=k2_seq_4(X6)|~v3_membered(X6)|~v1_xreal_0(X5)|~v4_xxreal_2(X6)|X3!=k7_relat_1(X1,X2)|~v4_relat_1(X1,X6)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_4078, negated_conjecture, (r1_xxreal_0(k1_funct_1(k6_weierstr(esk1_0,X1),X2),X3)|v1_xboole_0(k7_relat_1(k6_weierstr(esk1_0,X1),X1))|v1_xboole_0(X1)|X3!=k1_hausdorf(esk1_0,X1,X1)|~v1_xreal_0(k1_funct_1(k6_weierstr(esk1_0,X1),X2))|~v1_xreal_0(X3)|~r2_hidden(X2,X1)|~v2_compts_1(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3856, plain, (r1_xxreal_0(k6_numbers,X1)|v1_xboole_0(X2)|v2_struct_0(X3)|X4!=k7_relat_1(k6_weierstr(X3,X2),X5)|~r2_hidden(X1,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)|~v4_relat_1(k6_weierstr(X3,X2),u1_struct_0(k3_pcomps_1(X3)))|~v1_funct_1(k6_weierstr(X3,X2))|~v1_relat_1(k6_weierstr(X3,X2)))).#trainneg
cnf(i_0_4057, plain, (m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~r2_hidden(X1,u1_pre_topc(esk18_0)))).#trainneg
cnf(i_0_4053, plain, (u1_pre_topc(esk18_0)=k1_zfmisc_1(k6_numbers)|~m1_subset_1(k1_zfmisc_1(k6_numbers),k1_zfmisc_1(u1_pre_topc(esk18_0))))).#trainneg
cnf(i_0_518, plain, (v7_ordinal1(k1_relset_1(X1,X2))|~v6_membered(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_3767, negated_conjecture, (r1_xxreal_0(esk2_0,k6_numbers)|r1_xxreal_0(k6_numbers,esk2_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4049, plain, (v3_membered(u1_pre_topc(esk18_0))|~v3_membered(k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_4050, plain, (v2_membered(u1_pre_topc(esk18_0))|~v2_membered(k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_3290, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|r1_xxreal_0(esk3_0,X1)|~v3_membered(k1_zfmisc_1(X1))|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_519, plain, (v1_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_520, plain, (v1_xboole_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_xboole_0(X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_3831, plain, (m1_subset_1(esk6_4(X1,X2,X3,X4),k1_zfmisc_1(k1_numbers))|v1_xboole_0(esk6_4(X1,X2,X3,X4))|X3!=k7_relat_1(X1,X2)|~v4_relat_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr)))|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_3857, plain, (X1=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X3)|X4!=k7_relat_1(k6_weierstr(X3,X2),X5)|~v2_valued_0(k6_weierstr(X3,X2))|~r1_xxreal_0(X1,k6_numbers)|~r2_hidden(X1,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)|~v4_relat_1(k6_weierstr(X3,X2),u1_struct_0(k3_pcomps_1(X3)))|~v1_funct_1(k6_weierstr(X3,X2))|~v1_relat_1(k6_weierstr(X3,X2)))).#trainneg
cnf(i_0_3860, negated_conjecture, (k6_numbers=X1|v1_xboole_0(X2)|X3!=k7_relat_1(k6_weierstr(esk1_0,X2),esk3_0)|~v1_funct_1(k6_weierstr(esk1_0,X2))|~r2_hidden(X1,X3)|~v4_pre_topc(X2,k3_pcomps_1(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v4_relat_1(k6_weierstr(esk1_0,X2),X2)|~v1_relat_1(k6_weierstr(esk1_0,X2)))).#trainneg
cnf(i_0_893, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk19_1(X2))|~m1_subset_1(X1,esk19_1(X2)))).#trainneg
cnf(i_0_3473, plain, (esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=k6_numbers|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3241, plain, (v1_xxreal_0(esk14_1(k4_ordinal1))|v1_xboole_0(k4_ordinal1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_521, plain, (v1_xboole_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_xboole_0(X1)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_3770, negated_conjecture, (r1_xxreal_0(esk2_0,esk3_0)|r1_xxreal_0(esk3_0,esk2_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3956, plain, (v1_xxreal_0(u1_pre_topc(esk20_0))|~v2_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers))))).#trainneg
cnf(i_0_3957, plain, (v1_xreal_0(u1_pre_topc(esk20_0))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers))))).#trainneg
cnf(i_0_3291, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|r1_xxreal_0(esk2_0,X1)|~v3_membered(k1_zfmisc_1(X1))|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_522, plain, (v4_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_523, plain, (v3_valued_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v3_membered(X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_3996, negated_conjecture, (k2_square_1(k1_hausdorf(X1,X2,X2),k6_numbers)=k1_hausdorf(X1,X2,X2)|v2_struct_0(X1)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(k1_hausdorf(X1,X2,X2),k1_numbers)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X2,X2)!=k6_numbers|~v2_compts_1(X2,k3_pcomps_1(X1)))).#trainneg
cnf(i_0_4075, negated_conjecture, (r1_xxreal_0(k1_funct_1(k6_weierstr(esk1_0,X1),X2),k8_weierstr(esk1_0,X1,X3))|v1_xboole_0(k7_relat_1(k6_weierstr(esk1_0,X1),X3))|v1_xboole_0(X1)|~v1_xreal_0(k1_funct_1(k6_weierstr(esk1_0,X1),X2))|~v1_xreal_0(k8_weierstr(esk1_0,X1,X3))|~r2_hidden(X2,X3)|~v2_compts_1(X3,k3_pcomps_1(esk1_0))|~v2_compts_1(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_4133, plain, (k6_numbers=X1|v2_struct_0(X2)|X3!=k7_relat_1(k6_weierstr(X2,X4),X4)|~v4_relat_1(k6_weierstr(X2,X4),u1_struct_0(X2))|~r2_hidden(X1,X3)|~v4_pre_topc(X4,k3_pcomps_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v1_funct_1(k6_weierstr(X2,X4))|~v1_relat_1(k6_weierstr(X2,X4)))).#trainneg
cnf(i_0_4281, plain, (r1_xxreal_0(k6_numbers,X1)|v1_xboole_0(X2)|v2_struct_0(X3)|~r2_hidden(X1,k7_relat_1(k6_weierstr(X3,X2),X4))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_4158, plain, (k4_xxreal_0(X2,X1)=k4_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3321, plain, (esk19_1(k2_zfmisc_1(X1,X2))=k6_numbers|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3243, plain, (v1_xcmplx_0(esk14_1(k4_ordinal1))|v1_xboole_0(k4_ordinal1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_524, plain, (v2_valued_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v2_membered(X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_3606, plain, (X1=esk31_1(esk20_0)|~v1_xboole_0(X1)|~v2_pre_topc(esk20_0))).#trainneg
cnf(i_0_3958, plain, (v1_xcmplx_0(u1_pre_topc(esk20_0))|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers))))).#trainneg
cnf(i_0_3959, plain, (v7_ordinal1(u1_pre_topc(esk20_0))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers))))).#trainneg
cnf(i_0_525, plain, (v5_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_526, plain, (v1_xboole_0(k1_relset_1(X1,X2))|~v1_xboole_0(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_4228, plain, (k8_weierstr(X1,X2,X3)=k1_funct_1(k6_weierstr(X1,X2),X4)|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X3))|v1_xboole_0(X2)|v2_struct_0(X1)|~r1_xxreal_0(k8_weierstr(X1,X2,X3),k1_funct_1(k6_weierstr(X1,X2),X4))|~v1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X4))|~v1_xxreal_0(k8_weierstr(X1,X2,X3))|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X4))|~r2_hidden(X4,X3)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_4170, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),X2))|v1_xboole_0(X2)|X3!=k7_relat_1(X4,esk3_0)|~v1_funct_1(X4)|~v1_relat_1(X4)|~r2_hidden(esk6_4(X4,esk3_0,X3,X5),X1)|~r2_hidden(X5,X3)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk6_4(X4,esk3_0,X3,X5),X2))).#trainneg
cnf(i_0_4401, negated_conjecture, (k8_weierstr(esk1_0,X1,X2)=k1_funct_1(k6_weierstr(esk1_0,X1),X3)|v1_xboole_0(k7_relat_1(k6_weierstr(esk1_0,X1),X2))|v1_xboole_0(X1)|~r1_xxreal_0(k8_weierstr(esk1_0,X1,X2),k1_funct_1(k6_weierstr(esk1_0,X1),X3))|~v1_xxreal_0(k1_funct_1(k6_weierstr(esk1_0,X1),X3))|~v1_xxreal_0(k8_weierstr(esk1_0,X1,X2))|~v1_xreal_0(k1_funct_1(k6_weierstr(esk1_0,X1),X3))|~r2_hidden(X3,X2)|~v2_compts_1(X2,k3_pcomps_1(esk1_0))|~v2_compts_1(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_4244, plain, (m1_subset_1(k9_xtuple_0(X1),k1_zfmisc_1(X2))|v1_xboole_0(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4230, plain, (r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),k6_numbers)|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X4))|v1_xboole_0(X2)|v2_struct_0(X1)|v1_xboole_0(X4)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~r2_hidden(X3,X4)|~v2_compts_1(X4,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X4,X2)!=k6_numbers)).#trainneg
cnf(i_0_4190, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),X2))|v1_xboole_0(X2)|X3!=k7_relat_1(X4,esk2_0)|~v1_funct_1(X4)|~v1_relat_1(X4)|~r2_hidden(esk6_4(X4,esk2_0,X3,X5),X1)|~r2_hidden(X5,X3)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk6_4(X4,esk2_0,X3,X5),X2))).#trainneg
cnf(i_0_4131, plain, (k6_numbers=X1|v2_struct_0(X2)|X3!=k7_relat_1(k6_weierstr(X2,u1_struct_0(k3_pcomps_1(X2))),X4)|~v4_relat_1(k6_weierstr(X2,u1_struct_0(k3_pcomps_1(X2))),u1_struct_0(X2))|~r2_hidden(X1,X3)|~v4_pre_topc(u1_struct_0(k3_pcomps_1(X2)),k3_pcomps_1(X2))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_1232, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1486, plain, (v1_xreal_0(esk14_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_3244, plain, (v7_ordinal1(esk14_1(k4_ordinal1))|v1_xboole_0(k4_ordinal1)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_4163, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,esk2_0),k7_relset_1(X3,esk2_0,X4,X5))|~v1_funct_1(k6_weierstr(esk1_0,esk2_0))|~v1_relat_1(k6_weierstr(esk1_0,esk2_0))|~r2_hidden(X1,X2)|~v4_pre_topc(esk2_0,k3_pcomps_1(esk1_0))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,esk2_0))))).#trainneg
cnf(i_0_527, plain, (v3_membered(k1_relset_1(X1,X2))|~v3_membered(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_3447, plain, (r2_hidden(X1,k6_numbers)|~r2_hidden(X1,esk31_1(esk20_0))|~v2_pre_topc(esk20_0))).#trainneg
cnf(i_0_4044, plain, (v1_xxreal_0(u1_pre_topc(esk18_0))|~v2_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers))))).#trainneg
cnf(i_0_4045, plain, (v1_xreal_0(u1_pre_topc(esk18_0))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers))))).#trainneg
cnf(i_0_528, plain, (v2_membered(k1_relset_1(X1,X2))|~v2_membered(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_529, plain, (v3_valued_0(k1_relset_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_4229, plain, (r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),k6_numbers)|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X4))|v1_xboole_0(X2)|v2_struct_0(X1)|v1_xboole_0(X4)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~r2_hidden(X3,X4)|~v2_compts_1(X4,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X2,X4)!=k6_numbers)).#trainneg
cnf(i_0_4174, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),esk3_0))|X2!=k7_relat_1(X3,esk3_0)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(esk6_4(X3,esk3_0,X2,X4),X1)|~r2_hidden(X4,X2)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_4231, negated_conjecture, (r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),k1_hausdorf(X1,X2,X2))|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X2))|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~r2_hidden(X3,X2)|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_4238, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),u1_struct_0(esk1_0)))|r1_tarski(esk2_0,X2)|~r2_hidden(esk5_2(esk2_0,X2),X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(esk5_2(esk2_0,X2),esk2_0))).#trainneg
cnf(i_0_532, plain, (m1_subset_1(k9_xtuple_0(X2),k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_3719, negated_conjecture, (r1_xxreal_0(esk3_0,k6_numbers)|r1_xxreal_0(k6_numbers,esk3_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_4046, plain, (v1_xcmplx_0(u1_pre_topc(esk18_0))|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers))))).#trainneg
cnf(i_0_4047, plain, (v7_ordinal1(u1_pre_topc(esk18_0))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers))))).#trainneg
cnf(i_0_533, plain, (v1_xxreal_0(k2_square_1(X1,X2))|~v2_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_534, plain, (v1_xreal_0(k2_square_1(X1,X2))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4476, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,X3),X3)|~r2_hidden(X1,X2)|~v4_pre_topc(X3,k3_pcomps_1(esk1_0))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_4576, plain, (k8_weierstr(X1,X2,X3)=X6|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X3))|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_xreal_0(X6)|~r1_xxreal_0(k8_weierstr(X1,X2,X3),X6)|~v1_xxreal_0(k8_weierstr(X1,X2,X3))|~r2_hidden(esk6_4(k6_weierstr(X1,X2),X4,X5,X6),X3)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|X5!=k7_relat_1(k6_weierstr(X1,X2),X4)|~v1_funct_1(k6_weierstr(X1,X2))|~v1_relat_1(k6_weierstr(X1,X2))|~r2_hidden(X6,X5))).#trainneg
cnf(i_0_4750, negated_conjecture, (r1_xxreal_0(X5,k1_hausdorf(X1,X2,X2))|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X2))|v2_struct_0(X1)|~v1_xreal_0(X5)|~r2_hidden(esk6_4(k6_weierstr(X1,X2),X3,X4,X5),X2)|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|X4!=k7_relat_1(k6_weierstr(X1,X2),X3)|~v1_funct_1(k6_weierstr(X1,X2))|~v1_relat_1(k6_weierstr(X1,X2))|~r2_hidden(X5,X4))).#trainneg
cnf(i_0_3246, plain, (v3_valued_0(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))))).#trainneg
cnf(i_0_3438, plain, (v1_xcmplx_0(esk31_1(esk20_0))|~v1_membered(k1_zfmisc_1(k6_numbers))|~v2_pre_topc(esk20_0))).#trainneg
cnf(i_0_535, plain, (v1_xcmplx_0(k2_square_1(X1,X2))|~v1_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3769, negated_conjecture, (r1_xxreal_0(esk2_0,k6_numbers)|r1_xxreal_0(k6_numbers,esk2_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_4553, plain, (v1_xreal_0(u1_pre_topc(esk20_0))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers))))).#trainneg
cnf(i_0_3439, plain, (v7_ordinal1(esk31_1(esk20_0))|~v6_membered(k1_zfmisc_1(k6_numbers))|~v2_pre_topc(esk20_0))).#trainneg
cnf(i_0_2175, plain, (r1_xxreal_0(k6_numbers,k1_hausdorf(X1,X2,X3))|r1_xxreal_0(k1_hausdorf(X1,X2,X3),k6_numbers)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_536, plain, (v7_ordinal1(k2_square_1(X1,X2))|~v6_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4578, plain, (k8_weierstr(X1,X2,X3)=k6_numbers|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X3))|v1_xboole_0(X2)|v2_struct_0(X1)|~r1_xxreal_0(k8_weierstr(X1,X2,X3),k6_numbers)|~v1_xxreal_0(k8_weierstr(X1,X2,X3))|~r2_hidden(X4,X3)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~r2_hidden(X4,X2)|~v4_pre_topc(X2,k3_pcomps_1(X1))|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_4581, negated_conjecture, (k1_hausdorf(X1,X2,X2)=k1_funct_1(k6_weierstr(X1,X2),X3)|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X2))|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~r1_xxreal_0(k1_hausdorf(X1,X2,X2),k1_funct_1(k6_weierstr(X1,X2),X3))|~v1_xxreal_0(k1_hausdorf(X1,X2,X2))|~r2_hidden(X3,X2)|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_4584, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_numbers))|X2!=k7_relat_1(X3,esk3_0)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(esk6_4(X3,esk3_0,X2,X4),X1)|~r2_hidden(X4,X2)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_4565, plain, (m1_subset_1(k1_funct_1(k6_weierstr(X1,X2),X3),u1_struct_0(k3_topmetr))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_4583, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_numbers))|X2!=k7_relat_1(X3,esk3_0)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(esk6_4(X3,esk3_0,X2,X4),X1)|~r2_hidden(X4,X2)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k9_xtuple_0(X3),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_2105, plain, (v2_membered(X1)|v1_xboole_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_3366, plain, (esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=k6_numbers|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3267, plain, (r1_xxreal_0(esk14_1(k1_numbers),esk23_0)|r1_xxreal_0(esk23_0,esk14_1(k1_numbers)))).#trainneg
cnf(i_0_539, plain, (m1_subset_1(k4_xxreal_0(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3772, negated_conjecture, (r1_xxreal_0(esk2_0,esk3_0)|r1_xxreal_0(esk3_0,esk2_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_3963, plain, (v3_valued_0(u1_pre_topc(esk20_0))|~v3_valued_0(k1_zfmisc_1(k6_numbers))|~v1_relat_1(k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_4051, plain, (v3_valued_0(u1_pre_topc(esk18_0))|~v3_valued_0(k1_zfmisc_1(k6_numbers))|~v1_relat_1(k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2423, plain, (r1_xxreal_0(X1,k1_hausdorf(X2,X3,X4))|r1_xxreal_0(k1_hausdorf(X2,X3,X4),X1)|v2_struct_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_541, plain, (v1_xboole_0(X1)|v1_xboole_0(esk24_1(k9_xtuple_0(X1)))|v1_xboole_0(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_xboole_0(k7_relat_1(X1,esk24_1(k9_xtuple_0(X1)))))).#trainneg
cnf(i_0_4657, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_numbers))|X2!=k7_relat_1(X3,esk2_0)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(esk6_4(X3,esk2_0,X2,X4),X1)|~r2_hidden(X4,X2)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_4590, negated_conjecture, (k8_weierstr(esk1_0,X1,X2)=k6_numbers|v1_xboole_0(k7_relat_1(k6_weierstr(esk1_0,X1),X2))|v1_xboole_0(X1)|~r1_xxreal_0(k8_weierstr(esk1_0,X1,X2),k6_numbers)|~v1_xxreal_0(k8_weierstr(esk1_0,X1,X2))|~r2_hidden(X3,X2)|~v2_compts_1(X2,k3_pcomps_1(esk1_0))|~v2_compts_1(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(X3,X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4871, negated_conjecture, (k6_numbers=X1|~r2_hidden(X1,k7_relat_1(k6_weierstr(esk1_0,X2),X2))|~v4_pre_topc(X2,k3_pcomps_1(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_4987, plain, (v2_membered(X1)|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_4814, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k9_xtuple_0(X3))|~v4_relat_1(X3,X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_4169, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),u1_struct_0(esk1_0)))|X2!=k7_relat_1(X3,esk3_0)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(esk6_4(X3,esk3_0,X2,X4),X1)|~r2_hidden(X4,X2)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(esk6_4(X3,esk3_0,X2,X4),esk2_0))).#trainneg
cnf(i_0_4594, negated_conjecture, (k1_hausdorf(esk1_0,X1,X1)=k1_funct_1(k6_weierstr(esk1_0,X1),X2)|v1_xboole_0(k7_relat_1(k6_weierstr(esk1_0,X1),X1))|v1_xboole_0(X1)|~v1_xreal_0(k1_funct_1(k6_weierstr(esk1_0,X1),X2))|~r1_xxreal_0(k1_hausdorf(esk1_0,X1,X1),k1_funct_1(k6_weierstr(esk1_0,X1),X2))|~v1_xxreal_0(k1_hausdorf(esk1_0,X1,X1))|~r2_hidden(X2,X1)|~v2_compts_1(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1403, plain, (m1_subset_1(k2_pcomps_1(esk20_0),k1_zfmisc_1(k1_zfmisc_1(k6_numbers)))|~l1_metric_1(esk20_0))).#trainneg
cnf(i_0_542, plain, (v1_xboole_0(X1)|v1_xboole_0(esk19_1(k9_xtuple_0(X1)))|~v1_relat_1(X1)|~v1_xboole_0(k7_relat_1(X1,esk19_1(k9_xtuple_0(X1)))))).#trainneg
cnf(i_0_4821, plain, (v1_xreal_0(u1_pre_topc(esk18_0))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers))))).#trainneg
cnf(i_0_543, plain, (v1_xboole_0(X1)|v1_xboole_0(esk16_1(k9_xtuple_0(X1)))|v1_xboole_0(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_xboole_0(k7_relat_1(X1,esk16_1(k9_xtuple_0(X1)))))).#trainneg
cnf(i_0_544, plain, (v1_xboole_0(X1)|v1_xboole_0(esk14_1(k1_zfmisc_1(k9_xtuple_0(X1))))|~v1_relat_1(X1)|~v1_xboole_0(k7_relat_1(X1,esk14_1(k1_zfmisc_1(k9_xtuple_0(X1))))))).#trainneg
cnf(i_0_4256, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),u1_struct_0(esk1_0)))|r1_tarski(esk3_0,X2)|~r2_hidden(esk5_2(esk3_0,X2),X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(esk5_2(esk3_0,X2),esk2_0))).#trainneg
cnf(i_0_4255, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),u1_struct_0(esk1_0)))|r1_tarski(esk3_0,X2)|~r2_hidden(esk5_2(esk3_0,X2),X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(esk5_2(esk3_0,X2),esk3_0))).#trainneg
cnf(i_0_4882, plain, (k8_weierstr(X1,X2,X3)=X4|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X3))|v1_xboole_0(X2)|v2_struct_0(X1)|X5!=k7_relat_1(k6_weierstr(X1,X2),X3)|~v1_xreal_0(X4)|~r1_xxreal_0(k8_weierstr(X1,X2,X3),X4)|~v1_xxreal_0(k8_weierstr(X1,X2,X3))|~r2_hidden(X4,X5)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v1_funct_1(k6_weierstr(X1,X2))|~v1_relat_1(k6_weierstr(X1,X2)))).#trainneg
cnf(i_0_4129, plain, (k6_numbers=X1|v2_struct_0(X2)|v1_xboole_0(X4)|X3!=k7_relat_1(k6_weierstr(X2,X4),X5)|~v4_relat_1(k6_weierstr(X2,X4),u1_struct_0(X2))|~r2_hidden(X1,X3)|~v4_pre_topc(X4,k3_pcomps_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~m1_subset_1(esk6_4(k6_weierstr(X2,X4),X5,X3,X1),X4))).#trainneg
cnf(i_0_4402, negated_conjecture, (r1_xxreal_0(k1_funct_1(k6_weierstr(esk1_0,X1),X2),k6_numbers)|v1_xboole_0(k7_relat_1(k6_weierstr(esk1_0,X1),X3))|v1_xboole_0(X1)|v1_xboole_0(X3)|~v1_xreal_0(k1_funct_1(k6_weierstr(esk1_0,X1),X2))|~r2_hidden(X2,X3)|~v2_compts_1(X3,k3_pcomps_1(esk1_0))|~v2_compts_1(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|k1_hausdorf(esk1_0,X1,X3)!=k6_numbers)).#trainneg
cnf(i_0_4149, negated_conjecture, (k4_xxreal_0(k1_hausdorf(X1,X2,X2),k8_weierstr(X1,X2,X2))=k1_hausdorf(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(k8_weierstr(X1,X2,X2),k1_numbers)|~m1_subset_1(k1_hausdorf(X1,X2,X2),k1_numbers))).#trainneg
cnf(i_0_4195, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),esk2_0))|X2!=k7_relat_1(X3,esk2_0)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(esk6_4(X3,esk2_0,X2,X4),X1)|~r2_hidden(X4,X2)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_4589, negated_conjecture, (k8_weierstr(esk1_0,X1,X2)=X5|v1_xboole_0(k7_relat_1(k6_weierstr(esk1_0,X1),X2))|v1_xboole_0(X1)|~v1_xreal_0(X5)|~r1_xxreal_0(k8_weierstr(esk1_0,X1,X2),X5)|~v1_xxreal_0(k8_weierstr(esk1_0,X1,X2))|~r2_hidden(esk6_4(k6_weierstr(esk1_0,X1),X3,X4,X5),X2)|~v2_compts_1(X2,k3_pcomps_1(esk1_0))|~v2_compts_1(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|X4!=k7_relat_1(k6_weierstr(esk1_0,X1),X3)|~v1_funct_1(k6_weierstr(esk1_0,X1))|~v1_relat_1(k6_weierstr(esk1_0,X1))|~r2_hidden(X5,X4))).#trainneg
cnf(i_0_4168, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),u1_struct_0(esk1_0)))|X2!=k7_relat_1(X3,esk3_0)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(esk6_4(X3,esk3_0,X2,X4),X1)|~r2_hidden(X4,X2)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(esk6_4(X3,esk3_0,X2,X4),esk3_0))).#trainneg
cnf(i_0_4958, negated_conjecture, (k1_hausdorf(X1,X2,X2)=k6_numbers|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X2))|v2_struct_0(X1)|~r1_xxreal_0(k1_hausdorf(X1,X2,X2),k6_numbers)|~v1_xxreal_0(k1_hausdorf(X1,X2,X2))|~r2_hidden(X3,X2)|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v4_pre_topc(X2,k3_pcomps_1(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_590, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relset_1(X3,X2,X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_1471, plain, (m1_subset_1(k2_pcomps_1(esk18_0),k1_zfmisc_1(k1_zfmisc_1(k6_numbers)))|~l1_metric_1(esk18_0))).#trainneg
cnf(i_0_4924, plain, (v1_xreal_0(esk31_1(esk20_0))|~v6_membered(k1_zfmisc_1(k6_numbers))|~v2_pre_topc(esk20_0))).#trainneg
cnf(i_0_546, plain, (v1_xboole_0(X1)|v1_xboole_0(k1_relset_1(k9_xtuple_0(X1),X2))|~v1_relat_1(X1)|~v1_xboole_0(k7_relat_1(X1,k1_relset_1(k9_xtuple_0(X1),X2)))|~v4_relat_1(X2,k9_xtuple_0(X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_4804, plain, (k9_xtuple_0(X1)=k6_numbers|~v4_relat_1(X1,k6_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_757, negated_conjecture, (v2_struct_0(k3_pcomps_1(esk1_0))|~l1_pre_topc(k3_pcomps_1(esk1_0))|~v2_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(k3_pcomps_1(esk1_0)))))).#trainneg
cnf(i_0_853, plain, (v2_struct_0(k3_pcomps_1(esk22_0))|~l1_pre_topc(k3_pcomps_1(esk22_0))|~v2_struct_0(g1_pre_topc(u1_struct_0(esk22_0),u1_pre_topc(k3_pcomps_1(esk22_0)))))).#trainneg
cnf(i_0_5055, plain, (r1_xxreal_0(k1_hausdorf(X1,X2,k6_numbers),X3)|r1_xxreal_0(X3,k1_hausdorf(X1,X2,k6_numbers))|v2_struct_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_5056, plain, (r1_xxreal_0(k1_hausdorf(X1,X2,u1_struct_0(k3_pcomps_1(X1))),X3)|r1_xxreal_0(X3,k1_hausdorf(X1,X2,u1_struct_0(k3_pcomps_1(X1))))|v2_struct_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_547, plain, (v1_xxreal_0(esk28_1(X1))|v2_struct_0(X1)|~v2_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_548, plain, (v1_xreal_0(esk28_1(X1))|v2_struct_0(X1)|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_4656, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_numbers))|X2!=k7_relat_1(X3,esk2_0)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(esk6_4(X3,esk2_0,X2,X4),X1)|~r2_hidden(X4,X2)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k9_xtuple_0(X3),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_5083, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,X1),X2)=k6_numbers|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(X2,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5112, plain, (r2_hidden(esk6_4(X1,X2,X3,X4),X5)|~v4_relat_1(X1,X5)|~v1_relat_1(X1)|X3!=k7_relat_1(X1,X2)|~v1_funct_1(X1)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_5281, plain, (k9_xtuple_0(k6_numbers)=k6_numbers)).#trainneg
cnf(i_0_5081, negated_conjecture, (esk5_2(k7_relat_1(k6_weierstr(esk1_0,X1),X1),X2)=k6_numbers|r1_tarski(k7_relat_1(k6_weierstr(esk1_0,X1),X1),X2)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_5100, plain, (r1_xxreal_0(esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4),k8_weierstr(X1,X2,X3))|v1_xboole_0(X2)|v2_struct_0(X1)|r1_tarski(k7_relat_1(k6_weierstr(X1,X2),X3),X4)|~v1_xreal_0(esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4))|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_5094, plain, (r1_xxreal_0(esk9_3(k7_relat_1(k6_weierstr(X1,X2),X3),X4,X5),k8_weierstr(X1,X2,X3))|v1_xboole_0(X2)|v2_struct_0(X1)|r1_xxreal_0(X5,k6_numbers)|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X3))|~v1_xreal_0(esk9_3(k7_relat_1(k6_weierstr(X1,X2),X3),X4,X5))|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|X4!=k2_seq_4(k7_relat_1(k6_weierstr(X1,X2),X3))|~v3_membered(k7_relat_1(k6_weierstr(X1,X2),X3))|~v1_xreal_0(X5)|~v1_xreal_0(X4)|~v4_xxreal_2(k7_relat_1(k6_weierstr(X1,X2),X3)))).#trainneg
cnf(i_0_5173, negated_conjecture, (k6_numbers=X1|v2_struct_0(X2)|X3!=k7_relat_1(k6_weierstr(X2,u1_struct_0(esk1_0)),esk3_0)|~v4_relat_1(k6_weierstr(X2,u1_struct_0(esk1_0)),u1_struct_0(X2))|~r2_hidden(X1,X3)|~v4_pre_topc(u1_struct_0(esk1_0),k3_pcomps_1(X2))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v1_funct_1(k6_weierstr(X2,u1_struct_0(esk1_0)))|~v1_relat_1(k6_weierstr(X2,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_5138, plain, (m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~r2_hidden(X1,k2_pcomps_1(esk20_0))|~l1_metric_1(esk20_0))).#trainneg
cnf(i_0_5280, plain, (k9_xtuple_0(k2_zfmisc_1(k6_numbers,X1))=k6_numbers)).#trainneg
cnf(i_0_5086, plain, (v1_xxreal_0(k4_ordinal1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_4672, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,u1_struct_0(esk1_0)),X3)|~r2_hidden(X1,X2)|~v4_pre_topc(u1_struct_0(esk1_0),k3_pcomps_1(esk1_0)))).#trainneg
cnf(i_0_549, plain, (v1_xcmplx_0(esk28_1(X1))|v2_struct_0(X1)|~v1_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_5090, negated_conjecture, (v2_membered(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_5091, negated_conjecture, (v2_membered(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_5301, plain, (r1_xxreal_0(X1,k1_hausdorf(X2,k6_numbers,k6_numbers))|r1_xxreal_0(k1_hausdorf(X2,k6_numbers,k6_numbers),X1)|v2_struct_0(X2)|~v1_xreal_0(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_5302, plain, (r1_xxreal_0(X1,k1_hausdorf(X2,u1_struct_0(k3_pcomps_1(X2)),k6_numbers))|r1_xxreal_0(k1_hausdorf(X2,u1_struct_0(k3_pcomps_1(X2)),k6_numbers),X1)|v2_struct_0(X2)|~v1_xreal_0(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_550, plain, (v7_ordinal1(esk28_1(X1))|v2_struct_0(X1)|~v6_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_552, plain, (v3_membered(esk28_1(X1))|v2_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_5075, negated_conjecture, (esk9_3(k7_relat_1(k6_weierstr(esk1_0,X1),X1),X2,X3)=k6_numbers|r1_xxreal_0(X3,k6_numbers)|v1_xboole_0(k7_relat_1(k6_weierstr(esk1_0,X1),X1))|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|X2!=k2_seq_4(k7_relat_1(k6_weierstr(esk1_0,X1),X1))|~v3_membered(k7_relat_1(k6_weierstr(esk1_0,X1),X1))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v4_xxreal_2(k7_relat_1(k6_weierstr(esk1_0,X1),X1)))).#trainneg
cnf(i_0_5164, negated_conjecture, (k1_hausdorf(X1,X2,X2)=X3|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X2))|v1_xboole_0(X2)|v2_struct_0(X1)|X4!=k7_relat_1(k6_weierstr(X1,X2),X2)|~v1_xreal_0(X3)|~r1_xxreal_0(k1_hausdorf(X1,X2,X2),X3)|~v1_xxreal_0(k1_hausdorf(X1,X2,X2))|~r2_hidden(X3,X4)|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_5175, plain, (k6_numbers=X1|v1_xboole_0(X2)|v2_struct_0(X3)|X4!=k7_relat_1(k6_weierstr(X3,X2),X5)|~v4_relat_1(k6_weierstr(X3,X2),u1_struct_0(X3))|~r2_hidden(X1,X4)|~v4_pre_topc(X2,k3_pcomps_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)|~v4_relat_1(k6_weierstr(X3,X2),X2)|~v1_funct_1(k6_weierstr(X3,X2))|~v1_relat_1(k6_weierstr(X3,X2)))).#trainneg
cnf(i_0_5174, negated_conjecture, (k6_numbers=X1|v2_struct_0(X2)|X3!=k7_relat_1(k6_weierstr(X2,u1_struct_0(esk1_0)),esk2_0)|~v4_relat_1(k6_weierstr(X2,u1_struct_0(esk1_0)),u1_struct_0(X2))|~r2_hidden(X1,X3)|~v4_pre_topc(u1_struct_0(esk1_0),k3_pcomps_1(X2))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v1_funct_1(k6_weierstr(X2,u1_struct_0(esk1_0)))|~v1_relat_1(k6_weierstr(X2,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_5343, negated_conjecture, (k6_numbers=X4|X3!=k7_relat_1(k6_weierstr(esk1_0,X1),X2)|~v1_funct_1(k6_weierstr(esk1_0,X1))|~v1_relat_1(k6_weierstr(esk1_0,X1))|~r2_hidden(X4,X3)|~r2_hidden(esk6_4(k6_weierstr(esk1_0,X1),X2,X3,X4),X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_5444, plain, (k8_weierstr(X1,X2,X3)=k1_funct_1(k6_weierstr(X1,X2),X4)|v1_xboole_0(X2)|v2_struct_0(X1)|~r1_xxreal_0(k8_weierstr(X1,X2,X3),k1_funct_1(k6_weierstr(X1,X2),X4))|~v1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X4))|~v1_xxreal_0(k8_weierstr(X1,X2,X3))|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X4))|~r2_hidden(X4,X3)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_5449, plain, (r1_xxreal_0(X5,k8_weierstr(X1,X2,X6))|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_xreal_0(X5)|~r2_hidden(esk6_4(k6_weierstr(X1,X2),X3,X4,X5),X6)|~v2_compts_1(X6,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|X4!=k7_relat_1(k6_weierstr(X1,X2),X3)|~v1_funct_1(k6_weierstr(X1,X2))|~v1_relat_1(k6_weierstr(X1,X2))|~r2_hidden(X5,X4))).#trainneg
cnf(i_0_5404, plain, (r1_xxreal_0(esk6_4(X1,X2,X3,X4),X5)|X5!=k2_seq_4(X6)|~v3_membered(X6)|~v1_xreal_0(X5)|~v1_xreal_0(esk6_4(X1,X2,X3,X4))|~v4_xxreal_2(X6)|X3!=k7_relat_1(X1,X2)|~v4_relat_1(X1,X6)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_5452, plain, (r1_xxreal_0(k6_numbers,k8_weierstr(X1,X2,X4))|v1_xboole_0(X2)|v2_struct_0(X1)|~r2_hidden(X3,X4)|~v2_compts_1(X4,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~r2_hidden(X3,X2)|~v4_pre_topc(X2,k3_pcomps_1(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_5470, plain, (r1_xxreal_0(esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4),k6_numbers)|r1_tarski(k7_relat_1(k6_weierstr(X1,X2),X3),X4)|v1_xboole_0(X2)|v2_struct_0(X1)|v1_xboole_0(X3)|~v1_xreal_0(esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4))|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X2,X3)!=k6_numbers)).#trainneg
cnf(i_0_5469, plain, (esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4)=k8_weierstr(X1,X2,X3)|r1_tarski(k7_relat_1(k6_weierstr(X1,X2),X3),X4)|v1_xboole_0(X2)|v2_struct_0(X1)|~r1_xxreal_0(k8_weierstr(X1,X2,X3),esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4))|~v1_xxreal_0(esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4))|~v1_xxreal_0(k8_weierstr(X1,X2,X3))|~v1_xreal_0(esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4))|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_5276, plain, (k9_xtuple_0(esk19_1(k2_zfmisc_1(k6_numbers,X1)))=k6_numbers)).#trainneg
cnf(i_0_553, plain, (v2_membered(esk28_1(X1))|v2_struct_0(X1)|~v2_membered(u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_5087, negated_conjecture, (v1_xxreal_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_5088, negated_conjecture, (v1_xxreal_0(esk3_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_5564, plain, (r1_xxreal_0(k1_hausdorf(X1,k6_numbers,k6_numbers),k6_numbers)|r1_xxreal_0(k6_numbers,k1_hausdorf(X1,k6_numbers,k6_numbers))|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_5321, plain, (r1_xxreal_0(X1,k1_hausdorf(X2,k6_numbers,u1_struct_0(k3_pcomps_1(X2))))|r1_xxreal_0(k1_hausdorf(X2,k6_numbers,u1_struct_0(k3_pcomps_1(X2))),X1)|v2_struct_0(X2)|~v1_xreal_0(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_554, plain, (v3_valued_0(esk28_1(X1))|v2_struct_0(X1)|~v3_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_556, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,esk28_1(X2))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_5472, negated_conjecture, (r1_xxreal_0(esk5_2(k7_relat_1(k6_weierstr(X1,X2),X2),X3),k1_hausdorf(X1,X2,X2))|r1_tarski(k7_relat_1(k6_weierstr(X1,X2),X2),X3)|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_xreal_0(esk5_2(k7_relat_1(k6_weierstr(X1,X2),X2),X3))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_5471, plain, (r1_xxreal_0(esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4),k6_numbers)|r1_tarski(k7_relat_1(k6_weierstr(X1,X2),X3),X4)|v1_xboole_0(X2)|v2_struct_0(X1)|v1_xboole_0(X3)|~v1_xreal_0(esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4))|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X3,X2)!=k6_numbers)).#trainneg
cnf(i_0_5354, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),X3))|~r2_hidden(X2,u1_struct_0(esk1_0))|~r2_hidden(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(X2,X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0)))).#trainneg
cnf(i_0_5459, negated_conjecture, (r1_tarski(k7_relat_1(k6_weierstr(esk1_0,X1),X1),X2)|~r2_hidden(k6_numbers,X2)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_5353, negated_conjecture, (k1_hausdorf(esk1_0,X1,X1)=k6_numbers|v1_xboole_0(k7_relat_1(k6_weierstr(esk1_0,X1),X1))|~r1_xxreal_0(k1_hausdorf(esk1_0,X1,X1),k6_numbers)|~v1_xxreal_0(k1_hausdorf(esk1_0,X1,X1))|~r2_hidden(X2,X1)|~v2_compts_1(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v4_pre_topc(X1,k3_pcomps_1(esk1_0)))).#trainneg
cnf(i_0_5119, negated_conjecture, (k1_hausdorf(esk1_0,X1,X1)=X4|v1_xboole_0(k7_relat_1(k6_weierstr(esk1_0,X1),X1))|~v1_xreal_0(X4)|~r1_xxreal_0(k1_hausdorf(esk1_0,X1,X1),X4)|~v1_xxreal_0(k1_hausdorf(esk1_0,X1,X1))|~r2_hidden(esk6_4(k6_weierstr(esk1_0,X1),X2,X3,X4),X1)|~v2_compts_1(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|X3!=k7_relat_1(k6_weierstr(esk1_0,X1),X2)|~v1_funct_1(k6_weierstr(esk1_0,X1))|~v1_relat_1(k6_weierstr(esk1_0,X1))|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_5114, plain, (r2_hidden(esk5_2(k9_xtuple_0(X1),X2),X3)|r1_tarski(k9_xtuple_0(X1),X2)|~v4_relat_1(X1,X3)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5480, negated_conjecture, (r1_xxreal_0(esk9_3(k7_relat_1(k6_weierstr(X1,X2),X2),X3,X4),k1_hausdorf(X1,X2,X2))|r1_xxreal_0(X4,k6_numbers)|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X2))|v1_xboole_0(X2)|v2_struct_0(X1)|X3!=k2_seq_4(k7_relat_1(k6_weierstr(X1,X2),X2))|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v4_xxreal_2(k7_relat_1(k6_weierstr(X1,X2),X2))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_557, plain, (v1_xxreal_0(esk21_1(X1))|v2_struct_0(X1)|~v2_membered(k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1008, plain, (X1=esk19_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_5130, plain, (v3_membered(k2_pcomps_1(esk20_0))|~v3_membered(k1_zfmisc_1(k6_numbers))|~l1_metric_1(esk20_0))).#trainneg
cnf(i_0_5131, plain, (v2_membered(k2_pcomps_1(esk20_0))|~v2_membered(k1_zfmisc_1(k6_numbers))|~l1_metric_1(esk20_0))).#trainneg
cnf(i_0_558, plain, (v1_xreal_0(esk21_1(X1))|v2_struct_0(X1)|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_559, plain, (v1_xcmplx_0(esk21_1(X1))|v2_struct_0(X1)|~v1_membered(k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5661, negated_conjecture, (k1_hausdorf(X1,X2,X2)=k1_funct_1(k6_weierstr(X1,X2),X3)|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~r1_xxreal_0(k1_hausdorf(X1,X2,X2),k1_funct_1(k6_weierstr(X1,X2),X3))|~v1_xxreal_0(k1_hausdorf(X1,X2,X2))|~r2_hidden(X3,X2)|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_5916, plain, (r1_tarski(k9_xtuple_0(X1),X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5671, plain, (r1_xxreal_0(X1,k8_weierstr(X2,X3,X4))|v1_xboole_0(X3)|v2_struct_0(X2)|X5!=k7_relat_1(k6_weierstr(X2,X3),X4)|~v1_xreal_0(X1)|~r2_hidden(X1,X5)|~v2_compts_1(X4,k3_pcomps_1(X2))|~v2_compts_1(X3,k3_pcomps_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v1_funct_1(k6_weierstr(X2,X3))|~v1_relat_1(k6_weierstr(X2,X3)))).#trainneg
cnf(i_0_5655, plain, (k8_weierstr(X1,X2,X3)=X6|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_xreal_0(X6)|~r1_xxreal_0(k8_weierstr(X1,X2,X3),X6)|~v1_xxreal_0(k8_weierstr(X1,X2,X3))|~r2_hidden(esk6_4(k6_weierstr(X1,X2),X4,X5,X6),X3)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|X5!=k7_relat_1(k6_weierstr(X1,X2),X4)|~v1_funct_1(k6_weierstr(X1,X2))|~v1_relat_1(k6_weierstr(X1,X2))|~r2_hidden(X6,X5))).#trainneg
cnf(i_0_5646, negated_conjecture, (k6_numbers=X1|v1_xboole_0(X3)|X2!=k7_relat_1(k6_weierstr(esk1_0,X3),X4)|~v1_funct_1(k6_weierstr(esk1_0,X3))|~r2_hidden(X1,X2)|~v4_pre_topc(X3,k3_pcomps_1(esk1_0))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk6_4(k6_weierstr(esk1_0,X3),X4,X2,X1),X3))).#trainneg
cnf(i_0_5840, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),X2))|v1_xboole_0(X2)|~r2_hidden(X3,X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_5845, negated_conjecture, (k7_relat_1(k6_weierstr(esk1_0,X1),X1)=X2|~m1_subset_1(X2,k1_zfmisc_1(k7_relat_1(k6_weierstr(esk1_0,X1),X1)))|~r2_hidden(k6_numbers,X2)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_5658, plain, (k8_weierstr(X1,X2,X3)=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|~r1_xxreal_0(k8_weierstr(X1,X2,X3),k6_numbers)|~v1_xxreal_0(k8_weierstr(X1,X2,X3))|~r2_hidden(X4,X3)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~r2_hidden(X4,X2)|~v4_pre_topc(X2,k3_pcomps_1(X1))|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_5659, plain, (k1_funct_1(k6_weierstr(X1,X2),X4)=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|v1_xboole_0(X3)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X4))|~r1_xxreal_0(k6_numbers,k1_funct_1(k6_weierstr(X1,X2),X4))|~r2_hidden(X4,X3)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X2,X3)!=k6_numbers)).#trainneg
cnf(i_0_5400, plain, (~v1_xboole_0(X5)|X3!=k7_relat_1(X1,X2)|~v4_relat_1(X1,X5)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_560, plain, (v7_ordinal1(esk21_1(X1))|v2_struct_0(X1)|~v6_membered(k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1356, plain, (r1_xxreal_0(X1,k4_ordinal1)|r1_xxreal_0(k4_ordinal1,X1)|~v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5255, plain, (v3_membered(k2_pcomps_1(esk18_0))|~v3_membered(k1_zfmisc_1(k6_numbers))|~l1_metric_1(esk18_0))).#trainneg
cnf(i_0_5256, plain, (v2_membered(k2_pcomps_1(esk18_0))|~v2_membered(k1_zfmisc_1(k6_numbers))|~l1_metric_1(esk18_0))).#trainneg
cnf(i_0_5322, plain, (r1_xxreal_0(X1,k1_hausdorf(X2,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X2))))|r1_xxreal_0(k1_hausdorf(X2,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X2))),X1)|v2_struct_0(X2)|~v1_xreal_0(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_5605, plain, (r1_xxreal_0(k1_hausdorf(X1,u1_struct_0(k3_pcomps_1(X1)),k6_numbers),k6_numbers)|r1_xxreal_0(k6_numbers,k1_hausdorf(X1,u1_struct_0(k3_pcomps_1(X1)),k6_numbers))|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_562, plain, (v3_membered(esk21_1(X1))|v2_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5651, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,X3),X4)|~v1_funct_1(k6_weierstr(esk1_0,X3))|~r2_hidden(X1,X2)|~v4_pre_topc(X3,k3_pcomps_1(esk1_0))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v4_relat_1(k6_weierstr(esk1_0,X3),X3)|~v1_relat_1(k6_weierstr(esk1_0,X3)))).#trainneg
cnf(i_0_5696, negated_conjecture, (r1_xxreal_0(k6_numbers,k8_weierstr(esk1_0,X1,X2))|~r2_hidden(esk6_4(X3,esk3_0,X4,X5),X2)|~r2_hidden(esk6_4(X3,esk3_0,X4,X5),X1)|~v2_compts_1(X2,k3_pcomps_1(esk1_0))|~v2_compts_1(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|X4!=k7_relat_1(X3,esk3_0)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(X5,X4))).#trainneg
cnf(i_0_5451, negated_conjecture, (r1_xxreal_0(k6_numbers,k8_weierstr(esk1_0,X1,X3))|v1_xboole_0(X1)|~r2_hidden(X2,X3)|~v2_compts_1(X3,k3_pcomps_1(esk1_0))|~v2_compts_1(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(X2,X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0)))).#trainneg
cnf(i_0_5236, negated_conjecture, (k1_hausdorf(X1,X2,X2)=k6_numbers|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X2))|v2_struct_0(X1)|~r1_xxreal_0(k1_hausdorf(X1,X2,X2),k6_numbers)|~r2_hidden(X3,X2)|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v1_xreal_0(k1_hausdorf(X1,X2,X2)))).#trainneg
cnf(i_0_6070, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_numbers))|~r2_hidden(k6_numbers,X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_5939, plain, (esk19_1(X1)=esk19_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_563, plain, (v2_membered(esk21_1(X1))|v2_struct_0(X1)|~v2_membered(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5085, plain, (v1_xxreal_0(u1_struct_0(k3_topmetr))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_5125, plain, (v1_xxreal_0(k2_pcomps_1(esk20_0))|~v2_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers)))|~l1_metric_1(esk20_0))).#trainneg
cnf(i_0_5792, plain, (r1_xxreal_0(k1_hausdorf(X1,k6_numbers,u1_struct_0(k3_pcomps_1(X1))),k6_numbers)|r1_xxreal_0(k6_numbers,k1_hausdorf(X1,k6_numbers,u1_struct_0(k3_pcomps_1(X1))))|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_564, plain, (v3_valued_0(esk21_1(X1))|v2_struct_0(X1)|~v3_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_566, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,esk21_1(X2))|~v2_pre_topc(X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_6105, plain, (k1_funct_1(k6_weierstr(X1,X2),X3)=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|k1_hausdorf(X1,X2,X4)!=k6_numbers|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~r2_hidden(X3,X4)|~v2_compts_1(X4,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1))))).#trainneg
cnf(i_0_6066, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_numbers))|~r2_hidden(k2_weierstr(X2),X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_6064, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_numbers))|v1_xboole_0(X2)|~r2_hidden(k4_seq_4(X2),X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_6071, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_zfmisc_1(X2)))|~r2_hidden(k6_numbers,X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_5987, plain, (k8_weierstr(X1,X2,X3)=X4|v1_xboole_0(X2)|v2_struct_0(X1)|X5!=k7_relat_1(k6_weierstr(X1,X2),X3)|~v1_xreal_0(X4)|~r1_xxreal_0(k8_weierstr(X1,X2,X3),X4)|~v1_xxreal_0(k8_weierstr(X1,X2,X3))|~r2_hidden(X4,X5)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v1_funct_1(k6_weierstr(X1,X2))|~v1_relat_1(k6_weierstr(X1,X2)))).#trainneg
cnf(i_0_5958, negated_conjecture, (k1_hausdorf(X1,X2,X2)=X5|v2_struct_0(X1)|~v1_xreal_0(X5)|~r1_xxreal_0(k1_hausdorf(X1,X2,X2),X5)|~v1_xxreal_0(k1_hausdorf(X1,X2,X2))|~r2_hidden(esk6_4(k6_weierstr(X1,X2),X3,X4,X5),X2)|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|X4!=k7_relat_1(k6_weierstr(X1,X2),X3)|~v1_funct_1(k6_weierstr(X1,X2))|~v1_relat_1(k6_weierstr(X1,X2))|~r2_hidden(X5,X4))).#trainneg
cnf(i_0_6038, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0))))|~r2_hidden(esk2_0,X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_5263, plain, (m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~r2_hidden(X1,k2_pcomps_1(esk18_0))|~l1_metric_1(esk18_0))).#trainneg
cnf(i_0_5282, plain, (k9_xtuple_0(X1)=k6_numbers|~v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_573, plain, (v1_xxreal_0(k7_relset_1(X1,X2,X3,X4))|~v2_membered(k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5126, plain, (v1_xreal_0(k2_pcomps_1(esk20_0))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers)))|~l1_metric_1(esk20_0))).#trainneg
cnf(i_0_5127, plain, (v1_xcmplx_0(k2_pcomps_1(esk20_0))|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers)))|~l1_metric_1(esk20_0))).#trainneg
cnf(i_0_6151, plain, (r1_xxreal_0(k1_hausdorf(X1,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X1))),k6_numbers)|r1_xxreal_0(k6_numbers,k1_hausdorf(X1,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X1))))|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_574, plain, (v1_xreal_0(k7_relset_1(X1,X2,X3,X4))|~v3_membered(k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_575, plain, (v1_xcmplx_0(k7_relset_1(X1,X2,X3,X4))|~v1_membered(k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6037, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0))))|~r2_hidden(esk3_0,X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_5974, plain, (r1_tarski(u1_struct_0(k3_pcomps_1(X1)),X3)|v2_struct_0(X1)|~v4_relat_1(k6_weierstr(X1,X2),X3)|~v1_relat_1(k6_weierstr(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_6041, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_numbers))|v1_xboole_0(X2)|~r2_hidden(k2_seq_4(X2),X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_6065, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_numbers))|~r2_hidden(k4_seq_4(X2),X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6326, plain, (k1_funct_1(k6_weierstr(X1,X2),X3)=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|k1_hausdorf(X1,X2,X4)!=k6_numbers|~r2_hidden(X3,X4)|~v2_compts_1(X4,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(k1_funct_1(k6_weierstr(X1,X2),X3),k1_numbers))).#trainneg
cnf(i_0_6263, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk19_1(X3))|~v1_xboole_0(X2)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_5128, plain, (v7_ordinal1(k2_pcomps_1(esk20_0))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers)))|~l1_metric_1(esk20_0))).#trainneg
cnf(i_0_6409, plain, (X1!=k7_relat_1(X2,X3)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X4,X1)|~m1_subset_1(X2,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_576, plain, (v7_ordinal1(k7_relset_1(X1,X2,X3,X4))|~v6_membered(k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5250, plain, (v1_xxreal_0(k2_pcomps_1(esk18_0))|~v2_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers)))|~l1_metric_1(esk18_0))).#trainneg
cnf(i_0_5251, plain, (v1_xreal_0(k2_pcomps_1(esk18_0))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers)))|~l1_metric_1(esk18_0))).#trainneg
cnf(i_0_577, plain, (v1_xboole_0(X1)|v1_xboole_0(k7_relset_1(X2,k9_xtuple_0(X1),X3,X4))|~v1_relat_1(X1)|~v1_xboole_0(k7_relat_1(X1,k7_relset_1(X2,k9_xtuple_0(X1),X3,X4)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k9_xtuple_0(X1)))))).#trainneg
cnf(i_0_578, plain, (v1_relat_1(k7_relset_1(X1,k2_zfmisc_1(X2,X3),X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_6236, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,esk3_0),X1)=k6_numbers|~r2_hidden(X1,esk2_0)|~v3_valued_0(k6_weierstr(esk1_0,esk3_0))|~v1_funct_1(k6_weierstr(esk1_0,esk3_0))|~v1_relat_1(k6_weierstr(esk1_0,esk3_0)))).#trainneg
cnf(i_0_6365, plain, (k6_numbers=X4|v1_xboole_0(X2)|v2_struct_0(X1)|v1_xboole_0(X3)|X5!=k7_relat_1(k6_weierstr(X1,X2),X3)|~v1_xreal_0(X4)|~r1_xxreal_0(k6_numbers,X4)|~r2_hidden(X4,X5)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X3,X2)!=k6_numbers)).#trainneg
cnf(i_0_6328, plain, (k1_funct_1(k6_weierstr(X1,X2),X3)=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|k1_hausdorf(X1,X2,X4)!=k6_numbers|~r2_hidden(X3,X4)|~v2_compts_1(X4,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v3_valued_0(k6_weierstr(X1,X2))|~v1_funct_1(k6_weierstr(X1,X2))|~v1_relat_1(k6_weierstr(X1,X2)))).#trainneg
cnf(i_0_6325, plain, (X5=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|k1_hausdorf(X1,X2,X6)!=k6_numbers|~v1_xreal_0(X5)|~r2_hidden(esk6_4(k6_weierstr(X1,X2),X3,X4,X5),X6)|~v2_compts_1(X6,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|X4!=k7_relat_1(k6_weierstr(X1,X2),X3)|~v1_funct_1(k6_weierstr(X1,X2))|~v1_relat_1(k6_weierstr(X1,X2))|~r2_hidden(X5,X4))).#trainneg
cnf(i_0_6364, plain, (k6_numbers=X4|v1_xboole_0(X2)|v2_struct_0(X1)|v1_xboole_0(X3)|X5!=k7_relat_1(k6_weierstr(X1,X2),X3)|~v1_xreal_0(X4)|~r1_xxreal_0(k6_numbers,X4)|~r2_hidden(X4,X5)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X2,X3)!=k6_numbers)).#trainneg
cnf(i_0_6231, negated_conjecture, (k1_funct_1(k6_weierstr(esk1_0,esk2_0),X1)=k6_numbers|~r2_hidden(X1,esk3_0)|~v3_valued_0(k6_weierstr(esk1_0,esk2_0))|~v1_funct_1(k6_weierstr(esk1_0,esk2_0))|~v1_relat_1(k6_weierstr(esk1_0,esk2_0)))).#trainneg
cnf(i_0_6214, negated_conjecture, (m1_subset_1(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_numbers))|~r2_hidden(k6_numbers,X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_6233, negated_conjecture, (X3=k6_numbers|~v1_xreal_0(X3)|~r2_hidden(esk6_4(k6_weierstr(esk1_0,esk3_0),X1,X2,X3),esk2_0)|X2!=k7_relat_1(k6_weierstr(esk1_0,esk3_0),X1)|~v1_funct_1(k6_weierstr(esk1_0,esk3_0))|~v1_relat_1(k6_weierstr(esk1_0,esk3_0))|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_6213, negated_conjecture, (~v1_xboole_0(k7_relat_1(k6_weierstr(esk1_0,X1),k1_numbers))|~r2_hidden(k6_numbers,X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_6228, negated_conjecture, (X3=k6_numbers|~v1_xreal_0(X3)|~r2_hidden(esk6_4(k6_weierstr(esk1_0,esk2_0),X1,X2,X3),esk3_0)|X2!=k7_relat_1(k6_weierstr(esk1_0,esk2_0),X1)|~v1_funct_1(k6_weierstr(esk1_0,esk2_0))|~v1_relat_1(k6_weierstr(esk1_0,esk2_0))|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_5252, plain, (v1_xcmplx_0(k2_pcomps_1(esk18_0))|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers)))|~l1_metric_1(esk18_0))).#trainneg
cnf(i_0_5502, plain, (X1!=k7_relat_1(k2_zfmisc_1(k6_numbers,X2),X3)|~v1_funct_1(k2_zfmisc_1(k6_numbers,X2))|~r2_hidden(X4,X1))).#trainneg
cnf(i_0_579, plain, (v1_xboole_0(k7_relset_1(X1,k2_zfmisc_1(X2,X3),X4,X5))|~v1_xboole_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_5253, plain, (v7_ordinal1(k2_pcomps_1(esk18_0))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers)))|~l1_metric_1(esk18_0))).#trainneg
cnf(i_0_5521, plain, (r1_xxreal_0(k4_ordinal1,k4_ordinal1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_580, plain, (v1_xboole_0(k7_relset_1(X1,k2_zfmisc_1(X2,X3),X4,X5))|~v1_xboole_0(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_581, plain, (v4_relat_1(k7_relset_1(X1,k2_zfmisc_1(X2,X3),X4,X5),X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_6496, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_numbers))|~r2_hidden(k2_seq_4(X2),X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6081, negated_conjecture, (r1_xxreal_0(X3,k6_numbers)|~v1_xreal_0(X3)|~r2_hidden(esk6_4(k6_weierstr(esk1_0,esk3_0),X1,X2,X3),esk2_0)|X2!=k7_relat_1(k6_weierstr(esk1_0,esk3_0),X1)|~v1_funct_1(k6_weierstr(esk1_0,esk3_0))|~v1_relat_1(k6_weierstr(esk1_0,esk3_0))|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_5960, negated_conjecture, (k1_hausdorf(esk1_0,X1,X1)=k6_numbers|~r1_xxreal_0(k1_hausdorf(esk1_0,X1,X1),k6_numbers)|~v1_xxreal_0(k1_hausdorf(esk1_0,X1,X1))|~r2_hidden(X2,X1)|~v2_compts_1(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v4_pre_topc(X1,k3_pcomps_1(esk1_0)))).#trainneg
cnf(i_0_6024, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_numbers))|~r2_hidden(esk6_4(X2,X3,X4,X5),X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|X4!=k7_relat_1(X2,X3)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X5,X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_6072, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_numbers))|v2_struct_0(X2)|~r2_hidden(k1_hausdorf(X2,X3,X4),X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_6031, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_numbers))|r1_tarski(X2,X3)|~r2_hidden(esk5_2(X2,X3),X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_6100, plain, (k8_weierstr(X1,X2,X3)=k6_numbers|v2_struct_0(X1)|~r1_xxreal_0(k8_weierstr(X1,X2,X3),k6_numbers)|~r2_hidden(X4,X3)|~r2_hidden(X4,X2)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v1_xreal_0(k8_weierstr(X1,X2,X3)))).#trainneg
cnf(i_0_6023, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_numbers))|~r2_hidden(esk6_4(X2,X3,X4,X5),X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|X4!=k7_relat_1(X2,X3)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X5,X4)|~m1_subset_1(k9_xtuple_0(X2),k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_6250, plain, (m1_subset_1(esk19_1(X2),k1_zfmisc_1(X1))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3245, plain, (r1_xxreal_0(esk14_1(k4_ordinal1),k6_numbers)|r1_xxreal_0(k6_numbers,esk14_1(k4_ordinal1))|v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_582, plain, (v3_valued_0(k7_relset_1(X1,k2_zfmisc_1(X2,X3),X4,X5))|~v3_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_5565, plain, (r1_xxreal_0(k1_hausdorf(X1,k6_numbers,k6_numbers),k1_hausdorf(X2,X3,X4))|r1_xxreal_0(k1_hausdorf(X2,X3,X4),k1_hausdorf(X1,k6_numbers,k6_numbers))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_583, plain, (v2_valued_0(k7_relset_1(X1,k2_zfmisc_1(X2,X3),X4,X5))|~v2_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_584, plain, (v5_relat_1(k7_relset_1(X1,k2_zfmisc_1(X2,X3),X4,X5),X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_6507, plain, (k1_funct_1(k6_weierstr(X1,X2),X3)=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|k1_hausdorf(X1,X2,X4)!=k6_numbers|~r2_hidden(X3,X4)|~v2_compts_1(X4,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1))))).#trainneg
cnf(i_0_6586, plain, (k6_numbers=X1|v1_xboole_0(X2)|v1_xboole_0(X3)|v2_struct_0(X4)|k1_hausdorf(X4,X3,X2)!=k6_numbers|~v1_xreal_0(X1)|~r2_hidden(X1,k7_relat_1(k6_weierstr(X4,X3),X2))|~v2_compts_1(X2,k3_pcomps_1(X4))|~v2_compts_1(X3,k3_pcomps_1(X4))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~l1_metric_1(X4)|~v9_metric_1(X4)|~v8_metric_1(X4)|~v7_metric_1(X4)|~v6_metric_1(X4))).#trainneg
cnf(i_0_6566, plain, (k6_numbers=X1|v1_xboole_0(X2)|v1_xboole_0(X3)|v2_struct_0(X4)|k1_hausdorf(X4,X2,X3)!=k6_numbers|~v1_xreal_0(X1)|~r2_hidden(X1,k7_relat_1(k6_weierstr(X4,X3),X2))|~v2_compts_1(X2,k3_pcomps_1(X4))|~v2_compts_1(X3,k3_pcomps_1(X4))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~l1_metric_1(X4)|~v9_metric_1(X4)|~v8_metric_1(X4)|~v7_metric_1(X4)|~v6_metric_1(X4))).#trainneg
cnf(i_0_6689, negated_conjecture, (k1_hausdorf(X1,X2,X2)=k6_numbers|v2_struct_0(X1)|~r1_xxreal_0(k1_hausdorf(X1,X2,X2),k6_numbers)|~r2_hidden(X3,X2)|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_6655, negated_conjecture, (k1_hausdorf(esk1_0,X1,X1)=k6_numbers|~r1_xxreal_0(k1_hausdorf(esk1_0,X1,X1),k6_numbers)|~r2_hidden(X2,X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~v2_compts_1(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_xreal_0(k1_hausdorf(esk1_0,X1,X1)))).#trainneg
cnf(i_0_6688, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_numbers))|r1_tarski(X1,X2)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_6525, plain, (v1_xreal_0(k2_pcomps_1(esk20_0))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers)))|~l1_metric_1(esk20_0))).#trainneg
cnf(i_0_6600, negated_conjecture, (X1=k6_numbers|X2!=k7_relat_1(k6_weierstr(esk1_0,esk3_0),esk2_0)|~v1_xreal_0(X1)|~v1_funct_1(k6_weierstr(esk1_0,esk3_0))|~v1_relat_1(k6_weierstr(esk1_0,esk3_0))|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_585, plain, (v1_xboole_0(k7_relset_1(X1,X2,X3,X4))|~v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6621, plain, (v1_xreal_0(k2_pcomps_1(esk18_0))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers)))|~l1_metric_1(esk18_0))).#trainneg
cnf(i_0_5132, plain, (v3_valued_0(k2_pcomps_1(esk20_0))|~v3_valued_0(k1_zfmisc_1(k6_numbers))|~v1_relat_1(k1_zfmisc_1(k6_numbers))|~l1_metric_1(esk20_0))).#trainneg
cnf(i_0_6774, plain, (r1_xxreal_0(k1_hausdorf(X1,X2,k6_numbers),k1_hausdorf(X3,k6_numbers,k6_numbers))|r1_xxreal_0(k1_hausdorf(X3,k6_numbers,k6_numbers),k1_hausdorf(X1,X2,k6_numbers))|v2_struct_0(X1)|v2_struct_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X3)|~l1_metric_1(X1)|~v9_metric_1(X3)|~v9_metric_1(X1)|~v8_metric_1(X3)|~v8_metric_1(X1)|~v7_metric_1(X3)|~v7_metric_1(X1)|~v6_metric_1(X3)|~v6_metric_1(X1))).#trainneg
cnf(i_0_6775, plain, (r1_xxreal_0(k1_hausdorf(X1,X2,u1_struct_0(k3_pcomps_1(X1))),k1_hausdorf(X3,k6_numbers,k6_numbers))|r1_xxreal_0(k1_hausdorf(X3,k6_numbers,k6_numbers),k1_hausdorf(X1,X2,u1_struct_0(k3_pcomps_1(X1))))|v2_struct_0(X1)|v2_struct_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X3)|~l1_metric_1(X1)|~v9_metric_1(X3)|~v9_metric_1(X1)|~v8_metric_1(X3)|~v8_metric_1(X1)|~v7_metric_1(X3)|~v7_metric_1(X1)|~v6_metric_1(X3)|~v6_metric_1(X1))).#trainneg
cnf(i_0_586, plain, (v3_membered(k7_relset_1(X1,X2,X3,X4))|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_587, plain, (v2_membered(k7_relset_1(X1,X2,X3,X4))|~v2_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6640, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_numbers))|v1_xboole_0(k7_relat_1(k6_weierstr(X2,X3),X4))|v2_struct_0(X2)|~r2_hidden(k8_weierstr(X2,X3,X4),X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k7_relat_1(k6_weierstr(X2,X3),X4),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k7_relat_1(k6_weierstr(X2,X3),X4),k1_zfmisc_1(u1_struct_0(k3_topmetr)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_6667, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_numbers))|X2!=k7_relat_1(X3,X1)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(X4,X2)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_6579, plain, (X1=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X3)|X4!=k7_relat_1(k6_weierstr(X3,X2),X5)|k1_hausdorf(X3,X2,X5)!=k6_numbers|~v1_xreal_0(X1)|~r2_hidden(X1,X4)|~v2_compts_1(X5,k3_pcomps_1(X3))|~v2_compts_1(X2,k3_pcomps_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)|~v1_funct_1(k6_weierstr(X3,X2))|~v1_relat_1(k6_weierstr(X3,X2)))).#trainneg
cnf(i_0_5419, plain, (k6_numbers=X1|v2_struct_0(X2)|X3!=k7_relat_1(k6_weierstr(X2,X4),X5)|~v4_relat_1(k6_weierstr(X2,X4),u1_struct_0(X2))|~r2_hidden(X1,X3)|~v4_pre_topc(X4,k3_pcomps_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v4_relat_1(k6_weierstr(X2,X4),X4)|~v1_funct_1(k6_weierstr(X2,X4))|~v1_relat_1(k6_weierstr(X2,X4)))).#trainneg
cnf(i_0_5660, plain, (k1_funct_1(k6_weierstr(X1,X2),X4)=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|v1_xboole_0(X3)|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X4))|~r1_xxreal_0(k6_numbers,k1_funct_1(k6_weierstr(X1,X2),X4))|~r2_hidden(X4,X3)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X3,X2)!=k6_numbers)).#trainneg
cnf(i_0_5867, plain, (r1_xxreal_0(k1_funct_1(k6_weierstr(X1,X2),X3),k6_numbers)|v1_xboole_0(X2)|v2_struct_0(X1)|k1_hausdorf(X1,X4,X2)!=k6_numbers|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~r2_hidden(X3,X4)|~v2_compts_1(X4,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_6036, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_zfmisc_1(k1_numbers)))|v1_xboole_0(X2)|~r2_hidden(esk4_1(X2),X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_6052, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_numbers))|~r2_hidden(k8_weierstr(esk1_0,X2,X3),X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_5479, plain, (r1_xxreal_0(esk9_3(k7_relat_1(k6_weierstr(X1,X2),X3),X4,X5),k6_numbers)|r1_xxreal_0(X5,k6_numbers)|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X3))|v1_xboole_0(X2)|v2_struct_0(X1)|v1_xboole_0(X3)|X4!=k2_seq_4(k7_relat_1(k6_weierstr(X1,X2),X3))|~v1_xreal_0(X5)|~v1_xreal_0(X4)|~v4_xxreal_2(k7_relat_1(k6_weierstr(X1,X2),X3))|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X3,X2)!=k6_numbers)).#trainneg
cnf(i_0_3268, plain, (r1_xxreal_0(esk14_1(k4_ordinal1),esk23_0)|r1_xxreal_0(esk23_0,esk14_1(k4_ordinal1))|v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_6606, negated_conjecture, (X1=k6_numbers|X2!=k7_relat_1(k6_weierstr(esk1_0,esk2_0),esk3_0)|~v1_xreal_0(X1)|~v1_funct_1(k6_weierstr(esk1_0,esk2_0))|~v1_relat_1(k6_weierstr(esk1_0,esk2_0))|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_588, plain, (v3_valued_0(k7_relset_1(X1,X2,X3,X4))|~v3_valued_0(X2)|~v1_relat_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5257, plain, (v3_valued_0(k2_pcomps_1(esk18_0))|~v3_valued_0(k1_zfmisc_1(k6_numbers))|~v1_relat_1(k1_zfmisc_1(k6_numbers))|~l1_metric_1(esk18_0))).#trainneg
cnf(i_0_5606, plain, (r1_xxreal_0(k1_hausdorf(X1,u1_struct_0(k3_pcomps_1(X1)),k6_numbers),k1_hausdorf(X2,X3,X4))|r1_xxreal_0(k1_hausdorf(X2,X3,X4),k1_hausdorf(X1,u1_struct_0(k3_pcomps_1(X1)),k6_numbers))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_592, plain, (m1_subset_1(esk29_2(X1,X2),X1)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_593, plain, (m1_subset_1(esk29_2(X1,X2),X2)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_6962, plain, (v3_membered(k7_relat_1(X3,X4))|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6846, plain, (esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4)=k6_numbers|v1_xboole_0(X2)|v1_xboole_0(X3)|v2_struct_0(X1)|r1_tarski(k7_relat_1(k6_weierstr(X1,X2),X3),X4)|k1_hausdorf(X1,X2,X3)!=k6_numbers|~v1_xreal_0(esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4))|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_7011, negated_conjecture, (X1=k6_numbers|~v1_xreal_0(X1)|~v1_funct_1(k6_weierstr(esk1_0,esk2_0))|~v1_relat_1(k6_weierstr(esk1_0,esk2_0))|~r2_hidden(X1,k7_relat_1(k6_weierstr(esk1_0,esk2_0),esk3_0)))).#trainneg
cnf(i_0_6862, plain, (k1_funct_1(k6_weierstr(X1,X2),X3)=k6_numbers|v1_xboole_0(X2)|v1_xboole_0(X4)|v2_struct_0(X1)|k1_hausdorf(X1,X4,X2)!=k6_numbers|~v1_xreal_0(k1_funct_1(k6_weierstr(X1,X2),X3))|~v2_compts_1(X4,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(k3_pcomps_1(X1))))).#trainneg
cnf(i_0_6860, plain, (esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4)=k6_numbers|v1_xboole_0(X2)|v1_xboole_0(X3)|v2_struct_0(X1)|r1_tarski(k7_relat_1(k6_weierstr(X1,X2),X3),X4)|k1_hausdorf(X1,X3,X2)!=k6_numbers|~v1_xreal_0(esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4))|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_6809, negated_conjecture, (k6_numbers=X3|X2!=k7_relat_1(k6_weierstr(esk1_0,esk3_0),X1)|~v1_funct_1(k6_weierstr(esk1_0,esk3_0))|~v1_relat_1(k6_weierstr(esk1_0,esk3_0))|~r2_hidden(X3,X2)|~r2_hidden(esk6_4(k6_weierstr(esk1_0,esk3_0),X1,X2,X3),esk2_0))).#trainneg
cnf(i_0_6834, plain, (k1_funct_1(k6_weierstr(X1,X2),X3)=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|k1_hausdorf(X1,X4,X2)!=k6_numbers|~r2_hidden(X3,X4)|~v2_compts_1(X4,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_6903, negated_conjecture, (X1=k6_numbers|~v1_xreal_0(X1)|~v1_funct_1(k6_weierstr(esk1_0,esk3_0))|~v1_relat_1(k6_weierstr(esk1_0,esk3_0))|~r2_hidden(X1,k7_relat_1(k6_weierstr(esk1_0,esk3_0),esk2_0)))).#trainneg
cnf(i_0_5716, plain, (X1!=k7_relat_1(esk19_1(k2_zfmisc_1(k6_numbers,X2)),X3)|~v1_funct_1(esk19_1(k2_zfmisc_1(k6_numbers,X2)))|~r2_hidden(X4,X1))).#trainneg
cnf(i_0_594, plain, (X1=k6_numbers|v1_funct_2(X2,X3,X1)|k9_xtuple_0(X2)!=X3|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~v4_relat_1(X2,X3)|~v1_relat_1(X2))).#trainneg
cnf(i_0_5949, plain, (v3_membered(u1_pre_topc(k3_pcomps_1(esk20_0)))|~v3_membered(k1_zfmisc_1(k6_numbers))|~l1_metric_1(esk20_0))).#trainneg
cnf(i_0_5950, plain, (v2_membered(u1_pre_topc(k3_pcomps_1(esk20_0)))|~v2_membered(k1_zfmisc_1(k6_numbers))|~l1_metric_1(esk20_0))).#trainneg
cnf(i_0_6939, plain, (r1_xxreal_0(k1_hausdorf(X1,k6_numbers,k6_numbers),k1_hausdorf(X2,k6_numbers,k6_numbers))|r1_xxreal_0(k1_hausdorf(X2,k6_numbers,k6_numbers),k1_hausdorf(X1,k6_numbers,k6_numbers))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_metric_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X1)|~v9_metric_1(X2)|~v8_metric_1(X1)|~v8_metric_1(X2)|~v7_metric_1(X1)|~v7_metric_1(X2)|~v6_metric_1(X1)|~v6_metric_1(X2))).#trainneg
cnf(i_0_6940, plain, (r1_xxreal_0(k1_hausdorf(X1,k6_numbers,k6_numbers),k1_hausdorf(X2,u1_struct_0(k3_pcomps_1(X2)),k6_numbers))|r1_xxreal_0(k1_hausdorf(X2,u1_struct_0(k3_pcomps_1(X2)),k6_numbers),k1_hausdorf(X1,k6_numbers,k6_numbers))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_metric_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X1)|~v9_metric_1(X2)|~v8_metric_1(X1)|~v8_metric_1(X2)|~v7_metric_1(X1)|~v7_metric_1(X2)|~v6_metric_1(X1)|~v6_metric_1(X2))).#trainneg
cnf(i_0_597, plain, (m1_subset_1(esk6_4(X1,X2,X3,X4),X2)|X3!=k7_relat_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_599, plain, (v1_xxreal_0(X4)|~v2_valued_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|X3!=k7_relat_1(X1,X2)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_6976, plain, (X1=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X3)|k1_hausdorf(X3,X2,X4)!=k6_numbers|~v1_xreal_0(X1)|~r2_hidden(X1,k7_relat_1(k6_weierstr(X3,X2),X4))|~v2_compts_1(X4,k3_pcomps_1(X3))|~v2_compts_1(X2,k3_pcomps_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_6863, plain, (k6_numbers=X1|v1_xboole_0(X2)|v1_xboole_0(X3)|v2_struct_0(X4)|v1_xboole_0(k7_relat_1(k6_weierstr(X4,X2),X3))|k1_hausdorf(X4,X3,X2)!=k6_numbers|~v1_xreal_0(X1)|~v2_compts_1(X3,k3_pcomps_1(X4))|~v2_compts_1(X2,k3_pcomps_1(X4))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~l1_metric_1(X4)|~v9_metric_1(X4)|~v8_metric_1(X4)|~v7_metric_1(X4)|~v6_metric_1(X4)|~m1_subset_1(X1,k7_relat_1(k6_weierstr(X4,X2),X3)))).#trainneg
cnf(i_0_6852, plain, (esk9_3(k7_relat_1(k6_weierstr(X1,X2),X3),X4,X5)=k6_numbers|v1_xboole_0(X2)|v1_xboole_0(X3)|v2_struct_0(X1)|r1_xxreal_0(X5,k6_numbers)|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X3))|k1_hausdorf(X1,X3,X2)!=k6_numbers|~v1_xreal_0(esk9_3(k7_relat_1(k6_weierstr(X1,X2),X3),X4,X5))|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|X4!=k2_seq_4(k7_relat_1(k6_weierstr(X1,X2),X3))|~v3_membered(k7_relat_1(k6_weierstr(X1,X2),X3))|~v1_xreal_0(X5)|~v1_xreal_0(X4)|~v4_xxreal_2(k7_relat_1(k6_weierstr(X1,X2),X3)))).#trainneg
cnf(i_0_7084, plain, (r1_xxreal_0(esk29_2(X1,X2),X3)|v1_xboole_0(X1)|v1_xboole_0(X2)|X3!=k2_seq_4(X1)|~v3_membered(X1)|~v1_xreal_0(X3)|~v4_xxreal_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_6986, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,X3),X4)|~v4_relat_1(k6_weierstr(esk1_0,X3),X3)|~r2_hidden(X1,X2)|~v4_pre_topc(X3,k3_pcomps_1(esk1_0))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7002, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_numbers))|~r2_hidden(k1_hausdorf(esk1_0,X2,X2),X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_6995, negated_conjecture, (r2_hidden(k6_numbers,k7_relat_1(k6_weierstr(esk1_0,X1),k1_zfmisc_1(k1_numbers)))|v1_xboole_0(X2)|~r2_hidden(X2,X1)|~v4_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_5278, plain, (k9_xtuple_0(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k6_numbers,X1))))=k6_numbers)).#trainneg
cnf(i_0_600, plain, (v1_xreal_0(X4)|~v3_valued_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|X3!=k7_relat_1(X1,X2)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_7231, plain, (X1=k6_numbers|v1_funct_2(k6_numbers,X2,X1)|k6_numbers!=X2)).#trainneg
cnf(i_0_6113, plain, (v3_membered(u1_pre_topc(k3_pcomps_1(esk18_0)))|~v3_membered(k1_zfmisc_1(k6_numbers))|~l1_metric_1(esk18_0))).#trainneg
cnf(i_0_6114, plain, (v2_membered(u1_pre_topc(k3_pcomps_1(esk18_0)))|~v2_membered(k1_zfmisc_1(k6_numbers))|~l1_metric_1(esk18_0))).#trainneg
cnf(i_0_6960, plain, (r1_xxreal_0(k1_hausdorf(X1,k6_numbers,k6_numbers),k1_hausdorf(X2,k6_numbers,u1_struct_0(k3_pcomps_1(X2))))|r1_xxreal_0(k1_hausdorf(X2,k6_numbers,u1_struct_0(k3_pcomps_1(X2))),k1_hausdorf(X1,k6_numbers,k6_numbers))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_metric_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X1)|~v9_metric_1(X2)|~v8_metric_1(X1)|~v8_metric_1(X2)|~v7_metric_1(X1)|~v7_metric_1(X2)|~v6_metric_1(X1)|~v6_metric_1(X2))).#trainneg
cnf(i_0_6961, plain, (r1_xxreal_0(k1_hausdorf(X1,k6_numbers,k6_numbers),k1_hausdorf(X2,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X2))))|r1_xxreal_0(k1_hausdorf(X2,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X2))),k1_hausdorf(X1,k6_numbers,k6_numbers))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_metric_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X1)|~v9_metric_1(X2)|~v8_metric_1(X1)|~v8_metric_1(X2)|~v7_metric_1(X1)|~v7_metric_1(X2)|~v6_metric_1(X1)|~v6_metric_1(X2))).#trainneg
cnf(i_0_601, plain, (v1_xcmplx_0(X4)|~v1_valued_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|X3!=k7_relat_1(X1,X2)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_603, plain, (m1_subset_1(esk6_4(X1,X2,X3,X4),k9_xtuple_0(X1))|X3!=k7_relat_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_7167, plain, (esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4)=k6_numbers|r1_tarski(k7_relat_1(k6_weierstr(X1,X2),X3),X4)|v1_xboole_0(X3)|v1_xboole_0(X2)|v2_struct_0(X1)|k1_hausdorf(X1,X2,X3)!=k6_numbers|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4),k1_numbers))).#trainneg
cnf(i_0_7297, plain, (r2_hidden(esk6_4(k6_weierstr(X1,X2),u1_struct_0(X1),X3,X4),X2)|v1_xboole_0(X2)|v2_struct_0(X1)|X3!=k7_relat_1(k6_weierstr(X1,X2),u1_struct_0(X1))|X4!=k6_numbers|~r2_hidden(X4,X3)|~v4_pre_topc(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v1_funct_1(k6_weierstr(X1,X2))|~v1_relat_1(k6_weierstr(X1,X2)))).#trainneg
cnf(i_0_7314, plain, (esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4)=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|r1_tarski(k7_relat_1(k6_weierstr(X1,X2),X3),X4)|k1_hausdorf(X1,X2,X3)!=k6_numbers|~v1_xreal_0(esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4))|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_7190, plain, (esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4)=k6_numbers|r1_tarski(k7_relat_1(k6_weierstr(X1,X2),X3),X4)|v1_xboole_0(X3)|v1_xboole_0(X2)|v2_struct_0(X1)|k1_hausdorf(X1,X3,X2)!=k6_numbers|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4),k1_numbers))).#trainneg
cnf(i_0_2452, plain, (v1_xboole_0(k7_relat_1(k2_zfmisc_1(X1,X2),X3))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_7198, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,esk3_0),esk2_0)|~v1_funct_1(k6_weierstr(esk1_0,esk3_0))|~v1_relat_1(k6_weierstr(esk1_0,esk3_0))|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_607, plain, (m1_subset_1(esk10_2(X1,X2),X1)|X2=k2_seq_4(X1)|v1_xreal_0(esk11_2(X1,X2))|v1_xboole_0(X1)|~v3_membered(X1)|~v1_xreal_0(X2)|~v4_xxreal_2(X1))).#trainneg
cnf(i_0_1686, plain, (v1_xboole_0(k1_relset_1(k6_numbers,X1))|~v1_xboole_0(X2)|~v4_relat_1(X1,k6_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7035, plain, (r1_xxreal_0(k1_hausdorf(X1,X2,k6_numbers),k1_hausdorf(X3,u1_struct_0(k3_pcomps_1(X3)),k6_numbers))|r1_xxreal_0(k1_hausdorf(X3,u1_struct_0(k3_pcomps_1(X3)),k6_numbers),k1_hausdorf(X1,X2,k6_numbers))|v2_struct_0(X1)|v2_struct_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X3)|~l1_metric_1(X1)|~v9_metric_1(X3)|~v9_metric_1(X1)|~v8_metric_1(X3)|~v8_metric_1(X1)|~v7_metric_1(X3)|~v7_metric_1(X1)|~v6_metric_1(X3)|~v6_metric_1(X1))).#trainneg
cnf(i_0_609, plain, (r2_hidden(X1,X2)|X5=k2_seq_4(k9_xtuple_0(X3))|v1_xreal_0(esk11_2(k9_xtuple_0(X3),X5))|v1_xboole_0(k9_xtuple_0(X3))|X2!=k7_relat_1(X3,X4)|X1!=k1_funct_1(X3,esk10_2(k9_xtuple_0(X3),X5))|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(esk10_2(k9_xtuple_0(X3),X5),X4)|~v3_membered(k9_xtuple_0(X3))|~v1_xreal_0(X5)|~v4_xxreal_2(k9_xtuple_0(X3)))).#trainneg
cnf(i_0_610, plain, (r2_hidden(X1,X2)|X2!=k7_relat_1(X3,X4)|X1!=k1_funct_1(X3,esk6_4(X5,k9_xtuple_0(X3),X6,X7))|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(esk6_4(X5,k9_xtuple_0(X3),X6,X7),X4)|X6!=k7_relat_1(X5,k9_xtuple_0(X3))|~v1_funct_1(X5)|~v1_relat_1(X5)|~r2_hidden(X7,X6))).#trainneg
cnf(i_0_7432, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,esk3_0),X3)|~v1_funct_1(k6_weierstr(esk1_0,esk3_0))|~r2_hidden(X1,X2)|~v4_pre_topc(esk3_0,k3_pcomps_1(esk1_0))|~r2_hidden(esk6_4(k6_weierstr(esk1_0,esk3_0),X3,X2,X1),esk2_0))).#trainneg
cnf(i_0_7564, plain, (esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4)=k6_numbers|r1_tarski(k7_relat_1(k6_weierstr(X1,X2),X3),X4)|v1_xboole_0(X2)|v2_struct_0(X1)|k1_hausdorf(X1,X2,X3)!=k6_numbers|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(esk5_2(k7_relat_1(k6_weierstr(X1,X2),X3),X4),k1_numbers))).#trainneg
cnf(i_0_7287, plain, (v1_xreal_0(esk6_4(X1,X2,X3,X4))|~v3_membered(X2)|X3!=k7_relat_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_7306, plain, (esk9_3(k7_relat_1(k6_weierstr(X1,X2),X3),X4,X5)=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X1)|r1_xxreal_0(X5,k6_numbers)|v1_xboole_0(k7_relat_1(k6_weierstr(X1,X2),X3))|k1_hausdorf(X1,X2,X3)!=k6_numbers|~v1_xreal_0(esk9_3(k7_relat_1(k6_weierstr(X1,X2),X3),X4,X5))|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|X4!=k2_seq_4(k7_relat_1(k6_weierstr(X1,X2),X3))|~v3_membered(k7_relat_1(k6_weierstr(X1,X2),X3))|~v1_xreal_0(X5)|~v1_xreal_0(X4)|~v4_xxreal_2(k7_relat_1(k6_weierstr(X1,X2),X3)))).#trainneg
cnf(i_0_7139, plain, (v1_xreal_0(esk29_2(X1,X2))|v1_xboole_0(X2)|~v3_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_7436, negated_conjecture, (k6_numbers=X1|v2_struct_0(X2)|X3!=k7_relat_1(k6_weierstr(X2,esk3_0),X4)|~v4_relat_1(k6_weierstr(X2,esk3_0),u1_struct_0(X2))|~r2_hidden(X1,X3)|~v4_pre_topc(esk3_0,k3_pcomps_1(X2))|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~r2_hidden(esk6_4(k6_weierstr(X2,esk3_0),X4,X3,X1),esk2_0))).#trainneg
cnf(i_0_7298, plain, (k6_numbers=X1|v2_struct_0(X2)|X3!=k7_relat_1(k6_weierstr(X2,X4),u1_struct_0(X2))|~v1_relat_1(k6_weierstr(X2,X4))|~r2_hidden(esk6_4(k6_weierstr(X2,X4),u1_struct_0(X2),X3,X1),X4)|~r2_hidden(X1,X3)|~v4_pre_topc(X4,k3_pcomps_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v1_funct_1(k6_weierstr(X2,X4)))).#trainneg
cnf(i_0_7590, plain, (k7_relat_1(k2_zfmisc_1(X1,X2),X3)=k6_numbers|~v1_xboole_0(X2))).#trainneg
cnf(i_0_7376, plain, (X1!=k7_relat_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k6_numbers,X2))),X3)|~v1_funct_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k6_numbers,X2))))|~r2_hidden(X4,X1))).#trainneg
cnf(i_0_611, plain, (r2_hidden(X1,X2)|X2!=k7_relat_1(X3,X4)|X1!=k1_funct_1(X3,esk6_4(X3,X5,X6,X7))|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(esk6_4(X3,X5,X6,X7),X4)|X6!=k7_relat_1(X3,X5)|~r2_hidden(X7,X6))).#trainneg
cnf(i_0_7404, plain, (k1_relset_1(X1,k6_numbers)=X1|X2=k6_numbers|k6_numbers!=X1)).#trainneg
cnf(i_0_6309, plain, (v1_xxreal_0(u1_pre_topc(k3_pcomps_1(esk20_0)))|~v2_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers)))|~l1_metric_1(esk20_0))).#trainneg
cnf(i_0_6442, plain, (v1_xreal_0(u1_pre_topc(k3_pcomps_1(esk20_0)))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers)))|~l1_metric_1(esk20_0))).#trainneg
cnf(i_0_7678, plain, (r1_xxreal_0(k1_hausdorf(X1,u1_struct_0(k3_pcomps_1(X1)),k6_numbers),k1_hausdorf(X2,u1_struct_0(k3_pcomps_1(X2)),k6_numbers))|r1_xxreal_0(k1_hausdorf(X2,u1_struct_0(k3_pcomps_1(X2)),k6_numbers),k1_hausdorf(X1,u1_struct_0(k3_pcomps_1(X1)),k6_numbers))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_metric_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X1)|~v9_metric_1(X2)|~v8_metric_1(X1)|~v8_metric_1(X2)|~v7_metric_1(X1)|~v7_metric_1(X2)|~v6_metric_1(X1)|~v6_metric_1(X2))).#trainneg
cnf(i_0_7036, plain, (r1_xxreal_0(k1_hausdorf(X1,X2,u1_struct_0(k3_pcomps_1(X1))),k1_hausdorf(X3,u1_struct_0(k3_pcomps_1(X3)),k6_numbers))|r1_xxreal_0(k1_hausdorf(X3,u1_struct_0(k3_pcomps_1(X3)),k6_numbers),k1_hausdorf(X1,X2,u1_struct_0(k3_pcomps_1(X1))))|v2_struct_0(X1)|v2_struct_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X3)|~l1_metric_1(X1)|~v9_metric_1(X3)|~v9_metric_1(X1)|~v8_metric_1(X3)|~v8_metric_1(X1)|~v7_metric_1(X3)|~v7_metric_1(X1)|~v6_metric_1(X3)|~v6_metric_1(X1))).#trainneg
cnf(i_0_612, plain, (r2_hidden(X1,X2)|r1_tarski(k9_xtuple_0(X3),X5)|X2!=k7_relat_1(X3,X4)|X1!=k1_funct_1(X3,esk5_2(k9_xtuple_0(X3),X5))|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(esk5_2(k9_xtuple_0(X3),X5),X4))).#trainneg
cnf(i_0_614, plain, (k1_relset_1(u1_struct_0(X1),esk17_2(X1,X2))=u1_struct_0(X1)|u1_struct_0(X2)=k6_numbers|v2_struct_0(X2)|~m1_subset_1(esk17_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_7573, plain, (r1_tarski(k7_relat_1(k6_weierstr(X1,X2),X3),X4)|v1_xboole_0(X2)|v1_xboole_0(X3)|v2_struct_0(X1)|~r2_hidden(k6_numbers,X4)|k1_hausdorf(X1,X3,X2)!=k6_numbers|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_7424, plain, (m1_subset_1(esk6_4(k6_weierstr(X1,X2),X3,X4,X5),u1_struct_0(k3_pcomps_1(X1)))|v2_struct_0(X1)|X4!=k7_relat_1(k6_weierstr(X1,X2),X3)|~v1_funct_1(k6_weierstr(X1,X2))|~v1_relat_1(k6_weierstr(X1,X2))|~r2_hidden(X5,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_7340, plain, (X1=esk29_2(X2,X3)|v1_xboole_0(X3)|~r1_xxreal_0(X1,esk29_2(X2,X3))|~v1_xxreal_0(esk29_2(X2,X3))|~v1_xxreal_0(X1)|X1!=k2_seq_4(X2)|~v3_membered(X2)|~v1_xreal_0(X1)|~v4_xxreal_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_7285, plain, (r1_xxreal_0(esk6_4(X1,X2,X3,X4),X5)|v1_xboole_0(X2)|X5!=k2_seq_4(X2)|~v3_membered(X2)|~v1_xreal_0(X5)|~v4_xxreal_2(X2)|X3!=k7_relat_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_7137, plain, (r1_xxreal_0(esk29_2(X1,X2),X3)|v1_xboole_0(X2)|X3!=k2_seq_4(X2)|~v3_membered(X2)|~v1_xreal_0(X3)|~v4_xxreal_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_7341, negated_conjecture, (k6_numbers=X1|~v4_relat_1(k6_weierstr(esk1_0,X2),X2)|~r2_hidden(X1,k7_relat_1(k6_weierstr(esk1_0,X2),X3))|~v4_pre_topc(X2,k3_pcomps_1(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7294, negated_conjecture, (r1_xxreal_0(k6_numbers,X1)|v1_xboole_0(X2)|X3!=k7_relat_1(k6_weierstr(esk1_0,X2),u1_struct_0(esk1_0))|~r2_hidden(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_1(k6_weierstr(esk1_0,X2))|~v1_relat_1(k6_weierstr(esk1_0,X2)))).#trainneg
cnf(i_0_7226, plain, (X1=k6_numbers|v1_funct_2(X2,k9_xtuple_0(X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k9_xtuple_0(X2),X1))))).#trainneg
cnf(i_0_5238, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relat_1(X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_615, plain, (esk17_2(X1,X2)=k6_numbers|v2_struct_0(X2)|u1_struct_0(X2)!=k6_numbers|~m1_subset_1(esk17_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_7960, plain, (X1=k6_numbers|v1_funct_2(k6_numbers,k6_numbers,X1))).#trainneg
cnf(i_0_6443, plain, (v1_xcmplx_0(u1_pre_topc(k3_pcomps_1(esk20_0)))|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers)))|~l1_metric_1(esk20_0))).#trainneg
cnf(i_0_6526, plain, (v7_ordinal1(u1_pre_topc(k3_pcomps_1(esk20_0)))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers)))|~l1_metric_1(esk20_0))).#trainneg
cnf(i_0_7864, plain, (r1_xxreal_0(k1_hausdorf(X1,u1_struct_0(k3_pcomps_1(X1)),k6_numbers),k1_hausdorf(X2,k6_numbers,u1_struct_0(k3_pcomps_1(X2))))|r1_xxreal_0(k1_hausdorf(X2,k6_numbers,u1_struct_0(k3_pcomps_1(X2))),k1_hausdorf(X1,u1_struct_0(k3_pcomps_1(X1)),k6_numbers))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_metric_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X1)|~v9_metric_1(X2)|~v8_metric_1(X1)|~v8_metric_1(X2)|~v7_metric_1(X1)|~v7_metric_1(X2)|~v6_metric_1(X1)|~v6_metric_1(X2))).#trainneg
cnf(i_0_616, plain, (X1=k7_relat_1(esk25_0,X2)|r2_hidden(esk8_3(esk25_0,X2,X1),X2)|r2_hidden(esk7_3(esk25_0,X2,X1),X1))).#trainneg
cnf(i_0_617, plain, (X1=k7_relat_1(esk17_2(X2,X3),X4)|r2_hidden(esk8_3(esk17_2(X2,X3),X4,X1),X4)|r2_hidden(esk7_3(esk17_2(X2,X3),X4,X1),X1)|v2_struct_0(X3)|~v1_relat_1(esk17_2(X2,X3))|~v2_pre_topc(X3)|~v2_pre_topc(X2)|~l1_pre_topc(X3)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_7718, plain, (r1_tarski(k7_relat_1(k6_weierstr(X1,X2),X3),X4)|v1_xboole_0(X2)|v2_struct_0(X1)|~r2_hidden(k6_numbers,X4)|k1_hausdorf(X1,X2,X3)!=k6_numbers|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_7905, plain, (r1_xxreal_0(esk6_4(k6_weierstr(X1,X2),X3,X4,X5),X6)|v1_xboole_0(u1_struct_0(k3_pcomps_1(X1)))|v2_struct_0(X1)|X6!=k2_seq_4(u1_struct_0(k3_pcomps_1(X1)))|~v3_membered(u1_struct_0(k3_pcomps_1(X1)))|~v1_xreal_0(X6)|~v4_xxreal_2(u1_struct_0(k3_pcomps_1(X1)))|X4!=k7_relat_1(k6_weierstr(X1,X2),X3)|~r2_hidden(X5,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_7928, plain, (X1=esk6_4(X2,X3,X4,X5)|~r1_xxreal_0(X1,esk6_4(X2,X3,X4,X5))|~v1_xxreal_0(esk6_4(X2,X3,X4,X5))|~v1_xxreal_0(X1)|X4!=k7_relat_1(X2,X3)|X1!=k2_seq_4(X3)|~v3_membered(X3)|~v1_xreal_0(X1)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X5,X4)|~v4_xxreal_2(X3))).#trainneg
cnf(i_0_7776, plain, (k6_numbers=X1|v2_struct_0(X2)|X3!=k7_relat_1(k6_weierstr(X2,X4),u1_struct_0(X2))|~r2_hidden(X1,X3)|~v4_pre_topc(X4,k3_pcomps_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v4_relat_1(k6_weierstr(X2,X4),X4)|~v1_funct_1(k6_weierstr(X2,X4))|~v1_relat_1(k6_weierstr(X2,X4)))).#trainneg
cnf(i_0_7916, plain, (X1=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X3)|X4!=k7_relat_1(k6_weierstr(X3,X2),X5)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k6_numbers)|~r2_hidden(X1,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_7727, plain, (r1_xxreal_0(k6_numbers,k8_weierstr(X1,X2,X3))|r1_tarski(k7_relat_1(k6_weierstr(X1,X2),X3),X4)|v1_xboole_0(X2)|v2_struct_0(X1)|~v2_compts_1(X3,k3_pcomps_1(X1))|~v2_compts_1(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|k1_hausdorf(X1,X2,X3)!=k6_numbers)).#trainneg
cnf(i_0_7915, plain, (X1=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X3)|X4!=k7_relat_1(k6_weierstr(X3,X2),X5)|~r1_xxreal_0(X1,k6_numbers)|~v1_xxreal_0(X1)|~r2_hidden(X1,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_6786, negated_conjecture, (k6_numbers=X3|X2!=k7_relat_1(k6_weierstr(esk1_0,esk2_0),X1)|~v1_funct_1(k6_weierstr(esk1_0,esk2_0))|~v1_relat_1(k6_weierstr(esk1_0,esk2_0))|~r2_hidden(X3,X2)|~r2_hidden(esk6_4(k6_weierstr(esk1_0,esk2_0),X1,X2,X3),esk3_0))).#trainneg
cnf(i_0_7760, negated_conjecture, (k6_numbers=X1|v2_struct_0(X2)|X3!=k7_relat_1(k6_weierstr(X2,esk3_0),esk2_0)|~v4_relat_1(k6_weierstr(X2,esk3_0),u1_struct_0(X2))|~r2_hidden(X1,X3)|~v4_pre_topc(esk3_0,k3_pcomps_1(X2))|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v1_funct_1(k6_weierstr(X2,esk3_0))|~v1_relat_1(k6_weierstr(X2,esk3_0)))).#trainneg
cnf(i_0_7866, plain, (r2_hidden(k1_funct_1(X1,esk5_2(k9_xtuple_0(X1),X2)),X3)|r1_tarski(k9_xtuple_0(X1),X2)|X3!=k7_relat_1(X1,X4)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(esk5_2(k9_xtuple_0(X1),X2),X4))).#trainneg
cnf(i_0_7967, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relat_1(esk19_1(k2_zfmisc_1(X3,X2)),X4)))).#trainneg
cnf(i_0_7971, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relat_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X3,X2))),X4)))).#trainneg
cnf(i_0_5994, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,k1_numbers),X3)|~v1_funct_1(k6_weierstr(esk1_0,k1_numbers))|~r2_hidden(X1,X2)|~v4_pre_topc(k1_numbers,k3_pcomps_1(esk1_0))|~m1_subset_1(k1_numbers,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_relat_1(k6_weierstr(esk1_0,k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_topmetr))))).#trainneg
cnf(i_0_618, plain, (X1=k7_relat_1(k6_weierstr(X2,X3),X4)|r2_hidden(esk8_3(k6_weierstr(X2,X3),X4,X1),X4)|r2_hidden(esk7_3(k6_weierstr(X2,X3),X4,X1),X1)|v2_struct_0(X2)|~v1_relat_1(k6_weierstr(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_7824, plain, (k1_relset_1(k6_numbers,k6_numbers)=k6_numbers|X1=k6_numbers)).#trainneg
cnf(i_0_619, plain, (X1=k7_relat_1(esk25_0,X2)|r2_hidden(esk8_3(esk25_0,X2,X1),k9_xtuple_0(esk25_0))|r2_hidden(esk7_3(esk25_0,X2,X1),X1))).#trainneg
cnf(i_0_620, plain, (X1=k7_relat_1(esk17_2(X2,X3),X4)|r2_hidden(esk8_3(esk17_2(X2,X3),X4,X1),k9_xtuple_0(esk17_2(X2,X3)))|r2_hidden(esk7_3(esk17_2(X2,X3),X4,X1),X1)|v2_struct_0(X3)|~v1_relat_1(esk17_2(X2,X3))|~v2_pre_topc(X3)|~v2_pre_topc(X2)|~l1_pre_topc(X3)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_8251, plain, (k1_relset_1(k6_numbers,k6_numbers)=k6_numbers)).#trainneg
cnf(i_0_8170, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,esk3_0),esk2_0)|~r2_hidden(X1,X2)|~v4_pre_topc(esk3_0,k3_pcomps_1(esk1_0)))).#trainneg
cnf(i_0_7988, negated_conjecture, (~m1_subset_1(esk3_0,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_8158, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,esk2_0),X3)|~v1_funct_1(k6_weierstr(esk1_0,esk2_0))|~v1_relat_1(k6_weierstr(esk1_0,esk2_0))|~r2_hidden(X1,X2)|~r2_hidden(esk6_4(k6_weierstr(esk1_0,esk2_0),X3,X2,X1),esk2_0))).#trainneg
cnf(i_0_8160, negated_conjecture, (k6_numbers=X1|X2!=k7_relat_1(k6_weierstr(esk1_0,esk2_0),esk3_0)|~v1_funct_1(k6_weierstr(esk1_0,esk2_0))|~v1_relat_1(k6_weierstr(esk1_0,esk2_0))|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_7317, plain, (X1=k6_numbers|v1_xboole_0(X2)|v2_struct_0(X3)|v1_xboole_0(k7_relat_1(k6_weierstr(X3,X2),X4))|k1_hausdorf(X3,X2,X4)!=k6_numbers|~v1_xreal_0(X1)|~v2_compts_1(X4,k3_pcomps_1(X3))|~v2_compts_1(X2,k3_pcomps_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)|~m1_subset_1(X1,k7_relat_1(k6_weierstr(X3,X2),X4)))).#trainneg
cnf(i_0_7959, plain, (X1=k6_numbers|v1_funct_2(k6_weierstr(X2,X3),u1_struct_0(k3_pcomps_1(X2)),X1)|v2_struct_0(X2)|~m1_subset_1(k6_weierstr(X2,X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_7784, negated_conjecture, (r1_tarski(k9_xtuple_0(X1),esk3_0)|~v4_relat_1(X1,esk2_0)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7962, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relat_1(X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k6_numbers))|~v1_xboole_0(X5))).#trainneg
cnf(i_0_7978, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_relat_1(k7_relat_1(k2_zfmisc_1(X3,k2_zfmisc_1(X4,X2)),X5),X6)))).#trainneg
cnf(i_0_7987, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_621, plain, (X1=k7_relat_1(k6_weierstr(X2,X3),X4)|r2_hidden(esk8_3(k6_weierstr(X2,X3),X4,X1),k9_xtuple_0(k6_weierstr(X2,X3)))|r2_hidden(esk7_3(k6_weierstr(X2,X3),X4,X1),X1)|v2_struct_0(X2)|~v1_relat_1(k6_weierstr(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_4300, plain, (v1_relat_1(k1_relset_1(k6_numbers,X3))|~v4_relat_1(X3,k6_numbers)|~v1_relat_1(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6541, plain, (v1_xxreal_0(u1_pre_topc(k3_pcomps_1(esk18_0)))|~v2_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers)))|~l1_metric_1(esk18_0))).#trainneg
cnf(i_0_6549, plain, (v1_xreal_0(u1_pre_topc(k3_pcomps_1(esk18_0)))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(k6_numbers)))|~l1_metric_1(esk18_0))).#trainneg
cnf(i_0_7992, plain, (v1_funct_1(k6_numbers)|v2_struct_0(X2)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|u1_struct_0(X2)!=k6_numbers)).#trainneg
cnf(i_0_7995, plain, (v5_pre_topc(k6_numbers,X1,X2)|v2_struct_0(X2)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|u1_struct_0(X2)!=k6_numbers)).#trainneg
cnf(i_0_622, plain, (v1_xxreal_0(esk17_2(X1,X2))|v2_struct_0(X2)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_623, plain, (v1_xreal_0(esk17_2(X1,X2))|v2_struct_0(X2)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 160
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 259
# Removed in clause preprocessing      : 17
# Initial clauses in saturation        : 242
# Processed clauses                    : 1901
# ...of these trivial                  : 34
# ...subsumed                          : 554
# ...remaining for further processing  : 1313
# Other redundant clauses eliminated   : 12
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 121
# Backward-rewritten                   : 29
# Generated clauses                    : 9664
# ...of the previous two non-trivial   : 8376
# Contextual simplify-reflections      : 270
# Paramodulations                      : 9574
# Factorizations                       : 3
# Equation resolutions                 : 86
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1160
#    Positive orientable unit clauses  : 106
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 22
#    Non-unit-clauses                  : 1032
# Current number of unprocessed clauses: 6226
# ...number of literals in the above   : 50414
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1464
# Proof object given clauses           : 189
# Proof search given clauses           : 1313
# Clause-clause subsumption calls (NU) : 639059
# Rec. Clause-clause subsumption calls : 35950
# Non-unit clause-clause subsumptions  : 818
# Unit Clause-clause subsumption calls : 9800
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 60
# BW rewrite match successes           : 18
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 286914

# -------------------------------------------------
# User time                : 24.207 s
# System time              : 1.005 s
# Total time               : 25.212 s
# Maximum resident set size: 375796 pages
24.20user 1.11system 0:25.92elapsed 97%CPU (0avgtext+0avgdata 376268maxresident)k
0inputs+768outputs (0major+217110minor)pagefaults 0swaps
