2020-05-31 04:12:25.832714: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_15
2020-05-31 04:12:25.972054: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-05-31 04:12:26.358666: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-05-31 04:12:27.742069: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-05-31 04:12:32.794813: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 6962130 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_15' loaded (query=256; context=512; weigths=real; len_mult=0.000000).
# Preprocessing time       : 8.670 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t87_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', t87_zfmisc_1)).
fof(d5_tarski, axiom, ![X1, X2]:k4_tarski(X1,X2)=k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', d5_tarski)).
fof(d5_matrix_1, axiom, ![X1, X2]:((v1_matrix_1(X2)&m2_finseq_1(X2,k3_finseq_2(X1)))=>![X3]:(v7_ordinal1(X3)=>![X4]:(v7_ordinal1(X4)=>(r2_hidden(k4_tarski(X3,X4),k2_matrix_1(X2))=>![X5]:(m1_subset_1(X5,X1)=>(X5=k3_matrix_1(X1,X2,X3,X4)<=>?[X6]:((m2_finseq_1(X6,X1)&X6=k1_funct_1(X2,X3))&X5=k1_funct_1(X6,X4)))))))), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', d5_matrix_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', commutativity_k2_tarski)).
fof(t37_matrixr2, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_matrix_1(X2,X1,np__3,np__3)=>X2=k11_finseq_1(k3_finseq_4(X1,k3_matrix_1(X1,X2,np__1,np__1),k3_matrix_1(X1,X2,np__1,np__2),k3_matrix_1(X1,X2,np__1,np__3)),k3_finseq_4(X1,k3_matrix_1(X1,X2,np__2,np__1),k3_matrix_1(X1,X2,np__2,np__2),k3_matrix_1(X1,X2,np__2,np__3)),k3_finseq_4(X1,k3_matrix_1(X1,X2,np__3,np__1),k3_matrix_1(X1,X2,np__3,np__2),k3_matrix_1(X1,X2,np__3,np__3))))), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', t37_matrixr2)).
fof(d4_matrix_1, axiom, ![X1]:((((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))&v1_matrix_1(X1))=>k2_matrix_1(X1)=k2_zfmisc_1(k4_finseq_1(X1),k2_finseq_1(k1_matrix_1(X1)))), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', d4_matrix_1)).
fof(t24_matrix_1, axiom, ![X1]:(v7_ordinal1(X1)=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_matrix_1(X3,X2,X1,X1)=>((k3_finseq_1(X3)=X1&k1_matrix_1(X3)=X1)&k2_matrix_1(X3)=k2_zfmisc_1(k2_finseq_1(X1),k2_finseq_1(X1)))))), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', t24_matrix_1)).
fof(cc17_membered, axiom, ![X1]:(v6_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v7_ordinal1(X2))), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', cc17_membered)).
fof(fc6_membered, axiom, v6_membered(k4_ordinal1), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', fc6_membered)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', redefinition_k5_numbers)).
fof(dt_m1_finseq_1, axiom, ![X1, X2]:(m1_finseq_1(X2,X1)=>((v1_relat_1(X2)&v1_funct_1(X2))&v1_finseq_1(X2))), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', dt_m1_finseq_1)).
fof(redefinition_m2_finseq_1, axiom, ![X1, X2]:(m2_finseq_1(X2,X1)<=>m1_finseq_1(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', redefinition_m2_finseq_1)).
fof(fraenkel_a_1_0_finseq_1, axiom, ![X1, X2]:(v7_ordinal1(X2)=>(r2_hidden(X1,a_1_0_finseq_1(X2))<=>?[X3]:(((m2_subset_1(X3,k1_numbers,k5_numbers)&X1=X3)&r1_xxreal_0(np__1,X3))&r1_xxreal_0(X3,X2)))), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', fraenkel_a_1_0_finseq_1)).
fof(redefinition_k2_finseq_1, axiom, ![X1]:(v7_ordinal1(X1)=>k2_finseq_1(X1)=k1_finseq_1(X1)), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', redefinition_k2_finseq_1)).
fof(d1_finseq_1, axiom, ![X1]:(v7_ordinal1(X1)=>k1_finseq_1(X1)=a_1_0_finseq_1(X1)), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', d1_finseq_1)).
fof(reflexivity_r1_xxreal_0, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>r1_xxreal_0(X1,X1)), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', reflexivity_r1_xxreal_0)).
fof(cc2_xxreal_0, axiom, ![X1]:(v7_ordinal1(X1)=>v1_xxreal_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', cc2_xxreal_0)).
fof(spc3_numerals, axiom, (((v2_xxreal_0(np__3)&m2_subset_1(np__3,k1_numbers,k5_numbers))&m1_subset_1(np__3,k5_numbers))&m1_subset_1(np__3,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', spc3_numerals)).
fof(dt_m1_matrix_1, axiom, ![X1, X2, X3]:(((~(v1_xboole_0(X1))&v7_ordinal1(X2))&v7_ordinal1(X3))=>![X4]:(m1_matrix_1(X4,X1,X2,X3)=>(v1_matrix_1(X4)&m2_finseq_1(X4,k3_finseq_2(X1))))), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', dt_m1_matrix_1)).
fof(dt_m2_finseq_1, axiom, ![X1, X2]:(m2_finseq_1(X2,X1)=>((v1_funct_1(X2)&v1_finseq_1(X2))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', dt_m2_finseq_1)).
fof(t1_numerals, axiom, m1_subset_1(k1_xboole_0,k4_ordinal1), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', t1_numerals)).
fof(d1_enumset1, axiom, ![X1, X2, X3, X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~(((X5!=X1&X5!=X2)&X5!=X3)))), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', d1_enumset1)).
fof(t1_finseq_3, axiom, k2_finseq_1(np__3)=k1_enumset1(np__1,np__2,np__3), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', t1_finseq_3)).
fof(t7_boole, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', t7_boole)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', t3_subset)).
fof(reflexivity_r1_tarski, axiom, ![X1, X2]:r1_tarski(X1,X1), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', reflexivity_r1_tarski)).
fof(dt_k3_matrix_1, axiom, ![X1, X2, X3, X4]:((((v1_matrix_1(X2)&m1_finseq_1(X2,k3_finseq_2(X1)))&v7_ordinal1(X3))&v7_ordinal1(X4))=>m1_subset_1(k3_matrix_1(X1,X2,X3,X4),X1)), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', dt_k3_matrix_1)).
fof(cc3_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_xboole_0(X3))), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', cc3_relset_1)).
fof(rqLessOrEqual__r1_xxreal_0__r1_r3, axiom, r1_xxreal_0(np__1,np__3), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', rqLessOrEqual__r1_xxreal_0__r1_r3)).
fof(redefinition_k3_finseq_4, axiom, ![X1, X2, X3, X4]:((((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))&m1_subset_1(X3,X1))&m1_subset_1(X4,X1))=>k3_finseq_4(X1,X2,X3,X4)=k11_finseq_1(X2,X3,X4)), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', redefinition_k3_finseq_4)).
fof(t45_finseq_1, axiom, ![X1, X2, X3, X4]:(((v1_relat_1(X4)&v1_funct_1(X4))&v1_finseq_1(X4))=>(X4=k11_finseq_1(X1,X2,X3)<=>(((k3_finseq_1(X4)=np__3&k1_funct_1(X4,np__1)=X1)&k1_funct_1(X4,np__2)=X2)&k1_funct_1(X4,np__3)=X3))), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', t45_finseq_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', t2_subset)).
fof(spc2_numerals, axiom, (((v2_xxreal_0(np__2)&m2_subset_1(np__2,k1_numbers,k5_numbers))&m1_subset_1(np__2,k5_numbers))&m1_subset_1(np__2,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', spc2_numerals)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', spc1_numerals)).
fof(t36_matrixr2, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,k5_numbers)=>![X3]:(m1_matrix_1(X3,X1,X2,X2)=>![X4]:(m2_finseq_1(X4,X1)=>![X5]:(v7_ordinal1(X5)=>((X4=k1_funct_1(X3,X5)&r2_hidden(X5,k2_finseq_1(X2)))=>k3_finseq_1(X4)=X2)))))), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', t36_matrixr2)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', t1_subset)).
fof(rqLessOrEqual__r1_xxreal_0__r1_r1, axiom, r1_xxreal_0(np__1,np__1), file('/scratch/mptp/allpred/preds__0.005/t37_matrixr2', rqLessOrEqual__r1_xxreal_0__r1_r1)).
fof(c_0_37, plain, ![X27, X28, X29, X30]:(((r2_hidden(X27,X29)|~r2_hidden(k4_tarski(X27,X28),k2_zfmisc_1(X29,X30)))&(r2_hidden(X28,X30)|~r2_hidden(k4_tarski(X27,X28),k2_zfmisc_1(X29,X30))))&(~r2_hidden(X27,X29)|~r2_hidden(X28,X30)|r2_hidden(k4_tarski(X27,X28),k2_zfmisc_1(X29,X30)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t87_zfmisc_1])])])).
fof(c_0_38, plain, ![X130, X131]:k4_tarski(X130,X131)=k2_tarski(k2_tarski(X130,X131),k1_tarski(X130)), inference(variable_rename,[status(thm)],[d5_tarski])).
fof(c_0_39, plain, ![X54, X55, X56, X57, X58, X60]:((((m2_finseq_1(esk4_5(X54,X55,X56,X57,X58),X54)|X58!=k3_matrix_1(X54,X55,X56,X57)|~m1_subset_1(X58,X54)|~r2_hidden(k4_tarski(X56,X57),k2_matrix_1(X55))|~v7_ordinal1(X57)|~v7_ordinal1(X56)|(~v1_matrix_1(X55)|~m2_finseq_1(X55,k3_finseq_2(X54))))&(esk4_5(X54,X55,X56,X57,X58)=k1_funct_1(X55,X56)|X58!=k3_matrix_1(X54,X55,X56,X57)|~m1_subset_1(X58,X54)|~r2_hidden(k4_tarski(X56,X57),k2_matrix_1(X55))|~v7_ordinal1(X57)|~v7_ordinal1(X56)|(~v1_matrix_1(X55)|~m2_finseq_1(X55,k3_finseq_2(X54)))))&(X58=k1_funct_1(esk4_5(X54,X55,X56,X57,X58),X57)|X58!=k3_matrix_1(X54,X55,X56,X57)|~m1_subset_1(X58,X54)|~r2_hidden(k4_tarski(X56,X57),k2_matrix_1(X55))|~v7_ordinal1(X57)|~v7_ordinal1(X56)|(~v1_matrix_1(X55)|~m2_finseq_1(X55,k3_finseq_2(X54)))))&(~m2_finseq_1(X60,X54)|X60!=k1_funct_1(X55,X56)|X58!=k1_funct_1(X60,X57)|X58=k3_matrix_1(X54,X55,X56,X57)|~m1_subset_1(X58,X54)|~r2_hidden(k4_tarski(X56,X57),k2_matrix_1(X55))|~v7_ordinal1(X57)|~v7_ordinal1(X56)|(~v1_matrix_1(X55)|~m2_finseq_1(X55,k3_finseq_2(X54))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_matrix_1])])])])])).
cnf(c_0_40, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_41, plain, (k4_tarski(X1,X2)=k2_tarski(k2_tarski(X1,X2),k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_38])).
fof(c_0_42, plain, ![X143, X144]:k2_tarski(X143,X144)=k2_tarski(X144,X143), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
fof(c_0_43, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_matrix_1(X2,X1,np__3,np__3)=>X2=k11_finseq_1(k3_finseq_4(X1,k3_matrix_1(X1,X2,np__1,np__1),k3_matrix_1(X1,X2,np__1,np__2),k3_matrix_1(X1,X2,np__1,np__3)),k3_finseq_4(X1,k3_matrix_1(X1,X2,np__2,np__1),k3_matrix_1(X1,X2,np__2,np__2),k3_matrix_1(X1,X2,np__2,np__3)),k3_finseq_4(X1,k3_matrix_1(X1,X2,np__3,np__1),k3_matrix_1(X1,X2,np__3,np__2),k3_matrix_1(X1,X2,np__3,np__3)))))), inference(assume_negation,[status(cth)],[t37_matrixr2])).
cnf(c_0_44, plain, (esk4_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|X5!=k3_matrix_1(X1,X2,X3,X4)|~m1_subset_1(X5,X1)|~r2_hidden(k4_tarski(X3,X4),k2_matrix_1(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_45, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),k2_zfmisc_1(X2,X4))|~r2_hidden(X3,X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_46, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_47, plain, ![X69]:(~v1_relat_1(X69)|~v1_funct_1(X69)|~v1_finseq_1(X69)|~v1_matrix_1(X69)|k2_matrix_1(X69)=k2_zfmisc_1(k4_finseq_1(X69),k2_finseq_1(k1_matrix_1(X69)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_matrix_1])])).
fof(c_0_48, plain, ![X24, X25, X26]:(((k3_finseq_1(X26)=X24|~m1_matrix_1(X26,X25,X24,X24)|v1_xboole_0(X25)|~v7_ordinal1(X24))&(k1_matrix_1(X26)=X24|~m1_matrix_1(X26,X25,X24,X24)|v1_xboole_0(X25)|~v7_ordinal1(X24)))&(k2_matrix_1(X26)=k2_zfmisc_1(k2_finseq_1(X24),k2_finseq_1(X24))|~m1_matrix_1(X26,X25,X24,X24)|v1_xboole_0(X25)|~v7_ordinal1(X24))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_matrix_1])])])])])).
fof(c_0_49, negated_conjecture, (~v1_xboole_0(esk1_0)&(m1_matrix_1(esk2_0,esk1_0,np__3,np__3)&esk2_0!=k11_finseq_1(k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_43])])])])).
fof(c_0_50, plain, ![X86, X87]:(~v6_membered(X86)|(~m1_subset_1(X87,X86)|v7_ordinal1(X87))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc17_membered])])])).
cnf(c_0_51, plain, (v6_membered(k4_ordinal1)), inference(split_conjunct,[status(thm)],[fc6_membered])).
cnf(c_0_52, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
cnf(c_0_53, plain, (X1=k1_funct_1(esk4_5(X2,X3,X4,X5,X1),X5)|X1!=k3_matrix_1(X2,X3,X4,X5)|~m1_subset_1(X1,X2)|~r2_hidden(k4_tarski(X4,X5),k2_matrix_1(X3))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X2))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_54, plain, (esk4_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m1_subset_1(X5,X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~r2_hidden(k2_tarski(k2_tarski(X3,X4),k1_tarski(X3)),k2_matrix_1(X2))), inference(rw,[status(thm)],[c_0_44, c_0_41])).
cnf(c_0_55, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),k2_zfmisc_1(X2,X4))|~r2_hidden(X3,X4)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_45])).
cnf(c_0_56, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(evalgc,[status(thm)],[c_0_46])).
cnf(c_0_57, plain, (k2_matrix_1(X1)=k2_zfmisc_1(k4_finseq_1(X1),k2_finseq_1(k1_matrix_1(X1)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)|~v1_matrix_1(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_58, plain, ![X93, X94]:(((v1_relat_1(X94)|~m1_finseq_1(X94,X93))&(v1_funct_1(X94)|~m1_finseq_1(X94,X93)))&(v1_finseq_1(X94)|~m1_finseq_1(X94,X93))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_finseq_1])])])).
fof(c_0_59, plain, ![X76, X77]:((~m2_finseq_1(X77,X76)|m1_finseq_1(X77,X76))&(~m1_finseq_1(X77,X76)|m2_finseq_1(X77,X76))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_m2_finseq_1])])).
fof(c_0_60, plain, ![X159, X160, X162]:(((((m2_subset_1(esk8_2(X159,X160),k1_numbers,k5_numbers)|~r2_hidden(X159,a_1_0_finseq_1(X160))|~v7_ordinal1(X160))&(X159=esk8_2(X159,X160)|~r2_hidden(X159,a_1_0_finseq_1(X160))|~v7_ordinal1(X160)))&(r1_xxreal_0(np__1,esk8_2(X159,X160))|~r2_hidden(X159,a_1_0_finseq_1(X160))|~v7_ordinal1(X160)))&(r1_xxreal_0(esk8_2(X159,X160),X160)|~r2_hidden(X159,a_1_0_finseq_1(X160))|~v7_ordinal1(X160)))&(~m2_subset_1(X162,k1_numbers,k5_numbers)|X159!=X162|~r1_xxreal_0(np__1,X162)|~r1_xxreal_0(X162,X160)|r2_hidden(X159,a_1_0_finseq_1(X160))|~v7_ordinal1(X160))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fraenkel_a_1_0_finseq_1])])])])])).
fof(c_0_61, plain, ![X193]:(~v7_ordinal1(X193)|k2_finseq_1(X193)=k1_finseq_1(X193)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_finseq_1])])).
fof(c_0_62, plain, ![X213]:(~v7_ordinal1(X213)|k1_finseq_1(X213)=a_1_0_finseq_1(X213)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_finseq_1])])).
fof(c_0_63, plain, ![X127, X128]:(~v1_xxreal_0(X127)|~v1_xxreal_0(X128)|r1_xxreal_0(X127,X127)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[reflexivity_r1_xxreal_0])])).
fof(c_0_64, plain, ![X190]:(~v7_ordinal1(X190)|v1_xxreal_0(X190)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_xxreal_0])])).
cnf(c_0_65, plain, (k1_matrix_1(X1)=X2|v1_xboole_0(X3)|~m1_matrix_1(X1,X3,X2,X2)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_66, negated_conjecture, (m1_matrix_1(esk2_0,esk1_0,np__3,np__3)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_67, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_68, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_69, plain, (m1_subset_1(np__3,k5_numbers)), inference(split_conjunct,[status(thm)],[spc3_numerals])).
cnf(c_0_70, plain, (v6_membered(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_51])).
cnf(c_0_71, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_52])).
fof(c_0_72, plain, ![X61, X62, X63, X64]:((v1_matrix_1(X64)|~m1_matrix_1(X64,X61,X62,X63)|(v1_xboole_0(X61)|~v7_ordinal1(X62)|~v7_ordinal1(X63)))&(m2_finseq_1(X64,k3_finseq_2(X61))|~m1_matrix_1(X64,X61,X62,X63)|(v1_xboole_0(X61)|~v7_ordinal1(X62)|~v7_ordinal1(X63)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m1_matrix_1])])])])])).
cnf(c_0_73, plain, (X1=k1_funct_1(esk4_5(X2,X3,X4,X5,X1),X5)|X1!=k3_matrix_1(X2,X3,X4,X5)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v1_matrix_1(X3)|~m1_subset_1(X1,X2)|~m2_finseq_1(X3,k3_finseq_2(X2))|~r2_hidden(k2_tarski(k2_tarski(X4,X5),k1_tarski(X4)),k2_matrix_1(X3))), inference(rw,[status(thm)],[c_0_53, c_0_41])).
cnf(c_0_74, plain, (esk4_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m1_subset_1(X5,X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~r2_hidden(k2_tarski(k2_tarski(X3,X4),k1_tarski(X3)),k2_matrix_1(X2))), inference(evalgc,[status(thm)],[c_0_54])).
cnf(c_0_75, plain, (r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X1)),k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56])])).
cnf(c_0_76, plain, (k2_zfmisc_1(k4_finseq_1(X1),k2_finseq_1(k1_matrix_1(X1)))=k2_matrix_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)|~v1_matrix_1(X1)), inference(evalgc,[status(thm)],[c_0_57])).
cnf(c_0_77, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_78, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
fof(c_0_79, plain, ![X70, X71]:(((v1_funct_1(X71)|~m2_finseq_1(X71,X70))&(v1_finseq_1(X71)|~m2_finseq_1(X71,X70)))&(m1_subset_1(X71,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X70)))|~m2_finseq_1(X71,X70))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m2_finseq_1])])])).
cnf(c_0_80, plain, (r2_hidden(X2,a_1_0_finseq_1(X3))|~m2_subset_1(X1,k1_numbers,k5_numbers)|X2!=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X3)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_81, plain, (k2_finseq_1(X1)=k1_finseq_1(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_82, plain, (k1_finseq_1(X1)=a_1_0_finseq_1(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_83, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_84, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_85, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1)), inference(split_conjunct,[status(thm)],[t1_numerals])).
cnf(c_0_86, plain, (k1_matrix_1(X1)=X2|v1_xboole_0(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X1,X3,X2,X2)), inference(evalgc,[status(thm)],[c_0_65])).
cnf(c_0_87, negated_conjecture, (m1_matrix_1(esk2_0,esk1_0,np__3,np__3)), inference(evalgc,[status(thm)],[c_0_66])).
cnf(c_0_88, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_67])).
cnf(c_0_89, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_68])).
cnf(c_0_90, plain, (m1_subset_1(np__3,k5_numbers)), inference(evalgc,[status(thm)],[c_0_69])).
cnf(c_0_91, plain, (v6_membered(k5_numbers)), inference(rw,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_92, plain, (v1_matrix_1(X1)|v1_xboole_0(X2)|~m1_matrix_1(X1,X2,X3,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X4)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_93, plain, (m2_finseq_1(X1,k3_finseq_2(X2))|v1_xboole_0(X2)|~m1_matrix_1(X1,X2,X3,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X4)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_94, plain, (k1_funct_1(esk4_5(X2,X3,X4,X5,X1),X5)=X1|X1!=k3_matrix_1(X2,X3,X4,X5)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v1_matrix_1(X3)|~m1_subset_1(X1,X2)|~m2_finseq_1(X3,k3_finseq_2(X2))|~r2_hidden(k2_tarski(k2_tarski(X4,X5),k1_tarski(X4)),k2_matrix_1(X3))), inference(evalgc,[status(thm)],[c_0_73])).
cnf(c_0_95, plain, (esk4_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(k2_tarski(k2_tarski(X4,X3),k1_tarski(X3)),k2_matrix_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_56])])).
cnf(c_0_96, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X2)),k2_matrix_1(X3))|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(X3)))|~r2_hidden(X2,k4_finseq_1(X3))|~v1_matrix_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76])])).
cnf(c_0_97, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_77])).
cnf(c_0_98, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_78])).
cnf(c_0_99, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_100, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_101, plain, (r2_hidden(X2,a_1_0_finseq_1(X3))|X1!=X2|~v7_ordinal1(X3)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_80])).
cnf(c_0_102, plain, (k1_finseq_1(X1)=k2_finseq_1(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_81])).
cnf(c_0_103, plain, (k1_finseq_1(X1)=a_1_0_finseq_1(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_82])).
cnf(c_0_104, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_83])).
cnf(c_0_105, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_84])).
cnf(c_0_106, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_85])).
fof(c_0_107, plain, ![X13, X14, X15, X16, X17, X18, X19, X20, X21, X22]:(((~r2_hidden(X17,X16)|(X17=X13|X17=X14|X17=X15)|X16!=k1_enumset1(X13,X14,X15))&(((X18!=X13|r2_hidden(X18,X16)|X16!=k1_enumset1(X13,X14,X15))&(X18!=X14|r2_hidden(X18,X16)|X16!=k1_enumset1(X13,X14,X15)))&(X18!=X15|r2_hidden(X18,X16)|X16!=k1_enumset1(X13,X14,X15))))&((((esk3_4(X19,X20,X21,X22)!=X19|~r2_hidden(esk3_4(X19,X20,X21,X22),X22)|X22=k1_enumset1(X19,X20,X21))&(esk3_4(X19,X20,X21,X22)!=X20|~r2_hidden(esk3_4(X19,X20,X21,X22),X22)|X22=k1_enumset1(X19,X20,X21)))&(esk3_4(X19,X20,X21,X22)!=X21|~r2_hidden(esk3_4(X19,X20,X21,X22),X22)|X22=k1_enumset1(X19,X20,X21)))&(r2_hidden(esk3_4(X19,X20,X21,X22),X22)|(esk3_4(X19,X20,X21,X22)=X19|esk3_4(X19,X20,X21,X22)=X20|esk3_4(X19,X20,X21,X22)=X21)|X22=k1_enumset1(X19,X20,X21)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_enumset1])])])])])])).
cnf(c_0_108, negated_conjecture, (k1_matrix_1(esk2_0)=np__3|~v7_ordinal1(np__3)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_88])])).
cnf(c_0_109, plain, (v7_ordinal1(np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_91])])])).
cnf(c_0_110, plain, (k2_finseq_1(np__3)=k1_enumset1(np__1,np__2,np__3)), inference(split_conjunct,[status(thm)],[t1_finseq_3])).
cnf(c_0_111, plain, (v1_xboole_0(X2)|v1_matrix_1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4)), inference(evalgc,[status(thm)],[c_0_92])).
cnf(c_0_112, plain, (v1_xboole_0(X2)|m2_finseq_1(X1,k3_finseq_2(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4)), inference(evalgc,[status(thm)],[c_0_93])).
fof(c_0_113, plain, ![X166, X167]:(~r2_hidden(X166,X167)|~v1_xboole_0(X167)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
cnf(c_0_114, plain, (k2_matrix_1(X1)=k2_zfmisc_1(k2_finseq_1(X2),k2_finseq_1(X2))|v1_xboole_0(X3)|~m1_matrix_1(X1,X3,X2,X2)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_115, plain, (k1_funct_1(esk4_5(X1,X2,X3,X4,X5),X4)=X5|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(k2_tarski(k2_tarski(X4,X3),k1_tarski(X3)),k2_matrix_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_56])])).
cnf(c_0_116, plain, (esk4_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X4,k2_finseq_1(k1_matrix_1(X2)))|~r2_hidden(X3,k4_finseq_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96])])).
cnf(c_0_117, plain, (v1_relat_1(X1)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98])])).
cnf(c_0_118, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_99])).
cnf(c_0_119, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_100])).
cnf(c_0_120, plain, (r2_hidden(X1,a_1_0_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_101])])).
cnf(c_0_121, plain, (a_1_0_finseq_1(X1)=k2_finseq_1(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103])])).
cnf(c_0_122, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105])])).
cnf(c_0_123, plain, (m1_subset_1(k1_xboole_0,k5_numbers)), inference(rw,[status(thm)],[c_0_106, c_0_71])).
cnf(c_0_124, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X5,X2)), inference(split_conjunct,[status(thm)],[c_0_107])).
fof(c_0_125, plain, ![X171, X172]:((~m1_subset_1(X171,k1_zfmisc_1(X172))|r1_tarski(X171,X172))&(~r1_tarski(X171,X172)|m1_subset_1(X171,k1_zfmisc_1(X172)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_126, plain, ![X216]:r1_tarski(X216,X216), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r1_tarski])])).
cnf(c_0_127, negated_conjecture, (k1_matrix_1(esk2_0)=np__3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_108, c_0_109])])).
cnf(c_0_128, plain, (k2_finseq_1(np__3)=k1_enumset1(np__1,np__2,np__3)), inference(evalgc,[status(thm)],[c_0_110])).
cnf(c_0_129, negated_conjecture, (v1_matrix_1(esk2_0)|~v7_ordinal1(np__3)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_87]), c_0_88])])).
cnf(c_0_130, negated_conjecture, (m2_finseq_1(esk2_0,k3_finseq_2(esk1_0))|~v7_ordinal1(np__3)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_87]), c_0_88])])).
cnf(c_0_131, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_132, plain, (k2_matrix_1(X1)=k2_zfmisc_1(k2_finseq_1(X2),k2_finseq_1(X2))|v1_xboole_0(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X1,X3,X2,X2)), inference(evalgc,[status(thm)],[c_0_114])).
cnf(c_0_133, plain, (k1_funct_1(esk4_5(X1,X2,X3,X4,X5),X4)=X5|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X4,k2_finseq_1(k1_matrix_1(X2)))|~r2_hidden(X3,k4_finseq_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_96])])).
cnf(c_0_134, plain, (esk4_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X4,k2_finseq_1(k1_matrix_1(X2)))|~r2_hidden(X3,k4_finseq_1(X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_116, c_0_117]), c_0_118]), c_0_119])).
cnf(c_0_135, plain, (r2_hidden(X1,k2_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121])])).
cnf(c_0_136, plain, (r1_xxreal_0(X1,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_105])])).
cnf(c_0_137, plain, (v7_ordinal1(k1_xboole_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_123]), c_0_91])])])).
cnf(c_0_138, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X5,X2)), inference(evalgc,[status(thm)],[c_0_124])).
fof(c_0_139, plain, ![X78, X79, X80, X81]:(~v1_matrix_1(X79)|~m1_finseq_1(X79,k3_finseq_2(X78))|~v7_ordinal1(X80)|~v7_ordinal1(X81)|m1_subset_1(k3_matrix_1(X78,X79,X80,X81),X78)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_matrix_1])])).
fof(c_0_140, plain, ![X185, X186, X187]:(~v1_xboole_0(X185)|(~m1_subset_1(X187,k1_zfmisc_1(k2_zfmisc_1(X185,X186)))|v1_xboole_0(X187))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_relset_1])])])).
cnf(c_0_141, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_125])).
cnf(c_0_142, plain, (r1_tarski(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_126])).
cnf(c_0_143, negated_conjecture, (k2_zfmisc_1(k4_finseq_1(esk2_0),k1_enumset1(np__1,np__2,np__3))=k2_matrix_1(esk2_0)|~v1_matrix_1(esk2_0)|~v1_finseq_1(esk2_0)|~v1_funct_1(esk2_0)|~v1_relat_1(esk2_0)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_127]), c_0_128])])).
cnf(c_0_144, negated_conjecture, (v1_matrix_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_129, c_0_109])])).
cnf(c_0_145, negated_conjecture, (m2_finseq_1(esk2_0,k3_finseq_2(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_130, c_0_109])])).
cnf(c_0_146, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_131])).
cnf(c_0_147, negated_conjecture, (k2_zfmisc_1(k1_enumset1(np__1,np__2,np__3),k1_enumset1(np__1,np__2,np__3))=k2_matrix_1(esk2_0)|~v7_ordinal1(np__3)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_87]), c_0_128]), c_0_128]), c_0_88])])).
cnf(c_0_148, plain, (k1_funct_1(esk4_5(X1,X2,X3,X4,X5),X4)=X5|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X4,k2_finseq_1(k1_matrix_1(X2)))|~r2_hidden(X3,k4_finseq_1(X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_133, c_0_117]), c_0_118]), c_0_119])).
cnf(c_0_149, plain, (esk4_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,k1_matrix_1(X2))|~v7_ordinal1(k1_matrix_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135])])).
cnf(c_0_150, plain, (r1_xxreal_0(X1,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_137])])).
cnf(c_0_151, plain, (r2_hidden(X1,X2)|X2!=k1_enumset1(X3,X4,X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_138])])).
cnf(c_0_152, plain, (m2_subset_1(np__3,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[spc3_numerals])).
cnf(c_0_153, plain, (r1_xxreal_0(np__1,np__3)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r1_r3])).
cnf(c_0_154, plain, (m1_subset_1(k3_matrix_1(X2,X1,X3,X4),X2)|~v1_matrix_1(X1)|~m1_finseq_1(X1,k3_finseq_2(X2))|~v7_ordinal1(X3)|~v7_ordinal1(X4)), inference(split_conjunct,[status(thm)],[c_0_139])).
cnf(c_0_155, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(split_conjunct,[status(thm)],[c_0_140])).
cnf(c_0_156, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_141])).
cnf(c_0_157, plain, (r1_tarski(X1,X1)), inference(evalgc,[status(thm)],[c_0_142])).
cnf(c_0_158, negated_conjecture, (k2_zfmisc_1(k4_finseq_1(esk2_0),k1_enumset1(np__1,np__2,np__3))=k2_matrix_1(esk2_0)|~v1_finseq_1(esk2_0)|~v1_funct_1(esk2_0)|~v1_relat_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_143, c_0_144])])).
cnf(c_0_159, negated_conjecture, (v1_funct_1(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_145])])).
cnf(c_0_160, plain, (~v1_xboole_0(k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_55])])).
cnf(c_0_161, negated_conjecture, (k2_zfmisc_1(k1_enumset1(np__1,np__2,np__3),k1_enumset1(np__1,np__2,np__3))=k2_matrix_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_147, c_0_109])])).
cnf(c_0_162, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_163, plain, ![X50, X51, X52, X53]:(v1_xboole_0(X50)|~m1_subset_1(X51,X50)|~m1_subset_1(X52,X50)|~m1_subset_1(X53,X50)|k3_finseq_4(X50,X51,X52,X53)=k11_finseq_1(X51,X52,X53)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_finseq_4])])])).
cnf(c_0_164, negated_conjecture, (k1_funct_1(esk4_5(X1,esk2_0,X2,X3,X4),X3)=X4|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k4_finseq_1(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_127]), c_0_144]), c_0_128])])])).
cnf(c_0_165, plain, (esk4_5(X1,X2,X3,k1_matrix_1(X2),X4)=k1_funct_1(X2,X3)|X4!=k3_matrix_1(X1,X2,X3,k1_matrix_1(X2))|~m2_subset_1(k1_matrix_1(X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(X2))|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(k1_matrix_1(X2))|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_150])])).
cnf(c_0_166, plain, (r2_hidden(X1,k1_enumset1(X2,X3,X1))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_151])])).
cnf(c_0_167, plain, (m2_subset_1(np__3,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_152])).
cnf(c_0_168, plain, (r1_xxreal_0(np__1,np__3)), inference(evalgc,[status(thm)],[c_0_153])).
cnf(c_0_169, plain, (m1_subset_1(k3_matrix_1(X2,X1,X3,X4),X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X1)|~m1_finseq_1(X1,k3_finseq_2(X2))), inference(evalgc,[status(thm)],[c_0_154])).
cnf(c_0_170, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(evalgc,[status(thm)],[c_0_155])).
cnf(c_0_171, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_157])])).
cnf(c_0_172, negated_conjecture, (k2_zfmisc_1(k4_finseq_1(esk2_0),k1_enumset1(np__1,np__2,np__3))=k2_matrix_1(esk2_0)|~v1_finseq_1(esk2_0)|~v1_relat_1(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_158, c_0_159])])])).
cnf(c_0_173, negated_conjecture, (v1_finseq_1(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_145])])).
cnf(c_0_174, negated_conjecture, (v1_relat_1(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_145])])).
cnf(c_0_175, negated_conjecture, (~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~v1_xboole_0(k2_matrix_1(esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_161])])).
cnf(c_0_176, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),k2_zfmisc_1(X2,X4))), inference(rw,[status(thm)],[c_0_162, c_0_41])).
cnf(c_0_177, negated_conjecture, (esk2_0!=k11_finseq_1(k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_178, plain, (v1_xboole_0(X1)|k3_finseq_4(X1,X2,X3,X4)=k11_finseq_1(X2,X3,X4)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_163])).
fof(c_0_179, plain, ![X31, X32, X33, X34]:(((((k3_finseq_1(X34)=np__3|X34!=k11_finseq_1(X31,X32,X33)|(~v1_relat_1(X34)|~v1_funct_1(X34)|~v1_finseq_1(X34)))&(k1_funct_1(X34,np__1)=X31|X34!=k11_finseq_1(X31,X32,X33)|(~v1_relat_1(X34)|~v1_funct_1(X34)|~v1_finseq_1(X34))))&(k1_funct_1(X34,np__2)=X32|X34!=k11_finseq_1(X31,X32,X33)|(~v1_relat_1(X34)|~v1_funct_1(X34)|~v1_finseq_1(X34))))&(k1_funct_1(X34,np__3)=X33|X34!=k11_finseq_1(X31,X32,X33)|(~v1_relat_1(X34)|~v1_funct_1(X34)|~v1_finseq_1(X34))))&(k3_finseq_1(X34)!=np__3|k1_funct_1(X34,np__1)!=X31|k1_funct_1(X34,np__2)!=X32|k1_funct_1(X34,np__3)!=X33|X34=k11_finseq_1(X31,X32,X33)|(~v1_relat_1(X34)|~v1_funct_1(X34)|~v1_finseq_1(X34)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t45_finseq_1])])])).
cnf(c_0_180, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X2),np__3)=X3|X3!=k3_matrix_1(X1,esk2_0,X2,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_165]), c_0_127]), c_0_127]), c_0_127]), c_0_109]), c_0_127]), c_0_166]), c_0_127]), c_0_167]), c_0_127]), c_0_168]), c_0_144])])])).
cnf(c_0_181, plain, (m1_subset_1(k3_matrix_1(X1,X2,X3,X4),X1)|~v1_matrix_1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m2_finseq_1(X2,k3_finseq_2(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_98])])).
fof(c_0_182, plain, ![X175, X176]:(~m1_subset_1(X175,X176)|(v1_xboole_0(X176)|r2_hidden(X175,X176))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_183, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_171])])).
cnf(c_0_184, negated_conjecture, (k2_zfmisc_1(k4_finseq_1(esk2_0),k1_enumset1(np__1,np__2,np__3))=k2_matrix_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_172, c_0_173]), c_0_174])])).
cnf(c_0_185, negated_conjecture, (~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~v1_xboole_0(k2_matrix_1(esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_166])])).
cnf(c_0_186, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),k2_zfmisc_1(X2,X4))), inference(evalgc,[status(thm)],[c_0_176])).
cnf(c_0_187, negated_conjecture, (k11_finseq_1(k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0), inference(evalgc,[status(thm)],[c_0_177])).
cnf(c_0_188, plain, (k3_finseq_4(X1,X2,X3,X4)=k11_finseq_1(X2,X3,X4)|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_178])).
cnf(c_0_189, plain, (X1=k11_finseq_1(X2,X3,X4)|k3_finseq_1(X1)!=np__3|k1_funct_1(X1,np__1)!=X2|k1_funct_1(X1,np__2)!=X3|k1_funct_1(X1,np__3)!=X4|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_179])).
cnf(c_0_190, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),np__3)=k3_matrix_1(X2,esk2_0,X1,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m1_subset_1(k3_matrix_1(X2,esk2_0,X1,np__3),X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_180])])).
cnf(c_0_191, negated_conjecture, (m1_subset_1(k3_matrix_1(esk1_0,esk2_0,X1,X2),esk1_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_145]), c_0_144])])])).
cnf(c_0_192, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_182])).
cnf(c_0_193, negated_conjecture, (v1_xboole_0(k2_matrix_1(esk2_0))|~v1_xboole_0(k4_finseq_1(esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_184])])).
cnf(c_0_194, negated_conjecture, (~v1_xboole_0(k2_matrix_1(esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_166])])).
cnf(c_0_195, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k2_tarski(X3,X1),k1_tarski(X1)),k2_zfmisc_1(X2,X4))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_56])])).
cnf(c_0_196, negated_conjecture, (k11_finseq_1(k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_188]), c_0_88])])).
cnf(c_0_197, plain, (X1=k11_finseq_1(X2,X3,X4)|k3_finseq_1(X1)!=np__3|k1_funct_1(X1,np__3)!=X4|k1_funct_1(X1,np__1)!=X2|k1_funct_1(X1,np__2)!=X3|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_189])).
cnf(c_0_198, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),np__3)=k3_matrix_1(esk1_0,esk2_0,X1,np__3)|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_190, c_0_191]), c_0_145]), c_0_109])])])).
cnf(c_0_199, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_192])).
cnf(c_0_200, negated_conjecture, (~v1_xboole_0(k4_finseq_1(esk2_0))), inference(sr,[status(thm)],[c_0_193, c_0_194])).
cnf(c_0_201, plain, (m1_subset_1(np__2,k5_numbers)), inference(split_conjunct,[status(thm)],[spc2_numerals])).
cnf(c_0_202, plain, (r2_hidden(X1,k4_finseq_1(X2))|~r2_hidden(k2_tarski(k2_tarski(X3,X1),k1_tarski(X1)),k2_matrix_1(X2))|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_76])])).
cnf(c_0_203, negated_conjecture, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(esk2_0))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_161])])).
cnf(c_0_204, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X2,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_205, plain, (X1=esk8_2(X1,X2)|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_206, plain, (m2_finseq_1(esk4_5(X1,X2,X3,X4,X5),X1)|X5!=k3_matrix_1(X1,X2,X3,X4)|~m1_subset_1(X5,X1)|~r2_hidden(k4_tarski(X3,X4),k2_matrix_1(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_207, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X2,X5)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_208, negated_conjecture, (k11_finseq_1(k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_188]), c_0_88])])).
cnf(c_0_209, plain, (k11_finseq_1(X2,X3,k1_funct_1(X1,np__3))=X1|k1_funct_1(X1,np__1)!=X2|k1_funct_1(X1,np__2)!=X3|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_197])])).
cnf(c_0_210, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),np__3)=k3_matrix_1(esk1_0,esk2_0,X1,np__3)|~v7_ordinal1(X1)|~m1_subset_1(X1,k4_finseq_1(esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_199]), c_0_200])])).
cnf(c_0_211, plain, (m1_subset_1(np__2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_201])).
cnf(c_0_212, negated_conjecture, (r2_hidden(X1,k4_finseq_1(esk2_0))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_203]), c_0_144]), c_0_173]), c_0_159]), c_0_174])])])).
cnf(c_0_213, plain, (m1_subset_1(np__1,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_214, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X2,X4,X5)), inference(evalgc,[status(thm)],[c_0_204])).
fof(c_0_215, plain, ![X35, X36, X37, X38, X39]:(v1_xboole_0(X35)|(~m1_subset_1(X36,k5_numbers)|(~m1_matrix_1(X37,X35,X36,X36)|(~m2_finseq_1(X38,X35)|(~v7_ordinal1(X39)|(X38!=k1_funct_1(X37,X39)|~r2_hidden(X39,k2_finseq_1(X36))|k3_finseq_1(X38)=X36)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t36_matrixr2])])])])).
cnf(c_0_216, plain, (esk8_2(X1,X2)=X1|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2))), inference(evalgc,[status(thm)],[c_0_205])).
cnf(c_0_217, plain, (m2_finseq_1(esk4_5(X1,X2,X3,X4,X5),X1)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m1_subset_1(X5,X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~r2_hidden(k2_tarski(k2_tarski(X3,X4),k1_tarski(X3)),k2_matrix_1(X2))), inference(rw,[status(thm)],[c_0_206, c_0_41])).
cnf(c_0_218, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X2,X5)), inference(evalgc,[status(thm)],[c_0_207])).
cnf(c_0_219, negated_conjecture, (k11_finseq_1(k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_188]), c_0_88])])).
cnf(c_0_220, negated_conjecture, (k11_finseq_1(X1,X2,k3_matrix_1(esk1_0,esk2_0,X3,np__3))=k1_funct_1(esk2_0,X3)|k1_funct_1(k1_funct_1(esk2_0,X3),np__1)!=X1|k1_funct_1(k1_funct_1(esk2_0,X3),np__2)!=X2|k3_finseq_1(k1_funct_1(esk2_0,X3))!=np__3|~v1_finseq_1(k1_funct_1(esk2_0,X3))|~v1_funct_1(k1_funct_1(esk2_0,X3))|~v1_relat_1(k1_funct_1(esk2_0,X3))|~m1_subset_1(X3,k4_finseq_1(esk2_0))|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_210])])).
cnf(c_0_221, plain, (v7_ordinal1(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_211]), c_0_91])])])).
cnf(c_0_222, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),np__3)=k3_matrix_1(esk1_0,esk2_0,X1,np__3)|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_212])])).
cnf(c_0_223, plain, (m1_subset_1(np__1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_213])).
cnf(c_0_224, plain, (r2_hidden(X1,X2)|X2!=k1_enumset1(X1,X3,X4)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_214])])).
cnf(c_0_225, plain, (k3_finseq_1(X1)=X2|v1_xboole_0(X3)|~m1_matrix_1(X1,X3,X2,X2)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_226, plain, (v1_xboole_0(X1)|k3_finseq_1(X4)=X2|~m1_subset_1(X2,k5_numbers)|~m1_matrix_1(X3,X1,X2,X2)|~m2_finseq_1(X4,X1)|~v7_ordinal1(X5)|X4!=k1_funct_1(X3,X5)|~r2_hidden(X5,k2_finseq_1(X2))), inference(split_conjunct,[status(thm)],[c_0_215])).
cnf(c_0_227, plain, (m2_subset_1(esk8_2(X1,X2),k1_numbers,k5_numbers)|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_228, plain, (esk8_2(X1,X2)=X1|~v7_ordinal1(X2)|~r2_hidden(X1,k2_finseq_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_121])])).
cnf(c_0_229, plain, (m2_finseq_1(esk4_5(X1,X2,X3,X4,X5),X1)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m1_subset_1(X5,X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~r2_hidden(k2_tarski(k2_tarski(X3,X4),k1_tarski(X3)),k2_matrix_1(X2))), inference(evalgc,[status(thm)],[c_0_217])).
cnf(c_0_230, plain, (r2_hidden(X1,X2)|X2!=k1_enumset1(X3,X1,X4)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_218])])).
cnf(c_0_231, negated_conjecture, (k11_finseq_1(k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k1_funct_1(esk2_0,np__2),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_220]), c_0_221])])])).
cnf(c_0_232, negated_conjecture, (k11_finseq_1(X1,X2,k3_matrix_1(esk1_0,esk2_0,X3,np__3))=k1_funct_1(esk2_0,X3)|k1_funct_1(k1_funct_1(esk2_0,X3),np__1)!=X1|k1_funct_1(k1_funct_1(esk2_0,X3),np__2)!=X2|k3_finseq_1(k1_funct_1(esk2_0,X3))!=np__3|~v1_finseq_1(k1_funct_1(esk2_0,X3))|~v1_funct_1(k1_funct_1(esk2_0,X3))|~v1_relat_1(k1_funct_1(esk2_0,X3))|~v7_ordinal1(X3)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_222])])).
cnf(c_0_233, plain, (v7_ordinal1(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_223]), c_0_91])])])).
cnf(c_0_234, plain, (r2_hidden(X1,k1_enumset1(X1,X2,X3))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_224])])).
cnf(c_0_235, plain, (k3_finseq_1(X1)=X2|v1_xboole_0(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X1,X3,X2,X2)), inference(evalgc,[status(thm)],[c_0_225])).
cnf(c_0_236, negated_conjecture, (k1_funct_1(esk4_5(X1,esk2_0,X2,X3,X4),X3)=X4|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_203]), c_0_144])])])).
cnf(c_0_237, negated_conjecture, (esk4_5(X1,esk2_0,X2,X3,X4)=k1_funct_1(esk2_0,X2)|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_203]), c_0_144])])])).
cnf(c_0_238, plain, (k3_finseq_1(X4)=X2|v1_xboole_0(X1)|X4!=k1_funct_1(X3,X5)|~v7_ordinal1(X5)|~m2_finseq_1(X4,X1)|~m1_subset_1(X2,k5_numbers)|~r2_hidden(X5,k2_finseq_1(X2))|~m1_matrix_1(X3,X1,X2,X2)), inference(evalgc,[status(thm)],[c_0_226])).
cnf(c_0_239, plain, (m2_subset_1(esk8_2(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2))), inference(evalgc,[status(thm)],[c_0_227])).
cnf(c_0_240, plain, (esk8_2(X1,np__3)=X1|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_128]), c_0_109])])])).
cnf(c_0_241, plain, (r1_xxreal_0(esk8_2(X1,X2),X2)|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_242, plain, (r1_xxreal_0(np__1,esk8_2(X1,X2))|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_243, plain, (m2_finseq_1(esk4_5(X1,X2,X3,X3,X4),X1)|X4!=k3_matrix_1(X1,X2,X3,X3)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X3,k2_finseq_1(k1_matrix_1(X2)))|~r2_hidden(X3,k4_finseq_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_96])])).
cnf(c_0_244, plain, (r2_hidden(X1,k1_enumset1(X2,X1,X3))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_230])])).
cnf(c_0_245, negated_conjecture, (k11_finseq_1(k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__2),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_232]), c_0_233]), c_0_234])])])).
cnf(c_0_246, negated_conjecture, (k3_finseq_1(esk2_0)=np__3|~v7_ordinal1(np__3)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_87]), c_0_88])])).
cnf(c_0_247, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X2),X3)=X4|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_236, c_0_237])])).
cnf(c_0_248, plain, (k3_finseq_1(X1)=np__3|v1_xboole_0(X2)|X1!=k1_funct_1(X3,X4)|~m2_finseq_1(X1,X2)|~v7_ordinal1(X4)|~r2_hidden(X4,k1_enumset1(np__1,np__2,np__3))|~m1_matrix_1(X3,X2,np__3,np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_128]), c_0_90])])])).
cnf(c_0_249, plain, (r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_128]), c_0_109])])])).
cnf(c_0_250, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~r2_hidden(X1,a_1_0_finseq_1(np__3))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_240]), c_0_109])])])).
cnf(c_0_251, plain, (r1_xxreal_0(esk8_2(X1,X2),X2)|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2))), inference(evalgc,[status(thm)],[c_0_241])).
cnf(c_0_252, plain, (r1_xxreal_0(np__1,esk8_2(X1,X2))|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2))), inference(evalgc,[status(thm)],[c_0_242])).
cnf(c_0_253, plain, (m2_finseq_1(esk4_5(X1,X2,X3,X3,X4),X1)|X4!=k3_matrix_1(X1,X2,X3,X3)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~r2_hidden(X3,k2_finseq_1(k1_matrix_1(X2)))|~r2_hidden(X3,k4_finseq_1(X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_243, c_0_117]), c_0_118]), c_0_119])).
cnf(c_0_254, plain, (esk4_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|v1_xboole_0(k2_finseq_1(k1_matrix_1(X2)))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2))|~m1_subset_1(X4,k2_finseq_1(k1_matrix_1(X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_199])])).
cnf(c_0_255, plain, (~v1_xboole_0(k1_enumset1(X2,X1,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_244])])).
cnf(c_0_256, negated_conjecture, (k11_finseq_1(k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__2),k1_funct_1(esk2_0,np__3))!=esk2_0|k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_232]), c_0_109]), c_0_166])])])).
cnf(c_0_257, negated_conjecture, (k3_finseq_1(esk2_0)=np__3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_246, c_0_109])])).
cnf(c_0_258, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X2)=k3_matrix_1(X3,esk2_0,X1,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X3))|~m1_subset_1(k3_matrix_1(X3,esk2_0,X1,X2),X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_247])])).
cnf(c_0_259, plain, (k3_finseq_1(X1)=np__3|v1_xboole_0(X2)|X1!=k1_funct_1(X3,X4)|~m2_finseq_1(X1,X2)|~v7_ordinal1(X4)|~m1_matrix_1(X3,X2,np__3,np__3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,np__3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_248, c_0_249])])).
cnf(c_0_260, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_121]), c_0_128]), c_0_109])])])).
cnf(c_0_261, plain, (r1_xxreal_0(X1,np__3)|~r2_hidden(X1,a_1_0_finseq_1(np__3))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_240]), c_0_109])])])).
cnf(c_0_262, plain, (r1_xxreal_0(np__1,X1)|~r2_hidden(X1,a_1_0_finseq_1(np__3))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_240]), c_0_109])])])).
fof(c_0_263, plain, ![X163, X164]:(~r2_hidden(X163,X164)|m1_subset_1(X163,X164)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_264, negated_conjecture, (m2_finseq_1(esk4_5(X1,esk2_0,X2,X2,X3),X1)|X3!=k3_matrix_1(X1,esk2_0,X2,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k4_finseq_1(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_127]), c_0_144]), c_0_128])])])).
cnf(c_0_265, negated_conjecture, (esk4_5(X1,esk2_0,X2,X3,X4)=k1_funct_1(esk2_0,X2)|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_127]), c_0_128]), c_0_144]), c_0_128])]), c_0_255])])).
cnf(c_0_266, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_256, c_0_209]), c_0_257]), c_0_173]), c_0_159]), c_0_174])])])).
cnf(c_0_267, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X2)=k3_matrix_1(esk1_0,esk2_0,X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_191]), c_0_145])])])).
cnf(c_0_268, negated_conjecture, (k3_finseq_1(X1)=np__3|X1!=k1_funct_1(esk2_0,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(X1,esk1_0)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_259, c_0_87]), c_0_88])])).
cnf(c_0_269, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_199]), c_0_255])])).
cnf(c_0_270, plain, (r1_xxreal_0(X1,np__3)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_261, c_0_121]), c_0_128]), c_0_109])])])).
cnf(c_0_271, plain, (r1_xxreal_0(np__1,X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_121]), c_0_128]), c_0_109])])])).
cnf(c_0_272, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_263])).
cnf(c_0_273, negated_conjecture, (m2_finseq_1(esk4_5(X1,esk2_0,X2,X2,X3),X1)|X3!=k3_matrix_1(X1,esk2_0,X2,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))), inference(csr,[status(thm)],[c_0_264, c_0_212])).
cnf(c_0_274, negated_conjecture, (esk4_5(X1,esk2_0,X2,X3,X4)=k1_funct_1(esk2_0,X2)|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_212])])).
cnf(c_0_275, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_267]), c_0_221]), c_0_244])])])).
cnf(c_0_276, negated_conjecture, (k3_finseq_1(X1)=np__3|X1!=k1_funct_1(esk2_0,X2)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(X1,esk1_0)|~v7_ordinal1(X2)|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_269])])).
cnf(c_0_277, plain, (r1_xxreal_0(X1,np__3)|~m1_subset_1(X1,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_270, c_0_199]), c_0_255])])).
cnf(c_0_278, plain, (r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_199]), c_0_255])])).
cnf(c_0_279, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_272])).
cnf(c_0_280, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X2),X1)|X3!=k3_matrix_1(X1,esk2_0,X2,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_274])])).
cnf(c_0_281, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_267]), c_0_233]), c_0_221]), c_0_234]), c_0_244])])])).
cnf(c_0_282, negated_conjecture, (k3_finseq_1(X1)=np__3|X1!=k1_funct_1(esk2_0,X2)|~m2_finseq_1(X1,esk1_0)|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3))|~v7_ordinal1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_276, c_0_277]), c_0_278])).
cnf(c_0_283, plain, (m1_subset_1(X1,k1_enumset1(X2,X1,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_279, c_0_244])])).
cnf(c_0_284, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X1),X2)|X3!=k3_matrix_1(X2,esk2_0,X1,X1)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m1_subset_1(X3,X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))), inference(csr,[status(thm)],[c_0_280, c_0_279])).
cnf(c_0_285, plain, (k1_funct_1(esk4_5(X1,X2,X3,X4,X5),X4)=X5|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,k1_matrix_1(X2))|~v7_ordinal1(k1_matrix_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_135])])).
cnf(c_0_286, negated_conjecture, (esk4_5(X1,esk2_0,X2,X3,X4)=k1_funct_1(esk2_0,X2)|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X2,k4_finseq_1(esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_199]), c_0_200])])).
cnf(c_0_287, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_281, c_0_267]), c_0_221]), c_0_109]), c_0_244]), c_0_166])])])).
cnf(c_0_288, negated_conjecture, (k3_finseq_1(X1)=np__3|X1!=k1_funct_1(esk2_0,np__2)|~m2_finseq_1(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_283]), c_0_221])])])).
cnf(c_0_289, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X1),X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m1_subset_1(k3_matrix_1(X2,esk2_0,X1,X1),X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_284])])).
cnf(c_0_290, plain, (m1_subset_1(X1,k1_enumset1(X1,X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_279, c_0_234])])).
cnf(c_0_291, plain, (k3_finseq_1(X1)=np__3|v1_xboole_0(X2)|X1!=k1_funct_1(X3,np__3)|~m2_finseq_1(X1,X2)|~m1_matrix_1(X3,X2,np__3,np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_248, c_0_166]), c_0_109])])])).
cnf(c_0_292, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X2),X3)=X4|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X3,np__3)|~r1_xxreal_0(np__1,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0))|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X2,k4_finseq_1(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_285, c_0_286]), c_0_127]), c_0_144]), c_0_127]), c_0_109])])])).
cnf(c_0_293, plain, (m1_subset_1(X1,k1_enumset1(np__1,np__2,np__3))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_279, c_0_249])])).
cnf(c_0_294, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_287, c_0_267]), c_0_233]), c_0_109]), c_0_234]), c_0_166])])])).
cnf(c_0_295, negated_conjecture, (k3_finseq_1(k1_funct_1(esk2_0,np__2))=np__3|~m2_finseq_1(k1_funct_1(esk2_0,np__2),esk1_0)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_288])])).
cnf(c_0_296, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X1),esk1_0)|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_289, c_0_191]), c_0_145])])])).
cnf(c_0_297, negated_conjecture, (k3_finseq_1(X1)=np__3|X1!=k1_funct_1(esk2_0,np__1)|~m2_finseq_1(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_290]), c_0_233])])])).
cnf(c_0_298, negated_conjecture, (k3_finseq_1(X1)=np__3|X1!=k1_funct_1(esk2_0,np__3)|~m2_finseq_1(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_87]), c_0_88])])).
cnf(c_0_299, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X2)=X3|X3!=k3_matrix_1(X4,esk2_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X4))|~m1_subset_1(X3,X4)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_292, c_0_279]), c_0_293])).
cnf(c_0_300, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_294, c_0_267]), c_0_221]), c_0_233]), c_0_244]), c_0_234])])])).
cnf(c_0_301, negated_conjecture, (k3_finseq_1(k1_funct_1(esk2_0,np__2))=np__3), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_295, c_0_296]), c_0_221]), c_0_244])])])).
cnf(c_0_302, negated_conjecture, (k3_finseq_1(k1_funct_1(esk2_0,np__1))=np__3|~m2_finseq_1(k1_funct_1(esk2_0,np__1),esk1_0)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_297])])).
cnf(c_0_303, negated_conjecture, (k3_finseq_1(k1_funct_1(esk2_0,np__3))=np__3|~m2_finseq_1(k1_funct_1(esk2_0,np__3),esk1_0)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_298])])).
cnf(c_0_304, negated_conjecture, (v1_finseq_1(k1_funct_1(esk2_0,X1))|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_296])])).
cnf(c_0_305, negated_conjecture, (v1_funct_1(k1_funct_1(esk2_0,X1))|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_296])])).
cnf(c_0_306, negated_conjecture, (v1_relat_1(k1_funct_1(esk2_0,X1))|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_296])])).
cnf(c_0_307, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X2)=k3_matrix_1(X3,esk2_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X3))|~m1_subset_1(k3_matrix_1(X3,esk2_0,X1,X2),X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_299])])).
cnf(c_0_308, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_300, c_0_301])])])).
cnf(c_0_309, negated_conjecture, (k3_finseq_1(k1_funct_1(esk2_0,np__1))=np__3), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_296]), c_0_233]), c_0_234])])])).
cnf(c_0_310, negated_conjecture, (k3_finseq_1(k1_funct_1(esk2_0,np__3))=np__3), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_303, c_0_296]), c_0_109]), c_0_166])])])).
cnf(c_0_311, negated_conjecture, (v1_finseq_1(k1_funct_1(esk2_0,np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_304, c_0_244]), c_0_221])])])).
cnf(c_0_312, negated_conjecture, (v1_finseq_1(k1_funct_1(esk2_0,np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_304, c_0_234]), c_0_233])])])).
cnf(c_0_313, negated_conjecture, (v1_finseq_1(k1_funct_1(esk2_0,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_304, c_0_166]), c_0_109])])])).
cnf(c_0_314, negated_conjecture, (v1_funct_1(k1_funct_1(esk2_0,np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_244]), c_0_221])])])).
cnf(c_0_315, negated_conjecture, (v1_funct_1(k1_funct_1(esk2_0,np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_234]), c_0_233])])])).
cnf(c_0_316, negated_conjecture, (v1_funct_1(k1_funct_1(esk2_0,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_166]), c_0_109])])])).
cnf(c_0_317, negated_conjecture, (v1_relat_1(k1_funct_1(esk2_0,np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_306, c_0_244]), c_0_221])])])).
cnf(c_0_318, negated_conjecture, (v1_relat_1(k1_funct_1(esk2_0,np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_306, c_0_234]), c_0_233])])])).
cnf(c_0_319, negated_conjecture, (v1_relat_1(k1_funct_1(esk2_0,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_306, c_0_166]), c_0_109])])])).
cnf(c_0_320, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X2)=k3_matrix_1(esk1_0,esk2_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_307, c_0_191]), c_0_145])])])).
cnf(c_0_321, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_322, plain, (r1_xxreal_0(np__1,np__1)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r1_r1])).
cnf(c_0_323, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_308, c_0_309]), c_0_310]), c_0_311]), c_0_312]), c_0_313]), c_0_314]), c_0_315]), c_0_316]), c_0_317]), c_0_318]), c_0_319])])).
cnf(c_0_324, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X2)=k3_matrix_1(esk1_0,esk2_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_320, c_0_212])])).
cnf(c_0_325, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_321])).
cnf(c_0_326, plain, (r1_xxreal_0(np__1,np__1)), inference(evalgc,[status(thm)],[c_0_322])).
cnf(c_0_327, negated_conjecture, (~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_324]), c_0_325]), c_0_168]), c_0_326]), c_0_233]), c_0_234])])])).
cnf(c_0_328, negated_conjecture, (m1_subset_1(X1,k4_finseq_1(esk2_0))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_279, c_0_212])])).
cnf(c_0_329, negated_conjecture, (~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_328]), c_0_244])])])).
cnf(c_0_330, negated_conjecture, (~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_191]), c_0_221])])])).
cnf(c_0_331, negated_conjecture, (~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_330, c_0_191]), c_0_109]), c_0_221])])])).
cnf(c_0_332, negated_conjecture, (~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_331, c_0_191]), c_0_233]), c_0_221])])])).
cnf(c_0_333, negated_conjecture, (~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_191]), c_0_221]), c_0_109])])])).
cnf(c_0_334, negated_conjecture, (~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_333, c_0_191]), c_0_109])])])).
cnf(c_0_335, negated_conjecture, (~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_334, c_0_191]), c_0_233]), c_0_109])])])).
cnf(c_0_336, negated_conjecture, (~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_335, c_0_191]), c_0_221]), c_0_233])])])).
cnf(c_0_337, negated_conjecture, (~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_336, c_0_191]), c_0_109]), c_0_233])])])).
cnf(c_0_338, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_337, c_0_191]), c_0_233])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 339
# Proof object clause steps            : 273
# Proof object formula steps           : 66
# Proof object conjectures             : 111
# Proof object clause conjectures      : 108
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 54
# Proof object initial formulas used   : 37
# Proof object generating inferences   : 137
# Proof object simplifying inferences  : 254
# Training examples: 195 positive, 1328 negative
# Training: Positive examples begin
cnf(c_0_71, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_70, plain, (v6_membered(k4_ordinal1))).# trainpos
cnf(c_0_88, negated_conjecture, (~v1_xboole_0(esk1_0))).# trainpos
cnf(c_0_87, negated_conjecture, (m1_matrix_1(esk2_0,esk1_0,np__3,np__3))).# trainpos
cnf(c_0_187, negated_conjecture, (k11_finseq_1(k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0)).# trainpos
cnf(c_0_167, plain, (m2_subset_1(np__3,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_90, plain, (m1_subset_1(np__3,k5_numbers))).# trainpos
cnf(c_0_223, plain, (m1_subset_1(np__1,k5_numbers))).# trainpos
cnf(c_0_211, plain, (m1_subset_1(np__2,k5_numbers))).# trainpos
cnf(c_0_106, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).# trainpos
cnf(c_0_168, plain, (r1_xxreal_0(np__1,np__3))).# trainpos
cnf(c_0_326, plain, (r1_xxreal_0(np__1,np__1))).# trainpos
cnf(c_0_157, plain, (r1_tarski(X1,X1))).# trainpos
cnf(c_0_128, plain, (k2_finseq_1(np__3)=k1_enumset1(np__1,np__2,np__3))).# trainpos
cnf(c_0_105, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_325, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_56, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1))).# trainpos
cnf(c_0_102, plain, (k1_finseq_1(X1)=k2_finseq_1(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_103, plain, (k1_finseq_1(X1)=a_1_0_finseq_1(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_146, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_98, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_97, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2))).# trainpos
cnf(c_0_118, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_119, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_279, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_156, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_89, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_235, plain, (k3_finseq_1(X1)=X2|v1_xboole_0(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X1,X3,X2,X2))).# trainpos
cnf(c_0_104, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_199, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_138, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X5,X2))).# trainpos
cnf(c_0_218, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X2,X5))).# trainpos
cnf(c_0_214, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X2,X4,X5))).# trainpos
cnf(c_0_216, plain, (esk8_2(X1,X2)=X1|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2)))).# trainpos
cnf(c_0_170, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).# trainpos
cnf(c_0_86, plain, (k1_matrix_1(X1)=X2|v1_xboole_0(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X1,X3,X2,X2))).# trainpos
cnf(c_0_111, plain, (v1_xboole_0(X2)|v1_matrix_1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4))).# trainpos
cnf(c_0_252, plain, (r1_xxreal_0(np__1,esk8_2(X1,X2))|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2)))).# trainpos
cnf(c_0_251, plain, (r1_xxreal_0(esk8_2(X1,X2),X2)|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2)))).# trainpos
cnf(c_0_239, plain, (m2_subset_1(esk8_2(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2)))).# trainpos
cnf(c_0_76, plain, (k2_zfmisc_1(k4_finseq_1(X1),k2_finseq_1(k1_matrix_1(X1)))=k2_matrix_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)|~v1_matrix_1(X1))).# trainpos
cnf(c_0_186, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),k2_zfmisc_1(X2,X4)))).# trainpos
cnf(c_0_188, plain, (k3_finseq_4(X1,X2,X3,X4)=k11_finseq_1(X2,X3,X4)|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_55, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),k2_zfmisc_1(X2,X4))|~r2_hidden(X3,X4)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_132, plain, (k2_matrix_1(X1)=k2_zfmisc_1(k2_finseq_1(X2),k2_finseq_1(X2))|v1_xboole_0(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X1,X3,X2,X2))).# trainpos
cnf(c_0_112, plain, (v1_xboole_0(X2)|m2_finseq_1(X1,k3_finseq_2(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4))).# trainpos
cnf(c_0_169, plain, (m1_subset_1(k3_matrix_1(X2,X1,X3,X4),X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X1)|~m1_finseq_1(X1,k3_finseq_2(X2)))).# trainpos
cnf(c_0_101, plain, (r2_hidden(X2,a_1_0_finseq_1(X3))|X1!=X2|~v7_ordinal1(X3)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_238, plain, (k3_finseq_1(X4)=X2|v1_xboole_0(X1)|X4!=k1_funct_1(X3,X5)|~v7_ordinal1(X5)|~m2_finseq_1(X4,X1)|~m1_subset_1(X2,k5_numbers)|~r2_hidden(X5,k2_finseq_1(X2))|~m1_matrix_1(X3,X1,X2,X2))).# trainpos
cnf(c_0_197, plain, (X1=k11_finseq_1(X2,X3,X4)|k3_finseq_1(X1)!=np__3|k1_funct_1(X1,np__3)!=X4|k1_funct_1(X1,np__1)!=X2|k1_funct_1(X1,np__2)!=X3|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_74, plain, (esk4_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m1_subset_1(X5,X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~r2_hidden(k2_tarski(k2_tarski(X3,X4),k1_tarski(X3)),k2_matrix_1(X2)))).# trainpos
cnf(c_0_229, plain, (m2_finseq_1(esk4_5(X1,X2,X3,X4,X5),X1)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m1_subset_1(X5,X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~r2_hidden(k2_tarski(k2_tarski(X3,X4),k1_tarski(X3)),k2_matrix_1(X2)))).# trainpos
cnf(c_0_94, plain, (k1_funct_1(esk4_5(X2,X3,X4,X5,X1),X5)=X1|X1!=k3_matrix_1(X2,X3,X4,X5)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v1_matrix_1(X3)|~m1_subset_1(X1,X2)|~m2_finseq_1(X3,k3_finseq_2(X2))|~r2_hidden(k2_tarski(k2_tarski(X4,X5),k1_tarski(X4)),k2_matrix_1(X3)))).# trainpos
cnf(c_0_181, plain, (m1_subset_1(k3_matrix_1(X1,X2,X3,X4),X1)|~v1_matrix_1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m2_finseq_1(X2,k3_finseq_2(X1)))).# trainpos
cnf(c_0_109, plain, (v7_ordinal1(np__3))).# trainpos
cnf(c_0_221, plain, (v7_ordinal1(np__2))).# trainpos
cnf(c_0_233, plain, (v7_ordinal1(np__1))).# trainpos
cnf(c_0_151, plain, (r2_hidden(X1,X2)|X2!=k1_enumset1(X3,X4,X1))).# trainpos
cnf(c_0_137, plain, (v7_ordinal1(k1_xboole_0))).# trainpos
cnf(c_0_230, plain, (r2_hidden(X1,X2)|X2!=k1_enumset1(X3,X1,X4))).# trainpos
cnf(c_0_121, plain, (a_1_0_finseq_1(X1)=k2_finseq_1(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_75, plain, (r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X1)),k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3))).# trainpos
cnf(c_0_224, plain, (r2_hidden(X1,X2)|X2!=k1_enumset1(X1,X3,X4))).# trainpos
cnf(c_0_120, plain, (r2_hidden(X1,a_1_0_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_196, negated_conjecture, (k11_finseq_1(k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0))).# trainpos
cnf(c_0_195, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k2_tarski(X3,X1),k1_tarski(X1)),k2_zfmisc_1(X2,X4)))).# trainpos
cnf(c_0_117, plain, (v1_relat_1(X1)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_95, plain, (esk4_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(k2_tarski(k2_tarski(X4,X3),k1_tarski(X3)),k2_matrix_1(X2)))).# trainpos
cnf(c_0_108, negated_conjecture, (k1_matrix_1(esk2_0)=np__3|~v7_ordinal1(np__3))).# trainpos
cnf(c_0_129, negated_conjecture, (v1_matrix_1(esk2_0)|~v7_ordinal1(np__3))).# trainpos
cnf(c_0_171, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_122, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_115, plain, (k1_funct_1(esk4_5(X1,X2,X3,X4,X5),X4)=X5|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(k2_tarski(k2_tarski(X4,X3),k1_tarski(X3)),k2_matrix_1(X2)))).# trainpos
cnf(c_0_246, negated_conjecture, (k3_finseq_1(esk2_0)=np__3|~v7_ordinal1(np__3))).# trainpos
cnf(c_0_160, plain, (~v1_xboole_0(k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3))).# trainpos
cnf(c_0_130, negated_conjecture, (m2_finseq_1(esk2_0,k3_finseq_2(esk1_0))|~v7_ordinal1(np__3))).# trainpos
cnf(c_0_209, plain, (k11_finseq_1(X2,X3,k1_funct_1(X1,np__3))=X1|k1_funct_1(X1,np__1)!=X2|k1_funct_1(X1,np__2)!=X3|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_147, negated_conjecture, (k2_zfmisc_1(k1_enumset1(np__1,np__2,np__3),k1_enumset1(np__1,np__2,np__3))=k2_matrix_1(esk2_0)|~v7_ordinal1(np__3))).# trainpos
cnf(c_0_248, plain, (k3_finseq_1(X1)=np__3|v1_xboole_0(X2)|X1!=k1_funct_1(X3,X4)|~m2_finseq_1(X1,X2)|~v7_ordinal1(X4)|~r2_hidden(X4,k1_enumset1(np__1,np__2,np__3))|~m1_matrix_1(X3,X2,np__3,np__3))).# trainpos
cnf(c_0_244, plain, (r2_hidden(X1,k1_enumset1(X2,X1,X3)))).# trainpos
cnf(c_0_234, plain, (r2_hidden(X1,k1_enumset1(X1,X2,X3)))).# trainpos
cnf(c_0_143, negated_conjecture, (k2_zfmisc_1(k4_finseq_1(esk2_0),k1_enumset1(np__1,np__2,np__3))=k2_matrix_1(esk2_0)|~v1_matrix_1(esk2_0)|~v1_finseq_1(esk2_0)|~v1_funct_1(esk2_0)|~v1_relat_1(esk2_0))).# trainpos
cnf(c_0_166, plain, (r2_hidden(X1,k1_enumset1(X2,X3,X1)))).# trainpos
cnf(c_0_159, negated_conjecture, (v1_funct_1(esk2_0))).# trainpos
cnf(c_0_208, negated_conjecture, (k11_finseq_1(k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0))).# trainpos
cnf(c_0_173, negated_conjecture, (v1_finseq_1(esk2_0))).# trainpos
cnf(c_0_174, negated_conjecture, (v1_relat_1(esk2_0))).# trainpos
cnf(c_0_136, plain, (r1_xxreal_0(X1,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_135, plain, (r2_hidden(X1,k2_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_183, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).# trainpos
cnf(c_0_191, negated_conjecture, (m1_subset_1(k3_matrix_1(esk1_0,esk2_0,X1,X2),esk1_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_203, negated_conjecture, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(esk2_0))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_202, plain, (r2_hidden(X1,k4_finseq_1(X2))|~r2_hidden(k2_tarski(k2_tarski(X3,X1),k1_tarski(X1)),k2_matrix_1(X2))|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).# trainpos
cnf(c_0_228, plain, (esk8_2(X1,X2)=X1|~v7_ordinal1(X2)|~r2_hidden(X1,k2_finseq_1(X2)))).# trainpos
cnf(c_0_175, negated_conjecture, (~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~v1_xboole_0(k2_matrix_1(esk2_0)))).# trainpos
cnf(c_0_96, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X2)),k2_matrix_1(X3))|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(X3)))|~r2_hidden(X2,k4_finseq_1(X3))|~v1_matrix_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).# trainpos
cnf(c_0_255, plain, (~v1_xboole_0(k1_enumset1(X2,X1,X3)))).# trainpos
cnf(c_0_172, negated_conjecture, (k2_zfmisc_1(k4_finseq_1(esk2_0),k1_enumset1(np__1,np__2,np__3))=k2_matrix_1(esk2_0)|~v1_finseq_1(esk2_0)|~v1_relat_1(esk2_0))).# trainpos
cnf(c_0_219, negated_conjecture, (k11_finseq_1(k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0))).# trainpos
cnf(c_0_150, plain, (r1_xxreal_0(X1,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_237, negated_conjecture, (esk4_5(X1,esk2_0,X2,X3,X4)=k1_funct_1(esk2_0,X2)|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_116, plain, (esk4_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X4,k2_finseq_1(k1_matrix_1(X2)))|~r2_hidden(X3,k4_finseq_1(X2)))).# trainpos
cnf(c_0_283, plain, (m1_subset_1(X1,k1_enumset1(X2,X1,X3)))).# trainpos
cnf(c_0_212, negated_conjecture, (r2_hidden(X1,k4_finseq_1(esk2_0))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_236, negated_conjecture, (k1_funct_1(esk4_5(X1,esk2_0,X2,X3,X4),X3)=X4|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_290, plain, (m1_subset_1(X1,k1_enumset1(X1,X2,X3)))).# trainpos
cnf(c_0_133, plain, (k1_funct_1(esk4_5(X1,X2,X3,X4,X5),X4)=X5|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X4,k2_finseq_1(k1_matrix_1(X2)))|~r2_hidden(X3,k4_finseq_1(X2)))).# trainpos
cnf(c_0_249, plain, (r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3))).# trainpos
cnf(c_0_243, plain, (m2_finseq_1(esk4_5(X1,X2,X3,X3,X4),X1)|X4!=k3_matrix_1(X1,X2,X3,X3)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X3,k2_finseq_1(k1_matrix_1(X2)))|~r2_hidden(X3,k4_finseq_1(X2)))).# trainpos
cnf(c_0_240, plain, (esk8_2(X1,np__3)=X1|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_185, negated_conjecture, (~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~v1_xboole_0(k2_matrix_1(esk2_0)))).# trainpos
cnf(c_0_164, negated_conjecture, (k1_funct_1(esk4_5(X1,esk2_0,X2,X3,X4),X3)=X4|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k4_finseq_1(esk2_0)))).# trainpos
cnf(c_0_194, negated_conjecture, (~v1_xboole_0(k2_matrix_1(esk2_0)))).# trainpos
cnf(c_0_264, negated_conjecture, (m2_finseq_1(esk4_5(X1,esk2_0,X2,X2,X3),X1)|X3!=k3_matrix_1(X1,esk2_0,X2,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k4_finseq_1(esk2_0)))).# trainpos
cnf(c_0_261, plain, (r1_xxreal_0(X1,np__3)|~r2_hidden(X1,a_1_0_finseq_1(np__3))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_193, negated_conjecture, (v1_xboole_0(k2_matrix_1(esk2_0))|~v1_xboole_0(k4_finseq_1(esk2_0)))).# trainpos
cnf(c_0_247, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X2),X3)=X4|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_262, plain, (r1_xxreal_0(np__1,X1)|~r2_hidden(X1,a_1_0_finseq_1(np__3))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_149, plain, (esk4_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,k1_matrix_1(X2))|~v7_ordinal1(k1_matrix_1(X2)))).# trainpos
cnf(c_0_285, plain, (k1_funct_1(esk4_5(X1,X2,X3,X4,X5),X4)=X5|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,k1_matrix_1(X2))|~v7_ordinal1(k1_matrix_1(X2)))).# trainpos
cnf(c_0_270, plain, (r1_xxreal_0(X1,np__3)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_165, plain, (esk4_5(X1,X2,X3,k1_matrix_1(X2),X4)=k1_funct_1(X2,X3)|X4!=k3_matrix_1(X1,X2,X3,k1_matrix_1(X2))|~m2_subset_1(k1_matrix_1(X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(X2))|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(k1_matrix_1(X2))|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2)))).# trainpos
cnf(c_0_271, plain, (r1_xxreal_0(np__1,X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_293, plain, (m1_subset_1(X1,k1_enumset1(np__1,np__2,np__3))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3))).# trainpos
cnf(c_0_277, plain, (r1_xxreal_0(X1,np__3)|~m1_subset_1(X1,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_278, plain, (r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_180, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X2),np__3)=X3|X3!=k3_matrix_1(X1,esk2_0,X2,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0)))).# trainpos
cnf(c_0_250, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~r2_hidden(X1,a_1_0_finseq_1(np__3))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_190, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),np__3)=k3_matrix_1(X2,esk2_0,X1,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m1_subset_1(k3_matrix_1(X2,esk2_0,X1,np__3),X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0)))).# trainpos
cnf(c_0_260, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_198, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),np__3)=k3_matrix_1(esk1_0,esk2_0,X1,np__3)|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0)))).# trainpos
cnf(c_0_254, plain, (esk4_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|v1_xboole_0(k2_finseq_1(k1_matrix_1(X2)))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2))|~m1_subset_1(X4,k2_finseq_1(k1_matrix_1(X2))))).# trainpos
cnf(c_0_269, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_291, plain, (k3_finseq_1(X1)=np__3|v1_xboole_0(X2)|X1!=k1_funct_1(X3,np__3)|~m2_finseq_1(X1,X2)|~m1_matrix_1(X3,X2,np__3,np__3))).# trainpos
cnf(c_0_258, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X2)=k3_matrix_1(X3,esk2_0,X1,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X3))|~m1_subset_1(k3_matrix_1(X3,esk2_0,X1,X2),X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_259, plain, (k3_finseq_1(X1)=np__3|v1_xboole_0(X2)|X1!=k1_funct_1(X3,X4)|~m2_finseq_1(X1,X2)|~v7_ordinal1(X4)|~m1_matrix_1(X3,X2,np__3,np__3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,np__3))).# trainpos
cnf(c_0_222, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),np__3)=k3_matrix_1(esk1_0,esk2_0,X1,np__3)|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_210, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),np__3)=k3_matrix_1(esk1_0,esk2_0,X1,np__3)|~v7_ordinal1(X1)|~m1_subset_1(X1,k4_finseq_1(esk2_0)))).# trainpos
cnf(c_0_265, negated_conjecture, (esk4_5(X1,esk2_0,X2,X3,X4)=k1_funct_1(esk2_0,X2)|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0)))).# trainpos
cnf(c_0_267, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X2)=k3_matrix_1(esk1_0,esk2_0,X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_298, negated_conjecture, (k3_finseq_1(X1)=np__3|X1!=k1_funct_1(esk2_0,np__3)|~m2_finseq_1(X1,esk1_0))).# trainpos
cnf(c_0_268, negated_conjecture, (k3_finseq_1(X1)=np__3|X1!=k1_funct_1(esk2_0,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(X1,esk1_0)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_286, negated_conjecture, (esk4_5(X1,esk2_0,X2,X3,X4)=k1_funct_1(esk2_0,X2)|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X2,k4_finseq_1(esk2_0)))).# trainpos
cnf(c_0_274, negated_conjecture, (esk4_5(X1,esk2_0,X2,X3,X4)=k1_funct_1(esk2_0,X2)|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_220, negated_conjecture, (k11_finseq_1(X1,X2,k3_matrix_1(esk1_0,esk2_0,X3,np__3))=k1_funct_1(esk2_0,X3)|k1_funct_1(k1_funct_1(esk2_0,X3),np__1)!=X1|k1_funct_1(k1_funct_1(esk2_0,X3),np__2)!=X2|k3_finseq_1(k1_funct_1(esk2_0,X3))!=np__3|~v1_finseq_1(k1_funct_1(esk2_0,X3))|~v1_funct_1(k1_funct_1(esk2_0,X3))|~v1_relat_1(k1_funct_1(esk2_0,X3))|~m1_subset_1(X3,k4_finseq_1(esk2_0))|~v7_ordinal1(X3))).# trainpos
cnf(c_0_276, negated_conjecture, (k3_finseq_1(X1)=np__3|X1!=k1_funct_1(esk2_0,X2)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(X1,esk1_0)|~v7_ordinal1(X2)|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_231, negated_conjecture, (k11_finseq_1(k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k1_funct_1(esk2_0,np__2),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__2)))).# trainpos
cnf(c_0_232, negated_conjecture, (k11_finseq_1(X1,X2,k3_matrix_1(esk1_0,esk2_0,X3,np__3))=k1_funct_1(esk2_0,X3)|k1_funct_1(k1_funct_1(esk2_0,X3),np__1)!=X1|k1_funct_1(k1_funct_1(esk2_0,X3),np__2)!=X2|k3_finseq_1(k1_funct_1(esk2_0,X3))!=np__3|~v1_finseq_1(k1_funct_1(esk2_0,X3))|~v1_funct_1(k1_funct_1(esk2_0,X3))|~v1_relat_1(k1_funct_1(esk2_0,X3))|~v7_ordinal1(X3)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_288, negated_conjecture, (k3_finseq_1(X1)=np__3|X1!=k1_funct_1(esk2_0,np__2)|~m2_finseq_1(X1,esk1_0))).# trainpos
cnf(c_0_297, negated_conjecture, (k3_finseq_1(X1)=np__3|X1!=k1_funct_1(esk2_0,np__1)|~m2_finseq_1(X1,esk1_0))).# trainpos
cnf(c_0_245, negated_conjecture, (k11_finseq_1(k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__2),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__1)))).# trainpos
cnf(c_0_256, negated_conjecture, (k11_finseq_1(k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__2),k1_funct_1(esk2_0,np__3))!=esk2_0|k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).# trainpos
cnf(c_0_303, negated_conjecture, (k3_finseq_1(k1_funct_1(esk2_0,np__3))=np__3|~m2_finseq_1(k1_funct_1(esk2_0,np__3),esk1_0))).# trainpos
cnf(c_0_266, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).# trainpos
cnf(c_0_295, negated_conjecture, (k3_finseq_1(k1_funct_1(esk2_0,np__2))=np__3|~m2_finseq_1(k1_funct_1(esk2_0,np__2),esk1_0))).# trainpos
cnf(c_0_302, negated_conjecture, (k3_finseq_1(k1_funct_1(esk2_0,np__1))=np__3|~m2_finseq_1(k1_funct_1(esk2_0,np__1),esk1_0))).# trainpos
cnf(c_0_275, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).# trainpos
cnf(c_0_281, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).# trainpos
cnf(c_0_287, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).# trainpos
cnf(c_0_294, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).# trainpos
cnf(c_0_300, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).# trainpos
cnf(c_0_292, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X2),X3)=X4|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X3,np__3)|~r1_xxreal_0(np__1,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0))|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X2,k4_finseq_1(esk2_0)))).# trainpos
cnf(c_0_307, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X2)=k3_matrix_1(X3,esk2_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X3))|~m1_subset_1(k3_matrix_1(X3,esk2_0,X1,X2),X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0)))).# trainpos
cnf(c_0_328, negated_conjecture, (m1_subset_1(X1,k4_finseq_1(esk2_0))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_280, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X2),X1)|X3!=k3_matrix_1(X1,esk2_0,X2,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_289, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X1),X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m1_subset_1(k3_matrix_1(X2,esk2_0,X1,X1),X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_320, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X2)=k3_matrix_1(esk1_0,esk2_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0)))).# trainpos
cnf(c_0_296, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X1),esk1_0)|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_324, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X2)=k3_matrix_1(esk1_0,esk2_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_301, negated_conjecture, (k3_finseq_1(k1_funct_1(esk2_0,np__2))=np__3)).# trainpos
cnf(c_0_310, negated_conjecture, (k3_finseq_1(k1_funct_1(esk2_0,np__3))=np__3)).# trainpos
cnf(c_0_309, negated_conjecture, (k3_finseq_1(k1_funct_1(esk2_0,np__1))=np__3)).# trainpos
cnf(c_0_305, negated_conjecture, (v1_funct_1(k1_funct_1(esk2_0,X1))|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_306, negated_conjecture, (v1_relat_1(k1_funct_1(esk2_0,X1))|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_304, negated_conjecture, (v1_finseq_1(k1_funct_1(esk2_0,X1))|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).# trainpos
cnf(c_0_311, negated_conjecture, (v1_finseq_1(k1_funct_1(esk2_0,np__2)))).# trainpos
cnf(c_0_312, negated_conjecture, (v1_finseq_1(k1_funct_1(esk2_0,np__1)))).# trainpos
cnf(c_0_318, negated_conjecture, (v1_relat_1(k1_funct_1(esk2_0,np__1)))).# trainpos
cnf(c_0_317, negated_conjecture, (v1_relat_1(k1_funct_1(esk2_0,np__2)))).# trainpos
cnf(c_0_316, negated_conjecture, (v1_funct_1(k1_funct_1(esk2_0,np__3)))).# trainpos
cnf(c_0_315, negated_conjecture, (v1_funct_1(k1_funct_1(esk2_0,np__1)))).# trainpos
cnf(c_0_314, negated_conjecture, (v1_funct_1(k1_funct_1(esk2_0,np__2)))).# trainpos
cnf(c_0_313, negated_conjecture, (v1_finseq_1(k1_funct_1(esk2_0,np__3)))).# trainpos
cnf(c_0_319, negated_conjecture, (v1_relat_1(k1_funct_1(esk2_0,np__3)))).# trainpos
cnf(c_0_308, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).# trainpos
cnf(c_0_327, negated_conjecture, (~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0)))).# trainpos
cnf(c_0_329, negated_conjecture, (~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0))).# trainpos
cnf(c_0_330, negated_conjecture, (~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0))).# trainpos
cnf(c_0_331, negated_conjecture, (~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0))).# trainpos
cnf(c_0_332, negated_conjecture, (~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0))).# trainpos
cnf(c_0_333, negated_conjecture, (~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0))).# trainpos
cnf(c_0_334, negated_conjecture, (~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0))).# trainpos
cnf(c_0_335, negated_conjecture, (~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0))).# trainpos
cnf(c_0_336, negated_conjecture, (~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0))).# trainpos
cnf(c_0_337, negated_conjecture, (~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_80, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_148, plain, (v1_xboole_0(esk10_0))).#trainneg
cnf(i_0_49, plain, (v2_xxreal_0(np__3))).#trainneg
cnf(i_0_73, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_67, plain, (v2_xxreal_0(np__2))).#trainneg
cnf(i_0_76, plain, (v3_membered(k1_numbers))).#trainneg
cnf(i_0_147, plain, (v1_xxreal_0(esk10_0))).#trainneg
cnf(i_0_154, plain, (~v1_xboole_0(np__3))).#trainneg
cnf(i_0_179, plain, (~v1_xboole_0(np__2))).#trainneg
cnf(i_0_75, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_122, plain, (~v1_finset_1(k4_ordinal1))).#trainneg
cnf(i_0_123, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_100, plain, (~r1_xxreal_0(np__3,np__2))).#trainneg
cnf(i_0_174, plain, (~r1_xxreal_0(np__2,np__1))).#trainneg
cnf(i_0_168, plain, (v1_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_153, plain, (v3_ordinal1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_97, plain, (v1_card_1(k1_card_1(X1)))).#trainneg
cnf(i_0_46, plain, (m1_subset_1(np__3,k1_numbers))).#trainneg
cnf(i_0_70, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_64, plain, (m1_subset_1(np__2,k1_numbers))).#trainneg
cnf(i_0_140, plain, (r1_xxreal_0(np__1,np__2))).#trainneg
cnf(i_0_150, plain, (r1_xxreal_0(np__2,np__2))).#trainneg
cnf(i_0_89, plain, (~v1_xboole_0(k1_tarski(X1)))).#trainneg
cnf(i_0_125, plain, (~v1_xboole_0(k13_finseq_1(X1)))).#trainneg
cnf(i_0_91, plain, (v1_xboole_0(X1)|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_102, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_66, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_163, plain, (v2_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_105, plain, (v1_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_164, plain, (k1_card_1(k1_card_1(X1))=k1_card_1(X1))).#trainneg
cnf(i_0_136, plain, (m1_subset_1(esk9_1(X1),X1))).#trainneg
cnf(i_0_159, plain, (m1_finseq_2(k3_finseq_2(X1),X1))).#trainneg
cnf(i_0_149, plain, (~v1_xboole_0(k2_tarski(X1,X2)))).#trainneg
cnf(i_0_53, plain, (v1_relat_1(k11_finseq_1(X1,X2,X3)))).#trainneg
cnf(i_0_52, plain, (v1_funct_1(k11_finseq_1(X1,X2,X3)))).#trainneg
cnf(i_0_63, plain, (v1_finseq_1(k11_finseq_1(X1,X2,X3)))).#trainneg
cnf(i_0_84, plain, (v1_finset_1(k1_enumset1(X1,X2,X3)))).#trainneg
cnf(i_0_86, plain, (~v1_xboole_0(k11_finseq_1(X1,X2,X3)))).#trainneg
cnf(i_0_131, plain, (v1_xboole_0(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_95, plain, (v1_relat_1(k1_tarski(k2_tarski(k2_tarski(np__1,X1),k1_tarski(np__1)))))).#trainneg
cnf(i_0_35, plain, (X1=X2|k3_finseq_1(X1)!=k3_finseq_1(X2)|k1_funct_1(X1,esk5_2(X1,X2))!=k1_funct_1(X2,esk5_2(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_112, plain, (v1_finset_1(k1_card_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_32, plain, (v7_ordinal1(X1)|~r2_hidden(X1,k4_ordinal1))).#trainneg
cnf(i_0_94, plain, (v1_funct_1(k1_tarski(k2_tarski(k2_tarski(np__1,X1),k1_tarski(np__1)))))).#trainneg
cnf(i_0_170, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_51, plain, (v7_ordinal1(X1)|~v3_ordinal1(X1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_33, plain, (r2_hidden(X1,k4_ordinal1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_82, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_61, plain, (v1_funct_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_60, plain, (v1_finseq_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_172, plain, (v3_card_1(k11_finseq_1(X1,X2,X3),np__3))).#trainneg
cnf(i_0_44, plain, (m2_finseq_1(X1,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_167, plain, (v5_relat_1(X1,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_173, plain, (v1_xboole_0(k1_finseq_1(X1))|~v1_xboole_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_177, plain, (v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_83, plain, (v1_finseq_1(k1_tarski(k2_tarski(k2_tarski(np__1,X1),k1_tarski(np__1)))))).#trainneg
cnf(i_0_81, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_151, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_108, plain, (k1_card_1(X1)=k3_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_103, plain, (k9_xtuple_0(X1)=k4_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_128, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_69, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_139, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_74, plain, (v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_157, plain, (v7_ordinal1(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_40, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_135, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_134, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_137, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_129, plain, (m1_subset_1(k4_finseq_1(X1),k1_zfmisc_1(k5_numbers))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_106, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_92, plain, (m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_115, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_77, plain, (m1_subset_1(k1_matrix_1(X1),k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)|~v1_matrix_1(X1))).#trainneg
cnf(i_0_169, plain, (v1_funcop_1(k11_finseq_1(X1,X2,X3))|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_156, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_160, plain, (v1_xxreal_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v2_valued_0(X1))).#trainneg
cnf(i_0_101, plain, (k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_155, plain, (k1_nat_1(X1,X2)=k1_nat_1(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_126, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_87, plain, (m2_subset_1(k1_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_11, plain, (X1=X5|X1=X4|X1=X3|X2!=k1_enumset1(X3,X4,X5)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_113, plain, (v6_membered(k1_enumset1(X1,X2,X3))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_104, plain, (v3_membered(k1_enumset1(X1,X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_85, plain, (v2_membered(k1_enumset1(X1,X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_114, plain, (v1_membered(k1_enumset1(X1,X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_22, plain, (k3_finseq_1(X1)=np__3|X1!=k11_finseq_1(X2,X3,X4)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_19, plain, (k1_funct_1(X1,np__3)=X2|X1!=k11_finseq_1(X3,X4,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_5, plain, (X4=k1_enumset1(X1,X2,X3)|esk3_4(X1,X2,X3,X4)!=X3|~r2_hidden(esk3_4(X1,X2,X3,X4),X4))).#trainneg
cnf(i_0_145, plain, (v1_relat_1(k7_finseq_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_110, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_16, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k2_tarski(X3,X1),k1_tarski(X3)),k2_zfmisc_1(X4,X2)))).#trainneg
cnf(i_0_43, plain, (v1_xboole_0(X1)|m2_finseq_1(k3_finseq_4(X1,X2,X3,X4),X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_144, plain, (v1_funct_1(k7_finseq_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_143, plain, (v1_finseq_1(k7_finseq_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_21, plain, (k1_funct_1(X1,np__1)=X2|X1!=k11_finseq_1(X2,X3,X4)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_20, plain, (k1_funct_1(X1,np__2)=X2|X1!=k11_finseq_1(X3,X2,X4)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_6, plain, (X4=k1_enumset1(X1,X2,X3)|esk3_4(X1,X2,X3,X4)!=X2|~r2_hidden(esk3_4(X1,X2,X3,X4),X4))).#trainneg
cnf(i_0_7, plain, (X4=k1_enumset1(X1,X2,X3)|esk3_4(X1,X2,X3,X4)!=X1|~r2_hidden(esk3_4(X1,X2,X3,X4),X4))).#trainneg
cnf(i_0_109, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_161, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_166, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,X3),k3_xcmplx_0(X2,X3))=k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_88, plain, (k7_finseq_1(k7_finseq_1(k1_tarski(k2_tarski(k2_tarski(np__1,X1),k1_tarski(np__1))),k1_tarski(k2_tarski(k2_tarski(np__1,X2),k1_tarski(np__1)))),k1_tarski(k2_tarski(k2_tarski(np__1,X3),k1_tarski(np__1))))=k11_finseq_1(X1,X2,X3))).#trainneg
cnf(i_0_56, plain, (X2=X3|v1_xboole_0(X1)|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|k3_matrix_1(X1,X2,esk6_3(X1,X2,X3),esk7_3(X1,X2,X3))!=k3_matrix_1(X1,X3,esk6_3(X1,X2,X3),esk7_3(X1,X2,X3))|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_38, plain, (X1=X2|v7_ordinal1(esk5_2(X1,X2))|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_37, plain, (X1=X2|r1_xxreal_0(np__1,esk5_2(X1,X2))|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_36, plain, (X1=X2|r1_xxreal_0(esk5_2(X1,X2),k3_finseq_1(X1))|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_59, plain, (X2=X3|v1_xboole_0(X1)|v7_ordinal1(esk6_3(X1,X2,X3))|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_24, plain, (v1_xboole_0(X1)|m1_matrix_1(k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k3_finseq_4(X1,X8,X9,X10)),X1,np__3,np__3)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_4, plain, (X4=k1_enumset1(X1,X2,X3)|esk3_4(X1,X2,X3,X4)=X3|esk3_4(X1,X2,X3,X4)=X2|esk3_4(X1,X2,X3,X4)=X1|r2_hidden(esk3_4(X1,X2,X3,X4),X4))).#trainneg
cnf(i_0_58, plain, (X2=X3|v1_xboole_0(X1)|v7_ordinal1(esk7_3(X1,X2,X3))|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_26, plain, (X5=k3_matrix_1(X2,X3,X4,X6)|X5!=k1_funct_1(X1,X6)|X1!=k1_funct_1(X3,X4)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v1_matrix_1(X3)|~m1_subset_1(X5,X2)|~m2_finseq_1(X1,X2)|~m2_finseq_1(X3,k3_finseq_2(X2))|~r2_hidden(k2_tarski(k2_tarski(X4,X6),k1_tarski(X4)),k2_matrix_1(X3)))).#trainneg
cnf(i_0_57, plain, (X2=X3|v1_xboole_0(X1)|r2_hidden(k2_tarski(k2_tarski(esk6_3(X1,X2,X3),esk7_3(X1,X2,X3)),k1_tarski(esk6_3(X1,X2,X3))),k2_matrix_1(X2))|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_334, plain, (k1_funct_1(k11_finseq_1(X1,X2,X3),np__2)=X2)).#trainneg
cnf(i_0_333, plain, (k1_funct_1(k11_finseq_1(X1,X2,X3),np__1)=X1)).#trainneg
cnf(i_0_275, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_180, plain, (esk10_0=k1_xboole_0)).#trainneg
cnf(i_0_233, plain, (v1_xreal_0(np__2))).#trainneg
cnf(i_0_181, plain, (~v1_xboole_0(k5_numbers))).#trainneg
cnf(i_0_204, plain, (~m1_subset_1(k5_numbers,k5_numbers))).#trainneg
cnf(i_0_235, plain, (v1_xreal_0(np__1))).#trainneg
cnf(i_0_237, plain, (v1_xreal_0(np__3))).#trainneg
cnf(i_0_188, plain, (k1_card_1(X1)=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_191, plain, (v1_relat_1(k1_tarski(k2_tarski(k2_tarski(X1,np__1),k1_tarski(np__1)))))).#trainneg
cnf(i_0_318, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X3,X1)),k2_zfmisc_1(X4,X2)))).#trainneg
cnf(i_0_342, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_300, plain, (k1_funct_1(k11_finseq_1(X1,X2,X3),np__3)=X3)).#trainneg
cnf(i_0_324, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X3)),k2_zfmisc_1(X2,X4)))).#trainneg
cnf(i_0_208, plain, (m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_260, plain, (v3_membered(k5_numbers))).#trainneg
cnf(i_0_231, plain, (v1_xreal_0(k1_xboole_0)|~v3_membered(k5_numbers))).#trainneg
cnf(i_0_272, plain, (v5_relat_1(X1,X2)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_201, plain, (v7_ordinal1(X1)|~v1_finset_1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_193, plain, (v1_relat_1(k1_tarski(k2_tarski(k1_tarski(np__1),k2_tarski(np__1,X1)))))).#trainneg
cnf(i_0_299, plain, (k3_finseq_1(k11_finseq_1(X1,X2,X3))=np__3)).#trainneg
cnf(i_0_413, plain, (k1_funct_1(esk4_5(X1,X2,X3,X4,X5),X4)=X5|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X3,X4)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_298, plain, (X1=X2|X1=X3|X1=X4|~r2_hidden(X1,k1_enumset1(X2,X3,X4)))).#trainneg
cnf(i_0_256, plain, (v1_xboole_0(X1)|~v1_xboole_0(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_328, negated_conjecture, (k11_finseq_1(k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0))).#trainneg
cnf(i_0_197, plain, (v1_funct_1(k1_tarski(k2_tarski(k2_tarski(X1,np__1),k1_tarski(np__1)))))).#trainneg
cnf(i_0_243, plain, (v1_xcmplx_0(k1_xboole_0)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_245, plain, (v1_xcmplx_0(np__2)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_199, plain, (v1_funct_1(k1_tarski(k2_tarski(k1_tarski(np__1),k2_tarski(np__1,X1)))))).#trainneg
cnf(i_0_239, plain, (v7_ordinal1(X1)|v1_xboole_0(k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_403, plain, (esk4_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X3,X4)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_316, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X3)),k2_zfmisc_1(X4,X2)))).#trainneg
cnf(i_0_336, plain, (m1_subset_1(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_246, plain, (v1_xcmplx_0(np__2)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_212, plain, (v1_xboole_0(k2_finseq_1(X1))|~v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_247, plain, (v1_xcmplx_0(np__1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_329, negated_conjecture, (k11_finseq_1(k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0))).#trainneg
cnf(i_0_248, plain, (v1_xcmplx_0(np__1)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_249, plain, (v1_xcmplx_0(np__3)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_250, plain, (v1_xcmplx_0(np__3)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_284, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_286, plain, (v3_valued_0(X1)|~v3_membered(X2)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_211, plain, (k1_finseq_1(X1)=k1_xboole_0|~v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_213, plain, (v1_xboole_0(a_1_0_finseq_1(X1))|~v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_406, plain, (m2_finseq_1(esk4_5(X1,X2,X3,X4,X5),X1)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(k2_tarski(k2_tarski(X4,X3),k1_tarski(X3)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_408, plain, (m2_finseq_1(esk4_5(X1,X2,X3,X4,X5),X1)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X3,X4)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_252, plain, (v1_card_1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_326, plain, (r2_hidden(X1,X2)|v1_xboole_0(k2_zfmisc_1(X2,X4))|~m1_subset_1(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),k2_zfmisc_1(X2,X4)))).#trainneg
cnf(i_0_287, plain, (v7_ordinal1(k1_matrix_1(X1))|~v1_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_216, plain, (v7_ordinal1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_276, plain, (v1_xboole_0(esk9_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_232, plain, (v1_xreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_214, plain, (v7_ordinal1(esk9_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_244, plain, (v1_xcmplx_0(k5_numbers)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_226, plain, (v1_finseq_1(k1_tarski(k2_tarski(k2_tarski(X1,np__1),k1_tarski(np__1)))))).#trainneg
cnf(i_0_291, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v2_valued_0(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_330, plain, (v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_228, plain, (v1_finseq_1(k1_tarski(k2_tarski(k1_tarski(np__1),k2_tarski(np__1,X1)))))).#trainneg
cnf(i_0_320, plain, (r2_hidden(X1,X2)|v1_xboole_0(k2_zfmisc_1(X4,X2))|~m1_subset_1(k2_tarski(k2_tarski(X3,X1),k1_tarski(X3)),k2_zfmisc_1(X4,X2)))).#trainneg
cnf(i_0_315, plain, (r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2)))|~r2_hidden(k2_tarski(k2_tarski(X3,X1),k1_tarski(X3)),k2_matrix_1(X2))|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_398, plain, (X1=k3_matrix_1(X2,X3,X4,X5)|X1!=k1_funct_1(X6,X5)|X6!=k1_funct_1(X3,X4)|~v1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X6,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~r2_hidden(k2_tarski(k1_tarski(X4),k2_tarski(X4,X5)),k2_matrix_1(X3)))).#trainneg
cnf(i_0_362, plain, (k3_finseq_1(X1)=X2|v1_xboole_0(X3)|v1_xboole_0(k2_finseq_1(X2))|X1!=k1_funct_1(X4,X5)|~m2_finseq_1(X1,X3)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X5)|~m1_matrix_1(X4,X3,X2,X2)|~m1_subset_1(X5,k2_finseq_1(X2)))).#trainneg
cnf(i_0_254, plain, (v1_xboole_0(k3_finseq_1(X1))|~v1_xboole_0(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_296, plain, (m2_subset_1(k1_nat_1(X2,X1),k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_415, plain, (k1_funct_1(esk4_5(X1,X2,X3,X4,X5),X4)=X5|v1_xboole_0(k2_matrix_1(X2))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(k2_tarski(k2_tarski(X3,X4),k1_tarski(X3)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_396, plain, (X1=k3_matrix_1(X2,X3,X4,X5)|X1!=k1_funct_1(X6,X5)|X6!=k1_funct_1(X3,X4)|~v1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X6,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~r2_hidden(k2_tarski(k2_tarski(X5,X4),k1_tarski(X4)),k2_matrix_1(X3)))).#trainneg
cnf(i_0_386, plain, (m1_matrix_1(k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),k3_finseq_4(X1,X8,X9,X10)),X1,np__3,np__3)|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_416, plain, (X2=X3|v1_xboole_0(X1)|~v1_xboole_0(k2_matrix_1(X2))|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_230, plain, (v1_xreal_0(esk9_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_242, plain, (v1_xcmplx_0(esk9_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_278, plain, (v7_ordinal1(k4_finseq_1(X1))|~v6_membered(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_307, plain, (k1_enumset1(X1,X2,X3)=k5_numbers|esk3_4(X1,X2,X3,k5_numbers)!=X3|~v7_ordinal1(esk3_4(X1,X2,X3,k5_numbers)))).#trainneg
cnf(i_0_251, plain, (k3_finseq_1(k1_card_1(X1))=k1_card_1(X1)|~v1_finseq_1(k1_card_1(X1))|~v1_funct_1(k1_card_1(X1))|~v1_relat_1(k1_card_1(X1)))).#trainneg
cnf(i_0_147, plain, (v1_xxreal_0(k1_xboole_0))).#trainneg
cnf(i_0_544, plain, (v1_xcmplx_0(k1_xboole_0))).#trainneg
cnf(i_0_545, plain, (v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_592, plain, (v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_253, plain, (k1_card_1(k3_finseq_1(X1))=k3_finseq_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_255, plain, (v1_finset_1(k3_finseq_1(X1))|~v1_finset_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_596, plain, (v1_xcmplx_0(np__3))).#trainneg
cnf(i_0_657, plain, (v1_card_1(np__3))).#trainneg
cnf(i_0_258, plain, (v1_relat_1(esk9_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_430, plain, (v1_card_1(k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_599, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_259, plain, (v3_membered(esk9_1(k1_zfmisc_1(X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_261, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_433, plain, (v1_finset_1(k1_xboole_0)|~v1_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_587, negated_conjecture, (v1_xboole_0(esk2_0)|~v1_xboole_0(k3_finseq_2(esk1_0)))).#trainneg
cnf(i_0_495, negated_conjecture, (k11_finseq_1(k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0))).#trainneg
cnf(i_0_262, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_673, plain, (v1_xcmplx_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_705, negated_conjecture, (~v1_finseq_1(esk2_0)|~v1_funct_1(esk2_0)|~v1_relat_1(esk2_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_600, negated_conjecture, (v3_valued_0(esk2_0)|~v3_membered(k3_finseq_2(esk1_0)))).#trainneg
cnf(i_0_419, plain, (v1_xreal_0(X2)|~v3_valued_0(k11_finseq_1(X1,X2,X3)))).#trainneg
cnf(i_0_420, plain, (v1_xxreal_0(X2)|~v2_valued_0(k11_finseq_1(X1,X2,X3)))).#trainneg
cnf(i_0_263, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_264, plain, (v3_membered(X1)|~v3_membered(k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_461, plain, (r2_hidden(X1,k4_finseq_1(X2))|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X3)),k2_matrix_1(X2))|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_533, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_638, negated_conjecture, (r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(esk2_0)))).#trainneg
cnf(i_0_520, plain, (esk4_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X4,X3)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_534, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_538, negated_conjecture, (k11_finseq_1(k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0))).#trainneg
cnf(i_0_269, plain, (v4_relat_1(esk9_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_431, plain, (k1_card_1(k1_xboole_0)=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_671, negated_conjecture, (v1_xboole_0(esk9_1(k1_zfmisc_1(k2_matrix_1(esk2_0))))|~v1_xboole_0(k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_421, plain, (v1_xcmplx_0(X2)|~v1_valued_0(k11_finseq_1(X1,X2,X3)))).#trainneg
cnf(i_0_422, plain, (v1_xreal_0(X1)|~v3_valued_0(k11_finseq_1(X1,X2,X3)))).#trainneg
cnf(i_0_270, plain, (v4_relat_1(X1,k5_numbers)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_271, plain, (v5_relat_1(esk9_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_640, negated_conjecture, (r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k2_matrix_1(esk2_0)))).#trainneg
cnf(i_0_648, negated_conjecture, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k2_matrix_1(esk2_0))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_641, negated_conjecture, (r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X2)),k2_matrix_1(esk2_0)))).#trainneg
cnf(i_0_647, negated_conjecture, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X2)),k2_matrix_1(esk2_0))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_441, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X1,X3)),k2_zfmisc_1(X4,X2)))).#trainneg
cnf(i_0_528, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_532, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_649, negated_conjecture, (v3_valued_0(X1)|~v3_membered(k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X1,k1_zfmisc_1(k2_matrix_1(esk2_0))))).#trainneg
cnf(i_0_273, plain, (esk8_2(X1,X2)=X1|v1_xboole_0(a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,a_1_0_finseq_1(X2)))).#trainneg
cnf(i_0_739, plain, (k3_finseq_1(k1_xboole_0)=k1_xboole_0|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|~v1_relat_1(k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_438, plain, (v1_relat_1(k1_tarski(k2_tarski(k1_tarski(np__1),k2_tarski(X1,np__1)))))).#trainneg
cnf(i_0_423, plain, (v1_xxreal_0(X1)|~v2_valued_0(k11_finseq_1(X1,X2,X3)))).#trainneg
cnf(i_0_424, plain, (v1_xcmplx_0(X1)|~v1_valued_0(k11_finseq_1(X1,X2,X3)))).#trainneg
cnf(i_0_274, plain, (v1_xboole_0(esk9_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_279, plain, (v1_xreal_0(k4_finseq_1(X1))|~v3_membered(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_852, plain, (v1_relat_1(k2_matrix_1(X1))|~v1_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_853, negated_conjecture, (v1_relat_1(k2_matrix_1(esk2_0)))).#trainneg
cnf(i_0_754, plain, (k1_card_1(np__3)=np__3)).#trainneg
cnf(i_0_795, plain, (v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_833, plain, (v1_finset_1(k1_xboole_0))).#trainneg
cnf(i_0_280, plain, (v1_xcmplx_0(k4_finseq_1(X1))|~v1_membered(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_740, plain, (r1_xxreal_0(k1_xboole_0,k1_xboole_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_914, negated_conjecture, (v4_relat_1(esk2_0,k5_numbers))).#trainneg
cnf(i_0_755, negated_conjecture, (v1_finset_1(np__3)|~v1_finset_1(esk2_0)|~v1_finseq_1(esk2_0)|~v1_funct_1(esk2_0)|~v1_relat_1(esk2_0))).#trainneg
cnf(i_0_741, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_281, plain, (v3_membered(k4_finseq_1(X1))|~v3_membered(k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_285, plain, (v3_valued_0(esk9_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v3_membered(X2))).#trainneg
cnf(i_0_848, plain, (r2_hidden(X1,k4_finseq_1(X2))|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_865, negated_conjecture, (esk4_5(X1,esk2_0,X2,X2,X3)=k1_funct_1(esk2_0,X2)|X3!=k3_matrix_1(X1,esk2_0,X2,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_529, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_536, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_902, plain, (k1_card_1(k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_288, plain, (v1_xreal_0(k1_matrix_1(X1))|~v3_membered(k5_numbers)|~v1_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_742, plain, (r1_xxreal_0(X1,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_818, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_819, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_289, plain, (v1_xcmplx_0(k1_matrix_1(X1))|~v1_membered(k5_numbers)|~v1_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_292, plain, (r1_xxreal_0(X1,k1_funct_1(X2,X3))|r1_xxreal_0(k1_funct_1(X2,X3),X1)|~v1_xxreal_0(X1)|~v2_valued_0(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_867, negated_conjecture, (k1_funct_1(esk4_5(X1,esk2_0,X2,X2,X3),X2)=X3|X3!=k3_matrix_1(X1,esk2_0,X2,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_1013, plain, (esk4_5(X1,X2,X3,X3,X4)=k1_funct_1(X2,X3)|X4!=k3_matrix_1(X1,X2,X3,X3)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X3,k2_finseq_1(k1_matrix_1(X2)))|~r2_hidden(X3,k4_finseq_1(X2)))).#trainneg
cnf(i_0_530, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_293, plain, (k1_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_743, plain, (r1_xxreal_0(X1,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_790, negated_conjecture, (v1_relat_1(esk9_1(k1_zfmisc_1(k2_matrix_1(esk2_0)))))).#trainneg
cnf(i_0_989, plain, (k3_finseq_1(k1_xboole_0)=k1_xboole_0|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|~v1_relat_1(k1_xboole_0))).#trainneg
cnf(i_0_820, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__3,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_454, plain, (v1_xreal_0(X3)|~v3_valued_0(k11_finseq_1(X1,X2,X3)))).#trainneg
cnf(i_0_294, plain, (k1_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_295, plain, (m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_980, negated_conjecture, (v3_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_matrix_1(esk2_0))))).#trainneg
cnf(i_0_1015, plain, (k1_funct_1(esk4_5(X1,X2,X3,X3,X4),X3)=X4|X4!=k3_matrix_1(X1,X2,X3,X3)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X3,k2_finseq_1(k1_matrix_1(X2)))|~r2_hidden(X3,k4_finseq_1(X2)))).#trainneg
cnf(i_0_455, plain, (v1_xxreal_0(X3)|~v2_valued_0(k11_finseq_1(X1,X2,X3)))).#trainneg
cnf(i_0_785, plain, (m1_subset_1(X1,k1_enumset1(X2,X3,X1)))).#trainneg
cnf(i_0_614, plain, (~v1_xboole_0(X3)|k1_funct_1(X3,np__1)!=X1|k1_funct_1(X3,np__2)!=X2|k3_finseq_1(X3)!=np__3|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_979, negated_conjecture, (v5_relat_1(k2_matrix_1(esk2_0),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_456, plain, (v1_xcmplx_0(X3)|~v1_valued_0(k11_finseq_1(X1,X2,X3)))).#trainneg
cnf(i_0_757, plain, (v1_finset_1(np__3)|~v1_finset_1(k11_finseq_1(X1,X2,X3)))).#trainneg
cnf(i_0_301, plain, (v3_valued_0(X1)|~v3_membered(k2_finseq_1(k1_matrix_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(k2_matrix_1(X2)))|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_302, plain, (v5_relat_1(X1,k2_finseq_1(k1_matrix_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(k2_matrix_1(X2)))|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_846, plain, (r2_hidden(X1,k4_finseq_1(X2))|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X3,X1)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_912, plain, (v1_xcmplx_0(X1)|~v3_valued_0(k11_finseq_1(X2,X1,X3)))).#trainneg
cnf(i_0_531, plain, (v3_valued_0(k2_zfmisc_1(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_897, negated_conjecture, (v4_relat_1(esk9_1(k1_zfmisc_1(k2_matrix_1(esk2_0))),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_303, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_matrix_1(X2)))|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_917, negated_conjecture, (v5_relat_1(esk9_1(k1_zfmisc_1(k2_matrix_1(esk2_0))),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_1000, plain, (v1_xcmplx_0(X1)|~v3_valued_0(k11_finseq_1(X1,X2,X3)))).#trainneg
cnf(i_0_331, plain, (v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_304, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_matrix_1(X2)))|~v1_xboole_0(k2_finseq_1(k1_matrix_1(X2)))|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_305, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_matrix_1(X2)))|~v1_xboole_0(k4_finseq_1(X2))|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1022, plain, (r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X2)),k2_matrix_1(X3))|~v1_matrix_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(X3)))|~r2_hidden(X2,k4_finseq_1(X3)))).#trainneg
cnf(i_0_642, negated_conjecture, (r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X2)),k2_matrix_1(esk2_0)))).#trainneg
cnf(i_0_450, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X2,X1)),k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_490, plain, (~v1_xboole_0(a_1_0_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1287, plain, (m2_finseq_1(esk4_5(X1,X2,X3,X4,X5),X1)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X4,k2_finseq_1(k1_matrix_1(X2)))|~r2_hidden(X3,k4_finseq_1(X2)))).#trainneg
cnf(i_0_1116, plain, (r1_xxreal_0(np__2,np__3))).#trainneg
cnf(i_0_1069, plain, (r1_xxreal_0(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_1226, plain, (k1_funct_1(X1,np__2)!=X2|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_306, plain, (v4_relat_1(X1,k4_finseq_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_matrix_1(X2)))|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1185, plain, (r1_xxreal_0(np__3,np__3))).#trainneg
cnf(i_0_308, plain, (X1=k1_enumset1(X2,X3,X4)|v1_xboole_0(X1)|esk3_4(X2,X3,X4,X1)!=X4|~m1_subset_1(esk3_4(X2,X3,X4,X1),X1))).#trainneg
cnf(i_0_1131, negated_conjecture, (esk4_5(X1,esk2_0,X2,X3,X4)=k1_funct_1(esk2_0,X2)|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_1147, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_332, plain, (m2_finseq_1(k11_finseq_1(X2,X3,X4),X1)|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1111, plain, (r1_xxreal_0(np__2,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,np__2))).#trainneg
cnf(i_0_1077, negated_conjecture, (v1_finset_1(np__3)|~m1_subset_1(esk2_0,k5_numbers))).#trainneg
cnf(i_0_1210, negated_conjecture, (v3_valued_0(k2_matrix_1(esk2_0)))).#trainneg
cnf(i_0_1235, plain, (v1_xcmplx_0(X1)|~v3_valued_0(k11_finseq_1(X2,X3,X1)))).#trainneg
cnf(i_0_345, plain, (k1_enumset1(X1,X2,X3)=k5_numbers|esk3_4(X1,X2,X3,k5_numbers)!=X2|~v7_ordinal1(esk3_4(X1,X2,X3,k5_numbers)))).#trainneg
cnf(i_0_1035, negated_conjecture, (r2_hidden(X1,k4_finseq_1(esk2_0))|~r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(esk2_0)))).#trainneg
cnf(i_0_1135, negated_conjecture, (esk4_5(X1,esk2_0,np__3,np__3,X2)=k11_finseq_1(X3,X4,X2)|k1_funct_1(esk4_5(X1,esk2_0,np__3,np__3,X2),np__1)!=X3|k1_funct_1(esk4_5(X1,esk2_0,np__3,np__3,X2),np__2)!=X4|k3_finseq_1(esk4_5(X1,esk2_0,np__3,np__3,X2))!=np__3|~v1_finseq_1(esk4_5(X1,esk2_0,np__3,np__3,X2))|~v1_funct_1(esk4_5(X1,esk2_0,np__3,np__3,X2))|~v1_relat_1(esk4_5(X1,esk2_0,np__3,np__3,X2))|X2!=k3_matrix_1(X1,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1043, negated_conjecture, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(esk2_0))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_1148, plain, (v7_ordinal1(X1)|~v6_membered(k1_enumset1(X2,X1,X3)))).#trainneg
cnf(i_0_672, plain, (m1_subset_1(esk9_1(X1),k5_numbers)|~v6_membered(X1))).#trainneg
cnf(i_0_1123, plain, (r1_xxreal_0(np__1,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,np__1))).#trainneg
cnf(i_0_346, plain, (X1=k1_enumset1(X2,X3,X4)|v1_xboole_0(X1)|esk3_4(X2,X3,X4,X1)!=X3|~m1_subset_1(esk3_4(X2,X3,X4,X1),X1))).#trainneg
cnf(i_0_517, plain, (v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_601, plain, (v3_valued_0(k3_finseq_4(X1,X2,X3,X4))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_347, plain, (k1_enumset1(X1,X2,X3)=k5_numbers|esk3_4(X1,X2,X3,k5_numbers)!=X1|~v7_ordinal1(esk3_4(X1,X2,X3,k5_numbers)))).#trainneg
cnf(i_0_348, plain, (X1=k1_enumset1(X2,X3,X4)|v1_xboole_0(X1)|esk3_4(X2,X3,X4,X1)!=X2|~m1_subset_1(esk3_4(X2,X3,X4,X1),X1))).#trainneg
cnf(i_0_1060, plain, (k3_finseq_1(np__3)=np__3|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_1050, negated_conjecture, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_matrix_1(esk2_0))))).#trainneg
cnf(i_0_1045, negated_conjecture, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X2)),k2_matrix_1(esk2_0))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_1037, negated_conjecture, (r2_hidden(X1,k4_finseq_1(esk2_0))|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k2_matrix_1(esk2_0)))).#trainneg
cnf(i_0_1149, plain, (v1_xreal_0(X1)|~v3_membered(k1_enumset1(X2,X1,X3)))).#trainneg
cnf(i_0_356, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_1166, plain, (r1_xxreal_0(k1_xboole_0,np__3)|r1_xxreal_0(np__3,k1_xboole_0))).#trainneg
cnf(i_0_724, plain, (v1_matrix_1(k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),k3_finseq_4(X1,X8,X9,X10)))|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_726, plain, (m1_matrix_1(k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X1,np__3,np__3)|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_358, plain, (v7_ordinal1(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2))|~v7_ordinal1(k3_xcmplx_0(X3,X2))|~v7_ordinal1(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_367, plain, (k7_finseq_1(k7_finseq_1(k1_tarski(k2_tarski(k2_tarski(np__1,X1),k1_tarski(np__1))),k1_tarski(k2_tarski(k2_tarski(np__1,X2),k1_tarski(np__1)))),k1_tarski(k2_tarski(k2_tarski(X3,np__1),k1_tarski(np__1))))=k11_finseq_1(X1,X2,X3))).#trainneg
cnf(i_0_1283, negated_conjecture, (r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_1217, plain, (k1_funct_1(esk4_5(X1,X2,X3,X3,X4),X3)=X4|X4!=k3_matrix_1(X1,X2,X3,X3)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,k1_matrix_1(X2))|~v7_ordinal1(k1_matrix_1(X2)))).#trainneg
cnf(i_0_1186, plain, (r1_xxreal_0(np__3,X1)|r1_xxreal_0(X1,np__3)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1146, plain, (esk4_5(X1,X2,X3,X3,X4)=k1_funct_1(X2,X3)|X4!=k3_matrix_1(X1,X2,X3,X3)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,k1_matrix_1(X2))|~v7_ordinal1(k1_matrix_1(X2)))).#trainneg
cnf(i_0_1046, negated_conjecture, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k2_matrix_1(esk2_0))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_1152, plain, (k1_nat_1(X2,X1)=k2_xcmplx_0(X2,X1)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1150, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_enumset1(X2,X1,X3)))).#trainneg
cnf(i_0_1326, plain, (~r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1)|~v1_xboole_0(a_1_0_finseq_1(X1)))).#trainneg
cnf(i_0_1168, plain, (m2_subset_1(k1_xboole_0,k1_numbers,k5_numbers)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|~v1_relat_1(k1_xboole_0))).#trainneg
cnf(i_0_1093, negated_conjecture, (v4_relat_1(k2_matrix_1(esk2_0),k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_1209, negated_conjecture, (v3_valued_0(esk9_1(k1_zfmisc_1(k2_matrix_1(esk2_0)))))).#trainneg
cnf(i_0_728, plain, (m1_matrix_1(k11_finseq_1(k11_finseq_1(X2,X3,X4),k11_finseq_1(X5,X6,X7),k3_finseq_4(X1,X8,X9,X10)),X1,np__3,np__3)|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_383, plain, (v1_matrix_1(k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k3_finseq_4(X1,X8,X9,X10)))|v1_xboole_0(X1)|~v7_ordinal1(np__3)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_369, plain, (k7_finseq_1(k7_finseq_1(k1_tarski(k2_tarski(k2_tarski(np__1,X1),k1_tarski(np__1))),k1_tarski(k2_tarski(k2_tarski(np__1,X2),k1_tarski(np__1)))),k1_tarski(k2_tarski(k1_tarski(np__1),k2_tarski(np__1,X3))))=k11_finseq_1(X1,X2,X3))).#trainneg
cnf(i_0_449, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k2_matrix_1(X3))|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(X3)))|~r2_hidden(X1,k4_finseq_1(X3))|~v1_matrix_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_1223, plain, (v7_ordinal1(X1)|~v6_membered(k1_enumset1(X2,X3,X1)))).#trainneg
cnf(i_0_462, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X3,X1)),k2_zfmisc_1(X2,X4)))).#trainneg
cnf(i_0_861, negated_conjecture, (m2_finseq_1(esk4_5(X1,esk2_0,X2,X3,X4),X1)|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_1211, plain, (v7_ordinal1(X1)|~v6_membered(k1_enumset1(X1,X2,X3)))).#trainneg
cnf(i_0_542, plain, (v1_funct_1(k1_tarski(k2_tarski(k1_tarski(np__1),k2_tarski(X1,np__1)))))).#trainneg
cnf(i_0_1327, plain, (~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v1_xboole_0(a_1_0_finseq_1(X1)))).#trainneg
cnf(i_0_371, plain, (k7_finseq_1(k7_finseq_1(k1_tarski(k2_tarski(k2_tarski(np__1,X1),k1_tarski(np__1))),k1_tarski(k2_tarski(k2_tarski(X2,np__1),k1_tarski(np__1)))),k1_tarski(k2_tarski(k2_tarski(np__1,X3),k1_tarski(np__1))))=k11_finseq_1(X1,X2,X3))).#trainneg
cnf(i_0_385, plain, (m1_matrix_1(k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k11_finseq_1(X8,X9,X10)),X1,np__3,np__3)|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_387, plain, (m1_matrix_1(k11_finseq_1(k11_finseq_1(X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k3_finseq_4(X1,X8,X9,X10)),X1,np__3,np__3)|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_373, plain, (k7_finseq_1(k7_finseq_1(k1_tarski(k2_tarski(k2_tarski(np__1,X1),k1_tarski(np__1))),k1_tarski(k2_tarski(k1_tarski(np__1),k2_tarski(np__1,X2)))),k1_tarski(k2_tarski(k2_tarski(np__1,X3),k1_tarski(np__1))))=k11_finseq_1(X1,X2,X3))).#trainneg
cnf(i_0_1510, negated_conjecture, (r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_1429, negated_conjecture, (esk4_5(X1,esk2_0,X2,X3,X4)=k1_funct_1(esk2_0,X2)|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X3,np__3)|~r1_xxreal_0(np__1,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_1212, plain, (v1_xreal_0(X1)|~v3_membered(k1_enumset1(X1,X2,X3)))).#trainneg
cnf(i_0_1344, plain, (k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1247, plain, (X1=np__3|X1=np__2|X1=np__1|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3))).#trainneg
cnf(i_0_1569, plain, (m1_subset_1(k1_xboole_0,k1_numbers)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|~v1_relat_1(k1_xboole_0))).#trainneg
cnf(i_0_1364, plain, (v3_valued_0(k11_finseq_1(X1,X2,X3))|v1_xboole_0(X4)|~v3_membered(X4)|~m1_subset_1(X3,X4)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X4))).#trainneg
cnf(i_0_375, plain, (k7_finseq_1(k7_finseq_1(k1_tarski(k2_tarski(k2_tarski(X1,np__1),k1_tarski(np__1))),k1_tarski(k2_tarski(k2_tarski(np__1,X2),k1_tarski(np__1)))),k1_tarski(k2_tarski(k2_tarski(np__1,X3),k1_tarski(np__1))))=k11_finseq_1(X1,X2,X3))).#trainneg
cnf(i_0_1541, negated_conjecture, (esk4_5(X1,esk2_0,X2,X2,X3)=k1_funct_1(esk2_0,X2)|X3!=k3_matrix_1(X1,esk2_0,X2,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_1543, plain, (esk4_5(X1,X2,k1_matrix_1(X2),k1_matrix_1(X2),X3)=k1_funct_1(X2,k1_matrix_1(X2))|X3!=k3_matrix_1(X1,X2,k1_matrix_1(X2),k1_matrix_1(X2))|~m2_subset_1(k1_matrix_1(X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(X2))|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(k1_matrix_1(X2))|~r2_hidden(k1_matrix_1(X2),k4_finseq_1(X2)))).#trainneg
cnf(i_0_1628, negated_conjecture, (v1_relat_1(esk4_5(X1,esk2_0,X2,X3,X4))|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_1372, negated_conjecture, (v1_relat_1(esk4_5(X1,esk2_0,X2,X2,X3))|X3!=k3_matrix_1(X1,esk2_0,X2,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_1213, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_enumset1(X1,X2,X3)))).#trainneg
cnf(i_0_1417, plain, (v1_xreal_0(esk9_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_1534, plain, (r1_xxreal_0(esk9_1(k5_numbers),np__3)|r1_xxreal_0(np__3,esk9_1(k5_numbers)))).#trainneg
cnf(i_0_1392, negated_conjecture, (k11_finseq_1(X3,X4,X2)=k1_funct_1(esk2_0,np__3)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=X3|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=X4|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|X2!=k3_matrix_1(X1,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X2,X1)|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_377, plain, (k7_finseq_1(k7_finseq_1(k1_tarski(k2_tarski(k1_tarski(np__1),k2_tarski(np__1,X1))),k1_tarski(k2_tarski(k2_tarski(np__1,X2),k1_tarski(np__1)))),k1_tarski(k2_tarski(k2_tarski(np__1,X3),k1_tarski(np__1))))=k11_finseq_1(X1,X2,X3))).#trainneg
cnf(i_0_1475, plain, (v1_matrix_1(k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)))|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1477, plain, (v1_matrix_1(k11_finseq_1(k11_finseq_1(X2,X3,X4),k11_finseq_1(X5,X6,X7),k3_finseq_4(X1,X8,X9,X10)))|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1626, negated_conjecture, (v1_funct_1(esk4_5(X1,esk2_0,X2,X3,X4))|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_1224, plain, (v1_xreal_0(X1)|~v3_membered(k1_enumset1(X2,X3,X1)))).#trainneg
cnf(i_0_1027, negated_conjecture, (v4_relat_1(esk9_1(k1_zfmisc_1(k2_matrix_1(esk2_0))),k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_381, plain, (k1_matrix_1(k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k3_finseq_4(X1,X8,X9,X10)))=np__3|v1_xboole_0(X1)|~v7_ordinal1(np__3)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1094, negated_conjecture, (v4_relat_1(k2_matrix_1(esk2_0),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_382, plain, (k2_matrix_1(k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k3_finseq_4(X1,X8,X9,X10)))=k2_zfmisc_1(k1_enumset1(np__1,np__2,np__3),k1_enumset1(np__1,np__2,np__3))|v1_xboole_0(X1)|~v7_ordinal1(np__3)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_384, plain, (m2_finseq_1(k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k3_finseq_4(X1,X8,X9,X10)),k3_finseq_2(X1))|v1_xboole_0(X1)|~v7_ordinal1(np__3)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1370, negated_conjecture, (v1_funct_1(esk4_5(X1,esk2_0,X2,X2,X3))|X3!=k3_matrix_1(X1,esk2_0,X2,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_1627, negated_conjecture, (v1_finseq_1(esk4_5(X1,esk2_0,X2,X3,X4))|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_1371, negated_conjecture, (v1_finseq_1(esk4_5(X1,esk2_0,X2,X2,X3))|X3!=k3_matrix_1(X1,esk2_0,X2,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_1225, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_enumset1(X2,X3,X1)))).#trainneg
cnf(i_0_1328, plain, (~r1_xxreal_0(np__3,X1)|~v7_ordinal1(X1)|~v1_xboole_0(a_1_0_finseq_1(X1)))).#trainneg
cnf(i_0_388, plain, (esk8_2(esk3_4(X1,X2,X3,a_1_0_finseq_1(X4)),X4)=esk3_4(X1,X2,X3,a_1_0_finseq_1(X4))|esk3_4(X1,X2,X3,a_1_0_finseq_1(X4))=X1|esk3_4(X1,X2,X3,a_1_0_finseq_1(X4))=X2|esk3_4(X1,X2,X3,a_1_0_finseq_1(X4))=X3|a_1_0_finseq_1(X4)=k1_enumset1(X1,X2,X3)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_1238, plain, (v1_finset_1(np__3)|~m1_subset_1(k11_finseq_1(X1,X2,X3),k5_numbers))).#trainneg
cnf(i_0_1580, plain, (v1_matrix_1(k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k11_finseq_1(X8,X9,X10)))|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_389, plain, (v7_ordinal1(esk3_4(X1,X2,X3,k5_numbers))|esk3_4(X1,X2,X3,k5_numbers)=X1|esk3_4(X1,X2,X3,k5_numbers)=X2|esk3_4(X1,X2,X3,k5_numbers)=X3|k1_enumset1(X1,X2,X3)=k5_numbers)).#trainneg
cnf(i_0_390, plain, (k3_finseq_1(X1)=X2|v1_xboole_0(X3)|esk3_4(X5,X6,X7,k2_finseq_1(X2))=X5|esk3_4(X5,X6,X7,k2_finseq_1(X2))=X6|esk3_4(X5,X6,X7,k2_finseq_1(X2))=X7|k2_finseq_1(X2)=k1_enumset1(X5,X6,X7)|X1!=k1_funct_1(X4,esk3_4(X5,X6,X7,k2_finseq_1(X2)))|~m2_finseq_1(X1,X3)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(esk3_4(X5,X6,X7,k2_finseq_1(X2)))|~m1_matrix_1(X4,X3,X2,X2))).#trainneg
cnf(i_0_1602, plain, (X1=k3_matrix_1(X2,X3,X4,X5)|X1!=k1_funct_1(X6,X5)|X6!=k1_funct_1(X3,X4)|~v1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X6,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(X5,k2_finseq_1(k1_matrix_1(X3)))|~r2_hidden(X4,k4_finseq_1(X3)))).#trainneg
cnf(i_0_1386, negated_conjecture, (r2_hidden(X1,k4_finseq_1(esk2_0))|~m1_subset_1(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(esk2_0)))).#trainneg
cnf(i_0_1333, plain, (m2_finseq_1(esk4_5(X1,X2,X3,X4,X5),X1)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,k1_matrix_1(X2))|~v7_ordinal1(k1_matrix_1(X2)))).#trainneg
cnf(i_0_1616, negated_conjecture, (r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X2,X1)),k2_matrix_1(esk2_0)))).#trainneg
cnf(i_0_1382, negated_conjecture, (r2_hidden(X1,k4_finseq_1(esk2_0))|~r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X1)),k2_matrix_1(esk2_0)))).#trainneg
cnf(i_0_1565, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_enumset1(X2,X1,X3)))).#trainneg
cnf(i_0_1567, plain, (~r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1)|~v1_xboole_0(k2_finseq_1(X1)))).#trainneg
cnf(i_0_392, plain, (m1_subset_1(esk3_4(X1,X2,X3,X4),X4)|esk3_4(X1,X2,X3,X4)=X1|esk3_4(X1,X2,X3,X4)=X2|esk3_4(X1,X2,X3,X4)=X3|X4=k1_enumset1(X1,X2,X3))).#trainneg
cnf(i_0_1708, plain, (k1_xboole_0=np__1|k1_xboole_0=np__2|k1_xboole_0=np__3|~r1_xxreal_0(np__1,k1_xboole_0)|~r1_xxreal_0(k1_xboole_0,np__3)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|~v1_relat_1(k1_xboole_0))).#trainneg
cnf(i_0_1965, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1800, plain, (v1_matrix_1(k11_finseq_1(k11_finseq_1(X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)))|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_400, plain, (X1=k3_matrix_1(X2,X3,X4,X5)|v1_xboole_0(k2_matrix_1(X3))|X1!=k1_funct_1(X6,X5)|X6!=k1_funct_1(X3,X4)|~v1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X6,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(k2_tarski(k2_tarski(X4,X5),k1_tarski(X4)),k2_matrix_1(X3)))).#trainneg
cnf(i_0_405, plain, (esk4_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|v1_xboole_0(k2_matrix_1(X2))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(k2_tarski(k2_tarski(X3,X4),k1_tarski(X3)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_1694, negated_conjecture, (r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_1697, negated_conjecture, (esk4_5(X1,esk2_0,X2,X3,X4)=k1_funct_1(esk2_0,X2)|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X3,np__3)|~r1_xxreal_0(np__1,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_1698, negated_conjecture, (esk4_5(X1,esk2_0,X2,X3,X4)=k1_funct_1(esk2_0,X2)|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X3,np__3)|~r1_xxreal_0(np__1,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_1747, negated_conjecture, (esk4_5(X1,esk2_0,X2,X2,X3)=k1_funct_1(esk2_0,X2)|X3!=k3_matrix_1(X1,esk2_0,X2,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~m1_subset_1(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_1765, negated_conjecture, (esk4_5(X1,esk2_0,X2,np__3,X3)=k1_funct_1(esk2_0,X2)|X3!=k3_matrix_1(X1,esk2_0,X2,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_1756, negated_conjecture, (esk4_5(X1,esk2_0,np__3,np__3,X2)=k1_funct_1(esk2_0,np__3)|X2!=k3_matrix_1(X1,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X2,X1)|~r2_hidden(np__3,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_1774, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_enumset1(X1,X2,X3)))).#trainneg
cnf(i_0_676, plain, (v1_finseq_1(k1_tarski(k2_tarski(k1_tarski(np__1),k2_tarski(X1,np__1)))))).#trainneg
cnf(i_0_1702, plain, (k1_xboole_0!=np__3|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|~v1_relat_1(k1_xboole_0))).#trainneg
cnf(i_0_1776, negated_conjecture, (k11_finseq_1(k1_funct_1(k1_funct_1(esk2_0,np__3),np__1),X1,X2)=k1_funct_1(esk2_0,np__3)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=X1|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|X2!=k3_matrix_1(X3,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X3))|~m1_subset_1(X2,X3)|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_410, plain, (m2_finseq_1(esk4_5(X1,X2,X3,X4,X5),X1)|v1_xboole_0(k2_matrix_1(X2))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(k2_tarski(k2_tarski(X3,X4),k1_tarski(X3)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_1831, plain, (m1_subset_1(k1_xboole_0,k1_enumset1(np__1,np__2,np__3))|~r1_xxreal_0(np__1,k1_xboole_0)|~r1_xxreal_0(k1_xboole_0,np__3)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|~v1_relat_1(k1_xboole_0))).#trainneg
cnf(i_0_417, plain, (m1_subset_1(k2_tarski(k1_tarski(esk6_3(X1,X2,X3)),k2_tarski(esk6_3(X1,X2,X3),esk7_3(X1,X2,X3))),k2_matrix_1(X2))|X2=X3|v1_xboole_0(X1)|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_418, plain, (m1_subset_1(k3_matrix_1(X1,k3_finseq_4(k3_finseq_2(X1),X2,X3,X4),X5,X6),X1)|v1_xboole_0(k3_finseq_2(X1))|~v1_matrix_1(k3_finseq_4(k3_finseq_2(X1),X2,X3,X4))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m1_subset_1(X4,k3_finseq_2(X1))|~m1_subset_1(X3,k3_finseq_2(X1))|~m1_subset_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_1948, negated_conjecture, (r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X1)),k2_matrix_1(esk2_0)))).#trainneg
cnf(i_0_1711, plain, (X1=np__1|X1=np__2|X1=np__3|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1837, plain, (m1_subset_1(X1,k1_enumset1(np__1,np__2,np__3))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1932, negated_conjecture, (r2_hidden(X1,k4_finseq_1(esk2_0))|~m1_subset_1(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k2_matrix_1(esk2_0)))).#trainneg
cnf(i_0_1956, negated_conjecture, (r2_hidden(X1,k4_finseq_1(esk2_0))|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X2,X1)),k2_matrix_1(esk2_0)))).#trainneg
cnf(i_0_1892, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_enumset1(X2,X3,X1)))).#trainneg
cnf(i_0_435, plain, (k3_finseq_1(X1)=k1_xboole_0|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1568, plain, (~r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1640, plain, (~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1903, plain, (v1_matrix_1(k11_finseq_1(k11_finseq_1(X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k11_finseq_1(X8,X9,X10)))|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1582, plain, (v1_matrix_1(k11_finseq_1(k11_finseq_1(X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k3_finseq_4(X1,X8,X9,X10)))|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_440, plain, (r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2)))|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X3,X1)),k2_matrix_1(X2))|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_445, plain, (r2_hidden(X1,X2)|v1_xboole_0(k2_zfmisc_1(X4,X2))|~m1_subset_1(k2_tarski(k1_tarski(X3),k2_tarski(X3,X1)),k2_zfmisc_1(X4,X2)))).#trainneg
cnf(i_0_1750, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__3),np__3)=X2|X2!=k3_matrix_1(X1,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X2,X1)|~r2_hidden(np__3,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_1950, negated_conjecture, (r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(k2_tarski(k1_tarski(X1),k2_tarski(X2,X1)),k2_matrix_1(esk2_0)))).#trainneg
cnf(i_0_1925, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,X4)|X1!=k1_funct_1(X5,X4)|X5!=k1_funct_1(esk2_0,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X4,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X3,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_1938, plain, (v1_relat_1(esk4_5(X1,X2,X3,X4,X5))|X5!=k3_matrix_1(X1,X2,X3,X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X4,k1_matrix_1(X2))|~r1_xxreal_0(np__1,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(k1_matrix_1(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2)))).#trainneg
cnf(i_0_1355, negated_conjecture, (esk4_5(X1,esk2_0,X2,X3,X4)=k1_funct_1(esk2_0,X2)|v1_xboole_0(k4_finseq_1(esk2_0))|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_496, plain, (v7_ordinal1(k1_card_1(X1))|~v1_finset_1(k1_card_1(X1)))).#trainneg
cnf(i_0_854, negated_conjecture, (v7_ordinal1(k3_matrix_1(esk1_0,esk2_0,X1,X2))|~v6_membered(esk1_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_447, plain, (m1_subset_1(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_590, plain, (k2_finseq_1(X1)=k1_xboole_0|~v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1994, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_466, plain, (r2_hidden(X1,X2)|v1_xboole_0(k2_zfmisc_1(X2,X4))|~m1_subset_1(k2_tarski(k1_tarski(X1),k2_tarski(X1,X3)),k2_zfmisc_1(X2,X4)))).#trainneg
cnf(i_0_469, plain, (m1_subset_1(k2_tarski(k2_tarski(X1,X2),k1_tarski(X2)),k2_zfmisc_1(X3,X4))|~r2_hidden(X1,X4)|~r2_hidden(X2,X3))).#trainneg
cnf(i_0_2079, negated_conjecture, (esk4_5(X1,esk2_0,X2,np__3,X3)=k1_funct_1(esk2_0,X2)|X3!=k3_matrix_1(X1,esk2_0,X2,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~m1_subset_1(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_2082, negated_conjecture, (esk4_5(X1,esk2_0,np__3,np__3,X2)=k1_funct_1(esk2_0,np__3)|X2!=k3_matrix_1(X1,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_2265, plain, (v7_ordinal1(k1_card_1(X1))|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_2199, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_2229, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__3),np__3)=X1|X1!=k3_matrix_1(X2,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_477, plain, (m1_subset_1(esk7_3(X1,X2,X3),k5_numbers)|X2=X3|v1_xboole_0(X1)|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_1991, plain, (k1_xboole_0=np__3|k1_xboole_0=np__2|k1_xboole_0=np__1|r1_xxreal_0(np__3,k1_xboole_0)|~r1_xxreal_0(np__1,k1_xboole_0)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|~v1_relat_1(k1_xboole_0))).#trainneg
cnf(i_0_478, plain, (m1_subset_1(esk6_3(X1,X2,X3),k5_numbers)|X2=X3|v1_xboole_0(X1)|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_479, plain, (m1_subset_1(esk5_2(X1,X2),k5_numbers)|X1=X2|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2080, negated_conjecture, (esk4_5(X1,esk2_0,X2,np__3,X3)=k1_funct_1(esk2_0,X2)|X3!=k3_matrix_1(X1,esk2_0,X2,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_2237, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,np__2)|X4!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X4,np__2)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_2236, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,np__3)|X4!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X4,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_2238, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,np__1)|X4!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X4,np__1)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_2235, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,X4)|X5!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X5,X4)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(esk2_0))|~m1_subset_1(X4,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_2029, plain, (esk4_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|v1_xboole_0(k2_matrix_1(X2))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(k2_tarski(k1_tarski(X3),k2_tarski(X3,X4)),k2_matrix_1(X2))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_836, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_480, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2184, plain, (m2_subset_1(k1_xboole_0,k1_numbers,k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2139, plain, (r1_xxreal_0(esk9_1(k1_enumset1(np__1,np__2,np__3)),np__3))).#trainneg
cnf(i_0_2148, plain, (r1_xxreal_0(np__1,esk9_1(k1_enumset1(np__1,np__2,np__3))))).#trainneg
cnf(i_0_482, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_487, plain, (a_1_0_finseq_1(X1)=k1_enumset1(X2,X3,X4)|esk3_4(X2,X3,X4,a_1_0_finseq_1(X1))!=X4|~m2_subset_1(esk3_4(X2,X3,X4,a_1_0_finseq_1(X1)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk3_4(X2,X3,X4,a_1_0_finseq_1(X1)))|~r1_xxreal_0(esk3_4(X2,X3,X4,a_1_0_finseq_1(X1)),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2230, negated_conjecture, (r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k2_matrix_1(esk2_0)))).#trainneg
cnf(i_0_2239, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,X4)|X5!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X5,X4)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(esk2_0))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,np__3))).#trainneg
cnf(i_0_1330, negated_conjecture, (m2_finseq_1(esk4_5(X1,esk2_0,X2,X3,X4),X1)|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_2267, plain, (v7_ordinal1(k1_card_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_2031, negated_conjecture, (r2_hidden(esk9_1(k4_finseq_1(esk2_0)),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_488, plain, (a_1_0_finseq_1(X1)=k1_enumset1(X2,X3,X4)|esk3_4(X2,X3,X4,a_1_0_finseq_1(X1))!=X3|~m2_subset_1(esk3_4(X2,X3,X4,a_1_0_finseq_1(X1)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk3_4(X2,X3,X4,a_1_0_finseq_1(X1)))|~r1_xxreal_0(esk3_4(X2,X3,X4,a_1_0_finseq_1(X1)),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2296, plain, (r2_hidden(X1,k1_xboole_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1895, plain, (~r1_xxreal_0(np__3,X1)|~v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_489, plain, (a_1_0_finseq_1(X1)=k1_enumset1(X2,X3,X4)|esk3_4(X2,X3,X4,a_1_0_finseq_1(X1))!=X2|~m2_subset_1(esk3_4(X2,X3,X4,a_1_0_finseq_1(X1)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk3_4(X2,X3,X4,a_1_0_finseq_1(X1)))|~r1_xxreal_0(esk3_4(X2,X3,X4,a_1_0_finseq_1(X1)),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_491, plain, (m1_subset_1(X1,a_1_0_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2541, plain, (v7_ordinal1(k1_card_1(k1_enumset1(X1,X2,X3))))).#trainneg
cnf(i_0_2506, negated_conjecture, (r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(esk2_0)))).#trainneg
cnf(i_0_2360, negated_conjecture, (k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,np__3)=k1_funct_1(k1_funct_1(esk2_0,X1),np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(k1_enumset1(np__1,np__2,np__3)))|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0))|~r1_xxreal_0(np__1,k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,np__3))|~r1_xxreal_0(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,np__3),np__3)|~m1_subset_1(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,np__3),k5_numbers))).#trainneg
cnf(i_0_2408, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,np__2)|X3!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X3,np__2)|~m2_finseq_1(X3,esk1_0)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_2473, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2448, plain, (v4_relat_1(k1_xboole_0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_492, plain, (esk8_2(X1,X2)=X1|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_2445, plain, (v1_relat_1(k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2494, plain, (~v7_ordinal1(esk9_1(k1_enumset1(np__1,np__2,np__3)))|~v1_xboole_0(esk9_1(k1_enumset1(np__1,np__2,np__3))))).#trainneg
cnf(i_0_1485, plain, (m1_matrix_1(k11_finseq_1(k11_finseq_1(X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X1,np__3,np__3)|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_502, plain, (k11_finseq_1(X1,X2,X3)=X4|r1_xxreal_0(esk5_2(k11_finseq_1(X1,X2,X3),X4),np__3)|k3_finseq_1(X4)!=np__3|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4))).#trainneg
cnf(i_0_508, plain, (r2_hidden(X1,X2)|v1_xboole_0(k2_zfmisc_1(X2,X4))|~m1_subset_1(k2_tarski(k2_tarski(X3,X1),k1_tarski(X1)),k2_zfmisc_1(X2,X4)))).#trainneg
cnf(i_0_2409, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,np__3)|X3!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X3,np__3)|~m2_finseq_1(X3,esk1_0)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_2410, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,np__1)|X3!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X3,np__1)|~m2_finseq_1(X3,esk1_0)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_2374, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__3),np__3)=k3_matrix_1(X1,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(k3_matrix_1(X1,esk2_0,np__3,np__3),X1))).#trainneg
cnf(i_0_2543, plain, (v7_ordinal1(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_511, plain, (k1_funct_1(esk4_5(X1,X2,X3,X4,X5),X4)=X5|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X4,X3)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_2450, plain, (v5_relat_1(k1_xboole_0,X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2550, negated_conjecture, (r1_xxreal_0(np__1,esk9_1(k4_finseq_1(esk2_0))))).#trainneg
cnf(i_0_2553, negated_conjecture, (r1_xxreal_0(esk9_1(k4_finseq_1(esk2_0)),np__3))).#trainneg
cnf(i_0_1663, plain, (m1_matrix_1(k11_finseq_1(k11_finseq_1(X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k11_finseq_1(X8,X9,X10)),X1,np__3,np__3)|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1140, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X2),X2)=X3|X3!=k3_matrix_1(X1,esk2_0,X2,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_1710, plain, (k3_finseq_1(X1)=np__1|k3_finseq_1(X1)=np__2|k3_finseq_1(X1)=np__3|~r1_xxreal_0(np__1,k3_finseq_1(X1))|~r1_xxreal_0(k3_finseq_1(X1),np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2449, plain, (v3_valued_0(k1_xboole_0)|~v3_membered(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2440, plain, (v4_relat_1(esk9_1(k1_zfmisc_1(k1_xboole_0)),X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2547, negated_conjecture, (esk9_1(k4_finseq_1(esk2_0))=np__3|esk9_1(k4_finseq_1(esk2_0))=np__2|esk9_1(k4_finseq_1(esk2_0))=np__1)).#trainneg
cnf(i_0_2552, negated_conjecture, (m2_subset_1(esk9_1(k4_finseq_1(esk2_0)),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2546, negated_conjecture, (m1_subset_1(esk9_1(k4_finseq_1(esk2_0)),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_515, plain, (k1_funct_1(esk4_5(X1,X2,X3,X4,X5),X4)=X5|v1_xboole_0(k2_matrix_1(X2))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(k2_tarski(k1_tarski(X3),k2_tarski(X3,X4)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_516, plain, (k1_funct_1(esk4_5(X1,X2,esk6_3(X3,X2,X4),esk7_3(X3,X2,X4),X5),esk7_3(X3,X2,X4))=X5|X2=X4|v1_xboole_0(X3)|X5!=k3_matrix_1(X1,X2,esk6_3(X3,X2,X4),esk7_3(X3,X2,X4))|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(esk7_3(X3,X2,X4))|~v7_ordinal1(esk6_3(X3,X2,X4))|k3_finseq_1(X2)!=k3_finseq_1(X4)|k1_matrix_1(X2)!=k1_matrix_1(X4)|~v1_matrix_1(X4)|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X2,k3_finseq_2(X3)))).#trainneg
cnf(i_0_1201, negated_conjecture, (esk4_5(X1,esk2_0,X2,np__3,X3)=k11_finseq_1(X4,X5,X3)|k1_funct_1(esk4_5(X1,esk2_0,X2,np__3,X3),np__1)!=X4|k1_funct_1(esk4_5(X1,esk2_0,X2,np__3,X3),np__2)!=X5|k3_finseq_1(esk4_5(X1,esk2_0,X2,np__3,X3))!=np__3|~v1_finseq_1(esk4_5(X1,esk2_0,X2,np__3,X3))|~v1_funct_1(esk4_5(X1,esk2_0,X2,np__3,X3))|~v1_relat_1(esk4_5(X1,esk2_0,X2,np__3,X3))|X3!=k3_matrix_1(X1,esk2_0,X2,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_524, plain, (X1=X2|X1=X3|X1=X4|v1_xboole_0(k1_enumset1(X4,X3,X2))|~m1_subset_1(X1,k1_enumset1(X4,X3,X2)))).#trainneg
cnf(i_0_1153, plain, (m2_subset_1(k2_xcmplx_0(X2,X1),k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_643, negated_conjecture, (r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X2,X1)),k2_matrix_1(esk2_0)))).#trainneg
cnf(i_0_859, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,X4)|X1!=k1_funct_1(X5,X4)|X5!=k1_funct_1(esk2_0,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X4,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_2443, plain, (v1_xboole_0(esk9_1(k1_zfmisc_1(k1_xboole_0)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2567, plain, (k1_xboole_0=np__1|k1_xboole_0=np__2|~r1_xxreal_0(np__1,k1_xboole_0)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|~v1_relat_1(k1_xboole_0))).#trainneg
cnf(i_0_522, plain, (esk4_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|v1_xboole_0(k2_matrix_1(X2))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(k2_tarski(k2_tarski(X4,X3),k1_tarski(X3)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_2493, plain, (esk9_1(k1_enumset1(np__1,np__2,np__3))=np__3|esk9_1(k1_enumset1(np__1,np__2,np__3))=np__2|esk9_1(k1_enumset1(np__1,np__2,np__3))=np__1|~m1_subset_1(esk9_1(k1_enumset1(np__1,np__2,np__3)),k5_numbers))).#trainneg
cnf(i_0_2439, plain, (v1_relat_1(esk9_1(k1_zfmisc_1(k1_xboole_0)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1875, plain, (m1_subset_1(k3_matrix_1(X1,k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k3_finseq_4(X1,X8,X9,X10)),X11,X12),X1)|v1_xboole_0(X1)|~v1_matrix_1(k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k3_finseq_4(X1,X8,X9,X10)))|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_523, plain, (esk3_4(X1,X2,X3,k1_enumset1(X4,X5,X6))=X6|esk3_4(X1,X2,X3,k1_enumset1(X4,X5,X6))=X5|esk3_4(X1,X2,X3,k1_enumset1(X4,X5,X6))=X4|esk3_4(X1,X2,X3,k1_enumset1(X4,X5,X6))=X1|esk3_4(X1,X2,X3,k1_enumset1(X4,X5,X6))=X2|esk3_4(X1,X2,X3,k1_enumset1(X4,X5,X6))=X3|k1_enumset1(X4,X5,X6)=k1_enumset1(X1,X2,X3))).#trainneg
cnf(i_0_2647, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__3),np__3)=k3_matrix_1(esk1_0,esk2_0,np__3,np__3))).#trainneg
cnf(i_0_2639, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,np__3)|X3!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X3,np__3)|~m2_finseq_1(X3,esk1_0)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~m1_subset_1(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_2597, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,np__2)|X3!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X3,np__2)|~m2_finseq_1(X3,esk1_0)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~m1_subset_1(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_2598, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,np__2)|X3!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X3,np__2)|~m2_finseq_1(X3,esk1_0)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_2612, plain, (v1_relat_1(k1_xboole_0))).#trainneg
cnf(i_0_2675, plain, (v5_relat_1(k1_xboole_0,X1))).#trainneg
cnf(i_0_2824, plain, (v1_xboole_0(esk9_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_2724, plain, (v3_valued_0(k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2681, negated_conjecture, (~v7_ordinal1(esk9_1(k4_finseq_1(esk2_0)))|~v1_xboole_0(esk9_1(k4_finseq_1(esk2_0))))).#trainneg
cnf(i_0_2618, plain, (~v1_xboole_0(esk9_1(k1_enumset1(np__1,np__2,np__3)))|~v6_membered(k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_2642, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,np__1)|X3!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X3,np__1)|~m2_finseq_1(X3,esk1_0)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~m1_subset_1(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_2640, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,np__3)|X3!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X3,np__3)|~m2_finseq_1(X3,esk1_0)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_2701, plain, (m1_subset_1(X1,k5_numbers)|~m1_subset_1(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_2643, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,np__1)|X3!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X3,np__1)|~m2_finseq_1(X3,esk1_0)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_2841, plain, (v1_relat_1(esk9_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_2728, negated_conjecture, (esk9_1(k4_finseq_1(esk2_0))=np__1|esk9_1(k4_finseq_1(esk2_0))=np__3|np__2!=np__1)).#trainneg
cnf(i_0_556, plain, (k1_funct_1(esk4_5(X1,X2,X3,X4,X5),X4)=X5|v1_xboole_0(k2_matrix_1(X2))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(k2_tarski(k2_tarski(X4,X3),k1_tarski(X3)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_2718, plain, (k1_xboole_0=np__3|k1_xboole_0=np__2|k1_xboole_0=np__1|~r1_xxreal_0(np__1,k1_xboole_0)|~r1_xxreal_0(k1_xboole_0,np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2617, plain, (~v1_xboole_0(esk9_1(k1_enumset1(np__1,np__2,np__3)))|~m1_subset_1(esk9_1(k1_enumset1(np__1,np__2,np__3)),k5_numbers))).#trainneg
cnf(i_0_2739, negated_conjecture, (m1_subset_1(esk9_1(k4_finseq_1(esk2_0)),k1_numbers))).#trainneg
cnf(i_0_2874, plain, (m1_subset_1(k3_matrix_1(X1,k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12),X1)|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_2875, plain, (m1_subset_1(k3_matrix_1(X1,k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),k3_finseq_4(X1,X8,X9,X10)),X11,X12),X1)|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_563, plain, (esk4_5(X1,X2,esk6_3(X3,X2,X4),esk7_3(X3,X2,X4),X5)=k1_funct_1(X2,esk6_3(X3,X2,X4))|X2=X4|v1_xboole_0(X3)|X5!=k3_matrix_1(X1,X2,esk6_3(X3,X2,X4),esk7_3(X3,X2,X4))|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(esk7_3(X3,X2,X4))|~v7_ordinal1(esk6_3(X3,X2,X4))|k3_finseq_1(X2)!=k3_finseq_1(X4)|k1_matrix_1(X2)!=k1_matrix_1(X4)|~v1_matrix_1(X4)|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X2,k3_finseq_2(X3)))).#trainneg
cnf(i_0_3050, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,np__2)|X3!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X3,np__2)|~m2_finseq_1(X3,esk1_0)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3))).#trainneg
cnf(i_0_3046, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,np__2)|X3!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X3,np__2)|~m2_finseq_1(X3,esk1_0)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_3057, plain, (v3_valued_0(k1_xboole_0))).#trainneg
cnf(i_0_2649, plain, (m1_subset_1(k1_card_1(X1),k5_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3051, plain, (esk9_1(k1_zfmisc_1(k1_xboole_0))=k1_xboole_0)).#trainneg
cnf(i_0_2229, negated_conjecture, (k3_matrix_1(esk1_0,esk2_0,np__3,np__3)=X1|X1!=k3_matrix_1(X2,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_564, plain, (r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2)))|~r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X3)),k2_matrix_1(X2))|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1168, plain, (m2_subset_1(k1_xboole_0,k1_numbers,k5_numbers)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_1702, plain, (k1_xboole_0!=np__3|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_3068, negated_conjecture, (~v1_xboole_0(esk9_1(k4_finseq_1(esk2_0)))|~v6_membered(k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_3155, plain, (m1_subset_1(k3_matrix_1(X1,k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12),X1)|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_3156, plain, (m1_subset_1(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12),X1)|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_569, plain, (r2_hidden(X1,X2)|v1_xboole_0(k2_zfmisc_1(X4,X2))|~m1_subset_1(k2_tarski(k2_tarski(X1,X3),k1_tarski(X3)),k2_zfmisc_1(X4,X2)))).#trainneg
cnf(i_0_576, plain, (v7_ordinal1(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)))|~v6_membered(k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_3095, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,np__1)|X3!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X3,np__1)|~m2_finseq_1(X3,esk1_0)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3))).#trainneg
cnf(i_0_3079, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,np__3)|X3!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X3,np__3)|~m2_finseq_1(X3,esk1_0)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3))).#trainneg
cnf(i_0_3075, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,np__3)|X3!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X3,np__3)|~m2_finseq_1(X3,esk1_0)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_3091, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,np__1)|X3!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X3,np__1)|~m2_finseq_1(X3,esk1_0)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_2374, negated_conjecture, (k3_matrix_1(esk1_0,esk2_0,np__3,np__3)=k3_matrix_1(X1,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(k3_matrix_1(X1,esk2_0,np__3,np__3),X1))).#trainneg
cnf(i_0_3028, negated_conjecture, (X1=k1_funct_1(esk2_0,X2)|k1_funct_1(X1,esk5_2(X1,k1_funct_1(esk2_0,X2)))!=k3_matrix_1(esk1_0,esk2_0,X2,esk5_2(X1,k1_funct_1(esk2_0,X2)))|k3_finseq_1(X1)!=k3_finseq_1(k1_funct_1(esk2_0,X2))|~v1_finseq_1(k1_funct_1(esk2_0,X2))|~v1_finseq_1(X1)|~v1_funct_1(k1_funct_1(esk2_0,X2))|~v1_funct_1(X1)|~v1_relat_1(k1_funct_1(esk2_0,X2))|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,k1_funct_1(esk2_0,X2)))|~v7_ordinal1(X2)|~r2_hidden(esk5_2(X1,k1_funct_1(esk2_0,X2)),k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_3017, negated_conjecture, (k1_funct_1(esk2_0,X1)=k11_finseq_1(X2,X3,X4)|k3_matrix_1(esk1_0,esk2_0,X1,np__3)!=X4|k1_funct_1(k1_funct_1(esk2_0,X1),np__1)!=X2|k1_funct_1(k1_funct_1(esk2_0,X1),np__2)!=X3|k3_finseq_1(k1_funct_1(esk2_0,X1))!=np__3|~v1_finseq_1(k1_funct_1(esk2_0,X1))|~v1_funct_1(k1_funct_1(esk2_0,X1))|~v1_relat_1(k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k4_finseq_1(esk2_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3092, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|X2!=k1_funct_1(esk2_0,np__3)|X1!=k1_funct_1(X2,np__1)|~m2_finseq_1(X2,esk1_0)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_3076, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__3,np__3)|X2!=k1_funct_1(esk2_0,np__3)|X1!=k1_funct_1(X2,np__3)|~m2_finseq_1(X2,esk1_0)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_3274, plain, (v1_xreal_0(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_989, plain, (k3_finseq_1(k1_xboole_0)=k1_xboole_0|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_2567, plain, (k1_xboole_0=np__2|k1_xboole_0=np__1|~r1_xxreal_0(np__1,k1_xboole_0)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_577, plain, (v1_xreal_0(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)))|~v3_membered(k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_3067, negated_conjecture, (~v1_xboole_0(esk9_1(k4_finseq_1(esk2_0)))|~m1_subset_1(esk9_1(k4_finseq_1(esk2_0)),k5_numbers))).#trainneg
cnf(i_0_1569, plain, (m1_subset_1(k1_xboole_0,k1_numbers)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_3351, plain, (m1_subset_1(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12),X1)|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_578, plain, (v1_xcmplx_0(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)))|~v1_membered(k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_579, plain, (m1_subset_1(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(X3))|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(X3)))|~r2_hidden(X1,k4_finseq_1(X3))|~v1_matrix_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_3456, negated_conjecture, (k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,np__3,np__3)=k3_matrix_1(esk1_0,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(k1_enumset1(np__1,np__2,np__3)))|~r1_xxreal_0(np__1,k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,np__3,np__3))|~r1_xxreal_0(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,np__3,np__3),np__3)|~m1_subset_1(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,np__3,np__3),k5_numbers))).#trainneg
cnf(i_0_3272, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,np__2)|X3!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X3,np__2)|~m2_finseq_1(X3,esk1_0)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3413, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,np__1)|X3!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X3,np__1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(X3,esk1_0)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3116, negated_conjecture, (v1_xreal_0(esk9_1(k4_finseq_1(esk2_0))))).#trainneg
cnf(i_0_1831, plain, (m1_subset_1(k1_xboole_0,k1_enumset1(np__1,np__2,np__3))|~r1_xxreal_0(np__1,k1_xboole_0)|~r1_xxreal_0(k1_xboole_0,np__3)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_3115, negated_conjecture, (v7_ordinal1(esk9_1(k4_finseq_1(esk2_0)))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_3344, plain, (v7_ordinal1(k3_matrix_1(X1,k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v6_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_3345, plain, (v1_xreal_0(k3_matrix_1(X1,k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_3423, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__2,np__3)|X2!=k1_funct_1(esk2_0,np__2)|X1!=k1_funct_1(X2,np__3)|~m2_finseq_1(X2,esk1_0)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_3460, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|X1!=k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)|~m2_finseq_1(k1_funct_1(esk2_0,np__3),esk1_0)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_3426, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,np__3)|X3!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X3,np__3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(X3,esk1_0)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3259, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|X2!=k1_funct_1(esk2_0,np__3)|X1!=k1_funct_1(X2,np__2)|~m2_finseq_1(X2,esk1_0)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_2471, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3275, plain, (v1_xcmplx_0(k1_card_1(X1))|~v1_membered(k5_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3096, negated_conjecture, (esk9_1(k4_finseq_1(esk2_0))=np__1|np__3!=np__1|np__2!=np__1)).#trainneg
cnf(i_0_3464, negated_conjecture, (k11_finseq_1(k1_funct_1(esk2_0,np__1),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|~m1_subset_1(np__1,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__1)))).#trainneg
cnf(i_0_3117, negated_conjecture, (v1_xcmplx_0(esk9_1(k4_finseq_1(esk2_0)))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_2740, negated_conjecture, (m1_subset_1(esk9_1(k4_finseq_1(esk2_0)),k5_numbers))).#trainneg
cnf(i_0_3346, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v1_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_3380, plain, (v7_ordinal1(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v6_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_598, plain, (r1_xxreal_0(X1,k1_funct_1(X2,X3))|r1_xxreal_0(k1_funct_1(X2,X3),X1)|~v7_ordinal1(X1)|~v2_valued_0(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_3027, negated_conjecture, (k1_funct_1(esk2_0,X1)=X2|k3_matrix_1(esk1_0,esk2_0,X1,esk5_2(k1_funct_1(esk2_0,X1),X2))!=k1_funct_1(X2,esk5_2(k1_funct_1(esk2_0,X1),X2))|k3_finseq_1(k1_funct_1(esk2_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(k1_funct_1(esk2_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k1_funct_1(esk2_0,X1))|~v1_relat_1(X2)|~v1_relat_1(k1_funct_1(esk2_0,X1))|~v7_ordinal1(esk5_2(k1_funct_1(esk2_0,X1),X2))|~v7_ordinal1(X1)|~r2_hidden(esk5_2(k1_funct_1(esk2_0,X1),X2),k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_3257, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|X2!=k1_funct_1(esk2_0,np__2)|X1!=k1_funct_1(X2,np__2)|~m2_finseq_1(X2,esk1_0)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_3410, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|X2!=k1_funct_1(esk2_0,np__2)|X1!=k1_funct_1(X2,np__1)|~m2_finseq_1(X2,esk1_0)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_3012, negated_conjecture, (k1_funct_1(esk2_0,X1)=k11_finseq_1(X2,X3,X4)|k3_matrix_1(esk1_0,esk2_0,X1,np__3)!=X4|k1_funct_1(k1_funct_1(esk2_0,X1),np__1)!=X2|k1_funct_1(k1_funct_1(esk2_0,X1),np__2)!=X3|k3_finseq_1(k1_funct_1(esk2_0,X1))!=np__3|~v1_finseq_1(k1_funct_1(esk2_0,X1))|~v1_funct_1(k1_funct_1(esk2_0,X1))|~v1_relat_1(k1_funct_1(esk2_0,X1))|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_3078, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__1,np__3)|X2!=k1_funct_1(esk2_0,np__1)|X1!=k1_funct_1(X2,np__3)|~m2_finseq_1(X2,esk1_0)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_2544, plain, (v7_ordinal1(k1_card_1(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_603, plain, (a_1_0_finseq_1(X1)=k1_xboole_0|~v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3307, plain, (r2_hidden(k1_xboole_0,k1_xboole_0)|~r1_xxreal_0(np__1,k1_xboole_0)|~r1_xxreal_0(k1_xboole_0,X1)|~v7_ordinal1(X1)|~v1_xboole_0(X1)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_2747, negated_conjecture, (v7_ordinal1(esk9_1(k4_finseq_1(esk2_0)))|~v6_membered(k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_3381, plain, (v1_xreal_0(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_3382, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v1_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_612, plain, (v3_card_1(X3,np__3)|k1_funct_1(X3,np__1)!=X1|k1_funct_1(X3,np__2)!=X2|k3_finseq_1(X3)!=np__3|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_3454, negated_conjecture, (k3_matrix_1(a_1_0_finseq_1(X1),esk2_0,np__3,np__3)=k3_matrix_1(esk1_0,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(a_1_0_finseq_1(X1)))|~m2_subset_1(k3_matrix_1(a_1_0_finseq_1(X1),esk2_0,np__3,np__3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_matrix_1(a_1_0_finseq_1(X1),esk2_0,np__3,np__3))|~r1_xxreal_0(k3_matrix_1(a_1_0_finseq_1(X1),esk2_0,np__3,np__3),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3455, negated_conjecture, (k3_matrix_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)),esk2_0,np__3,np__3)=k3_matrix_1(esk1_0,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))|~m2_finseq_1(k3_matrix_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)),esk2_0,np__3,np__3),X1))).#trainneg
cnf(i_0_3606, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,np__1)|X1!=k1_funct_1(k1_funct_1(esk2_0,X2),np__1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,X2),esk1_0)|~m1_subset_1(X1,esk1_0)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3605, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,np__2)|X1!=k1_funct_1(k1_funct_1(esk2_0,X2),np__2)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,X2),esk1_0)|~m1_subset_1(X1,esk1_0)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3652, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,np__3)|X1!=k1_funct_1(k1_funct_1(esk2_0,X2),np__3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,X2),esk1_0)|~m1_subset_1(X1,esk1_0)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3719, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|X1!=k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)|~m2_finseq_1(k1_funct_1(esk2_0,np__2),esk1_0)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_3649, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__2,np__3)|X1!=k1_funct_1(k1_funct_1(esk2_0,np__2),np__3)|~m2_finseq_1(k1_funct_1(esk2_0,np__2),esk1_0)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_3532, plain, (v1_xreal_0(k1_card_1(X1))|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_2585, plain, (m1_subset_1(k1_card_1(k1_enumset1(X1,X2,X3)),k5_numbers))).#trainneg
cnf(i_0_3067, negated_conjecture, (~v1_xboole_0(esk9_1(k4_finseq_1(esk2_0))))).#trainneg
cnf(i_0_613, plain, (v1_funcop_1(X3)|~v1_funct_1(k1_funct_1(X3,np__3))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(k1_funct_1(X3,np__3))|~v1_relat_1(X2)|~v1_relat_1(X1)|k1_funct_1(X3,np__1)!=X1|k1_funct_1(X3,np__2)!=X2|k3_finseq_1(X3)!=np__3|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_3643, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v1_xreal_0(X13)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_3779, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v1_xreal_0(X13)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_3722, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__1,np__3)|X1!=k1_funct_1(k1_funct_1(esk2_0,np__1),np__3)|~m2_finseq_1(k1_funct_1(esk2_0,np__1),esk1_0)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_3653, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|X1!=k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)|~m2_finseq_1(k1_funct_1(esk2_0,np__3),esk1_0)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_3720, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|X1!=k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)|~m2_finseq_1(k1_funct_1(esk2_0,np__2),esk1_0)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_3650, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|~m2_finseq_1(k1_funct_1(esk2_0,np__3),esk1_0)|~m1_subset_1(k1_funct_1(k1_funct_1(esk2_0,np__3),np__1),esk1_0))).#trainneg
cnf(i_0_3737, negated_conjecture, (k3_finseq_1(X1)=np__3|X1!=k1_funct_1(esk2_0,X2)|~m2_finseq_1(X1,esk1_0)|~v7_ordinal1(X2)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3803, negated_conjecture, (k3_matrix_1(k2_finseq_1(X1),esk2_0,np__3,np__3)=k3_matrix_1(esk1_0,esk2_0,np__3,np__3)|~m2_subset_1(k3_matrix_1(k2_finseq_1(X1),esk2_0,np__3,np__3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_matrix_1(k2_finseq_1(X1),esk2_0,np__3,np__3))|~r1_xxreal_0(k3_matrix_1(k2_finseq_1(X1),esk2_0,np__3,np__3),X1)|~m2_finseq_1(esk2_0,k3_finseq_2(k2_finseq_1(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3457, negated_conjecture, (X1=k1_funct_1(esk2_0,X2)|k1_funct_1(X1,esk5_2(X1,k1_funct_1(esk2_0,X2)))!=k3_matrix_1(esk1_0,esk2_0,X2,esk5_2(X1,k1_funct_1(esk2_0,X2)))|k3_finseq_1(X1)!=k3_finseq_1(k1_funct_1(esk2_0,X2))|~v1_finseq_1(k1_funct_1(esk2_0,X2))|~v1_finseq_1(X1)|~v1_funct_1(k1_funct_1(esk2_0,X2))|~v1_funct_1(X1)|~v1_relat_1(k1_funct_1(esk2_0,X2))|~v1_relat_1(X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(esk5_2(X1,k1_funct_1(esk2_0,X2)),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_2813, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,X4)|X5!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X5,X4)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X4,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_2815, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,np__2)|X4!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X4,np__2)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_3531, plain, (v1_xcmplx_0(k1_card_1(X1))|~v1_xreal_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3739, plain, (m1_subset_1(k1_card_1(X1),k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3692, negated_conjecture, (v7_ordinal1(esk9_1(k4_finseq_1(esk2_0))))).#trainneg
cnf(i_0_629, plain, (m2_finseq_1(esk4_5(X1,X2,X3,X4,X5),X1)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X4,X3)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_3694, negated_conjecture, (v1_xcmplx_0(esk9_1(k4_finseq_1(esk2_0)))|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_3632, plain, (r1_xxreal_0(X1,k3_matrix_1(X2,k11_finseq_1(k3_finseq_4(X2,X3,X4,X5),k11_finseq_1(X6,X7,X8),k11_finseq_1(X9,X10,X11)),X12,X13))|r1_xxreal_0(k3_matrix_1(X2,k11_finseq_1(k3_finseq_4(X2,X3,X4,X5),k11_finseq_1(X6,X7,X8),k11_finseq_1(X9,X10,X11)),X12,X13),X1)|v1_xboole_0(X2)|~v7_ordinal1(X1)|~v6_membered(X2)|~m1_subset_1(X11,X2)|~m1_subset_1(X10,X2)|~m1_subset_1(X9,X2)|~m1_subset_1(X8,X2)|~m1_subset_1(X7,X2)|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v7_ordinal1(X13)|~v7_ordinal1(X12))).#trainneg
cnf(i_0_3706, plain, (r1_xxreal_0(X1,k3_matrix_1(X2,k11_finseq_1(k11_finseq_1(X3,X4,X5),k3_finseq_4(X2,X6,X7,X8),k11_finseq_1(X9,X10,X11)),X12,X13))|r1_xxreal_0(k3_matrix_1(X2,k11_finseq_1(k11_finseq_1(X3,X4,X5),k3_finseq_4(X2,X6,X7,X8),k11_finseq_1(X9,X10,X11)),X12,X13),X1)|v1_xboole_0(X2)|~v7_ordinal1(X1)|~v6_membered(X2)|~m1_subset_1(X11,X2)|~m1_subset_1(X10,X2)|~m1_subset_1(X9,X2)|~m1_subset_1(X8,X2)|~m1_subset_1(X7,X2)|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v7_ordinal1(X13)|~v7_ordinal1(X12))).#trainneg
cnf(i_0_631, plain, (m2_finseq_1(esk4_5(X1,X2,X3,X4,X5),X1)|v1_xboole_0(k2_matrix_1(X2))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(k2_tarski(k2_tarski(X4,X3),k1_tarski(X3)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_2814, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,np__3)|X4!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X4,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_2807, negated_conjecture, (r2_hidden(esk7_3(X1,esk2_0,X2),k1_enumset1(np__1,np__2,np__3))|esk2_0=X2|v1_xboole_0(X1)|k3_finseq_1(X2)!=np__3|k1_matrix_1(X2)!=np__3|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m2_finseq_1(esk2_0,k3_finseq_2(X1)))).#trainneg
cnf(i_0_3389, negated_conjecture, (r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(k2_tarski(k2_tarski(X1,X2),k1_tarski(X2)),k2_matrix_1(esk2_0)))).#trainneg
cnf(i_0_3258, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|X2!=k1_funct_1(esk2_0,np__1)|X1!=k1_funct_1(X2,np__2)|~m2_finseq_1(X2,esk1_0)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_2731, negated_conjecture, (m1_subset_1(np__2,k4_finseq_1(esk2_0))|esk9_1(k4_finseq_1(esk2_0))=np__1|esk9_1(k4_finseq_1(esk2_0))=np__3)).#trainneg
cnf(i_0_2816, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,np__1)|X4!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X4,np__1)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_2817, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,X4)|X5!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X5,X4)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,np__3))).#trainneg
cnf(i_0_3899, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),np__2)=k3_matrix_1(esk1_0,esk2_0,X1,np__2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,X1),esk1_0)|~m1_subset_1(k1_funct_1(k1_funct_1(esk2_0,X1),np__2),esk1_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3896, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),np__1)=k3_matrix_1(esk1_0,esk2_0,X1,np__1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,X1),esk1_0)|~m1_subset_1(k1_funct_1(k1_funct_1(esk2_0,X1),np__1),esk1_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4015, plain, (v1_xcmplx_0(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3764, plain, (r2_hidden(k1_xboole_0,k1_xboole_0)|~r1_xxreal_0(np__1,k1_xboole_0)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_2749, negated_conjecture, (v1_xcmplx_0(esk9_1(k4_finseq_1(esk2_0)))|~v1_membered(k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_3949, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v1_xboole_0(X13))).#trainneg
cnf(i_0_3901, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),np__3)=k3_matrix_1(esk1_0,esk2_0,X1,np__3)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,X1),esk1_0)|~m1_subset_1(k1_funct_1(k1_funct_1(esk2_0,X1),np__3),esk1_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3973, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__3)=k3_matrix_1(esk1_0,esk2_0,np__1,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,np__1),esk1_0)|~m1_subset_1(k1_funct_1(k1_funct_1(esk2_0,np__1),np__3),esk1_0))).#trainneg
cnf(i_0_3979, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|~m2_finseq_1(k1_funct_1(esk2_0,np__2),esk1_0)|~m1_subset_1(k1_funct_1(k1_funct_1(esk2_0,np__2),np__1),esk1_0))).#trainneg
cnf(i_0_3977, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|~m2_finseq_1(k1_funct_1(esk2_0,np__3),esk1_0)|~m1_subset_1(k1_funct_1(k1_funct_1(esk2_0,np__3),np__2),esk1_0))).#trainneg
cnf(i_0_4178, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|X1!=k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)|~m2_finseq_1(k1_funct_1(esk2_0,np__1),esk1_0)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_3920, plain, (v1_xreal_0(k1_card_1(k1_enumset1(X1,X2,X3))))).#trainneg
cnf(i_0_3671, negated_conjecture, (m1_subset_1(np__1,k4_finseq_1(esk2_0))|np__3!=np__1|np__2!=np__1)).#trainneg
cnf(i_0_3080, plain, (m1_subset_1(esk9_1(k1_enumset1(np__1,np__2,np__3)),k5_numbers))).#trainneg
cnf(i_0_3963, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v1_xboole_0(X13))).#trainneg
cnf(i_0_3946, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v6_membered(X13))).#trainneg
cnf(i_0_636, plain, (m2_finseq_1(esk4_5(X1,X2,X3,X4,X5),X1)|v1_xboole_0(k2_matrix_1(X2))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(k2_tarski(k1_tarski(X3),k2_tarski(X3,X4)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_637, plain, (m2_finseq_1(esk4_5(X1,X2,esk6_3(X3,X2,X4),esk7_3(X3,X2,X4),X5),X1)|X2=X4|v1_xboole_0(X3)|X5!=k3_matrix_1(X1,X2,esk6_3(X3,X2,X4),esk7_3(X3,X2,X4))|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(esk7_3(X3,X2,X4))|~v7_ordinal1(esk6_3(X3,X2,X4))|k3_finseq_1(X2)!=k3_finseq_1(X4)|k1_matrix_1(X2)!=k1_matrix_1(X4)|~v1_matrix_1(X4)|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X2,k3_finseq_2(X3)))).#trainneg
cnf(i_0_4154, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,np__3)|X4!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X4,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_3906, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|~m2_finseq_1(k1_funct_1(esk2_0,np__2),esk1_0)|~m1_subset_1(k1_funct_1(k1_funct_1(esk2_0,np__2),np__2),esk1_0))).#trainneg
cnf(i_0_4203, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,np__1)|X4!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X4,np__1)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_3908, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__2),np__3)=k3_matrix_1(esk1_0,esk2_0,np__2,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,np__2),esk1_0)|~m1_subset_1(k1_funct_1(k1_funct_1(esk2_0,np__2),np__3),esk1_0))).#trainneg
cnf(i_0_4005, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,np__2)|X4!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X4,np__2)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_4168, negated_conjecture, (r1_xxreal_0(esk7_3(X1,esk2_0,X2),np__3)|esk2_0=X2|v1_xboole_0(X1)|k3_finseq_1(X2)!=np__3|k1_matrix_1(X2)!=np__3|~v1_matrix_1(X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_3981, negated_conjecture, (k3_finseq_1(k1_funct_1(esk2_0,X1))=np__3|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,X1),esk1_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4161, negated_conjecture, (esk7_3(X1,esk2_0,X2)=np__3|esk7_3(X1,esk2_0,X2)=np__2|esk7_3(X1,esk2_0,X2)=np__1|esk2_0=X2|v1_xboole_0(X1)|k3_finseq_1(X2)!=np__3|k1_matrix_1(X2)!=np__3|~v1_matrix_1(X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_4157, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,np__1,np__3)|X3!=k1_funct_1(esk2_0,np__1)|X1!=k1_funct_1(X3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X3,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_4206, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,np__1,np__1)|X3!=k1_funct_1(esk2_0,np__1)|X1!=k1_funct_1(X3,np__1)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X3,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_4025, plain, (v1_xreal_0(k1_card_1(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2729, negated_conjecture, (esk9_1(k4_finseq_1(esk2_0))=np__1|esk9_1(k4_finseq_1(esk2_0))=np__3|np__2!=np__3)).#trainneg
cnf(i_0_655, plain, (v7_ordinal1(k3_finseq_1(X1))|~v1_finset_1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3114, plain, (k1_xboole_0=np__1|k1_xboole_0=np__2|k1_xboole_0=np__3|r1_xxreal_0(np__3,k1_xboole_0)|~r1_xxreal_0(np__1,k1_xboole_0)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3947, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v3_membered(X13))).#trainneg
cnf(i_0_3960, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v6_membered(X13))).#trainneg
cnf(i_0_658, plain, (k3_finseq_1(X1)=np__3|v1_xboole_0(X2)|esk3_4(X4,X5,X6,k1_enumset1(np__1,np__2,np__3))=X4|esk3_4(X4,X5,X6,k1_enumset1(np__1,np__2,np__3))=X5|esk3_4(X4,X5,X6,k1_enumset1(np__1,np__2,np__3))=X6|k1_enumset1(np__1,np__2,np__3)=k1_enumset1(X4,X5,X6)|X1!=k1_funct_1(X3,esk3_4(X4,X5,X6,k1_enumset1(np__1,np__2,np__3)))|~m2_finseq_1(X1,X2)|~v7_ordinal1(esk3_4(X4,X5,X6,k1_enumset1(np__1,np__2,np__3)))|~m1_matrix_1(X3,X2,np__3,np__3))).#trainneg
cnf(i_0_659, plain, (k3_finseq_1(X1)=np__3|v1_xboole_0(X2)|v1_xboole_0(k1_enumset1(np__1,np__2,np__3))|X1!=k1_funct_1(X3,X4)|~m2_finseq_1(X1,X2)|~v7_ordinal1(X4)|~m1_matrix_1(X3,X2,np__3,np__3)|~m1_subset_1(X4,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_4008, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,np__1,np__2)|X3!=k1_funct_1(esk2_0,np__1)|X1!=k1_funct_1(X3,np__2)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X3,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_4158, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,np__3)|X4!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X4,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3))).#trainneg
cnf(i_0_4214, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,np__1,X3)|X4!=k1_funct_1(esk2_0,np__1)|X1!=k1_funct_1(X4,X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_4156, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,np__2,np__3)|X3!=k1_funct_1(esk2_0,np__2)|X1!=k1_funct_1(X3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X3,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_4007, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,np__2,np__2)|X3!=k1_funct_1(esk2_0,np__2)|X1!=k1_funct_1(X3,np__2)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X3,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_4205, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,np__2,np__1)|X3!=k1_funct_1(esk2_0,np__2)|X1!=k1_funct_1(X3,np__1)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X3,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_4207, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,np__1)|X4!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X4,np__1)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3))).#trainneg
cnf(i_0_4204, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,np__3,np__1)|X3!=k1_funct_1(esk2_0,np__3)|X1!=k1_funct_1(X3,np__1)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X3,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_4006, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,np__3,np__2)|X3!=k1_funct_1(esk2_0,np__3)|X1!=k1_funct_1(X3,np__2)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X3,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_4213, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,np__2,X3)|X4!=k1_funct_1(esk2_0,np__2)|X1!=k1_funct_1(X4,X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_4155, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,np__3,np__3)|X3!=k1_funct_1(esk2_0,np__3)|X1!=k1_funct_1(X3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X3,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_4212, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,np__3,X3)|X4!=k1_funct_1(esk2_0,np__3)|X1!=k1_funct_1(X4,X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_3098, negated_conjecture, (m1_subset_1(np__3,k4_finseq_1(esk2_0))|esk9_1(k4_finseq_1(esk2_0))=np__1|np__2!=np__1)).#trainneg
cnf(i_0_3816, negated_conjecture, (k11_finseq_1(k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__2),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__1,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__2)))).#trainneg
cnf(i_0_660, plain, (r2_hidden(X1,k4_finseq_1(X2))|v1_xboole_0(k2_matrix_1(X2))|~m1_subset_1(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),k2_matrix_1(X2))|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_4357, plain, (v1_xcmplx_0(k1_card_1(X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2617, plain, (~v1_xboole_0(esk9_1(k1_enumset1(np__1,np__2,np__3))))).#trainneg
cnf(i_0_4237, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|~m2_finseq_1(k1_funct_1(esk2_0,np__1),esk1_0)|~m1_subset_1(k1_funct_1(k1_funct_1(esk2_0,np__1),np__2),esk1_0))).#trainneg
cnf(i_0_4312, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,np__3)|X4!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X4,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_4342, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,np__2)|X4!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X4,np__2)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_4327, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,np__1)|X4!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X4,np__1)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_4498, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,np__2,X3)|X4!=k1_funct_1(esk2_0,np__2)|X1!=k1_funct_1(X4,X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_4216, plain, (v1_xcmplx_0(k1_card_1(X1))|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_4360, plain, (v1_xreal_0(k1_card_1(X1))|~m1_subset_1(k1_card_1(X1),k5_numbers))).#trainneg
cnf(i_0_4366, negated_conjecture, (esk9_1(k4_finseq_1(esk2_0))=np__1|np__3!=np__1|np__2!=np__3)).#trainneg
cnf(i_0_2493, plain, (esk9_1(k1_enumset1(np__1,np__2,np__3))=np__1|esk9_1(k1_enumset1(np__1,np__2,np__3))=np__2|esk9_1(k1_enumset1(np__1,np__2,np__3))=np__3)).#trainneg
cnf(i_0_2468, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_4258, plain, (v7_ordinal1(esk9_1(k1_enumset1(np__1,np__2,np__3))))).#trainneg
cnf(i_0_4259, plain, (v1_xreal_0(esk9_1(k1_enumset1(np__1,np__2,np__3))))).#trainneg
cnf(i_0_4499, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,np__2,X3)|X4!=k1_funct_1(esk2_0,np__2)|X1!=k1_funct_1(X4,X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_4467, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,np__1,X3)|X4!=k1_funct_1(esk2_0,np__1)|X1!=k1_funct_1(X4,X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_4009, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,np__2)|X4!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X4,np__2)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3))).#trainneg
cnf(i_0_3983, negated_conjecture, (k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,np__3,np__3)=k3_matrix_1(esk1_0,esk2_0,np__3,np__3)|~m2_subset_1(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,np__3,np__3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,np__3,np__3))|~r1_xxreal_0(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,np__3,np__3),np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(k1_enumset1(np__1,np__2,np__3))))).#trainneg
cnf(i_0_4356, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|X2!=k1_funct_1(esk2_0,np__1)|X1!=k1_funct_1(X2,np__1)|~m2_finseq_1(X2,esk1_0)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_1418, plain, (v1_xcmplx_0(esk9_1(X1))|~v1_membered(k5_numbers)|~v6_membered(X1))).#trainneg
cnf(i_0_669, plain, (esk9_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2298, plain, (v1_xcmplx_0(esk9_1(X1))|~v1_xreal_0(X2)|~v6_membered(X1))).#trainneg
cnf(i_0_4260, plain, (v1_xcmplx_0(esk9_1(k1_enumset1(np__1,np__2,np__3)))|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_670, plain, (v1_xboole_0(esk9_1(k1_zfmisc_1(k2_matrix_1(X1))))|~v1_xboole_0(k4_finseq_1(X1))|~v1_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_683, plain, (r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2)))|v1_xboole_0(k2_matrix_1(X2))|~m1_subset_1(k2_tarski(k2_tarski(X3,X1),k1_tarski(X3)),k2_matrix_1(X2))|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_4468, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,np__1,X3)|X4!=k1_funct_1(esk2_0,np__1)|X1!=k1_funct_1(X4,X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_4396, plain, (v7_ordinal1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_4174, negated_conjecture, (r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(k2_tarski(k2_tarski(X2,X1),k1_tarski(X2)),k2_matrix_1(esk2_0)))).#trainneg
cnf(i_0_3921, plain, (v1_xcmplx_0(k1_card_1(k1_enumset1(X1,X2,X3)))|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_4179, negated_conjecture, (esk9_1(k4_finseq_1(esk2_0))=np__1|m1_subset_1(np__2,k4_finseq_1(esk2_0))|np__3!=np__1)).#trainneg
cnf(i_0_4534, negated_conjecture, (k11_finseq_1(k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__2),k1_funct_1(esk2_0,np__3))!=esk2_0|k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__1,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_2299, plain, (v1_xcmplx_0(esk9_1(X1))|~v1_xreal_0(X2)|~v3_membered(X1))).#trainneg
cnf(i_0_3112, plain, (k1_funct_1(esk4_5(X1,X2,X3,X4,X5),X4)=X5|v1_xboole_0(k2_matrix_1(X2))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(k2_tarski(k1_tarski(X3),k2_tarski(X4,X3)),k2_matrix_1(X2))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_760, plain, (k3_finseq_1(X1)=np__3|v1_xboole_0(X2)|X1!=k1_funct_1(X3,np__2)|~m2_finseq_1(X1,X2)|~m1_matrix_1(X3,X2,np__3,np__3))).#trainneg
cnf(i_0_2467, plain, (v3_valued_0(X1)|~v3_membered(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_4368, negated_conjecture, (m1_subset_1(np__3,k4_finseq_1(esk2_0))|esk9_1(k4_finseq_1(esk2_0))=np__1|np__2!=np__3)).#trainneg
cnf(i_0_4515, negated_conjecture, (m1_subset_1(np__1,k4_finseq_1(esk2_0))|m1_subset_1(np__3,k4_finseq_1(esk2_0))|np__2!=np__1)).#trainneg
cnf(i_0_4441, plain, (k3_finseq_1(X1)=np__3|v1_xboole_0(X2)|X1!=k1_funct_1(X3,X4)|~m2_finseq_1(X1,X2)|~v7_ordinal1(X4)|~m1_matrix_1(X3,X2,np__3,np__3)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,np__3)|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_3806, negated_conjecture, (k3_matrix_1(a_1_0_finseq_1(X1),esk2_0,np__3,np__3)=k3_matrix_1(esk1_0,esk2_0,np__3,np__3)|~r1_xxreal_0(np__1,k3_matrix_1(a_1_0_finseq_1(X1),esk2_0,np__3,np__3))|~r1_xxreal_0(k3_matrix_1(a_1_0_finseq_1(X1),esk2_0,np__3,np__3),X1)|~m2_finseq_1(esk2_0,k3_finseq_2(a_1_0_finseq_1(X1)))|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(a_1_0_finseq_1(X1),esk2_0,np__3,np__3),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_1705, plain, (esk8_2(X1,X2)=np__1|esk8_2(X1,X2)=np__2|esk8_2(X1,X2)=np__3|~r1_xxreal_0(np__1,esk8_2(X1,X2))|~r1_xxreal_0(esk8_2(X1,X2),np__3)|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2)))).#trainneg
cnf(i_0_1936, plain, (v1_funct_1(esk4_5(X1,X2,X3,X4,X5))|X5!=k3_matrix_1(X1,X2,X3,X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X4,k1_matrix_1(X2))|~r1_xxreal_0(np__1,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(k1_matrix_1(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2)))).#trainneg
cnf(i_0_1117, plain, (r1_xxreal_0(np__2,X1)|r1_xxreal_0(X1,np__2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_845, plain, (v1_xcmplx_0(X1)|~m2_finseq_1(X1,X2)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_2478, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3961, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v3_membered(X13))).#trainneg
cnf(i_0_4649, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__1,X2)|X3!=k1_funct_1(esk2_0,np__1)|X1!=k1_funct_1(X3,X2)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(X3,esk1_0)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_4665, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|X1!=k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)|~m2_finseq_1(k1_funct_1(esk2_0,np__1),esk1_0)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_4560, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__2,X2)|X3!=k1_funct_1(esk2_0,np__2)|X1!=k1_funct_1(X3,X2)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(X3,esk1_0)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_3296, plain, (r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2)))|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X1,X3)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_4683, plain, (v3_membered(k1_xboole_0)|~v3_membered(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4698, plain, (v1_xcmplx_0(esk9_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_4739, plain, (v7_ordinal1(k3_finseq_1(k1_xboole_0))|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_695, plain, (X1=k3_matrix_1(X2,X3,X4,X5)|X1!=k1_funct_1(X6,X5)|X6!=k1_funct_1(X3,X4)|~v1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X6,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~r2_hidden(k2_tarski(k1_tarski(X4),k2_tarski(X5,X4)),k2_matrix_1(X3)))).#trainneg
cnf(i_0_4806, plain, (v3_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_4742, negated_conjecture, (v7_ordinal1(k3_finseq_1(np__3))|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(esk2_0,k5_numbers))).#trainneg
cnf(i_0_699, plain, (X1=k3_matrix_1(X2,X3,X4,X5)|v1_xboole_0(k2_matrix_1(X3))|X1!=k1_funct_1(X6,X5)|X6!=k1_funct_1(X3,X4)|~v1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X6,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(k2_tarski(k1_tarski(X4),k2_tarski(X4,X5)),k2_matrix_1(X3)))).#trainneg
cnf(i_0_1253, plain, (m2_finseq_1(esk4_5(X1,X2,X3,X3,X4),X1)|X4!=k3_matrix_1(X1,X2,X3,X3)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,k1_matrix_1(X2))|~v7_ordinal1(k1_matrix_1(X2)))).#trainneg
cnf(i_0_4215, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,X4)|X5!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X5,X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3))).#trainneg
cnf(i_0_4777, plain, (v1_xcmplx_0(esk9_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_4863, plain, (r1_xxreal_0(esk9_1(k5_numbers),np__2)|r1_xxreal_0(np__2,esk9_1(k5_numbers)))).#trainneg
cnf(i_0_4772, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__1,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_4545, plain, (v1_xcmplx_0(k1_card_1(X1))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4578, negated_conjecture, (m1_subset_1(np__1,k4_finseq_1(esk2_0))|np__3!=np__1|np__2!=np__3)).#trainneg
cnf(i_0_700, plain, (X1=k3_matrix_1(X2,X3,esk6_3(X4,X3,X5),esk7_3(X4,X3,X5))|X3=X5|v1_xboole_0(X4)|X1!=k1_funct_1(X6,esk7_3(X4,X3,X5))|X6!=k1_funct_1(X3,esk6_3(X4,X3,X5))|~v1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X6,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(esk7_3(X4,X3,X5))|~v7_ordinal1(esk6_3(X4,X3,X5))|k3_finseq_1(X3)!=k3_finseq_1(X5)|k1_matrix_1(X3)!=k1_matrix_1(X5)|~v1_matrix_1(X5)|~m2_finseq_1(X5,k3_finseq_2(X4))|~m2_finseq_1(X3,k3_finseq_2(X4)))).#trainneg
cnf(i_0_702, plain, (k3_finseq_1(X1)=X2|v1_xboole_0(k2_finseq_1(X2))|v1_xboole_0(X3)|X1!=k1_funct_1(X4,esk9_1(k2_finseq_1(X2)))|~m2_finseq_1(X1,X3)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(esk9_1(k2_finseq_1(X2)))|~m1_matrix_1(X4,X3,X2,X2))).#trainneg
cnf(i_0_1639, plain, (~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v1_xboole_0(k2_finseq_1(X1)))).#trainneg
cnf(i_0_851, plain, (r2_hidden(esk6_3(X1,X2,X3),k4_finseq_1(X2))|X2=X3|v1_xboole_0(X1)|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_4513, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,np__3,X3)|X4!=k1_funct_1(esk2_0,np__3)|X1!=k1_funct_1(X4,X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_4544, plain, (v1_xcmplx_0(k1_card_1(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4678, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3822, negated_conjecture, (k11_finseq_1(k1_funct_1(esk2_0,np__1),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k1_funct_1(esk2_0,np__3))!=esk2_0|k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__1,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_1894, plain, (~r1_xxreal_0(np__3,X1)|~v7_ordinal1(X1)|~v1_xboole_0(k2_finseq_1(X1)))).#trainneg
cnf(i_0_1220, plain, (k1_funct_1(esk4_5(X1,X2,X3,X4,X5),X4)=X5|v1_xboole_0(k2_finseq_1(k1_matrix_1(X2)))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2))|~m1_subset_1(X4,k2_finseq_1(k1_matrix_1(X2))))).#trainneg
cnf(i_0_1460, negated_conjecture, (r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X1,X2)),k2_matrix_1(esk2_0))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_4211, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,X4)|X5!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X5,X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_2721, negated_conjecture, (k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,X2)=k1_funct_1(k1_funct_1(esk2_0,X1),X2)|~m2_finseq_1(esk2_0,k3_finseq_2(k1_enumset1(np__1,np__2,np__3)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r1_xxreal_0(np__1,k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,X2))|~r1_xxreal_0(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,X2),np__3)|~m1_subset_1(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,X2),k5_numbers))).#trainneg
cnf(i_0_1937, plain, (v1_finseq_1(esk4_5(X1,X2,X3,X4,X5))|X5!=k3_matrix_1(X1,X2,X3,X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X4,k1_matrix_1(X2))|~r1_xxreal_0(np__1,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(k1_matrix_1(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2)))).#trainneg
cnf(i_0_4903, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__2,X2)|X1!=k1_funct_1(k1_funct_1(esk2_0,np__2),X2)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,np__2),esk1_0)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_4954, plain, (v3_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_4930, plain, (m1_subset_1(k3_finseq_1(k1_xboole_0),k5_numbers)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_4893, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v7_ordinal1(X14)|~v7_ordinal1(X13))).#trainneg
cnf(i_0_4900, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__1,X2)|X1!=k1_funct_1(k1_funct_1(esk2_0,np__1),X2)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,np__1),esk1_0)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_4901, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|~m2_finseq_1(k1_funct_1(esk2_0,np__1),esk1_0)|~m1_subset_1(k1_funct_1(k1_funct_1(esk2_0,np__1),np__1),esk1_0))).#trainneg
cnf(i_0_5114, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,np__3,X3)|X4!=k1_funct_1(esk2_0,np__3)|X1!=k1_funct_1(X4,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_4181, negated_conjecture, (m1_subset_1(np__3,k4_finseq_1(esk2_0))|esk9_1(k4_finseq_1(esk2_0))=np__1|m1_subset_1(np__2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_5019, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__1,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_4894, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v7_ordinal1(X14)|~v7_ordinal1(X13))).#trainneg
cnf(i_0_3952, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v3_valued_0(X13)|~v1_funct_1(X13)|~v1_relat_1(X13))).#trainneg
cnf(i_0_5215, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,X4)|X5!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X5,X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_5006, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,X4)|X5!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X5,X4)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,np__3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X4,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_4617, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_4753, negated_conjecture, (m1_subset_1(np__1,k4_finseq_1(esk2_0))|m1_subset_1(np__2,k4_finseq_1(esk2_0))|np__3!=np__1)).#trainneg
cnf(i_0_720, plain, (X1=k3_matrix_1(X2,X3,X4,X5)|v1_xboole_0(k2_matrix_1(X3))|X1!=k1_funct_1(X6,X5)|X6!=k1_funct_1(X3,X4)|~v1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X6,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(k2_tarski(k2_tarski(X5,X4),k1_tarski(X4)),k2_matrix_1(X3)))).#trainneg
cnf(i_0_4960, negated_conjecture, (m1_subset_1(k3_finseq_1(np__3),k5_numbers)|~m1_subset_1(esk2_0,k5_numbers)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_3966, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v3_valued_0(X13)|~v1_funct_1(X13)|~v1_relat_1(X13))).#trainneg
cnf(i_0_722, plain, (k1_matrix_1(k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),k3_finseq_4(X1,X8,X9,X10)))=np__3|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_725, plain, (m2_finseq_1(k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),k3_finseq_4(X1,X8,X9,X10)),k3_finseq_2(X1))|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_4980, plain, (v3_valued_0(esk4_5(X1,X2,X3,X3,X4))|~v3_membered(X1)|X4!=k3_matrix_1(X1,X2,X3,X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X3,k1_matrix_1(X2))|~r1_xxreal_0(np__1,X3)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(k1_matrix_1(X2))|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2)))).#trainneg
cnf(i_0_5005, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,X4)|X5!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X5,X4)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,np__3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_4977, plain, (v1_relat_1(esk4_5(X1,X2,X3,X3,X4))|X4!=k3_matrix_1(X1,X2,X3,X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X3,k1_matrix_1(X2))|~r1_xxreal_0(np__1,X3)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(k1_matrix_1(X2))|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2)))).#trainneg
cnf(i_0_4946, plain, (X1=k3_matrix_1(X2,X3,X4,X4)|X1!=k1_funct_1(X5,X4)|X5!=k1_funct_1(X3,X4)|~v1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X4)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(X4,k2_finseq_1(k1_matrix_1(X3)))|~r2_hidden(X4,k4_finseq_1(X3)))).#trainneg
cnf(i_0_4975, plain, (v1_funct_1(esk4_5(X1,X2,X3,X3,X4))|X4!=k3_matrix_1(X1,X2,X3,X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X3,k1_matrix_1(X2))|~r1_xxreal_0(np__1,X3)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(k1_matrix_1(X2))|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2)))).#trainneg
cnf(i_0_5115, plain, (v1_xcmplx_0(k1_card_1(X1))|~m1_subset_1(k1_card_1(X1),k5_numbers))).#trainneg
cnf(i_0_4812, negated_conjecture, (m1_subset_1(np__1,k4_finseq_1(esk2_0))|m1_subset_1(np__3,k4_finseq_1(esk2_0))|np__2!=np__3)).#trainneg
cnf(i_0_732, plain, (m2_subset_1(k1_card_1(X1),k1_numbers,k5_numbers)|~v1_finseq_1(k1_card_1(X1))|~v1_funct_1(k1_card_1(X1))|~v1_relat_1(k1_card_1(X1)))).#trainneg
cnf(i_0_4409, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v1_finseq_1(X13)|~v1_funct_1(X13)|~v1_relat_1(X13))).#trainneg
cnf(i_0_5263, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v7_ordinal1(X16))).#trainneg
cnf(i_0_727, plain, (m1_matrix_1(k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),X7,k3_finseq_4(X1,X8,X9,X10)),X1,np__3,np__3)|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(k1_funct_1(X7,np__3),X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|k1_funct_1(X7,np__1)!=X5|k1_funct_1(X7,np__2)!=X6|k3_finseq_1(X7)!=np__3|~v1_finseq_1(X7)|~v1_funct_1(X7)|~v1_relat_1(X7))).#trainneg
cnf(i_0_733, plain, (k1_card_1(X1)=X2|r1_xxreal_0(esk5_2(k1_card_1(X1),X2),k1_card_1(X1))|k1_card_1(X1)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(k1_card_1(X1))|~v1_funct_1(X2)|~v1_funct_1(k1_card_1(X1))|~v1_relat_1(X2)|~v1_relat_1(k1_card_1(X1)))).#trainneg
cnf(i_0_5282, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__3,X2)|X3!=k1_funct_1(esk2_0,np__3)|X1!=k1_funct_1(X3,X2)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(X3,esk1_0)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_5278, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),X1)=k3_matrix_1(esk1_0,esk2_0,np__1,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,np__1),esk1_0)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(k1_funct_1(k1_funct_1(esk2_0,np__1),X1),esk1_0))).#trainneg
cnf(i_0_5232, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__2),X1)=k3_matrix_1(esk1_0,esk2_0,np__2,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,np__2),esk1_0)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(k1_funct_1(k1_funct_1(esk2_0,np__2),X1),esk1_0))).#trainneg
cnf(i_0_5286, negated_conjecture, (m1_subset_1(np__1,k4_finseq_1(esk2_0))|m1_subset_1(np__2,k4_finseq_1(esk2_0))|m1_subset_1(np__3,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_5464, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,X4)|X5!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X5,X4)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,np__3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_5236, plain, (v1_xreal_0(X1)|~m1_subset_1(k11_finseq_1(X1,X2,X3),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_5246, plain, (v1_xreal_0(k3_finseq_1(k1_xboole_0))|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_5247, plain, (v1_xcmplx_0(k3_finseq_1(k1_xboole_0))|~v1_membered(k5_numbers)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_5315, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v7_ordinal1(X16))).#trainneg
cnf(i_0_738, plain, (k3_finseq_1(k3_finseq_1(X1))=k3_finseq_1(X1)|~v1_finseq_1(k3_finseq_1(X1))|~v1_funct_1(k3_finseq_1(X1))|~v1_relat_1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5465, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,X4)|X5!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X5,X4)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,np__3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_5349, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,X4)|X5!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X5,X4)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,np__3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X1,X2)|~v7_ordinal1(X4)|~m1_subset_1(X4,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_5466, negated_conjecture, (v1_relat_1(esk4_5(X1,esk2_0,X2,X2,X3))|X3!=k3_matrix_1(X1,esk2_0,X2,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_5471, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,X3)|X1!=k1_funct_1(X4,X3)|X4!=k1_funct_1(esk2_0,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X3,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_5478, negated_conjecture, (v1_funct_1(esk4_5(X1,esk2_0,X2,X2,X3))|X3!=k3_matrix_1(X1,esk2_0,X2,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_5237, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k11_finseq_1(X1,X2,X3),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_2496, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5305, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__1,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_5384, negated_conjecture, (v1_xreal_0(k3_finseq_1(np__3))|~m1_subset_1(esk2_0,k5_numbers)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_4897, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v1_finseq_1(X13)|~v1_funct_1(X13)|~v1_relat_1(X13))).#trainneg
cnf(i_0_4976, plain, (v1_finseq_1(esk4_5(X1,X2,X3,X3,X4))|X4!=k3_matrix_1(X1,X2,X3,X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X3,k1_matrix_1(X2))|~r1_xxreal_0(np__1,X3)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(k1_matrix_1(X2))|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2)))).#trainneg
cnf(i_0_825, plain, (~v1_xboole_0(k2_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2307, plain, (r2_hidden(X1,k4_finseq_1(X2))|v1_xboole_0(k2_matrix_1(X2))|~m1_subset_1(k2_tarski(k1_tarski(X1),k2_tarski(X1,X3)),k2_matrix_1(X2))|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_2147, plain, (esk5_2(X1,X2)=np__3|esk5_2(X1,X2)=np__2|esk5_2(X1,X2)=np__1|X1=X2|~r1_xxreal_0(esk5_2(X1,X2),np__3)|~m1_subset_1(esk5_2(X1,X2),k5_numbers)|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5238, plain, (v1_xreal_0(X1)|~m1_subset_1(k11_finseq_1(X2,X1,X3),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_889, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_5142, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__1,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__3))|k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3))!=k1_funct_1(esk2_0,np__2))).#trainneg
cnf(i_0_5790, plain, (v4_relat_1(k1_xboole_0,X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_763, plain, (k3_finseq_1(X1)=np__3|v1_xboole_0(X2)|X1!=k1_funct_1(X3,np__1)|~m2_finseq_1(X1,X2)|~m1_matrix_1(X3,X2,np__3,np__3))).#trainneg
cnf(i_0_789, plain, (v1_relat_1(esk9_1(k1_zfmisc_1(k2_matrix_1(X1))))|~v1_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5602, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,X3)|X4!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X4,X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(X4,esk1_0)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_5590, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__3,X2)|X1!=k1_funct_1(k1_funct_1(esk2_0,np__3),X2)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,np__3),esk1_0)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_5697, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,X3)|X4!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X4,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3))).#trainneg
cnf(i_0_810, plain, (r1_xxreal_0(X1,esk9_1(X2))|r1_xxreal_0(esk9_1(X2),X1)|~v7_ordinal1(X1)|~v6_membered(X2))).#trainneg
cnf(i_0_811, plain, (r1_xxreal_0(X1,esk7_3(X2,X3,X4))|r1_xxreal_0(esk7_3(X2,X3,X4),X1)|X3=X4|v1_xboole_0(X2)|~v7_ordinal1(X1)|k3_finseq_1(X3)!=k3_finseq_1(X4)|k1_matrix_1(X3)!=k1_matrix_1(X4)|~v1_matrix_1(X4)|~v1_matrix_1(X3)|~m2_finseq_1(X4,k3_finseq_2(X2))|~m2_finseq_1(X3,k3_finseq_2(X2)))).#trainneg
cnf(i_0_812, plain, (r1_xxreal_0(X1,esk6_3(X2,X3,X4))|r1_xxreal_0(esk6_3(X2,X3,X4),X1)|X3=X4|v1_xboole_0(X2)|~v7_ordinal1(X1)|k3_finseq_1(X3)!=k3_finseq_1(X4)|k1_matrix_1(X3)!=k1_matrix_1(X4)|~v1_matrix_1(X4)|~v1_matrix_1(X3)|~m2_finseq_1(X4,k3_finseq_2(X2))|~m2_finseq_1(X3,k3_finseq_2(X2)))).#trainneg
cnf(i_0_5693, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,X3)|X4!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X4,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_5688, negated_conjecture, (v1_relat_1(esk4_5(X1,esk2_0,X2,X2,X3))|X3!=k3_matrix_1(X1,esk2_0,X2,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~m1_subset_1(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_5718, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__1,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_813, plain, (r1_xxreal_0(X1,esk5_2(X2,X3))|r1_xxreal_0(esk5_2(X2,X3),X1)|X2=X3|~v7_ordinal1(X1)|k3_finseq_1(X2)!=k3_finseq_1(X3)|~v1_finseq_1(X3)|~v1_finseq_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X2))).#trainneg
cnf(i_0_5783, plain, (v3_membered(X1)|~v3_membered(k1_xboole_0)|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_5385, negated_conjecture, (v1_xcmplx_0(k3_finseq_1(np__3))|~v1_membered(k5_numbers)|~m1_subset_1(esk2_0,k5_numbers)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_814, plain, (r1_xxreal_0(X1,k3_xcmplx_0(X2,X3))|r1_xxreal_0(k3_xcmplx_0(X2,X3),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_816, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|r1_xxreal_0(k2_xcmplx_0(X2,X3),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3805, negated_conjecture, (k3_matrix_1(a_1_0_finseq_1(X1),esk2_0,np__3,np__3)=k3_matrix_1(esk1_0,esk2_0,np__3,np__3)|~r1_xxreal_0(np__1,k3_matrix_1(a_1_0_finseq_1(X1),esk2_0,np__3,np__3))|~r1_xxreal_0(k3_matrix_1(a_1_0_finseq_1(X1),esk2_0,np__3,np__3),X1)|~m2_finseq_1(esk2_0,k3_finseq_2(a_1_0_finseq_1(X1)))|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(a_1_0_finseq_1(X1),esk2_0,np__3,np__3),k5_numbers))).#trainneg
cnf(i_0_4297, plain, (X1=X2|m2_finseq_1(k1_funct_1(X1,esk6_3(X4,X1,X2)),X3)|v1_xboole_0(X4)|X5!=k3_matrix_1(X3,X1,esk6_3(X4,X1,X2),esk7_3(X4,X1,X2))|k3_finseq_1(X1)!=k3_finseq_1(X2)|k1_matrix_1(X1)!=k1_matrix_1(X2)|~v1_matrix_1(X1)|~v1_matrix_1(X2)|~m2_finseq_1(X1,k3_finseq_2(X3))|~m2_finseq_1(X2,k3_finseq_2(X4))|~m2_finseq_1(X1,k3_finseq_2(X4))|~m1_subset_1(X5,X3))).#trainneg
cnf(i_0_2708, plain, (k3_finseq_1(X1)=np__3|v1_xboole_0(X2)|X1!=k11_finseq_1(X9,X10,X11)|~m2_finseq_1(X1,X2)|~m1_subset_1(X11,X2)|~m1_subset_1(X10,X2)|~m1_subset_1(X9,X2)|~m1_subset_1(X8,X2)|~m1_subset_1(X7,X2)|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_3388, plain, (r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2)))|v1_xboole_0(k2_matrix_1(X2))|~m1_subset_1(k2_tarski(k2_tarski(X1,X3),k1_tarski(X3)),k2_matrix_1(X2))|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_5239, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k11_finseq_1(X2,X1,X3),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_2474, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_817, plain, (r1_xxreal_0(X1,k1_matrix_1(X2))|r1_xxreal_0(k1_matrix_1(X2),X1)|~v7_ordinal1(X1)|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_5607, plain, (v1_xcmplx_0(k3_finseq_1(k1_xboole_0))|~v1_xreal_0(X1)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_3019, negated_conjecture, (v1_xreal_0(k3_matrix_1(esk1_0,esk2_0,np__3,np__3))|~v3_valued_0(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_3020, negated_conjecture, (v1_xxreal_0(k3_matrix_1(esk1_0,esk2_0,np__3,np__3))|~v2_valued_0(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_822, plain, (k2_finseq_1(X1)=k1_enumset1(X2,X3,X4)|esk3_4(X2,X3,X4,k2_finseq_1(X1))!=X4|~m2_subset_1(esk3_4(X2,X3,X4,k2_finseq_1(X1)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk3_4(X2,X3,X4,k2_finseq_1(X1)))|~r1_xxreal_0(esk3_4(X2,X3,X4,k2_finseq_1(X1)),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_823, plain, (k2_finseq_1(X1)=k1_enumset1(X2,X3,X4)|esk3_4(X2,X3,X4,k2_finseq_1(X1))!=X3|~m2_subset_1(esk3_4(X2,X3,X4,k2_finseq_1(X1)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk3_4(X2,X3,X4,k2_finseq_1(X1)))|~r1_xxreal_0(esk3_4(X2,X3,X4,k2_finseq_1(X1)),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_923, negated_conjecture, (r2_hidden(esk6_3(X1,esk2_0,X2),k1_enumset1(np__1,np__2,np__3))|esk2_0=X2|v1_xboole_0(X1)|k3_finseq_1(X2)!=np__3|k1_matrix_1(X2)!=np__3|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m2_finseq_1(esk2_0,k3_finseq_2(X1)))).#trainneg
cnf(i_0_5240, plain, (v1_xreal_0(X1)|~m1_subset_1(k11_finseq_1(X2,X3,X1),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_4038, negated_conjecture, (r1_xxreal_0(np__2,esk9_1(k4_finseq_1(esk2_0)))|r1_xxreal_0(esk9_1(k4_finseq_1(esk2_0)),np__2))).#trainneg
cnf(i_0_824, plain, (k2_finseq_1(X1)=k1_enumset1(X2,X3,X4)|esk3_4(X2,X3,X4,k2_finseq_1(X1))!=X2|~m2_subset_1(esk3_4(X2,X3,X4,k2_finseq_1(X1)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk3_4(X2,X3,X4,k2_finseq_1(X1)))|~r1_xxreal_0(esk3_4(X2,X3,X4,k2_finseq_1(X1)),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5828, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_3021, negated_conjecture, (v1_xcmplx_0(k3_matrix_1(esk1_0,esk2_0,np__3,np__3))|~v1_valued_0(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_4931, plain, (r1_xxreal_0(np__2,k3_finseq_1(k1_xboole_0))|r1_xxreal_0(k3_finseq_1(k1_xboole_0),np__2)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_826, plain, (m1_subset_1(X1,k2_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_827, plain, (k3_finseq_1(X1)=X2|v1_xboole_0(X3)|X1!=k1_funct_1(X4,X5)|~m2_finseq_1(X1,X3)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X5)|~m1_matrix_1(X4,X3,X2,X2)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5857, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__3),X1)=k3_matrix_1(esk1_0,esk2_0,np__3,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,np__3),esk1_0)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(k1_funct_1(k1_funct_1(esk2_0,np__3),X1),esk1_0))).#trainneg
cnf(i_0_5856, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,X3)|X1!=k1_funct_1(k1_funct_1(esk2_0,X2),X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,X2),esk1_0)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_5994, plain, (X1=X2|m2_finseq_1(k1_funct_1(X1,esk6_3(X3,X1,X2)),X4)|v1_xboole_0(X3)|k3_finseq_1(X1)!=k3_finseq_1(X2)|k1_matrix_1(X1)!=k1_matrix_1(X2)|~v1_matrix_1(X1)|~v1_matrix_1(X2)|~m2_finseq_1(X1,k3_finseq_2(X4))|~m2_finseq_1(X2,k3_finseq_2(X3))|~m2_finseq_1(X1,k3_finseq_2(X3))|~m1_subset_1(k3_matrix_1(X4,X1,esk6_3(X3,X1,X2),esk7_3(X3,X1,X2)),X4))).#trainneg
cnf(i_0_6101, negated_conjecture, (k3_matrix_1(k2_finseq_1(X2),esk2_0,X1,np__3)=k1_funct_1(k1_funct_1(esk2_0,X1),np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(k2_finseq_1(X2)))|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0))|~m2_subset_1(k3_matrix_1(k2_finseq_1(X2),esk2_0,X1,np__3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_matrix_1(k2_finseq_1(X2),esk2_0,X1,np__3))|~r1_xxreal_0(k3_matrix_1(k2_finseq_1(X2),esk2_0,X1,np__3),X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5851, plain, (k3_finseq_1(X1)=np__3|v1_xboole_0(X2)|X1!=k3_finseq_4(X2,X3,X4,X5)|~m2_finseq_1(X1,X2)|~m1_subset_1(X11,X2)|~m1_subset_1(X10,X2)|~m1_subset_1(X9,X2)|~m1_subset_1(X8,X2)|~m1_subset_1(X7,X2)|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_5841, plain, (v4_relat_1(k1_xboole_0,X1))).#trainneg
cnf(i_0_6090, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_6017, negated_conjecture, (v1_xcmplx_0(k3_finseq_1(k1_xboole_0))|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_840, plain, (v1_xboole_0(k2_matrix_1(X1))|~v1_xboole_0(k4_finseq_1(X1))|~v1_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4932, plain, (r1_xxreal_0(np__1,k3_finseq_1(k1_xboole_0))|r1_xxreal_0(k3_finseq_1(k1_xboole_0),np__1)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_843, plain, (v1_xreal_0(X1)|~v3_valued_0(X3)|k1_funct_1(X3,np__1)!=X2|k1_funct_1(X3,np__2)!=X1|k3_finseq_1(X3)!=np__3|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_844, plain, (v1_xxreal_0(X1)|~v2_valued_0(X3)|k1_funct_1(X3,np__1)!=X2|k1_funct_1(X3,np__2)!=X1|k3_finseq_1(X3)!=np__3|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_5874, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,X3)|X4!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X4,X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X4,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X3)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_6050, negated_conjecture, (esk6_3(X1,esk2_0,X2)=np__3|esk6_3(X1,esk2_0,X2)=np__2|esk6_3(X1,esk2_0,X2)=np__1|esk2_0=X2|v1_xboole_0(X1)|k3_finseq_1(X2)!=np__3|k1_matrix_1(X2)!=np__3|~v1_matrix_1(X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_5992, negated_conjecture, (k3_matrix_1(k2_finseq_1(X1),esk2_0,np__3,np__3)=k3_matrix_1(esk1_0,esk2_0,np__3,np__3)|~r1_xxreal_0(np__1,k3_matrix_1(k2_finseq_1(X1),esk2_0,np__3,np__3))|~r1_xxreal_0(k3_matrix_1(k2_finseq_1(X1),esk2_0,np__3,np__3),X1)|~m2_finseq_1(esk2_0,k3_finseq_2(k2_finseq_1(X1)))|~m1_subset_1(k3_matrix_1(k2_finseq_1(X1),esk2_0,np__3,np__3),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6057, negated_conjecture, (r1_xxreal_0(esk6_3(X1,esk2_0,X2),np__3)|esk2_0=X2|v1_xboole_0(X1)|k3_finseq_1(X2)!=np__3|k1_matrix_1(X2)!=np__3|~v1_matrix_1(X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_4858, plain, (esk8_2(X1,np__3)=np__3|esk8_2(X1,np__3)=np__2|esk8_2(X1,np__3)=np__1|~r2_hidden(X1,a_1_0_finseq_1(np__3)))).#trainneg
cnf(i_0_5241, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k11_finseq_1(X2,X3,X1),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_4041, negated_conjecture, (r1_xxreal_0(k1_xboole_0,esk9_1(k4_finseq_1(esk2_0)))|r1_xxreal_0(esk9_1(k4_finseq_1(esk2_0)),k1_xboole_0))).#trainneg
cnf(i_0_5929, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__1,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_4933, plain, (r1_xxreal_0(np__3,k3_finseq_1(k1_xboole_0))|r1_xxreal_0(k3_finseq_1(k1_xboole_0),np__3)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_4934, plain, (r1_xxreal_0(k1_xboole_0,k3_finseq_1(k1_xboole_0))|r1_xxreal_0(k3_finseq_1(k1_xboole_0),k1_xboole_0)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_882, plain, (esk4_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|v1_xboole_0(k2_matrix_1(X2))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(k2_tarski(k1_tarski(X3),k2_tarski(X4,X3)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_887, plain, (r2_hidden(X1,k4_finseq_1(X2))|v1_xboole_0(k2_matrix_1(X2))|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(k2_tarski(k2_tarski(X3,X1),k1_tarski(X1)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_1930, negated_conjecture, (r2_hidden(X1,k4_finseq_1(esk2_0))|~m1_subset_1(k2_tarski(k2_tarski(X2,X1),k1_tarski(X1)),k2_matrix_1(esk2_0)))).#trainneg
cnf(i_0_5780, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(k1_xboole_0))|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2658, plain, (v7_ordinal1(k3_finseq_1(X1))|~v1_xboole_0(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3823, negated_conjecture, (k11_finseq_1(k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k1_funct_1(esk2_0,np__2),k1_funct_1(esk2_0,np__3))!=esk2_0|k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_893, plain, (v1_xboole_0(k2_matrix_1(X1))|~v1_xboole_0(k2_finseq_1(k1_matrix_1(X1)))|~v1_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_896, plain, (v4_relat_1(esk9_1(k1_zfmisc_1(k2_matrix_1(X1))),k4_finseq_1(X1))|~v1_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_911, plain, (v1_xcmplx_0(X1)|~v1_valued_0(X3)|k1_funct_1(X3,np__1)!=X2|k1_funct_1(X3,np__2)!=X1|k3_finseq_1(X3)!=np__3|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_5511, plain, (m1_subset_1(k1_card_1(X1),k1_enumset1(np__1,np__2,np__3))|~r1_xxreal_0(np__1,k1_card_1(X1))|~r1_xxreal_0(k1_card_1(X1),np__3)|~v1_finseq_1(k1_card_1(X1))|~v1_funct_1(k1_card_1(X1))|~v1_relat_1(k1_card_1(X1)))).#trainneg
cnf(i_0_708, plain, (m1_subset_1(k1_nat_1(X1,X2),k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1929, plain, (X1=k3_matrix_1(X2,X3,X4,X5)|X1!=k1_funct_1(X6,X5)|X6!=k1_funct_1(X3,X4)|~v1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X6,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~r2_hidden(X4,k4_finseq_1(X3))|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,k1_matrix_1(X3))|~v7_ordinal1(k1_matrix_1(X3)))).#trainneg
cnf(i_0_1216, plain, (k1_funct_1(esk4_5(X1,X2,X3,X3,X4),X3)=X4|v1_xboole_0(k2_finseq_1(k1_matrix_1(X2)))|X4!=k3_matrix_1(X1,X2,X3,X3)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2))|~m1_subset_1(X3,k2_finseq_1(k1_matrix_1(X2))))).#trainneg
cnf(i_0_2696, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X1)=k3_matrix_1(X2,esk2_0,X1,X1)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m1_subset_1(k3_matrix_1(X2,esk2_0,X1,X1),X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_5781, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(k1_xboole_0))|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3536, plain, (v1_xreal_0(k3_finseq_1(X1))|~v1_xboole_0(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_913, plain, (v1_xreal_0(X1)|~v3_valued_0(X3)|k1_funct_1(X3,np__1)!=X1|k1_funct_1(X3,np__2)!=X2|k3_finseq_1(X3)!=np__3|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_3762, plain, (m1_subset_1(X1,k1_xboole_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_6095, negated_conjecture, (v1_xcmplx_0(k3_matrix_1(esk1_0,esk2_0,np__3,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3))|~v3_valued_0(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_3583, plain, (v7_ordinal1(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v6_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_915, plain, (v4_relat_1(k3_finseq_4(X1,X2,X3,X4),k5_numbers)|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_916, plain, (v5_relat_1(esk9_1(k1_zfmisc_1(k2_matrix_1(X1))),k2_finseq_1(k1_matrix_1(X1)))|~v1_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6129, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X2)=k3_matrix_1(esk1_0,esk2_0,X1,X2)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,X1),esk1_0)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(k1_funct_1(k1_funct_1(esk2_0,X1),X2),esk1_0))).#trainneg
cnf(i_0_6390, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X1)=k3_matrix_1(esk1_0,esk2_0,X1,X1)|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_6192, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,X2)|X3!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X3,X2)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(X3,esk1_0)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_6147, negated_conjecture, (k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,np__3)=k1_funct_1(k1_funct_1(esk2_0,X1),np__3)|~m2_subset_1(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,np__3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,np__3))|~r1_xxreal_0(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,np__3),np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(k1_enumset1(np__1,np__2,np__3)))|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_6151, plain, (k3_finseq_1(k3_finseq_4(X1,X2,X3,X4))=np__3|v1_xboole_0(X1)|~m2_finseq_1(k3_finseq_4(X1,X2,X3,X4),X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_6127, plain, (k3_finseq_1(X1)=X2|v1_xboole_0(X3)|X1!=k1_funct_1(X4,X5)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X2)|~m2_finseq_1(X1,X3)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X5)|~m1_matrix_1(X4,X3,X2,X2)|~m1_subset_1(X5,k5_numbers))).#trainneg
cnf(i_0_4220, plain, (v1_xcmplx_0(k3_finseq_1(X1))|~v1_xboole_0(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6228, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__1,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_944, plain, (esk8_2(esk3_4(X1,X2,X3,k2_finseq_1(X4)),X4)=esk3_4(X1,X2,X3,k2_finseq_1(X4))|esk3_4(X1,X2,X3,k2_finseq_1(X4))=X1|esk3_4(X1,X2,X3,k2_finseq_1(X4))=X2|esk3_4(X1,X2,X3,k2_finseq_1(X4))=X3|k2_finseq_1(X4)=k1_enumset1(X1,X2,X3)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_945, plain, (esk8_2(X1,X2)=X1|v1_xboole_0(k2_finseq_1(X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k2_finseq_1(X2)))).#trainneg
cnf(i_0_976, plain, (r2_hidden(X1,X2)|v1_xboole_0(k2_zfmisc_1(X4,X2))|~m1_subset_1(k2_tarski(k1_tarski(X3),k2_tarski(X1,X3)),k2_zfmisc_1(X4,X2)))).#trainneg
cnf(i_0_6389, negated_conjecture, (k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,X1)=k1_funct_1(k1_funct_1(esk2_0,X1),X1)|~m2_finseq_1(esk2_0,k3_finseq_2(k1_enumset1(np__1,np__2,np__3)))|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r1_xxreal_0(np__1,k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,X1))|~r1_xxreal_0(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,X1),np__3)|~m1_subset_1(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,X1),k5_numbers))).#trainneg
cnf(i_0_6214, plain, (esk8_2(X1,np__3)=np__1|esk8_2(X1,np__3)=np__2|esk8_2(X1,np__3)=np__3|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3))).#trainneg
cnf(i_0_6128, plain, (k3_finseq_1(X1)=X2|v1_xboole_0(X3)|X1!=k1_funct_1(X4,X5)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X2)|~m2_finseq_1(X1,X3)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X5)|~m1_matrix_1(X4,X3,X2,X2)|~m1_subset_1(X5,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_6210, plain, (esk8_2(X1,np__3)=np__1|esk8_2(X1,np__3)=np__2|esk8_2(X1,np__3)=np__3|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_6150, negated_conjecture, (k3_matrix_1(k2_finseq_1(X1),esk2_0,X2,np__3)=k1_funct_1(k1_funct_1(esk2_0,X2),np__3)|~r1_xxreal_0(np__1,k3_matrix_1(k2_finseq_1(X1),esk2_0,X2,np__3))|~r1_xxreal_0(k3_matrix_1(k2_finseq_1(X1),esk2_0,X2,np__3),X1)|~m2_finseq_1(esk2_0,k3_finseq_2(k2_finseq_1(X1)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X2,k4_finseq_1(esk2_0))|~m1_subset_1(k3_matrix_1(k2_finseq_1(X1),esk2_0,X2,np__3),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_6346, plain, (X1=k3_matrix_1(X2,X3,X4,k1_matrix_1(X3))|X1!=k1_funct_1(X5,k1_matrix_1(X3))|X5!=k1_funct_1(X3,X4)|~m2_subset_1(k1_matrix_1(X3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(X3))|~v1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(k1_matrix_1(X3))|~v7_ordinal1(X4)|~r2_hidden(X4,k4_finseq_1(X3)))).#trainneg
cnf(i_0_6138, plain, (X1=X2|m2_finseq_1(k1_funct_1(X1,esk6_3(X3,X1,X2)),k1_enumset1(np__1,np__2,np__3))|v1_xboole_0(X3)|k3_finseq_1(X1)!=k3_finseq_1(X2)|k1_matrix_1(X1)!=k1_matrix_1(X2)|~v1_matrix_1(X1)|~v1_matrix_1(X2)|~m2_finseq_1(X1,k3_finseq_2(k1_enumset1(np__1,np__2,np__3)))|~m2_finseq_1(X2,k3_finseq_2(X3))|~m2_finseq_1(X1,k3_finseq_2(X3))|~r1_xxreal_0(np__1,k3_matrix_1(k1_enumset1(np__1,np__2,np__3),X1,esk6_3(X3,X1,X2),esk7_3(X3,X1,X2)))|~r1_xxreal_0(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),X1,esk6_3(X3,X1,X2),esk7_3(X3,X1,X2)),np__3)|~m1_subset_1(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),X1,esk6_3(X3,X1,X2),esk7_3(X3,X1,X2)),k5_numbers))).#trainneg
cnf(i_0_6125, plain, (k3_finseq_1(X1)=X2|v1_xboole_0(X3)|X1!=k1_funct_1(X4,np__1)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,X3)|~m1_subset_1(X2,k5_numbers)|~m1_matrix_1(X4,X3,X2,X2))).#trainneg
cnf(i_0_6149, negated_conjecture, (k3_matrix_1(k2_finseq_1(X1),esk2_0,X2,np__3)=k1_funct_1(k1_funct_1(esk2_0,X2),np__3)|~r1_xxreal_0(np__1,k3_matrix_1(k2_finseq_1(X1),esk2_0,X2,np__3))|~r1_xxreal_0(k3_matrix_1(k2_finseq_1(X1),esk2_0,X2,np__3),X1)|~m2_finseq_1(esk2_0,k3_finseq_2(k2_finseq_1(X1)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X2,k4_finseq_1(esk2_0))|~m1_subset_1(k3_matrix_1(k2_finseq_1(X1),esk2_0,X2,np__3),k5_numbers))).#trainneg
cnf(i_0_978, plain, (v5_relat_1(k2_matrix_1(X1),k2_finseq_1(k1_matrix_1(X1)))|~v1_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6039, negated_conjecture, (v1_xreal_0(k3_matrix_1(esk1_0,esk2_0,np__3,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3))|~m1_subset_1(k1_funct_1(esk2_0,np__3),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_981, plain, (r1_xxreal_0(np__1,X1)|v1_xboole_0(a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2))|~m1_subset_1(X1,a_1_0_finseq_1(X2)))).#trainneg
cnf(i_0_982, plain, (r1_xxreal_0(X1,X2)|v1_xboole_0(a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2))|~m1_subset_1(X1,a_1_0_finseq_1(X2)))).#trainneg
cnf(i_0_4165, negated_conjecture, (r1_xxreal_0(np__1,esk7_3(X1,esk2_0,X2))|esk2_0=X2|v1_xboole_0(X1)|k3_finseq_1(X2)!=np__3|k1_matrix_1(X2)!=np__3|~v1_matrix_1(X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_6201, negated_conjecture, (esk2_0=X1|r2_hidden(np__2,k4_finseq_1(esk2_0))|v1_xboole_0(X2)|esk6_3(X2,esk2_0,X1)=np__1|esk6_3(X2,esk2_0,X1)=np__3|k3_finseq_1(X1)!=np__3|k1_matrix_1(X1)!=np__3|~v1_matrix_1(X1)|~m2_finseq_1(X1,k3_finseq_2(X2))|~m2_finseq_1(esk2_0,k3_finseq_2(X2)))).#trainneg
cnf(i_0_3986, negated_conjecture, (k3_matrix_1(k2_finseq_1(X1),esk2_0,np__3,np__3)=k3_matrix_1(esk1_0,esk2_0,np__3,np__3)|~r1_xxreal_0(np__1,k3_matrix_1(k2_finseq_1(X1),esk2_0,np__3,np__3))|~r1_xxreal_0(k3_matrix_1(k2_finseq_1(X1),esk2_0,np__3,np__3),X1)|~m2_finseq_1(esk2_0,k3_finseq_2(k2_finseq_1(X1)))|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(k2_finseq_1(X1),esk2_0,np__3,np__3),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_3236, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X2),X1)|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_1395, negated_conjecture, (m1_subset_1(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(esk2_0))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_1151, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5782, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(k1_xboole_0))|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1001, plain, (esk9_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=k1_xboole_0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_983, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|v1_xboole_0(a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2))|~m1_subset_1(X1,a_1_0_finseq_1(X2)))).#trainneg
cnf(i_0_6434, negated_conjecture, (v1_xcmplx_0(k3_matrix_1(esk1_0,esk2_0,np__3,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3))|~m1_subset_1(k1_funct_1(esk2_0,np__3),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_4961, negated_conjecture, (r1_xxreal_0(np__2,k3_finseq_1(np__3))|r1_xxreal_0(k3_finseq_1(np__3),np__2)|~m1_subset_1(esk2_0,k5_numbers)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_998, plain, (v1_xxreal_0(X1)|~v2_valued_0(X3)|k1_funct_1(X3,np__1)!=X1|k1_funct_1(X3,np__2)!=X2|k3_finseq_1(X3)!=np__3|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_999, plain, (v1_xcmplx_0(X1)|~v1_valued_0(X3)|k1_funct_1(X3,np__1)!=X1|k1_funct_1(X3,np__2)!=X2|k3_finseq_1(X3)!=np__3|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_6704, negated_conjecture, (X1=X2|m2_finseq_1(k1_funct_1(X1,esk6_3(X3,X1,X2)),k4_finseq_1(esk2_0))|v1_xboole_0(X3)|k3_finseq_1(X1)!=k3_finseq_1(X2)|k1_matrix_1(X1)!=k1_matrix_1(X2)|~v1_matrix_1(X1)|~v1_matrix_1(X2)|~m2_finseq_1(X1,k3_finseq_2(k4_finseq_1(esk2_0)))|~m2_finseq_1(X2,k3_finseq_2(X3))|~m2_finseq_1(X1,k3_finseq_2(X3))|~r2_hidden(k3_matrix_1(k4_finseq_1(esk2_0),X1,esk6_3(X3,X1,X2),esk7_3(X3,X1,X2)),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_6699, negated_conjecture, (k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,np__3)=k1_funct_1(k1_funct_1(esk2_0,X1),np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(k4_finseq_1(esk2_0)))|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0))|~r2_hidden(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,np__3),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_6700, negated_conjecture, (k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,np__3,np__3)=k3_matrix_1(esk1_0,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(k4_finseq_1(esk2_0)))|~r2_hidden(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,np__3,np__3),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_5784, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_zfmisc_1(k1_xboole_0))|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2479, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_membered(k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6268, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__3))|k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3))!=k1_funct_1(esk2_0,np__1))).#trainneg
cnf(i_0_1006, plain, (v1_xboole_0(esk9_1(k1_zfmisc_1(k2_matrix_1(X1))))|~v1_xboole_0(k2_finseq_1(k1_matrix_1(X1)))|~v1_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4962, negated_conjecture, (r1_xxreal_0(np__1,k3_finseq_1(np__3))|r1_xxreal_0(k3_finseq_1(np__3),np__1)|~m1_subset_1(esk2_0,k5_numbers)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_4963, negated_conjecture, (r1_xxreal_0(np__3,k3_finseq_1(np__3))|r1_xxreal_0(k3_finseq_1(np__3),np__3)|~m1_subset_1(esk2_0,k5_numbers)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_6435, plain, (r1_xxreal_0(X1,k3_matrix_1(X2,k11_finseq_1(k11_finseq_1(X3,X4,X5),k11_finseq_1(X6,X7,X8),k11_finseq_1(X9,X10,X11)),X12,X13))|r1_xxreal_0(k3_matrix_1(X2,k11_finseq_1(k11_finseq_1(X3,X4,X5),k11_finseq_1(X6,X7,X8),k11_finseq_1(X9,X10,X11)),X12,X13),X1)|v1_xboole_0(X2)|~v7_ordinal1(X1)|~v6_membered(X2)|~m1_subset_1(X11,X2)|~m1_subset_1(X10,X2)|~m1_subset_1(X9,X2)|~m1_subset_1(X8,X2)|~m1_subset_1(X7,X2)|~m1_subset_1(X6,X2)|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~v7_ordinal1(X13)|~v7_ordinal1(X12))).#trainneg
cnf(i_0_3584, plain, (v1_xreal_0(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_1009, plain, (m1_subset_1(k2_tarski(k2_tarski(X1,X2),k1_tarski(X2)),k2_matrix_1(X3))|~v1_matrix_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(X3)))|~r2_hidden(X2,k4_finseq_1(X3)))).#trainneg
cnf(i_0_6493, negated_conjecture, (k3_finseq_1(k3_finseq_4(k1_enumset1(np__1,np__2,np__3),X1,X2,X3))=np__3|~m1_subset_1(X4,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X5,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X6,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X7,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X8,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_6692, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X1),X2)|X3!=k3_matrix_1(X2,esk2_0,X1,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m1_subset_1(X3,X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_6486, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,X2)|X1!=k1_funct_1(k1_funct_1(esk2_0,X2),X2)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,X2),esk1_0)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_6691, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X1),X2)|X3!=k3_matrix_1(X2,esk2_0,X1,X4)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m1_subset_1(X3,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X4,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_6693, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X1),X2)|X3!=k3_matrix_1(X2,esk2_0,X1,np__2)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m1_subset_1(X3,X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_6694, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X1),X2)|X3!=k3_matrix_1(X2,esk2_0,X1,np__1)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m1_subset_1(X3,X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_6543, plain, (k3_finseq_1(k3_finseq_4(k1_enumset1(X1,X2,X3),X4,X5,X6))=np__3|~m1_subset_1(X7,k1_enumset1(X1,X2,X3))|~m1_subset_1(X8,k1_enumset1(X1,X2,X3))|~m1_subset_1(X9,k1_enumset1(X1,X2,X3))|~m1_subset_1(X10,k1_enumset1(X1,X2,X3))|~m1_subset_1(X11,k1_enumset1(X1,X2,X3))|~m1_subset_1(X6,k1_enumset1(X1,X2,X3))|~m1_subset_1(X5,k1_enumset1(X1,X2,X3))|~m1_subset_1(X4,k1_enumset1(X1,X2,X3)))).#trainneg
cnf(i_0_2497, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_membered(k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4964, negated_conjecture, (r1_xxreal_0(k1_xboole_0,k3_finseq_1(np__3))|r1_xxreal_0(k3_finseq_1(np__3),k1_xboole_0)|~m1_subset_1(esk2_0,k5_numbers)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_3585, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v1_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_1024, plain, (r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X1,X2)),k2_matrix_1(X3))|~v1_matrix_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(X3)))|~r2_hidden(X2,k4_finseq_1(X3)))).#trainneg
cnf(i_0_6702, negated_conjecture, (k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X2)=k1_funct_1(k1_funct_1(esk2_0,X1),X2)|~m2_finseq_1(esk2_0,k3_finseq_2(k4_finseq_1(esk2_0)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X2),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_6701, negated_conjecture, (k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X1)=k1_funct_1(k1_funct_1(esk2_0,X1),X1)|~m2_finseq_1(esk2_0,k3_finseq_2(k4_finseq_1(esk2_0)))|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X1),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_6705, negated_conjecture, (k3_finseq_1(k3_finseq_4(k4_finseq_1(esk2_0),X1,X2,X3))=np__3|~m1_subset_1(X5,k4_finseq_1(esk2_0))|~m1_subset_1(X6,k4_finseq_1(esk2_0))|~m1_subset_1(X7,k4_finseq_1(esk2_0))|~m1_subset_1(X8,k4_finseq_1(esk2_0))|~m1_subset_1(X9,k4_finseq_1(esk2_0))|~m1_subset_1(X3,k4_finseq_1(esk2_0))|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~m1_subset_1(X1,k4_finseq_1(esk2_0))|~r2_hidden(X4,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_2363, plain, (m1_subset_1(k1_card_1(X1),k5_numbers)|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_3283, plain, (m1_subset_1(k3_finseq_1(X1),k5_numbers)|~v1_xboole_0(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1079, plain, (r1_xxreal_0(k1_funct_1(X1,X2),k1_funct_1(X1,X2))|~v2_valued_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6425, plain, (m1_subset_1(k1_xboole_0,k1_xboole_0)|~r1_xxreal_0(np__1,k1_xboole_0)|~r1_xxreal_0(k1_xboole_0,X1)|~v7_ordinal1(X1)|~v1_xboole_0(X1)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_1081, plain, (v3_valued_0(esk9_1(k1_zfmisc_1(k2_matrix_1(X1))))|~v3_membered(k2_finseq_1(k1_matrix_1(X1)))|~v1_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1092, plain, (v4_relat_1(k2_matrix_1(X1),k4_finseq_1(X1))|~v1_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6840, negated_conjecture, (k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,np__3)=k1_funct_1(k1_funct_1(esk2_0,X1),np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(k4_finseq_1(esk2_0)))|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0))|~m1_subset_1(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,np__3),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_7028, negated_conjecture, (k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X2)=k1_funct_1(k1_funct_1(esk2_0,X1),X2)|~m2_finseq_1(esk2_0,k3_finseq_2(k4_finseq_1(esk2_0)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X2),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_6950, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X1)=k3_matrix_1(esk1_0,esk2_0,X1,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,X1),esk1_0)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(k1_funct_1(k1_funct_1(esk2_0,X1),X1),esk1_0))).#trainneg
cnf(i_0_7030, negated_conjecture, (k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X1)=k1_funct_1(k1_funct_1(esk2_0,X1),X1)|~m2_finseq_1(esk2_0,k3_finseq_2(k4_finseq_1(esk2_0)))|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X1),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_7038, negated_conjecture, (k3_finseq_1(k3_finseq_4(k4_finseq_1(esk2_0),X1,X2,X3))=np__3|~m1_subset_1(X4,k4_finseq_1(esk2_0))|~m1_subset_1(X5,k4_finseq_1(esk2_0))|~m1_subset_1(X6,k4_finseq_1(esk2_0))|~m1_subset_1(X7,k4_finseq_1(esk2_0))|~m1_subset_1(X8,k4_finseq_1(esk2_0))|~m1_subset_1(X3,k4_finseq_1(esk2_0))|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~m1_subset_1(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_6770, plain, (v3_membered(k1_xboole_0)|~v3_membered(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1104, plain, (r1_xxreal_0(k1_xboole_0,k1_funct_1(X1,X2))|r1_xxreal_0(k1_funct_1(X1,X2),k1_xboole_0)|~v2_valued_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7089, plain, (r1_xxreal_0(X2,X2)|~v2_valued_0(k11_finseq_1(X1,X2,X3)))).#trainneg
cnf(i_0_7090, plain, (r1_xxreal_0(X1,X1)|~v2_valued_0(k11_finseq_1(X1,X2,X3)))).#trainneg
cnf(i_0_1106, plain, (r1_xxreal_0(np__2,esk9_1(X1))|r1_xxreal_0(esk9_1(X1),np__2)|~v6_membered(X1))).#trainneg
cnf(i_0_1107, plain, (r1_xxreal_0(np__2,esk7_3(X1,X2,X3))|r1_xxreal_0(esk7_3(X1,X2,X3),np__2)|X2=X3|v1_xboole_0(X1)|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_6975, negated_conjecture, (k3_finseq_1(k3_finseq_4(k1_enumset1(np__1,np__2,np__3),X1,X2,X3))=np__3|~m1_subset_1(X4,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X5,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X6,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X7,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_7034, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X1),k4_finseq_1(esk2_0))|~m2_finseq_1(esk2_0,k3_finseq_2(k4_finseq_1(esk2_0)))|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X1),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_6974, plain, (k3_finseq_1(k3_finseq_4(k1_enumset1(X1,X2,X3),X4,X5,X6))=np__3|~m1_subset_1(X7,k1_enumset1(X1,X2,X3))|~m1_subset_1(X8,k1_enumset1(X1,X2,X3))|~m1_subset_1(X9,k1_enumset1(X1,X2,X3))|~m1_subset_1(X10,k1_enumset1(X1,X2,X3))|~m1_subset_1(X6,k1_enumset1(X1,X2,X3))|~m1_subset_1(X5,k1_enumset1(X1,X2,X3))|~m1_subset_1(X4,k1_enumset1(X1,X2,X3)))).#trainneg
cnf(i_0_7049, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X2)=k3_matrix_1(esk1_0,esk2_0,X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_6838, negated_conjecture, (X1=X2|m2_finseq_1(k1_funct_1(X1,esk6_3(X3,X1,X2)),k4_finseq_1(esk2_0))|v1_xboole_0(X3)|k3_finseq_1(X1)!=k3_finseq_1(X2)|k1_matrix_1(X1)!=k1_matrix_1(X2)|~v1_matrix_1(X1)|~v1_matrix_1(X2)|~m2_finseq_1(X1,k3_finseq_2(k4_finseq_1(esk2_0)))|~m2_finseq_1(X2,k3_finseq_2(X3))|~m2_finseq_1(X1,k3_finseq_2(X3))|~m1_subset_1(k3_matrix_1(k4_finseq_1(esk2_0),X1,esk6_3(X3,X1,X2),esk7_3(X3,X1,X2)),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_6841, negated_conjecture, (k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,np__3)=k1_funct_1(k1_funct_1(esk2_0,X1),np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(k4_finseq_1(esk2_0)))|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0))|~m2_subset_1(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,np__3),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,np__3))|~r1_xxreal_0(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,np__3),np__3))).#trainneg
cnf(i_0_6949, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X1),X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m1_subset_1(k3_matrix_1(X2,esk2_0,X1,np__3),X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_6765, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1108, plain, (r1_xxreal_0(np__2,esk6_3(X1,X2,X3))|r1_xxreal_0(esk6_3(X1,X2,X3),np__2)|X2=X3|v1_xboole_0(X1)|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_7091, plain, (r1_xxreal_0(X3,X3)|~v2_valued_0(k11_finseq_1(X1,X2,X3)))).#trainneg
cnf(i_0_1109, plain, (r1_xxreal_0(np__2,esk5_2(X1,X2))|r1_xxreal_0(esk5_2(X1,X2),np__2)|X1=X2|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1110, plain, (r1_xxreal_0(np__2,k3_xcmplx_0(X1,X2))|r1_xxreal_0(k3_xcmplx_0(X1,X2),np__2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6973, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X1),X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m1_subset_1(k3_matrix_1(X2,esk2_0,X1,np__1),X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_6972, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X1),X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m1_subset_1(k3_matrix_1(X2,esk2_0,X1,np__2),X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_6842, negated_conjecture, (k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,np__3,np__3)=k3_matrix_1(esk1_0,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(k4_finseq_1(esk2_0)))|~m1_subset_1(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,np__3,np__3),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_7029, negated_conjecture, (k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X2)=k1_funct_1(k1_funct_1(esk2_0,X1),X2)|~m2_finseq_1(esk2_0,k3_finseq_2(k4_finseq_1(esk2_0)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~m2_subset_1(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X2))|~r1_xxreal_0(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X2),np__3))).#trainneg
cnf(i_0_6571, plain, (r1_xxreal_0(X1,X2)|v1_xboole_0(k2_finseq_1(X2))|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2))|~m1_subset_1(X1,k2_finseq_1(X2)))).#trainneg
cnf(i_0_1112, plain, (r1_xxreal_0(np__2,k2_xcmplx_0(X1,X2))|r1_xxreal_0(k2_xcmplx_0(X1,X2),np__2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6914, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v1_xreal_0(X13)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_1113, plain, (r1_xxreal_0(np__2,k1_matrix_1(X1))|r1_xxreal_0(k1_matrix_1(X1),np__2)|~v1_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1118, plain, (r1_xxreal_0(np__1,esk9_1(X1))|r1_xxreal_0(esk9_1(X1),np__1)|~v6_membered(X1))).#trainneg
cnf(i_0_6668, plain, (r1_xxreal_0(X1,X2)|v1_xboole_0(k2_finseq_1(X2))|~m1_subset_1(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5764, plain, (~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~v1_xboole_0(k2_finseq_1(X2))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_646, negated_conjecture, (m1_subset_1(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(esk2_0))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_3810, negated_conjecture, (k11_finseq_1(k1_funct_1(esk2_0,np__1),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__1)))).#trainneg
cnf(i_0_1119, plain, (r1_xxreal_0(np__1,esk7_3(X1,X2,X3))|r1_xxreal_0(esk7_3(X1,X2,X3),np__1)|X2=X3|v1_xboole_0(X1)|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_7108, plain, (m1_subset_1(k1_xboole_0,k1_xboole_0)|~r1_xxreal_0(np__1,k1_xboole_0)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_7373, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v1_xboole_0(X13))).#trainneg
cnf(i_0_7366, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v6_membered(X13))).#trainneg
cnf(i_0_1120, plain, (r1_xxreal_0(np__1,esk6_3(X1,X2,X3))|r1_xxreal_0(esk6_3(X1,X2,X3),np__1)|X2=X3|v1_xboole_0(X1)|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_3815, negated_conjecture, (k11_finseq_1(k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k1_funct_1(esk2_0,np__2),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__2)))).#trainneg
cnf(i_0_1122, plain, (r1_xxreal_0(np__1,k3_xcmplx_0(X1,X2))|r1_xxreal_0(k3_xcmplx_0(X1,X2),np__1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7367, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v3_membered(X13))).#trainneg
cnf(i_0_7368, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v7_ordinal1(X14)|~v7_ordinal1(X13))).#trainneg
cnf(i_0_1124, plain, (r1_xxreal_0(np__1,k2_xcmplx_0(X1,X2))|r1_xxreal_0(k2_xcmplx_0(X1,X2),np__1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1125, plain, (r1_xxreal_0(np__1,k1_matrix_1(X1))|r1_xxreal_0(k1_matrix_1(X1),np__1)|~v1_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7206, plain, (k3_finseq_1(k3_finseq_4(k1_enumset1(X1,X2,X3),X4,X5,X6))=np__3|~m1_subset_1(X7,k1_enumset1(X1,X2,X3))|~m1_subset_1(X8,k1_enumset1(X1,X2,X3))|~m1_subset_1(X9,k1_enumset1(X1,X2,X3))|~m1_subset_1(X6,k1_enumset1(X1,X2,X3))|~m1_subset_1(X5,k1_enumset1(X1,X2,X3))|~m1_subset_1(X4,k1_enumset1(X1,X2,X3)))).#trainneg
cnf(i_0_7130, negated_conjecture, (k3_finseq_1(k3_finseq_4(k4_finseq_1(esk2_0),X1,X2,X3))=np__3|~m1_subset_1(X4,k4_finseq_1(esk2_0))|~m1_subset_1(X5,k4_finseq_1(esk2_0))|~m1_subset_1(X6,k4_finseq_1(esk2_0))|~m1_subset_1(X7,k4_finseq_1(esk2_0))|~m1_subset_1(X3,k4_finseq_1(esk2_0))|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~m1_subset_1(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_7349, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X1),k4_finseq_1(esk2_0))|~m2_finseq_1(esk2_0,k3_finseq_2(k4_finseq_1(esk2_0)))|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,np__2),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_7311, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X1),k4_finseq_1(esk2_0))|~m2_finseq_1(esk2_0,k3_finseq_2(k4_finseq_1(esk2_0)))|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,np__3),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_2370, plain, (v7_ordinal1(k3_finseq_1(X1))|~v1_xboole_0(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3821, negated_conjecture, (k11_finseq_1(k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k1_funct_1(esk2_0,np__3))!=esk2_0|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_1145, plain, (esk4_5(X1,X2,X3,X3,X4)=k1_funct_1(X2,X3)|v1_xboole_0(k2_finseq_1(k1_matrix_1(X2)))|X4!=k3_matrix_1(X1,X2,X3,X3)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2))|~m1_subset_1(X3,k2_finseq_1(k1_matrix_1(X2))))).#trainneg
cnf(i_0_4624, plain, (r1_xxreal_0(np__2,esk9_1(k1_enumset1(np__1,np__2,np__3)))|r1_xxreal_0(esk9_1(k1_enumset1(np__1,np__2,np__3)),np__2))).#trainneg
cnf(i_0_4627, plain, (r1_xxreal_0(k1_xboole_0,esk9_1(k1_enumset1(np__1,np__2,np__3)))|r1_xxreal_0(esk9_1(k1_enumset1(np__1,np__2,np__3)),k1_xboole_0))).#trainneg
cnf(i_0_7377, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v3_valued_0(X13)|~v1_funct_1(X13)|~v1_relat_1(X13))).#trainneg
cnf(i_0_1132, plain, (esk4_5(X1,X2,X3,esk3_4(X4,X5,X6,k2_finseq_1(k1_matrix_1(X2))),X7)=k1_funct_1(X2,X3)|esk3_4(X4,X5,X6,k2_finseq_1(k1_matrix_1(X2)))=X4|esk3_4(X4,X5,X6,k2_finseq_1(k1_matrix_1(X2)))=X5|esk3_4(X4,X5,X6,k2_finseq_1(k1_matrix_1(X2)))=X6|k2_finseq_1(k1_matrix_1(X2))=k1_enumset1(X4,X5,X6)|X7!=k3_matrix_1(X1,X2,X3,esk3_4(X4,X5,X6,k2_finseq_1(k1_matrix_1(X2))))|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X7,X1)|~v7_ordinal1(esk3_4(X4,X5,X6,k2_finseq_1(k1_matrix_1(X2))))|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2)))).#trainneg
cnf(i_0_1144, plain, (esk4_5(X1,X2,esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2))),esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2))),X6)=k1_funct_1(X2,esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2))))|esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2)))=X3|esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2)))=X4|esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2)))=X5|k2_finseq_1(k1_matrix_1(X2))=k1_enumset1(X3,X4,X5)|X6!=k3_matrix_1(X1,X2,esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2))),esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2))))|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X6,X1)|~v7_ordinal1(esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2))))|~r2_hidden(esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2))),k4_finseq_1(X2)))).#trainneg
cnf(i_0_7343, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X1),k4_finseq_1(esk2_0))|~m2_finseq_1(esk2_0,k3_finseq_2(k4_finseq_1(esk2_0)))|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,np__1),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_7204, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X1),k4_finseq_1(esk2_0))|~m2_finseq_1(esk2_0,k3_finseq_2(k4_finseq_1(esk2_0)))|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X1),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_7351, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X1),k1_enumset1(np__1,np__2,np__3))|~m2_finseq_1(esk2_0,k3_finseq_2(k1_enumset1(np__1,np__2,np__3)))|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r1_xxreal_0(np__1,k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,np__2))|~r1_xxreal_0(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,np__2),np__3)|~m1_subset_1(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,np__2),k5_numbers))).#trainneg
cnf(i_0_7313, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X1),k1_enumset1(np__1,np__2,np__3))|~m2_finseq_1(esk2_0,k3_finseq_2(k1_enumset1(np__1,np__2,np__3)))|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r1_xxreal_0(np__1,k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,np__3))|~r1_xxreal_0(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,np__3),np__3)|~m1_subset_1(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,np__3),k5_numbers))).#trainneg
cnf(i_0_7345, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X1),k1_enumset1(np__1,np__2,np__3))|~m2_finseq_1(esk2_0,k3_finseq_2(k1_enumset1(np__1,np__2,np__3)))|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~r1_xxreal_0(np__1,k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,np__1))|~r1_xxreal_0(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,np__1),np__3)|~m1_subset_1(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,np__1),k5_numbers))).#trainneg
cnf(i_0_7348, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(esk2_0,k3_finseq_2(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~m2_finseq_1(k3_matrix_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)),esk2_0,X1,np__2),X2))).#trainneg
cnf(i_0_6228, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__1,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_1156, plain, (r1_xxreal_0(k1_xboole_0,esk9_1(X1))|r1_xxreal_0(esk9_1(X1),k1_xboole_0)|~v6_membered(X1))).#trainneg
cnf(i_0_5779, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_7462, negated_conjecture, (~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_xboole_0(k1_funct_1(esk2_0,np__2)))).#trainneg
cnf(i_0_1157, plain, (r1_xxreal_0(k1_xboole_0,esk7_3(X1,X2,X3))|r1_xxreal_0(esk7_3(X1,X2,X3),k1_xboole_0)|X2=X3|v1_xboole_0(X1)|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_1158, plain, (r1_xxreal_0(k1_xboole_0,esk6_3(X1,X2,X3))|r1_xxreal_0(esk6_3(X1,X2,X3),k1_xboole_0)|X2=X3|v1_xboole_0(X1)|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_7653, negated_conjecture, (esk4_5(X1,esk2_0,X2,X2,X3)=k1_funct_1(esk2_0,X2)|X3!=k3_matrix_1(X1,esk2_0,X2,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_7577, negated_conjecture, (v1_relat_1(k1_funct_1(esk2_0,X1))|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_7618, negated_conjecture, (v1_finseq_1(k1_funct_1(esk2_0,X1))|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_7568, negated_conjecture, (v1_funct_1(k1_funct_1(esk2_0,X1))|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_7581, negated_conjecture, (v1_relat_1(k1_funct_1(esk2_0,X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3))).#trainneg
cnf(i_0_7598, negated_conjecture, (k3_finseq_1(k3_finseq_4(k4_finseq_1(esk2_0),X1,X2,X3))=np__3|~m1_subset_1(X4,k4_finseq_1(esk2_0))|~m1_subset_1(X5,k4_finseq_1(esk2_0))|~m1_subset_1(X6,k4_finseq_1(esk2_0))|~m1_subset_1(X3,k4_finseq_1(esk2_0))|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~m1_subset_1(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_7572, negated_conjecture, (v1_funct_1(k1_funct_1(esk2_0,X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3))).#trainneg
cnf(i_0_7622, negated_conjecture, (v1_finseq_1(k1_funct_1(esk2_0,X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3))).#trainneg
cnf(i_0_7564, negated_conjecture, (v1_funct_1(k1_funct_1(esk2_0,esk9_1(k4_finseq_1(esk2_0)))))).#trainneg
cnf(i_0_1159, plain, (r1_xxreal_0(k1_xboole_0,esk5_2(X1,X2))|r1_xxreal_0(esk5_2(X1,X2),k1_xboole_0)|X1=X2|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7498, negated_conjecture, (~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3))|~v1_xboole_0(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_1160, plain, (r1_xxreal_0(k1_xboole_0,k3_xcmplx_0(X1,X2))|r1_xxreal_0(k3_xcmplx_0(X1,X2),k1_xboole_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1162, plain, (r1_xxreal_0(k1_xboole_0,k2_xcmplx_0(X1,X2))|r1_xxreal_0(k2_xcmplx_0(X1,X2),k1_xboole_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3810, negated_conjecture, (k11_finseq_1(k1_funct_1(esk2_0,np__1),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__1)))).#trainneg
cnf(i_0_7582, plain, (k3_finseq_1(k3_finseq_4(k1_enumset1(X1,X2,X3),X4,X5,X6))=np__3|~m1_subset_1(X7,k1_enumset1(X1,X2,X3))|~m1_subset_1(X8,k1_enumset1(X1,X2,X3))|~m1_subset_1(X6,k1_enumset1(X1,X2,X3))|~m1_subset_1(X5,k1_enumset1(X1,X2,X3))|~m1_subset_1(X4,k1_enumset1(X1,X2,X3)))).#trainneg
cnf(i_0_3468, negated_conjecture, (k11_finseq_1(k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k1_funct_1(esk2_0,np__2),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__2)))).#trainneg
cnf(i_0_7513, negated_conjecture, (~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__1))|~v1_xboole_0(k1_funct_1(esk2_0,np__1)))).#trainneg
cnf(i_0_3224, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X2),X3)=X4|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k4_finseq_1(esk2_0))|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_3534, plain, (v1_xreal_0(k3_finseq_1(X1))|~v1_xboole_0(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7573, negated_conjecture, (v1_relat_1(k1_funct_1(esk2_0,esk9_1(k4_finseq_1(esk2_0)))))).#trainneg
cnf(i_0_6268, negated_conjecture, (k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3))!=k1_funct_1(esk2_0,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__3))!=np__3|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_1163, plain, (r1_xxreal_0(k1_xboole_0,k1_matrix_1(X1))|r1_xxreal_0(k1_matrix_1(X1),k1_xboole_0)|~v1_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4889, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7463, negated_conjecture, (v1_finset_1(np__3)|~v1_finset_1(k1_funct_1(esk2_0,np__2))|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__2)))).#trainneg
cnf(i_0_7499, negated_conjecture, (v1_finset_1(np__3)|~v1_finset_1(k1_funct_1(esk2_0,np__3))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_1175, plain, (r1_xxreal_0(np__3,esk9_1(X1))|r1_xxreal_0(esk9_1(X1),np__3)|~v6_membered(X1))).#trainneg
cnf(i_0_1176, plain, (r1_xxreal_0(np__3,esk7_3(X1,X2,X3))|r1_xxreal_0(esk7_3(X1,X2,X3),np__3)|X2=X3|v1_xboole_0(X1)|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_1129, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4346, negated_conjecture, (esk2_0=X1|r2_hidden(k2_tarski(k1_tarski(esk6_3(X2,esk2_0,X1)),k2_tarski(esk6_3(X2,esk2_0,X1),np__2)),k2_matrix_1(esk2_0))|v1_xboole_0(X2)|esk7_3(X2,esk2_0,X1)=np__1|esk7_3(X2,esk2_0,X1)=np__3|k3_finseq_1(X1)!=np__3|k1_matrix_1(X1)!=np__3|~v1_matrix_1(X1)|~m2_finseq_1(X1,k3_finseq_2(X2))|~m2_finseq_1(esk2_0,k3_finseq_2(X2)))).#trainneg
cnf(i_0_1827, plain, (k1_matrix_1(k11_finseq_1(k11_finseq_1(X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k3_finseq_4(X1,X8,X9,X10)))=np__3|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_6228, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(np__1,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_7963, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7983, negated_conjecture, (v1_finset_1(np__3)|~v1_xboole_0(k1_funct_1(esk2_0,np__2)))).#trainneg
cnf(i_0_1177, plain, (r1_xxreal_0(np__3,esk6_3(X1,X2,X3))|r1_xxreal_0(esk6_3(X1,X2,X3),np__3)|X2=X3|v1_xboole_0(X1)|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_7462, negated_conjecture, (~v1_xboole_0(k1_funct_1(esk2_0,np__2)))).#trainneg
cnf(i_0_7541, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v7_ordinal1(X16))).#trainneg
cnf(i_0_1178, plain, (r1_xxreal_0(np__3,esk5_2(X1,X2))|r1_xxreal_0(esk5_2(X1,X2),np__3)|X1=X2|k3_finseq_1(X1)!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1179, plain, (r1_xxreal_0(np__3,k3_xcmplx_0(X1,X2))|r1_xxreal_0(k3_xcmplx_0(X1,X2),np__3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7785, negated_conjecture, (esk4_5(X1,esk2_0,X2,X2,X3)=k1_funct_1(esk2_0,X2)|X3!=k3_matrix_1(X1,esk2_0,X2,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~m1_subset_1(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_7906, negated_conjecture, (k11_finseq_1(k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__2),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2))).#trainneg
cnf(i_0_7908, plain, (k3_finseq_1(k3_finseq_4(k1_enumset1(X1,X2,X3),X4,X5,X6))=np__3|~m1_subset_1(X7,k1_enumset1(X1,X2,X3))|~m1_subset_1(X6,k1_enumset1(X1,X2,X3))|~m1_subset_1(X5,k1_enumset1(X1,X2,X3))|~m1_subset_1(X4,k1_enumset1(X1,X2,X3)))).#trainneg
cnf(i_0_7904, negated_conjecture, (k11_finseq_1(k1_funct_1(esk2_0,np__1),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k1_funct_1(esk2_0,np__3))!=esk2_0|k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2))).#trainneg
cnf(i_0_3815, negated_conjecture, (k11_finseq_1(k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k1_funct_1(esk2_0,np__2),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__2)))).#trainneg
cnf(i_0_7843, negated_conjecture, (k3_finseq_1(k3_finseq_4(k4_finseq_1(esk2_0),X1,X2,X3))=np__3|~m1_subset_1(X4,k4_finseq_1(esk2_0))|~m1_subset_1(X5,k4_finseq_1(esk2_0))|~m1_subset_1(X3,k4_finseq_1(esk2_0))|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~m1_subset_1(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8037, negated_conjecture, (r2_hidden(esk6_3(X1,esk2_0,X2),k4_finseq_1(esk2_0))|esk7_3(X1,esk2_0,X2)=np__3|esk7_3(X1,esk2_0,X2)=np__1|esk2_0=X2|v1_xboole_0(X1)|k3_finseq_1(X2)!=np__3|k1_matrix_1(X2)!=np__3|~v1_matrix_1(X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_3821, negated_conjecture, (k11_finseq_1(k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k1_funct_1(esk2_0,np__3))!=esk2_0|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_7928, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X2)=k3_matrix_1(X3,esk2_0,X1,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X3))|~m1_subset_1(k3_matrix_1(X3,esk2_0,X1,X2),X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_1181, plain, (r1_xxreal_0(np__3,k2_xcmplx_0(X1,X2))|r1_xxreal_0(k2_xcmplx_0(X1,X2),np__3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5714, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7984, negated_conjecture, (v1_finset_1(np__3)|~m1_subset_1(k1_funct_1(esk2_0,np__2),k5_numbers))).#trainneg
cnf(i_0_7986, negated_conjecture, (v1_finset_1(np__3)|~m1_subset_1(k1_funct_1(esk2_0,np__3),k5_numbers))).#trainneg
cnf(i_0_1182, plain, (r1_xxreal_0(np__3,k1_matrix_1(X1))|r1_xxreal_0(k1_matrix_1(X1),np__3)|~v1_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1776, negated_conjecture, (k11_finseq_1(k1_funct_1(k1_funct_1(esk2_0,np__3),np__1),X1,X2)=k1_funct_1(esk2_0,np__3)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=X1|X2!=k3_matrix_1(X3,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X3))|~m1_subset_1(X2,X3)|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_5160, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X2),X3)=X4|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0))|~m1_subset_1(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_2518, negated_conjecture, (X1=k3_matrix_1(X2,esk2_0,X3,X4)|X5!=k1_funct_1(esk2_0,X3)|X1!=k1_funct_1(X5,X4)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~m1_subset_1(X1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(esk2_0))|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_4218, plain, (v1_xcmplx_0(k3_finseq_1(X1))|~v1_xboole_0(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4929, plain, (r1_xxreal_0(X1,k3_finseq_1(k1_xboole_0))|r1_xxreal_0(k3_finseq_1(k1_xboole_0),X1)|~v7_ordinal1(X1)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_7775, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(k3_finseq_2(esk1_0)))).#trainneg
cnf(i_0_7514, negated_conjecture, (v1_finset_1(np__3)|~v1_finset_1(k1_funct_1(esk2_0,np__1))|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__1)))).#trainneg
cnf(i_0_7529, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11)|~v1_finseq_1(X13)|~v1_funct_1(X13)|~v1_relat_1(X13))).#trainneg
cnf(i_0_1195, plain, (m2_subset_1(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),k1_numbers,k5_numbers)|~m1_subset_1(k3_xcmplx_0(X3,X2),k5_numbers)|~v7_ordinal1(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1215, plain, (k1_funct_1(esk4_5(X1,X2,esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2))),esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2))),X6),esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2))))=X6|esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2)))=X3|esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2)))=X4|esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2)))=X5|k2_finseq_1(k1_matrix_1(X2))=k1_enumset1(X3,X4,X5)|X6!=k3_matrix_1(X1,X2,esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2))),esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2))))|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X6,X1)|~v7_ordinal1(esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2))))|~r2_hidden(esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2))),k4_finseq_1(X2)))).#trainneg
cnf(i_0_6124, plain, (k3_finseq_1(X1)=X2|v1_xboole_0(X3)|X1!=k1_funct_1(X4,np__2)|~r1_xxreal_0(np__2,X2)|~m2_finseq_1(X1,X3)|~m1_subset_1(X2,k5_numbers)|~m1_matrix_1(X4,X3,X2,X2))).#trainneg
cnf(i_0_7310, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(esk2_0,k3_finseq_2(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~m2_finseq_1(k3_matrix_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)),esk2_0,X1,np__3),X2))).#trainneg
cnf(i_0_1763, plain, (k1_funct_1(k1_funct_1(X2,X3),k1_matrix_1(X2))=X4|X4!=k3_matrix_1(X1,X2,X3,k1_matrix_1(X2))|~m2_subset_1(k1_matrix_1(X2),k1_numbers,k5_numbers)|~r1_xxreal_0(k1_matrix_1(X2),k1_matrix_1(X2))|~r1_xxreal_0(np__1,k1_matrix_1(X2))|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(k1_matrix_1(X2))|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2)))).#trainneg
cnf(i_0_8254, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,X3)|X4!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X4,X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(X4,esk1_0)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8224, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2575, plain, (v7_ordinal1(X1)|~v6_membered(a_1_0_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_8266, negated_conjecture, (v1_finset_1(np__3)|~m1_subset_1(k1_funct_1(esk2_0,np__1),k5_numbers))).#trainneg
cnf(i_0_1219, plain, (k1_funct_1(esk4_5(X1,X2,X3,esk3_4(X4,X5,X6,k2_finseq_1(k1_matrix_1(X2))),X7),esk3_4(X4,X5,X6,k2_finseq_1(k1_matrix_1(X2))))=X7|esk3_4(X4,X5,X6,k2_finseq_1(k1_matrix_1(X2)))=X4|esk3_4(X4,X5,X6,k2_finseq_1(k1_matrix_1(X2)))=X5|esk3_4(X4,X5,X6,k2_finseq_1(k1_matrix_1(X2)))=X6|k2_finseq_1(k1_matrix_1(X2))=k1_enumset1(X4,X5,X6)|X7!=k3_matrix_1(X1,X2,X3,esk3_4(X4,X5,X6,k2_finseq_1(k1_matrix_1(X2))))|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X7,X1)|~v7_ordinal1(esk3_4(X4,X5,X6,k2_finseq_1(k1_matrix_1(X2))))|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2)))).#trainneg
cnf(i_0_8215, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X2)=k3_matrix_1(esk1_0,esk2_0,X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8212, negated_conjecture, (k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X2)=k1_funct_1(k1_funct_1(esk2_0,X1),X2)|~m2_finseq_1(esk2_0,k3_finseq_2(k4_finseq_1(esk2_0)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k4_finseq_1(esk2_0))|~r2_hidden(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X2),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_8150, plain, (k3_finseq_1(k3_finseq_4(k1_enumset1(X1,X2,X3),X4,X5,X6))=np__3|~m1_subset_1(X6,k1_enumset1(X1,X2,X3))|~m1_subset_1(X5,k1_enumset1(X1,X2,X3))|~m1_subset_1(X4,k1_enumset1(X1,X2,X3)))).#trainneg
cnf(i_0_8173, negated_conjecture, (k3_finseq_1(k3_finseq_4(k4_finseq_1(esk2_0),X1,X2,X3))=np__3|~m1_subset_1(X4,k4_finseq_1(esk2_0))|~m1_subset_1(X3,k4_finseq_1(esk2_0))|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~m1_subset_1(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8253, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X2)=k3_matrix_1(X3,esk2_0,X1,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X3))|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(k3_matrix_1(X3,esk2_0,X1,X2),X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8357, plain, (k1_funct_1(k1_funct_1(X1,X2),k1_matrix_1(X1))=k3_matrix_1(X3,X1,X2,k1_matrix_1(X1))|~m2_subset_1(k1_matrix_1(X1),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(X1))|~v1_matrix_1(X1)|~m2_finseq_1(X1,k3_finseq_2(X3))|~m1_subset_1(k3_matrix_1(X3,X1,X2,k1_matrix_1(X1)),X3)|~v7_ordinal1(k1_matrix_1(X1))|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(X1)))).#trainneg
cnf(i_0_8214, negated_conjecture, (k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,X2)=k1_funct_1(k1_funct_1(esk2_0,X1),X2)|~m2_finseq_1(esk2_0,k3_finseq_2(k1_enumset1(np__1,np__2,np__3)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k4_finseq_1(esk2_0))|~r1_xxreal_0(np__1,k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,X2))|~r1_xxreal_0(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,X2),np__3)|~m1_subset_1(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,X2),k5_numbers))).#trainneg
cnf(i_0_8126, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X2),X2)=X3|X3!=k3_matrix_1(X1,esk2_0,X2,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0))|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8002, plain, (r1_xxreal_0(esk9_1(k5_numbers),np__1)|r1_xxreal_0(np__1,esk9_1(k5_numbers)))).#trainneg
cnf(i_0_1236, plain, (v1_finset_1(np__3)|~v1_finset_1(X3)|k1_funct_1(X3,np__1)!=X1|k1_funct_1(X3,np__2)!=X2|k3_finseq_1(X3)!=np__3|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_2577, plain, (v1_xcmplx_0(X1)|~v1_membered(a_1_0_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_8244, negated_conjecture, (v7_ordinal1(k3_finseq_1(np__3))|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(k1_funct_1(esk2_0,np__2),k5_numbers))).#trainneg
cnf(i_0_8246, negated_conjecture, (v7_ordinal1(k3_finseq_1(np__3))|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(k1_funct_1(esk2_0,np__3),k5_numbers))).#trainneg
cnf(i_0_1240, plain, (v3_valued_0(k2_matrix_1(X1))|~v3_membered(k2_finseq_1(k1_matrix_1(X1)))|~v1_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8190, negated_conjecture, (m1_subset_1(esk6_3(X1,esk2_0,X2),k4_finseq_1(esk2_0))|esk7_3(X1,esk2_0,X2)=np__1|esk7_3(X1,esk2_0,X2)=np__3|esk2_0=X2|v1_xboole_0(X1)|k3_finseq_1(X2)!=np__3|k1_matrix_1(X2)!=np__3|~v1_matrix_1(X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_8191, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,esk6_3(X1,esk2_0,X2)),X3)=k3_matrix_1(esk1_0,esk2_0,esk6_3(X1,esk2_0,X2),X3)|esk7_3(X1,esk2_0,X2)=np__1|esk7_3(X1,esk2_0,X2)=np__3|esk2_0=X2|v1_xboole_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X3,np__3)|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~v7_ordinal1(esk6_3(X1,esk2_0,X2))|k3_finseq_1(X2)!=np__3|k1_matrix_1(X2)!=np__3|~v1_matrix_1(X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_5772, plain, (esk5_2(k11_finseq_1(X1,X2,X3),X4)=np__1|esk5_2(k11_finseq_1(X1,X2,X3),X4)=np__2|esk5_2(k11_finseq_1(X1,X2,X3),X4)=np__3|k11_finseq_1(X1,X2,X3)=X4|k3_finseq_1(X4)!=np__3|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4))).#trainneg
cnf(i_0_6505, plain, (k3_finseq_1(k3_finseq_4(k5_numbers,X1,X2,X3))=np__3|~m1_subset_1(X5,k5_numbers)|~m1_subset_1(X6,k5_numbers)|~m1_subset_1(X7,k5_numbers)|~m1_subset_1(X8,k5_numbers)|~m1_subset_1(X9,k5_numbers)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_2540, plain, (v7_ordinal1(k1_card_1(k3_finseq_1(X1)))|~v1_finset_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7614, negated_conjecture, (v1_finseq_1(k1_funct_1(esk2_0,esk9_1(k4_finseq_1(esk2_0)))))).#trainneg
cnf(i_0_1392, negated_conjecture, (k11_finseq_1(X1,X2,X3)=k1_funct_1(esk2_0,np__3)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=X1|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=X2|X3!=k3_matrix_1(X4,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X4))|~m1_subset_1(X3,X4)|~v1_finseq_1(k1_funct_1(esk2_0,np__3))|~v1_funct_1(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_1242, plain, (k1_enumset1(np__1,np__2,np__3)=k1_enumset1(X1,X2,X3)|esk3_4(X1,X2,X3,k1_enumset1(np__1,np__2,np__3))!=X3|~m2_subset_1(esk3_4(X1,X2,X3,k1_enumset1(np__1,np__2,np__3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk3_4(X1,X2,X3,k1_enumset1(np__1,np__2,np__3)))|~r1_xxreal_0(esk3_4(X1,X2,X3,k1_enumset1(np__1,np__2,np__3)),np__3))).#trainneg
cnf(i_0_6106, plain, (v7_ordinal1(X1)|~v6_membered(k2_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_8386, negated_conjecture, (v7_ordinal1(k3_finseq_1(np__3))|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(k1_funct_1(esk2_0,np__1),k5_numbers))).#trainneg
cnf(i_0_1243, plain, (k1_enumset1(np__1,np__2,np__3)=k1_enumset1(X1,X2,X3)|esk3_4(X1,X2,X3,k1_enumset1(np__1,np__2,np__3))!=X2|~m2_subset_1(esk3_4(X1,X2,X3,k1_enumset1(np__1,np__2,np__3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk3_4(X1,X2,X3,k1_enumset1(np__1,np__2,np__3)))|~r1_xxreal_0(esk3_4(X1,X2,X3,k1_enumset1(np__1,np__2,np__3)),np__3))).#trainneg
cnf(i_0_1244, plain, (k1_enumset1(np__1,np__2,np__3)=k1_enumset1(X1,X2,X3)|esk3_4(X1,X2,X3,k1_enumset1(np__1,np__2,np__3))!=X1|~m2_subset_1(esk3_4(X1,X2,X3,k1_enumset1(np__1,np__2,np__3)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk3_4(X1,X2,X3,k1_enumset1(np__1,np__2,np__3)))|~r1_xxreal_0(esk3_4(X1,X2,X3,k1_enumset1(np__1,np__2,np__3)),np__3))).#trainneg
cnf(i_0_8503, negated_conjecture, (k3_matrix_1(X2,esk2_0,X1,np__3)=k3_matrix_1(esk1_0,esk2_0,X1,np__3)|~m1_subset_1(X1,k4_finseq_1(esk2_0))|~v7_ordinal1(X1)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~r2_hidden(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8366, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,X3)|X4!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X4,X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(X4,esk1_0)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~m1_subset_1(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8506, negated_conjecture, (k3_matrix_1(X1,esk2_0,np__3,np__3)=k3_matrix_1(esk1_0,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~r2_hidden(np__3,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8511, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),np__3)=k3_matrix_1(X2,esk2_0,X1,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8367, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,X3)|X4!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X4,X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(X4,esk1_0)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_8462, negated_conjecture, (k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,X2)=k1_funct_1(k1_funct_1(esk2_0,X1),X2)|~m2_finseq_1(esk2_0,k3_finseq_2(k1_enumset1(np__1,np__2,np__3)))|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0))|~r1_xxreal_0(np__1,k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,X2))|~r1_xxreal_0(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,X2),np__3)|~m1_subset_1(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,X2),k5_numbers))).#trainneg
cnf(i_0_3748, plain, (v7_ordinal1(k3_finseq_1(X1))|~m1_subset_1(X1,k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3021, negated_conjecture, (v1_xcmplx_0(k3_matrix_1(esk1_0,esk2_0,np__3,np__3))|~v1_valued_0(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_1252, plain, (m2_finseq_1(esk4_5(X1,X2,X3,X3,X4),X1)|v1_xboole_0(k2_finseq_1(k1_matrix_1(X2)))|X4!=k3_matrix_1(X1,X2,X3,X3)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2))|~m1_subset_1(X3,k2_finseq_1(k1_matrix_1(X2))))).#trainneg
cnf(i_0_6108, plain, (v1_xcmplx_0(X1)|~v1_membered(k2_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1251, plain, (m2_finseq_1(esk4_5(X1,X2,esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2))),esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2))),X6),X1)|esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2)))=X3|esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2)))=X4|esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2)))=X5|k2_finseq_1(k1_matrix_1(X2))=k1_enumset1(X3,X4,X5)|X6!=k3_matrix_1(X1,X2,esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2))),esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2))))|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X6,X1)|~v7_ordinal1(esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2))))|~r2_hidden(esk3_4(X3,X4,X5,k2_finseq_1(k1_matrix_1(X2))),k4_finseq_1(X2)))).#trainneg
cnf(i_0_1258, plain, (r2_hidden(X1,k4_finseq_1(X2))|v1_xboole_0(k2_matrix_1(X2))|~v1_matrix_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(k2_tarski(k1_tarski(X1),k2_tarski(X3,X1)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_8419, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X2)=k3_matrix_1(esk1_0,esk2_0,X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8463, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X2)=k3_matrix_1(esk1_0,esk2_0,X1,X2)|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8509, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,np__3)|X1!=k3_matrix_1(X3,esk2_0,X2,np__3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,X2),esk1_0)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,esk1_0)|~m2_finseq_1(esk2_0,k3_finseq_2(X3))|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8442, negated_conjecture, (k3_finseq_1(k3_finseq_4(k4_finseq_1(esk2_0),X1,X2,X3))=np__3|~m1_subset_1(X3,k4_finseq_1(esk2_0))|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~m1_subset_1(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8499, negated_conjecture, (k3_matrix_1(X1,esk2_0,np__1,np__3)=k3_matrix_1(esk1_0,esk2_0,np__1,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(X1,esk2_0,np__1,np__3),esk1_0)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~r2_hidden(np__1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8497, negated_conjecture, (k3_matrix_1(X1,esk2_0,np__2,np__3)=k3_matrix_1(esk1_0,esk2_0,np__2,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(X1,esk2_0,np__2,np__3),esk1_0)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~r2_hidden(np__2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8422, negated_conjecture, (k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X2)=k1_funct_1(k1_funct_1(esk2_0,X1),X2)|~m2_finseq_1(esk2_0,k3_finseq_2(k4_finseq_1(esk2_0)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k4_finseq_1(esk2_0))|~m2_subset_1(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X2))|~r1_xxreal_0(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X2),np__3))).#trainneg
cnf(i_0_8498, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__1,np__3)|X1!=k3_matrix_1(X2,esk2_0,np__1,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,np__1),esk1_0)|~m1_subset_1(X1,esk1_0)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~r2_hidden(np__1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8496, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__2,np__3)|X1!=k3_matrix_1(X2,esk2_0,np__2,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,np__2),esk1_0)|~m1_subset_1(X1,esk1_0)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~r2_hidden(np__2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8460, negated_conjecture, (k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X2)=k1_funct_1(k1_funct_1(esk2_0,X1),X2)|~m2_finseq_1(esk2_0,k3_finseq_2(k4_finseq_1(esk2_0)))|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0))|~r2_hidden(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X2),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_2451, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),k1_xboole_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_4364, plain, (v1_xreal_0(k3_finseq_1(X1))|~m1_subset_1(X1,k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3020, negated_conjecture, (v1_xxreal_0(k3_matrix_1(esk1_0,esk2_0,np__3,np__3))|~v2_valued_0(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_1260, plain, (v1_xcmplx_0(X1)|~v3_valued_0(X3)|k1_funct_1(X3,np__1)!=X2|k1_funct_1(X3,np__2)!=X1|k3_finseq_1(X3)!=np__3|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_5877, plain, (r1_xxreal_0(esk9_1(k5_numbers),X1)|r1_xxreal_0(X1,esk9_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8699, negated_conjecture, (m2_subset_1(k1_xboole_0,k1_numbers,k5_numbers)|~v1_xboole_0(k1_funct_1(esk2_0,esk9_1(k4_finseq_1(esk2_0)))))).#trainneg
cnf(i_0_8777, negated_conjecture, (k3_matrix_1(X1,esk2_0,np__3,np__3)=k3_matrix_1(esk1_0,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1)))).#trainneg
cnf(i_0_8947, negated_conjecture, (k3_matrix_1(X1,esk2_0,np__1,np__3)=k3_matrix_1(esk1_0,esk2_0,np__1,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,np__1),esk1_0)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(k3_matrix_1(X1,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(np__1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8949, negated_conjecture, (k3_matrix_1(X1,esk2_0,np__2,np__3)=k3_matrix_1(esk1_0,esk2_0,np__2,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,np__2),esk1_0)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(k3_matrix_1(X1,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(np__2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8948, negated_conjecture, (k3_matrix_1(X1,esk2_0,np__1,np__3)=k3_matrix_1(esk1_0,esk2_0,np__1,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,np__1),esk1_0)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(k3_matrix_1(X1,esk2_0,np__1,np__3),esk1_0))).#trainneg
cnf(i_0_8950, negated_conjecture, (k3_matrix_1(X1,esk2_0,np__2,np__3)=k3_matrix_1(esk1_0,esk2_0,np__2,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,np__2),esk1_0)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(k3_matrix_1(X1,esk2_0,np__2,np__3),esk1_0))).#trainneg
cnf(i_0_8954, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__1,np__3)|X1!=k3_matrix_1(X2,esk2_0,np__1,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,np__1),esk1_0)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_8925, negated_conjecture, (k3_matrix_1(X1,esk2_0,X2,np__3)=k3_matrix_1(esk1_0,esk2_0,X2,np__3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,X2),esk1_0)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(k3_matrix_1(X1,esk2_0,X2,np__3),esk1_0)|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8956, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__2,np__3)|X1!=k3_matrix_1(X2,esk2_0,np__2,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,np__2),esk1_0)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_8811, negated_conjecture, (m2_finseq_1(esk4_5(X1,esk2_0,X2,X2,X3),X1)|X3!=k3_matrix_1(X1,esk2_0,X2,X2)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_9018, plain, (r1_xxreal_0(k1_xboole_0,esk9_1(k5_numbers))|r1_xxreal_0(esk9_1(k5_numbers),k1_xboole_0))).#trainneg
cnf(i_0_3019, negated_conjecture, (v1_xreal_0(k3_matrix_1(esk1_0,esk2_0,np__3,np__3))|~v3_valued_0(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_1275, plain, (v1_xcmplx_0(X1)|~v3_valued_0(X3)|k1_funct_1(X3,np__1)!=X1|k1_funct_1(X3,np__2)!=X2|k3_finseq_1(X3)!=np__3|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_7791, negated_conjecture, (v1_relat_1(k1_funct_1(esk2_0,k1_xboole_0))|~r1_xxreal_0(np__1,k1_xboole_0)|~r1_xxreal_0(k1_xboole_0,np__3)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_7817, negated_conjecture, (v1_funct_1(k1_funct_1(esk2_0,k1_xboole_0))|~r1_xxreal_0(np__1,k1_xboole_0)|~r1_xxreal_0(k1_xboole_0,np__3)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_8923, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X2)=k3_matrix_1(esk1_0,esk2_0,X1,X2)|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8924, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X2)=k3_matrix_1(esk1_0,esk2_0,X1,X2)|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_8880, negated_conjecture, (k1_funct_1(esk2_0,X1)=X2|k3_matrix_1(esk1_0,esk2_0,X1,esk5_2(k1_funct_1(esk2_0,X1),X2))!=k1_funct_1(X2,esk5_2(k1_funct_1(esk2_0,X1),X2))|k3_finseq_1(k1_funct_1(esk2_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(k1_funct_1(esk2_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k1_funct_1(esk2_0,X1))|~v1_relat_1(X2)|~v1_relat_1(k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k4_finseq_1(esk2_0))|~v7_ordinal1(esk5_2(k1_funct_1(esk2_0,X1),X2))|~v7_ordinal1(X1)|~r2_hidden(esk5_2(k1_funct_1(esk2_0,X1),X2),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_8421, negated_conjecture, (k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X2)=k1_funct_1(k1_funct_1(esk2_0,X1),X2)|~m2_finseq_1(esk2_0,k3_finseq_2(k4_finseq_1(esk2_0)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k4_finseq_1(esk2_0))|~m1_subset_1(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X2),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_8783, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,X3)|X4!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X4,X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(X4,esk1_0)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_8881, negated_conjecture, (X1=k1_funct_1(esk2_0,X2)|k1_funct_1(X1,esk5_2(X1,k1_funct_1(esk2_0,X2)))!=k3_matrix_1(esk1_0,esk2_0,X2,esk5_2(X1,k1_funct_1(esk2_0,X2)))|k3_finseq_1(X1)!=k3_finseq_1(k1_funct_1(esk2_0,X2))|~v1_finseq_1(k1_funct_1(esk2_0,X2))|~v1_finseq_1(X1)|~v1_funct_1(k1_funct_1(esk2_0,X2))|~v1_funct_1(X1)|~v1_relat_1(k1_funct_1(esk2_0,X2))|~v1_relat_1(X1)|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~v7_ordinal1(esk5_2(X1,k1_funct_1(esk2_0,X2)))|~v7_ordinal1(X2)|~r2_hidden(esk5_2(X1,k1_funct_1(esk2_0,X2)),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_2452, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k2_tarski(X3,X1),k1_tarski(X1)),k1_xboole_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_5119, plain, (v1_xcmplx_0(k3_finseq_1(X1))|~m1_subset_1(X1,k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6095, negated_conjecture, (v1_xcmplx_0(k3_matrix_1(esk1_0,esk2_0,np__3,np__3))|~v3_valued_0(k1_funct_1(esk2_0,np__3))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_8787, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,X2,X3)|X4!=k1_funct_1(esk2_0,X2)|X1!=k1_funct_1(X4,X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(X4,esk1_0)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X1,esk1_0)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3))).#trainneg
cnf(i_0_8642, plain, (esk5_2(X3,X4)=np__3|esk5_2(X3,X4)=np__2|esk5_2(X3,X4)=np__1|X3=X4|k3_finseq_1(X4)!=np__3|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4)|k1_funct_1(X3,np__1)!=X1|k1_funct_1(X3,np__2)!=X2|k3_finseq_1(X3)!=np__3|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_8500, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__3,np__3)|X1!=k3_matrix_1(X2,esk2_0,np__3,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,np__3),esk1_0)|~m1_subset_1(X1,esk1_0)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~r2_hidden(np__3,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8510, plain, (k3_matrix_1(X4,X1,X2,k1_matrix_1(X1))=k3_matrix_1(X3,X1,X2,k1_matrix_1(X1))|~m2_subset_1(k1_matrix_1(X1),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k1_matrix_1(X1))|~v1_matrix_1(X1)|~m2_finseq_1(X1,k3_finseq_2(X3))|~v7_ordinal1(k1_matrix_1(X1))|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(X1))|~m2_finseq_1(X1,k3_finseq_2(X4)))).#trainneg
cnf(i_0_8539, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X1)=k3_matrix_1(X2,esk2_0,X1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m1_subset_1(k3_matrix_1(X2,esk2_0,X1,X1),X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_8461, negated_conjecture, (k3_matrix_1(k2_finseq_1(X3),esk2_0,X1,X2)=k1_funct_1(k1_funct_1(esk2_0,X1),X2)|~m2_finseq_1(esk2_0,k3_finseq_2(k2_finseq_1(X3)))|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0))|~m2_subset_1(k3_matrix_1(k2_finseq_1(X3),esk2_0,X1,X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_matrix_1(k2_finseq_1(X3),esk2_0,X1,X2))|~r1_xxreal_0(k3_matrix_1(k2_finseq_1(X3),esk2_0,X1,X2),X3)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_8645, plain, (k3_finseq_1(k3_finseq_4(k5_numbers,X1,X2,X3))=np__3|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X5,k5_numbers)|~m1_subset_1(X6,k5_numbers)|~m1_subset_1(X7,k5_numbers)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X8,k5_numbers))).#trainneg
cnf(i_0_8458, negated_conjecture, (k3_matrix_1(a_1_0_finseq_1(X3),esk2_0,X1,X2)=k1_funct_1(k1_funct_1(esk2_0,X1),X2)|~m2_finseq_1(esk2_0,k3_finseq_2(a_1_0_finseq_1(X3)))|~m1_subset_1(X2,k1_enumset1(np__1,np__2,np__3))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0))|~m2_subset_1(k3_matrix_1(a_1_0_finseq_1(X3),esk2_0,X1,X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,k3_matrix_1(a_1_0_finseq_1(X3),esk2_0,X1,X2))|~r1_xxreal_0(k3_matrix_1(a_1_0_finseq_1(X3),esk2_0,X1,X2),X3)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2453, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X3)),k1_xboole_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_8675, plain, (m1_subset_1(k1_card_1(k3_finseq_1(X1)),k5_numbers)|~v1_finset_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6434, negated_conjecture, (v1_xcmplx_0(k3_matrix_1(esk1_0,esk2_0,np__3,np__3))|~m1_subset_1(k1_funct_1(esk2_0,np__3),k1_zfmisc_1(k1_xboole_0))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_7878, negated_conjecture, (v1_finseq_1(k1_funct_1(esk2_0,k1_xboole_0))|~r1_xxreal_0(np__1,k1_xboole_0)|~r1_xxreal_0(k1_xboole_0,np__3)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_1310, plain, (m1_subset_1(k2_tarski(k1_tarski(X1),k2_tarski(X2,X1)),k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_9277, negated_conjecture, (k3_matrix_1(X1,esk2_0,X2,np__3)=k3_matrix_1(X3,esk2_0,X2,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X3))|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_9053, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__2,np__3)|X1!=k3_matrix_1(X2,esk2_0,np__2,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_9052, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__1,np__3)|X1!=k3_matrix_1(X2,esk2_0,np__1,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_2454, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X3,X1)),k1_xboole_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_6039, negated_conjecture, (v1_xreal_0(k3_matrix_1(esk1_0,esk2_0,np__3,np__3))|~m1_subset_1(k1_funct_1(esk2_0,np__3),k1_zfmisc_1(k1_xboole_0))|~v1_relat_1(k1_funct_1(esk2_0,np__3)))).#trainneg
cnf(i_0_9358, plain, (m1_subset_1(k3_finseq_1(X1),k5_numbers)|~v1_finset_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1331, plain, (m2_finseq_1(esk4_5(X1,X2,X3,esk3_4(X4,X5,X6,k2_finseq_1(k1_matrix_1(X2))),X7),X1)|esk3_4(X4,X5,X6,k2_finseq_1(k1_matrix_1(X2)))=X4|esk3_4(X4,X5,X6,k2_finseq_1(k1_matrix_1(X2)))=X5|esk3_4(X4,X5,X6,k2_finseq_1(k1_matrix_1(X2)))=X6|k2_finseq_1(k1_matrix_1(X2))=k1_enumset1(X4,X5,X6)|X7!=k3_matrix_1(X1,X2,X3,esk3_4(X4,X5,X6,k2_finseq_1(k1_matrix_1(X2))))|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X7,X1)|~v7_ordinal1(esk3_4(X4,X5,X6,k2_finseq_1(k1_matrix_1(X2))))|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2)))).#trainneg
cnf(i_0_1332, plain, (m2_finseq_1(esk4_5(X1,X2,X3,X4,X5),X1)|v1_xboole_0(k2_finseq_1(k1_matrix_1(X2)))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2))|~m1_subset_1(X4,k2_finseq_1(k1_matrix_1(X2))))).#trainneg
cnf(i_0_9276, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__3,np__3)|X1!=k3_matrix_1(X2,esk2_0,np__3,np__3)|~m2_finseq_1(k1_funct_1(esk2_0,np__3),esk1_0)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_9051, negated_conjecture, (k3_matrix_1(X1,esk2_0,np__2,np__3)=k3_matrix_1(esk1_0,esk2_0,np__2,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(k3_matrix_1(X1,esk2_0,np__2,np__3),esk1_0))).#trainneg
cnf(i_0_9050, negated_conjecture, (k3_matrix_1(X1,esk2_0,np__1,np__3)=k3_matrix_1(esk1_0,esk2_0,np__1,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(k3_matrix_1(X1,esk2_0,np__1,np__3),esk1_0))).#trainneg
cnf(i_0_9194, negated_conjecture, (k1_funct_1(esk2_0,X1)=X2|k3_matrix_1(esk1_0,esk2_0,X1,esk5_2(k1_funct_1(esk2_0,X1),X2))!=k1_funct_1(X2,esk5_2(k1_funct_1(esk2_0,X1),X2))|k3_finseq_1(k1_funct_1(esk2_0,X1))!=k3_finseq_1(X2)|~m1_subset_1(X1,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,X1))|~v1_finseq_1(X2)|~v1_funct_1(k1_funct_1(esk2_0,X1))|~v1_funct_1(X2)|~v1_relat_1(k1_funct_1(esk2_0,X1))|~v1_relat_1(X2)|~v7_ordinal1(X1)|~m1_subset_1(esk5_2(k1_funct_1(esk2_0,X1),X2),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_9195, negated_conjecture, (k1_funct_1(esk2_0,X1)=X2|k3_matrix_1(esk1_0,esk2_0,X1,esk5_2(k1_funct_1(esk2_0,X1),X2))!=k1_funct_1(X2,esk5_2(k1_funct_1(esk2_0,X1),X2))|k3_finseq_1(k1_funct_1(esk2_0,X1))!=k3_finseq_1(X2)|~m1_subset_1(X1,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,X1))|~v1_finseq_1(X2)|~v1_funct_1(k1_funct_1(esk2_0,X1))|~v1_funct_1(X2)|~v1_relat_1(k1_funct_1(esk2_0,X1))|~v1_relat_1(X2)|~v7_ordinal1(X1)|~m2_subset_1(esk5_2(k1_funct_1(esk2_0,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk5_2(k1_funct_1(esk2_0,X1),X2))|~r1_xxreal_0(esk5_2(k1_funct_1(esk2_0,X1),X2),np__3))).#trainneg
cnf(i_0_9183, negated_conjecture, (k11_finseq_1(X1,X2,X3)=k1_funct_1(esk2_0,np__3)|k3_matrix_1(esk1_0,esk2_0,np__3,np__1)!=X1|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=X2|X3!=k3_matrix_1(X4,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X4))|~m1_subset_1(X3,X4))).#trainneg
cnf(i_0_9152, negated_conjecture, (k1_funct_1(esk2_0,X1)=X2|k3_matrix_1(esk1_0,esk2_0,X1,esk5_2(k1_funct_1(esk2_0,X1),X2))!=k1_funct_1(X2,esk5_2(k1_funct_1(esk2_0,X1),X2))|k3_finseq_1(k1_funct_1(esk2_0,X1))!=k3_finseq_1(X2)|~v1_finseq_1(X2)|~v1_finseq_1(k1_funct_1(esk2_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k1_funct_1(esk2_0,X1))|~v1_relat_1(X2)|~v1_relat_1(k1_funct_1(esk2_0,X1))|~m1_subset_1(esk5_2(k1_funct_1(esk2_0,X1),X2),k1_enumset1(np__1,np__2,np__3))|~v7_ordinal1(esk5_2(k1_funct_1(esk2_0,X1),X2))|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_9214, negated_conjecture, (X1=k1_funct_1(esk2_0,X2)|k1_funct_1(X1,esk5_2(X1,k1_funct_1(esk2_0,X2)))!=k3_matrix_1(esk1_0,esk2_0,X2,esk5_2(X1,k1_funct_1(esk2_0,X2)))|k3_finseq_1(X1)!=k3_finseq_1(k1_funct_1(esk2_0,X2))|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,X2))|~v1_finseq_1(X1)|~v1_funct_1(k1_funct_1(esk2_0,X2))|~v1_funct_1(X1)|~v1_relat_1(k1_funct_1(esk2_0,X2))|~v1_relat_1(X1)|~v7_ordinal1(X2)|~m1_subset_1(esk5_2(X1,k1_funct_1(esk2_0,X2)),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_9336, plain, (v1_xreal_0(k1_card_1(k3_finseq_1(X1)))|~v1_finset_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3812, negated_conjecture, (k11_finseq_1(k1_funct_1(esk2_0,np__1),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__1))!=np__3|~v1_finseq_1(k1_funct_1(esk2_0,np__1))|~v1_funct_1(k1_funct_1(esk2_0,np__1))|~v1_relat_1(k1_funct_1(esk2_0,np__1)))).#trainneg
cnf(i_0_1353, plain, (k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))=k1_enumset1(X2,X3,X4)|v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|esk3_4(X2,X3,X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))!=X4|~m2_finseq_1(esk3_4(X2,X3,X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))),X1))).#trainneg
cnf(i_0_2302, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_xreal_0(X3)|~v3_valued_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1357, plain, (m1_subset_1(k3_matrix_1(X1,k11_finseq_1(X2,X3,X4),X5,X6),X1)|v1_xboole_0(k3_finseq_2(X1))|~v1_matrix_1(k11_finseq_1(X2,X3,X4))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m1_subset_1(X4,k3_finseq_2(X1))|~m1_subset_1(X3,k3_finseq_2(X1))|~m1_subset_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_1362, plain, (v4_relat_1(k11_finseq_1(X1,X2,X3),k5_numbers)|v1_xboole_0(X4)|~m1_subset_1(X3,X4)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X4))).#trainneg
cnf(i_0_9452, negated_conjecture, (X1=k3_matrix_1(esk1_0,esk2_0,np__3,np__3)|X1!=k3_matrix_1(X2,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X2))|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_9437, negated_conjecture, (m2_finseq_1(esk4_5(X1,esk2_0,X2,X3,X4),X1)|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_9487, negated_conjecture, (k11_finseq_1(k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__2),k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|~m1_subset_1(np__2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_9486, negated_conjecture, (k11_finseq_1(k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__2),k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|k1_funct_1(k1_funct_1(esk2_0,np__1),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__1),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__1,np__2)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2))).#trainneg
cnf(i_0_9215, negated_conjecture, (X1=k1_funct_1(esk2_0,X2)|k1_funct_1(X1,esk5_2(X1,k1_funct_1(esk2_0,X2)))!=k3_matrix_1(esk1_0,esk2_0,X2,esk5_2(X1,k1_funct_1(esk2_0,X2)))|k3_finseq_1(X1)!=k3_finseq_1(k1_funct_1(esk2_0,X2))|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,X2))|~v1_finseq_1(X1)|~v1_funct_1(k1_funct_1(esk2_0,X2))|~v1_funct_1(X1)|~v1_relat_1(k1_funct_1(esk2_0,X2))|~v1_relat_1(X1)|~v7_ordinal1(X2)|~m2_subset_1(esk5_2(X1,k1_funct_1(esk2_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk5_2(X1,k1_funct_1(esk2_0,X2)))|~r1_xxreal_0(esk5_2(X1,k1_funct_1(esk2_0,X2)),np__3))).#trainneg
cnf(i_0_9496, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v3_valued_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3813, negated_conjecture, (k11_finseq_1(k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k1_funct_1(esk2_0,np__2),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__2)))).#trainneg
cnf(i_0_1365, plain, (m2_finseq_1(X3,X4)|v1_xboole_0(X4)|~m1_subset_1(k1_funct_1(X3,np__3),X4)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X4)|k1_funct_1(X3,np__1)!=X1|k1_funct_1(X3,np__2)!=X2|k3_finseq_1(X3)!=np__3|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_4239, plain, (v1_xcmplx_0(k1_card_1(k1_enumset1(X1,X2,X3)))|~v1_xreal_0(X4))).#trainneg
cnf(i_0_1419, plain, (k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))=k1_enumset1(X2,X3,X4)|v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|esk3_4(X2,X3,X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))!=X3|~m2_finseq_1(esk3_4(X2,X3,X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))),X1))).#trainneg
cnf(i_0_9454, negated_conjecture, (k1_funct_1(esk2_0,X1)=X2|k3_matrix_1(esk1_0,esk2_0,X1,esk5_2(k1_funct_1(esk2_0,X1),X2))!=k1_funct_1(X2,esk5_2(k1_funct_1(esk2_0,X1),X2))|k3_finseq_1(k1_funct_1(esk2_0,X1))!=k3_finseq_1(X2)|~r1_xxreal_0(esk5_2(k1_funct_1(esk2_0,X1),X2),np__3)|~m1_subset_1(X1,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,X1))|~v1_finseq_1(X2)|~v1_funct_1(k1_funct_1(esk2_0,X1))|~v1_funct_1(X2)|~v1_relat_1(k1_funct_1(esk2_0,X1))|~v1_relat_1(X2)|~v7_ordinal1(X1)|~m1_subset_1(esk5_2(k1_funct_1(esk2_0,X1),X2),k5_numbers))).#trainneg
cnf(i_0_9465, negated_conjecture, (X1=k1_funct_1(esk2_0,X2)|k1_funct_1(X1,esk5_2(X1,k1_funct_1(esk2_0,X2)))!=k3_matrix_1(esk1_0,esk2_0,X2,esk5_2(X1,k1_funct_1(esk2_0,X2)))|k3_finseq_1(X1)!=k3_finseq_1(k1_funct_1(esk2_0,X2))|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,X2))|~v1_finseq_1(X1)|~v1_funct_1(k1_funct_1(esk2_0,X2))|~v1_funct_1(X1)|~v1_relat_1(k1_funct_1(esk2_0,X2))|~v1_relat_1(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(np__1,esk5_2(X1,k1_funct_1(esk2_0,X2)))|~r1_xxreal_0(esk5_2(X1,k1_funct_1(esk2_0,X2)),np__3)|~m1_subset_1(esk5_2(X1,k1_funct_1(esk2_0,X2)),k5_numbers))).#trainneg
cnf(i_0_9288, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X1)=k3_matrix_1(esk1_0,esk2_0,X1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_9293, plain, (k3_finseq_1(k3_finseq_4(k5_numbers,X1,X2,X3))=np__3|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X5,k5_numbers)|~m1_subset_1(X6,k5_numbers)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X7,k5_numbers))).#trainneg
cnf(i_0_9196, negated_conjecture, (k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X2)=k1_funct_1(k1_funct_1(esk2_0,X1),X2)|~m2_finseq_1(esk2_0,k3_finseq_2(k4_finseq_1(esk2_0)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k4_finseq_1(esk2_0))|~r1_xxreal_0(np__1,k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X2))|~r1_xxreal_0(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X2),np__3)|~m1_subset_1(k3_matrix_1(k4_finseq_1(esk2_0),esk2_0,X1,X2),k5_numbers))).#trainneg
cnf(i_0_9464, negated_conjecture, (k11_finseq_1(X1,X2,X3)=k1_funct_1(esk2_0,X4)|esk5_2(k11_finseq_1(X1,X2,X3),k1_funct_1(esk2_0,X4))=np__3|esk5_2(k11_finseq_1(X1,X2,X3),k1_funct_1(esk2_0,X4))=np__1|X2!=k3_matrix_1(esk1_0,esk2_0,X4,np__2)|k3_finseq_1(k1_funct_1(esk2_0,X4))!=np__3|~m1_subset_1(X4,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,X4))|~v1_funct_1(k1_funct_1(esk2_0,X4))|~v1_relat_1(k1_funct_1(esk2_0,X4))|~v7_ordinal1(X4))).#trainneg
cnf(i_0_9249, plain, (esk5_2(X1,X2)=np__1|esk5_2(X1,X2)=np__2|esk5_2(X1,X2)=np__3|X1=X2|k1_funct_1(X1,np__2)!=X3|k3_finseq_1(X2)!=np__3|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9484, plain, (v1_xreal_0(k3_finseq_1(X1))|~v1_finset_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3817, negated_conjecture, (k11_finseq_1(k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k1_funct_1(esk2_0,np__2),k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__3,np__1),k3_matrix_1(esk1_0,esk2_0,np__3,np__2),k3_matrix_1(esk1_0,esk2_0,np__3,np__3)))!=esk2_0|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__1,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|k3_finseq_1(k1_funct_1(esk2_0,np__2))!=np__3|~v1_finseq_1(k1_funct_1(esk2_0,np__2))|~v1_funct_1(k1_funct_1(esk2_0,np__2))|~v1_relat_1(k1_funct_1(esk2_0,np__2)))).#trainneg
cnf(i_0_1423, plain, (k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))=k1_enumset1(X2,X3,X4)|v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|esk3_4(X2,X3,X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))!=X2|~m2_finseq_1(esk3_4(X2,X3,X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))),X1))).#trainneg
cnf(i_0_2303, plain, (v1_xcmplx_0(k1_matrix_1(X1))|~v1_xreal_0(X2)|~v1_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9403, negated_conjecture, (m1_subset_1(k3_finseq_1(np__3),k5_numbers)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(k1_funct_1(esk2_0,np__1),k5_numbers))).#trainneg
cnf(i_0_9405, negated_conjecture, (m1_subset_1(k3_finseq_1(np__3),k5_numbers)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(k1_funct_1(esk2_0,np__2),k5_numbers))).#trainneg
cnf(i_0_7369, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_1430, plain, (esk4_5(X1,X2,X3,esk8_2(X4,k1_matrix_1(X2)),X5)=k1_funct_1(X2,X3)|X5!=k3_matrix_1(X1,X2,X3,esk8_2(X4,k1_matrix_1(X2)))|~m2_subset_1(esk8_2(X4,k1_matrix_1(X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk8_2(X4,k1_matrix_1(X2)))|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~m1_subset_1(X5,X1)|~v7_ordinal1(k1_matrix_1(X2))|~v7_ordinal1(esk8_2(X4,k1_matrix_1(X2)))|~v7_ordinal1(X3)|~r2_hidden(X3,k4_finseq_1(X2))|~r2_hidden(X4,a_1_0_finseq_1(k1_matrix_1(X2))))).#trainneg
cnf(i_0_1433, plain, (np__3=X1|r1_xxreal_0(esk5_2(np__3,X1),np__3)|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_finseq_1(np__3)|~v1_funct_1(X1)|~v1_funct_1(np__3)|~v1_relat_1(X1)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_9801, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,X1),X1)=k3_matrix_1(esk1_0,esk2_0,X1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~m1_subset_1(X1,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_9671, negated_conjecture, (v3_valued_0(esk4_5(X1,esk2_0,X2,X3,X4))|~v3_membered(X1)|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_9666, negated_conjecture, (v1_funct_1(esk4_5(X1,esk2_0,X2,X3,X4))|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_9668, negated_conjecture, (v1_relat_1(esk4_5(X1,esk2_0,X2,X3,X4))|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_9758, plain, (v1_xcmplx_0(k1_card_1(k1_enumset1(X1,X2,X3))))).#trainneg
cnf(i_0_9719, negated_conjecture, (k11_finseq_1(k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k1_funct_1(esk2_0,np__2),k1_funct_1(esk2_0,np__3))!=esk2_0|k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2))).#trainneg
cnf(i_0_1470, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk9_1(k1_zfmisc_1(X2)))|v1_xboole_0(X2)|~m1_subset_1(X1,esk9_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_9406, negated_conjecture, (m1_subset_1(k3_finseq_1(np__3),k5_numbers)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(k1_funct_1(esk2_0,np__3),k5_numbers))).#trainneg
cnf(i_0_1473, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(X3)|v1_xboole_0(k2_zfmisc_1(k5_numbers,X2))|~m1_subset_1(X1,X3)|~m2_finseq_1(X3,X2))).#trainneg
cnf(i_0_1476, plain, (v1_matrix_1(k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),X7,k3_finseq_4(X1,X8,X9,X10)))|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(k1_funct_1(X7,np__3),X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|k1_funct_1(X7,np__1)!=X5|k1_funct_1(X7,np__2)!=X6|k3_finseq_1(X7)!=np__3|~v1_finseq_1(X7)|~v1_funct_1(X7)|~v1_relat_1(X7))).#trainneg
cnf(i_0_9788, negated_conjecture, (k1_funct_1(esk2_0,X1)=k1_funct_1(esk2_0,X2)|k3_matrix_1(esk1_0,esk2_0,X1,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))!=k3_matrix_1(esk1_0,esk2_0,X2,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|k3_finseq_1(k1_funct_1(esk2_0,X1))!=k3_finseq_1(k1_funct_1(esk2_0,X2))|~r1_xxreal_0(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)),np__3)|~m1_subset_1(X1,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,X1))|~v1_finseq_1(k1_funct_1(esk2_0,X2))|~v1_funct_1(k1_funct_1(esk2_0,X1))|~v1_funct_1(k1_funct_1(esk2_0,X2))|~v1_relat_1(k1_funct_1(esk2_0,X1))|~v1_relat_1(k1_funct_1(esk2_0,X2))|~v7_ordinal1(X1)|~v7_ordinal1(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|~v7_ordinal1(X2)|~r2_hidden(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)),k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_9667, negated_conjecture, (v1_finseq_1(esk4_5(X1,esk2_0,X2,X3,X4))|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_9784, negated_conjecture, (k1_funct_1(esk2_0,X1)=k1_funct_1(esk2_0,X2)|k3_matrix_1(esk1_0,esk2_0,X1,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))!=k3_matrix_1(esk1_0,esk2_0,X2,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|k3_finseq_1(k1_funct_1(esk2_0,X1))!=k3_finseq_1(k1_funct_1(esk2_0,X2))|~r1_xxreal_0(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)),np__3)|~m1_subset_1(X1,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,X1))|~v1_finseq_1(k1_funct_1(esk2_0,X2))|~v1_funct_1(k1_funct_1(esk2_0,X1))|~v1_funct_1(k1_funct_1(esk2_0,X2))|~v1_relat_1(k1_funct_1(esk2_0,X1))|~v1_relat_1(k1_funct_1(esk2_0,X2))|~v7_ordinal1(X1)|~m2_subset_1(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|~v7_ordinal1(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_9783, negated_conjecture, (k1_funct_1(esk2_0,X1)=k1_funct_1(esk2_0,X2)|k3_matrix_1(esk1_0,esk2_0,X1,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))!=k3_matrix_1(esk1_0,esk2_0,X2,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|k3_finseq_1(k1_funct_1(esk2_0,X1))!=k3_finseq_1(k1_funct_1(esk2_0,X2))|~r1_xxreal_0(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)),np__3)|~m1_subset_1(X1,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,X1))|~v1_finseq_1(k1_funct_1(esk2_0,X2))|~v1_funct_1(k1_funct_1(esk2_0,X1))|~v1_funct_1(k1_funct_1(esk2_0,X2))|~v1_relat_1(k1_funct_1(esk2_0,X1))|~v1_relat_1(k1_funct_1(esk2_0,X2))|~v7_ordinal1(X1)|~m2_subset_1(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~v7_ordinal1(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_9793, negated_conjecture, (k1_funct_1(esk2_0,X1)=k1_funct_1(esk2_0,X2)|k3_matrix_1(esk1_0,esk2_0,X1,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))!=k3_matrix_1(esk1_0,esk2_0,X2,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|k3_finseq_1(k1_funct_1(esk2_0,X1))!=k3_finseq_1(k1_funct_1(esk2_0,X2))|~r1_xxreal_0(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)),np__3)|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,X2))|~v1_finseq_1(k1_funct_1(esk2_0,X1))|~v1_funct_1(k1_funct_1(esk2_0,X2))|~v1_funct_1(k1_funct_1(esk2_0,X1))|~v1_relat_1(k1_funct_1(esk2_0,X2))|~v1_relat_1(k1_funct_1(esk2_0,X1))|~v7_ordinal1(X2)|~m2_subset_1(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)),k1_numbers,k5_numbers)|~r1_xxreal_0(np__1,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|~v7_ordinal1(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_9852, plain, (esk5_2(X1,X2)=np__3|esk5_2(X1,X2)=np__2|esk5_2(X1,X2)=np__1|X1=X2|k3_finseq_1(X2)!=np__3|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9794, negated_conjecture, (k1_funct_1(esk2_0,X1)=k1_funct_1(esk2_0,X2)|k3_matrix_1(esk1_0,esk2_0,X1,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))!=k3_matrix_1(esk1_0,esk2_0,X2,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|k3_finseq_1(k1_funct_1(esk2_0,X1))!=k3_finseq_1(k1_funct_1(esk2_0,X2))|~r1_xxreal_0(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)),np__3)|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,X2))|~v1_finseq_1(k1_funct_1(esk2_0,X1))|~v1_funct_1(k1_funct_1(esk2_0,X2))|~v1_funct_1(k1_funct_1(esk2_0,X1))|~v1_relat_1(k1_funct_1(esk2_0,X2))|~v1_relat_1(k1_funct_1(esk2_0,X1))|~v7_ordinal1(X2)|~m1_subset_1(X1,k4_finseq_1(esk2_0))|~v7_ordinal1(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|~v7_ordinal1(X1)|~r2_hidden(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_9905, plain, (v1_xcmplx_0(k1_matrix_1(X1))|~v1_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4034, plain, (m1_subset_1(k3_finseq_1(X1),k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1479, plain, (k1_matrix_1(k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)))=np__3|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_9877, plain, (v1_xcmplx_0(k3_finseq_1(X1))|~v1_xreal_0(X2)|~v1_finset_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1482, plain, (m2_finseq_1(k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),k3_finseq_2(X1))|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1483, plain, (m1_matrix_1(k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),X10),X1,np__3,np__3)|v1_xboole_0(X1)|~m1_subset_1(k1_funct_1(X10,np__3),X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|k1_funct_1(X10,np__1)!=X8|k1_funct_1(X10,np__2)!=X9|k3_finseq_1(X10)!=np__3|~v1_finseq_1(X10)|~v1_funct_1(X10)|~v1_relat_1(X10))).#trainneg
cnf(i_0_10115, negated_conjecture, (k1_funct_1(esk2_0,X1)=k1_funct_1(esk2_0,X2)|k3_matrix_1(esk1_0,esk2_0,X1,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))!=k3_matrix_1(esk1_0,esk2_0,X2,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|k3_finseq_1(k1_funct_1(esk2_0,X1))!=k3_finseq_1(k1_funct_1(esk2_0,X2))|~m1_subset_1(X1,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,X1))|~v1_funct_1(k1_funct_1(esk2_0,X1))|~v1_relat_1(k1_funct_1(esk2_0,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_10142, negated_conjecture, (k1_funct_1(esk2_0,X1)=k1_funct_1(esk2_0,X2)|k3_matrix_1(esk1_0,esk2_0,X1,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))!=k3_matrix_1(esk1_0,esk2_0,X2,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|k3_finseq_1(k1_funct_1(esk2_0,X1))!=k3_finseq_1(k1_funct_1(esk2_0,X2))|~r1_xxreal_0(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)),np__3)|~m1_subset_1(X1,k4_finseq_1(esk2_0))|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,X1))|~v1_finseq_1(k1_funct_1(esk2_0,X2))|~v1_funct_1(k1_funct_1(esk2_0,X1))|~v1_funct_1(k1_funct_1(esk2_0,X2))|~v1_relat_1(k1_funct_1(esk2_0,X1))|~v1_relat_1(k1_funct_1(esk2_0,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)),k5_numbers))).#trainneg
cnf(i_0_10140, negated_conjecture, (k1_funct_1(esk2_0,X1)=k1_funct_1(esk2_0,X2)|k3_matrix_1(esk1_0,esk2_0,X1,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))!=k3_matrix_1(esk1_0,esk2_0,X2,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|k3_finseq_1(k1_funct_1(esk2_0,X1))!=k3_finseq_1(k1_funct_1(esk2_0,X2))|~r1_xxreal_0(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)),np__3)|~r1_xxreal_0(np__1,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|~m1_subset_1(X1,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,X1))|~v1_funct_1(k1_funct_1(esk2_0,X1))|~v1_relat_1(k1_funct_1(esk2_0,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)),k5_numbers))).#trainneg
cnf(i_0_10053, negated_conjecture, (v1_relat_1(esk4_5(X1,esk2_0,X2,X3,X4))|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_10054, negated_conjecture, (v1_relat_1(esk4_5(X1,esk2_0,X2,X3,X4))|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r2_hidden(X2,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_9993, negated_conjecture, (k11_finseq_1(X1,X2,k3_matrix_1(esk1_0,esk2_0,np__3,np__3))=k1_funct_1(esk2_0,np__3)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=X1|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=X2|~m1_subset_1(np__3,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_10175, negated_conjecture, (k1_funct_1(esk2_0,X1)=k1_funct_1(esk2_0,X2)|k3_matrix_1(esk1_0,esk2_0,X1,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))!=k3_matrix_1(esk1_0,esk2_0,X2,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|k3_finseq_1(k1_funct_1(esk2_0,X1))!=k3_finseq_1(k1_funct_1(esk2_0,X2))|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~m1_subset_1(X1,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,X2))|~v1_finseq_1(k1_funct_1(esk2_0,X1))|~v1_funct_1(k1_funct_1(esk2_0,X2))|~v1_funct_1(k1_funct_1(esk2_0,X1))|~v1_relat_1(k1_funct_1(esk2_0,X2))|~v1_relat_1(k1_funct_1(esk2_0,X1))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_10268, plain, (v1_xcmplx_0(k3_finseq_1(X1))|~v1_finset_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10058, negated_conjecture, (k1_funct_1(k1_funct_1(esk2_0,np__2),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__2),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__2,np__2)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3))!=k1_funct_1(esk2_0,np__1))).#trainneg
cnf(i_0_1486, plain, (r1_xxreal_0(X1,k3_xcmplx_0(k2_xcmplx_0(X2,X3),X4))|r1_xxreal_0(k3_xcmplx_0(k2_xcmplx_0(X2,X3),X4),X1)|~v7_ordinal1(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v7_ordinal1(k3_xcmplx_0(X3,X4))|~v7_ordinal1(k3_xcmplx_0(X2,X4)))).#trainneg
cnf(i_0_7164, plain, (r1_xxreal_0(X2,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,X2)|~v2_valued_0(k11_finseq_1(X1,X2,X3)))).#trainneg
cnf(i_0_3950, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),k11_finseq_1(X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_3964, plain, (v1_xcmplx_0(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k3_finseq_4(X1,X5,X6,X7),k11_finseq_1(X8,X9,X10)),X11,X12))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_1484, plain, (m1_matrix_1(k11_finseq_1(k3_finseq_4(X1,X2,X3,X4),X7,k11_finseq_1(X8,X9,X10)),X1,np__3,np__3)|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(k1_funct_1(X7,np__3),X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|k1_funct_1(X7,np__1)!=X5|k1_funct_1(X7,np__2)!=X6|k3_finseq_1(X7)!=np__3|~v1_finseq_1(X7)|~v1_funct_1(X7)|~v1_relat_1(X7))).#trainneg
cnf(i_0_1487, plain, (m1_subset_1(k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3),k5_numbers)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v7_ordinal1(k3_xcmplx_0(X2,X3))|~v7_ordinal1(k3_xcmplx_0(X1,X3)))).#trainneg
cnf(i_0_9991, negated_conjecture, (k11_finseq_1(X1,X2,X3)=k1_funct_1(esk2_0,np__3)|k3_matrix_1(esk1_0,esk2_0,np__3,np__3)!=X3|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=X1|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=X2|~m1_subset_1(np__3,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_10144, negated_conjecture, (k1_funct_1(esk2_0,X1)=k1_funct_1(esk2_0,X2)|k3_matrix_1(esk1_0,esk2_0,X1,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))!=k3_matrix_1(esk1_0,esk2_0,X2,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|k3_finseq_1(k1_funct_1(esk2_0,X1))!=k3_finseq_1(k1_funct_1(esk2_0,X2))|~r1_xxreal_0(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)),np__3)|~r1_xxreal_0(np__1,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,X2))|~v1_funct_1(k1_funct_1(esk2_0,X2))|~v1_relat_1(k1_funct_1(esk2_0,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)),k5_numbers))).#trainneg
cnf(i_0_9797, negated_conjecture, (k1_funct_1(esk2_0,X1)=k1_funct_1(esk2_0,X2)|k3_matrix_1(esk1_0,esk2_0,X1,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))!=k3_matrix_1(esk1_0,esk2_0,X2,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|k3_finseq_1(k1_funct_1(esk2_0,X1))!=k3_finseq_1(k1_funct_1(esk2_0,X2))|~r1_xxreal_0(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)),np__3)|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,X2))|~v1_finseq_1(k1_funct_1(esk2_0,X1))|~v1_funct_1(k1_funct_1(esk2_0,X2))|~v1_funct_1(k1_funct_1(esk2_0,X1))|~v1_relat_1(k1_funct_1(esk2_0,X2))|~v1_relat_1(k1_funct_1(esk2_0,X1))|~v7_ordinal1(X2)|~v7_ordinal1(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|~v7_ordinal1(X1)|~r2_hidden(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)),k1_enumset1(np__1,np__2,np__3))|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_10145, negated_conjecture, (k1_funct_1(esk2_0,X1)=k1_funct_1(esk2_0,X2)|k3_matrix_1(esk1_0,esk2_0,X1,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))!=k3_matrix_1(esk1_0,esk2_0,X2,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|k3_finseq_1(k1_funct_1(esk2_0,X1))!=k3_finseq_1(k1_funct_1(esk2_0,X2))|~r1_xxreal_0(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)),np__3)|~r1_xxreal_0(np__1,esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)))|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~v1_finseq_1(k1_funct_1(esk2_0,X2))|~v1_funct_1(k1_funct_1(esk2_0,X2))|~v1_relat_1(k1_funct_1(esk2_0,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r2_hidden(X1,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(esk5_2(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2)),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_9744, negated_conjecture, (m2_finseq_1(X1,k4_finseq_1(esk2_0))|k1_funct_1(X1,np__1)!=X2|k1_funct_1(X1,np__2)!=X3|k3_finseq_1(X1)!=np__3|~m1_subset_1(X3,k4_finseq_1(esk2_0))|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(k1_funct_1(X1,np__3),k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_10609, negated_conjecture, (m2_finseq_1(k11_finseq_1(X1,X2,X3),k4_finseq_1(esk2_0))|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~m1_subset_1(X1,k4_finseq_1(esk2_0))|~r2_hidden(X3,k1_enumset1(np__1,np__2,np__3)))).#trainneg
cnf(i_0_10363, negated_conjecture, (v1_relat_1(k1_funct_1(esk2_0,X2))|X3!=k3_matrix_1(X1,esk2_0,X2,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10355, negated_conjecture, (v1_relat_1(k1_funct_1(esk2_0,X2))|X4!=k3_matrix_1(X1,esk2_0,X2,X3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,k1_enumset1(np__1,np__2,np__3))|~m1_subset_1(X2,k4_finseq_1(esk2_0))|~m1_subset_1(X4,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1488, plain, (r1_xxreal_0(np__2,k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3))|r1_xxreal_0(k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3),np__2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v7_ordinal1(k3_xcmplx_0(X2,X3))|~v7_ordinal1(k3_xcmplx_0(X1,X3)))).#trainneg
cnf(i_0_7165, plain, (r1_xxreal_0(X1,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,X1)|~v2_valued_0(k11_finseq_1(X1,X2,X3)))).#trainneg
cnf(i_0_1489, plain, (r1_xxreal_0(np__1,k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3))|r1_xxreal_0(k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3),np__1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v7_ordinal1(k3_xcmplx_0(X2,X3))|~v7_ordinal1(k3_xcmplx_0(X1,X3)))).#trainneg
cnf(i_0_1490, plain, (r1_xxreal_0(np__3,k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3))|r1_xxreal_0(k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3),np__3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v7_ordinal1(k3_xcmplx_0(X2,X3))|~v7_ordinal1(k3_xcmplx_0(X1,X3)))).#trainneg
cnf(i_0_10575, negated_conjecture, (k11_finseq_1(k1_funct_1(k1_funct_1(esk2_0,np__3),np__1),X1,X2)=k1_funct_1(esk2_0,np__3)|k3_matrix_1(esk1_0,esk2_0,np__3,np__3)!=X2|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=X1|~m1_subset_1(np__3,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_1491, plain, (r1_xxreal_0(k1_xboole_0,k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3))|r1_xxreal_0(k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3),k1_xboole_0)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v7_ordinal1(k3_xcmplx_0(X2,X3))|~v7_ordinal1(k3_xcmplx_0(X1,X3)))).#trainneg
cnf(i_0_7166, plain, (r1_xxreal_0(X3,k1_xboole_0)|r1_xxreal_0(k1_xboole_0,X3)|~v2_valued_0(k11_finseq_1(X1,X2,X3)))).#trainneg
cnf(i_0_1492, plain, (v7_ordinal1(k3_xcmplx_0(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k3_xcmplx_0(X3,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v7_ordinal1(k3_xcmplx_0(k3_xcmplx_0(X3,X2),X4))|~v7_ordinal1(k3_xcmplx_0(k3_xcmplx_0(X1,X2),X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1498, plain, (k7_finseq_1(k7_finseq_1(k1_tarski(k2_tarski(k2_tarski(np__1,X1),k1_tarski(np__1))),k1_tarski(k2_tarski(k2_tarski(np__1,X2),k1_tarski(np__1)))),k1_tarski(k2_tarski(k1_tarski(np__1),k2_tarski(X3,np__1))))=k11_finseq_1(X1,X2,X3))).#trainneg
cnf(i_0_9672, negated_conjecture, (m2_finseq_1(k1_funct_1(esk2_0,X2),X1)|X3!=k3_matrix_1(X1,esk2_0,X2,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X1))|~m1_subset_1(X3,X1)|~v7_ordinal1(X2)|~r2_hidden(X2,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_9803, plain, (k3_finseq_1(k3_finseq_4(k5_numbers,X1,X2,X3))=np__3|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X5,k5_numbers)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X6,k5_numbers))).#trainneg
cnf(i_0_9457, negated_conjecture, (k11_finseq_1(X1,X2,X3)=k1_funct_1(esk2_0,np__3)|k3_matrix_1(esk1_0,esk2_0,np__3,np__1)!=X1|k3_matrix_1(esk1_0,esk2_0,np__3,np__2)!=X2|X3!=k3_matrix_1(X4,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X4))|~m1_subset_1(X3,X4)|~m1_subset_1(np__3,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_9458, negated_conjecture, (k11_finseq_1(X1,X2,X3)=k1_funct_1(esk2_0,np__3)|k3_matrix_1(esk1_0,esk2_0,np__3,np__1)!=X1|k3_matrix_1(esk1_0,esk2_0,np__3,np__2)!=X2|X3!=k3_matrix_1(X4,esk2_0,np__3,np__3)|~m2_finseq_1(esk2_0,k3_finseq_2(X4))|~m1_subset_1(X3,X4))).#trainneg
cnf(i_0_9287, negated_conjecture, (k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,X1)=k1_funct_1(k1_funct_1(esk2_0,X1),X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~m2_finseq_1(esk2_0,k3_finseq_2(k1_enumset1(np__1,np__2,np__3)))|~v7_ordinal1(X1)|~r2_hidden(X1,k4_finseq_1(esk2_0))|~r1_xxreal_0(np__1,k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,X1))|~r1_xxreal_0(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,X1),np__3)|~m1_subset_1(k3_matrix_1(k1_enumset1(np__1,np__2,np__3),esk2_0,X1,X1),k5_numbers))).#trainneg
cnf(i_0_10382, negated_conjecture, (k11_finseq_1(k3_finseq_4(esk1_0,k3_matrix_1(esk1_0,esk2_0,np__1,np__1),k3_matrix_1(esk1_0,esk2_0,np__1,np__2),k3_matrix_1(esk1_0,esk2_0,np__1,np__3)),k11_finseq_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),k3_matrix_1(esk1_0,esk2_0,np__2,np__2),k3_matrix_1(esk1_0,esk2_0,np__2,np__3)),k1_funct_1(esk2_0,np__3))!=esk2_0|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__2,np__1),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__3),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__2),esk1_0)|~m1_subset_1(k3_matrix_1(esk1_0,esk2_0,np__3,np__1),esk1_0)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__1)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__1)|k1_funct_1(k1_funct_1(esk2_0,np__3),np__2)!=k3_matrix_1(esk1_0,esk2_0,np__3,np__2)|~m1_subset_1(np__3,k4_finseq_1(esk2_0)))).#trainneg
cnf(i_0_7787, negated_conjecture, (v1_relat_1(k1_funct_1(esk2_0,esk9_1(k1_enumset1(np__1,np__2,np__3)))))).#trainneg
cnf(i_0_7813, negated_conjecture, (v1_funct_1(k1_funct_1(esk2_0,esk9_1(k1_enumset1(np__1,np__2,np__3)))))).#trainneg
cnf(i_0_3191, plain, (m1_subset_1(k3_matrix_1(X1,k11_finseq_1(k11_finseq_1(X2,X3,X4),k11_finseq_1(X5,X6,X7),k3_finseq_4(X1,X8,X9,X10)),X11,X12),X1)|v1_xboole_0(X1)|~m1_subset_1(X10,X1)|~m1_subset_1(X9,X1)|~m1_subset_1(X8,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X12)|~v7_ordinal1(X11))).#trainneg
cnf(i_0_1500, plain, (k7_finseq_1(k7_finseq_1(k1_tarski(k2_tarski(k2_tarski(np__1,X1),k1_tarski(np__1))),k1_tarski(k2_tarski(k2_tarski(X2,np__1),k1_tarski(np__1)))),k1_tarski(k2_tarski(k2_tarski(X3,np__1),k1_tarski(np__1))))=k11_finseq_1(X1,X2,X3))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 117
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 179
# Removed in clause preprocessing      : 7
# Initial clauses in saturation        : 172
# Processed clauses                    : 2761
# ...of these trivial                  : 61
# ...subsumed                          : 1177
# ...remaining for further processing  : 1523
# Other redundant clauses eliminated   : 283
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 166
# Backward-rewritten                   : 62
# Generated clauses                    : 10779
# ...of the previous two non-trivial   : 9069
# Contextual simplify-reflections      : 144
# Paramodulations                      : 10338
# Factorizations                       : 50
# Equation resolutions                 : 391
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1291
#    Positive orientable unit clauses  : 156
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 25
#    Non-unit-clauses                  : 1109
# Current number of unprocessed clauses: 5858
# ...number of literals in the above   : 67879
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1754
# Proof object given clauses           : 195
# Proof search given clauses           : 1523
# Clause-clause subsumption calls (NU) : 756427
# Rec. Clause-clause subsumption calls : 44703
# Non-unit clause-clause subsumptions  : 1349
# Unit Clause-clause subsumption calls : 8910
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 114
# BW rewrite match successes           : 32
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 386881

# -------------------------------------------------
# User time                : 44.426 s
# System time              : 1.377 s
# Total time               : 45.803 s
# Maximum resident set size: 390620 pages
44.42user 1.52system 0:46.69elapsed 98%CPU (0avgtext+0avgdata 391376maxresident)k
0inputs+944outputs (0major+327701minor)pagefaults 0swaps
