2020-09-01 13:32:35.488030: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_57
2020-09-01 13:32:35.586244: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-09-01 13:32:35.768362: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-09-01 13:32:36.211372: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-09-01 13:32:37.883650: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 2395640 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_57' loaded (query=768; context=1024; weigths=real; len_mult=0.000000).
# Preprocessing time       : 3.183 s
#TF# Flushing cached clauses: 0 unprocessed
#TF# Cached clauses flushed: 321 unprocessed

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_uniform1, conjecture, ![X1]:((((((~(v2_struct_0(X1))&v6_metric_1(X1))&v7_metric_1(X1))&v8_metric_1(X1))&v9_metric_1(X1))&l1_metric_1(X1))=>![X2]:((((((~(v2_struct_0(X2))&v6_metric_1(X2))&v7_metric_1(X2))&v8_metric_1(X2))&v9_metric_1(X2))&l1_metric_1(X2))=>![X3]:(((v1_funct_1(X3)&v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2))))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2))))))=>(![X4]:(v1_xreal_0(X4)=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>![X6]:(m1_subset_1(X6,u1_struct_0(X2))=>~(((~(r1_xxreal_0(X4,k6_numbers))&X6=k1_funct_1(X3,X5))&![X7]:(v1_xreal_0(X7)=>~((~(r1_xxreal_0(X7,k6_numbers))&![X8]:(m1_subset_1(X8,u1_struct_0(X1))=>![X9]:(m1_subset_1(X9,u1_struct_0(X2))=>~(((X9=k1_funct_1(X3,X8)&~(r1_xxreal_0(X7,k4_metric_1(X1,X5,X8))))&r1_xxreal_0(X4,k4_metric_1(X2,X6,X9))))))))))))))=>v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2)))))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', t3_uniform1)).
fof(t11_metric_1, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(l1_metric_1(X2)=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>(r2_hidden(X4,k9_metric_1(X2,X3,X1))<=>(~(v2_struct_0(X2))&~(r1_xxreal_0(X1,k2_metric_1(X2,X3,X4))))))))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', t11_metric_1)).
fof(redefinition_k4_metric_1, axiom, ![X1, X2, X3]:((((v8_metric_1(X1)&l1_metric_1(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>k4_metric_1(X1,X2,X3)=k2_metric_1(X1,X2,X3)), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', redefinition_k4_metric_1)).
fof(t12_topmetr, axiom, ![X1]:(l1_metric_1(X1)=>(u1_struct_0(X1)=u1_struct_0(k3_pcomps_1(X1))&u1_pre_topc(k3_pcomps_1(X1))=k2_pcomps_1(X1))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', t12_topmetr)).
fof(commutativity_k4_metric_1, axiom, ![X1, X2, X3]:((((v8_metric_1(X1)&l1_metric_1(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>k4_metric_1(X1,X2,X3)=k4_metric_1(X1,X3,X2)), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', commutativity_k4_metric_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', t3_subset)).
fof(reflexivity_r1_tarski, axiom, ![X1, X2]:r1_tarski(X1,X1), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', reflexivity_r1_tarski)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', t4_subset)).
fof(dt_k9_metric_1, axiom, ![X1, X2, X3]:(((l1_metric_1(X1)&m1_subset_1(X2,u1_struct_0(X1)))&v1_xreal_0(X3))=>m1_subset_1(k9_metric_1(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', dt_k9_metric_1)).
fof(cc4_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>v1_xboole_0(X3))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', cc4_relset_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', d3_tarski)).
fof(t15_topmetr, axiom, ![X1]:((((((~(v2_struct_0(X1))&v6_metric_1(X1))&v7_metric_1(X1))&v8_metric_1(X1))&v9_metric_1(X1))&l1_metric_1(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))=>(v3_pre_topc(X2,k3_pcomps_1(X1))<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>~((r2_hidden(X3,X2)&![X4]:(v1_xreal_0(X4)=>~((~(r1_xxreal_0(X4,k6_numbers))&r1_tarski(k9_metric_1(X1,X3,X4),X2)))))))))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', t15_topmetr)).
fof(d1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((X2!=k1_xboole_0=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X3)))&(X2=k1_xboole_0=>(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0)))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', d1_funct_2)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', t6_boole)).
fof(dt_k1_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>m1_subset_1(k1_relset_1(X1,X2),k1_zfmisc_1(X1))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', dt_k1_relset_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>k1_relset_1(X1,X2)=k9_xtuple_0(X2)), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', redefinition_k1_relset_1)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', cc2_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', cc1_relset_1)).
fof(dt_k2_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v5_relat_1(X2,X1))=>m1_subset_1(k2_relset_1(X1,X2),k1_zfmisc_1(X1))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', dt_k2_relset_1)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v5_relat_1(X2,X1))=>k2_relset_1(X1,X2)=k10_xtuple_0(X2)), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', redefinition_k2_relset_1)).
fof(d7_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2, X3]:(X3=k8_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,k9_xtuple_0(X1))&r2_hidden(k1_funct_1(X1,X4),X2))))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', d7_funct_1)).
fof(d3_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(X2=k10_xtuple_0(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X4,k9_xtuple_0(X1))&X3=k1_funct_1(X1,X4))))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', d3_funct_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', fc1_xboole_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', fc2_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', dt_l1_pre_topc)).
fof(t7_boole, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', t7_boole)).
fof(dt_k3_pcomps_1, axiom, ![X1]:(l1_metric_1(X1)=>l1_pre_topc(k3_pcomps_1(X1))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', dt_k3_pcomps_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', t2_subset)).
fof(fc1_struct_0, axiom, ![X1]:((v2_struct_0(X1)&l1_struct_0(X1))=>v1_xboole_0(u1_struct_0(X1))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', fc1_struct_0)).
fof(dt_l1_metric_1, axiom, ![X1]:(l1_metric_1(X1)=>l1_struct_0(X1)), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', dt_l1_metric_1)).
fof(fc2_pcomps_1, axiom, ![X1]:((~(v2_struct_0(X1))&l1_metric_1(X1))=>~(v2_struct_0(k3_pcomps_1(X1)))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', fc2_pcomps_1)).
fof(abstractness_v1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v1_pre_topc(X1)=>X1=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', abstractness_v1_pre_topc)).
fof(fc1_pcomps_1, axiom, ![X1]:(l1_metric_1(X1)=>(v1_pre_topc(k3_pcomps_1(X1))&v2_pre_topc(k3_pcomps_1(X1)))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', fc1_pcomps_1)).
fof(dt_u1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', dt_u1_pre_topc)).
fof(dt_g1_pre_topc, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>(v1_pre_topc(g1_pre_topc(X1,X2))&l1_pre_topc(g1_pre_topc(X1,X2)))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', dt_g1_pre_topc)).
fof(dt_k8_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>m1_subset_1(k8_relset_1(X1,X2,X3,X4),k1_zfmisc_1(X1))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', dt_k8_relset_1)).
fof(redefinition_k8_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k8_relset_1(X1,X2,X3,X4)=k8_relat_1(X3,X4)), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', redefinition_k8_relset_1)).
fof(t14_topmetr, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:((v9_metric_1(X2)&l1_metric_1(X2))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))=>(X4=k9_metric_1(X2,X3,X1)=>v3_pre_topc(X4,k3_pcomps_1(X2))))))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', t14_topmetr)).
fof(t16_jordan2b, axiom, ![X1]:(((~(v2_struct_0(X1))&v2_pre_topc(X1))&l1_pre_topc(X1))=>![X2]:((((((~(v2_struct_0(X2))&v6_metric_1(X2))&v7_metric_1(X2))&v8_metric_1(X2))&v9_metric_1(X2))&l1_metric_1(X2))=>![X3]:(((v1_funct_1(X3)&v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2))))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2))))))=>(![X4]:(v1_xreal_0(X4)=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>![X6]:(m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))=>(X6=k9_metric_1(X2,X5,X4)=>(r1_xxreal_0(X4,k6_numbers)|v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)),X3,X6),X1))))))=>v5_pre_topc(X3,X1,k3_pcomps_1(X2)))))), file('/scratch/mptp/hardpred/preds__0.005//t3_uniform1', t16_jordan2b)).
fof(c_0_39, negated_conjecture, ~(![X1]:((((((~(v2_struct_0(X1))&v6_metric_1(X1))&v7_metric_1(X1))&v8_metric_1(X1))&v9_metric_1(X1))&l1_metric_1(X1))=>![X2]:((((((~(v2_struct_0(X2))&v6_metric_1(X2))&v7_metric_1(X2))&v8_metric_1(X2))&v9_metric_1(X2))&l1_metric_1(X2))=>![X3]:(((v1_funct_1(X3)&v1_funct_2(X3,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2))))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2))))))=>(![X4]:(v1_xreal_0(X4)=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>![X6]:(m1_subset_1(X6,u1_struct_0(X2))=>~(((~(r1_xxreal_0(X4,k6_numbers))&X6=k1_funct_1(X3,X5))&![X7]:(v1_xreal_0(X7)=>~((~(r1_xxreal_0(X7,k6_numbers))&![X8]:(m1_subset_1(X8,u1_struct_0(X1))=>![X9]:(m1_subset_1(X9,u1_struct_0(X2))=>~(((X9=k1_funct_1(X3,X8)&~(r1_xxreal_0(X7,k4_metric_1(X1,X5,X8))))&r1_xxreal_0(X4,k4_metric_1(X2,X6,X9))))))))))))))=>v5_pre_topc(X3,k3_pcomps_1(X1),k3_pcomps_1(X2))))))), inference(assume_negation,[status(cth)],[t3_uniform1])).
fof(c_0_40, plain, ![X39, X40, X41, X42]:(((~v2_struct_0(X40)|~r2_hidden(X42,k9_metric_1(X40,X41,X39))|~m1_subset_1(X42,u1_struct_0(X40))|~m1_subset_1(X41,u1_struct_0(X40))|~l1_metric_1(X40)|~v1_xreal_0(X39))&(~r1_xxreal_0(X39,k2_metric_1(X40,X41,X42))|~r2_hidden(X42,k9_metric_1(X40,X41,X39))|~m1_subset_1(X42,u1_struct_0(X40))|~m1_subset_1(X41,u1_struct_0(X40))|~l1_metric_1(X40)|~v1_xreal_0(X39)))&(v2_struct_0(X40)|r1_xxreal_0(X39,k2_metric_1(X40,X41,X42))|r2_hidden(X42,k9_metric_1(X40,X41,X39))|~m1_subset_1(X42,u1_struct_0(X40))|~m1_subset_1(X41,u1_struct_0(X40))|~l1_metric_1(X40)|~v1_xreal_0(X39))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t11_metric_1])])])])])).
fof(c_0_41, plain, ![X95, X96, X97]:(~v8_metric_1(X95)|~l1_metric_1(X95)|~m1_subset_1(X96,u1_struct_0(X95))|~m1_subset_1(X97,u1_struct_0(X95))|k4_metric_1(X95,X96,X97)=k2_metric_1(X95,X96,X97)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_metric_1])])).
fof(c_0_42, negated_conjecture, ![X13, X14, X15, X17, X18]:((((((~v2_struct_0(esk1_0)&v6_metric_1(esk1_0))&v7_metric_1(esk1_0))&v8_metric_1(esk1_0))&v9_metric_1(esk1_0))&l1_metric_1(esk1_0))&((((((~v2_struct_0(esk2_0)&v6_metric_1(esk2_0))&v7_metric_1(esk2_0))&v8_metric_1(esk2_0))&v9_metric_1(esk2_0))&l1_metric_1(esk2_0))&(((v1_funct_1(esk3_0)&v1_funct_2(esk3_0,u1_struct_0(k3_pcomps_1(esk1_0)),u1_struct_0(k3_pcomps_1(esk2_0))))&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(esk1_0)),u1_struct_0(k3_pcomps_1(esk2_0))))))&(((v1_xreal_0(esk4_3(X13,X14,X15))|(r1_xxreal_0(X13,k6_numbers)|X15!=k1_funct_1(esk3_0,X14))|~m1_subset_1(X15,u1_struct_0(esk2_0))|~m1_subset_1(X14,u1_struct_0(esk1_0))|~v1_xreal_0(X13))&((~r1_xxreal_0(esk4_3(X13,X14,X15),k6_numbers)|(r1_xxreal_0(X13,k6_numbers)|X15!=k1_funct_1(esk3_0,X14))|~m1_subset_1(X15,u1_struct_0(esk2_0))|~m1_subset_1(X14,u1_struct_0(esk1_0))|~v1_xreal_0(X13))&(~m1_subset_1(X17,u1_struct_0(esk1_0))|(~m1_subset_1(X18,u1_struct_0(esk2_0))|(X18!=k1_funct_1(esk3_0,X17)|r1_xxreal_0(esk4_3(X13,X14,X15),k4_metric_1(esk1_0,X14,X17))|~r1_xxreal_0(X13,k4_metric_1(esk2_0,X15,X18))))|(r1_xxreal_0(X13,k6_numbers)|X15!=k1_funct_1(esk3_0,X14))|~m1_subset_1(X15,u1_struct_0(esk2_0))|~m1_subset_1(X14,u1_struct_0(esk1_0))|~v1_xreal_0(X13))))&~v5_pre_topc(esk3_0,k3_pcomps_1(esk1_0),k3_pcomps_1(esk2_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_39])])])])])])).
cnf(c_0_43, plain, (v2_struct_0(X1)|r1_xxreal_0(X2,k2_metric_1(X1,X3,X4))|r2_hidden(X4,k9_metric_1(X1,X3,X2))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_metric_1(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_44, plain, (k4_metric_1(X1,X2,X3)=k2_metric_1(X1,X2,X3)|~v8_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_45, negated_conjecture, (r1_xxreal_0(esk4_3(X3,X4,X5),k4_metric_1(esk1_0,X4,X1))|r1_xxreal_0(X3,k6_numbers)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|X2!=k1_funct_1(esk3_0,X1)|~r1_xxreal_0(X3,k4_metric_1(esk2_0,X5,X2))|X5!=k1_funct_1(esk3_0,X4)|~m1_subset_1(X5,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~v1_xreal_0(X3)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_46, plain, (v2_struct_0(X1)|r1_xxreal_0(X2,k2_metric_1(X1,X3,X4))|r2_hidden(X4,k9_metric_1(X1,X3,X2))|~l1_metric_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_43])).
cnf(c_0_47, plain, (k2_metric_1(X1,X2,X3)=k4_metric_1(X1,X2,X3)|~v8_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_44])).
cnf(c_0_48, negated_conjecture, (l1_metric_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_49, negated_conjecture, (v8_metric_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_50, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_51, plain, ![X25]:((u1_struct_0(X25)=u1_struct_0(k3_pcomps_1(X25))|~l1_metric_1(X25))&(u1_pre_topc(k3_pcomps_1(X25))=k2_pcomps_1(X25)|~l1_metric_1(X25))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_topmetr])])])).
fof(c_0_52, plain, ![X86, X87, X88]:(~v8_metric_1(X86)|~l1_metric_1(X86)|~m1_subset_1(X87,u1_struct_0(X86))|~m1_subset_1(X88,u1_struct_0(X86))|k4_metric_1(X86,X87,X88)=k4_metric_1(X86,X88,X87)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k4_metric_1])])).
cnf(c_0_53, plain, (~r1_xxreal_0(X1,k2_metric_1(X2,X3,X4))|~r2_hidden(X4,k9_metric_1(X2,X3,X1))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_54, negated_conjecture, (r1_xxreal_0(X3,k6_numbers)|r1_xxreal_0(esk4_3(X3,X4,X5),k4_metric_1(esk1_0,X4,X1))|X5!=k1_funct_1(esk3_0,X4)|X2!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(X3)|~m1_subset_1(X5,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_xxreal_0(X3,k4_metric_1(esk2_0,X5,X2))), inference(evalgc,[status(thm)],[c_0_45])).
cnf(c_0_55, plain, (r2_hidden(X1,k9_metric_1(X2,X3,X4))|r1_xxreal_0(X4,k4_metric_1(X2,X3,X1))|v2_struct_0(X2)|~v1_xreal_0(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)|~v8_metric_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47])])).
cnf(c_0_56, negated_conjecture, (l1_metric_1(esk2_0)), inference(evalgc,[status(thm)],[c_0_48])).
cnf(c_0_57, negated_conjecture, (v8_metric_1(esk2_0)), inference(evalgc,[status(thm)],[c_0_49])).
cnf(c_0_58, negated_conjecture, (~v2_struct_0(esk2_0)), inference(evalgc,[status(thm)],[c_0_50])).
cnf(c_0_59, plain, (u1_struct_0(X1)=u1_struct_0(k3_pcomps_1(X1))|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
fof(c_0_60, plain, ![X109, X110]:((~m1_subset_1(X109,k1_zfmisc_1(X110))|r1_tarski(X109,X110))&(~r1_tarski(X109,X110)|m1_subset_1(X109,k1_zfmisc_1(X110)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_61, plain, ![X123]:r1_tarski(X123,X123), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r1_tarski])])).
cnf(c_0_62, plain, (k4_metric_1(X1,X2,X3)=k4_metric_1(X1,X3,X2)|~v8_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_52])).
fof(c_0_63, plain, ![X102, X103, X104]:(~r2_hidden(X102,X103)|~m1_subset_1(X103,k1_zfmisc_1(X104))|m1_subset_1(X102,X104)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_64, plain, ![X128, X129, X130]:(~l1_metric_1(X128)|~m1_subset_1(X129,u1_struct_0(X128))|~v1_xreal_0(X130)|m1_subset_1(k9_metric_1(X128,X129,X130),k1_zfmisc_1(u1_struct_0(X128)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k9_metric_1])])).
cnf(c_0_65, plain, (~l1_metric_1(X2)|~v1_xreal_0(X1)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~r1_xxreal_0(X1,k2_metric_1(X2,X3,X4))|~r2_hidden(X4,k9_metric_1(X2,X3,X1))), inference(evalgc,[status(thm)],[c_0_53])).
cnf(c_0_66, negated_conjecture, (r1_xxreal_0(esk4_3(X1,X2,X3),k4_metric_1(esk1_0,X2,X4))|r1_xxreal_0(X1,k6_numbers)|r2_hidden(X5,k9_metric_1(esk2_0,X3,X1))|X3!=k1_funct_1(esk3_0,X2)|X5!=k1_funct_1(esk3_0,X4)|~v1_xreal_0(X1)|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X5,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_56]), c_0_57])]), c_0_58])])).
cnf(c_0_67, negated_conjecture, (l1_metric_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_68, negated_conjecture, (v8_metric_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_69, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(k3_pcomps_1(esk1_0)),u1_struct_0(k3_pcomps_1(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_70, plain, (u1_struct_0(X1)=u1_struct_0(k3_pcomps_1(X1))|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[c_0_59])).
fof(c_0_71, plain, ![X150, X151, X152]:(~v1_xboole_0(X150)|(~m1_subset_1(X152,k1_zfmisc_1(k2_zfmisc_1(X151,X150)))|v1_xboole_0(X152))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_relset_1])])])).
cnf(c_0_72, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_73, plain, (r1_tarski(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
fof(c_0_74, plain, ![X72, X73, X74, X75, X76]:((~r1_tarski(X72,X73)|(~r2_hidden(X74,X72)|r2_hidden(X74,X73)))&((r2_hidden(esk14_2(X75,X76),X75)|r1_tarski(X75,X76))&(~r2_hidden(esk14_2(X75,X76),X76)|r1_tarski(X75,X76)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_75, plain, (k4_metric_1(X1,X2,X3)=k4_metric_1(X1,X3,X2)|~v8_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_62])).
cnf(c_0_76, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_77, plain, (m1_subset_1(k9_metric_1(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_metric_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_xreal_0(X3)), inference(split_conjunct,[status(thm)],[c_0_64])).
fof(c_0_78, plain, ![X51, X52, X53, X56]:(((v1_xreal_0(esk11_3(X51,X52,X53))|~r2_hidden(X53,X52)|~m1_subset_1(X53,u1_struct_0(X51))|~v3_pre_topc(X52,k3_pcomps_1(X51))|~m1_subset_1(X52,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X51))))|(v2_struct_0(X51)|~v6_metric_1(X51)|~v7_metric_1(X51)|~v8_metric_1(X51)|~v9_metric_1(X51)|~l1_metric_1(X51)))&((~r1_xxreal_0(esk11_3(X51,X52,X53),k6_numbers)|~r2_hidden(X53,X52)|~m1_subset_1(X53,u1_struct_0(X51))|~v3_pre_topc(X52,k3_pcomps_1(X51))|~m1_subset_1(X52,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X51))))|(v2_struct_0(X51)|~v6_metric_1(X51)|~v7_metric_1(X51)|~v8_metric_1(X51)|~v9_metric_1(X51)|~l1_metric_1(X51)))&(r1_tarski(k9_metric_1(X51,X53,esk11_3(X51,X52,X53)),X52)|~r2_hidden(X53,X52)|~m1_subset_1(X53,u1_struct_0(X51))|~v3_pre_topc(X52,k3_pcomps_1(X51))|~m1_subset_1(X52,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X51))))|(v2_struct_0(X51)|~v6_metric_1(X51)|~v7_metric_1(X51)|~v8_metric_1(X51)|~v9_metric_1(X51)|~l1_metric_1(X51)))))&((m1_subset_1(esk12_2(X51,X52),u1_struct_0(X51))|v3_pre_topc(X52,k3_pcomps_1(X51))|~m1_subset_1(X52,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X51))))|(v2_struct_0(X51)|~v6_metric_1(X51)|~v7_metric_1(X51)|~v8_metric_1(X51)|~v9_metric_1(X51)|~l1_metric_1(X51)))&((r2_hidden(esk12_2(X51,X52),X52)|v3_pre_topc(X52,k3_pcomps_1(X51))|~m1_subset_1(X52,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X51))))|(v2_struct_0(X51)|~v6_metric_1(X51)|~v7_metric_1(X51)|~v8_metric_1(X51)|~v9_metric_1(X51)|~l1_metric_1(X51)))&(~v1_xreal_0(X56)|(r1_xxreal_0(X56,k6_numbers)|~r1_tarski(k9_metric_1(X51,esk12_2(X51,X52),X56),X52))|v3_pre_topc(X52,k3_pcomps_1(X51))|~m1_subset_1(X52,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X51))))|(v2_struct_0(X51)|~v6_metric_1(X51)|~v7_metric_1(X51)|~v8_metric_1(X51)|~v9_metric_1(X51)|~l1_metric_1(X51)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t15_topmetr])])])])])])).
cnf(c_0_79, plain, (~r2_hidden(X1,k9_metric_1(X2,X3,X4))|~r1_xxreal_0(X4,k4_metric_1(X2,X3,X1))|~v1_xreal_0(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)|~v8_metric_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_47])])).
cnf(c_0_80, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,X1),k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(esk4_3(X3,X4,X2),k4_metric_1(esk1_0,X4,X1))|r1_xxreal_0(X3,k6_numbers)|X2!=k1_funct_1(esk3_0,X4)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,X1),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_66])])).
cnf(c_0_81, negated_conjecture, (l1_metric_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_67])).
cnf(c_0_82, negated_conjecture, (v8_metric_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_68])).
cnf(c_0_83, negated_conjecture, (v1_xreal_0(esk4_3(X1,X2,X3))|r1_xxreal_0(X1,k6_numbers)|X3!=k1_funct_1(esk3_0,X2)|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_84, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(esk1_0)),u1_struct_0(k3_pcomps_1(esk2_0)))))), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_85, plain, ![X26, X27, X28]:(((~v1_funct_2(X28,X26,X27)|X26=k1_relset_1(X26,X28)|X27=k1_xboole_0|~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(X26,X27))))&(X26!=k1_relset_1(X26,X28)|v1_funct_2(X28,X26,X27)|X27=k1_xboole_0|~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(X26,X27)))))&((~v1_funct_2(X28,X26,X27)|X28=k1_xboole_0|X27!=k1_xboole_0|~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(X26,X27))))&(X28!=k1_xboole_0|v1_funct_2(X28,X26,X27)|X27!=k1_xboole_0|~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(X26,X27)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_2])])])).
cnf(c_0_86, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(k3_pcomps_1(esk1_0)),u1_struct_0(k3_pcomps_1(esk2_0)))), inference(evalgc,[status(thm)],[c_0_69])).
cnf(c_0_87, negated_conjecture, (u1_struct_0(k3_pcomps_1(esk2_0))=u1_struct_0(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_56])])).
fof(c_0_88, plain, ![X148]:(~v1_xboole_0(X148)|X148=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
cnf(c_0_89, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_90, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_72])).
cnf(c_0_91, plain, (r1_tarski(X1,X1)), inference(evalgc,[status(thm)],[c_0_73])).
cnf(c_0_92, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_93, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_94, plain, (r2_hidden(X1,k9_metric_1(X2,X3,X4))|r1_xxreal_0(X4,k4_metric_1(X2,X1,X3))|v2_struct_0(X2)|~v1_xreal_0(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)|~v8_metric_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_75])])).
cnf(c_0_95, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_76])).
cnf(c_0_96, plain, (m1_subset_1(k9_metric_1(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_metric_1(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_77])).
cnf(c_0_97, plain, (r1_tarski(k9_metric_1(X1,X2,esk11_3(X1,X3,X2)),X3)|v2_struct_0(X1)|~r2_hidden(X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_pre_topc(X3,k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_98, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,X1),k9_metric_1(esk2_0,X4,X3))|r1_xxreal_0(X3,k6_numbers)|~r2_hidden(X1,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)))|~v1_xreal_0(esk4_3(X3,X2,X4))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|X4!=k1_funct_1(esk3_0,X2)|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,X1),u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_81]), c_0_82])])])).
cnf(c_0_99, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|v1_xreal_0(esk4_3(X1,X2,X3))|X3!=k1_funct_1(esk3_0,X2)|~v1_xreal_0(X1)|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_83])).
cnf(c_0_100, plain, (r2_hidden(esk14_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_74])).
fof(c_0_101, plain, ![X120, X121]:(~v1_relat_1(X121)|~v4_relat_1(X121,X120)|m1_subset_1(k1_relset_1(X120,X121),k1_zfmisc_1(X120))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_relset_1])])).
fof(c_0_102, plain, ![X92, X93]:(~v1_relat_1(X93)|~v4_relat_1(X93,X92)|k1_relset_1(X92,X93)=k9_xtuple_0(X93)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
fof(c_0_103, plain, ![X89, X90, X91]:((v4_relat_1(X91,X89)|~m1_subset_1(X91,k1_zfmisc_1(k2_zfmisc_1(X89,X90))))&(v5_relat_1(X91,X90)|~m1_subset_1(X91,k1_zfmisc_1(k2_zfmisc_1(X89,X90))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
cnf(c_0_104, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(esk1_0)),u1_struct_0(k3_pcomps_1(esk2_0)))))), inference(evalgc,[status(thm)],[c_0_84])).
fof(c_0_105, plain, ![X81, X82, X83]:(~m1_subset_1(X83,k1_zfmisc_1(k2_zfmisc_1(X81,X82)))|v1_relat_1(X83)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_106, plain, (X2=k1_relset_1(X2,X1)|X3=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_107, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(k3_pcomps_1(esk1_0)),u1_struct_0(esk2_0))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_86, c_0_87])])).
cnf(c_0_108, negated_conjecture, (u1_struct_0(k3_pcomps_1(esk1_0))=u1_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_81])])).
cnf(c_0_109, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_110, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_111, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91])])).
cnf(c_0_112, plain, (r2_hidden(X3,X2)|~r2_hidden(X3,X1)|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_92])).
cnf(c_0_113, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_93])).
cnf(c_0_114, plain, (r2_hidden(X3,k9_metric_1(X2,X1,X4))|v2_struct_0(X2)|~r2_hidden(X1,k9_metric_1(X2,X3,X4))|~v1_xreal_0(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)|~v8_metric_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_94])])).
cnf(c_0_115, plain, (m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,k9_metric_1(X2,X3,X4))|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96])])).
cnf(c_0_116, plain, (v2_struct_0(X1)|r1_tarski(k9_metric_1(X1,X2,esk11_3(X1,X3,X2)),X3)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~r2_hidden(X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_pre_topc(X3,k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(evalgc,[status(thm)],[c_0_97])).
cnf(c_0_117, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,X1),k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(X3,k6_numbers)|X2!=k1_funct_1(esk3_0,X4)|~r2_hidden(X1,k9_metric_1(esk1_0,X4,esk4_3(X3,X4,X2)))|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,X1),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_118, plain, (r1_tarski(X1,X2)|r2_hidden(esk14_2(X1,X2),X1)), inference(evalgc,[status(thm)],[c_0_100])).
cnf(c_0_119, negated_conjecture, (v9_metric_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_120, negated_conjecture, (v7_metric_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_121, negated_conjecture, (v6_metric_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_122, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_123, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_124, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_125, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(esk1_0)),u1_struct_0(esk2_0))))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_104, c_0_87])])).
cnf(c_0_126, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_105])).
fof(c_0_127, plain, ![X126, X127]:(~v1_relat_1(X127)|~v5_relat_1(X127,X126)|m1_subset_1(k2_relset_1(X126,X127),k1_zfmisc_1(X126))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_relset_1])])).
fof(c_0_128, plain, ![X84, X85]:(~v1_relat_1(X85)|~v5_relat_1(X85,X84)|k2_relset_1(X84,X85)=k10_xtuple_0(X85)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
cnf(c_0_129, plain, (X3=k1_xboole_0|X2=k1_relset_1(X2,X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_106])).
cnf(c_0_130, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))), inference(rw,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_131, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_109])).
cnf(c_0_132, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111])])).
cnf(c_0_133, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113])])).
cnf(c_0_134, plain, (r2_hidden(X1,k9_metric_1(X2,X3,X4))|v2_struct_0(X2)|~r2_hidden(X3,k9_metric_1(X2,X1,X4))|~v1_xreal_0(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_metric_1(X2)|~v8_metric_1(X2)), inference(csr,[status(thm)],[c_0_114, c_0_115])).
fof(c_0_135, plain, ![X64, X65, X66, X67, X68, X69, X70]:((((r2_hidden(X67,k9_xtuple_0(X64))|~r2_hidden(X67,X66)|X66!=k8_relat_1(X64,X65)|(~v1_relat_1(X64)|~v1_funct_1(X64)))&(r2_hidden(k1_funct_1(X64,X67),X65)|~r2_hidden(X67,X66)|X66!=k8_relat_1(X64,X65)|(~v1_relat_1(X64)|~v1_funct_1(X64))))&(~r2_hidden(X68,k9_xtuple_0(X64))|~r2_hidden(k1_funct_1(X64,X68),X65)|r2_hidden(X68,X66)|X66!=k8_relat_1(X64,X65)|(~v1_relat_1(X64)|~v1_funct_1(X64))))&((~r2_hidden(esk13_3(X64,X69,X70),X70)|(~r2_hidden(esk13_3(X64,X69,X70),k9_xtuple_0(X64))|~r2_hidden(k1_funct_1(X64,esk13_3(X64,X69,X70)),X69))|X70=k8_relat_1(X64,X69)|(~v1_relat_1(X64)|~v1_funct_1(X64)))&((r2_hidden(esk13_3(X64,X69,X70),k9_xtuple_0(X64))|r2_hidden(esk13_3(X64,X69,X70),X70)|X70=k8_relat_1(X64,X69)|(~v1_relat_1(X64)|~v1_funct_1(X64)))&(r2_hidden(k1_funct_1(X64,esk13_3(X64,X69,X70)),X69)|r2_hidden(esk13_3(X64,X69,X70),X70)|X70=k8_relat_1(X64,X69)|(~v1_relat_1(X64)|~v1_funct_1(X64)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_funct_1])])])])])])).
cnf(c_0_136, plain, (r2_hidden(X1,X2)|v2_struct_0(X3)|~r2_hidden(X1,k9_metric_1(X3,X4,esk11_3(X3,X2,X4)))|~r2_hidden(X4,X2)|~v3_pre_topc(X2,k3_pcomps_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_116])])).
cnf(c_0_137, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)),k9_metric_1(esk2_0,X3,X2))|r1_xxreal_0(X2,k6_numbers)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)|X3!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(X2)|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118])])).
cnf(c_0_138, negated_conjecture, (v9_metric_1(esk2_0)), inference(evalgc,[status(thm)],[c_0_119])).
cnf(c_0_139, negated_conjecture, (v7_metric_1(esk2_0)), inference(evalgc,[status(thm)],[c_0_120])).
cnf(c_0_140, negated_conjecture, (v6_metric_1(esk2_0)), inference(evalgc,[status(thm)],[c_0_121])).
cnf(c_0_141, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_122])).
cnf(c_0_142, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_123])).
cnf(c_0_143, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_124])).
cnf(c_0_144, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))), inference(rw,[status(thm)],[c_0_125, c_0_108])).
cnf(c_0_145, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_126])).
cnf(c_0_146, plain, (m1_subset_1(k2_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_147, plain, (k2_relset_1(X2,X1)=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_128])).
cnf(c_0_148, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_103])).
fof(c_0_149, plain, ![X29, X30, X31, X33, X34, X35, X37]:((((r2_hidden(esk8_3(X29,X30,X31),k9_xtuple_0(X29))|~r2_hidden(X31,X30)|X30!=k10_xtuple_0(X29)|(~v1_relat_1(X29)|~v1_funct_1(X29)))&(X31=k1_funct_1(X29,esk8_3(X29,X30,X31))|~r2_hidden(X31,X30)|X30!=k10_xtuple_0(X29)|(~v1_relat_1(X29)|~v1_funct_1(X29))))&(~r2_hidden(X34,k9_xtuple_0(X29))|X33!=k1_funct_1(X29,X34)|r2_hidden(X33,X30)|X30!=k10_xtuple_0(X29)|(~v1_relat_1(X29)|~v1_funct_1(X29))))&((~r2_hidden(esk9_2(X29,X35),X35)|(~r2_hidden(X37,k9_xtuple_0(X29))|esk9_2(X29,X35)!=k1_funct_1(X29,X37))|X35=k10_xtuple_0(X29)|(~v1_relat_1(X29)|~v1_funct_1(X29)))&((r2_hidden(esk10_2(X29,X35),k9_xtuple_0(X29))|r2_hidden(esk9_2(X29,X35),X35)|X35=k10_xtuple_0(X29)|(~v1_relat_1(X29)|~v1_funct_1(X29)))&(esk9_2(X29,X35)=k1_funct_1(X29,esk10_2(X29,X35))|r2_hidden(esk9_2(X29,X35),X35)|X35=k10_xtuple_0(X29)|(~v1_relat_1(X29)|~v1_funct_1(X29)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_funct_1])])])])])])).
cnf(c_0_150, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)=u1_struct_0(esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130])])).
cnf(c_0_151, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_132])])).
cnf(c_0_152, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_153, plain, (r2_hidden(X1,u1_struct_0(X2))|~r2_hidden(X1,k9_metric_1(X2,X3,X4))|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_96])])).
cnf(c_0_154, plain, (r2_hidden(X1,k9_metric_1(X2,esk14_2(k9_metric_1(X2,X1,X3),X4),X3))|v2_struct_0(X2)|r1_tarski(k9_metric_1(X2,X1,X3),X4)|~v1_xreal_0(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_metric_1(X2)|~v8_metric_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_118])])).
fof(c_0_155, plain, ![X62]:(v2_struct_0(X62)|~l1_struct_0(X62)|~v1_xboole_0(u1_struct_0(X62))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_156, plain, ![X99]:(~l1_pre_topc(X99)|l1_struct_0(X99)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_157, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,k9_xtuple_0(X2))|~r2_hidden(k1_funct_1(X2,X1),X3)|X4!=k8_relat_1(X2,X3)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_135])).
cnf(c_0_158, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)),X2)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X3,u1_struct_0(esk2_0))|X3!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_137]), c_0_87]), c_0_56]), c_0_138]), c_0_57]), c_0_139]), c_0_140])]), c_0_58])])).
cnf(c_0_159, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_160, plain, (m1_subset_1(k9_xtuple_0(X2),k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_142])])).
cnf(c_0_161, negated_conjecture, (v4_relat_1(esk3_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144])])).
cnf(c_0_162, negated_conjecture, (v1_relat_1(esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_144])])).
cnf(c_0_163, plain, (m1_subset_1(k2_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_146])).
cnf(c_0_164, plain, (k2_relset_1(X2,X1)=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_147])).
cnf(c_0_165, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(evalgc,[status(thm)],[c_0_148])).
cnf(c_0_166, plain, (r2_hidden(X3,X4)|~r2_hidden(X1,k9_xtuple_0(X2))|X3!=k1_funct_1(X2,X1)|X4!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_149])).
cnf(c_0_167, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)=u1_struct_0(esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_150, c_0_144])])).
fof(c_0_168, plain, ![X111, X112]:(~r2_hidden(X111,X112)|~v1_xboole_0(X112)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
cnf(c_0_169, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_151])])).
cnf(c_0_170, plain, (v1_xboole_0(k1_xboole_0)), inference(evalgc,[status(thm)],[c_0_152])).
cnf(c_0_171, plain, (r2_hidden(X1,u1_struct_0(X2))|r1_tarski(k9_metric_1(X2,X1,X3),X4)|v2_struct_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(esk14_2(k9_metric_1(X2,X1,X3),X4),u1_struct_0(X2))|~l1_metric_1(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~v8_metric_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_154])])).
cnf(c_0_172, plain, (m1_subset_1(esk14_2(k9_metric_1(X1,X2,X3),X4),u1_struct_0(X1))|r1_tarski(k9_metric_1(X1,X2,X3),X4)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_118])])).
cnf(c_0_173, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_155])).
cnf(c_0_174, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_156])).
fof(c_0_175, plain, ![X61]:(~l1_metric_1(X61)|l1_pre_topc(k3_pcomps_1(X61))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_pcomps_1])])).
cnf(c_0_176, plain, (r2_hidden(X1,X4)|X4!=k8_relat_1(X2,X3)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k9_xtuple_0(X2))|~r2_hidden(k1_funct_1(X2,X1),X3)), inference(evalgc,[status(thm)],[c_0_157])).
cnf(c_0_177, negated_conjecture, (r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)|r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)),X2)|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[c_0_158, c_0_95])).
cnf(c_0_178, negated_conjecture, (v1_funct_1(esk3_0)), inference(evalgc,[status(thm)],[c_0_159])).
cnf(c_0_179, negated_conjecture, (m1_subset_1(k9_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_161]), c_0_162])])])).
cnf(c_0_180, plain, (m1_subset_1(k10_xtuple_0(X2),k1_zfmisc_1(X1))|~v5_relat_1(X2,X1)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_164])])).
cnf(c_0_181, negated_conjecture, (v5_relat_1(esk3_0,u1_struct_0(esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_144])])).
cnf(c_0_182, plain, (r2_hidden(X3,X4)|X4!=k10_xtuple_0(X2)|X3!=k1_funct_1(X2,X1)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k9_xtuple_0(X2))), inference(evalgc,[status(thm)],[c_0_166])).
cnf(c_0_183, negated_conjecture, (k9_xtuple_0(esk3_0)=u1_struct_0(esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_167]), c_0_161]), c_0_162])])])).
cnf(c_0_184, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_168])).
cnf(c_0_185, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_170])])).
cnf(c_0_186, plain, (r1_tarski(k9_metric_1(X1,X2,X3),X4)|r2_hidden(X2,u1_struct_0(X1))|v2_struct_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v8_metric_1(X1)), inference(csr,[status(thm)],[c_0_171, c_0_172])).
fof(c_0_187, plain, ![X118, X119]:(~m1_subset_1(X118,X119)|(v1_xboole_0(X119)|r2_hidden(X118,X119))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_188, plain, ![X149]:(~v2_struct_0(X149)|~l1_struct_0(X149)|v1_xboole_0(u1_struct_0(X149))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_struct_0])])).
fof(c_0_189, plain, ![X79]:(~l1_metric_1(X79)|l1_struct_0(X79)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_metric_1])])).
cnf(c_0_190, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_173])).
cnf(c_0_191, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[c_0_174])).
cnf(c_0_192, plain, (l1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_175])).
cnf(c_0_193, negated_conjecture, (r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),X5)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|X5!=k8_relat_1(esk3_0,X2)|~r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),k9_xtuple_0(esk3_0))|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_177]), c_0_162]), c_0_178])])])).
cnf(c_0_194, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k9_xtuple_0(esk3_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_179])])).
cnf(c_0_195, negated_conjecture, (m1_subset_1(k10_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_181]), c_0_162])])])).
cnf(c_0_196, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k1_funct_1(esk3_0,X3)|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(X3,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_183]), c_0_162]), c_0_178])])])).
cnf(c_0_197, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_184])).
cnf(c_0_198, plain, (X1=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_185])])).
cnf(c_0_199, plain, (m1_subset_1(k9_metric_1(X1,X2,X3),k1_zfmisc_1(X4))|r2_hidden(X2,u1_struct_0(X1))|v2_struct_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v8_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_186])])).
cnf(c_0_200, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_187])).
cnf(c_0_201, plain, (m1_subset_1(esk12_2(X1,X2),u1_struct_0(X1))|v3_pre_topc(X2,k3_pcomps_1(X1))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_202, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_188])).
cnf(c_0_203, plain, (l1_struct_0(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_189])).
fof(c_0_204, plain, ![X122]:(v2_struct_0(X122)|~l1_metric_1(X122)|~v2_struct_0(k3_pcomps_1(X122))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_pcomps_1])])])).
cnf(c_0_205, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_190, c_0_191])])).
cnf(c_0_206, plain, (l1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[c_0_192])).
cnf(c_0_207, negated_conjecture, (r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)|r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),X5)|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|X5!=k8_relat_1(esk3_0,X2)|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)),u1_struct_0(esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[c_0_193, c_0_194])).
cnf(c_0_208, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,k10_xtuple_0(esk3_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_195])])).
cnf(c_0_209, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(esk3_0,X1),X2)|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_196])])).
cnf(c_0_210, plain, (r1_tarski(k9_metric_1(X1,X2,X3),X4)|v2_struct_0(X1)|~v1_xboole_0(k9_metric_1(X1,esk14_2(k9_metric_1(X1,X2,X3),X4),X3))|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v8_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_154])])).
cnf(c_0_211, plain, (k9_metric_1(X1,X2,X3)=k1_xboole_0|r2_hidden(X2,u1_struct_0(X1))|v2_struct_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v8_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_199])])).
cnf(c_0_212, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,X3))|~v4_relat_1(X3,X2)|~v1_relat_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_141])])).
cnf(c_0_213, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_200])).
cnf(c_0_214, plain, (v2_struct_0(X1)|v3_pre_topc(X2,k3_pcomps_1(X1))|m1_subset_1(esk12_2(X1,X2),u1_struct_0(X1))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(evalgc,[status(thm)],[c_0_201])).
cnf(c_0_215, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_202])).
cnf(c_0_216, plain, (l1_struct_0(X1)|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[c_0_203])).
cnf(c_0_217, plain, (v2_struct_0(X1)|~l1_metric_1(X1)|~v2_struct_0(k3_pcomps_1(X1))), inference(split_conjunct,[status(thm)],[c_0_204])).
cnf(c_0_218, plain, (v2_struct_0(k3_pcomps_1(X1))|~v1_xboole_0(u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_206])])).
cnf(c_0_219, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_220, negated_conjecture, (r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)|r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),X5)|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|X5!=k8_relat_1(esk3_0,X2)|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)),u1_struct_0(esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_183])])).
cnf(c_0_221, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,X1),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_209])])).
cnf(c_0_222, plain, (r1_tarski(k9_metric_1(X1,X2,X3),X4)|v2_struct_0(X1)|r2_hidden(esk14_2(k9_metric_1(X1,X2,X3),X4),u1_struct_0(X1))|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(esk14_2(k9_metric_1(X1,X2,X3),X4),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_211]), c_0_170])])])).
cnf(c_0_223, plain, (r2_hidden(X1,X2)|~v4_relat_1(X3,X2)|~r2_hidden(X1,k9_xtuple_0(X3))|~v1_relat_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_142])])).
cnf(c_0_224, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(esk12_2(X1,X2),u1_struct_0(X1))|v3_pre_topc(X2,k3_pcomps_1(X1))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_214])])).
cnf(c_0_225, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_216])])).
cnf(c_0_226, negated_conjecture, (v9_metric_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_227, negated_conjecture, (v7_metric_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_228, negated_conjecture, (v6_metric_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_229, plain, (v2_struct_0(X1)|~l1_metric_1(X1)|~v2_struct_0(k3_pcomps_1(X1))), inference(evalgc,[status(thm)],[c_0_217])).
cnf(c_0_230, negated_conjecture, (v2_struct_0(k3_pcomps_1(esk1_0))|~v1_xboole_0(u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_108]), c_0_81])])])).
cnf(c_0_231, negated_conjecture, (~v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_219])).
cnf(c_0_232, plain, (r1_tarski(X1,X2)|~r2_hidden(esk14_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_233, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)|r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),X5)|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|X5!=k8_relat_1(esk3_0,X2)|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[c_0_220, c_0_221])).
cnf(c_0_234, plain, (r1_tarski(k9_metric_1(X1,X2,X3),X4)|r2_hidden(esk14_2(k9_metric_1(X1,X2,X3),X4),u1_struct_0(X1))|v2_struct_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v8_metric_1(X1)), inference(csr,[status(thm)],[c_0_222, c_0_172])).
cnf(c_0_235, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|~v4_relat_1(esk3_0,X2)|~r2_hidden(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_183]), c_0_162])])])).
cnf(c_0_236, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(esk12_2(X1,X2),u1_struct_0(X1))|v3_pre_topc(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(csr,[status(thm)],[c_0_224, c_0_225])).
cnf(c_0_237, negated_conjecture, (v9_metric_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_226])).
cnf(c_0_238, negated_conjecture, (v7_metric_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_227])).
cnf(c_0_239, negated_conjecture, (v6_metric_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_228])).
cnf(c_0_240, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_230]), c_0_81])]), c_0_231])])).
cnf(c_0_241, plain, (r1_xxreal_0(X1,k6_numbers)|v3_pre_topc(X3,k3_pcomps_1(X2))|v2_struct_0(X2)|~v1_xreal_0(X1)|~r1_tarski(k9_metric_1(X2,esk12_2(X2,X3),X1),X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_242, plain, (r1_tarski(X1,X2)|~r2_hidden(esk14_2(X1,X2),X2)), inference(evalgc,[status(thm)],[c_0_232])).
cnf(c_0_243, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)|r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),X5)|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|X5!=k8_relat_1(esk3_0,X2)|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v1_xreal_0(esk4_3(esk11_3(esk2_0,X2,X3),X1,X3))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_234]), c_0_81]), c_0_82])]), c_0_231])])).
cnf(c_0_244, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk12_2(esk1_0,X1),X2)|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|~v4_relat_1(esk3_0,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_236]), c_0_108]), c_0_81]), c_0_237]), c_0_82]), c_0_238]), c_0_239])]), c_0_240])])).
cnf(c_0_245, plain, (v2_struct_0(X2)|r1_xxreal_0(X1,k6_numbers)|v3_pre_topc(X3,k3_pcomps_1(X2))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v1_xreal_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~r1_tarski(k9_metric_1(X2,esk12_2(X2,X3),X1),X3)), inference(evalgc,[status(thm)],[c_0_241])).
cnf(c_0_246, negated_conjecture, (r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)|u1_struct_0(esk2_0)=k1_xboole_0|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|X4!=k8_relat_1(esk3_0,X2)|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk4_3(esk11_3(esk2_0,X2,X3),X1,X3))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_243])])).
cnf(c_0_247, negated_conjecture, (m1_subset_1(esk12_2(esk1_0,X1),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|~v4_relat_1(esk3_0,k9_xtuple_0(esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_194, c_0_244])])).
cnf(c_0_248, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(esk4_3(X1,X2,X3),k6_numbers)|X3!=k1_funct_1(esk3_0,X2)|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_249, negated_conjecture, (v3_pre_topc(X1,k3_pcomps_1(esk1_0))|r1_xxreal_0(esk4_3(esk11_3(esk2_0,X2,X3),esk12_2(esk1_0,X1),X3),k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|~v1_xreal_0(esk4_3(esk11_3(esk2_0,X2,X3),esk12_2(esk1_0,X1),X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|X1!=k8_relat_1(esk3_0,X2)|X3!=k1_funct_1(esk3_0,esk12_2(esk1_0,X1))|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(esk12_2(esk1_0,X1),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_246]), c_0_108]), c_0_81]), c_0_237]), c_0_82]), c_0_238]), c_0_239])]), c_0_231])])).
cnf(c_0_250, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|m1_subset_1(esk12_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_183]), c_0_161])])])).
cnf(c_0_251, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|X3!=k1_funct_1(esk3_0,X2)|~v1_xreal_0(X1)|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r1_xxreal_0(esk4_3(X1,X2,X3),k6_numbers)), inference(evalgc,[status(thm)],[c_0_248])).
cnf(c_0_252, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|r1_xxreal_0(esk4_3(esk11_3(esk2_0,X2,X3),esk12_2(esk1_0,X1),X3),k6_numbers)|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|X3!=k1_funct_1(esk3_0,esk12_2(esk1_0,X1))|X1!=k8_relat_1(esk3_0,X2)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk4_3(esk11_3(esk2_0,X2,X3),esk12_2(esk1_0,X1),X3))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(csr,[status(thm)],[c_0_249, c_0_250])).
cnf(c_0_253, negated_conjecture, (r1_xxreal_0(esk11_3(esk2_0,X1,X2),k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X3,k3_pcomps_1(esk1_0))|X2!=k1_funct_1(esk3_0,esk12_2(esk1_0,X3))|~v1_xreal_0(esk11_3(esk2_0,X1,X2))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(esk12_2(esk1_0,X3),u1_struct_0(esk1_0))|X3!=k8_relat_1(esk3_0,X1)|~r2_hidden(X2,X1)|~v3_pre_topc(X1,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk4_3(esk11_3(esk2_0,X1,X2),esk12_2(esk1_0,X3),X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_252])])).
fof(c_0_254, plain, ![X98]:(~l1_pre_topc(X98)|(~v1_pre_topc(X98)|X98=g1_pre_topc(u1_struct_0(X98),u1_pre_topc(X98)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[abstractness_v1_pre_topc])])).
cnf(c_0_255, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|X3!=k1_funct_1(esk3_0,esk12_2(esk1_0,X1))|X1!=k8_relat_1(esk3_0,X2)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk4_3(esk11_3(esk2_0,X2,X3),esk12_2(esk1_0,X1),X3))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_253, c_0_95]), c_0_250])).
cnf(c_0_256, plain, (X1=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X1)|~v1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_254])).
cnf(c_0_257, plain, (u1_pre_topc(k3_pcomps_1(X1))=k2_pcomps_1(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
fof(c_0_258, plain, ![X78]:((v1_pre_topc(k3_pcomps_1(X78))|~l1_metric_1(X78))&(v2_pre_topc(k3_pcomps_1(X78))|~l1_metric_1(X78))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_pcomps_1])])])).
fof(c_0_259, plain, ![X94]:(~l1_pre_topc(X94)|m1_subset_1(u1_pre_topc(X94),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X94))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_pre_topc])])).
cnf(c_0_260, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|X3!=k1_funct_1(esk3_0,esk12_2(esk1_0,X1))|X1!=k8_relat_1(esk3_0,X2)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(esk12_2(esk1_0,X1),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_99])])).
cnf(c_0_261, plain, (X1=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X1)|~v1_pre_topc(X1)), inference(evalgc,[status(thm)],[c_0_256])).
cnf(c_0_262, plain, (u1_pre_topc(k3_pcomps_1(X1))=k2_pcomps_1(X1)|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[c_0_257])).
cnf(c_0_263, plain, (v1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_258])).
cnf(c_0_264, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_259])).
cnf(c_0_265, plain, (v2_struct_0(X1)|~r1_xxreal_0(esk11_3(X1,X2,X3),k6_numbers)|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~v3_pre_topc(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_266, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|X3!=k1_funct_1(esk3_0,esk12_2(esk1_0,X1))|X1!=k8_relat_1(esk3_0,X2)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_260, c_0_95]), c_0_250])).
fof(c_0_267, plain, ![X124, X125]:((v1_pre_topc(g1_pre_topc(X124,X125))|~m1_subset_1(X125,k1_zfmisc_1(k1_zfmisc_1(X124))))&(l1_pre_topc(g1_pre_topc(X124,X125))|~m1_subset_1(X125,k1_zfmisc_1(k1_zfmisc_1(X124))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_g1_pre_topc])])])).
cnf(c_0_268, plain, (g1_pre_topc(u1_struct_0(k3_pcomps_1(X1)),k2_pcomps_1(X1))=k3_pcomps_1(X1)|~v1_pre_topc(k3_pcomps_1(X1))|~l1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_261, c_0_262])])).
cnf(c_0_269, plain, (v1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[c_0_263])).
cnf(c_0_270, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[c_0_264])).
cnf(c_0_271, plain, (v2_struct_0(X1)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~v3_pre_topc(X2,k3_pcomps_1(X1))|~r1_xxreal_0(esk11_3(X1,X2,X3),k6_numbers)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(evalgc,[status(thm)],[c_0_265])).
cnf(c_0_272, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|r1_xxreal_0(esk11_3(esk2_0,X2,k1_funct_1(esk3_0,esk12_2(esk1_0,X1))),k6_numbers)|X1!=k8_relat_1(esk3_0,X2)|~r2_hidden(k1_funct_1(esk3_0,esk12_2(esk1_0,X1)),X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,k1_funct_1(esk3_0,esk12_2(esk1_0,X1))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_266])])).
fof(c_0_273, plain, ![X105, X106, X107, X108]:(~m1_subset_1(X107,k1_zfmisc_1(k2_zfmisc_1(X105,X106)))|m1_subset_1(k8_relset_1(X105,X106,X107,X108),k1_zfmisc_1(X105))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_relset_1])])).
fof(c_0_274, plain, ![X43, X44, X45, X46]:(~m1_subset_1(X45,k1_zfmisc_1(k2_zfmisc_1(X43,X44)))|k8_relset_1(X43,X44,X45,X46)=k8_relat_1(X45,X46)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_relset_1])])).
cnf(c_0_275, plain, (l1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_267])).
cnf(c_0_276, plain, (g1_pre_topc(u1_struct_0(k3_pcomps_1(X1)),k2_pcomps_1(X1))=k3_pcomps_1(X1)|~l1_metric_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_268, c_0_206]), c_0_269])).
cnf(c_0_277, plain, (m1_subset_1(k2_pcomps_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))|~l1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_270, c_0_262])])).
fof(c_0_278, plain, ![X47, X48, X49, X50]:(~v1_xreal_0(X47)|(~v9_metric_1(X48)|~l1_metric_1(X48)|(~m1_subset_1(X49,u1_struct_0(X48))|(~m1_subset_1(X50,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X48))))|(X50!=k9_metric_1(X48,X49,X47)|v3_pre_topc(X50,k3_pcomps_1(X48))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t14_topmetr])])])).
fof(c_0_279, plain, ![X19, X20, X21]:((v1_xreal_0(esk5_3(X19,X20,X21))|v5_pre_topc(X21,X19,k3_pcomps_1(X20))|(~v1_funct_1(X21)|~v1_funct_2(X21,u1_struct_0(X19),u1_struct_0(k3_pcomps_1(X20)))|~m1_subset_1(X21,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X19),u1_struct_0(k3_pcomps_1(X20))))))|(v2_struct_0(X20)|~v6_metric_1(X20)|~v7_metric_1(X20)|~v8_metric_1(X20)|~v9_metric_1(X20)|~l1_metric_1(X20))|(v2_struct_0(X19)|~v2_pre_topc(X19)|~l1_pre_topc(X19)))&((m1_subset_1(esk6_3(X19,X20,X21),u1_struct_0(X20))|v5_pre_topc(X21,X19,k3_pcomps_1(X20))|(~v1_funct_1(X21)|~v1_funct_2(X21,u1_struct_0(X19),u1_struct_0(k3_pcomps_1(X20)))|~m1_subset_1(X21,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X19),u1_struct_0(k3_pcomps_1(X20))))))|(v2_struct_0(X20)|~v6_metric_1(X20)|~v7_metric_1(X20)|~v8_metric_1(X20)|~v9_metric_1(X20)|~l1_metric_1(X20))|(v2_struct_0(X19)|~v2_pre_topc(X19)|~l1_pre_topc(X19)))&((m1_subset_1(esk7_3(X19,X20,X21),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X20))))|v5_pre_topc(X21,X19,k3_pcomps_1(X20))|(~v1_funct_1(X21)|~v1_funct_2(X21,u1_struct_0(X19),u1_struct_0(k3_pcomps_1(X20)))|~m1_subset_1(X21,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X19),u1_struct_0(k3_pcomps_1(X20))))))|(v2_struct_0(X20)|~v6_metric_1(X20)|~v7_metric_1(X20)|~v8_metric_1(X20)|~v9_metric_1(X20)|~l1_metric_1(X20))|(v2_struct_0(X19)|~v2_pre_topc(X19)|~l1_pre_topc(X19)))&((esk7_3(X19,X20,X21)=k9_metric_1(X20,esk6_3(X19,X20,X21),esk5_3(X19,X20,X21))|v5_pre_topc(X21,X19,k3_pcomps_1(X20))|(~v1_funct_1(X21)|~v1_funct_2(X21,u1_struct_0(X19),u1_struct_0(k3_pcomps_1(X20)))|~m1_subset_1(X21,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X19),u1_struct_0(k3_pcomps_1(X20))))))|(v2_struct_0(X20)|~v6_metric_1(X20)|~v7_metric_1(X20)|~v8_metric_1(X20)|~v9_metric_1(X20)|~l1_metric_1(X20))|(v2_struct_0(X19)|~v2_pre_topc(X19)|~l1_pre_topc(X19)))&((~r1_xxreal_0(esk5_3(X19,X20,X21),k6_numbers)|v5_pre_topc(X21,X19,k3_pcomps_1(X20))|(~v1_funct_1(X21)|~v1_funct_2(X21,u1_struct_0(X19),u1_struct_0(k3_pcomps_1(X20)))|~m1_subset_1(X21,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X19),u1_struct_0(k3_pcomps_1(X20))))))|(v2_struct_0(X20)|~v6_metric_1(X20)|~v7_metric_1(X20)|~v8_metric_1(X20)|~v9_metric_1(X20)|~l1_metric_1(X20))|(v2_struct_0(X19)|~v2_pre_topc(X19)|~l1_pre_topc(X19)))&(~v3_pre_topc(k8_relset_1(u1_struct_0(X19),u1_struct_0(k3_pcomps_1(X20)),X21,esk7_3(X19,X20,X21)),X19)|v5_pre_topc(X21,X19,k3_pcomps_1(X20))|(~v1_funct_1(X21)|~v1_funct_2(X21,u1_struct_0(X19),u1_struct_0(k3_pcomps_1(X20)))|~m1_subset_1(X21,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X19),u1_struct_0(k3_pcomps_1(X20))))))|(v2_struct_0(X20)|~v6_metric_1(X20)|~v7_metric_1(X20)|~v8_metric_1(X20)|~v9_metric_1(X20)|~l1_metric_1(X20))|(v2_struct_0(X19)|~v2_pre_topc(X19)|~l1_pre_topc(X19)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t16_jordan2b])])])])])])).
cnf(c_0_280, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|~r2_hidden(k1_funct_1(esk3_0,esk12_2(esk1_0,X1)),X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(k1_funct_1(esk3_0,esk12_2(esk1_0,X1)),u1_struct_0(esk2_0))|X1!=k8_relat_1(esk3_0,X2)|~v1_xreal_0(esk11_3(esk2_0,X2,k1_funct_1(esk3_0,esk12_2(esk1_0,X1))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_272]), c_0_87]), c_0_56]), c_0_138]), c_0_57]), c_0_139]), c_0_140])]), c_0_58])])).
cnf(c_0_281, plain, (v1_xreal_0(esk11_3(X1,X2,X3))|v2_struct_0(X1)|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~v3_pre_topc(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_282, plain, (r2_hidden(k1_funct_1(X1,X2),X3)|~r2_hidden(X2,X4)|X4!=k8_relat_1(X1,X3)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_135])).
cnf(c_0_283, plain, (m1_subset_1(k8_relset_1(X2,X3,X1,X4),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_273])).
cnf(c_0_284, plain, (k8_relset_1(X2,X3,X1,X4)=k8_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_274])).
cnf(c_0_285, plain, (l1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(evalgc,[status(thm)],[c_0_275])).
cnf(c_0_286, negated_conjecture, (g1_pre_topc(u1_struct_0(esk1_0),k2_pcomps_1(esk1_0))=k3_pcomps_1(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_276, c_0_108]), c_0_81])])])).
cnf(c_0_287, plain, (m1_subset_1(k2_pcomps_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))|~l1_metric_1(X1)), inference(csr,[status(thm)],[c_0_277, c_0_206])).
cnf(c_0_288, plain, (v3_pre_topc(X4,k3_pcomps_1(X2))|~v1_xreal_0(X1)|~v9_metric_1(X2)|~l1_metric_1(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|X4!=k9_metric_1(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_278])).
cnf(c_0_289, plain, (esk7_3(X1,X2,X3)=k9_metric_1(X2,esk6_3(X1,X2,X3),esk5_3(X1,X2,X3))|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_279])).
cnf(c_0_290, plain, (v5_pre_topc(X3,X1,k3_pcomps_1(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)),X3,esk7_3(X1,X2,X3)),X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_279])).
cnf(c_0_291, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|X1!=k8_relat_1(esk3_0,X2)|~r2_hidden(k1_funct_1(esk3_0,esk12_2(esk1_0,X1)),X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,k1_funct_1(esk3_0,esk12_2(esk1_0,X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[c_0_280, c_0_95])).
cnf(c_0_292, plain, (v2_struct_0(X1)|v1_xreal_0(esk11_3(X1,X2,X3))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~v3_pre_topc(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(evalgc,[status(thm)],[c_0_281])).
cnf(c_0_293, plain, (r2_hidden(k1_funct_1(X1,X2),X3)|X4!=k8_relat_1(X1,X3)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X2,X4)), inference(evalgc,[status(thm)],[c_0_282])).
cnf(c_0_294, plain, (r2_hidden(esk12_2(X1,X2),X2)|v3_pre_topc(X2,k3_pcomps_1(X1))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_295, plain, (m1_subset_1(k8_relset_1(X2,X3,X1,X4),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_283])).
cnf(c_0_296, plain, (k8_relset_1(X2,X3,X1,X4)=k8_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_284])).
cnf(c_0_297, negated_conjecture, (l1_pre_topc(k3_pcomps_1(esk1_0))|~m1_subset_1(k2_pcomps_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_285, c_0_286])])).
cnf(c_0_298, negated_conjecture, (m1_subset_1(k2_pcomps_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_287, c_0_108]), c_0_81])])])).
cnf(c_0_299, plain, (v3_pre_topc(X4,k3_pcomps_1(X2))|X4!=k9_metric_1(X2,X3,X1)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v1_xreal_0(X1)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))), inference(evalgc,[status(thm)],[c_0_288])).
cnf(c_0_300, plain, (esk7_3(X1,X2,X3)=k9_metric_1(X2,esk6_3(X1,X2,X3),esk5_3(X1,X2,X3))|v2_struct_0(X2)|v2_struct_0(X1)|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v1_funct_1(X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))), inference(evalgc,[status(thm)],[c_0_289])).
cnf(c_0_301, plain, (m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X2))|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_279])).
cnf(c_0_302, plain, (v1_xreal_0(esk5_3(X1,X2,X3))|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_279])).
cnf(c_0_303, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v1_funct_1(X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))|~v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)),X3,esk7_3(X1,X2,X3)),X1)), inference(evalgc,[status(thm)],[c_0_290])).
cnf(c_0_304, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|X1!=k8_relat_1(esk3_0,X2)|~r2_hidden(k1_funct_1(esk3_0,esk12_2(esk1_0,X1)),X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_funct_1(esk3_0,esk12_2(esk1_0,X1)),u1_struct_0(esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_292]), c_0_87]), c_0_56]), c_0_138]), c_0_57]), c_0_139]), c_0_140])]), c_0_58])])).
cnf(c_0_305, plain, (r2_hidden(k1_funct_1(X1,X2),X3)|~r2_hidden(X2,k8_relat_1(X1,X3))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_293])])).
cnf(c_0_306, plain, (v2_struct_0(X1)|v3_pre_topc(X2,k3_pcomps_1(X1))|r2_hidden(esk12_2(X1,X2),X2)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))), inference(evalgc,[status(thm)],[c_0_294])).
cnf(c_0_307, plain, (m1_subset_1(k8_relat_1(X3,X4),k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_295, c_0_296])])).
cnf(c_0_308, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_191])])).
cnf(c_0_309, negated_conjecture, (l1_pre_topc(k3_pcomps_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_297, c_0_298])])).
cnf(c_0_310, plain, (v3_pre_topc(X1,k3_pcomps_1(X2))|v5_pre_topc(X4,X3,k3_pcomps_1(X2))|v2_struct_0(X3)|v2_struct_0(X2)|X1!=esk7_3(X3,X2,X4)|~v1_xreal_0(esk5_3(X3,X2,X4))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(esk6_3(X3,X2,X4),u1_struct_0(X2))|~l1_metric_1(X2)|~v9_metric_1(X2)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X4)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_299, c_0_300])])).
cnf(c_0_311, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X2))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v1_funct_1(X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))), inference(evalgc,[status(thm)],[c_0_301])).
cnf(c_0_312, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v1_xreal_0(esk5_3(X1,X2,X3))|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v1_funct_1(X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))), inference(evalgc,[status(thm)],[c_0_302])).
cnf(c_0_313, negated_conjecture, (v5_pre_topc(X1,X2,k3_pcomps_1(esk2_0))|v2_struct_0(X2)|~v3_pre_topc(k8_relset_1(u1_struct_0(X2),u1_struct_0(esk2_0),X1,esk7_3(X2,esk2_0,X1)),X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_303, c_0_87]), c_0_56]), c_0_138]), c_0_57]), c_0_139]), c_0_140])]), c_0_58])])).
cnf(c_0_314, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|X1!=k8_relat_1(esk3_0,X2)|~r2_hidden(k1_funct_1(esk3_0,esk12_2(esk1_0,X1)),X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[c_0_304, c_0_95])).
cnf(c_0_315, plain, (r2_hidden(k1_funct_1(X1,esk12_2(X2,k8_relat_1(X1,X3))),X3)|v3_pre_topc(k8_relat_1(X1,X3),k3_pcomps_1(X2))|v2_struct_0(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(k8_relat_1(X1,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_306])])).
cnf(c_0_316, negated_conjecture, (m1_subset_1(k8_relat_1(esk3_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_307, c_0_144])])).
cnf(c_0_317, negated_conjecture, (~v5_pre_topc(esk3_0,k3_pcomps_1(esk1_0),k3_pcomps_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_318, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~v2_struct_0(k3_pcomps_1(esk1_0))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_308, c_0_309]), c_0_108])])).
cnf(c_0_319, plain, (v3_pre_topc(X1,k3_pcomps_1(X2))|v5_pre_topc(X3,X4,k3_pcomps_1(X2))|v2_struct_0(X2)|v2_struct_0(X4)|X1!=esk7_3(X4,X2,X3)|~l1_pre_topc(X4)|~v2_pre_topc(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_310, c_0_311]), c_0_312])).
cnf(c_0_320, plain, (m1_subset_1(esk7_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_279])).
cnf(c_0_321, negated_conjecture, (v5_pre_topc(X1,X2,k3_pcomps_1(esk2_0))|v2_struct_0(X2)|~v3_pre_topc(k8_relat_1(X1,esk7_3(X2,esk2_0,X1)),X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_313, c_0_296])])).
cnf(c_0_322, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X1),k3_pcomps_1(esk1_0))|~v3_pre_topc(X1,k3_pcomps_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_314, c_0_315]), c_0_316]), c_0_162]), c_0_108]), c_0_316]), c_0_178]), c_0_81]), c_0_237]), c_0_82]), c_0_238]), c_0_239])]), c_0_231])])).
cnf(c_0_323, negated_conjecture, (~v5_pre_topc(esk3_0,k3_pcomps_1(esk1_0),k3_pcomps_1(esk2_0))), inference(evalgc,[status(thm)],[c_0_317])).
cnf(c_0_324, negated_conjecture, (~v2_struct_0(k3_pcomps_1(esk1_0))), inference(sr,[status(thm)],[c_0_318, c_0_240])).
cnf(c_0_325, plain, (v3_pre_topc(esk7_3(X1,X2,X3),k3_pcomps_1(X2))|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(esk7_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_319])])).
cnf(c_0_326, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|m1_subset_1(esk7_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v1_funct_1(X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))), inference(evalgc,[status(thm)],[c_0_320])).
cnf(c_0_327, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|~v2_pre_topc(k3_pcomps_1(esk1_0))|~v3_pre_topc(esk7_3(k3_pcomps_1(esk1_0),esk2_0,esk3_0),k3_pcomps_1(esk2_0))|~m1_subset_1(esk7_3(k3_pcomps_1(esk1_0),esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_321, c_0_322]), c_0_309]), c_0_108]), c_0_144]), c_0_108]), c_0_130]), c_0_178])]), c_0_323]), c_0_324])])).
cnf(c_0_328, plain, (v3_pre_topc(esk7_3(X1,X2,X3),k3_pcomps_1(X2))|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)), inference(csr,[status(thm)],[c_0_325, c_0_326])).
cnf(c_0_329, plain, (m1_subset_1(esk7_3(X2,X1,X3),k1_zfmisc_1(u1_struct_0(X1)))|v5_pre_topc(X3,X2,k3_pcomps_1(X1))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_xreal_0(esk5_3(X2,X1,X3))|~m1_subset_1(esk6_3(X2,X1,X3),u1_struct_0(X1))|~l1_metric_1(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X1)))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X3)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_300])])).
cnf(c_0_330, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|~v2_pre_topc(k3_pcomps_1(esk1_0))|~m1_subset_1(esk7_3(k3_pcomps_1(esk1_0),esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_328]), c_0_309]), c_0_108]), c_0_87]), c_0_144]), c_0_108]), c_0_87]), c_0_130]), c_0_178]), c_0_56]), c_0_138]), c_0_57]), c_0_139]), c_0_140])]), c_0_323]), c_0_324]), c_0_58])])).
cnf(c_0_331, plain, (v5_pre_topc(X1,X2,k3_pcomps_1(X3))|m1_subset_1(esk7_3(X2,X3,X1),k1_zfmisc_1(u1_struct_0(X3)))|v2_struct_0(X3)|v2_struct_0(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X3)))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X3)))|~v1_funct_1(X1)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_329, c_0_311]), c_0_312])).
cnf(c_0_332, plain, (v2_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_258])).
cnf(c_0_333, negated_conjecture, (v2_struct_0(k3_pcomps_1(esk2_0))|~v1_xboole_0(u1_struct_0(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_87]), c_0_56])])])).
cnf(c_0_334, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|~v2_pre_topc(k3_pcomps_1(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_330, c_0_331]), c_0_309]), c_0_108]), c_0_87]), c_0_144]), c_0_108]), c_0_87]), c_0_130]), c_0_178]), c_0_56]), c_0_138]), c_0_57]), c_0_139]), c_0_140])]), c_0_323]), c_0_58]), c_0_324])])).
cnf(c_0_335, plain, (v2_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[c_0_332])).
cnf(c_0_336, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_333]), c_0_56])]), c_0_58])])).
cnf(c_0_337, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_334, c_0_335]), c_0_81])])])).
cnf(c_0_338, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_336, c_0_337]), c_0_170])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 339
# Proof object clause steps            : 261
# Proof object formula steps           : 78
# Proof object conjectures             : 107
# Proof object clause conjectures      : 104
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 74
# Proof object initial formulas used   : 39
# Proof object generating inferences   : 87
# Proof object simplifying inferences  : 184
# Training examples: 163 positive, 1073 negative
# Training: Positive examples begin
cnf(c_0_239, negated_conjecture, (v6_metric_1(esk1_0))).# trainpos
cnf(c_0_140, negated_conjecture, (v6_metric_1(esk2_0))).# trainpos
cnf(c_0_238, negated_conjecture, (v7_metric_1(esk1_0))).# trainpos
cnf(c_0_139, negated_conjecture, (v7_metric_1(esk2_0))).# trainpos
cnf(c_0_82, negated_conjecture, (v8_metric_1(esk1_0))).# trainpos
cnf(c_0_57, negated_conjecture, (v8_metric_1(esk2_0))).# trainpos
cnf(c_0_237, negated_conjecture, (v9_metric_1(esk1_0))).# trainpos
cnf(c_0_138, negated_conjecture, (v9_metric_1(esk2_0))).# trainpos
cnf(c_0_81, negated_conjecture, (l1_metric_1(esk1_0))).# trainpos
cnf(c_0_56, negated_conjecture, (l1_metric_1(esk2_0))).# trainpos
cnf(c_0_178, negated_conjecture, (v1_funct_1(esk3_0))).# trainpos
cnf(c_0_170, plain, (v1_xboole_0(k1_xboole_0))).# trainpos
cnf(c_0_231, negated_conjecture, (~v2_struct_0(esk1_0))).# trainpos
cnf(c_0_58, negated_conjecture, (~v2_struct_0(esk2_0))).# trainpos
cnf(c_0_91, plain, (r1_tarski(X1,X1))).# trainpos
cnf(c_0_131, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_216, plain, (l1_struct_0(X1)|~l1_metric_1(X1))).# trainpos
cnf(c_0_191, plain, (l1_struct_0(X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_335, plain, (v2_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1))).# trainpos
cnf(c_0_206, plain, (l1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1))).# trainpos
cnf(c_0_269, plain, (v1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1))).# trainpos
cnf(c_0_197, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_323, negated_conjecture, (~v5_pre_topc(esk3_0,k3_pcomps_1(esk1_0),k3_pcomps_1(esk2_0)))).# trainpos
cnf(c_0_70, plain, (u1_struct_0(X1)=u1_struct_0(k3_pcomps_1(X1))|~l1_metric_1(X1))).# trainpos
cnf(c_0_262, plain, (u1_pre_topc(k3_pcomps_1(X1))=k2_pcomps_1(X1)|~l1_metric_1(X1))).# trainpos
cnf(c_0_90, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_113, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_86, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(k3_pcomps_1(esk1_0)),u1_struct_0(k3_pcomps_1(esk2_0))))).# trainpos
cnf(c_0_215, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1))).# trainpos
cnf(c_0_229, plain, (v2_struct_0(X1)|~l1_metric_1(X1)|~v2_struct_0(k3_pcomps_1(X1)))).# trainpos
cnf(c_0_190, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).# trainpos
cnf(c_0_118, plain, (r1_tarski(X1,X2)|r2_hidden(esk14_2(X1,X2),X1))).# trainpos
cnf(c_0_242, plain, (r1_tarski(X1,X2)|~r2_hidden(esk14_2(X1,X2),X2))).# trainpos
cnf(c_0_145, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_213, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_270, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).# trainpos
cnf(c_0_285, plain, (l1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).# trainpos
cnf(c_0_165, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).# trainpos
cnf(c_0_143, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_104, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(esk1_0)),u1_struct_0(k3_pcomps_1(esk2_0))))))).# trainpos
cnf(c_0_261, plain, (X1=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X1)|~v1_pre_topc(X1))).# trainpos
cnf(c_0_142, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).# trainpos
cnf(c_0_164, plain, (k2_relset_1(X2,X1)=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).# trainpos
cnf(c_0_112, plain, (r2_hidden(X3,X2)|~r2_hidden(X3,X1)|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_95, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_110, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).# trainpos
cnf(c_0_141, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).# trainpos
cnf(c_0_163, plain, (m1_subset_1(k2_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).# trainpos
cnf(c_0_296, plain, (k8_relset_1(X2,X3,X1,X4)=k8_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_295, plain, (m1_subset_1(k8_relset_1(X2,X3,X1,X4),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_129, plain, (X3=k1_xboole_0|X2=k1_relset_1(X2,X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_293, plain, (r2_hidden(k1_funct_1(X1,X2),X3)|X4!=k8_relat_1(X1,X3)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X2,X4))).# trainpos
cnf(c_0_96, plain, (m1_subset_1(k9_metric_1(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_metric_1(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_182, plain, (r2_hidden(X3,X4)|X4!=k10_xtuple_0(X2)|X3!=k1_funct_1(X2,X1)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k9_xtuple_0(X2)))).# trainpos
cnf(c_0_75, plain, (k4_metric_1(X1,X2,X3)=k4_metric_1(X1,X3,X2)|~v8_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_47, plain, (k2_metric_1(X1,X2,X3)=k4_metric_1(X1,X2,X3)|~v8_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_176, plain, (r2_hidden(X1,X4)|X4!=k8_relat_1(X2,X3)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k9_xtuple_0(X2))|~r2_hidden(k1_funct_1(X2,X1),X3))).# trainpos
cnf(c_0_99, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|v1_xreal_0(esk4_3(X1,X2,X3))|X3!=k1_funct_1(esk3_0,X2)|~v1_xreal_0(X1)|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_251, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|X3!=k1_funct_1(esk3_0,X2)|~v1_xreal_0(X1)|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r1_xxreal_0(esk4_3(X1,X2,X3),k6_numbers))).# trainpos
cnf(c_0_65, plain, (~l1_metric_1(X2)|~v1_xreal_0(X1)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~r1_xxreal_0(X1,k2_metric_1(X2,X3,X4))|~r2_hidden(X4,k9_metric_1(X2,X3,X1)))).# trainpos
cnf(c_0_299, plain, (v3_pre_topc(X4,k3_pcomps_1(X2))|X4!=k9_metric_1(X2,X3,X1)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v1_xreal_0(X1)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)))))).# trainpos
cnf(c_0_46, plain, (v2_struct_0(X1)|r1_xxreal_0(X2,k2_metric_1(X1,X3,X4))|r2_hidden(X4,k9_metric_1(X1,X3,X2))|~l1_metric_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).# trainpos
cnf(c_0_306, plain, (v2_struct_0(X1)|v3_pre_topc(X2,k3_pcomps_1(X1))|r2_hidden(esk12_2(X1,X2),X2)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))).# trainpos
cnf(c_0_214, plain, (v2_struct_0(X1)|v3_pre_topc(X2,k3_pcomps_1(X1))|m1_subset_1(esk12_2(X1,X2),u1_struct_0(X1))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))).# trainpos
cnf(c_0_292, plain, (v2_struct_0(X1)|v1_xreal_0(esk11_3(X1,X2,X3))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~v3_pre_topc(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))).# trainpos
cnf(c_0_271, plain, (v2_struct_0(X1)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~v3_pre_topc(X2,k3_pcomps_1(X1))|~r1_xxreal_0(esk11_3(X1,X2,X3),k6_numbers)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))).# trainpos
cnf(c_0_245, plain, (v2_struct_0(X2)|r1_xxreal_0(X1,k6_numbers)|v3_pre_topc(X3,k3_pcomps_1(X2))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v1_xreal_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~r1_tarski(k9_metric_1(X2,esk12_2(X2,X3),X1),X3))).# trainpos
cnf(c_0_116, plain, (v2_struct_0(X1)|r1_tarski(k9_metric_1(X1,X2,esk11_3(X1,X3,X2)),X3)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~r2_hidden(X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_pre_topc(X3,k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))).# trainpos
cnf(c_0_54, negated_conjecture, (r1_xxreal_0(X3,k6_numbers)|r1_xxreal_0(esk4_3(X3,X4,X5),k4_metric_1(esk1_0,X4,X1))|X5!=k1_funct_1(esk3_0,X4)|X2!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(X3)|~m1_subset_1(X5,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_xxreal_0(X3,k4_metric_1(esk2_0,X5,X2)))).# trainpos
cnf(c_0_312, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v1_xreal_0(esk5_3(X1,X2,X3))|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v1_funct_1(X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2))))))).# trainpos
cnf(c_0_311, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X2))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v1_funct_1(X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2))))))).# trainpos
cnf(c_0_326, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|m1_subset_1(esk7_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v1_funct_1(X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2))))))).# trainpos
cnf(c_0_300, plain, (esk7_3(X1,X2,X3)=k9_metric_1(X2,esk6_3(X1,X2,X3),esk5_3(X1,X2,X3))|v2_struct_0(X2)|v2_struct_0(X1)|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v1_funct_1(X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2))))))).# trainpos
cnf(c_0_303, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v1_funct_1(X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))|~v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)),X3,esk7_3(X1,X2,X3)),X1))).# trainpos
cnf(c_0_87, negated_conjecture, (u1_struct_0(k3_pcomps_1(esk2_0))=u1_struct_0(esk2_0))).# trainpos
cnf(c_0_108, negated_conjecture, (u1_struct_0(k3_pcomps_1(esk1_0))=u1_struct_0(esk1_0))).# trainpos
cnf(c_0_55, plain, (r2_hidden(X1,k9_metric_1(X2,X3,X4))|r1_xxreal_0(X4,k4_metric_1(X2,X3,X1))|v2_struct_0(X2)|~v1_xreal_0(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)|~v8_metric_1(X2))).# trainpos
cnf(c_0_310, plain, (v3_pre_topc(X1,k3_pcomps_1(X2))|v5_pre_topc(X4,X3,k3_pcomps_1(X2))|v2_struct_0(X3)|v2_struct_0(X2)|X1!=esk7_3(X3,X2,X4)|~v1_xreal_0(esk5_3(X3,X2,X4))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(esk6_3(X3,X2,X4),u1_struct_0(X2))|~l1_metric_1(X2)|~v9_metric_1(X2)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X4)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).# trainpos
cnf(c_0_305, plain, (r2_hidden(k1_funct_1(X1,X2),X3)|~r2_hidden(X2,k8_relat_1(X1,X3))|~v1_relat_1(X1)|~v1_funct_1(X1))).# trainpos
cnf(c_0_79, plain, (~r2_hidden(X1,k9_metric_1(X2,X3,X4))|~r1_xxreal_0(X4,k4_metric_1(X2,X3,X1))|~v1_xreal_0(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)|~v8_metric_1(X2))).# trainpos
cnf(c_0_329, plain, (m1_subset_1(esk7_3(X2,X1,X3),k1_zfmisc_1(u1_struct_0(X1)))|v5_pre_topc(X3,X2,k3_pcomps_1(X1))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_xreal_0(esk5_3(X2,X1,X3))|~m1_subset_1(esk6_3(X2,X1,X3),u1_struct_0(X1))|~l1_metric_1(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X1)))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X3)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_225, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_metric_1(X1))).# trainpos
cnf(c_0_224, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(esk12_2(X1,X2),u1_struct_0(X1))|v3_pre_topc(X2,k3_pcomps_1(X1))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_205, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1))).# trainpos
cnf(c_0_133, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_111, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_160, plain, (m1_subset_1(k9_xtuple_0(X2),k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).# trainpos
cnf(c_0_308, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_115, plain, (m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,k9_metric_1(X2,X3,X4))|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2))).# trainpos
cnf(c_0_307, plain, (m1_subset_1(k8_relat_1(X3,X4),k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_268, plain, (g1_pre_topc(u1_struct_0(k3_pcomps_1(X1)),k2_pcomps_1(X1))=k3_pcomps_1(X1)|~v1_pre_topc(k3_pcomps_1(X1))|~l1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1))).# trainpos
cnf(c_0_180, plain, (m1_subset_1(k10_xtuple_0(X2),k1_zfmisc_1(X1))|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).# trainpos
cnf(c_0_136, plain, (r2_hidden(X1,X2)|v2_struct_0(X3)|~r2_hidden(X1,k9_metric_1(X3,X4,esk11_3(X3,X2,X4)))|~r2_hidden(X4,X2)|~v3_pre_topc(X2,k3_pcomps_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).# trainpos
cnf(c_0_277, plain, (m1_subset_1(k2_pcomps_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))|~l1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1))).# trainpos
cnf(c_0_313, negated_conjecture, (v5_pre_topc(X1,X2,k3_pcomps_1(esk2_0))|v2_struct_0(X2)|~v3_pre_topc(k8_relset_1(u1_struct_0(X2),u1_struct_0(esk2_0),X1,esk7_3(X2,esk2_0,X1)),X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X1))).# trainpos
cnf(c_0_107, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(k3_pcomps_1(esk1_0)),u1_struct_0(esk2_0)))).# trainpos
cnf(c_0_94, plain, (r2_hidden(X1,k9_metric_1(X2,X3,X4))|r1_xxreal_0(X4,k4_metric_1(X2,X1,X3))|v2_struct_0(X2)|~v1_xreal_0(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)|~v8_metric_1(X2))).# trainpos
cnf(c_0_212, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,X3))|~v4_relat_1(X3,X2)|~v1_relat_1(X3))).# trainpos
cnf(c_0_325, plain, (v3_pre_topc(esk7_3(X1,X2,X3),k3_pcomps_1(X2))|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))|~m1_subset_1(esk7_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).# trainpos
cnf(c_0_286, negated_conjecture, (g1_pre_topc(u1_struct_0(esk1_0),k2_pcomps_1(esk1_0))=k3_pcomps_1(esk1_0))).# trainpos
cnf(c_0_315, plain, (r2_hidden(k1_funct_1(X1,esk12_2(X2,k8_relat_1(X1,X3))),X3)|v3_pre_topc(k8_relat_1(X1,X3),k3_pcomps_1(X2))|v2_struct_0(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(k8_relat_1(X1,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).# trainpos
cnf(c_0_125, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(esk1_0)),u1_struct_0(esk2_0)))))).# trainpos
cnf(c_0_66, negated_conjecture, (r1_xxreal_0(esk4_3(X1,X2,X3),k4_metric_1(esk1_0,X2,X4))|r1_xxreal_0(X1,k6_numbers)|r2_hidden(X5,k9_metric_1(esk2_0,X3,X1))|X3!=k1_funct_1(esk3_0,X2)|X5!=k1_funct_1(esk3_0,X4)|~v1_xreal_0(X1)|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X5,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_218, plain, (v2_struct_0(k3_pcomps_1(X1))|~v1_xboole_0(u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1))).# trainpos
cnf(c_0_298, negated_conjecture, (m1_subset_1(k2_pcomps_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).# trainpos
cnf(c_0_172, plain, (m1_subset_1(esk14_2(k9_metric_1(X1,X2,X3),X4),u1_struct_0(X1))|r1_tarski(k9_metric_1(X1,X2,X3),X4)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1))).# trainpos
cnf(c_0_132, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).# trainpos
cnf(c_0_153, plain, (r2_hidden(X1,u1_struct_0(X2))|~r2_hidden(X1,k9_metric_1(X2,X3,X4))|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2))).# trainpos
cnf(c_0_162, negated_conjecture, (v1_relat_1(esk3_0))).# trainpos
cnf(c_0_161, negated_conjecture, (v4_relat_1(esk3_0,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_181, negated_conjecture, (v5_relat_1(esk3_0,u1_struct_0(esk2_0)))).# trainpos
cnf(c_0_321, negated_conjecture, (v5_pre_topc(X1,X2,k3_pcomps_1(esk2_0))|v2_struct_0(X2)|~v3_pre_topc(k8_relat_1(X1,esk7_3(X2,esk2_0,X1)),X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X1))).# trainpos
cnf(c_0_80, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,X1),k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(esk4_3(X3,X4,X2),k4_metric_1(esk1_0,X4,X1))|r1_xxreal_0(X3,k6_numbers)|X2!=k1_funct_1(esk3_0,X4)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,X1),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_114, plain, (r2_hidden(X3,k9_metric_1(X2,X1,X4))|v2_struct_0(X2)|~r2_hidden(X1,k9_metric_1(X2,X3,X4))|~v1_xreal_0(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)|~v8_metric_1(X2))).# trainpos
cnf(c_0_316, negated_conjecture, (m1_subset_1(k8_relat_1(esk3_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_150, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)=u1_struct_0(esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))))).# trainpos
cnf(c_0_333, negated_conjecture, (v2_struct_0(k3_pcomps_1(esk2_0))|~v1_xboole_0(u1_struct_0(esk2_0)))).# trainpos
cnf(c_0_223, plain, (r2_hidden(X1,X2)|~v4_relat_1(X3,X2)|~r2_hidden(X1,k9_xtuple_0(X3))|~v1_relat_1(X3))).# trainpos
cnf(c_0_230, negated_conjecture, (v2_struct_0(k3_pcomps_1(esk1_0))|~v1_xboole_0(u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_297, negated_conjecture, (l1_pre_topc(k3_pcomps_1(esk1_0))|~m1_subset_1(k2_pcomps_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).# trainpos
cnf(c_0_151, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X2))).# trainpos
cnf(c_0_336, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk2_0)))).# trainpos
cnf(c_0_240, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_154, plain, (r2_hidden(X1,k9_metric_1(X2,esk14_2(k9_metric_1(X2,X1,X3),X4),X3))|v2_struct_0(X2)|r1_tarski(k9_metric_1(X2,X1,X3),X4)|~v1_xreal_0(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_metric_1(X2)|~v8_metric_1(X2))).# trainpos
cnf(c_0_98, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,X1),k9_metric_1(esk2_0,X4,X3))|r1_xxreal_0(X3,k6_numbers)|~r2_hidden(X1,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)))|~v1_xreal_0(esk4_3(X3,X2,X4))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|X4!=k1_funct_1(esk3_0,X2)|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,X1),u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk2_0)))).# trainpos
cnf(c_0_183, negated_conjecture, (k9_xtuple_0(esk3_0)=u1_struct_0(esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0)).# trainpos
cnf(c_0_195, negated_conjecture, (m1_subset_1(k10_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_179, negated_conjecture, (m1_subset_1(k9_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_169, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).# trainpos
cnf(c_0_318, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~v2_struct_0(k3_pcomps_1(esk1_0)))).# trainpos
cnf(c_0_208, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,k10_xtuple_0(esk3_0)))).# trainpos
cnf(c_0_196, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k1_funct_1(esk3_0,X3)|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(X3,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_194, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k9_xtuple_0(esk3_0)))).# trainpos
cnf(c_0_137, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)),k9_metric_1(esk2_0,X3,X2))|r1_xxreal_0(X2,k6_numbers)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)|X3!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(X2)|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk2_0)))).# trainpos
cnf(c_0_235, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|~v4_relat_1(esk3_0,X2)|~r2_hidden(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_185, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).# trainpos
cnf(c_0_210, plain, (r1_tarski(k9_metric_1(X1,X2,X3),X4)|v2_struct_0(X1)|~v1_xboole_0(k9_metric_1(X1,esk14_2(k9_metric_1(X1,X2,X3),X4),X3))|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v8_metric_1(X1))).# trainpos
cnf(c_0_171, plain, (r2_hidden(X1,u1_struct_0(X2))|r1_tarski(k9_metric_1(X2,X1,X3),X4)|v2_struct_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(esk14_2(k9_metric_1(X2,X1,X3),X4),u1_struct_0(X2))|~l1_metric_1(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~v8_metric_1(X2))).# trainpos
cnf(c_0_199, plain, (m1_subset_1(k9_metric_1(X1,X2,X3),k1_zfmisc_1(X4))|r2_hidden(X2,u1_struct_0(X1))|v2_struct_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v8_metric_1(X1))).# trainpos
cnf(c_0_209, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(esk3_0,X1),X2)|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_158, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)),X2)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X3,u1_struct_0(esk2_0))|X3!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_244, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk12_2(esk1_0,X1),X2)|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|~v4_relat_1(esk3_0,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_198, plain, (X1=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).# trainpos
cnf(c_0_221, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,X1),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_211, plain, (k9_metric_1(X1,X2,X3)=k1_xboole_0|r2_hidden(X2,u1_struct_0(X1))|v2_struct_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v8_metric_1(X1))).# trainpos
cnf(c_0_193, negated_conjecture, (r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),X5)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|X5!=k8_relat_1(esk3_0,X2)|~r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),k9_xtuple_0(esk3_0))|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_220, negated_conjecture, (r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)|r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),X5)|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|X5!=k8_relat_1(esk3_0,X2)|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)),u1_struct_0(esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_222, plain, (r1_tarski(k9_metric_1(X1,X2,X3),X4)|v2_struct_0(X1)|r2_hidden(esk14_2(k9_metric_1(X1,X2,X3),X4),u1_struct_0(X1))|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(esk14_2(k9_metric_1(X1,X2,X3),X4),u1_struct_0(X1)))).# trainpos
cnf(c_0_243, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)|r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),X5)|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|X5!=k8_relat_1(esk3_0,X2)|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v1_xreal_0(esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)))).# trainpos
cnf(c_0_247, negated_conjecture, (m1_subset_1(esk12_2(esk1_0,X1),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|~v4_relat_1(esk3_0,k9_xtuple_0(esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_246, negated_conjecture, (r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)|u1_struct_0(esk2_0)=k1_xboole_0|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|X4!=k8_relat_1(esk3_0,X2)|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk4_3(esk11_3(esk2_0,X2,X3),X1,X3))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_250, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|m1_subset_1(esk12_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_249, negated_conjecture, (v3_pre_topc(X1,k3_pcomps_1(esk1_0))|r1_xxreal_0(esk4_3(esk11_3(esk2_0,X2,X3),esk12_2(esk1_0,X1),X3),k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|~v1_xreal_0(esk4_3(esk11_3(esk2_0,X2,X3),esk12_2(esk1_0,X1),X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|X1!=k8_relat_1(esk3_0,X2)|X3!=k1_funct_1(esk3_0,esk12_2(esk1_0,X1))|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(esk12_2(esk1_0,X1),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_253, negated_conjecture, (r1_xxreal_0(esk11_3(esk2_0,X1,X2),k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X3,k3_pcomps_1(esk1_0))|X2!=k1_funct_1(esk3_0,esk12_2(esk1_0,X3))|~v1_xreal_0(esk11_3(esk2_0,X1,X2))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(esk12_2(esk1_0,X3),u1_struct_0(esk1_0))|X3!=k8_relat_1(esk3_0,X1)|~r2_hidden(X2,X1)|~v3_pre_topc(X1,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk4_3(esk11_3(esk2_0,X1,X2),esk12_2(esk1_0,X3),X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_260, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|X3!=k1_funct_1(esk3_0,esk12_2(esk1_0,X1))|X1!=k8_relat_1(esk3_0,X2)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(esk12_2(esk1_0,X1),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_272, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|r1_xxreal_0(esk11_3(esk2_0,X2,k1_funct_1(esk3_0,esk12_2(esk1_0,X1))),k6_numbers)|X1!=k8_relat_1(esk3_0,X2)|~r2_hidden(k1_funct_1(esk3_0,esk12_2(esk1_0,X1)),X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,k1_funct_1(esk3_0,esk12_2(esk1_0,X1))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_280, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|~r2_hidden(k1_funct_1(esk3_0,esk12_2(esk1_0,X1)),X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(k1_funct_1(esk3_0,esk12_2(esk1_0,X1)),u1_struct_0(esk2_0))|X1!=k8_relat_1(esk3_0,X2)|~v1_xreal_0(esk11_3(esk2_0,X2,k1_funct_1(esk3_0,esk12_2(esk1_0,X1))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_304, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|X1!=k8_relat_1(esk3_0,X2)|~r2_hidden(k1_funct_1(esk3_0,esk12_2(esk1_0,X1)),X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_funct_1(esk3_0,esk12_2(esk1_0,X1)),u1_struct_0(esk2_0)))).# trainpos
cnf(c_0_322, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X1),k3_pcomps_1(esk1_0))|~v3_pre_topc(X1,k3_pcomps_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_327, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|~v2_pre_topc(k3_pcomps_1(esk1_0))|~v3_pre_topc(esk7_3(k3_pcomps_1(esk1_0),esk2_0,esk3_0),k3_pcomps_1(esk2_0))|~m1_subset_1(esk7_3(k3_pcomps_1(esk1_0),esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_330, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|~v2_pre_topc(k3_pcomps_1(esk1_0))|~m1_subset_1(esk7_3(k3_pcomps_1(esk1_0),esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_334, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|~v2_pre_topc(k3_pcomps_1(esk1_0)))).# trainpos
cnf(c_0_337, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0)).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_109, plain, (v1_xreal_0(esk15_0))).#trainneg
cnf(i_0_112, plain, (v1_xboole_0(esk15_0))).#trainneg
cnf(i_0_111, plain, (v1_xcmplx_0(esk15_0))).#trainneg
cnf(i_0_110, plain, (v1_xxreal_0(esk15_0))).#trainneg
cnf(i_0_124, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_129, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_125, plain, (~v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_126, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_127, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_108, plain, (~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_128, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_54, plain, (r2_hidden(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_53, plain, (v1_xreal_0(X1)|~r2_hidden(X1,k1_numbers))).#trainneg
cnf(i_0_78, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_113, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_86, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_123, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_67, plain, (m1_subset_1(k2_pcomps_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_metric_1(X1))).#trainneg
cnf(i_0_92, plain, (v1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_132, plain, (v4_relat_1(X1,k5_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_107, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_102, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_98, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_49, plain, (X1=X2|g1_pre_topc(X3,X1)!=g1_pre_topc(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X3))))).#trainneg
cnf(i_0_50, plain, (X1=X2|g1_pre_topc(X1,X3)!=g1_pre_topc(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_28, plain, (v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_100, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_29, plain, (X1=k1_xboole_0|X3!=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_116, plain, (v1_partfun1(X1,X2)|~v1_funct_2(X1,X2,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2))))).#trainneg
cnf(i_0_30, plain, (X3=k1_xboole_0|v1_funct_2(X2,X1,X3)|X1!=k1_relset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_60, plain, (r2_hidden(X1,k9_xtuple_0(X2))|X3!=k8_relat_1(X2,X4)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_99, plain, (v1_xboole_0(X1)|v1_partfun1(X2,X3)|~v1_funct_2(X2,X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_84, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_85, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_115, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_36, plain, (X1=k1_funct_1(X2,esk8_3(X2,X3,X1))|X3!=k10_xtuple_0(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_114, plain, (m1_subset_1(k2_metric_1(X1,X2,X3),k1_numbers)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_117, plain, (v2_struct_0(X1)|v11_metric_1(X1)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1))).#trainneg
cnf(i_0_37, plain, (r2_hidden(esk8_3(X1,X2,X3),k9_xtuple_0(X1))|X2!=k10_xtuple_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_33, plain, (X2=k10_xtuple_0(X1)|r2_hidden(esk9_2(X1,X2),X2)|r2_hidden(esk10_2(X1,X2),k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_32, plain, (X2=k10_xtuple_0(X1)|esk9_2(X1,X2)=k1_funct_1(X1,esk10_2(X1,X2))|r2_hidden(esk9_2(X1,X2),X2)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_103, plain, (m1_subset_1(k4_metric_1(X1,X2,X3),k1_numbers)|~v8_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_96, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|~v1_funct_1(X1)|~v1_xboole_0(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_56, plain, (X3=k8_relat_1(X1,X2)|r2_hidden(esk13_3(X1,X2,X3),X3)|r2_hidden(esk13_3(X1,X2,X3),k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_34, plain, (X2=k10_xtuple_0(X1)|esk9_2(X1,X2)!=k1_funct_1(X1,X3)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X3,k9_xtuple_0(X1))|~r2_hidden(esk9_2(X1,X2),X2))).#trainneg
cnf(i_0_40, plain, (~v2_struct_0(X1)|~l1_metric_1(X1)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X2,k9_metric_1(X1,X3,X4)))).#trainneg
cnf(i_0_55, plain, (X3=k8_relat_1(X1,X2)|r2_hidden(esk13_3(X1,X2,X3),X3)|r2_hidden(k1_funct_1(X1,esk13_3(X1,X2,X3)),X2)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_57, plain, (X3=k8_relat_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(esk13_3(X1,X2,X3),X3)|~r2_hidden(esk13_3(X1,X2,X3),k9_xtuple_0(X1))|~r2_hidden(k1_funct_1(X1,esk13_3(X1,X2,X3)),X2))).#trainneg
cnf(i_0_21, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v1_funct_1(X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~r1_xxreal_0(esk5_3(X1,X2,X3),k6_numbers)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2))))))).#trainneg
cnf(i_0_266, negated_conjecture, (r1_xxreal_0(esk4_3(X1,X2,X3),k4_metric_1(esk1_0,X2,X4))|r1_xxreal_0(X1,k6_numbers)|X3!=k1_funct_1(esk3_0,X2)|X5!=k1_funct_1(esk3_0,X4)|~r1_xxreal_0(X1,k4_metric_1(esk2_0,X5,X3))|~v1_xreal_0(X1)|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X5,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_278, plain, (v5_pre_topc(X1,X2,k3_pcomps_1(X3))|v2_struct_0(X2)|v2_struct_0(X3)|~v3_pre_topc(k8_relat_1(X1,esk7_3(X2,X3,X1)),X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X3)))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X3)))|~v1_funct_1(X1)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_143, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_262, plain, (v3_pre_topc(X1,k3_pcomps_1(X2))|r1_xxreal_0(X3,k6_numbers)|v2_struct_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~m1_subset_1(k9_metric_1(X2,esk12_2(X2,X1),X3),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_270, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(esk6_3(X2,X1,X3),u1_struct_0(X1))|v5_pre_topc(X3,X2,k3_pcomps_1(X1))|v2_struct_0(X2)|v2_struct_0(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X1)))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X3)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_186, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,X3))|~v4_relat_1(X3,X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_134, plain, (esk15_0=k1_xboole_0)).#trainneg
cnf(i_0_258, plain, (v3_pre_topc(X1,k3_pcomps_1(X2))|v2_struct_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_257, plain, (m1_subset_1(esk12_2(X1,X2),X2)|v3_pre_topc(X2,k3_pcomps_1(X1))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_252, plain, (v3_pre_topc(k9_metric_1(X1,X2,X3),k3_pcomps_1(X1))|~v1_xreal_0(X3)|~m1_subset_1(k9_metric_1(X1,X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1))).#trainneg
cnf(i_0_222, plain, (r2_hidden(X1,k9_xtuple_0(X2))|~r2_hidden(X1,k8_relat_1(X2,X3))|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_274, plain, (m1_subset_1(X1,u1_struct_0(k3_pcomps_1(X2)))|v5_pre_topc(X4,X3,k3_pcomps_1(X2))|v2_struct_0(X3)|v2_struct_0(X2)|~r2_hidden(X1,esk7_3(X3,X2,X4))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X4)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_255, plain, (r2_hidden(X1,X2)|v3_pre_topc(k9_xtuple_0(X3),k3_pcomps_1(X4))|v2_struct_0(X4)|X1!=k1_funct_1(X3,esk12_2(X4,k9_xtuple_0(X3)))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~m1_subset_1(k9_xtuple_0(X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~l1_metric_1(X4)|~v9_metric_1(X4)|~v8_metric_1(X4)|~v7_metric_1(X4)|~v6_metric_1(X4))).#trainneg
cnf(i_0_273, plain, (v1_xboole_0(esk7_3(X1,X2,X3))|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v1_xboole_0(u1_struct_0(k3_pcomps_1(X2)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_229, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(X3))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_148, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_197, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k2_relset_1(X2,X3))|~v5_relat_1(X3,X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_239, plain, (r2_hidden(X1,X2)|r1_tarski(k9_xtuple_0(X3),X4)|X1!=k1_funct_1(X3,esk14_2(k9_xtuple_0(X3),X4))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3))).#trainneg
cnf(i_0_238, plain, (~v1_xboole_0(k9_xtuple_0(X1))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_240, plain, (r2_hidden(X1,X2)|X1!=k1_funct_1(X3,esk8_3(X3,X4,X5))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|X4!=k10_xtuple_0(X3)|~r2_hidden(X5,X4))).#trainneg
cnf(i_0_217, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k8_relset_1(X2,X3,X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_250, negated_conjecture, (r2_hidden(esk4_3(X1,X2,X3),k1_numbers)|r1_xxreal_0(X1,k6_numbers)|X3!=k1_funct_1(esk3_0,X2)|~v1_xreal_0(X1)|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_147, plain, (m1_subset_1(esk14_2(X1,X2),X1)|r1_tarski(X1,X2))).#trainneg
cnf(i_0_233, plain, (v1_xboole_0(k9_metric_1(X1,X2,X3))|~v1_xboole_0(u1_struct_0(X1))|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_248, plain, (r2_hidden(esk8_3(X1,X2,X3),X4)|X4!=k8_relat_1(X1,X5)|~r2_hidden(X3,X5)|~r2_hidden(esk8_3(X1,X2,X3),k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_203, plain, (X1=u1_struct_0(X2)|g1_pre_topc(X1,X3)!=X2|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v1_pre_topc(X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_204, plain, (u1_struct_0(X1)=X2|X1!=g1_pre_topc(X2,X3)|~m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_216, plain, (v1_xboole_0(k8_relset_1(X1,X2,X3,X4))|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_272, plain, (r2_hidden(esk7_3(X2,X1,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|v5_pre_topc(X3,X2,k3_pcomps_1(X1))|v2_struct_0(X2)|v2_struct_0(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X1)))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X3)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_263, plain, (m1_subset_1(k9_metric_1(X1,X2,esk11_3(X1,X3,X2)),k1_zfmisc_1(X3))|v2_struct_0(X1)|~r2_hidden(X2,X3)|~v3_pre_topc(X3,k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_230, plain, (v1_xreal_0(X3)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_178, plain, (r2_hidden(k1_relset_1(X1,X2),k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_268, plain, (r2_hidden(esk5_3(X1,X2,X3),k1_numbers)|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_213, plain, (v1_xboole_0(k8_relset_1(k2_zfmisc_1(X1,X2),X3,X4,X5))|~v1_xboole_0(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_209, plain, (r2_hidden(k8_relset_1(X1,X2,X3,X4),k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_219, plain, (X1=k1_xboole_0|v1_funct_2(X2,X3,X1)|k9_xtuple_0(X2)!=X3|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~v4_relat_1(X2,X3)|~v1_relat_1(X2))).#trainneg
cnf(i_0_185, plain, (v1_xboole_0(k1_relset_1(X1,X2))|~v1_xboole_0(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_249, plain, (r1_tarski(k9_metric_1(X1,X2,X3),X4)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(esk14_2(k9_metric_1(X1,X2,X3),X4),u1_struct_0(X1))|~l1_metric_1(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_207, plain, (v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|X2!=k1_xboole_0|X1!=k1_xboole_0)).#trainneg
cnf(i_0_232, plain, (r2_hidden(k9_metric_1(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_182, plain, (v1_xboole_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_xboole_0(X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_189, plain, (r2_hidden(k2_relset_1(X1,X2),k1_zfmisc_1(X1))|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_172, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,u1_pre_topc(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_261, plain, (r2_hidden(esk11_3(X1,X2,X3),k1_numbers)|v2_struct_0(X1)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_201, plain, (u1_pre_topc(X1)=X2|X1!=g1_pre_topc(X3,X2)|~m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_256, plain, (v3_pre_topc(k9_metric_1(X2,X3,X4),k3_pcomps_1(X1))|v2_struct_0(X1)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(esk12_2(X1,k9_metric_1(X2,X3,X4)),u1_struct_0(X2))|~l1_metric_1(X2)|~v2_struct_0(X2)|~m1_subset_1(k9_metric_1(X2,X3,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_237, plain, (m1_subset_1(esk8_3(X1,X2,X3),k9_xtuple_0(X1))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_211, plain, (v1_xboole_0(k8_relset_1(k2_zfmisc_1(X1,X2),X3,X4,X5))|~v1_xboole_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_153, plain, (r2_hidden(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_176, plain, (r2_hidden(k1_funct_1(X1,X2),k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_210, plain, (v1_relat_1(k8_relset_1(k2_zfmisc_1(X1,X2),X3,X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_171, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,k2_pcomps_1(X2))|~l1_metric_1(X2))).#trainneg
cnf(i_0_180, plain, (v1_xboole_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_xboole_0(X1)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_196, plain, (v1_xboole_0(k2_relset_1(X1,X2))|~v1_xboole_0(X1)|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_242, plain, (v1_xboole_0(k1_numbers)|r2_hidden(k4_metric_1(X1,X2,X3),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v8_metric_1(X1))).#trainneg
cnf(i_0_193, plain, (v1_xboole_0(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_xboole_0(X2)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_157, plain, (r2_hidden(k2_pcomps_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_metric_1(X1))).#trainneg
cnf(i_0_236, plain, (v1_xboole_0(k1_numbers)|r2_hidden(k2_metric_1(X1,X2,X3),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_215, plain, (v5_relat_1(k8_relset_1(k2_zfmisc_1(X1,X2),X3,X4,X5),X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_179, plain, (v1_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_191, plain, (v1_xboole_0(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_xboole_0(X1)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_224, plain, (v1_partfun1(X1,X2)|v1_xboole_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|X3!=k1_xboole_0|X1!=k1_xboole_0)).#trainneg
cnf(i_0_190, plain, (v1_relat_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_192, plain, (v4_relat_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_184, plain, (v5_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_151, plain, (v1_xboole_0(k1_numbers)|r2_hidden(np__1,k1_numbers))).#trainneg
cnf(i_0_195, plain, (v5_relat_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X2)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_181, plain, (v4_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_212, plain, (v4_relat_1(k8_relset_1(k2_zfmisc_1(X1,X2),X3,X4,X5),X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_183, plain, (v3_valued_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v3_membered(X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_150, plain, (v1_xboole_0(k5_numbers)|r2_hidden(np__1,k5_numbers))).#trainneg
cnf(i_0_194, plain, (v3_valued_0(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v3_membered(X2)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_241, plain, (v3_ordinal1(k4_metric_1(X1,X2,X3))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v8_metric_1(X1))).#trainneg
cnf(i_0_235, plain, (v3_ordinal1(k2_metric_1(X1,X2,X3))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_259, plain, (v3_ordinal1(esk12_2(X1,X2))|v3_pre_topc(X2,k3_pcomps_1(X1))|v2_struct_0(X1)|~v3_ordinal1(u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_269, plain, (v3_ordinal1(esk6_3(X1,X2,X3))|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v3_ordinal1(u1_struct_0(X2))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_145, plain, (v3_ordinal1(np__1)|~v3_ordinal1(k5_numbers))).#trainneg
cnf(i_0_146, plain, (v3_ordinal1(np__1)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_188, plain, (v3_ordinal1(k2_relset_1(X1,X2))|~v3_ordinal1(k1_zfmisc_1(X1))|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_214, plain, (v3_valued_0(k8_relset_1(k2_zfmisc_1(X1,X2),X3,X4,X5))|~v3_membered(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_208, plain, (v3_ordinal1(k8_relset_1(X1,X2,X3,X4))|~v3_ordinal1(k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_231, plain, (v3_ordinal1(k9_metric_1(X1,X2,X3))|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(X1)))|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_177, plain, (v3_ordinal1(k1_relset_1(X1,X2))|~v3_ordinal1(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_156, plain, (v3_ordinal1(k2_pcomps_1(X1))|~v3_ordinal1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_metric_1(X1))).#trainneg
cnf(i_0_152, plain, (v3_ordinal1(u1_pre_topc(X1))|~v3_ordinal1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_271, plain, (v3_ordinal1(esk7_3(X1,X2,X3))|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_296, negated_conjecture, (v5_pre_topc(X1,X2,k3_pcomps_1(esk1_0))|v2_struct_0(X2)|~v3_pre_topc(k8_relset_1(u1_struct_0(X2),u1_struct_0(esk1_0),X1,esk7_3(X2,esk1_0,X1)),X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_348, plain, (r2_hidden(esk14_2(k8_relat_1(X1,X2),X3),k9_xtuple_0(X1))|r1_tarski(k8_relat_1(X1,X2),X3)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_301, plain, (r2_hidden(k1_funct_1(X1,esk14_2(k8_relat_1(X1,X2),X3)),X2)|r1_tarski(k8_relat_1(X1,X2),X3)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_329, plain, (m1_subset_1(esk12_2(X1,X2),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v3_pre_topc(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_349, plain, (r2_hidden(esk12_2(X1,k8_relat_1(X2,X3)),k9_xtuple_0(X2))|v3_pre_topc(k8_relat_1(X2,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~m1_subset_1(k8_relat_1(X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_384, plain, (v3_pre_topc(X1,k3_pcomps_1(X2))|r1_xxreal_0(X3,k6_numbers)|v2_struct_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v1_xboole_0(k9_metric_1(X2,esk12_2(X2,X1),X3)))).#trainneg
cnf(i_0_303, plain, (~r2_hidden(X1,k9_metric_1(X2,X3,X4))|~r1_xxreal_0(X4,k4_metric_1(X2,X1,X3))|~v1_xreal_0(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)|~v8_metric_1(X2))).#trainneg
cnf(i_0_308, negated_conjecture, (r1_xxreal_0(esk4_3(X1,X2,X3),k4_metric_1(esk1_0,X2,X4))|r1_xxreal_0(X1,k6_numbers)|r2_hidden(X3,k9_metric_1(esk2_0,X5,X1))|X3!=k1_funct_1(esk3_0,X2)|X5!=k1_funct_1(esk3_0,X4)|~v1_xreal_0(X1)|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X5,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_350, plain, (v5_pre_topc(X1,X2,k3_pcomps_1(X3))|m1_subset_1(esk14_2(esk7_3(X2,X3,X1),X4),u1_struct_0(k3_pcomps_1(X3)))|v2_struct_0(X3)|v2_struct_0(X2)|r1_tarski(esk7_3(X2,X3,X1),X4)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X3)))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X3)))|~v1_funct_1(X1)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_428, negated_conjecture, (g1_pre_topc(u1_struct_0(esk2_0),k2_pcomps_1(esk2_0))=k3_pcomps_1(esk2_0))).#trainneg
cnf(i_0_295, negated_conjecture, (v5_pre_topc(X1,X2,k3_pcomps_1(esk1_0))|m1_subset_1(esk7_3(X2,esk1_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|v2_struct_0(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_355, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k2_relset_1(X2,X3))|~v5_relat_1(X3,X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_335, plain, (v5_pre_topc(X1,k3_pcomps_1(X2),k3_pcomps_1(X3))|v2_struct_0(X3)|v2_struct_0(k3_pcomps_1(X2))|v2_struct_0(X2)|~l1_pre_topc(k3_pcomps_1(X2))|~v2_pre_topc(k3_pcomps_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X3)))))|~v1_funct_2(X1,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X3)))|~v1_funct_1(X1)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)|~v1_xboole_0(k8_relat_1(X1,esk7_3(k3_pcomps_1(X2),X3,X1)))|~m1_subset_1(k8_relat_1(X1,esk7_3(k3_pcomps_1(X2),X3,X1)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_286, negated_conjecture, (v5_pre_topc(X1,X2,k3_pcomps_1(esk2_0))|m1_subset_1(esk7_3(X2,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|v2_struct_0(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_336, plain, (v5_pre_topc(X1,k3_pcomps_1(X2),k3_pcomps_1(X3))|v2_struct_0(k3_pcomps_1(X2))|v2_struct_0(X3)|v2_struct_0(X2)|~l1_pre_topc(k3_pcomps_1(X2))|~v2_pre_topc(k3_pcomps_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X3)))))|~v1_funct_2(X1,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X3)))|~v1_funct_1(X1)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)|~v1_xboole_0(k8_relset_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X3)),X1,esk7_3(k3_pcomps_1(X2),X3,X1)))|~m1_subset_1(k8_relset_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X3)),X1,esk7_3(k3_pcomps_1(X2),X3,X1)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_322, plain, (m1_subset_1(X1,X2)|~v4_relat_1(X3,X2)|~r2_hidden(X1,k9_xtuple_0(X3))|~v1_relat_1(X3))).#trainneg
cnf(i_0_365, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_387, plain, (m1_subset_1(X1,X2)|~v5_relat_1(X3,X2)|~r2_hidden(X1,k10_xtuple_0(X3))|~v1_relat_1(X3))).#trainneg
cnf(i_0_109, plain, (v1_xreal_0(k1_xboole_0))).#trainneg
cnf(i_0_401, plain, (m1_subset_1(esk12_2(X1,k8_relset_1(X2,X3,X4,X5)),X2)|v3_pre_topc(k8_relset_1(X2,X3,X4,X5),k3_pcomps_1(X1))|v2_struct_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k8_relset_1(X2,X3,X4,X5),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_324, plain, (m1_subset_1(esk12_2(X1,k1_relset_1(X2,X3)),X2)|v3_pre_topc(k1_relset_1(X2,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~m1_subset_1(k1_relset_1(X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_443, plain, (r2_hidden(esk14_2(k9_metric_1(X1,X2,esk11_3(X1,X3,X2)),X4),X3)|v2_struct_0(X1)|r1_tarski(k9_metric_1(X1,X2,esk11_3(X1,X3,X2)),X4)|~r2_hidden(X2,X3)|~v3_pre_topc(X3,k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_403, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v1_xboole_0(k1_numbers)|X3!=k1_funct_1(esk3_0,X2)|~v1_xreal_0(X1)|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_359, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k8_relset_1(X2,X3,X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_485, negated_conjecture, (m1_subset_1(k2_pcomps_1(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_523, plain, (u1_pre_topc(X1)=k2_pcomps_1(X2)|X1!=k3_pcomps_1(X2)|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~l1_metric_1(X2))).#trainneg
cnf(i_0_323, plain, (m1_subset_1(esk14_2(k1_relset_1(X1,X2),X3),X1)|r1_tarski(k1_relset_1(X1,X2),X3)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_395, plain, (m1_subset_1(esk12_2(X1,k9_metric_1(X2,X3,X4)),u1_struct_0(X2))|v3_pre_topc(k9_metric_1(X2,X3,X4),k3_pcomps_1(X1))|v2_struct_0(X1)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)|~m1_subset_1(k9_metric_1(X2,X3,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_437, plain, (u1_struct_0(g1_pre_topc(X1,X2))=X1|~v1_pre_topc(g1_pre_topc(X1,X2))|~l1_pre_topc(g1_pre_topc(X1,X2)))).#trainneg
cnf(i_0_400, plain, (m1_subset_1(esk14_2(k8_relset_1(X1,X2,X3,X4),X5),X1)|r1_tarski(k8_relset_1(X1,X2,X3,X4),X5)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_389, plain, (m1_subset_1(esk12_2(X1,k2_relset_1(X2,X3)),X2)|v3_pre_topc(k2_relset_1(X2,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|~v5_relat_1(X3,X2)|~v1_relat_1(X3)|~m1_subset_1(k2_relset_1(X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_280, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk2_0),k3_pcomps_1(X2))|v2_struct_0(k3_pcomps_1(esk2_0))|v2_struct_0(X2)|~v3_pre_topc(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(k3_pcomps_1(X2)),X1,esk7_3(k3_pcomps_1(esk2_0),X2,X1)),k3_pcomps_1(esk2_0))|~l1_pre_topc(k3_pcomps_1(esk2_0))|~v2_pre_topc(k3_pcomps_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_385, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_378, plain, (v1_xboole_0(k1_relset_1(X1,X2))|m1_subset_1(X3,X1)|~m1_subset_1(X3,k1_relset_1(X1,X2))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_354, plain, (r2_hidden(X1,u1_struct_0(X2))|v5_pre_topc(X4,X3,k3_pcomps_1(X2))|v2_struct_0(X2)|v2_struct_0(X3)|~r2_hidden(X1,esk7_3(X3,X2,X4))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X4)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_493, plain, (v1_xboole_0(k9_xtuple_0(X2))|~v4_relat_1(X2,X1)|~v1_xboole_0(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_474, plain, (m1_subset_1(esk5_3(X1,X2,X3),k1_numbers)|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_408, plain, (m1_subset_1(k8_relat_1(k2_zfmisc_1(X1,X2),X3),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_476, plain, (v5_pre_topc(X3,X1,k3_pcomps_1(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_xboole_0(k1_numbers)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_289, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk1_0),k3_pcomps_1(X2))|v2_struct_0(k3_pcomps_1(esk1_0))|v2_struct_0(X2)|~v3_pre_topc(k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(k3_pcomps_1(X2)),X1,esk7_3(k3_pcomps_1(esk1_0),X2,X1)),k3_pcomps_1(esk1_0))|~l1_pre_topc(k3_pcomps_1(esk1_0))|~v2_pre_topc(k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_444, plain, (r2_hidden(esk12_2(X1,k9_metric_1(X2,X3,esk11_3(X2,X4,X3))),X4)|v2_struct_0(X2)|v3_pre_topc(k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),k3_pcomps_1(X1))|v2_struct_0(X1)|~r2_hidden(X3,X4)|~v3_pre_topc(X4,k3_pcomps_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~m1_subset_1(k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_402, negated_conjecture, (m1_subset_1(esk4_3(X1,X2,X3),k1_numbers)|r1_xxreal_0(X1,k6_numbers)|X3!=k1_funct_1(esk3_0,X2)|~v1_xreal_0(X1)|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_522, plain, (u1_pre_topc(g1_pre_topc(X1,X2))=X2|~v1_pre_topc(g1_pre_topc(X1,X2))|~l1_pre_topc(g1_pre_topc(X1,X2)))).#trainneg
cnf(i_0_377, plain, (v1_xboole_0(k2_relset_1(X1,X2))|m1_subset_1(X3,X1)|~m1_subset_1(X3,k2_relset_1(X1,X2))|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_491, plain, (X1=k1_xboole_0|v1_funct_2(X2,k9_xtuple_0(X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k9_xtuple_0(X2),X1))))).#trainneg
cnf(i_0_432, plain, (r2_hidden(esk8_3(X1,X2,X3),k8_relat_1(X1,X4))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X4)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_550, plain, (v1_xboole_0(k10_xtuple_0(X2))|~v5_relat_1(X2,X1)|~v1_xboole_0(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_521, plain, (v2_struct_0(X1)|~v1_xboole_0(k1_numbers)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_367, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_317, plain, (v3_pre_topc(u1_struct_0(X1),k3_pcomps_1(X1))|r1_xxreal_0(X2,k6_numbers)|v2_struct_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(esk12_2(X1,u1_struct_0(X1)),u1_struct_0(X1)))).#trainneg
cnf(i_0_392, plain, (v1_xboole_0(u1_struct_0(k3_pcomps_1(X1)))|~v2_struct_0(k3_pcomps_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_370, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_390, plain, (r1_tarski(k9_xtuple_0(X1),X2)|r2_hidden(k1_funct_1(X1,esk14_2(k9_xtuple_0(X1),X2)),X3)|X3!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_441, plain, (k8_relset_1(X1,X2,X3,X4)=k1_xboole_0|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_546, plain, (m1_subset_1(esk12_2(X1,k2_pcomps_1(X2)),k1_zfmisc_1(u1_struct_0(X2)))|v3_pre_topc(k2_pcomps_1(X2),k3_pcomps_1(X1))|v2_struct_0(X1)|~l1_metric_1(X2)|~m1_subset_1(k2_pcomps_1(X2),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_545, plain, (m1_subset_1(esk14_2(k2_pcomps_1(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|r1_tarski(k2_pcomps_1(X1),X2)|~l1_metric_1(X1))).#trainneg
cnf(i_0_364, plain, (r2_hidden(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_399, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k8_relat_1(X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_388, plain, (m1_subset_1(esk14_2(k2_relset_1(X1,X2),X3),X1)|r1_tarski(k2_relset_1(X1,X2),X3)|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_430, plain, (k9_metric_1(X1,X2,X3)=k1_xboole_0|~v1_xboole_0(u1_struct_0(X1))|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_448, negated_conjecture, (r2_hidden(esk7_3(X1,esk1_0,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_312, plain, (m1_subset_1(X1,u1_struct_0(X2))|v5_pre_topc(X4,X3,k3_pcomps_1(X2))|v2_struct_0(X2)|v2_struct_0(X3)|~r2_hidden(X1,esk7_3(X3,X2,X4))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X4)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_290, negated_conjecture, (g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(k3_pcomps_1(esk1_0)))=k3_pcomps_1(esk1_0)|~v1_pre_topc(k3_pcomps_1(esk1_0))|~l1_pre_topc(k3_pcomps_1(esk1_0)))).#trainneg
cnf(i_0_351, plain, (v5_pre_topc(X1,X2,k3_pcomps_1(X3))|m1_subset_1(esk12_2(X4,esk7_3(X2,X3,X1)),u1_struct_0(k3_pcomps_1(X3)))|v2_struct_0(X3)|v2_struct_0(X2)|v3_pre_topc(esk7_3(X2,X3,X1),k3_pcomps_1(X4))|v2_struct_0(X4)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X3)))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X3)))|~v1_funct_1(X1)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)|~m1_subset_1(esk7_3(X2,X3,X1),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~l1_metric_1(X4)|~v9_metric_1(X4)|~v8_metric_1(X4)|~v7_metric_1(X4)|~v6_metric_1(X4))).#trainneg
cnf(i_0_447, negated_conjecture, (r2_hidden(esk7_3(X1,esk2_0,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|v5_pre_topc(X2,X1,k3_pcomps_1(esk2_0))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk2_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_407, plain, (m1_subset_1(k8_relat_1(k8_relset_1(k2_zfmisc_1(X1,X2),X3,X4,X5),X6),k1_zfmisc_1(X1))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_482, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~r2_hidden(X1,k2_pcomps_1(X2))|~l1_metric_1(X2))).#trainneg
cnf(i_0_281, negated_conjecture, (g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(k3_pcomps_1(esk2_0)))=k3_pcomps_1(esk2_0)|~v1_pre_topc(k3_pcomps_1(esk2_0))|~l1_pre_topc(k3_pcomps_1(esk2_0)))).#trainneg
cnf(i_0_490, plain, (r2_hidden(k8_relat_1(X3,X4),k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_511, plain, (v1_xboole_0(k9_xtuple_0(X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_463, plain, (v1_xreal_0(esk14_2(X1,X2))|r1_tarski(X1,X2)|X1!=k10_xtuple_0(X3)|~v3_valued_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3))).#trainneg
cnf(i_0_381, plain, (v1_xboole_0(k8_relset_1(X1,X2,X3,X4))|m1_subset_1(X5,X1)|~m1_subset_1(X5,k8_relset_1(X1,X2,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_427, plain, (X1=u1_struct_0(k3_pcomps_1(X2))|g1_pre_topc(X1,X3)!=k3_pcomps_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~l1_metric_1(X2))).#trainneg
cnf(i_0_492, plain, (k1_relset_1(X1,X2)=k1_xboole_0|~v4_relat_1(X2,X1)|~v1_xboole_0(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_518, plain, (m1_subset_1(esk12_2(X1,u1_pre_topc(X2)),k1_zfmisc_1(u1_struct_0(X2)))|v3_pre_topc(u1_pre_topc(X2),k3_pcomps_1(X1))|v2_struct_0(X1)|~l1_pre_topc(X2)|~m1_subset_1(u1_pre_topc(X2),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_461, plain, (v1_xboole_0(k9_metric_1(X1,X2,esk11_3(X1,X3,X2)))|m1_subset_1(X4,X3)|v2_struct_0(X1)|~m1_subset_1(X4,k9_metric_1(X1,X2,esk11_3(X1,X3,X2)))|~r2_hidden(X2,X3)|~v3_pre_topc(X3,k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_465, plain, (v1_xreal_0(esk12_2(X1,X2))|v3_pre_topc(X2,k3_pcomps_1(X1))|v2_struct_0(X1)|X2!=k10_xtuple_0(X3)|~v3_valued_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_557, plain, (v1_xboole_0(k10_xtuple_0(X3))|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_549, plain, (k2_relset_1(X1,X2)=k1_xboole_0|~v5_relat_1(X2,X1)|~v1_xboole_0(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_544, plain, (v1_relat_1(k8_relat_1(X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_464, plain, (v1_xreal_0(esk12_2(X1,X2))|v1_xboole_0(u1_struct_0(X1))|v3_pre_topc(X2,k3_pcomps_1(X1))|u1_struct_0(X1)!=k10_xtuple_0(X3)|~v3_valued_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_361, plain, (r2_hidden(k1_funct_1(X1,esk12_2(X2,k9_xtuple_0(X1))),X3)|v3_pre_topc(k9_xtuple_0(X1),k3_pcomps_1(X2))|v2_struct_0(X2)|X3!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~m1_subset_1(k9_xtuple_0(X1),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_597, plain, (m1_subset_1(k9_xtuple_0(k8_relset_1(k2_zfmisc_1(X1,X2),X3,X4,X5)),k1_zfmisc_1(X1))|~v1_relat_1(k8_relset_1(k2_zfmisc_1(X1,X2),X3,X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_509, plain, (r2_hidden(esk7_3(X2,X1,X3),k1_zfmisc_1(u1_struct_0(X1)))|v5_pre_topc(X3,X2,k3_pcomps_1(X1))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_xreal_0(esk5_3(X2,X1,X3))|~m1_subset_1(esk6_3(X2,X1,X3),u1_struct_0(X1))|~l1_metric_1(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X1)))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X3)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_459, plain, (m1_subset_1(X1,X2)|v2_struct_0(X3)|~r2_hidden(X1,k9_metric_1(X3,X4,esk11_3(X3,X2,X4)))|~r2_hidden(X4,X2)|~v3_pre_topc(X2,k3_pcomps_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_396, plain, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),X4)|X4!=k10_xtuple_0(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_616, negated_conjecture, (v5_pre_topc(X1,X2,k3_pcomps_1(esk1_0))|v2_struct_0(X2)|~v3_pre_topc(k8_relat_1(X1,esk7_3(X2,esk1_0,X1)),X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_1015, plain, (r1_tarski(esk7_3(X2,X1,X3),X4)|m1_subset_1(esk14_2(esk7_3(X2,X1,X3),X4),u1_struct_0(X1))|v5_pre_topc(X3,X2,k3_pcomps_1(X1))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_xreal_0(esk5_3(X2,X1,X3))|~m1_subset_1(esk6_3(X2,X1,X3),u1_struct_0(X1))|~l1_metric_1(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X1)))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X3)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_837, plain, (r2_hidden(esk12_2(X1,k8_relset_1(X2,X3,X4,X5)),X2)|v3_pre_topc(k8_relset_1(X2,X3,X4,X5),k3_pcomps_1(X1))|v2_struct_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k8_relset_1(X2,X3,X4,X5),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_936, plain, (r2_hidden(esk14_2(esk7_3(X1,X2,X3),X4),u1_struct_0(X2))|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|v2_struct_0(X1)|v2_struct_0(X2)|r1_tarski(esk7_3(X1,X2,X3),X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_635, plain, (r1_tarski(k8_relat_1(X1,X2),k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1429, plain, (r2_hidden(esk8_3(X1,X2,X3),X4)|X4!=k8_relat_1(X1,X5)|~r2_hidden(esk8_3(X1,X2,X3),k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|X5!=k10_xtuple_0(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_1062, plain, (r2_hidden(esk14_2(k9_xtuple_0(X1),X2),X3)|r1_tarski(k9_xtuple_0(X1),X2)|X3!=k8_relat_1(X1,X4)|~r2_hidden(esk14_2(k9_xtuple_0(X1),X2),k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|X4!=k10_xtuple_0(X1))).#trainneg
cnf(i_0_742, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(esk4_3(X3,X4,X2),k4_metric_1(esk1_0,X4,esk8_3(esk3_0,X5,X1)))|r1_xxreal_0(X3,k6_numbers)|X2!=k1_funct_1(esk3_0,X4)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(esk8_3(esk3_0,X5,X1),u1_struct_0(esk1_0))|X5!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,X5)|~v1_relat_1(esk3_0))).#trainneg
cnf(i_0_921, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_705, negated_conjecture, (u1_struct_0(esk2_0)=X1|g1_pre_topc(X1,X2)!=k3_pcomps_1(esk2_0)|~m1_subset_1(k2_pcomps_1(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_1030, negated_conjecture, (v3_pre_topc(u1_struct_0(esk1_0),k3_pcomps_1(esk1_0))|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_659, negated_conjecture, (u1_struct_0(esk1_0)=X1|g1_pre_topc(X1,X2)!=k3_pcomps_1(esk1_0)|~m1_subset_1(k2_pcomps_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_704, negated_conjecture, (X1=k2_pcomps_1(esk2_0)|g1_pre_topc(X2,X1)!=k3_pcomps_1(esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_728, negated_conjecture, (r2_hidden(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_744, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(X2),k3_pcomps_1(esk2_0))|v2_struct_0(X2)|~v1_xboole_0(k8_relset_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(esk2_0),X1,esk7_3(k3_pcomps_1(X2),esk2_0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(esk2_0))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_696, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(esk4_3(X3,X4,X1),k4_metric_1(esk1_0,X4,esk8_3(esk3_0,X5,X2)))|r1_xxreal_0(X3,k6_numbers)|X1!=k1_funct_1(esk3_0,X4)|~v1_xreal_0(X3)|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(esk8_3(esk3_0,X5,X2),u1_struct_0(esk1_0))|X5!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X5)|~v1_relat_1(esk3_0))).#trainneg
cnf(i_0_708, negated_conjecture, (u1_pre_topc(X1)=k2_pcomps_1(esk2_0)|X1!=k3_pcomps_1(esk2_0)|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_824, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(X1,k2_pcomps_1(esk1_0)))).#trainneg
cnf(i_0_658, negated_conjecture, (X1=k2_pcomps_1(esk1_0)|g1_pre_topc(X2,X1)!=k3_pcomps_1(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_841, negated_conjecture, (r2_hidden(k2_pcomps_1(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_745, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(X2),k3_pcomps_1(esk1_0))|v2_struct_0(X2)|~v1_xboole_0(k8_relset_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(esk1_0),X1,esk7_3(k3_pcomps_1(X2),esk1_0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_747, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk1_0),k3_pcomps_1(X2))|v2_struct_0(X2)|~v1_xboole_0(k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(k3_pcomps_1(X2)),X1,esk7_3(k3_pcomps_1(esk1_0),X2,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_709, negated_conjecture, (u1_struct_0(X1)=u1_struct_0(esk2_0)|X1!=k3_pcomps_1(esk2_0)|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_746, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk2_0),k3_pcomps_1(X2))|v2_struct_0(X2)|~v1_xboole_0(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(k3_pcomps_1(X2)),X1,esk7_3(k3_pcomps_1(esk2_0),X2,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_1029, negated_conjecture, (v3_pre_topc(u1_struct_0(esk2_0),k3_pcomps_1(esk2_0))|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_716, plain, (r2_hidden(X1,X2)|~v5_relat_1(X3,X2)|~r2_hidden(X1,k10_xtuple_0(X3))|~v1_relat_1(X3))).#trainneg
cnf(i_0_992, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(g1_pre_topc(X2,X3))))|~r2_hidden(X1,X3)|~l1_pre_topc(g1_pre_topc(X2,X3))|~v1_pre_topc(g1_pre_topc(X2,X3)))).#trainneg
cnf(i_0_822, negated_conjecture, (r2_hidden(k2_pcomps_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_843, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,k2_pcomps_1(esk2_0)))).#trainneg
cnf(i_0_662, negated_conjecture, (u1_pre_topc(X1)=k2_pcomps_1(esk1_0)|X1!=k3_pcomps_1(esk1_0)|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_889, plain, (v1_xboole_0(X1)|r2_hidden(esk14_2(k8_relset_1(X1,X2,X3,X4),X5),X1)|r1_tarski(k8_relset_1(X1,X2,X3,X4),X5)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1111, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k8_relat_1(esk3_0,X2)))).#trainneg
cnf(i_0_754, plain, (m1_subset_1(esk8_3(X1,X2,X3),X4)|~v4_relat_1(X1,X4)|~v1_relat_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_654, plain, (r2_hidden(esk12_2(X1,k1_relset_1(X2,X3)),X2)|v3_pre_topc(k1_relset_1(X2,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~m1_subset_1(k1_relset_1(X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_836, plain, (r2_hidden(esk14_2(k8_relset_1(X1,X2,X3,X4),X5),X1)|r1_tarski(k8_relset_1(X1,X2,X3,X4),X5)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1007, plain, (v1_partfun1(X1,k9_xtuple_0(X1))|k9_xtuple_0(X1)=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k9_xtuple_0(X1),k9_xtuple_0(X1)))))).#trainneg
cnf(i_0_614, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk2_0),k3_pcomps_1(esk1_0))|v2_struct_0(k3_pcomps_1(esk2_0))|~v3_pre_topc(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),X1,esk7_3(k3_pcomps_1(esk2_0),esk1_0,X1)),k3_pcomps_1(esk2_0))|~l1_pre_topc(k3_pcomps_1(esk2_0))|~v2_pre_topc(k3_pcomps_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_805, plain, (r2_hidden(k1_funct_1(X1,esk14_2(k9_metric_1(X2,X3,esk11_3(X2,k8_relat_1(X1,X4),X3)),X5)),X4)|r1_tarski(k9_metric_1(X2,X3,esk11_3(X2,k8_relat_1(X1,X4),X3)),X5)|v2_struct_0(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X3,k8_relat_1(X1,X4))|~v3_pre_topc(k8_relat_1(X1,X4),k3_pcomps_1(X2))|~m1_subset_1(k8_relat_1(X1,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_663, negated_conjecture, (u1_struct_0(X1)=u1_struct_0(esk1_0)|X1!=k3_pcomps_1(esk1_0)|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_660, negated_conjecture, (k2_pcomps_1(esk1_0)=X1|g1_pre_topc(X2,X1)!=k3_pcomps_1(esk1_0)|~m1_subset_1(k2_pcomps_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1301, plain, (v1_xboole_0(k9_metric_1(X1,X2,esk11_3(X1,X3,X2)))|m1_subset_1(esk14_2(k9_metric_1(X1,X2,esk11_3(X1,X3,X2)),X4),X3)|v2_struct_0(X1)|r1_tarski(k9_metric_1(X1,X2,esk11_3(X1,X3,X2)),X4)|~r2_hidden(X2,X3)|~v3_pre_topc(X3,k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_1210, plain, (m1_subset_1(esk12_2(X1,k2_pcomps_1(X2)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|v3_pre_topc(k2_pcomps_1(X2),k3_pcomps_1(X1))|v2_struct_0(X1)|~l1_metric_1(X2)|~m1_subset_1(k2_pcomps_1(X2),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_695, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,k1_funct_1(esk3_0,X2),X3))|r1_xxreal_0(esk4_3(X3,X4,X1),k4_metric_1(esk1_0,X4,X2))|r1_xxreal_0(X3,k6_numbers)|X1!=k1_funct_1(esk3_0,X4)|~v1_xreal_0(X3)|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_624, negated_conjecture, (v1_partfun1(esk3_0,u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk2_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_718, plain, (r2_hidden(esk12_2(X1,k2_relset_1(X2,X3)),X2)|v3_pre_topc(k2_relset_1(X2,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|~v5_relat_1(X3,X2)|~v1_relat_1(X3)|~m1_subset_1(k2_relset_1(X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_1020, plain, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),X4)|~v1_relat_1(X1)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X4)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_751, plain, (m1_subset_1(esk14_2(k8_relat_1(X1,X2),X3),X4)|r1_tarski(k8_relat_1(X1,X2),X3)|~v4_relat_1(X1,X4)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_618, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk2_0),k3_pcomps_1(esk2_0))|v2_struct_0(k3_pcomps_1(esk2_0))|~v3_pre_topc(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),X1,esk7_3(k3_pcomps_1(esk2_0),esk2_0,X1)),k3_pcomps_1(esk2_0))|~l1_pre_topc(k3_pcomps_1(esk2_0))|~v2_pre_topc(k3_pcomps_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(esk2_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_706, negated_conjecture, (k2_pcomps_1(esk2_0)=X1|g1_pre_topc(X2,X1)!=k3_pcomps_1(esk2_0)|~m1_subset_1(k2_pcomps_1(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_737, negated_conjecture, (v1_xboole_0(esk3_0)|~v1_xboole_0(u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_734, negated_conjecture, (v1_xboole_0(esk3_0)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_651, plain, (r2_hidden(esk14_2(k8_relat_1(X1,X2),X3),X4)|r1_tarski(k8_relat_1(X1,X2),X3)|X4!=k8_relat_1(X1,X2)|~r2_hidden(esk14_2(k8_relat_1(X1,X2),X3),k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_634, plain, (r2_hidden(X1,X2)|r1_tarski(k8_relat_1(X3,X4),X5)|X1!=k1_funct_1(X3,esk14_2(k8_relat_1(X3,X4),X5))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3))).#trainneg
cnf(i_0_725, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|v5_pre_topc(X3,X2,k3_pcomps_1(esk2_0))|v2_struct_0(X2)|~r2_hidden(X1,esk7_3(X2,esk2_0,X3))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3))).#trainneg
cnf(i_0_753, plain, (m1_subset_1(esk12_2(X1,k8_relat_1(X2,X3)),X4)|v3_pre_topc(k8_relat_1(X2,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|~v4_relat_1(X2,X4)|~v1_relat_1(X2)|~m1_subset_1(k8_relat_1(X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~v1_funct_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_749, plain, (v5_pre_topc(X1,k3_pcomps_1(X2),k3_pcomps_1(X3))|v2_struct_0(X2)|v2_struct_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X3)))))|~v1_funct_2(X1,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(k3_pcomps_1(X3)))|~v1_funct_1(X1)|~l1_metric_1(X3)|~l1_metric_1(X2)|~v9_metric_1(X3)|~v9_metric_1(X2)|~v8_metric_1(X3)|~v8_metric_1(X2)|~v7_metric_1(X3)|~v7_metric_1(X2)|~v6_metric_1(X3)|~v6_metric_1(X2)|~v1_xboole_0(u1_struct_0(k3_pcomps_1(X2))))).#trainneg
cnf(i_0_724, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk2_0))|v5_pre_topc(X3,X2,k3_pcomps_1(esk2_0))|v2_struct_0(X2)|~r2_hidden(X1,esk7_3(X2,esk2_0,X3))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1027, plain, (m1_subset_1(k9_xtuple_0(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_648, plain, (m1_subset_1(k1_funct_1(X1,esk14_2(k8_relat_1(X1,X2),X3)),X2)|r1_tarski(k8_relat_1(X1,X2),X3)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1154, plain, (v3_pre_topc(X1,k3_pcomps_1(X2))|X1!=k1_xboole_0|~v1_xreal_0(X4)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v1_xboole_0(u1_struct_0(X2)))).#trainneg
cnf(i_0_1295, plain, (v3_pre_topc(X2,k3_pcomps_1(X3))|m1_subset_1(esk12_2(X3,X2),k1_zfmisc_1(u1_struct_0(g1_pre_topc(X1,X2))))|v2_struct_0(X3)|~l1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)|~v1_pre_topc(g1_pre_topc(X1,X2)))).#trainneg
cnf(i_0_672, plain, (r2_hidden(X1,X2)|v3_pre_topc(k8_relat_1(X3,X5),k3_pcomps_1(X4))|v2_struct_0(X4)|X1!=k1_funct_1(X3,esk12_2(X4,k8_relat_1(X3,X5)))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~m1_subset_1(k8_relat_1(X3,X5),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~l1_metric_1(X4)|~v9_metric_1(X4)|~v8_metric_1(X4)|~v7_metric_1(X4)|~v6_metric_1(X4))).#trainneg
cnf(i_0_702, negated_conjecture, (l1_pre_topc(k3_pcomps_1(esk2_0))|~m1_subset_1(k2_pcomps_1(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_1743, negated_conjecture, (m1_subset_1(esk8_3(esk3_0,X1,X2),u1_struct_0(esk1_0))|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X3)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_1485, plain, (r2_hidden(esk8_3(X1,X2,X3),k8_relat_1(X1,X4))|X4!=k10_xtuple_0(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1483, plain, (m1_subset_1(k8_relat_1(X1,X2),k1_zfmisc_1(k9_xtuple_0(X1)))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_289, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk1_0),k3_pcomps_1(X2))|v2_struct_0(k3_pcomps_1(esk1_0))|v2_struct_0(X2)|~v3_pre_topc(k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(k3_pcomps_1(X2)),X1,esk7_3(k3_pcomps_1(esk1_0),X2,X1)),k3_pcomps_1(esk1_0))|~v2_pre_topc(k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_1613, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk1_0),k3_pcomps_1(esk1_0))|~v1_xboole_0(k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1,esk7_3(k3_pcomps_1(esk1_0),esk1_0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_1563, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk1_0),k3_pcomps_1(esk2_0))|~v1_xboole_0(k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),X1,esk7_3(k3_pcomps_1(esk1_0),esk2_0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk2_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_1505, negated_conjecture, (r2_hidden(X2,k9_metric_1(esk2_0,X5,X4))|r1_xxreal_0(X4,k6_numbers)|~r2_hidden(esk8_3(esk3_0,X1,X2),k9_metric_1(esk1_0,X3,esk4_3(X4,X3,X5)))|~v1_xreal_0(esk4_3(X4,X3,X5))|~m1_subset_1(esk8_3(esk3_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|X5!=k1_funct_1(esk3_0,X3)|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1)|~v1_xreal_0(X4)|~m1_subset_1(X5,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_1521, plain, (m1_subset_1(k8_relat_1(k1_xboole_0,X1),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1564, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(X2),k3_pcomps_1(esk2_0))|v2_struct_0(X2)|~v1_xboole_0(k8_relat_1(X1,esk7_3(k3_pcomps_1(X2),esk2_0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(esk2_0))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_1614, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(X2),k3_pcomps_1(esk1_0))|v2_struct_0(X2)|~v1_xboole_0(k8_relat_1(X1,esk7_3(k3_pcomps_1(X2),esk1_0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_1567, negated_conjecture, (r2_hidden(X5,k9_metric_1(esk2_0,X2,X4))|r1_xxreal_0(X4,k6_numbers)|~r2_hidden(esk8_3(esk3_0,X1,X2),k9_metric_1(esk1_0,X3,esk4_3(X4,X3,X5)))|~v1_xreal_0(esk4_3(X4,X3,X5))|~m1_subset_1(esk8_3(esk3_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|X5!=k1_funct_1(esk3_0,X3)|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1)|~v1_xreal_0(X4)|~m1_subset_1(X5,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_1662, plain, (r2_hidden(esk8_3(X1,X2,X3),X4)|~v4_relat_1(X1,X4)|~v1_relat_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_618, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk2_0),k3_pcomps_1(esk2_0))|v2_struct_0(k3_pcomps_1(esk2_0))|~v3_pre_topc(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),X1,esk7_3(k3_pcomps_1(esk2_0),esk2_0,X1)),k3_pcomps_1(esk2_0))|~v2_pre_topc(k3_pcomps_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(esk2_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_1492, plain, (r2_hidden(X1,k9_metric_1(X2,esk12_2(X3,k9_metric_1(X2,X1,X4)),X4))|v2_struct_0(X2)|v3_pre_topc(k9_metric_1(X2,X1,X4),k3_pcomps_1(X3))|v2_struct_0(X3)|~v1_xreal_0(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_metric_1(X2)|~v8_metric_1(X2)|~m1_subset_1(k9_metric_1(X2,X1,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_614, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk2_0),k3_pcomps_1(esk1_0))|v2_struct_0(k3_pcomps_1(esk2_0))|~v3_pre_topc(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),X1,esk7_3(k3_pcomps_1(esk2_0),esk1_0,X1)),k3_pcomps_1(esk2_0))|~v2_pre_topc(k3_pcomps_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_1481, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,X1),k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(esk4_3(X3,X4,X2),k4_metric_1(esk1_0,X1,X4))|r1_xxreal_0(X3,k6_numbers)|X2!=k1_funct_1(esk3_0,X4)|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,X1),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1453, plain, (r2_hidden(X1,X2)|v3_pre_topc(k8_relset_1(k9_xtuple_0(X3),X5,X6,X7),k3_pcomps_1(X4))|v2_struct_0(X4)|X1!=k1_funct_1(X3,esk12_2(X4,k8_relset_1(k9_xtuple_0(X3),X5,X6,X7)))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~m1_subset_1(k8_relset_1(k9_xtuple_0(X3),X5,X6,X7),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k9_xtuple_0(X3),X5)))|~l1_metric_1(X4)|~v9_metric_1(X4)|~v8_metric_1(X4)|~v7_metric_1(X4)|~v6_metric_1(X4))).#trainneg
cnf(i_0_280, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk2_0),k3_pcomps_1(X2))|v2_struct_0(k3_pcomps_1(esk2_0))|v2_struct_0(X2)|~v3_pre_topc(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(k3_pcomps_1(X2)),X1,esk7_3(k3_pcomps_1(esk2_0),X2,X1)),k3_pcomps_1(esk2_0))|~v2_pre_topc(k3_pcomps_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_1480, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,X2),k9_metric_1(esk2_0,X4,X3))|r1_xxreal_0(X3,k6_numbers)|~r2_hidden(X1,k9_metric_1(esk1_0,X2,esk4_3(X3,X1,X4)))|~v1_xreal_0(esk4_3(X3,X1,X4))|~m1_subset_1(X2,u1_struct_0(esk1_0))|X4!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,X2),u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1626, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk2_0),k3_pcomps_1(X2))|v2_struct_0(X2)|~v1_xboole_0(k8_relat_1(X1,esk7_3(k3_pcomps_1(esk2_0),X2,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_1562, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk2_0),k3_pcomps_1(esk2_0))|~v1_xboole_0(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),X1,esk7_3(k3_pcomps_1(esk2_0),esk2_0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(esk2_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_1619, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk1_0),k3_pcomps_1(X2))|v2_struct_0(X2)|~v1_xboole_0(k8_relat_1(X1,esk7_3(k3_pcomps_1(esk1_0),X2,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_1503, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k8_relat_1(esk3_0,X2)))).#trainneg
cnf(i_0_1612, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk2_0),k3_pcomps_1(esk1_0))|~v1_xboole_0(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),X1,esk7_3(k3_pcomps_1(esk2_0),esk1_0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_1498, plain, (r1_tarski(k9_xtuple_0(X1),X2)|r2_hidden(esk14_2(k9_xtuple_0(X1),X2),k8_relat_1(X1,X3))|X3!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1623, negated_conjecture, (u1_struct_0(k3_pcomps_1(X1))=u1_struct_0(esk2_0)|k3_pcomps_1(X1)!=k3_pcomps_1(esk2_0)|~l1_pre_topc(k3_pcomps_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_1506, negated_conjecture, (r2_hidden(X3,k9_metric_1(esk2_0,X5,X4))|r1_xxreal_0(X4,k6_numbers)|~r2_hidden(X1,k9_metric_1(esk1_0,esk8_3(esk3_0,X2,X3),esk4_3(X4,X1,X5)))|~v1_xreal_0(esk4_3(X4,X1,X5))|~m1_subset_1(esk8_3(esk3_0,X2,X3),u1_struct_0(esk1_0))|X5!=k1_funct_1(esk3_0,X1)|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(X3,X2)|~v1_xreal_0(X4)|~m1_subset_1(X5,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_1569, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(esk4_3(X3,X4,X1),k4_metric_1(esk1_0,esk8_3(esk3_0,X5,X2),X4))|r1_xxreal_0(X3,k6_numbers)|X1!=k1_funct_1(esk3_0,X4)|X5!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X5)|~v1_xreal_0(X3)|~m1_subset_1(esk8_3(esk3_0,X5,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_1592, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,k2_pcomps_1(esk2_0))|~l1_pre_topc(X2)|X2!=k3_pcomps_1(esk2_0)|~v1_pre_topc(X2))).#trainneg
cnf(i_0_2069, negated_conjecture, (r2_hidden(esk12_2(X1,esk7_3(X2,esk2_0,X3)),u1_struct_0(esk2_0))|v5_pre_topc(X3,X2,k3_pcomps_1(esk2_0))|v2_struct_0(X2)|v3_pre_topc(esk7_3(X2,esk2_0,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~m1_subset_1(esk7_3(X2,esk2_0,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_1568, negated_conjecture, (r2_hidden(X5,k9_metric_1(esk2_0,X3,X4))|r1_xxreal_0(X4,k6_numbers)|~r2_hidden(X1,k9_metric_1(esk1_0,esk8_3(esk3_0,X2,X3),esk4_3(X4,X1,X5)))|~v1_xreal_0(esk4_3(X4,X1,X5))|~m1_subset_1(esk8_3(esk3_0,X2,X3),u1_struct_0(esk1_0))|X5!=k1_funct_1(esk3_0,X1)|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(X3,X2)|~v1_xreal_0(X4)|~m1_subset_1(X5,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_1922, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,k1_funct_1(esk3_0,X2),X3))|r1_xxreal_0(esk4_3(X3,X4,X1),k4_metric_1(esk1_0,X2,X4))|r1_xxreal_0(X3,k6_numbers)|X1!=k1_funct_1(esk3_0,X4)|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,X2),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1507, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(esk4_3(X3,X4,X2),k4_metric_1(esk1_0,esk8_3(esk3_0,X5,X1),X4))|r1_xxreal_0(X3,k6_numbers)|X2!=k1_funct_1(esk3_0,X4)|X5!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,X5)|~v1_xreal_0(X3)|~m1_subset_1(esk8_3(esk3_0,X5,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_1476, plain, (r1_tarski(esk7_3(X1,X2,X3),u1_struct_0(X2))|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_1471, plain, (r2_hidden(esk12_2(X1,k8_relat_1(X4,X5)),X2)|v3_pre_topc(k8_relat_1(X4,X5),k3_pcomps_1(X1))|v2_struct_0(X1)|~m1_subset_1(k8_relat_1(X4,X5),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_1920, negated_conjecture, (r2_hidden(X4,k9_metric_1(esk2_0,k1_funct_1(esk3_0,X1),X3))|r1_xxreal_0(X3,k6_numbers)|~r2_hidden(X1,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)))|~v1_xreal_0(esk4_3(X3,X2,X4))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|X4!=k1_funct_1(esk3_0,X2)|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,X1),u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_1837, plain, (r1_tarski(k8_relset_1(X1,X2,X3,X4),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1703, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,k2_pcomps_1(esk1_0))|~l1_pre_topc(X2)|X2!=k3_pcomps_1(esk1_0)|~v1_pre_topc(X2))).#trainneg
cnf(i_0_1921, negated_conjecture, (r2_hidden(X4,k9_metric_1(esk2_0,k1_funct_1(esk3_0,X2),X3))|r1_xxreal_0(X3,k6_numbers)|~r2_hidden(X1,k9_metric_1(esk1_0,X2,esk4_3(X3,X1,X4)))|~v1_xreal_0(esk4_3(X3,X1,X4))|~m1_subset_1(X2,u1_struct_0(esk1_0))|X4!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,X2),u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2100, plain, (v1_xboole_0(X1)|r2_hidden(esk12_2(X2,k8_relat_1(X3,X4)),X1)|v3_pre_topc(k8_relat_1(X3,X4),k3_pcomps_1(X2))|v2_struct_0(X2)|~v4_relat_1(X3,X1)|~v1_relat_1(X3)|~m1_subset_1(k8_relat_1(X3,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~v1_funct_1(X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_2193, plain, (r2_hidden(k1_funct_1(X1,esk12_2(X2,k8_relat_1(X1,X3))),X4)|v3_pre_topc(k8_relat_1(X1,X3),k3_pcomps_1(X2))|v2_struct_0(X2)|X4!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~m1_subset_1(k8_relat_1(X1,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_1615, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(X2),k3_pcomps_1(esk1_0))|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v1_xboole_0(u1_struct_0(k3_pcomps_1(X2))))).#trainneg
cnf(i_0_2032, plain, (r2_hidden(esk14_2(k8_relat_1(X1,X2),X3),k9_xtuple_0(X4))|r1_tarski(k8_relat_1(X1,X2),X3)|~v1_relat_1(X4)|~v1_funct_1(X4)|k8_relat_1(X4,X5)!=k8_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1565, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(X2),k3_pcomps_1(esk2_0))|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(k3_pcomps_1(X2)),u1_struct_0(esk2_0))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v1_xboole_0(u1_struct_0(k3_pcomps_1(X2))))).#trainneg
cnf(i_0_1538, negated_conjecture, (u1_struct_0(k3_pcomps_1(X1))=u1_struct_0(esk1_0)|k3_pcomps_1(X1)!=k3_pcomps_1(esk1_0)|~l1_metric_1(X1))).#trainneg
cnf(i_0_1571, negated_conjecture, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk2_0))))|X1!=k3_pcomps_1(esk2_0)|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1739, negated_conjecture, (m1_subset_1(esk14_2(k8_relat_1(esk3_0,X1),X2),u1_struct_0(esk1_0))|r1_tarski(k8_relat_1(esk3_0,X1),X2))).#trainneg
cnf(i_0_1912, plain, (v1_xboole_0(esk12_2(X1,k2_pcomps_1(X2)))|v3_pre_topc(k2_pcomps_1(X2),k3_pcomps_1(X1))|v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(k3_pcomps_1(X2)))|~m1_subset_1(k2_pcomps_1(X2),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_1964, plain, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),k9_xtuple_0(X4))|~v1_relat_1(X4)|~v1_funct_1(X4)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k8_relat_1(X4,X5))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2064, negated_conjecture, (r2_hidden(esk14_2(esk7_3(X1,esk2_0,X2),X3),u1_struct_0(esk2_0))|v5_pre_topc(X2,X1,k3_pcomps_1(esk2_0))|v2_struct_0(X1)|r1_tarski(esk7_3(X1,esk2_0,X2),X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk2_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1588, negated_conjecture, (m1_subset_1(k2_pcomps_1(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)|X1!=k3_pcomps_1(esk2_0)|~v1_pre_topc(X1))).#trainneg
cnf(i_0_1486, plain, (r2_hidden(esk6_3(X1,X2,X3),k9_metric_1(X2,X4,esk5_3(X1,X2,X3)))|v2_struct_0(X2)|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|v2_struct_0(X1)|~r2_hidden(X4,esk7_3(X1,X2,X3))|~v1_xreal_0(esk5_3(X1,X2,X3))|~m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X2))|~l1_metric_1(X2)|~v8_metric_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~v9_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_2217, plain, (v1_xboole_0(k8_relset_1(k1_xboole_0,X3,X4,X5))|~v1_xboole_0(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,X3))))).#trainneg
cnf(i_0_1976, plain, (m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~l1_metric_1(X4)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k2_pcomps_1(X4))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1742, negated_conjecture, (m1_subset_1(esk12_2(X1,k8_relat_1(esk3_0,X2)),u1_struct_0(esk1_0))|v3_pre_topc(k8_relat_1(esk3_0,X2),k3_pcomps_1(X1))|v2_struct_0(X1)|~m1_subset_1(k8_relat_1(esk3_0,X2),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_2051, plain, (r2_hidden(esk14_2(k8_relat_1(X1,X2),X3),X4)|v2_struct_0(X5)|r1_tarski(k8_relat_1(X1,X2),X3)|~r2_hidden(X6,X4)|~v3_pre_topc(X4,k3_pcomps_1(X5))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X5))))|~m1_subset_1(X6,u1_struct_0(X5))|~l1_metric_1(X5)|~v9_metric_1(X5)|~v8_metric_1(X5)|~v7_metric_1(X5)|~v6_metric_1(X5)|k9_metric_1(X5,X6,esk11_3(X5,X4,X6))!=k8_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2195, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|~v2_struct_0(k3_pcomps_1(esk2_0)))).#trainneg
cnf(i_0_1975, plain, (m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(X4)))|~l1_metric_1(X4)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k2_pcomps_1(X4))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1525, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_1629, negated_conjecture, (v3_pre_topc(u1_struct_0(esk2_0),k3_pcomps_1(esk2_0))|r1_xxreal_0(esk11_3(X1,X2,X3),k6_numbers)|v2_struct_0(X1)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_1679, negated_conjecture, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|X1!=k3_pcomps_1(esk1_0)|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1451, plain, (r2_hidden(esk12_2(X1,k8_relset_1(k8_relat_1(X2,X3),X4,X5,X6)),k9_xtuple_0(X2))|v3_pre_topc(k8_relset_1(k8_relat_1(X2,X3),X4,X5,X6),k3_pcomps_1(X1))|v2_struct_0(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~m1_subset_1(k8_relset_1(k8_relat_1(X2,X3),X4,X5,X6),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k8_relat_1(X2,X3),X4)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_1838, plain, (r1_tarski(k8_relat_1(X3,X4),X5)|r2_hidden(esk14_2(k8_relat_1(X3,X4),X5),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1983, plain, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),X4)|v2_struct_0(X5)|~r2_hidden(X6,X4)|~v3_pre_topc(X4,k3_pcomps_1(X5))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X5))))|~m1_subset_1(X6,u1_struct_0(X5))|~l1_metric_1(X5)|~v9_metric_1(X5)|~v8_metric_1(X5)|~v7_metric_1(X5)|~v6_metric_1(X5)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k9_metric_1(X5,X6,esk11_3(X5,X4,X6)))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1466, plain, (r2_hidden(esk12_2(X1,k8_relset_1(k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5,X6,X7)),X4)|v2_struct_0(X2)|v3_pre_topc(k8_relset_1(k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5,X6,X7),k3_pcomps_1(X1))|v2_struct_0(X1)|~r2_hidden(X3,X4)|~v3_pre_topc(X4,k3_pcomps_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~m1_subset_1(k8_relset_1(k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5,X6,X7),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_1575, negated_conjecture, (g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(X1))=k3_pcomps_1(esk2_0)|X1!=k3_pcomps_1(esk2_0)|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1491, plain, (r2_hidden(X1,k9_metric_1(X2,esk12_2(X3,k9_metric_1(X4,X5,esk11_3(X4,k9_metric_1(X2,X1,X6),X5))),X6))|v2_struct_0(X2)|v3_pre_topc(k9_metric_1(X4,X5,esk11_3(X4,k9_metric_1(X2,X1,X6),X5)),k3_pcomps_1(X3))|v2_struct_0(X3)|v2_struct_0(X4)|~v1_xreal_0(X6)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_metric_1(X2)|~v8_metric_1(X2)|~r2_hidden(X5,k9_metric_1(X2,X1,X6))|~v3_pre_topc(k9_metric_1(X2,X1,X6),k3_pcomps_1(X4))|~m1_subset_1(k9_metric_1(X4,X5,esk11_3(X4,k9_metric_1(X2,X1,X6),X5)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(k9_metric_1(X2,X1,X6),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~m1_subset_1(X5,u1_struct_0(X4))|~l1_metric_1(X4)|~l1_metric_1(X3)|~v9_metric_1(X4)|~v9_metric_1(X3)|~v8_metric_1(X4)|~v8_metric_1(X3)|~v7_metric_1(X4)|~v7_metric_1(X3)|~v6_metric_1(X4)|~v6_metric_1(X3))).#trainneg
cnf(i_0_2284, negated_conjecture, (m1_subset_1(esk8_3(esk3_0,X1,X2),u1_struct_0(esk1_0))|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_2486, negated_conjecture, (r2_hidden(X1,X2)|v3_pre_topc(k8_relset_1(u1_struct_0(esk1_0),X3,X4,X5),k3_pcomps_1(X6))|v2_struct_0(X6)|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k1_funct_1(esk3_0,esk12_2(X6,k8_relset_1(u1_struct_0(esk1_0),X3,X4,X5)))|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(k8_relset_1(u1_struct_0(esk1_0),X3,X4,X5),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X6))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X3)))|~l1_metric_1(X6)|~v9_metric_1(X6)|~v8_metric_1(X6)|~v7_metric_1(X6)|~v6_metric_1(X6))).#trainneg
cnf(i_0_2523, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,k10_xtuple_0(esk3_0)))).#trainneg
cnf(i_0_2778, negated_conjecture, (m1_subset_1(esk14_2(esk7_3(X1,esk2_0,X2),X3),u1_struct_0(esk2_0))|r1_tarski(esk7_3(X1,esk2_0,X2),X3)|v5_pre_topc(X2,X1,k3_pcomps_1(esk2_0))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk2_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_2782, negated_conjecture, (r1_tarski(esk7_3(X1,esk2_0,X2),u1_struct_0(esk2_0))|v5_pre_topc(X2,X1,k3_pcomps_1(esk2_0))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk2_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_2422, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk2_0),k3_pcomps_1(esk1_0))|v2_struct_0(k3_pcomps_1(esk2_0))|~v3_pre_topc(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),X1,esk7_3(k3_pcomps_1(esk2_0),esk1_0,X1)),k3_pcomps_1(esk2_0))|~v2_pre_topc(k3_pcomps_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_2525, negated_conjecture, (m1_subset_1(esk12_2(X1,esk7_3(X2,esk2_0,X3)),u1_struct_0(esk2_0))|v3_pre_topc(esk7_3(X2,esk2_0,X3),k3_pcomps_1(X1))|v5_pre_topc(X3,X2,k3_pcomps_1(esk2_0))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(esk7_3(X2,esk2_0,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_2441, negated_conjecture, (r2_hidden(esk14_2(k8_relat_1(esk3_0,X1),X2),u1_struct_0(esk1_0))|r1_tarski(k8_relat_1(esk3_0,X1),X2))).#trainneg
cnf(i_0_2357, plain, (m1_subset_1(esk8_3(X1,X2,X3),X4)|~v4_relat_1(X5,X4)|~v1_relat_1(X5)|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,k1_relset_1(X4,X5))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2578, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X1)),X4)),X3))|r1_xxreal_0(X3,k6_numbers)|r1_tarski(k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X1)),X4)|X1!=k1_funct_1(esk3_0,X2)|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X1)),X4)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X1)),X4),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_2299, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk1_0),k3_pcomps_1(esk1_0))|v2_struct_0(k3_pcomps_1(esk1_0))|~v3_pre_topc(k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1,esk7_3(k3_pcomps_1(esk1_0),esk1_0,X1)),k3_pcomps_1(esk1_0))|~v2_pre_topc(k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_2376, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(X3,k6_numbers)|X2!=k1_funct_1(esk3_0,X4)|X5!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,X5)|~v1_xreal_0(X3)|~m1_subset_1(esk8_3(esk3_0,X5,X1),u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~v4_relat_1(esk3_0,k9_metric_1(esk1_0,X4,esk4_3(X3,X4,X2))))).#trainneg
cnf(i_0_2298, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk1_0),k3_pcomps_1(esk2_0))|v2_struct_0(k3_pcomps_1(esk1_0))|~v3_pre_topc(k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),X1,esk7_3(k3_pcomps_1(esk1_0),esk2_0,X1)),k3_pcomps_1(esk1_0))|~v2_pre_topc(k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk2_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_2473, negated_conjecture, (r2_hidden(esk8_3(esk3_0,X1,X2),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_2421, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk2_0),k3_pcomps_1(esk2_0))|v2_struct_0(k3_pcomps_1(esk2_0))|~v3_pre_topc(k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),X1,esk7_3(k3_pcomps_1(esk2_0),esk2_0,X1)),k3_pcomps_1(esk2_0))|~v2_pre_topc(k3_pcomps_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(esk2_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_2524, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk3_0))|m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X1,k10_xtuple_0(esk3_0)))).#trainneg
cnf(i_0_624, negated_conjecture, (v1_partfun1(esk3_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2479, negated_conjecture, (r1_tarski(k8_relat_1(esk3_0,X1),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0)).#trainneg
cnf(i_0_2799, plain, (r2_hidden(X1,k9_metric_1(X2,esk6_3(X3,X2,X4),esk5_3(X3,X2,X4)))|v2_struct_0(X2)|v5_pre_topc(X4,X3,k3_pcomps_1(X2))|v2_struct_0(X3)|~v1_xreal_0(esk5_3(X3,X2,X4))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_metric_1(X2)|~v8_metric_1(X2)|~r2_hidden(X1,esk7_3(X3,X2,X4))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X4)|~v9_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_2716, negated_conjecture, (v1_xboole_0(k9_xtuple_0(esk3_0))|m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,k9_xtuple_0(esk3_0)))).#trainneg
cnf(i_0_2938, plain, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,esk14_2(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6))),X5)|v2_struct_0(X3)|r1_tarski(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6)|X2!=k10_xtuple_0(X1)|~r2_hidden(X4,X5)|~r2_hidden(esk14_2(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2)|~v1_relat_1(X1)|~v3_pre_topc(X5,k3_pcomps_1(X3))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X4,u1_struct_0(X3))|~v1_funct_1(X1)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_2715, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k9_xtuple_0(esk3_0)))).#trainneg
cnf(i_0_2377, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(X3,k6_numbers)|X1!=k1_funct_1(esk3_0,X4)|X5!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X5)|~v1_xreal_0(X3)|~m1_subset_1(esk8_3(esk3_0,X5,X2),u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~v4_relat_1(esk3_0,k9_metric_1(esk1_0,X4,esk4_3(X3,X4,X1))))).#trainneg
cnf(i_0_2468, plain, (r1_tarski(k9_xtuple_0(X1),k8_relat_1(X1,X2))|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2370, plain, (m1_subset_1(esk8_3(X1,X2,X3),X4)|v2_struct_0(X5)|~r2_hidden(X6,X4)|~v3_pre_topc(X4,k3_pcomps_1(X5))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X5))))|~m1_subset_1(X6,u1_struct_0(X5))|~l1_metric_1(X5)|~v9_metric_1(X5)|~v8_metric_1(X5)|~v7_metric_1(X5)|~v6_metric_1(X5)|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,k9_metric_1(X5,X6,esk11_3(X5,X4,X6)))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2362, plain, (m1_subset_1(esk8_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~l1_metric_1(X4)|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,k2_pcomps_1(X4))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2449, negated_conjecture, (r2_hidden(esk8_3(esk3_0,X1,X2),u1_struct_0(esk1_0))|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X3)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_2303, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk1_0),k3_pcomps_1(esk1_0))|~v1_xboole_0(k8_relat_1(X1,esk7_3(k3_pcomps_1(esk1_0),esk1_0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_2300, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk1_0),k3_pcomps_1(X2))|v2_struct_0(k3_pcomps_1(esk1_0))|v2_struct_0(X2)|~v3_pre_topc(k8_relat_1(X1,esk7_3(k3_pcomps_1(esk1_0),X2,X1)),k3_pcomps_1(esk1_0))|~v2_pre_topc(k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_2448, negated_conjecture, (r2_hidden(esk8_3(esk3_0,X1,X2),u1_struct_0(esk1_0))|X3!=k10_xtuple_0(esk3_0)|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_3004, negated_conjecture, (X1=u1_pre_topc(X2)|g1_pre_topc(X3,X1)!=k3_pcomps_1(esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X3)))|X2!=k3_pcomps_1(esk2_0)|~v1_pre_topc(X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_2827, negated_conjecture, (m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(esk1_0)))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k2_pcomps_1(esk1_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2435, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk2_0),k3_pcomps_1(esk2_0))|~v1_xboole_0(k8_relat_1(X1,esk7_3(k3_pcomps_1(esk2_0),esk2_0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(esk2_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_2306, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk1_0),k3_pcomps_1(esk2_0))|~v1_xboole_0(k8_relat_1(X1,esk7_3(k3_pcomps_1(esk1_0),esk2_0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk2_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_2616, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,k1_funct_1(esk3_0,X2),X3))|r1_xxreal_0(X3,k6_numbers)|X1!=k1_funct_1(esk3_0,esk8_3(X4,X5,X6))|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(esk8_3(X4,X5,X6),u1_struct_0(esk1_0))|X5!=k10_xtuple_0(X4)|~v4_relat_1(X4,k9_metric_1(esk1_0,X2,esk4_3(X3,esk8_3(X4,X5,X6),X1)))|~r2_hidden(X6,X5)|~v1_relat_1(X4)|~v1_funct_1(X4))).#trainneg
cnf(i_0_2361, plain, (m1_subset_1(esk8_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X4)))|~l1_metric_1(X4)|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,k2_pcomps_1(X4))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2826, negated_conjecture, (m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(esk2_0)))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k2_pcomps_1(esk2_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2387, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)))),k9_metric_1(esk2_0,X4,X3))|r1_xxreal_0(X3,k6_numbers)|v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),k3_pcomps_1(X1))|v2_struct_0(X1)|X4!=k1_funct_1(esk3_0,X2)|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)))),u1_struct_0(esk2_0))|~m1_subset_1(esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4))),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_2457, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk2_0),k3_pcomps_1(esk1_0))|~v1_xboole_0(k8_relat_1(X1,esk7_3(k3_pcomps_1(esk2_0),esk1_0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_2423, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk2_0),k3_pcomps_1(X2))|v2_struct_0(k3_pcomps_1(esk2_0))|v2_struct_0(X2)|~v3_pre_topc(k8_relat_1(X1,esk7_3(k3_pcomps_1(esk2_0),X2,X1)),k3_pcomps_1(esk2_0))|~v2_pre_topc(k3_pcomps_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_2408, plain, (r2_hidden(esk6_3(X1,X2,X3),k9_metric_1(X2,esk12_2(X4,esk7_3(X1,X2,X3)),esk5_3(X1,X2,X3)))|v3_pre_topc(esk7_3(X1,X2,X3),k3_pcomps_1(X4))|v2_struct_0(X4)|v2_struct_0(X2)|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|v2_struct_0(X1)|~v1_xreal_0(esk5_3(X1,X2,X3))|~m1_subset_1(esk7_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X2))|~l1_metric_1(X2)|~l1_metric_1(X4)|~v9_metric_1(X4)|~v8_metric_1(X2)|~v8_metric_1(X4)|~v7_metric_1(X4)|~v6_metric_1(X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~v9_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_1780, plain, (r2_hidden(X1,X2)|v3_pre_topc(k1_relset_1(k9_xtuple_0(X3),X5),k3_pcomps_1(X4))|v2_struct_0(X4)|X1!=k1_funct_1(X3,esk12_2(X4,k1_relset_1(k9_xtuple_0(X3),X5)))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X5,k9_xtuple_0(X3))|~v1_relat_1(X5)|~m1_subset_1(k1_relset_1(k9_xtuple_0(X3),X5),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~l1_metric_1(X4)|~v9_metric_1(X4)|~v8_metric_1(X4)|~v7_metric_1(X4)|~v6_metric_1(X4))).#trainneg
cnf(i_0_2134, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k9_xtuple_0(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1961, plain, (m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),X4)|~v5_relat_1(X5,X4)|~v1_relat_1(X5)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k2_relset_1(X4,X5))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1888, plain, (v1_xboole_0(k9_metric_1(X1,X2,esk11_3(X1,X3,X2)))|m1_subset_1(esk14_2(k9_metric_1(X4,X5,esk11_3(X4,k9_metric_1(X1,X2,esk11_3(X1,X3,X2)),X5)),X6),X3)|v2_struct_0(X1)|r1_tarski(k9_metric_1(X4,X5,esk11_3(X4,k9_metric_1(X1,X2,esk11_3(X1,X3,X2)),X5)),X6)|v2_struct_0(X4)|~r2_hidden(X2,X3)|~v3_pre_topc(X3,k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~r2_hidden(X5,k9_metric_1(X1,X2,esk11_3(X1,X3,X2)))|~v3_pre_topc(k9_metric_1(X1,X2,esk11_3(X1,X3,X2)),k3_pcomps_1(X4))|~m1_subset_1(k9_metric_1(X1,X2,esk11_3(X1,X3,X2)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~m1_subset_1(X5,u1_struct_0(X4))|~l1_metric_1(X4)|~v9_metric_1(X4)|~v8_metric_1(X4)|~v7_metric_1(X4)|~v6_metric_1(X4))).#trainneg
cnf(i_0_2014, plain, (v1_xboole_0(X1)|r2_hidden(esk14_2(k8_relat_1(X2,X3),X4),X1)|r1_tarski(k8_relat_1(X2,X3),X4)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2063, negated_conjecture, (r2_hidden(esk14_2(k9_metric_1(X1,X2,esk11_3(X1,esk7_3(X3,esk2_0,X4),X2)),X5),u1_struct_0(esk2_0))|v5_pre_topc(X4,X3,k3_pcomps_1(esk2_0))|v2_struct_0(X3)|r1_tarski(k9_metric_1(X1,X2,esk11_3(X1,esk7_3(X3,esk2_0,X4),X2)),X5)|v2_struct_0(X1)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk2_0))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(esk2_0))|~v1_funct_1(X4)|~r2_hidden(X2,esk7_3(X3,esk2_0,X4))|~v3_pre_topc(esk7_3(X3,esk2_0,X4),k3_pcomps_1(X1))|~m1_subset_1(esk7_3(X3,esk2_0,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_2406, plain, (r2_hidden(X1,u1_struct_0(X2))|v3_pre_topc(k9_metric_1(X2,X1,X4),k3_pcomps_1(X3))|v2_struct_0(X3)|v2_struct_0(X2)|~v1_xreal_0(X4)|~m1_subset_1(esk12_2(X3,k9_metric_1(X2,X1,X4)),u1_struct_0(X2))|~l1_metric_1(X2)|~m1_subset_1(k9_metric_1(X2,X1,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X2)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_2369, plain, (r2_hidden(esk8_3(X1,X2,X3),X4)|v2_struct_0(X5)|~r2_hidden(X6,X4)|~v3_pre_topc(X4,k3_pcomps_1(X5))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X5))))|~m1_subset_1(X6,u1_struct_0(X5))|~l1_metric_1(X5)|~v9_metric_1(X5)|~v8_metric_1(X5)|~v7_metric_1(X5)|~v6_metric_1(X5)|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,k9_metric_1(X5,X6,esk11_3(X5,X4,X6)))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2464, negated_conjecture, (m1_subset_1(esk14_2(k9_xtuple_0(esk3_0),X1),u1_struct_0(esk1_0))|r1_tarski(k9_xtuple_0(esk3_0),X1)|X2!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_768, plain, (r2_hidden(k1_xboole_0,k1_numbers))).#trainneg
cnf(i_0_2428, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,X1),k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(X3,k6_numbers)|X2!=k1_funct_1(esk3_0,esk8_3(X4,X5,X6))|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,X1),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(esk8_3(X4,X5,X6),u1_struct_0(esk1_0))|X5!=k10_xtuple_0(X4)|~v4_relat_1(X4,k9_metric_1(esk1_0,X1,esk4_3(X3,esk8_3(X4,X5,X6),X2)))|~r2_hidden(X6,X5)|~v1_relat_1(X4)|~v1_funct_1(X4))).#trainneg
cnf(i_0_1971, plain, (m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),X4)|~v4_relat_1(X5,X4)|~v1_relat_1(X5)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k1_relset_1(X4,X5))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2349, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,k8_relat_1(esk3_0,X4))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_714, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|v5_pre_topc(X3,X2,k3_pcomps_1(esk1_0))|v2_struct_0(X2)|~r2_hidden(X1,esk7_3(X2,esk1_0,X3))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1452, plain, (r2_hidden(k1_funct_1(X1,esk12_2(X2,k8_relset_1(k8_relat_1(X1,X3),X4,X5,X6))),X3)|v3_pre_topc(k8_relset_1(k8_relat_1(X1,X3),X4,X5,X6),k3_pcomps_1(X2))|v2_struct_0(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(k8_relset_1(k8_relat_1(X1,X3),X4,X5,X6),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k8_relat_1(X1,X3),X4)))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_2505, negated_conjecture, (m1_subset_1(esk14_2(k2_pcomps_1(esk2_0),X1),k1_zfmisc_1(u1_struct_0(X2)))|r1_tarski(k2_pcomps_1(esk2_0),X1)|X2!=k3_pcomps_1(esk2_0)|~v1_pre_topc(X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1208, plain, (m1_subset_1(esk14_2(k2_pcomps_1(X1),X2),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|r1_tarski(k2_pcomps_1(X1),X2)|~l1_metric_1(X1))).#trainneg
cnf(i_0_2058, plain, (r1_tarski(k8_relat_1(X1,X2),X3)|r2_hidden(k1_funct_1(X1,esk14_2(k8_relat_1(X1,X2),X3)),X4)|X4!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3189, plain, (r1_tarski(X1,k9_metric_1(X2,esk6_3(X3,X2,X4),esk5_3(X3,X2,X4)))|v5_pre_topc(X4,X3,k3_pcomps_1(X2))|v2_struct_0(X3)|v2_struct_0(X2)|~r2_hidden(esk14_2(X1,k9_metric_1(X2,esk6_3(X3,X2,X4),esk5_3(X3,X2,X4))),esk7_3(X3,X2,X4))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X4)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_3036, negated_conjecture, (r2_hidden(esk8_3(esk3_0,X1,X2),u1_struct_0(esk1_0))|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_3475, plain, (v3_pre_topc(X1,k3_pcomps_1(X2))|r1_xxreal_0(X3,k6_numbers)|v2_struct_0(X2)|r2_hidden(esk12_2(X2,X1),u1_struct_0(X2))|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~m1_subset_1(esk12_2(X2,X1),u1_struct_0(X2)))).#trainneg
cnf(i_0_3318, plain, (r2_hidden(X1,k8_relat_1(X2,X3))|~r2_hidden(X1,k9_xtuple_0(X2))|X3!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_3037, negated_conjecture, (v3_pre_topc(u1_struct_0(esk2_0),k3_pcomps_1(esk2_0))|~r2_hidden(esk8_3(esk3_0,X1,X2),X3)|~v3_pre_topc(X3,k3_pcomps_1(esk1_0))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_3120, negated_conjecture, (r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4),X5)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)|r1_xxreal_0(X2,k6_numbers)|X5!=k8_relat_1(esk3_0,k9_metric_1(esk2_0,X3,X2))|~r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4),k9_xtuple_0(esk3_0))|X3!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(X2)|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_3186, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk1_0),X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X1))))).#trainneg
cnf(i_0_3050, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(esk2_0))|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,k10_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3463, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)))),X3)|v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),k3_pcomps_1(X1))|r1_xxreal_0(esk11_3(esk2_0,X3,X4),k6_numbers)|v2_struct_0(X1)|~r2_hidden(X4,X3)|~v3_pre_topc(X3,k3_pcomps_1(esk2_0))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X4,u1_struct_0(esk2_0))|X4!=k1_funct_1(esk3_0,X2)|~v1_xreal_0(esk11_3(esk2_0,X3,X4))|~m1_subset_1(k1_funct_1(esk3_0,esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)))),u1_struct_0(esk2_0))|~m1_subset_1(esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4))),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3118, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,X4,X5),X3,X5)),X6)))),X4)|r1_tarski(k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,X4,X5),X3,X5)),X6)|r1_xxreal_0(esk11_3(esk2_0,X4,X5),k6_numbers)|X2!=k10_xtuple_0(X1)|~r2_hidden(X5,X4)|~r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,X4,X5),X3,X5)),X6)),X2)|~v1_relat_1(X1)|~v3_pre_topc(X4,k3_pcomps_1(esk2_0))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X5,u1_struct_0(esk2_0))|~v1_funct_1(X1)|X5!=k1_funct_1(esk3_0,X3)|~v1_xreal_0(esk11_3(esk2_0,X4,X5))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,X4,X5),X3,X5)),X6)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,X4,X5),X3,X5)),X6),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3057, negated_conjecture, (m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),u1_struct_0(esk2_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k10_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3170, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk2_0),k3_pcomps_1(esk2_0))|~v3_pre_topc(k8_relat_1(X1,esk7_3(k3_pcomps_1(esk2_0),esk2_0,X1)),k3_pcomps_1(esk2_0))|~v2_pre_topc(k3_pcomps_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(esk2_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_3108, negated_conjecture, (m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3142, negated_conjecture, (r1_tarski(k8_relat_1(esk3_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3136, negated_conjecture, (v3_pre_topc(u1_struct_0(esk2_0),k3_pcomps_1(esk2_0))|v3_pre_topc(esk7_3(X2,esk2_0,X3),k3_pcomps_1(X1))|v5_pre_topc(X3,X2,k3_pcomps_1(esk2_0))|v2_struct_0(X2)|v2_struct_0(X1)|~r2_hidden(esk12_2(X1,esk7_3(X2,esk2_0,X3)),X4)|~v3_pre_topc(X4,k3_pcomps_1(esk2_0))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(esk7_3(X2,esk2_0,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3464, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)))),X3)|v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),k3_pcomps_1(X1))|r1_xxreal_0(esk11_3(esk2_0,X3,X4),k6_numbers)|v2_struct_0(X1)|~r2_hidden(X4,X3)|~v3_pre_topc(X3,k3_pcomps_1(esk2_0))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X4,u1_struct_0(esk2_0))|X4!=k1_funct_1(esk3_0,X2)|~v1_xreal_0(esk11_3(esk2_0,X3,X4))|~m1_subset_1(k1_funct_1(esk3_0,esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)))),u1_struct_0(esk2_0))|~m1_subset_1(esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4))),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3117, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)),X2)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X3,u1_struct_0(esk2_0))|X3!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3068, negated_conjecture, (r2_hidden(esk14_2(k10_xtuple_0(esk3_0),X1),u1_struct_0(esk2_0))|r1_tarski(k10_xtuple_0(esk3_0),X1))).#trainneg
cnf(i_0_3131, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk2_0),k3_pcomps_1(esk1_0))|~v3_pre_topc(k8_relat_1(X1,esk7_3(k3_pcomps_1(esk2_0),esk1_0,X1)),k3_pcomps_1(esk2_0))|~v2_pre_topc(k3_pcomps_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_3164, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk1_0),k3_pcomps_1(esk2_0))|~v3_pre_topc(k8_relat_1(X1,esk7_3(k3_pcomps_1(esk1_0),esk2_0,X1)),k3_pcomps_1(esk1_0))|~v2_pre_topc(k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk2_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_3153, negated_conjecture, (r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)|r1_xxreal_0(X2,k6_numbers)|~v1_xboole_0(k9_metric_1(esk2_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)),X2))|X3!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(X2)|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_3465, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,k1_funct_1(esk3_0,esk12_2(X3,k9_metric_1(esk1_0,X4,esk4_3(esk11_3(esk2_0,X5,X6),X4,X6)))))),X5)|v3_pre_topc(k9_metric_1(esk1_0,X4,esk4_3(esk11_3(esk2_0,X5,X6),X4,X6)),k3_pcomps_1(X3))|r1_xxreal_0(esk11_3(esk2_0,X5,X6),k6_numbers)|v2_struct_0(X3)|X2!=k10_xtuple_0(X1)|~r2_hidden(X6,X5)|~r2_hidden(k1_funct_1(esk3_0,esk12_2(X3,k9_metric_1(esk1_0,X4,esk4_3(esk11_3(esk2_0,X5,X6),X4,X6)))),X2)|~v1_relat_1(X1)|~v3_pre_topc(X5,k3_pcomps_1(esk2_0))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X6,u1_struct_0(esk2_0))|~v1_funct_1(X1)|X6!=k1_funct_1(esk3_0,X4)|~v1_xreal_0(esk11_3(esk2_0,X5,X6))|~m1_subset_1(k1_funct_1(esk3_0,esk12_2(X3,k9_metric_1(esk1_0,X4,esk4_3(esk11_3(esk2_0,X5,X6),X4,X6)))),u1_struct_0(esk2_0))|~m1_subset_1(esk12_2(X3,k9_metric_1(esk1_0,X4,esk4_3(esk11_3(esk2_0,X5,X6),X4,X6))),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X4,esk4_3(esk11_3(esk2_0,X5,X6),X4,X6)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_3076, negated_conjecture, (r2_hidden(esk8_3(X1,X2,X3),u1_struct_0(esk2_0))|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,k10_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3160, negated_conjecture, (v5_pre_topc(X1,k3_pcomps_1(esk1_0),k3_pcomps_1(esk1_0))|~v3_pre_topc(k8_relat_1(X1,esk7_3(k3_pcomps_1(esk1_0),esk1_0,X1)),k3_pcomps_1(esk1_0))|~v2_pre_topc(k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_3319, negated_conjecture, (r1_tarski(u1_struct_0(esk1_0),k8_relat_1(esk3_0,X1))|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_3083, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),u1_struct_0(esk2_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k10_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3625, plain, (r1_tarski(k8_relat_1(X1,X2),X3)|v1_xboole_0(X3)|~v4_relat_1(X1,X3)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3258, plain, (m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,esk14_2(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6))),X5)|r1_tarski(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6)|v2_struct_0(X3)|X2!=k10_xtuple_0(X1)|~r2_hidden(esk14_2(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2)|~r2_hidden(X4,X5)|~v1_relat_1(X1)|~v3_pre_topc(X5,k3_pcomps_1(X3))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X4,u1_struct_0(X3))|~v1_funct_1(X1)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_3042, negated_conjecture, (m1_subset_1(esk14_2(k10_xtuple_0(esk3_0),X1),u1_struct_0(esk2_0))|r1_tarski(k10_xtuple_0(esk3_0),X1))).#trainneg
cnf(i_0_3126, negated_conjecture, (v3_pre_topc(u1_struct_0(esk2_0),k3_pcomps_1(esk2_0))|r1_tarski(esk7_3(X1,esk2_0,X2),X3)|v5_pre_topc(X2,X1,k3_pcomps_1(esk2_0))|v2_struct_0(X1)|~r2_hidden(esk14_2(esk7_3(X1,esk2_0,X2),X3),X4)|~v3_pre_topc(X4,k3_pcomps_1(esk2_0))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk2_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_3825, plain, (r2_hidden(esk14_2(k8_relat_1(X1,X2),X3),X4)|r1_tarski(k8_relat_1(X1,X2),X3)|X4!=k8_relat_1(X1,X5)|~r2_hidden(esk14_2(k8_relat_1(X1,X2),X3),k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|X5!=k10_xtuple_0(X1))).#trainneg
cnf(i_0_3097, negated_conjecture, (m1_subset_1(esk12_2(X1,k9_xtuple_0(esk3_0)),u1_struct_0(esk1_0))|v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(X1))|v2_struct_0(X1)|~m1_subset_1(k9_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3468, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk5_3(X3,esk2_0,X4),X2,esk6_3(X3,esk2_0,X4))))),esk7_3(X3,esk2_0,X4))|v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(esk5_3(X3,esk2_0,X4),X2,esk6_3(X3,esk2_0,X4))),k3_pcomps_1(X1))|r1_xxreal_0(esk5_3(X3,esk2_0,X4),k6_numbers)|v2_struct_0(X1)|v5_pre_topc(X4,X3,k3_pcomps_1(esk2_0))|v2_struct_0(X3)|esk6_3(X3,esk2_0,X4)!=k1_funct_1(esk3_0,X2)|~v1_xreal_0(esk5_3(X3,esk2_0,X4))|~m1_subset_1(k1_funct_1(esk3_0,esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk5_3(X3,esk2_0,X4),X2,esk6_3(X3,esk2_0,X4))))),u1_struct_0(esk2_0))|~m1_subset_1(esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk5_3(X3,esk2_0,X4),X2,esk6_3(X3,esk2_0,X4)))),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(esk5_3(X3,esk2_0,X4),X2,esk6_3(X3,esk2_0,X4))),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(esk6_3(X3,esk2_0,X4),u1_struct_0(esk2_0))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk2_0))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(esk2_0))|~v1_funct_1(X4))).#trainneg
cnf(i_0_3686, negated_conjecture, (r2_hidden(esk12_2(X1,esk7_3(X2,esk1_0,X3)),u1_struct_0(esk1_0))|v5_pre_topc(X3,X2,k3_pcomps_1(esk1_0))|v2_struct_0(X2)|v3_pre_topc(esk7_3(X2,esk1_0,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X3)|~m1_subset_1(esk7_3(X2,esk1_0,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3049, negated_conjecture, (m1_subset_1(esk12_2(X1,k10_xtuple_0(esk3_0)),u1_struct_0(esk2_0))|v3_pre_topc(k10_xtuple_0(esk3_0),k3_pcomps_1(X1))|v2_struct_0(X1)|~m1_subset_1(k10_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3089, negated_conjecture, (m1_subset_1(esk14_2(k9_xtuple_0(esk3_0),X1),u1_struct_0(esk1_0))|r1_tarski(k9_xtuple_0(esk3_0),X1))).#trainneg
cnf(i_0_3467, negated_conjecture, (r2_hidden(esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4))),X5)|v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),k3_pcomps_1(X1))|r1_xxreal_0(X3,k6_numbers)|v2_struct_0(X1)|X5!=k8_relat_1(esk3_0,k9_metric_1(esk2_0,X4,X3))|~r2_hidden(esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4))),k9_xtuple_0(esk3_0))|X4!=k1_funct_1(esk3_0,X2)|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)))),u1_struct_0(esk2_0))|~m1_subset_1(esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4))),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3646, negated_conjecture, (m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3100, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3072, negated_conjecture, (r2_hidden(esk12_2(X1,k8_relset_1(k10_xtuple_0(esk3_0),X2,X3,X4)),u1_struct_0(esk2_0))|v3_pre_topc(k8_relset_1(k10_xtuple_0(esk3_0),X2,X3,X4),k3_pcomps_1(X1))|v2_struct_0(X1)|~m1_subset_1(k8_relset_1(k10_xtuple_0(esk3_0),X2,X3,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k10_xtuple_0(esk3_0),X2)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3121, negated_conjecture, (r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk5_3(X2,esk2_0,X3),X1,esk6_3(X2,esk2_0,X3))),X4)|r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk5_3(X2,esk2_0,X3),X1,esk6_3(X2,esk2_0,X3))),X4)),esk7_3(X2,esk2_0,X3))|r1_xxreal_0(esk5_3(X2,esk2_0,X3),k6_numbers)|v5_pre_topc(X3,X2,k3_pcomps_1(esk2_0))|v2_struct_0(X2)|esk6_3(X2,esk2_0,X3)!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(esk5_3(X2,esk2_0,X3))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk5_3(X2,esk2_0,X3),X1,esk6_3(X2,esk2_0,X3))),X4)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk5_3(X2,esk2_0,X3),X1,esk6_3(X2,esk2_0,X3))),X4),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(esk6_3(X2,esk2_0,X3),u1_struct_0(esk2_0))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3))).#trainneg
cnf(i_0_3112, negated_conjecture, (r1_tarski(k9_metric_1(esk1_0,X3,esk4_3(X2,X3,X1)),X4)|r1_xxreal_0(X2,k6_numbers)|~v1_xboole_0(k9_metric_1(esk2_0,X1,X2))|X1!=k1_funct_1(esk3_0,X3)|~v1_xreal_0(X2)|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X3,esk4_3(X2,X3,X1)),X4)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X3,esk4_3(X2,X3,X1)),X4),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_3268, negated_conjecture, (r2_hidden(esk14_2(k9_xtuple_0(esk3_0),X1),u1_struct_0(esk1_0))|r1_tarski(k9_xtuple_0(esk3_0),X1))).#trainneg
cnf(i_0_3288, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3277, negated_conjecture, (r2_hidden(esk12_2(X1,k8_relat_1(esk3_0,X2)),u1_struct_0(esk1_0))|v3_pre_topc(k8_relat_1(esk3_0,X2),k3_pcomps_1(X1))|v2_struct_0(X1)|~m1_subset_1(k8_relat_1(esk3_0,X2),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3082, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk12_2(X2,k8_relat_1(X1,X3))),u1_struct_0(esk2_0))|v3_pre_topc(k8_relat_1(X1,X3),k3_pcomps_1(X2))|v2_struct_0(X2)|k10_xtuple_0(esk3_0)!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~m1_subset_1(k8_relat_1(X1,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_3488, negated_conjecture, (r2_hidden(X1,X2)|v3_pre_topc(k1_relset_1(u1_struct_0(esk1_0),X3),k3_pcomps_1(X4))|v2_struct_0(X4)|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k1_funct_1(esk3_0,esk12_2(X4,k1_relset_1(u1_struct_0(esk1_0),X3)))|X2!=k10_xtuple_0(esk3_0)|~v4_relat_1(X3,u1_struct_0(esk1_0))|~v1_relat_1(X3)|~m1_subset_1(k1_relset_1(u1_struct_0(esk1_0),X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~l1_metric_1(X4)|~v9_metric_1(X4)|~v8_metric_1(X4)|~v7_metric_1(X4)|~v6_metric_1(X4))).#trainneg
cnf(i_0_3459, negated_conjecture, (v3_pre_topc(k9_metric_1(esk1_0,X4,esk4_3(X2,X4,X1)),k3_pcomps_1(X3))|r1_xxreal_0(X2,k6_numbers)|v2_struct_0(X3)|~v1_xboole_0(k9_metric_1(esk2_0,X1,X2))|X1!=k1_funct_1(esk3_0,X4)|~v1_xreal_0(X2)|~m1_subset_1(k1_funct_1(esk3_0,esk12_2(X3,k9_metric_1(esk1_0,X4,esk4_3(X2,X4,X1)))),u1_struct_0(esk2_0))|~m1_subset_1(esk12_2(X3,k9_metric_1(esk1_0,X4,esk4_3(X2,X4,X1))),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X4,esk4_3(X2,X4,X1)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_3075, negated_conjecture, (r2_hidden(esk12_2(X1,k10_xtuple_0(esk3_0)),u1_struct_0(esk2_0))|v3_pre_topc(k10_xtuple_0(esk3_0),k3_pcomps_1(X1))|v2_struct_0(X1)|~m1_subset_1(k10_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3046, negated_conjecture, (m1_subset_1(esk12_2(X1,k8_relset_1(k10_xtuple_0(esk3_0),X2,X3,X4)),u1_struct_0(esk2_0))|v3_pre_topc(k8_relset_1(k10_xtuple_0(esk3_0),X2,X3,X4),k3_pcomps_1(X1))|v2_struct_0(X1)|~m1_subset_1(k8_relset_1(k10_xtuple_0(esk3_0),X2,X3,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k10_xtuple_0(esk3_0),X2)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3296, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk14_2(k8_relat_1(esk3_0,X1),X2),X3)|r1_tarski(k8_relat_1(esk3_0,X1),X2)|~v4_relat_1(esk3_0,X3))).#trainneg
cnf(i_0_3261, plain, (r2_hidden(esk8_3(X1,X2,esk14_2(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6)),X7)|r1_tarski(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6)|v2_struct_0(X3)|X7!=k8_relat_1(X1,X5)|~r2_hidden(esk8_3(X1,X2,esk14_2(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6)),k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(esk14_2(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2)|~r2_hidden(X4,X5)|~v3_pre_topc(X5,k3_pcomps_1(X3))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_3081, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk12_2(X2,k9_xtuple_0(X1))),u1_struct_0(esk2_0))|v3_pre_topc(k9_xtuple_0(X1),k3_pcomps_1(X2))|v2_struct_0(X2)|k10_xtuple_0(esk3_0)!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~m1_subset_1(k9_xtuple_0(X1),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_3317, plain, (m1_subset_1(k9_xtuple_0(X1),k1_zfmisc_1(k8_relat_1(X1,X2)))|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3094, negated_conjecture, (m1_subset_1(esk12_2(X1,k8_relset_1(k9_xtuple_0(esk3_0),X2,X3,X4)),u1_struct_0(esk1_0))|v3_pre_topc(k8_relset_1(k9_xtuple_0(esk3_0),X2,X3,X4),k3_pcomps_1(X1))|v2_struct_0(X1)|~m1_subset_1(k8_relset_1(k9_xtuple_0(esk3_0),X2,X3,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k9_xtuple_0(esk3_0),X2)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3481, plain, (v3_pre_topc(esk7_3(X3,X1,X4),k3_pcomps_1(X2))|v5_pre_topc(X4,X3,k3_pcomps_1(X1))|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|~v1_xboole_0(k9_metric_1(X1,esk12_2(X2,esk7_3(X3,X1,X4)),esk5_3(X3,X1,X4)))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k3_pcomps_1(X1)))))|~m1_subset_1(esk7_3(X3,X1,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X4)|~l1_metric_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3766, negated_conjecture, (r1_tarski(k2_pcomps_1(esk2_0),X1)|m1_subset_1(esk14_2(k2_pcomps_1(esk2_0),X1),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_3886, plain, (r2_hidden(k1_funct_1(X1,X2),X3)|~v1_relat_1(X1)|~v1_funct_1(X1)|X3!=k10_xtuple_0(X1)|~r2_hidden(X2,k9_xtuple_0(X1)))).#trainneg
cnf(i_0_3933, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,X1),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3832, plain, (r1_tarski(esk7_3(X1,X2,X3),k9_metric_1(X2,esk6_3(X1,X2,X3),esk5_3(X1,X2,X3)))|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_3967, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k8_relat_1(esk3_0,X3)|~r2_hidden(X1,k9_xtuple_0(esk3_0))|X3!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3837, negated_conjecture, (r2_hidden(esk12_2(esk2_0,X1),u1_struct_0(esk2_0))|v3_pre_topc(X1,k3_pcomps_1(esk2_0))|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_3838, negated_conjecture, (r2_hidden(esk12_2(esk1_0,X1),u1_struct_0(esk1_0))|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_4299, plain, (r1_tarski(k8_relat_1(X1,X2),X3)|r2_hidden(esk14_2(k8_relat_1(X1,X2),X3),k8_relat_1(X1,X4))|X4!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4250, plain, (m1_subset_1(k8_relat_1(X1,X2),k1_zfmisc_1(X3))|v1_xboole_0(X3)|~v4_relat_1(X1,X3)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4248, negated_conjecture, (r2_hidden(esk8_3(X1,X2,X3),X4)|X4!=k8_relat_1(X1,u1_struct_0(esk2_0))|~r2_hidden(esk8_3(X1,X2,X3),k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k10_xtuple_0(esk3_0))|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_4367, negated_conjecture, (r2_hidden(esk8_3(X1,X2,X3),X4)|X4!=k8_relat_1(X1,u1_struct_0(esk1_0))|~r2_hidden(esk8_3(X1,X2,X3),k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_4178, negated_conjecture, (r1_tarski(k10_xtuple_0(esk3_0),u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_4558, negated_conjecture, (r2_hidden(esk14_2(k9_metric_1(X1,X2,esk11_3(X1,X3,X2)),X4),k9_metric_1(esk2_0,X5,X6))|r1_xxreal_0(X6,k6_numbers)|r1_tarski(k9_metric_1(X1,X2,esk11_3(X1,X3,X2)),X4)|v2_struct_0(X1)|X5!=k1_funct_1(esk3_0,X7)|X8!=k10_xtuple_0(esk3_0)|~r2_hidden(esk14_2(k9_metric_1(X1,X2,esk11_3(X1,X3,X2)),X4),X8)|~v1_xreal_0(X6)|~m1_subset_1(esk8_3(esk3_0,X8,esk14_2(k9_metric_1(X1,X2,esk11_3(X1,X3,X2)),X4)),u1_struct_0(esk1_0))|~m1_subset_1(X7,u1_struct_0(esk1_0))|~m1_subset_1(X5,u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(X1,X2,esk11_3(X1,X3,X2)),X4),u1_struct_0(esk2_0))|k9_metric_1(esk1_0,X7,esk4_3(X6,X7,X5))!=k8_relat_1(esk3_0,X3)|~r2_hidden(X2,X3)|~v3_pre_topc(X3,k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_4361, negated_conjecture, (r1_tarski(k9_xtuple_0(esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4244, negated_conjecture, (r2_hidden(X1,k8_relat_1(esk3_0,X2))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_3981, negated_conjecture, (r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)|r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4),X5)|r1_xxreal_0(X2,k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|X5!=k8_relat_1(esk3_0,k9_metric_1(esk2_0,X3,X2))|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4),u1_struct_0(esk1_0))|~v1_xreal_0(X2)|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_3996, negated_conjecture, (r2_hidden(k1_funct_1(X1,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,X3),X4),X2,X4)),X5))),X3)|r1_tarski(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,X3),X4),X2,X4)),X5)|r1_xxreal_0(esk11_3(esk2_0,k8_relat_1(X1,X3),X4),k6_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|X4!=k1_funct_1(esk3_0,X2)|~r2_hidden(X4,k8_relat_1(X1,X3))|~v3_pre_topc(k8_relat_1(X1,X3),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k8_relat_1(X1,X3),X4))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,X3),X4),X2,X4)),X5)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,X3),X4),X2,X4)),X5),u1_struct_0(esk1_0))|~m1_subset_1(k8_relat_1(X1,X3),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4605, plain, (k8_relat_1(k1_xboole_0,X1)=k1_xboole_0)).#trainneg
cnf(i_0_4398, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X4,k3_pcomps_1(esk1_0))|X1!=k1_funct_1(X3,esk12_2(esk1_0,X4))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v4_relat_1(esk3_0,k9_xtuple_0(X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3963, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,X1),X2)|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4251, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,k8_relat_1(X3,X4))|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~v1_funct_1(X3))).#trainneg
cnf(i_0_3978, negated_conjecture, (v3_pre_topc(u1_struct_0(esk2_0),k3_pcomps_1(esk2_0))|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1)|~v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(esk1_0)))).#trainneg
cnf(i_0_3875, plain, (r1_tarski(X1,k8_relat_1(X2,X3))|X3!=k10_xtuple_0(X2)|~r2_hidden(esk14_2(X1,k8_relat_1(X2,X3)),k9_xtuple_0(X2))|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_3945, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,X1),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|u1_struct_0(esk2_0)=k1_xboole_0|~l1_metric_1(X2)|k2_pcomps_1(X2)!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4532, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,esk14_2(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6)),k1_zfmisc_1(u1_struct_0(esk2_0)))|r1_tarski(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6)|v2_struct_0(X3)|k2_pcomps_1(esk2_0)!=k8_relat_1(X1,X5)|X2!=k10_xtuple_0(X1)|~r2_hidden(esk14_2(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2)|~r2_hidden(X4,X5)|~v1_relat_1(X1)|~v3_pre_topc(X5,k3_pcomps_1(X3))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X4,u1_struct_0(X3))|~v1_funct_1(X1)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_3921, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,X1),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|k8_relat_1(esk3_0,X2)!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4315, negated_conjecture, (m1_subset_1(esk12_2(X1,esk7_3(X2,esk1_0,X3)),u1_struct_0(esk1_0))|v3_pre_topc(esk7_3(X2,esk1_0,X3),k3_pcomps_1(X1))|v5_pre_topc(X3,X2,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(esk7_3(X2,esk1_0,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X3)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_4503, negated_conjecture, (r1_tarski(k8_relat_1(esk3_0,X1),X2)|u1_struct_0(esk2_0)=k1_xboole_0|~v4_relat_1(esk3_0,X2))).#trainneg
cnf(i_0_4541, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(X3,k6_numbers)|r1_tarski(k9_metric_1(X6,X7,esk11_3(X6,X8,X7)),X9)|v2_struct_0(X6)|X2!=k1_funct_1(esk3_0,esk8_3(X4,X5,esk14_2(k9_metric_1(X6,X7,esk11_3(X6,X8,X7)),X9)))|X10!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,X10)|~v1_xreal_0(X3)|~m1_subset_1(esk8_3(esk3_0,X10,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(esk8_3(X4,X5,esk14_2(k9_metric_1(X6,X7,esk11_3(X6,X8,X7)),X9)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|k9_metric_1(esk1_0,esk8_3(esk3_0,X10,X1),esk4_3(X3,esk8_3(X4,X5,esk14_2(k9_metric_1(X6,X7,esk11_3(X6,X8,X7)),X9)),X2))!=k8_relat_1(X4,X8)|X5!=k10_xtuple_0(X4)|~r2_hidden(esk14_2(k9_metric_1(X6,X7,esk11_3(X6,X8,X7)),X9),X5)|~r2_hidden(X7,X8)|~v1_relat_1(X4)|~v3_pre_topc(X8,k3_pcomps_1(X6))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X6))))|~m1_subset_1(X7,u1_struct_0(X6))|~v1_funct_1(X4)|~l1_metric_1(X6)|~v9_metric_1(X6)|~v8_metric_1(X6)|~v7_metric_1(X6)|~v6_metric_1(X6))).#trainneg
cnf(i_0_4366, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),X4)|~v4_relat_1(esk3_0,X4)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4534, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,esk14_2(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6)),k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_tarski(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6)|v2_struct_0(X3)|k2_pcomps_1(esk1_0)!=k8_relat_1(X1,X5)|X2!=k10_xtuple_0(X1)|~r2_hidden(esk14_2(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2)|~r2_hidden(X4,X5)|~v1_relat_1(X1)|~v3_pre_topc(X5,k3_pcomps_1(X3))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X4,u1_struct_0(X3))|~v1_funct_1(X1)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_4432, negated_conjecture, (m1_subset_1(esk12_2(esk1_0,X1),X2)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|~v4_relat_1(esk3_0,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_4026, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5),X1,X5)),X6)),X4)|v2_struct_0(X2)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5),X1,X5)),X6)|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5),k6_numbers)|~r2_hidden(X3,X4)|~v3_pre_topc(X4,k3_pcomps_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|X5!=k1_funct_1(esk3_0,X1)|~r2_hidden(X5,k9_metric_1(X2,X3,esk11_3(X2,X4,X3)))|~v3_pre_topc(k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5),X1,X5)),X6)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5),X1,X5)),X6),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4401, negated_conjecture, (m1_subset_1(esk12_2(esk1_0,X1),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|~v4_relat_1(esk3_0,k10_xtuple_0(esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3053, negated_conjecture, (m1_subset_1(k1_funct_1(X1,esk14_2(k9_xtuple_0(X1),X2)),u1_struct_0(esk2_0))|r1_tarski(k9_xtuple_0(X1),X2)|k10_xtuple_0(esk3_0)!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3647, plain, (m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),X4)|X2!=k10_xtuple_0(X1)|~v4_relat_1(X5,X4)|~r2_hidden(X3,k9_xtuple_0(X5))|~r2_hidden(X3,X2)|~v1_relat_1(X5)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4755, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X2,u1_struct_0(esk2_0))|~r2_hidden(esk8_3(esk3_0,X1,X2),u1_struct_0(esk1_0))|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_4625, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|m1_subset_1(X2,u1_struct_0(esk2_0))|~r2_hidden(esk8_3(esk3_0,X1,X2),u1_struct_0(esk1_0))|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_4761, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,X2)|X2!=k8_relat_1(esk3_0,X3)|X3!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4882, plain, (r1_tarski(k8_relat_1(X1,X2),k8_relat_1(X1,X3))|X3!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4756, plain, (m1_subset_1(esk7_3(X1,X2,X3),k1_zfmisc_1(k9_metric_1(X2,esk6_3(X1,X2,X3),esk5_3(X1,X2,X3))))|v5_pre_topc(X3,X1,k3_pcomps_1(X2))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k3_pcomps_1(X2)))|~v1_funct_1(X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_4941, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(esk2_0)))|k2_pcomps_1(esk2_0)!=k8_relat_1(X1,u1_struct_0(esk2_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k10_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5000, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(esk2_0)))|k2_pcomps_1(esk2_0)!=k8_relat_1(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4782, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X2),X1,X2)),X3)),X4)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X2),X1,X2)),X3)|r1_xxreal_0(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X2),k6_numbers)|X4!=k8_relat_1(esk3_0,X5)|X5!=k10_xtuple_0(esk3_0)|X2!=k1_funct_1(esk3_0,X1)|~r2_hidden(X2,k9_xtuple_0(esk3_0))|~v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X2))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X2),X1,X2)),X3)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X2),X1,X2)),X3),u1_struct_0(esk1_0))|~m1_subset_1(k9_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4621, negated_conjecture, (r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)),X2)|u1_struct_0(esk2_0)=k1_xboole_0|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5002, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(esk1_0)))|k2_pcomps_1(esk1_0)!=k8_relat_1(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4943, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(esk1_0)))|k2_pcomps_1(esk1_0)!=k8_relat_1(X1,u1_struct_0(esk2_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k10_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4781, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(esk3_0,X1,X2),X3)|X3!=k8_relat_1(esk3_0,X4)|X4!=k10_xtuple_0(esk3_0)|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_4767, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk14_2(k9_xtuple_0(esk3_0),X1),X2)|r1_tarski(k9_xtuple_0(esk3_0),X1)|X2!=k8_relat_1(esk3_0,X3)|X3!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_5048, negated_conjecture, (r1_tarski(X1,k8_relat_1(esk3_0,X2))|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(esk14_2(X1,k8_relat_1(esk3_0,X2)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4658, plain, (r1_tarski(k9_metric_1(X1,X2,X3),X4)|r2_hidden(X2,k1_xboole_0)|v2_struct_0(X1)|r2_hidden(esk14_2(k9_metric_1(X1,X2,X3),X4),u1_struct_0(X1))|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(esk14_2(k9_metric_1(X1,X2,X3),X4),u1_struct_0(X1)))).#trainneg
cnf(i_0_5058, negated_conjecture, (r2_hidden(X1,k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5053, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k8_relat_1(X3,k8_relat_1(esk3_0,X4))|~r2_hidden(X1,k9_xtuple_0(X3))|~v1_relat_1(X3)|~v1_funct_1(X3)|X4!=k10_xtuple_0(esk3_0)|~r2_hidden(k1_funct_1(X3,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5057, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|X4!=k10_xtuple_0(esk3_0)|~r2_hidden(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4768, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk14_2(k8_relat_1(esk3_0,X1),X2),X3)|r1_tarski(k8_relat_1(esk3_0,X1),X2)|X3!=k8_relat_1(esk3_0,X4)|X4!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_5181, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|m1_subset_1(k1_funct_1(esk3_0,X1),k10_xtuple_0(esk3_0))|~r2_hidden(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4795, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),X4)|X4!=k8_relat_1(esk3_0,X5)|X5!=k10_xtuple_0(esk3_0)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4950, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(X3,k6_numbers)|X2!=k1_funct_1(esk3_0,esk8_3(X4,X5,X6))|X7!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,X7)|~v1_xreal_0(X3)|~m1_subset_1(esk8_3(esk3_0,X7,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(esk8_3(X4,X5,X6),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|k9_metric_1(esk1_0,esk8_3(esk3_0,X7,X1),esk4_3(X3,esk8_3(X4,X5,X6),X2))!=k8_relat_1(X4,u1_struct_0(esk2_0))|X5!=k10_xtuple_0(X4)|~r2_hidden(X6,k10_xtuple_0(esk3_0))|~r2_hidden(X6,X5)|~v1_relat_1(X4)|~v1_funct_1(X4))).#trainneg
cnf(i_0_5009, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(X3,k6_numbers)|X2!=k1_funct_1(esk3_0,esk8_3(X4,X5,X6))|X7!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,X7)|~v1_xreal_0(X3)|~m1_subset_1(esk8_3(esk3_0,X7,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(esk8_3(X4,X5,X6),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|k9_metric_1(esk1_0,esk8_3(esk3_0,X7,X1),esk4_3(X3,esk8_3(X4,X5,X6),X2))!=k8_relat_1(X4,u1_struct_0(esk1_0))|X5!=k10_xtuple_0(X4)|~r2_hidden(X6,k9_xtuple_0(esk3_0))|~r2_hidden(X6,X5)|~v1_relat_1(X4)|~v1_funct_1(X4))).#trainneg
cnf(i_0_5400, plain, (m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,esk8_3(X3,X4,X5))),X6)|X2!=k10_xtuple_0(X1)|~v4_relat_1(X3,X6)|~r2_hidden(esk8_3(X3,X4,X5),X2)|~v1_relat_1(X3)|~v1_relat_1(X1)|~v1_funct_1(X1)|X4!=k10_xtuple_0(X3)|~r2_hidden(X5,X4)|~v1_funct_1(X3))).#trainneg
cnf(i_0_5051, negated_conjecture, (v3_pre_topc(u1_struct_0(esk2_0),k3_pcomps_1(esk2_0))|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1)|~v3_pre_topc(k8_relat_1(esk3_0,X3),k3_pcomps_1(esk1_0))|X3!=k10_xtuple_0(esk3_0)|~r2_hidden(esk8_3(esk3_0,X1,X2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4860, negated_conjecture, (r2_hidden(esk12_2(esk1_0,k9_xtuple_0(esk3_0)),u1_struct_0(esk1_0))|v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(esk1_0))|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4682, plain, (r1_tarski(k1_xboole_0,X3)|v2_struct_0(X1)|r2_hidden(X2,u1_struct_0(X1))|~r2_hidden(X2,X3)|~v3_pre_topc(X3,k3_pcomps_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v1_xreal_0(esk11_3(X1,X3,X2)))).#trainneg
cnf(i_0_4777, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk12_2(X1,k8_relat_1(esk3_0,X2)),X3)|v3_pre_topc(k8_relat_1(esk3_0,X2),k3_pcomps_1(X1))|v2_struct_0(X1)|X3!=k8_relat_1(esk3_0,X4)|X4!=k10_xtuple_0(esk3_0)|~m1_subset_1(k8_relat_1(esk3_0,X2),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_4624, negated_conjecture, (v3_pre_topc(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),k3_pcomps_1(X4))|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|m1_subset_1(k1_funct_1(esk3_0,esk12_2(X4,k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)))),X2)|v2_struct_0(X4)|u1_struct_0(esk2_0)=k1_xboole_0|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(esk12_2(X4,k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3))),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~l1_metric_1(X4)|~v9_metric_1(X4)|~v8_metric_1(X4)|~v7_metric_1(X4)|~v6_metric_1(X4)|~r2_hidden(esk12_2(X4,k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3))),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4927, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),X4)|X4!=k8_relat_1(esk3_0,X5)|X5!=k10_xtuple_0(esk3_0)|k8_relat_1(X1,u1_struct_0(esk2_0))!=k9_xtuple_0(esk3_0)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k10_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4964, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,X3),X4)|X4!=k8_relat_1(X1,u1_struct_0(esk2_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k10_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4824, negated_conjecture, (r2_hidden(esk12_2(esk2_0,k10_xtuple_0(esk3_0)),u1_struct_0(esk2_0))|v3_pre_topc(k10_xtuple_0(esk3_0),k3_pcomps_1(esk2_0))|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4925, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(esk1_0))|k8_relat_1(X1,u1_struct_0(esk2_0))!=k9_xtuple_0(esk3_0)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k10_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5087, negated_conjecture, (r2_hidden(k1_funct_1(X1,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,k10_xtuple_0(esk3_0)),X3),X2,X3)),X4))),u1_struct_0(esk2_0))|r1_tarski(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,k10_xtuple_0(esk3_0)),X3),X2,X3)),X4)|r1_xxreal_0(esk11_3(esk2_0,k8_relat_1(X1,k10_xtuple_0(esk3_0)),X3),k6_numbers)|X3!=k1_funct_1(esk3_0,X2)|~r2_hidden(X3,k8_relat_1(X1,k10_xtuple_0(esk3_0)))|~v1_relat_1(X1)|~v3_pre_topc(k8_relat_1(X1,k10_xtuple_0(esk3_0)),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k8_relat_1(X1,k10_xtuple_0(esk3_0)),X3))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,k10_xtuple_0(esk3_0)),X3),X2,X3)),X4)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,k10_xtuple_0(esk3_0)),X3),X2,X3)),X4),u1_struct_0(esk1_0))|~m1_subset_1(k8_relat_1(X1,k10_xtuple_0(esk3_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_4783, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(esk3_0,esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X3),X2,X3)))),X4)|v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X3),X2,X3)),k3_pcomps_1(X1))|r1_xxreal_0(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X3),k6_numbers)|v2_struct_0(X1)|X4!=k8_relat_1(esk3_0,X5)|X5!=k10_xtuple_0(esk3_0)|X3!=k1_funct_1(esk3_0,X2)|~r2_hidden(X3,k9_xtuple_0(esk3_0))|~v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X3))|~m1_subset_1(k1_funct_1(esk3_0,esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X3),X2,X3)))),u1_struct_0(esk2_0))|~m1_subset_1(esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X3),X2,X3))),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X3),X2,X3)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(k9_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_4779, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),X4)|X4!=k8_relat_1(esk3_0,X5)|X5!=k10_xtuple_0(esk3_0)|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4797, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X4),X3,X4)),X5)))),X6)|r1_tarski(k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X4),X3,X4)),X5)|r1_xxreal_0(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X4),k6_numbers)|X6!=k8_relat_1(esk3_0,X7)|X7!=k10_xtuple_0(esk3_0)|X4!=k1_funct_1(esk3_0,X3)|X2!=k10_xtuple_0(X1)|~r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X4),X3,X4)),X5)),X2)|~r2_hidden(X4,k9_xtuple_0(esk3_0))|~v1_relat_1(X1)|~v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X4))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X4),X3,X4)),X5)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X4),X3,X4)),X5),u1_struct_0(esk1_0))|~m1_subset_1(k9_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_5023, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,X3),X4)|X4!=k8_relat_1(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4663, plain, (r2_hidden(X1,k1_xboole_0)|v3_pre_topc(k9_metric_1(X2,X1,X4),k3_pcomps_1(X3))|v2_struct_0(X3)|v2_struct_0(X2)|r2_hidden(esk12_2(X3,k9_metric_1(X2,X1,X4)),u1_struct_0(X2))|~v1_xreal_0(X4)|~m1_subset_1(k9_metric_1(X2,X1,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_metric_1(X2)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X2)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)|~m1_subset_1(esk12_2(X3,k9_metric_1(X2,X1,X4)),u1_struct_0(X2)))).#trainneg
cnf(i_0_4920, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(esk1_0))|k8_relat_1(esk3_0,X4)!=k8_relat_1(X1,u1_struct_0(esk2_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k10_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4626, negated_conjecture, (r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)|r1_xxreal_0(X2,k6_numbers)|r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)),u1_struct_0(esk2_0))|X3!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(X2)|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_3485, plain, (r2_hidden(esk12_2(X1,esk7_3(X2,X3,X4)),k9_metric_1(X3,esk6_3(X2,X3,X4),esk5_3(X2,X3,X4)))|v2_struct_0(X3)|v3_pre_topc(esk7_3(X2,X3,X4),k3_pcomps_1(X1))|v5_pre_topc(X4,X2,k3_pcomps_1(X3))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_xreal_0(esk5_3(X2,X3,X4))|~m1_subset_1(esk12_2(X1,esk7_3(X2,X3,X4)),u1_struct_0(X3))|~l1_metric_1(X3)|~v8_metric_1(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X3)))))|~m1_subset_1(esk7_3(X2,X3,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X3)))|~v1_funct_1(X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v9_metric_1(X3)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v7_metric_1(X3)|~v6_metric_1(X1)|~v6_metric_1(X3))).#trainneg
cnf(i_0_3284, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk12_2(X2,k8_relat_1(X1,k9_xtuple_0(esk3_0)))),u1_struct_0(esk1_0))|v3_pre_topc(k8_relat_1(X1,k9_xtuple_0(esk3_0)),k3_pcomps_1(X2))|v2_struct_0(X2)|~v1_relat_1(X1)|~m1_subset_1(k8_relat_1(X1,k9_xtuple_0(esk3_0)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_3630, negated_conjecture, (r1_tarski(k9_metric_1(X1,X2,esk11_3(X1,esk7_3(X3,esk2_0,X4),X2)),u1_struct_0(esk2_0))|v5_pre_topc(X4,X3,k3_pcomps_1(esk2_0))|v2_struct_0(X1)|v2_struct_0(X3)|~r2_hidden(X2,esk7_3(X3,esk2_0,X4))|~v3_pre_topc(esk7_3(X3,esk2_0,X4),k3_pcomps_1(X1))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(esk7_3(X3,esk2_0,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk2_0))))|~m1_subset_1(X2,u1_struct_0(X1))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(esk2_0))|~v1_funct_1(X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_4025, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5),X1,X5)),X6)),X4)|v2_struct_0(X2)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5),X1,X5)),X6)|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5),k6_numbers)|~r2_hidden(X3,X4)|~v3_pre_topc(X4,k3_pcomps_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|X5!=k1_funct_1(esk3_0,X1)|~r2_hidden(X5,k9_metric_1(X2,X3,esk11_3(X2,X4,X3)))|~v3_pre_topc(k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5),X1,X5)),X6)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5),X1,X5)),X6),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5393, plain, (m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,esk12_2(X3,k9_xtuple_0(X4)))),X5)|v3_pre_topc(k9_xtuple_0(X4),k3_pcomps_1(X3))|v2_struct_0(X3)|X2!=k10_xtuple_0(X1)|~v4_relat_1(X4,X5)|~r2_hidden(esk12_2(X3,k9_xtuple_0(X4)),X2)|~v1_relat_1(X4)|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(k9_xtuple_0(X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_4378, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,X2)|v3_pre_topc(k9_xtuple_0(X3),k3_pcomps_1(X4))|v2_struct_0(X4)|X1!=k1_funct_1(esk3_0,esk12_2(X4,k9_xtuple_0(X3)))|X2!=k10_xtuple_0(esk3_0)|~v4_relat_1(X3,u1_struct_0(esk1_0))|~v1_relat_1(X3)|~m1_subset_1(k9_xtuple_0(X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~l1_metric_1(X4)|~v9_metric_1(X4)|~v8_metric_1(X4)|~v7_metric_1(X4)|~v6_metric_1(X4))).#trainneg
cnf(i_0_4709, plain, (r2_hidden(X1,k9_metric_1(X2,X3,X4))|v2_struct_0(X2)|r2_hidden(X1,u1_struct_0(X2))|~r2_hidden(X3,k1_xboole_0)|~v1_xreal_0(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_metric_1(X2)|~v8_metric_1(X2))).#trainneg
cnf(i_0_4603, plain, (esk8_3(X1,X2,X3)=k1_xboole_0|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,k1_zfmisc_1(k1_xboole_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5113, negated_conjecture, (m1_subset_1(k1_funct_1(X1,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,X3),X4),X2,X4)),X5))),X3)|r1_tarski(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,X3),X4),X2,X4)),X5)|r1_xxreal_0(esk11_3(esk2_0,k8_relat_1(X1,X3),X4),k6_numbers)|X4!=k1_funct_1(esk3_0,X2)|~r2_hidden(X4,k8_relat_1(X1,X3))|~v1_relat_1(X1)|~v3_pre_topc(k8_relat_1(X1,X3),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k8_relat_1(X1,X3),X4))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,X3),X4),X2,X4)),X5)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,X3),X4),X2,X4)),X5),u1_struct_0(esk1_0))|~m1_subset_1(k8_relat_1(X1,X3),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_3276, negated_conjecture, (r2_hidden(esk12_2(X1,k9_xtuple_0(esk3_0)),u1_struct_0(esk1_0))|v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(X1))|v2_struct_0(X1)|~m1_subset_1(k9_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_5231, plain, (r1_tarski(k8_relset_1(k9_xtuple_0(X1),X2,X3,X4),k8_relat_1(X1,X5))|X5!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k9_xtuple_0(X1),X2))))).#trainneg
cnf(i_0_4951, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(X3,k6_numbers)|X1!=k1_funct_1(esk3_0,esk8_3(X4,X5,X6))|X7!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X7)|~v1_xreal_0(X3)|~m1_subset_1(esk8_3(esk3_0,X7,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(esk8_3(X4,X5,X6),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|k9_metric_1(esk1_0,esk8_3(esk3_0,X7,X2),esk4_3(X3,esk8_3(X4,X5,X6),X1))!=k8_relat_1(X4,u1_struct_0(esk2_0))|X5!=k10_xtuple_0(X4)|~r2_hidden(X6,k10_xtuple_0(esk3_0))|~r2_hidden(X6,X5)|~v1_relat_1(X4)|~v1_funct_1(X4))).#trainneg
cnf(i_0_3195, plain, (v5_pre_topc(X3,X2,k3_pcomps_1(X1))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_xboole_0(k9_metric_1(X1,esk6_3(X2,X1,X3),esk5_3(X2,X1,X3)))|~r2_hidden(X4,esk7_3(X2,X1,X3))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X1)))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X3)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3062, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(esk3_0,esk12_2(X1,k8_relset_1(u1_struct_0(esk1_0),X2,X3,X4))),X5)|v3_pre_topc(k8_relset_1(u1_struct_0(esk1_0),X2,X3,X4),k3_pcomps_1(X1))|v2_struct_0(X1)|X5!=k10_xtuple_0(esk3_0)|~m1_subset_1(k8_relset_1(u1_struct_0(esk1_0),X2,X3,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X2)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3080, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk12_2(X2,k8_relat_1(X1,k10_xtuple_0(esk3_0)))),u1_struct_0(esk2_0))|v3_pre_topc(k8_relat_1(X1,k10_xtuple_0(esk3_0)),k3_pcomps_1(X2))|v2_struct_0(X2)|~v1_relat_1(X1)|~m1_subset_1(k8_relat_1(X1,k10_xtuple_0(esk3_0)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_4986, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),X4)|X4!=k8_relat_1(esk3_0,X5)|X5!=k10_xtuple_0(esk3_0)|k8_relat_1(X1,u1_struct_0(esk1_0))!=k9_xtuple_0(esk3_0)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4094, negated_conjecture, (r2_hidden(esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4))),X5)|v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),k3_pcomps_1(X1))|r1_xxreal_0(esk11_3(esk2_0,X3,X4),k6_numbers)|v2_struct_0(X1)|X5!=k8_relat_1(esk3_0,X3)|~r2_hidden(esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4))),k9_xtuple_0(esk3_0))|X4!=k1_funct_1(esk3_0,X2)|~r2_hidden(X4,X3)|~v3_pre_topc(X3,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X3,X4))|~m1_subset_1(k1_funct_1(esk3_0,esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)))),u1_struct_0(esk2_0))|~m1_subset_1(esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4))),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3949, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(X3,k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k1_funct_1(esk3_0,k1_funct_1(esk3_0,X4))|X5!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,X5)|~v1_xreal_0(X3)|~m1_subset_1(esk8_3(esk3_0,X5,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,X4),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|k9_metric_1(esk1_0,esk8_3(esk3_0,X5,X1),esk4_3(X3,k1_funct_1(esk3_0,X4),X2))!=k10_xtuple_0(esk3_0)|~r2_hidden(X4,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3644, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_5318, negated_conjecture, (m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),X4)|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(X1)|~v4_relat_1(esk3_0,X4)|~r2_hidden(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4020, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6))),k9_metric_1(esk2_0,X4,X3))|r1_xxreal_0(X3,k6_numbers)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6)|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),k6_numbers)|X4!=k1_funct_1(esk3_0,X2)|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6))),u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6)),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))|X5!=k1_funct_1(esk3_0,X1)|~r2_hidden(X5,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)))|~v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5573, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,k9_xtuple_0(esk3_0))|~r2_hidden(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5557, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_xtuple_0(esk3_0),k8_relat_1(esk3_0,X1))|X1!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_5444, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),X4)|X4!=k8_relat_1(esk3_0,X5)|X5!=k10_xtuple_0(esk3_0)|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5606, plain, (r1_tarski(k9_metric_1(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v8_metric_1(X1))).#trainneg
cnf(i_0_5594, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),k9_xtuple_0(esk3_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5699, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(k1_funct_1(esk3_0,esk12_2(X2,k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,X1,X4),X3,X4)))),X1)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,X1,X4),X3,X4)),k3_pcomps_1(X2))|r1_xxreal_0(esk11_3(esk2_0,X1,X4),k6_numbers)|v2_struct_0(X2)|X4!=k1_funct_1(esk3_0,X3)|~r2_hidden(esk12_2(X2,k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,X1,X4),X3,X4))),u1_struct_0(esk1_0))|~r2_hidden(X4,X1)|~v3_pre_topc(X1,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X1,X4))|~m1_subset_1(k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,X1,X4),X3,X4)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_5528, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X1,X3),X2,X3)),X4)),X1)|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X1,X3),X2,X3)),X4)|r1_xxreal_0(esk11_3(esk2_0,X1,X3),k6_numbers)|X3!=k1_funct_1(esk3_0,X2)|~r2_hidden(esk14_2(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X1,X3),X2,X3)),X4),u1_struct_0(esk1_0))|~r2_hidden(X3,X1)|~v3_pre_topc(X1,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X1,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5448, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk6_3(X1,esk1_0,X2),X3)|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|X3!=k8_relat_1(esk3_0,X4)|X4!=k10_xtuple_0(esk3_0)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_5597, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk3_0))|r2_hidden(k1_funct_1(esk3_0,X1),k10_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5605, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,X2),X3),X4)|r1_tarski(k9_metric_1(esk1_0,X1,X2),X3)|X4!=k8_relat_1(esk3_0,X5)|X5!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5609, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)|r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4),X5)|r1_xxreal_0(X2,k6_numbers)|X5!=k8_relat_1(esk3_0,k9_metric_1(esk2_0,X3,X2))|X3!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~v1_xreal_0(esk4_3(X2,X1,X3)))).#trainneg
cnf(i_0_5595, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relat_1(esk3_0,X1),X2)|r2_hidden(esk14_2(k8_relat_1(esk3_0,X1),X2),k8_relat_1(esk3_0,X3))|X3!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_5462, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),X4)|X4!=k8_relat_1(esk3_0,X5)|X5!=k10_xtuple_0(esk3_0)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5432, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk12_2(esk1_0,X1),X2)|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|X2!=k8_relat_1(esk3_0,X3)|X3!=k10_xtuple_0(esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_5608, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,X2),k8_relat_1(esk3_0,X3))|X3!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5550, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relat_1(esk3_0,X1),k8_relat_1(esk3_0,X2))|X2!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_5574, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),X4)|X4!=k8_relat_1(X1,k8_relat_1(esk3_0,X5))|X5!=k10_xtuple_0(esk3_0)|~r2_hidden(X3,u1_struct_0(esk1_0))|~r2_hidden(esk8_3(X1,X2,X3),k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_5548, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_xtuple_0(esk3_0),X1)|r2_hidden(esk14_2(k9_xtuple_0(esk3_0),X1),k8_relat_1(esk3_0,X2))|X2!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_5473, plain, (r2_hidden(esk7_3(X2,X1,X3),k1_zfmisc_1(k9_metric_1(X1,esk6_3(X2,X1,X3),esk5_3(X2,X1,X3))))|v5_pre_topc(X3,X2,k3_pcomps_1(X1))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X1)))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X3)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_5590, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),X4)|X4!=k8_relat_1(X1,k8_relat_1(esk3_0,X5))|X5!=k10_xtuple_0(esk3_0)|~r2_hidden(esk8_3(X1,X2,X3),k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_6203, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6))),X3)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6)|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),k6_numbers)|r1_xxreal_0(esk11_3(esk2_0,X3,X4),k6_numbers)|~r2_hidden(X4,X3)|~v3_pre_topc(X3,k3_pcomps_1(esk2_0))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X4,u1_struct_0(esk2_0))|X4!=k1_funct_1(esk3_0,X2)|X5!=k1_funct_1(esk3_0,X1)|~r2_hidden(X5,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)))|~v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5))|~v1_xreal_0(esk11_3(esk2_0,X3,X4))|~m1_subset_1(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6))),u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6)),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5469, plain, (r2_hidden(X1,k8_relat_1(X2,X3))|~r2_hidden(X1,k8_relat_1(X2,X4))|X3!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_5547, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(esk3_0,X1,X2),k8_relat_1(esk3_0,X3))|X3!=k10_xtuple_0(esk3_0)|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_5666, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk12_2(X1,k8_relat_1(esk3_0,X2)),k8_relat_1(esk3_0,X3))|v3_pre_topc(k8_relat_1(esk3_0,X2),k3_pcomps_1(X1))|v2_struct_0(X1)|X3!=k10_xtuple_0(esk3_0)|~m1_subset_1(k8_relat_1(esk3_0,X2),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_6071, negated_conjecture, (r1_tarski(k8_relset_1(u1_struct_0(esk1_0),X1,X2,X3),k8_relat_1(esk3_0,X4))|u1_struct_0(esk2_0)=k1_xboole_0|X4!=k10_xtuple_0(esk3_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X1))))).#trainneg
cnf(i_0_5982, plain, (k1_funct_1(X1,k1_xboole_0)=X3|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_5495, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X2),X1,X2)),X3)|r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X2),X1,X2)),X3)),k8_relat_1(esk3_0,X4))|r1_xxreal_0(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X2),k6_numbers)|X2!=k1_funct_1(esk3_0,X1)|X4!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,k9_xtuple_0(esk3_0))|~v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X2))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X2),X1,X2)),X3)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X2),X1,X2)),X3),u1_struct_0(esk1_0))|~m1_subset_1(k9_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5468, plain, (m1_subset_1(k8_relat_1(X1,X2),k1_zfmisc_1(k8_relat_1(X1,X3)))|X3!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6204, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6))),X3)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6)|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),k6_numbers)|r1_xxreal_0(esk11_3(esk2_0,X3,X4),k6_numbers)|~r2_hidden(X4,X3)|~v3_pre_topc(X3,k3_pcomps_1(esk2_0))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X4,u1_struct_0(esk2_0))|X4!=k1_funct_1(esk3_0,X2)|X5!=k1_funct_1(esk3_0,X1)|~r2_hidden(X5,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)))|~v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5))|~v1_xreal_0(esk11_3(esk2_0,X3,X4))|~m1_subset_1(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6))),u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6)),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6206, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X3,esk4_3(X4,X3,X1)),X5),X2,X5)),X6))),X4))|r1_tarski(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X3,esk4_3(X4,X3,X1)),X5),X2,X5)),X6)|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X3,esk4_3(X4,X3,X1)),X5),k6_numbers)|r1_xxreal_0(X4,k6_numbers)|~v1_xreal_0(X4)|~m1_subset_1(X1,u1_struct_0(esk2_0))|X1!=k1_funct_1(esk3_0,X3)|X5!=k1_funct_1(esk3_0,X2)|~r2_hidden(X5,k9_metric_1(esk1_0,X3,esk4_3(X4,X3,X1)))|~v3_pre_topc(k9_metric_1(esk1_0,X3,esk4_3(X4,X3,X1)),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X3,esk4_3(X4,X3,X1)),X5))|~m1_subset_1(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X3,esk4_3(X4,X3,X1)),X5),X2,X5)),X6))),u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X3,esk4_3(X4,X3,X1)),X5),X2,X5)),X6)),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X3,esk4_3(X4,X3,X1)),X5),X2,X5)),X6)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X3,esk4_3(X4,X3,X1)),X5),X2,X5)),X6),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X3,esk4_3(X4,X3,X1)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5951, plain, (r2_hidden(X4,u1_struct_0(X1))|v2_struct_0(X1)|~v1_xboole_0(k9_metric_1(X1,X2,X3))|~r2_hidden(X2,k1_xboole_0)|~v1_xreal_0(X3)|~m1_subset_1(X4,u1_struct_0(X1))|~l1_metric_1(X1)|~v8_metric_1(X1))).#trainneg
cnf(i_0_6207, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6)),X7)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6)|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),k6_numbers)|r1_xxreal_0(X3,k6_numbers)|X7!=k8_relat_1(esk3_0,k9_metric_1(esk2_0,X4,X3))|~r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6)),k9_xtuple_0(esk3_0))|X4!=k1_funct_1(esk3_0,X2)|X5!=k1_funct_1(esk3_0,X1)|~r2_hidden(X5,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)))|~v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5))|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6))),u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6)),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5599, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|m1_subset_1(X2,k10_xtuple_0(esk3_0))|~r2_hidden(esk8_3(esk3_0,X1,X2),u1_struct_0(esk1_0))|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_6108, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk12_2(X1,k8_relset_1(u1_struct_0(esk1_0),X2,X3,X4))),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k8_relset_1(u1_struct_0(esk1_0),X2,X3,X4),k3_pcomps_1(X1))|v2_struct_0(X1)|~m1_subset_1(k8_relset_1(u1_struct_0(esk1_0),X2,X3,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X2)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_5600, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),k8_relat_1(esk3_0,X4))|X4!=k10_xtuple_0(esk3_0)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5822, negated_conjecture, (v3_pre_topc(esk7_3(X1,esk2_0,X2),k3_pcomps_1(X3))|v3_pre_topc(u1_struct_0(esk2_0),k3_pcomps_1(esk2_0))|v5_pre_topc(X2,X1,k3_pcomps_1(esk2_0))|v2_struct_0(X3)|v2_struct_0(X1)|~v3_pre_topc(k9_metric_1(esk2_0,esk6_3(X1,esk2_0,X2),esk5_3(X1,esk2_0,X2)),k3_pcomps_1(esk2_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk7_3(X1,esk2_0,X2),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(k9_metric_1(esk2_0,esk6_3(X1,esk2_0,X2),esk5_3(X1,esk2_0,X2)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk2_0))|~v1_funct_1(X2)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)|~m1_subset_1(esk12_2(X3,esk7_3(X1,esk2_0,X2)),u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_5801, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk12_2(X1,k9_metric_1(esk1_0,X2,X3)),X4)|r2_hidden(X2,k1_xboole_0)|v3_pre_topc(k9_metric_1(esk1_0,X2,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|X4!=k8_relat_1(esk3_0,X5)|X5!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X3)|~m1_subset_1(k9_metric_1(esk1_0,X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_5010, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(X3,k6_numbers)|X1!=k1_funct_1(esk3_0,esk8_3(X4,X5,X6))|X7!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X7)|~v1_xreal_0(X3)|~m1_subset_1(esk8_3(esk3_0,X7,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(esk8_3(X4,X5,X6),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|k9_metric_1(esk1_0,esk8_3(esk3_0,X7,X2),esk4_3(X3,esk8_3(X4,X5,X6),X1))!=k8_relat_1(X4,u1_struct_0(esk1_0))|X5!=k10_xtuple_0(X4)|~r2_hidden(X6,k9_xtuple_0(esk3_0))|~r2_hidden(X6,X5)|~v1_relat_1(X4)|~v1_funct_1(X4))).#trainneg
cnf(i_0_3976, negated_conjecture, (v3_pre_topc(u1_struct_0(esk2_0),k3_pcomps_1(esk2_0))|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1)|~v3_pre_topc(k8_relat_1(esk3_0,X3),k3_pcomps_1(esk1_0))|X3!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_4984, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(esk1_0))|k8_relat_1(X1,u1_struct_0(esk1_0))!=k9_xtuple_0(esk3_0)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5817, plain, (v3_pre_topc(esk7_3(X2,X1,X3),k3_pcomps_1(X4))|v5_pre_topc(X3,X2,k3_pcomps_1(X1))|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X1)|~v1_xboole_0(k9_metric_1(X1,esk6_3(X2,X1,X3),esk5_3(X2,X1,X3)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X1)))))|~m1_subset_1(esk7_3(X2,X1,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~m1_subset_1(esk12_2(X4,esk7_3(X2,X1,X3)),u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X1)))|~v1_funct_1(X3)|~l1_metric_1(X1)|~l1_metric_1(X4)|~v9_metric_1(X4)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v8_metric_1(X4)|~v7_metric_1(X4)|~v7_metric_1(X1)|~v6_metric_1(X4)|~v6_metric_1(X1))).#trainneg
cnf(i_0_5258, negated_conjecture, (m1_subset_1(k8_relat_1(esk3_0,X1),k1_zfmisc_1(X2))|u1_struct_0(esk2_0)=k1_xboole_0|~v4_relat_1(esk3_0,X2))).#trainneg
cnf(i_0_4776, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk12_2(X1,k9_xtuple_0(esk3_0)),X2)|v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(X1))|v2_struct_0(X1)|X2!=k8_relat_1(esk3_0,X3)|X3!=k10_xtuple_0(esk3_0)|~m1_subset_1(k9_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_5338, negated_conjecture, (esk12_2(esk1_0,X1)=k1_xboole_0|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|~v4_relat_1(esk3_0,k1_zfmisc_1(k1_xboole_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_4979, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(esk1_0))|k8_relat_1(esk3_0,X4)!=k8_relat_1(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5938, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,X2)|v3_pre_topc(u1_struct_0(esk1_0),k3_pcomps_1(X3))|v2_struct_0(X3)|X1!=k1_funct_1(esk3_0,esk12_2(X3,u1_struct_0(esk1_0)))|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_6152, negated_conjecture, (r2_hidden(esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4))),X5)|v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),k3_pcomps_1(X1))|r1_xxreal_0(esk11_3(esk2_0,X3,X4),k6_numbers)|v2_struct_0(X1)|u1_struct_0(esk2_0)=k1_xboole_0|X5!=k8_relat_1(esk3_0,X3)|X4!=k1_funct_1(esk3_0,X2)|~r2_hidden(esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4))),u1_struct_0(esk1_0))|~r2_hidden(X4,X3)|~v3_pre_topc(X3,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X3,X4))|~m1_subset_1(k1_funct_1(esk3_0,esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)))),u1_struct_0(esk2_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3071, negated_conjecture, (r2_hidden(esk12_2(X1,k1_relset_1(k10_xtuple_0(esk3_0),X2)),u1_struct_0(esk2_0))|v3_pre_topc(k1_relset_1(k10_xtuple_0(esk3_0),X2),k3_pcomps_1(X1))|v2_struct_0(X1)|~v4_relat_1(X2,k10_xtuple_0(esk3_0))|~v1_relat_1(X2)|~m1_subset_1(k1_relset_1(k10_xtuple_0(esk3_0),X2),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3074, negated_conjecture, (r2_hidden(esk12_2(X1,k9_metric_1(X2,X3,esk11_3(X2,k10_xtuple_0(esk3_0),X3))),u1_struct_0(esk2_0))|v3_pre_topc(k9_metric_1(X2,X3,esk11_3(X2,k10_xtuple_0(esk3_0),X3)),k3_pcomps_1(X1))|v2_struct_0(X1)|v2_struct_0(X2)|~r2_hidden(X3,k10_xtuple_0(esk3_0))|~v3_pre_topc(k10_xtuple_0(esk3_0),k3_pcomps_1(X2))|~m1_subset_1(k9_metric_1(X2,X3,esk11_3(X2,k10_xtuple_0(esk3_0),X3)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(k10_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)|~l1_metric_1(X1)|~v9_metric_1(X2)|~v9_metric_1(X1)|~v8_metric_1(X2)|~v8_metric_1(X1)|~v7_metric_1(X2)|~v7_metric_1(X1)|~v6_metric_1(X2)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3767, negated_conjecture, (r1_tarski(k2_pcomps_1(esk1_0),X1)|m1_subset_1(esk14_2(k2_pcomps_1(esk1_0),X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3679, negated_conjecture, (r2_hidden(esk14_2(esk7_3(X1,esk1_0,X2),X3),u1_struct_0(esk1_0))|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|r1_tarski(esk7_3(X1,esk1_0,X2),X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_3152, negated_conjecture, (v1_xreal_0(X1)|r1_tarski(k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X1)),X4)|r1_xxreal_0(X3,k6_numbers)|k9_metric_1(esk2_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X1)),X4)),X3)!=k10_xtuple_0(X5)|~v3_valued_0(X5)|~v1_relat_1(X5)|~v1_funct_1(X5)|X1!=k1_funct_1(esk3_0,X2)|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X1)),X4)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X1)),X4),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_5127, negated_conjecture, (v5_pre_topc(k1_xboole_0,k3_pcomps_1(esk1_0),k3_pcomps_1(esk1_0))|~v1_funct_2(k1_xboole_0,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_4559, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,esk14_2(k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5),X6))|r1_xxreal_0(X6,k6_numbers)|r1_tarski(k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5)|v2_struct_0(X2)|X1!=k1_funct_1(esk3_0,X7)|X8!=k10_xtuple_0(esk3_0)|~r2_hidden(esk14_2(k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5),X8)|~v1_xreal_0(X6)|~m1_subset_1(esk8_3(esk3_0,X8,esk14_2(k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5)),u1_struct_0(esk1_0))|~m1_subset_1(X7,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5),u1_struct_0(esk2_0))|k9_metric_1(esk1_0,X7,esk4_3(X6,X7,X1))!=k8_relat_1(esk3_0,X4)|~r2_hidden(X3,X4)|~v3_pre_topc(X4,k3_pcomps_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_4354, negated_conjecture, (r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk5_3(X2,esk2_0,X3),X1,esk6_3(X2,esk2_0,X3))),X4),X5)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk5_3(X2,esk2_0,X3),X1,esk6_3(X2,esk2_0,X3))),X4)|v5_pre_topc(X3,X2,k3_pcomps_1(esk2_0))|r1_xxreal_0(esk5_3(X2,esk2_0,X3),k6_numbers)|v2_struct_0(X2)|X5!=k8_relat_1(esk3_0,esk7_3(X2,esk2_0,X3))|~r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk5_3(X2,esk2_0,X3),X1,esk6_3(X2,esk2_0,X3))),X4),k9_xtuple_0(esk3_0))|esk6_3(X2,esk2_0,X3)!=k1_funct_1(esk3_0,X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xreal_0(esk5_3(X2,esk2_0,X3))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk5_3(X2,esk2_0,X3),X1,esk6_3(X2,esk2_0,X3))),X4)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk5_3(X2,esk2_0,X3),X1,esk6_3(X2,esk2_0,X3))),X4),u1_struct_0(esk1_0))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~m1_subset_1(esk6_3(X2,esk2_0,X3),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3))).#trainneg
cnf(i_0_3926, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,X1),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|k10_xtuple_0(esk3_0)!=k9_xtuple_0(esk3_0)|~r2_hidden(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4513, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,esk14_2(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6)),u1_struct_0(esk1_0))|r1_tarski(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6)|v2_struct_0(X3)|k8_relat_1(esk3_0,X7)!=k8_relat_1(X1,X5)|X2!=k10_xtuple_0(X1)|~r2_hidden(esk14_2(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2)|~r2_hidden(X4,X5)|~v1_relat_1(X1)|~v3_pre_topc(X5,k3_pcomps_1(X3))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X4,u1_struct_0(X3))|~v1_funct_1(X1)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_4830, negated_conjecture, (r2_hidden(esk12_2(esk2_0,k9_metric_1(esk2_0,X1,X2)),u1_struct_0(esk2_0))|v3_pre_topc(k9_metric_1(esk2_0,X1,X2),k3_pcomps_1(esk2_0))|r1_xxreal_0(X3,k6_numbers)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_3487, plain, (r2_hidden(k1_funct_1(X1,esk12_2(X2,k1_relset_1(k9_xtuple_0(X1),X3))),X4)|v3_pre_topc(k1_relset_1(k9_xtuple_0(X1),X3),k3_pcomps_1(X2))|v2_struct_0(X2)|X4!=k10_xtuple_0(X1)|~v4_relat_1(X3,k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_relat_1(X3)|~m1_subset_1(k1_relset_1(k9_xtuple_0(X1),X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~v1_funct_1(X1)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_6227, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k8_relat_1(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X1,k9_xtuple_0(X3))|~v1_relat_1(X3)|~v1_funct_1(X3)|~r2_hidden(k1_funct_1(X3,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6217, negated_conjecture, (r1_tarski(X1,k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(esk14_2(X1,k9_xtuple_0(esk3_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6263, negated_conjecture, (r2_hidden(esk8_3(X1,X2,X3),X4)|u1_struct_0(esk2_0)=k1_xboole_0|X4!=k8_relat_1(X1,k9_xtuple_0(esk3_0))|~r2_hidden(esk8_3(X1,X2,X3),k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,u1_struct_0(esk1_0))|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_6549, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(esk7_3(X1,esk1_0,X2),k8_relat_1(esk3_0,X3))|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|X3!=k10_xtuple_0(esk3_0)|~v1_xreal_0(esk5_3(X1,esk1_0,X2))|~m1_subset_1(esk6_3(X1,esk1_0,X2),u1_struct_0(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_6694, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk8_3(esk3_0,X1,X2)),X3)|u1_struct_0(esk2_0)=k1_xboole_0|X3!=k10_xtuple_0(esk3_0)|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_6385, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk6_3(X1,esk1_0,X2),k8_relat_1(esk3_0,X3))|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|X3!=k10_xtuple_0(esk3_0)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_6243, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),k8_relat_1(esk3_0,X4))|X4!=k10_xtuple_0(esk3_0)|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6546, negated_conjecture, (m1_subset_1(k9_metric_1(esk1_0,X1,X2),k1_zfmisc_1(k8_relat_1(esk3_0,X3)))|u1_struct_0(esk2_0)=k1_xboole_0|X3!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6255, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k1_funct_1(esk3_0,k1_funct_1(X3,esk8_3(X3,X4,X5)))|X2!=k10_xtuple_0(esk3_0)|X4!=k10_xtuple_0(X3)|~r2_hidden(X5,u1_struct_0(esk1_0))|~r2_hidden(X5,X4)|~v1_relat_1(X3)|~v1_funct_1(X3))).#trainneg
cnf(i_0_6372, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5),X1,X5)),X6)),X4)|v2_struct_0(X2)|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5),X1,X5)),X6)|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5),k6_numbers)|~r2_hidden(X3,X4)|~v3_pre_topc(X4,k3_pcomps_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|X5!=k1_funct_1(esk3_0,X1)|~r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5),X1,X5)),X6),u1_struct_0(esk1_0))|~r2_hidden(X5,k9_metric_1(X2,X3,esk11_3(X2,X4,X3)))|~v3_pre_topc(k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),X5))|~m1_subset_1(k9_metric_1(X2,X3,esk11_3(X2,X4,X3)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6562, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_xtuple_0(esk3_0),X1)),X2)|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_xtuple_0(esk3_0),X1)|X2!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_6465, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,X2),X3)|r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,X2),X3),k8_relat_1(esk3_0,X4))|X4!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6241, negated_conjecture, (r2_hidden(X1,k8_relat_1(esk3_0,X2))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,k9_xtuple_0(esk3_0))|X2!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_6573, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),k8_relat_1(X1,k8_relat_1(esk3_0,X4)))|X4!=k10_xtuple_0(esk3_0)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6554, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),k8_relat_1(X1,k8_relat_1(esk3_0,X4)))|X4!=k10_xtuple_0(esk3_0)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6840, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,k8_relat_1(esk3_0,X3))|~v4_relat_1(esk3_0,X2))).#trainneg
cnf(i_0_6540, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk14_2(k8_relat_1(esk3_0,X1),X2)),X3)|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relat_1(esk3_0,X1),X2)|X3!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_6547, negated_conjecture, (r2_hidden(X1,k8_relat_1(esk3_0,X2))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,k9_metric_1(esk1_0,X3,X4))|X2!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6252, negated_conjecture, (m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6553, negated_conjecture, (r2_hidden(X1,k8_relat_1(esk3_0,X2))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,k8_relat_1(esk3_0,X3))|X2!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_6705, negated_conjecture, (r2_hidden(esk12_2(X1,k8_relat_1(esk3_0,X2)),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X2),k3_pcomps_1(X1))|v2_struct_0(X1)|X3!=k10_xtuple_0(esk3_0)|~m1_subset_1(k8_relat_1(esk3_0,X2),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_6545, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk12_2(esk1_0,X1),k8_relat_1(esk3_0,X2))|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_6234, negated_conjecture, (m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),X4)|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(X1)|~v4_relat_1(esk3_0,X4)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6483, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4),X5)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)|r1_xxreal_0(X2,k6_numbers)|X5!=k8_relat_1(esk3_0,X6)|X6!=k10_xtuple_0(esk3_0)|k8_relat_1(esk3_0,k9_metric_1(esk2_0,X3,X2))!=k9_xtuple_0(esk3_0)|X3!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_6240, negated_conjecture, (m1_subset_1(k9_xtuple_0(esk3_0),k1_zfmisc_1(k8_relat_1(esk3_0,X1)))|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_6876, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(esk7_3(X1,esk1_0,X2),k8_relat_1(esk3_0,X3))|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|X3!=k10_xtuple_0(esk3_0)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_6621, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),X7),X1,X7)),X8))),X5)|v2_struct_0(X3)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),X7),X1,X7)),X8)|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),X7),k6_numbers)|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),k6_numbers)|~r2_hidden(X4,X5)|~v3_pre_topc(X5,k3_pcomps_1(X3))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)|X6!=k1_funct_1(esk3_0,X2)|X7!=k1_funct_1(esk3_0,X1)|~r2_hidden(X7,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)))|~r2_hidden(X6,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)))|~v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),k3_pcomps_1(esk2_0))|~v3_pre_topc(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),X7))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6))|~m1_subset_1(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),X7),X1,X7)),X8))),u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),X7),X1,X7)),X8)),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),X7),X1,X7)),X8)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),X7),X1,X7)),X8),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6531, negated_conjecture, (r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)|u1_struct_0(esk2_0)=k1_xboole_0|r1_xxreal_0(X2,k6_numbers)|X4!=k8_relat_1(esk3_0,k9_metric_1(esk2_0,X3,X2))|X3!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_6539, negated_conjecture, (r2_hidden(esk14_2(k8_relat_1(esk3_0,X1),X2),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relat_1(esk3_0,X1),X2)|X3!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_6374, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,k9_metric_1(X4,X5,esk11_3(X4,X6,X5)),X7),X3,X7)),X8)))),X6)|v2_struct_0(X4)|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,k9_metric_1(X4,X5,esk11_3(X4,X6,X5)),X7),X3,X7)),X8)|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(X4,X5,esk11_3(X4,X6,X5)),X7),k6_numbers)|X2!=k10_xtuple_0(X1)|~r2_hidden(X5,X6)|~r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,k9_metric_1(X4,X5,esk11_3(X4,X6,X5)),X7),X3,X7)),X8)),X2)|~v1_relat_1(X1)|~v3_pre_topc(X6,k3_pcomps_1(X4))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~m1_subset_1(X5,u1_struct_0(X4))|~v1_funct_1(X1)|~l1_metric_1(X4)|~v9_metric_1(X4)|~v8_metric_1(X4)|~v7_metric_1(X4)|~v6_metric_1(X4)|X7!=k1_funct_1(esk3_0,X3)|~r2_hidden(esk14_2(k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,k9_metric_1(X4,X5,esk11_3(X4,X6,X5)),X7),X3,X7)),X8),u1_struct_0(esk1_0))|~r2_hidden(X7,k9_metric_1(X4,X5,esk11_3(X4,X6,X5)))|~v3_pre_topc(k9_metric_1(X4,X5,esk11_3(X4,X6,X5)),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(X4,X5,esk11_3(X4,X6,X5)),X7))|~m1_subset_1(k9_metric_1(X4,X5,esk11_3(X4,X6,X5)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7153, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|m1_subset_1(esk7_3(X1,esk1_0,X2),k1_zfmisc_1(k8_relat_1(esk3_0,X3)))|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|X3!=k10_xtuple_0(esk3_0)|~v1_xreal_0(esk5_3(X1,esk1_0,X2))|~m1_subset_1(esk6_3(X1,esk1_0,X2),u1_struct_0(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_7327, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,X1),X2)|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,k9_xtuple_0(esk3_0)))).#trainneg
cnf(i_0_7617, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk12_2(X1,k8_relat_1(esk3_0,X2)),k9_xtuple_0(esk3_0))|v3_pre_topc(k8_relat_1(esk3_0,X2),k3_pcomps_1(X1))|v2_struct_0(X1)|~m1_subset_1(k8_relat_1(esk3_0,X2),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_7054, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),k8_relat_1(X1,k9_xtuple_0(esk3_0)))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7315, negated_conjecture, (r1_tarski(X1,k8_relat_1(esk3_0,X2))|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(esk14_2(X1,k8_relat_1(esk3_0,X2)),k9_xtuple_0(esk3_0)))).#trainneg
cnf(i_0_7042, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relat_1(esk3_0,X1),k9_xtuple_0(esk3_0)))).#trainneg
cnf(i_0_7542, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,k8_relat_1(esk3_0,X2))|v3_pre_topc(k9_metric_1(esk1_0,X1,X4),k3_pcomps_1(X3))|v2_struct_0(X3)|X2!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X4)|~m1_subset_1(esk12_2(X3,k9_metric_1(esk1_0,X1,X4)),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X1,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_7676, negated_conjecture, (m1_subset_1(esk7_3(X1,esk1_0,X2),k1_zfmisc_1(k8_relat_1(esk3_0,X3)))|u1_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|X3!=k10_xtuple_0(esk3_0)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_7748, negated_conjecture, (v3_pre_topc(X1,k3_pcomps_1(esk1_0))|r1_xxreal_0(esk4_3(X2,esk12_2(esk1_0,X1),X3),k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(esk4_3(X2,esk12_2(esk1_0,X1),X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|X1!=k8_relat_1(esk3_0,k9_metric_1(esk2_0,X3,X2))|X3!=k1_funct_1(esk3_0,esk12_2(esk1_0,X1))|~v1_xreal_0(X2)|~m1_subset_1(esk12_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_7053, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,X2),k9_xtuple_0(esk3_0))|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7751, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relat_1(esk3_0,X1),X2)|r2_hidden(esk14_2(k8_relat_1(esk3_0,X1),X2),k9_xtuple_0(esk3_0)))).#trainneg
cnf(i_0_7623, negated_conjecture, (r2_hidden(esk12_2(esk1_0,X1),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7504, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk12_2(X1,k9_metric_1(esk1_0,X2,X3)),k8_relat_1(esk3_0,X4))|v3_pre_topc(k9_metric_1(esk1_0,X2,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|X4!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_7132, negated_conjecture, (m1_subset_1(esk6_3(X1,esk1_0,X2),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|X3!=k10_xtuple_0(esk3_0)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_7129, negated_conjecture, (r2_hidden(esk6_3(X1,esk1_0,X2),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|X3!=k10_xtuple_0(esk3_0)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_7540, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,k8_relat_1(esk3_0,X2))|r1_tarski(k9_metric_1(esk1_0,X1,X3),X4)|X2!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X3)|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,X3),X4),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7047, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(X1,X2,esk11_3(X1,u1_struct_0(esk1_0),X2)),k9_xtuple_0(esk3_0))|v2_struct_0(X1)|~r2_hidden(X2,u1_struct_0(esk1_0))|~v3_pre_topc(u1_struct_0(esk1_0),k3_pcomps_1(X1))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_7049, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(esk7_3(X1,esk1_0,X2),k9_xtuple_0(esk3_0))|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_7139, negated_conjecture, (r2_hidden(esk8_3(X1,X2,X3),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|X4!=k10_xtuple_0(esk3_0)|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7304, negated_conjecture, (m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,X2),X3),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,X2),X3)|X4!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7665, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(k1_funct_1(X2,esk8_3(X2,X3,X4)),X1)|u1_struct_0(esk2_0)=k1_xboole_0|X3!=k10_xtuple_0(X2)|~v4_relat_1(esk3_0,X1)|~r2_hidden(X4,u1_struct_0(esk1_0))|~r2_hidden(X4,X3)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_7044, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relset_1(u1_struct_0(esk1_0),X1,X2,X3),k9_xtuple_0(esk3_0))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X1))))).#trainneg
cnf(i_0_7670, negated_conjecture, (r2_hidden(k9_xtuple_0(esk3_0),k1_zfmisc_1(k8_relat_1(esk3_0,X1)))|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_6313, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)))),X5)|v2_struct_0(X3)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),k3_pcomps_1(X1))|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),k6_numbers)|v2_struct_0(X1)|~r2_hidden(X4,X5)|~v3_pre_topc(X5,k3_pcomps_1(X3))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)|X6!=k1_funct_1(esk3_0,X2)|~r2_hidden(esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6))),u1_struct_0(esk1_0))|~r2_hidden(X6,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)))|~v3_pre_topc(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_6544, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),k8_relat_1(esk3_0,X4))|X4!=k10_xtuple_0(esk3_0)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6389, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,X1),X2)|u1_struct_0(esk2_0)=k1_xboole_0|v1_xboole_0(k10_xtuple_0(esk3_0))|~v5_relat_1(esk3_0,X2)|~r2_hidden(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7486, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,k8_relat_1(esk3_0,X2))|v5_pre_topc(X4,X3,k3_pcomps_1(esk1_0))|v2_struct_0(X3)|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,esk7_3(X3,esk1_0,X4))|~v1_xreal_0(esk5_3(X3,esk1_0,X4))|~m1_subset_1(esk6_3(X3,esk1_0,X4),u1_struct_0(esk1_0))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk1_0))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(esk1_0))|~v1_funct_1(X4))).#trainneg
cnf(i_0_6693, negated_conjecture, (r2_hidden(esk8_3(esk3_0,X1,X2),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|X3!=k10_xtuple_0(esk3_0)|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_7041, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relat_1(X1,X2),k9_xtuple_0(esk3_0))|~v4_relat_1(X1,u1_struct_0(esk1_0))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7675, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k8_relat_1(esk3_0,X1)))|X1!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_7309, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(esk7_3(X1,esk1_0,X2),X3)|r2_hidden(esk14_2(esk7_3(X1,esk1_0,X2),X3),k8_relat_1(esk3_0,X4))|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|X4!=k10_xtuple_0(esk3_0)|~v1_xreal_0(esk5_3(X1,esk1_0,X2))|~m1_subset_1(esk6_3(X1,esk1_0,X2),u1_struct_0(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_6481, negated_conjecture, (m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)|r1_xxreal_0(X2,k6_numbers)|k8_relat_1(esk3_0,k9_metric_1(esk2_0,X3,X2))!=k9_xtuple_0(esk3_0)|X3!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_6428, negated_conjecture, (m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),k1_zfmisc_1(u1_struct_0(esk2_0)))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|k2_pcomps_1(esk2_0)!=k8_relat_1(esk3_0,X2)|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk4_3(esk11_3(esk2_0,X2,X3),X1,X3))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6819, negated_conjecture, (k8_relat_1(esk3_0,X1)=k1_xboole_0|u1_struct_0(esk2_0)=k1_xboole_0|~v4_relat_1(esk3_0,k1_xboole_0))).#trainneg
cnf(i_0_6552, negated_conjecture, (m1_subset_1(k8_relat_1(esk3_0,X1),k1_zfmisc_1(k8_relat_1(esk3_0,X2)))|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_7671, negated_conjecture, (v1_xboole_0(k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|~v1_xboole_0(k8_relat_1(esk3_0,X1))|X1!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_7048, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(u1_struct_0(esk1_0),k9_xtuple_0(esk3_0)))).#trainneg
cnf(i_0_6498, negated_conjecture, (m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4),k1_zfmisc_1(u1_struct_0(esk2_0)))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)|r1_xxreal_0(X2,k6_numbers)|k8_relat_1(esk3_0,k9_metric_1(esk2_0,X3,X2))!=k2_pcomps_1(esk2_0)|X3!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_6710, negated_conjecture, (m1_subset_1(k8_relset_1(u1_struct_0(esk1_0),X1,X2,X3),k1_zfmisc_1(k8_relat_1(esk3_0,X4)))|u1_struct_0(esk2_0)=k1_xboole_0|X4!=k10_xtuple_0(esk3_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X1))))).#trainneg
cnf(i_0_7165, negated_conjecture, (r2_hidden(esk12_2(esk1_0,X1),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_6842, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk12_2(X1,k9_xtuple_0(esk3_0)),k8_relat_1(esk3_0,X2))|v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(X1))|v2_struct_0(X1)|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(k9_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_6711, negated_conjecture, (r2_hidden(X1,k8_relat_1(esk3_0,X2))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,k8_relset_1(u1_struct_0(esk1_0),X3,X4,X5))|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X3))))).#trainneg
cnf(i_0_7147, negated_conjecture, (r2_hidden(k9_metric_1(esk1_0,X2,X3),k1_zfmisc_1(k8_relat_1(esk3_0,X1)))|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7624, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk12_2(esk1_0,X1)),X2)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7152, negated_conjecture, (v3_pre_topc(k8_relat_1(esk3_0,X1),k3_pcomps_1(esk1_0))|r1_xxreal_0(X2,k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|~v1_xreal_0(X2)|X1!=k10_xtuple_0(esk3_0)|~m1_subset_1(esk12_2(esk1_0,k8_relat_1(esk3_0,X1)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6500, negated_conjecture, (m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4),k1_zfmisc_1(u1_struct_0(esk1_0)))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)|r1_xxreal_0(X2,k6_numbers)|k8_relat_1(esk3_0,k9_metric_1(esk2_0,X3,X2))!=k2_pcomps_1(esk1_0)|X3!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_6703, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(esk12_2(X2,k8_relat_1(esk3_0,X3)),X1)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X3),k3_pcomps_1(X2))|v2_struct_0(X2)|~v4_relat_1(esk3_0,X1)|X4!=k10_xtuple_0(esk3_0)|~m1_subset_1(k8_relat_1(esk3_0,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_6560, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,esk14_2(k9_xtuple_0(esk3_0),X3))),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_xtuple_0(esk3_0),X3)|X2!=k10_xtuple_0(X1)|~r2_hidden(esk14_2(k9_xtuple_0(esk3_0),X3),X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|X4!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_6228, negated_conjecture, (m1_subset_1(X1,k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7301, negated_conjecture, (r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,X2),X3),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,X2),X3)|X4!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7133, negated_conjecture, (r2_hidden(esk6_3(X1,esk1_0,X2),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|X3!=k10_xtuple_0(esk3_0)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_6490, negated_conjecture, (m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)|r1_xxreal_0(X2,k6_numbers)|k8_relat_1(esk3_0,k9_metric_1(esk2_0,X3,X2))!=k10_xtuple_0(esk3_0)|X3!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_6370, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(X2,X3,X4),X5),X1,X5)),X6)),u1_struct_0(X2))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(X2,X3,X4),X5),X1,X5)),X6)|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(X2,X3,X4),X5),k6_numbers)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)|X5!=k1_funct_1(esk3_0,X1)|~r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(X2,X3,X4),X5),X1,X5)),X6),u1_struct_0(esk1_0))|~r2_hidden(X5,k9_metric_1(X2,X3,X4))|~v3_pre_topc(k9_metric_1(X2,X3,X4),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(X2,X3,X4),X5))|~m1_subset_1(k9_metric_1(X2,X3,X4),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6782, negated_conjecture, (m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),k8_relat_1(esk3_0,X4))|u1_struct_0(esk2_0)=k1_xboole_0|X4!=k10_xtuple_0(esk3_0)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6681, plain, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),k8_relat_1(X4,X5))|X5!=k10_xtuple_0(X4)|~v1_relat_1(X4)|~v1_funct_1(X4)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k8_relat_1(X4,X6))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7622, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,esk12_2(esk1_0,X3))),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X3,k3_pcomps_1(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(esk12_2(esk1_0,X3),X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|X4!=k10_xtuple_0(esk3_0)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_6706, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk12_2(X1,k8_relat_1(esk3_0,X2))),X3)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X2),k3_pcomps_1(X1))|v2_struct_0(X1)|X3!=k10_xtuple_0(esk3_0)|~m1_subset_1(k8_relat_1(esk3_0,X2),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_7998, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),k9_xtuple_0(esk3_0))|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7999, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,X2),X3)|m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,X2),X3),u1_struct_0(esk1_0))|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7930, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,k8_relat_1(esk3_0,X2))|v3_pre_topc(k9_metric_1(esk1_0,X1,X3),k3_pcomps_1(esk1_0))|X2!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X3)|~m1_subset_1(k9_metric_1(esk1_0,X1,X3),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8451, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,X2),X3)|r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,X2),X3),k9_xtuple_0(esk3_0))|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7942, negated_conjecture, (m1_subset_1(k9_metric_1(esk1_0,X1,X2),k1_zfmisc_1(k9_xtuple_0(esk3_0)))|u1_struct_0(esk2_0)=k1_xboole_0|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7981, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(X1,X2,k3_pcomps_1(esk1_0))|m1_subset_1(esk6_3(X2,esk1_0,X1),u1_struct_0(esk1_0))|v2_struct_0(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_7943, negated_conjecture, (r2_hidden(X1,k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,k9_metric_1(esk1_0,X2,X3))|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8213, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(esk3_0,X1,X2),k9_xtuple_0(esk3_0))|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_7965, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk12_2(esk1_0,X1),k9_xtuple_0(esk3_0))|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7995, negated_conjecture, (m1_subset_1(esk7_3(X1,esk1_0,X2),k1_zfmisc_1(k9_xtuple_0(esk3_0)))|u1_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_7945, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,X2),u1_struct_0(esk1_0))|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7905, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,esk8_3(X3,X4,X5))),k9_xtuple_0(X3))|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(X1)|~r2_hidden(esk8_3(X3,X4,X5),X2)|~v1_relat_1(X3)|~v1_relat_1(X1)|~v1_funct_1(X3)|~v1_funct_1(X1)|X4!=k10_xtuple_0(X3)|~r2_hidden(X5,u1_struct_0(esk1_0))|~r2_hidden(X5,X4))).#trainneg
cnf(i_0_8390, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,esk12_2(esk1_0,X1)),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_8077, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(u1_struct_0(esk1_0),k1_zfmisc_1(k8_relat_1(esk3_0,X1)))|X1!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_7941, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X3,k3_pcomps_1(esk1_0))|X2!=k1_funct_1(esk3_0,esk12_2(esk1_0,X3))|~v1_xreal_0(X1)|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(esk12_2(esk1_0,X3),u1_struct_0(esk1_0))|X3!=k8_relat_1(esk3_0,k9_metric_1(esk2_0,X2,X1))|~v1_xreal_0(esk4_3(X1,esk12_2(esk1_0,X3),X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7952, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relat_1(esk3_0,X3),X4)|X1!=k1_funct_1(esk3_0,esk14_2(k8_relat_1(esk3_0,X3),X4))|X2!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_8516, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,esk8_3(X3,X4,X5))),k8_relat_1(X3,X6))|u1_struct_0(esk2_0)=k1_xboole_0|X6!=k10_xtuple_0(X3)|X2!=k10_xtuple_0(X1)|~r2_hidden(esk8_3(X3,X4,X5),X2)|~v1_relat_1(X3)|~v1_relat_1(X1)|~v1_funct_1(X3)|~v1_funct_1(X1)|X4!=k10_xtuple_0(X3)|~r2_hidden(X5,u1_struct_0(esk1_0))|~r2_hidden(X5,X4))).#trainneg
cnf(i_0_7934, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,k8_relat_1(esk3_0,X2))|v3_pre_topc(k9_metric_1(esk1_0,X1,X3),k3_pcomps_1(X4))|v2_struct_0(X4)|X2!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X3)|~m1_subset_1(k9_metric_1(esk1_0,X1,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~l1_metric_1(X4)|~v9_metric_1(X4)|~v8_metric_1(X4)|~v7_metric_1(X4)|~v6_metric_1(X4))).#trainneg
cnf(i_0_8070, negated_conjecture, (r2_hidden(X1,k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,k8_relset_1(u1_struct_0(esk1_0),X2,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X2))))).#trainneg
cnf(i_0_7982, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk6_3(X1,esk1_0,X2),k9_xtuple_0(esk3_0))|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_7944, negated_conjecture, (v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(esk1_0))|r1_xxreal_0(X1,k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|~v1_xreal_0(X1)|~m1_subset_1(esk12_2(esk1_0,k9_xtuple_0(esk3_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7925, negated_conjecture, (r2_hidden(X1,k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,k8_relat_1(esk3_0,X2)))).#trainneg
cnf(i_0_8430, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X1),k3_pcomps_1(esk1_0))|r1_xxreal_0(X2,k6_numbers)|X1!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_7971, negated_conjecture, (r2_hidden(esk12_2(X1,k9_metric_1(esk1_0,X2,X3)),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k9_metric_1(esk1_0,X2,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|X4!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X3)|~m1_subset_1(k9_metric_1(esk1_0,X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_7921, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(X1,X2,esk11_3(X1,k9_xtuple_0(esk3_0),X2)),k8_relat_1(esk3_0,X3))|v2_struct_0(X1)|X3!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,k9_xtuple_0(esk3_0))|~v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(X1))|~m1_subset_1(k9_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_7988, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,X2),X3)|r2_hidden(X1,k8_relat_1(esk3_0,X4))|X4!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7949, negated_conjecture, (m1_subset_1(esk14_2(k8_relat_1(esk3_0,X1),X2),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relat_1(esk3_0,X1),X2))).#trainneg
cnf(i_0_7897, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,X3),k8_relat_1(X1,k9_xtuple_0(esk3_0)))|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8391, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk12_2(esk1_0,X1)),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7953, negated_conjecture, (m1_subset_1(esk14_2(k8_relat_1(esk3_0,X1),X2),X3)|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relat_1(esk3_0,X1),X2)|~v4_relat_1(esk3_0,X3))).#trainneg
cnf(i_0_8446, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_xtuple_0(esk3_0),X1)|r2_hidden(k1_funct_1(X2,esk8_3(X2,X3,esk14_2(k9_xtuple_0(esk3_0),X1))),k9_xtuple_0(esk3_0))|X3!=k10_xtuple_0(X2)|~r2_hidden(esk14_2(k9_xtuple_0(esk3_0),X1),X3)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_8223, negated_conjecture, (m1_subset_1(X1,k8_relat_1(esk3_0,X2))|u1_struct_0(esk2_0)=k1_xboole_0|~m1_subset_1(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_8559, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,esk12_2(esk1_0,X3))),k9_xtuple_0(esk3_0))|v3_pre_topc(X3,k3_pcomps_1(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(esk12_2(esk1_0,X3),X2)|~v1_relat_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_1(X1))).#trainneg
cnf(i_0_7919, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),k8_relat_1(esk3_0,X4))|r1_xxreal_0(X2,k6_numbers)|X4!=k10_xtuple_0(esk3_0)|k8_relat_1(esk3_0,k9_metric_1(esk2_0,X3,X2))!=k9_xtuple_0(esk3_0)|X3!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_7924, negated_conjecture, (m1_subset_1(k8_relat_1(esk3_0,X1),k1_zfmisc_1(k9_xtuple_0(esk3_0)))|u1_struct_0(esk2_0)=k1_xboole_0)).#trainneg
cnf(i_0_8409, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(X3,k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X4,k3_pcomps_1(esk1_0))|X1!=k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk12_2(esk1_0,X4)))|X5!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X5)|~v1_xreal_0(X3)|~m1_subset_1(esk8_3(esk3_0,X5,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk12_2(esk1_0,X4)),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|k9_metric_1(esk1_0,esk8_3(esk3_0,X5,X2),esk4_3(X3,k1_funct_1(esk3_0,esk12_2(esk1_0,X4)),X1))!=k10_xtuple_0(esk3_0)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_8452, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk6_3(X1,esk1_0,X2),u1_struct_0(esk1_0))|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_8221, negated_conjecture, (m1_subset_1(X1,k8_relat_1(esk3_0,X2))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_7918, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relset_1(k9_xtuple_0(esk3_0),X1,X2,X3),k8_relat_1(esk3_0,X4))|X4!=k10_xtuple_0(esk3_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k9_xtuple_0(esk3_0),X1))))).#trainneg
cnf(i_0_8408, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(X3,k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X4,k3_pcomps_1(esk1_0))|X2!=k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk12_2(esk1_0,X4)))|X5!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,X5)|~v1_xreal_0(X3)|~m1_subset_1(esk8_3(esk3_0,X5,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk12_2(esk1_0,X4)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|k9_metric_1(esk1_0,esk8_3(esk3_0,X5,X1),esk4_3(X3,k1_funct_1(esk3_0,esk12_2(esk1_0,X4)),X2))!=k10_xtuple_0(esk3_0)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7974, negated_conjecture, (m1_subset_1(esk12_2(X1,k9_metric_1(esk1_0,X2,X3)),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k9_metric_1(esk1_0,X2,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|X4!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X3)|~m1_subset_1(k9_metric_1(esk1_0,X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_8220, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|~v1_xboole_0(k8_relat_1(esk3_0,X1))|X1!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_8310, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk14_2(k8_relset_1(u1_struct_0(esk1_0),X1,X2,X3),X4),k8_relat_1(esk3_0,X5))|r1_tarski(k8_relset_1(u1_struct_0(esk1_0),X1,X2,X3),X4)|X5!=k10_xtuple_0(esk3_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X1))))).#trainneg
cnf(i_0_7505, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk12_2(X1,esk7_3(X2,esk1_0,X3)),k8_relat_1(esk3_0,X4))|v3_pre_topc(esk7_3(X2,esk1_0,X3),k3_pcomps_1(X1))|v5_pre_topc(X3,X2,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|v2_struct_0(X2)|X4!=k10_xtuple_0(esk3_0)|~v1_xreal_0(esk5_3(X2,esk1_0,X3))|~m1_subset_1(esk6_3(X2,esk1_0,X3),u1_struct_0(esk1_0))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~m1_subset_1(esk7_3(X2,esk1_0,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(esk12_2(X1,esk7_3(X2,esk1_0,X3)),u1_struct_0(esk1_0))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X3)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_7324, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(X2,X1)|u1_struct_0(esk2_0)=k1_xboole_0|~v4_relat_1(esk3_0,X1)|X3!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,k9_xtuple_0(esk3_0)))).#trainneg
cnf(i_0_6257, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),X4)|u1_struct_0(esk2_0)=k1_xboole_0|~v4_relat_1(esk3_0,X4)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6656, plain, (r2_hidden(esk12_2(X1,k8_relat_1(X2,X3)),k8_relat_1(X2,X4))|v3_pre_topc(k8_relat_1(X2,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|X4!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~m1_subset_1(k8_relat_1(X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_7887, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X4),k3_pcomps_1(X3))|v2_struct_0(X3)|X1!=k1_funct_1(esk3_0,esk12_2(X3,k8_relat_1(esk3_0,X4)))|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(k8_relat_1(esk3_0,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_8358, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk7_3(X1,esk1_0,X2),k1_zfmisc_1(k8_relat_1(esk3_0,X3)))|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|X3!=k10_xtuple_0(esk3_0)|~v1_xreal_0(esk5_3(X1,esk1_0,X2))|~m1_subset_1(esk6_3(X1,esk1_0,X2),u1_struct_0(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_6315, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,k1_funct_1(esk3_0,esk12_2(X3,k9_metric_1(esk1_0,X4,esk4_3(esk11_3(esk2_0,k9_metric_1(X5,X6,esk11_3(X5,X7,X6)),X8),X4,X8)))))),X7)|v2_struct_0(X5)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k9_metric_1(esk1_0,X4,esk4_3(esk11_3(esk2_0,k9_metric_1(X5,X6,esk11_3(X5,X7,X6)),X8),X4,X8)),k3_pcomps_1(X3))|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(X5,X6,esk11_3(X5,X7,X6)),X8),k6_numbers)|v2_struct_0(X3)|X2!=k10_xtuple_0(X1)|~r2_hidden(X6,X7)|~r2_hidden(k1_funct_1(esk3_0,esk12_2(X3,k9_metric_1(esk1_0,X4,esk4_3(esk11_3(esk2_0,k9_metric_1(X5,X6,esk11_3(X5,X7,X6)),X8),X4,X8)))),X2)|~v1_relat_1(X1)|~v3_pre_topc(X7,k3_pcomps_1(X5))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X5))))|~m1_subset_1(X6,u1_struct_0(X5))|~v1_funct_1(X1)|~l1_metric_1(X5)|~v9_metric_1(X5)|~v8_metric_1(X5)|~v7_metric_1(X5)|~v6_metric_1(X5)|X8!=k1_funct_1(esk3_0,X4)|~r2_hidden(esk12_2(X3,k9_metric_1(esk1_0,X4,esk4_3(esk11_3(esk2_0,k9_metric_1(X5,X6,esk11_3(X5,X7,X6)),X8),X4,X8))),u1_struct_0(esk1_0))|~r2_hidden(X8,k9_metric_1(X5,X6,esk11_3(X5,X7,X6)))|~v3_pre_topc(k9_metric_1(X5,X6,esk11_3(X5,X7,X6)),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(X5,X6,esk11_3(X5,X7,X6)),X8))|~m1_subset_1(k9_metric_1(esk1_0,X4,esk4_3(esk11_3(esk2_0,k9_metric_1(X5,X6,esk11_3(X5,X7,X6)),X8),X4,X8)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(k9_metric_1(X5,X6,esk11_3(X5,X7,X6)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_8860, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,k8_relat_1(esk3_0,X2))|v3_pre_topc(k9_metric_1(esk1_0,X1,X3),k3_pcomps_1(esk2_0))|X2!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X3)|~m1_subset_1(k9_metric_1(esk1_0,X1,X3),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8769, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k1_funct_1(esk3_0,esk8_3(esk3_0,X3,X4))|X2!=k10_xtuple_0(esk3_0)|X3!=k10_xtuple_0(esk3_0)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_8807, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(esk7_3(X1,esk1_0,X2),u1_struct_0(esk1_0))|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|~v1_xreal_0(esk5_3(X1,esk1_0,X2))|~m1_subset_1(esk6_3(X1,esk1_0,X2),u1_struct_0(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_8764, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,k9_xtuple_0(esk3_0))|v5_pre_topc(X3,X2,k3_pcomps_1(esk1_0))|v2_struct_0(X2)|~v1_xreal_0(esk5_3(X2,esk1_0,X3))|~m1_subset_1(esk6_3(X2,esk1_0,X3),u1_struct_0(esk1_0))|~r2_hidden(X1,esk7_3(X2,esk1_0,X3))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X3))).#trainneg
cnf(i_0_8691, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|m1_subset_1(k9_metric_1(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8681, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(esk7_3(X1,esk1_0,X2),X3)|r2_hidden(esk14_2(esk7_3(X1,esk1_0,X2),X3),k9_xtuple_0(esk3_0))|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|~v1_xreal_0(esk5_3(X1,esk1_0,X2))|~m1_subset_1(esk6_3(X1,esk1_0,X2),u1_struct_0(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_8784, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X3,k3_pcomps_1(esk1_0))|X1!=k1_funct_1(esk3_0,esk12_2(esk1_0,X3))|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_8659, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,k8_relat_1(esk3_0,X2))|v3_pre_topc(k9_metric_1(esk1_0,X1,X3),k3_pcomps_1(esk1_0))|X2!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X3)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8630, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,X3),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8651, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(esk7_3(X1,esk1_0,X2),X3)|m1_subset_1(esk14_2(esk7_3(X1,esk1_0,X2),X3),u1_struct_0(esk1_0))|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|~v1_xreal_0(esk5_3(X1,esk1_0,X2))|~m1_subset_1(esk6_3(X1,esk1_0,X2),u1_struct_0(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_8761, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,k9_xtuple_0(esk3_0))|r1_tarski(k9_metric_1(esk1_0,X1,X2),X3)|~v1_xreal_0(X2)|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,X2),X3),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8926, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(X4,X3,k3_pcomps_1(esk1_0))|v2_struct_0(X3)|X1!=k1_funct_1(esk3_0,esk6_3(X3,esk1_0,X4))|X2!=k10_xtuple_0(esk3_0)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk1_0))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(esk1_0))|~v1_funct_1(X4))).#trainneg
cnf(i_0_8663, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X3,X4),X5)|X1!=k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X3,X4),X5))|X2!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9144, negated_conjecture, (v1_xboole_0(k8_relat_1(esk3_0,X1))|r2_hidden(X2,k8_relat_1(esk3_0,X1))|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k10_xtuple_0(esk3_0)|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8633, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k1_funct_1(esk3_0,esk8_3(X3,X4,X5))|X2!=k10_xtuple_0(esk3_0)|X4!=k10_xtuple_0(X3)|~v4_relat_1(X3,u1_struct_0(esk1_0))|~r2_hidden(X5,X4)|~v1_relat_1(X3)|~v1_funct_1(X3))).#trainneg
cnf(i_0_9025, negated_conjecture, (v3_pre_topc(k8_relat_1(esk3_0,X1),k3_pcomps_1(X2))|r1_xxreal_0(esk11_3(X2,k9_xtuple_0(esk3_0),esk12_2(X2,k8_relat_1(esk3_0,X1))),k6_numbers)|v2_struct_0(X2)|u1_struct_0(esk2_0)=k1_xboole_0|~v1_xreal_0(esk11_3(X2,k9_xtuple_0(esk3_0),esk12_2(X2,k8_relat_1(esk3_0,X1))))|~m1_subset_1(k8_relat_1(esk3_0,X1),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(esk12_2(X2,k8_relat_1(esk3_0,X1)),k9_xtuple_0(esk3_0))|~v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(X2))|~m1_subset_1(k9_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(esk12_2(X2,k8_relat_1(esk3_0,X1)),u1_struct_0(X2)))).#trainneg
cnf(i_0_8923, negated_conjecture, (m1_subset_1(esk6_3(X1,esk1_0,X2),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_8970, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),k9_xtuple_0(esk3_0))|k8_relat_1(esk3_0,X4)!=k8_relat_1(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8720, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk12_2(X1,esk7_3(X2,esk1_0,X3)),k9_xtuple_0(esk3_0))|v3_pre_topc(esk7_3(X2,esk1_0,X3),k3_pcomps_1(X1))|v5_pre_topc(X3,X2,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|v2_struct_0(X2)|~v1_xreal_0(esk5_3(X2,esk1_0,X3))|~m1_subset_1(esk6_3(X2,esk1_0,X3),u1_struct_0(esk1_0))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~m1_subset_1(esk7_3(X2,esk1_0,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(esk12_2(X1,esk7_3(X2,esk1_0,X3)),u1_struct_0(esk1_0))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X3)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_8763, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,k9_xtuple_0(esk3_0))|v3_pre_topc(k9_metric_1(esk1_0,X1,X3),k3_pcomps_1(X2))|v2_struct_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(esk12_2(X2,k9_metric_1(esk1_0,X1,X3)),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X1,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_8679, negated_conjecture, (r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk5_3(X2,esk2_0,X3),X1,esk6_3(X2,esk2_0,X3))),X4)|r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk5_3(X2,esk2_0,X3),X1,esk6_3(X2,esk2_0,X3))),X4),X5)|v5_pre_topc(X3,X2,k3_pcomps_1(esk2_0))|r1_xxreal_0(esk5_3(X2,esk2_0,X3),k6_numbers)|v2_struct_0(X2)|u1_struct_0(esk2_0)=k1_xboole_0|X5!=k8_relat_1(esk3_0,esk7_3(X2,esk2_0,X3))|esk6_3(X2,esk2_0,X3)!=k1_funct_1(esk3_0,X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xreal_0(esk5_3(X2,esk2_0,X3))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk5_3(X2,esk2_0,X3),X1,esk6_3(X2,esk2_0,X3))),X4)),u1_struct_0(esk2_0))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~m1_subset_1(esk6_3(X2,esk2_0,X3),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~v1_xreal_0(esk4_3(esk5_3(X2,esk2_0,X3),X1,esk6_3(X2,esk2_0,X3))))).#trainneg
cnf(i_0_8766, negated_conjecture, (m1_subset_1(esk8_3(esk3_0,X1,X2),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_9008, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),k9_xtuple_0(esk3_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k8_relat_1(esk3_0,X4))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_9026, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(X1,X2,esk11_3(X1,u1_struct_0(esk1_0),X2)),k8_relat_1(esk3_0,X3))|v2_struct_0(X1)|X3!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,u1_struct_0(esk1_0))|~v3_pre_topc(u1_struct_0(esk1_0),k3_pcomps_1(X1))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_8840, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|r1_xxreal_0(X2,k6_numbers)|X1!=k8_relat_1(esk3_0,k9_metric_1(esk2_0,X3,X2))|X3!=k1_funct_1(esk3_0,esk12_2(esk1_0,X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(esk12_2(esk1_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8876, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk14_2(k8_relset_1(u1_struct_0(esk1_0),X1,X2,X3),X4),k9_xtuple_0(esk3_0))|r1_tarski(k8_relset_1(u1_struct_0(esk1_0),X1,X2,X3),X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X1))))).#trainneg
cnf(i_0_8927, negated_conjecture, (m1_subset_1(esk6_3(X1,esk1_0,X2),X3)|u1_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|~v4_relat_1(esk3_0,X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_8781, negated_conjecture, (m1_subset_1(esk12_2(esk1_0,X1),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_8706, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk14_2(k8_relat_1(X1,X2),X3),k9_xtuple_0(esk3_0))|r1_tarski(k8_relat_1(X1,X2),X3)|~v1_xreal_0(X5)|~m1_subset_1(X4,u1_struct_0(esk1_0))|k9_metric_1(esk1_0,X4,X5)!=k8_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8646, negated_conjecture, (r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,X2),X3),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,X2),X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8993, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(esk3_0,X3),X4),X2,X4)),X5),X1,X5)),X6))),k9_xtuple_0(esk3_0))|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(esk3_0,X3),X4),X2,X4)),X5),X1,X5)),X6)|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(esk3_0,X3),X4),X2,X4)),X5),k6_numbers)|r1_xxreal_0(esk11_3(esk2_0,k8_relat_1(esk3_0,X3),X4),k6_numbers)|X4!=k1_funct_1(esk3_0,X2)|X5!=k1_funct_1(esk3_0,X1)|~r2_hidden(X5,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(esk3_0,X3),X4),X2,X4)))|~r2_hidden(X4,k8_relat_1(esk3_0,X3))|~v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(esk3_0,X3),X4),X2,X4)),k3_pcomps_1(esk2_0))|~v3_pre_topc(k8_relat_1(esk3_0,X3),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(esk3_0,X3),X4),X2,X4)),X5))|~v1_xreal_0(esk11_3(esk2_0,k8_relat_1(esk3_0,X3),X4))|~m1_subset_1(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(esk3_0,X3),X4),X2,X4)),X5),X1,X5)),X6))),u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(esk3_0,X3),X4),X2,X4)),X5),X1,X5)),X6)),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(esk3_0,X3),X4),X2,X4)),X5),X1,X5)),X6)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(esk3_0,X3),X4),X2,X4)),X5),X1,X5)),X6),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(esk3_0,X3),X4),X2,X4)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(k8_relat_1(esk3_0,X3),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8823, negated_conjecture, (r2_hidden(esk8_3(X1,X2,esk8_3(X3,X4,X5)),X6)|u1_struct_0(esk2_0)=k1_xboole_0|X6!=k8_relat_1(X1,k9_xtuple_0(X3))|~r2_hidden(esk8_3(X1,X2,esk8_3(X3,X4,X5)),k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|X4!=k10_xtuple_0(X3)|~r2_hidden(esk8_3(X3,X4,X5),X2)|~r2_hidden(X5,u1_struct_0(esk1_0))|~r2_hidden(X5,X4)|~v1_relat_1(X3)|~v1_funct_1(X3))).#trainneg
cnf(i_0_8805, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,k9_metric_1(esk1_0,X2,X3))|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8755, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),k9_xtuple_0(esk3_0))|~v1_xreal_0(X5)|~m1_subset_1(X4,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k9_metric_1(esk1_0,X4,X5))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8699, negated_conjecture, (v3_pre_topc(u1_struct_0(esk2_0),k3_pcomps_1(esk2_0))|u1_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|~r2_hidden(esk6_3(X1,esk1_0,X2),X3)|~v3_pre_topc(X3,k3_pcomps_1(esk1_0))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_8634, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,X3),X4)|u1_struct_0(esk2_0)=k1_xboole_0|~v4_relat_1(esk3_0,X4)|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8980, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k8_relat_1(esk3_0,X2),X3),X1,X3)),X4)),k9_xtuple_0(esk3_0))|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k8_relat_1(esk3_0,X2),X3),X1,X3)),X4)|r1_xxreal_0(esk11_3(esk2_0,k8_relat_1(esk3_0,X2),X3),k6_numbers)|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(X3,k8_relat_1(esk3_0,X2))|~v3_pre_topc(k8_relat_1(esk3_0,X2),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k8_relat_1(esk3_0,X2),X3))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k8_relat_1(esk3_0,X2),X3),X1,X3)),X4)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k8_relat_1(esk3_0,X2),X3),X1,X3)),X4),u1_struct_0(esk1_0))|~m1_subset_1(k8_relat_1(esk3_0,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8801, negated_conjecture, (r2_hidden(X1,k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(X3,X2,k3_pcomps_1(esk1_0))|v2_struct_0(X2)|~r2_hidden(X1,esk7_3(X2,esk1_0,X3))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X3))).#trainneg
cnf(i_0_8835, negated_conjecture, (v3_pre_topc(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),k3_pcomps_1(esk1_0))|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|m1_subset_1(k1_funct_1(esk3_0,esk12_2(esk1_0,k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)))),X2)|u1_struct_0(esk2_0)=k1_xboole_0|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(esk12_2(esk1_0,k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3))),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9650, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,X1),X2)|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9638, negated_conjecture, (r1_tarski(X1,k8_relat_1(esk3_0,X2))|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(esk14_2(X1,k8_relat_1(esk3_0,X2)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9687, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,k9_xtuple_0(esk3_0))|v3_pre_topc(k9_metric_1(esk1_0,X1,X2),k3_pcomps_1(esk1_0))|~v1_xreal_0(X2)|~m1_subset_1(k9_metric_1(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9623, negated_conjecture, (m1_subset_1(k9_metric_1(esk1_0,X1,X2),k1_zfmisc_1(X3))|u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,k9_xtuple_0(esk3_0))|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9649, negated_conjecture, (r2_hidden(X1,k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9662, negated_conjecture, (v2_struct_0(X1)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X2),k3_pcomps_1(X1))|~r2_hidden(esk12_2(X1,k8_relat_1(esk3_0,X2)),k9_xtuple_0(esk3_0))|~v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(X1))|~m1_subset_1(k9_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(esk12_2(X1,k8_relat_1(esk3_0,X2)),u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|X2!=k10_xtuple_0(esk3_0)|~v1_xreal_0(esk11_3(X1,k9_xtuple_0(esk3_0),esk12_2(X1,k8_relat_1(esk3_0,X2))))|~m1_subset_1(k8_relat_1(esk3_0,X2),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)))))).#trainneg
cnf(i_0_10033, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),k9_xtuple_0(esk3_0))|v5_pre_topc(X5,X4,k3_pcomps_1(esk1_0))|v2_struct_0(X4)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,esk7_3(X4,esk1_0,X5))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_xreal_0(esk5_3(X4,esk1_0,X5))|~m1_subset_1(esk6_3(X4,esk1_0,X5),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~l1_pre_topc(X4)|~v2_pre_topc(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(esk1_0))))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(esk1_0))|~v1_funct_1(X5))).#trainneg
cnf(i_0_9673, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(esk7_3(X4,esk1_0,X5),k3_pcomps_1(X3))|v5_pre_topc(X5,X4,k3_pcomps_1(esk1_0))|v2_struct_0(X4)|v2_struct_0(X3)|X1!=k1_funct_1(esk3_0,esk12_2(X3,esk7_3(X4,esk1_0,X5)))|X2!=k10_xtuple_0(esk3_0)|~l1_pre_topc(X4)|~v2_pre_topc(X4)|~v1_xreal_0(esk5_3(X4,esk1_0,X5))|~m1_subset_1(esk7_3(X4,esk1_0,X5),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(esk1_0))))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(esk1_0))|~v1_funct_1(X5)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_9643, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k8_relat_1(X3,k8_relat_1(esk3_0,X4))|~r2_hidden(X1,k9_xtuple_0(X3))|~v1_relat_1(X3)|~v1_funct_1(X3)|X4!=k10_xtuple_0(esk3_0)|~m1_subset_1(k1_funct_1(X3,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10179, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk14_2(esk7_3(X1,esk1_0,X2),X3),k9_xtuple_0(esk3_0))|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|r1_tarski(esk7_3(X1,esk1_0,X2),X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_10187, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk12_2(X1,esk7_3(X2,esk1_0,X3)),k9_xtuple_0(esk3_0))|v5_pre_topc(X3,X2,k3_pcomps_1(esk1_0))|v2_struct_0(X2)|v3_pre_topc(esk7_3(X2,esk1_0,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X3)|~m1_subset_1(esk7_3(X2,esk1_0,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_10095, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),k9_xtuple_0(esk3_0))|r1_tarski(k9_metric_1(esk1_0,X3,X4),X5)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_xreal_0(X4)|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X3,X4),X5),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9658, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(esk3_0,esk8_3(X1,X2,X3)),X4)|X4!=k10_xtuple_0(esk3_0)|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_10218, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,esk14_2(k9_metric_1(X3,X4,esk11_3(X3,esk7_3(X5,esk1_0,X6),X4)),X7))),k9_xtuple_0(esk3_0))|v5_pre_topc(X6,X5,k3_pcomps_1(esk1_0))|v2_struct_0(X5)|r1_tarski(k9_metric_1(X3,X4,esk11_3(X3,esk7_3(X5,esk1_0,X6),X4)),X7)|v2_struct_0(X3)|~l1_pre_topc(X5)|~v2_pre_topc(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(esk1_0))))|~v1_funct_2(X6,u1_struct_0(X5),u1_struct_0(esk1_0))|~v1_funct_1(X6)|X2!=k10_xtuple_0(X1)|~r2_hidden(esk14_2(k9_metric_1(X3,X4,esk11_3(X3,esk7_3(X5,esk1_0,X6),X4)),X7),X2)|~r2_hidden(X4,esk7_3(X5,esk1_0,X6))|~v1_relat_1(X1)|~v3_pre_topc(esk7_3(X5,esk1_0,X6),k3_pcomps_1(X3))|~m1_subset_1(esk7_3(X5,esk1_0,X6),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X4,u1_struct_0(X3))|~v1_funct_1(X1)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_10051, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,esk12_2(X3,k9_metric_1(esk1_0,X4,X5)))),k9_xtuple_0(esk3_0))|v3_pre_topc(k9_metric_1(esk1_0,X4,X5),k3_pcomps_1(X3))|v2_struct_0(X3)|X2!=k10_xtuple_0(X1)|~r2_hidden(esk12_2(X3,k9_metric_1(esk1_0,X4,X5)),X2)|~v1_relat_1(X1)|~v1_xreal_0(X5)|~m1_subset_1(X4,u1_struct_0(esk1_0))|~v1_funct_1(X1)|~m1_subset_1(k9_metric_1(esk1_0,X4,X5),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_10220, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),k9_xtuple_0(esk3_0))|v5_pre_topc(X4,X5,k3_pcomps_1(esk1_0))|v2_struct_0(X5)|~l1_pre_topc(X5)|~v2_pre_topc(X5)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(esk1_0))))|~v1_funct_2(X4,u1_struct_0(X5),u1_struct_0(esk1_0))|~v1_funct_1(X4)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,esk7_3(X5,esk1_0,X4))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_9707, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,esk14_2(k8_relat_1(esk3_0,X3),X4))),k9_xtuple_0(esk3_0))|r1_tarski(k8_relat_1(esk3_0,X3),X4)|X2!=k10_xtuple_0(X1)|~r2_hidden(esk14_2(k8_relat_1(esk3_0,X3),X4),X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_10178, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk14_2(k9_metric_1(X1,X2,esk11_3(X1,esk7_3(X3,esk1_0,X4),X2)),X5),k9_xtuple_0(esk3_0))|v5_pre_topc(X4,X3,k3_pcomps_1(esk1_0))|v2_struct_0(X3)|r1_tarski(k9_metric_1(X1,X2,esk11_3(X1,esk7_3(X3,esk1_0,X4),X2)),X5)|v2_struct_0(X1)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk1_0))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(esk1_0))|~v1_funct_1(X4)|~r2_hidden(X2,esk7_3(X3,esk1_0,X4))|~v3_pre_topc(esk7_3(X3,esk1_0,X4),k3_pcomps_1(X1))|~m1_subset_1(esk7_3(X3,esk1_0,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_9536, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,k9_metric_1(esk1_0,X2,X3))|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9691, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,k9_xtuple_0(esk3_0))|v3_pre_topc(k9_metric_1(esk1_0,X1,X2),k3_pcomps_1(X3))|v2_struct_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(k9_metric_1(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_9952, negated_conjecture, (r2_hidden(esk8_3(X1,X2,esk8_3(X3,X4,X5)),X6)|v2_struct_0(X7)|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X8,X6)|~v3_pre_topc(X6,k3_pcomps_1(X7))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X7))))|~m1_subset_1(X8,u1_struct_0(X7))|~l1_metric_1(X7)|~v9_metric_1(X7)|~v8_metric_1(X7)|~v7_metric_1(X7)|~v6_metric_1(X7)|k9_metric_1(X7,X8,esk11_3(X7,X6,X8))!=k8_relat_1(X1,k9_xtuple_0(X3))|X2!=k10_xtuple_0(X1)|X4!=k10_xtuple_0(X3)|~r2_hidden(esk8_3(X3,X4,X5),X2)|~r2_hidden(X5,u1_struct_0(esk1_0))|~r2_hidden(X5,X4)|~v1_relat_1(X1)|~v1_relat_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_10052, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,esk12_2(X3,esk7_3(X4,esk1_0,X5)))),k9_xtuple_0(esk3_0))|v3_pre_topc(esk7_3(X4,esk1_0,X5),k3_pcomps_1(X3))|v5_pre_topc(X5,X4,k3_pcomps_1(esk1_0))|v2_struct_0(X3)|v2_struct_0(X4)|X2!=k10_xtuple_0(X1)|~r2_hidden(esk12_2(X3,esk7_3(X4,esk1_0,X5)),X2)|~v1_relat_1(X1)|~v1_xreal_0(esk5_3(X4,esk1_0,X5))|~m1_subset_1(esk6_3(X4,esk1_0,X5),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~l1_pre_topc(X4)|~v2_pre_topc(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(esk1_0))))|~m1_subset_1(esk7_3(X4,esk1_0,X5),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(esk12_2(X3,esk7_3(X4,esk1_0,X5)),u1_struct_0(esk1_0))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(esk1_0))|~v1_funct_1(X5)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_9632, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(esk7_3(X1,esk1_0,X2),X3)|r2_hidden(X4,X5)|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|X4!=k1_funct_1(esk3_0,esk14_2(esk7_3(X1,esk1_0,X2),X3))|X5!=k10_xtuple_0(esk3_0)|~v1_xreal_0(esk5_3(X1,esk1_0,X2))|~m1_subset_1(esk6_3(X1,esk1_0,X2),u1_struct_0(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_10097, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),k9_xtuple_0(esk3_0))|v3_pre_topc(k9_metric_1(esk1_0,X3,X5),k3_pcomps_1(X4))|v2_struct_0(X4)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_xreal_0(X5)|~m1_subset_1(esk12_2(X4,k9_metric_1(esk1_0,X3,X5)),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~m1_subset_1(k9_metric_1(esk1_0,X3,X5),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X4))))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~l1_metric_1(X4)|~v9_metric_1(X4)|~v8_metric_1(X4)|~v7_metric_1(X4)|~v6_metric_1(X4))).#trainneg
cnf(i_0_10043, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,esk14_2(k9_metric_1(esk1_0,X3,X4),X5))),k9_xtuple_0(esk3_0))|r1_tarski(k9_metric_1(esk1_0,X3,X4),X5)|X2!=k10_xtuple_0(X1)|~r2_hidden(esk14_2(k9_metric_1(esk1_0,X3,X4),X5),X2)|~v1_relat_1(X1)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_9783, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),k9_xtuple_0(esk3_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|X4!=k10_xtuple_0(esk3_0)|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10261, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(k1_funct_1(esk3_0,esk12_2(esk1_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X1,X3),X2,X3)))),X1)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X1,X3),X2,X3)),k3_pcomps_1(esk1_0))|r1_xxreal_0(esk11_3(esk2_0,X1,X3),k6_numbers)|X3!=k1_funct_1(esk3_0,X2)|~r2_hidden(X3,X1)|~v3_pre_topc(X1,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X1,X3))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X1,X3),X2,X3)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9629, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(esk3_0,esk6_3(X1,esk1_0,X2)),X3)|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|X3!=k10_xtuple_0(esk3_0)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_10300, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,X1),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)=k1_xboole_0|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10396, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,k9_xtuple_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10458, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,k8_relat_1(X2,k8_relat_1(esk3_0,X3)))|X3!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,k9_xtuple_0(X2))|~v1_relat_1(X2)|~m1_subset_1(k1_funct_1(X2,X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_10399, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X1),k3_pcomps_1(X2))|v2_struct_0(X2)|X1!=k10_xtuple_0(esk3_0)|~v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(X2))|~m1_subset_1(k8_relat_1(esk3_0,X1),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(k9_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~r2_hidden(esk12_2(X2,k8_relat_1(esk3_0,X1)),k9_xtuple_0(esk3_0))|~m1_subset_1(esk12_2(X2,k8_relat_1(esk3_0,X1)),u1_struct_0(X2)))).#trainneg
cnf(i_0_10885, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),k9_xtuple_0(esk3_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_10301, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,X1),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)=k1_xboole_0|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10779, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|m1_subset_1(esk12_2(X1,k9_metric_1(esk1_0,X2,X3)),u1_struct_0(esk1_0))|v3_pre_topc(k9_metric_1(esk1_0,X2,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_10348, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relat_1(X1,X2),k8_relat_1(esk3_0,X3))|X3!=k10_xtuple_0(esk3_0)|~v4_relat_1(X1,u1_struct_0(esk1_0))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_10347, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k2_relset_1(u1_struct_0(esk1_0),X1),k8_relat_1(esk3_0,X2))|X2!=k10_xtuple_0(esk3_0)|~v5_relat_1(X1,u1_struct_0(esk1_0))|~v1_relat_1(X1))).#trainneg
cnf(i_0_10469, negated_conjecture, (m1_subset_1(esk14_2(esk7_3(X1,esk1_0,X2),X3),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(esk7_3(X1,esk1_0,X2),X3)|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_10338, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,X1),X2)|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10464, negated_conjecture, (m1_subset_1(esk14_2(esk7_3(X1,esk1_0,X2),X3),X4)|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(esk7_3(X1,esk1_0,X2),X3)|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|~v4_relat_1(esk3_0,X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_10471, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk14_2(esk7_3(X1,esk1_0,X2),X3),X4)|r1_tarski(esk7_3(X1,esk1_0,X2),X3)|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|X4!=k8_relat_1(esk3_0,X5)|X5!=k10_xtuple_0(esk3_0)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_10153, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(esk8_3(X2,X3,X4),X1)|u1_struct_0(esk2_0)=k1_xboole_0|X3!=k10_xtuple_0(X2)|~v4_relat_1(X2,u1_struct_0(esk1_0))|~v4_relat_1(esk3_0,X1)|~r2_hidden(X4,X3)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_9846, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(esk6_3(X2,esk1_0,X3),X1)|u1_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(X3,X2,k3_pcomps_1(esk1_0))|v2_struct_0(X2)|~v4_relat_1(esk3_0,X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X3))).#trainneg
cnf(i_0_10480, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(esk7_3(X4,esk1_0,X5),k3_pcomps_1(X3))|v5_pre_topc(X5,X4,k3_pcomps_1(esk1_0))|v2_struct_0(X3)|v2_struct_0(X4)|X1!=k1_funct_1(esk3_0,esk12_2(X3,esk7_3(X4,esk1_0,X5)))|X2!=k10_xtuple_0(esk3_0)|~l1_pre_topc(X4)|~v2_pre_topc(X4)|~m1_subset_1(esk7_3(X4,esk1_0,X5),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(esk1_0))))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(esk1_0))|~v1_funct_1(X5)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_10194, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,esk7_3(X2,esk1_0,X3),X4),X1,X4)),X5)),k9_xtuple_0(esk3_0))|v5_pre_topc(X3,X2,k3_pcomps_1(esk1_0))|v2_struct_0(X2)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,esk7_3(X2,esk1_0,X3),X4),X1,X4)),X5)|r1_xxreal_0(esk11_3(esk2_0,esk7_3(X2,esk1_0,X3),X4),k6_numbers)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X3)|X4!=k1_funct_1(esk3_0,X1)|~r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,esk7_3(X2,esk1_0,X3),X4),X1,X4)),X5),u1_struct_0(esk1_0))|~r2_hidden(X4,esk7_3(X2,esk1_0,X3))|~v3_pre_topc(esk7_3(X2,esk1_0,X3),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,esk7_3(X2,esk1_0,X3),X4))|~m1_subset_1(esk7_3(X2,esk1_0,X3),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10460, negated_conjecture, (m1_subset_1(esk14_2(esk7_3(X1,esk1_0,X2),X3),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(esk7_3(X1,esk1_0,X2),X3)|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_9631, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,X2),X3)|r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,X2),X3)),X4)|X4!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10028, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,u1_struct_0(esk1_0))|r1_tarski(k9_metric_1(esk1_0,X1,X2),X3)|~v1_xreal_0(X2)|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,X2),X3),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10059, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,esk6_3(X3,esk1_0,X4))),k9_xtuple_0(esk3_0))|v5_pre_topc(X4,X3,k3_pcomps_1(esk1_0))|v2_struct_0(X3)|X2!=k10_xtuple_0(X1)|~r2_hidden(esk6_3(X3,esk1_0,X4),X2)|~v1_relat_1(X1)|~v1_xreal_0(esk5_3(X3,esk1_0,X4))|~m1_subset_1(X5,u1_struct_0(esk1_0))|~v1_funct_1(X1)|~r2_hidden(X5,esk7_3(X3,esk1_0,X4))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk1_0))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(esk1_0))|~v1_funct_1(X4))).#trainneg
cnf(i_0_10783, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(esk1_0))|~v1_xreal_0(X5)|~m1_subset_1(X4,u1_struct_0(esk1_0))|k9_metric_1(esk1_0,X4,X5)!=k8_relat_1(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_9984, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk12_2(X1,k9_metric_1(esk1_0,X2,X3)),u1_struct_0(esk1_0))|v3_pre_topc(k9_metric_1(esk1_0,X2,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_9961, negated_conjecture, (r2_hidden(esk8_3(X1,X2,X3),k9_metric_1(esk2_0,X4,X5))|r1_xxreal_0(X5,k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|X4!=k1_funct_1(esk3_0,X6)|X7!=k10_xtuple_0(esk3_0)|~r2_hidden(esk8_3(X1,X2,X3),X7)|~v1_xreal_0(X5)|~m1_subset_1(esk8_3(esk3_0,X7,esk8_3(X1,X2,X3)),u1_struct_0(esk1_0))|~m1_subset_1(X6,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))|~m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(esk2_0))|k9_metric_1(esk1_0,X6,esk4_3(X5,X6,X4))!=k8_relat_1(esk3_0,k9_xtuple_0(X1))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_10757, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(X1,X2,esk11_3(X1,esk7_3(X3,esk1_0,X4),X2)),k8_relat_1(esk3_0,X5))|v5_pre_topc(X4,X3,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|v2_struct_0(X3)|X5!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,esk7_3(X3,esk1_0,X4))|~v3_pre_topc(esk7_3(X3,esk1_0,X4),k3_pcomps_1(X1))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(esk7_3(X3,esk1_0,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk1_0))))|~m1_subset_1(X2,u1_struct_0(X1))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(esk1_0))|~v1_funct_1(X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_10193, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,esk7_3(X2,esk1_0,X3),X4),X1,X4)),X5)),k9_xtuple_0(esk3_0))|v5_pre_topc(X3,X2,k3_pcomps_1(esk1_0))|v2_struct_0(X2)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,esk7_3(X2,esk1_0,X3),X4),X1,X4)),X5)|r1_xxreal_0(esk11_3(esk2_0,esk7_3(X2,esk1_0,X3),X4),k6_numbers)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X3)|X4!=k1_funct_1(esk3_0,X1)|~r2_hidden(X4,esk7_3(X2,esk1_0,X3))|~v3_pre_topc(esk7_3(X2,esk1_0,X3),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,esk7_3(X2,esk1_0,X3),X4))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,esk7_3(X2,esk1_0,X3),X4),X1,X4)),X5)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,esk7_3(X2,esk1_0,X3),X4),X1,X4)),X5),u1_struct_0(esk1_0))|~m1_subset_1(esk7_3(X2,esk1_0,X3),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9736, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,esk8_3(esk3_0,X3,X4))),k9_xtuple_0(esk3_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(esk8_3(esk3_0,X3,X4),X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|X3!=k10_xtuple_0(esk3_0)|~r2_hidden(X4,u1_struct_0(esk1_0))|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_8968, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),k9_xtuple_0(esk3_0))|k8_relat_1(esk3_0,X4)!=k8_relat_1(X1,u1_struct_0(esk2_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k10_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_9168, negated_conjecture, (v1_xboole_0(k8_relat_1(esk3_0,X1))|u1_struct_0(esk2_0)=k1_xboole_0|~v1_xboole_0(k9_xtuple_0(esk3_0)))).#trainneg
cnf(i_0_9134, negated_conjecture, (m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,esk14_2(k9_xtuple_0(esk3_0),X3))),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_xtuple_0(esk3_0),X3)|X2!=k10_xtuple_0(X1)|~r2_hidden(esk14_2(k9_xtuple_0(esk3_0),X3),X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_10472, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(esk3_0,esk14_2(esk7_3(X1,esk1_0,X2),X3)),X4)|r1_tarski(esk7_3(X1,esk1_0,X2),X3)|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|X4!=k10_xtuple_0(esk3_0)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_8723, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),k9_xtuple_0(esk3_0))|~v1_xreal_0(X5)|~m1_subset_1(X4,u1_struct_0(esk1_0))|k9_metric_1(esk1_0,X4,X5)!=k8_relat_1(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6521, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4),X5)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)|r1_xxreal_0(X2,k6_numbers)|X5!=k8_relat_1(esk3_0,X6)|X6!=k10_xtuple_0(esk3_0)|k8_relat_1(esk3_0,k9_metric_1(esk2_0,X3,X2))!=u1_struct_0(esk1_0)|X3!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_8489, plain, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,esk14_2(k8_relat_1(X3,X4),X5))),k8_relat_1(X3,X6))|r1_tarski(k8_relat_1(X3,X4),X5)|X6!=k10_xtuple_0(X3)|X2!=k10_xtuple_0(X1)|~r2_hidden(esk14_2(k8_relat_1(X3,X4),X5),X2)|~v1_relat_1(X3)|~v1_relat_1(X1)|~v1_funct_1(X3)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7844, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(esk3_0,esk12_2(X1,k9_xtuple_0(esk3_0))),X2)|v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(X1))|v2_struct_0(X1)|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(k9_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_11079, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k1_funct_1(esk3_0,X3)|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11087, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,X2)|X2!=k8_relat_1(esk3_0,X3)|X3!=k10_xtuple_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11061, negated_conjecture, (r1_tarski(X1,k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|~m1_subset_1(esk14_2(X1,k9_xtuple_0(esk3_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11074, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k8_relat_1(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X1,k9_xtuple_0(X3))|~v1_relat_1(X3)|~v1_funct_1(X3)|~m1_subset_1(k1_funct_1(X3,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11093, negated_conjecture, (r1_tarski(X1,k8_relat_1(X2,k8_relat_1(esk3_0,X3)))|u1_struct_0(esk2_0)=k1_xboole_0|X3!=k10_xtuple_0(esk3_0)|~r2_hidden(esk14_2(X1,k8_relat_1(X2,k8_relat_1(esk3_0,X3))),k9_xtuple_0(X2))|~v1_relat_1(X2)|~m1_subset_1(k1_funct_1(X2,esk14_2(X1,k8_relat_1(X2,k8_relat_1(esk3_0,X3)))),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_11081, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|~v4_relat_1(esk3_0,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11121, negated_conjecture, (v5_pre_topc(esk3_0,k3_pcomps_1(esk2_0),k3_pcomps_1(X1))|v2_struct_0(X1)|u1_struct_0(esk2_0)=k1_xboole_0|~v2_pre_topc(k3_pcomps_1(esk2_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(k3_pcomps_1(X1)))))|~v1_funct_2(esk3_0,u1_struct_0(esk2_0),u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|esk7_3(k3_pcomps_1(esk2_0),X1,esk3_0)!=k10_xtuple_0(esk3_0)|~v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(esk2_0))|~m1_subset_1(k8_relat_1(esk3_0,esk7_3(k3_pcomps_1(esk2_0),X1,esk3_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(k9_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_11593, negated_conjecture, (v3_pre_topc(k8_relat_1(esk3_0,X1),k3_pcomps_1(X2))|r1_xxreal_0(esk11_3(X2,esk7_3(X3,esk1_0,X4),esk12_2(X2,k8_relat_1(esk3_0,X1))),k6_numbers)|v2_struct_0(X2)|u1_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(X4,X3,k3_pcomps_1(esk1_0))|v2_struct_0(X3)|~v1_xreal_0(esk11_3(X2,esk7_3(X3,esk1_0,X4),esk12_2(X2,k8_relat_1(esk3_0,X1))))|~m1_subset_1(k8_relat_1(esk3_0,X1),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(esk12_2(X2,k8_relat_1(esk3_0,X1)),esk7_3(X3,esk1_0,X4))|~v3_pre_topc(esk7_3(X3,esk1_0,X4),k3_pcomps_1(X2))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(esk7_3(X3,esk1_0,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk1_0))))|~m1_subset_1(esk12_2(X2,k8_relat_1(esk3_0,X1)),u1_struct_0(X2))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(esk1_0))|~v1_funct_1(X4))).#trainneg
cnf(i_0_11141, negated_conjecture, (r2_hidden(esk8_3(X1,X2,X3),X4)|u1_struct_0(esk2_0)=k1_xboole_0|X4!=k8_relat_1(X1,k9_xtuple_0(esk3_0))|~r2_hidden(esk8_3(X1,X2,X3),k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11137, negated_conjecture, (m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_11210, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(esk7_3(X1,esk1_0,X2),X3)|r2_hidden(esk14_2(esk7_3(X1,esk1_0,X2),X3),k8_relat_1(esk3_0,X4))|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|X4!=k10_xtuple_0(esk3_0)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_11173, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|m1_subset_1(k1_funct_1(esk3_0,X1),k10_xtuple_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11161, negated_conjecture, (m1_subset_1(k8_relat_1(X1,X2),k1_zfmisc_1(k8_relat_1(esk3_0,X3)))|u1_struct_0(esk2_0)=k1_xboole_0|X3!=k10_xtuple_0(esk3_0)|~v4_relat_1(X1,u1_struct_0(esk1_0))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_11086, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11131, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k1_funct_1(esk3_0,k1_funct_1(X3,esk8_3(X3,X4,X5)))|X2!=k10_xtuple_0(esk3_0)|X4!=k10_xtuple_0(X3)|~r2_hidden(X5,X4)|~v1_relat_1(X3)|~m1_subset_1(X5,u1_struct_0(esk1_0))|~v1_funct_1(X3))).#trainneg
cnf(i_0_11218, negated_conjecture, (v5_pre_topc(X1,X2,k3_pcomps_1(X3))|m1_subset_1(esk8_3(X4,X5,X6),u1_struct_0(k3_pcomps_1(X3)))|v2_struct_0(X3)|v2_struct_0(X2)|u1_struct_0(esk2_0)=k1_xboole_0|v1_xboole_0(esk7_3(X2,X3,X1))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X3)))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X3)))|~v1_funct_1(X1)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)|X5!=k10_xtuple_0(X4)|~v4_relat_1(X4,u1_struct_0(esk1_0))|~v4_relat_1(esk3_0,esk7_3(X2,X3,X1))|~r2_hidden(X6,X5)|~v1_relat_1(X4)|~v1_funct_1(X4))).#trainneg
cnf(i_0_11166, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k10_xtuple_0(X1),k8_relat_1(esk3_0,X2))|X2!=k10_xtuple_0(esk3_0)|~v5_relat_1(X1,u1_struct_0(esk1_0))|~v1_relat_1(X1))).#trainneg
cnf(i_0_11118, negated_conjecture, (v5_pre_topc(esk3_0,k3_pcomps_1(X1),k3_pcomps_1(X2))|v2_struct_0(X2)|v2_struct_0(k3_pcomps_1(X1))|u1_struct_0(esk2_0)=k1_xboole_0|v2_struct_0(X1)|~l1_pre_topc(k3_pcomps_1(X1))|~v2_pre_topc(k3_pcomps_1(X1))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))))|~v1_funct_2(esk3_0,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(k3_pcomps_1(X2)))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|esk7_3(k3_pcomps_1(X1),X2,esk3_0)!=k10_xtuple_0(esk3_0)|~v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(X1))|~m1_subset_1(k8_relat_1(esk3_0,esk7_3(k3_pcomps_1(X1),X2,esk3_0)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(k9_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_11139, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),X4)|X4!=k8_relat_1(esk3_0,X5)|X5!=k10_xtuple_0(esk3_0)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_11063, negated_conjecture, (r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)|r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),X5)|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|X5!=k8_relat_1(esk3_0,X2)|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)),u1_struct_0(esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11128, negated_conjecture, (m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_11108, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,X2)|~v4_relat_1(esk3_0,X2)|X3!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,k9_xtuple_0(esk3_0))|~m1_subset_1(k1_funct_1(esk3_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11082, negated_conjecture, (m1_subset_1(k1_funct_1(X1,esk8_3(X1,X2,X3)),X4)|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(X1)|~v4_relat_1(esk3_0,X4)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11076, negated_conjecture, (m1_subset_1(X1,k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11208, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(esk14_2(esk7_3(X2,esk1_0,X3),X4),X1)|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(esk7_3(X2,esk1_0,X3),X4)|v5_pre_topc(X3,X2,k3_pcomps_1(esk1_0))|v2_struct_0(X2)|~v4_relat_1(esk3_0,X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X3))).#trainneg
cnf(i_0_11136, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v1_xboole_0(X1)|r2_hidden(k1_funct_1(X2,esk8_3(X2,X3,X4)),X1)|X5!=k10_xtuple_0(esk3_0)|~v4_relat_1(esk3_0,X1)|X3!=k10_xtuple_0(X2)|~r2_hidden(X4,X3)|~v1_relat_1(X2)|~m1_subset_1(X4,u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_11138, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_11640, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|~v1_xboole_0(k9_xtuple_0(esk3_0)))).#trainneg
cnf(i_0_11119, negated_conjecture, (v5_pre_topc(esk3_0,k3_pcomps_1(X1),k3_pcomps_1(esk1_0))|v2_struct_0(k3_pcomps_1(X1))|u1_struct_0(esk2_0)=k1_xboole_0|v2_struct_0(X1)|~l1_pre_topc(k3_pcomps_1(X1))|~v2_pre_topc(k3_pcomps_1(X1))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(esk1_0))))|~v1_funct_2(esk3_0,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(esk1_0))|esk7_3(k3_pcomps_1(X1),esk1_0,esk3_0)!=k10_xtuple_0(esk3_0)|~v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(X1))|~m1_subset_1(k8_relat_1(esk3_0,esk7_3(k3_pcomps_1(X1),esk1_0,esk3_0)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(k9_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_11592, negated_conjecture, (r2_hidden(X1,k8_relat_1(esk3_0,X2))|u1_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(X6,X5,k3_pcomps_1(esk1_0))|v2_struct_0(X5)|v2_struct_0(X3)|~r2_hidden(X1,k9_metric_1(X3,X4,esk11_3(X3,esk7_3(X5,esk1_0,X6),X4)))|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(X4,esk7_3(X5,esk1_0,X6))|~v3_pre_topc(esk7_3(X5,esk1_0,X6),k3_pcomps_1(X3))|~l1_pre_topc(X5)|~v2_pre_topc(X5)|~m1_subset_1(esk7_3(X5,esk1_0,X6),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(esk1_0))))|~m1_subset_1(X4,u1_struct_0(X3))|~v1_funct_2(X6,u1_struct_0(X5),u1_struct_0(esk1_0))|~v1_funct_1(X6)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_11124, negated_conjecture, (v5_pre_topc(esk3_0,k3_pcomps_1(esk1_0),k3_pcomps_1(X1))|v2_struct_0(X1)|u1_struct_0(esk2_0)=k1_xboole_0|~v2_pre_topc(k3_pcomps_1(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_pcomps_1(X1)))))|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|esk7_3(k3_pcomps_1(esk1_0),X1,esk3_0)!=k10_xtuple_0(esk3_0)|~v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(esk1_0)))).#trainneg
cnf(i_0_11050, negated_conjecture, (r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)),X2)|u1_struct_0(esk2_0)=k1_xboole_0|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11120, negated_conjecture, (v5_pre_topc(esk3_0,k3_pcomps_1(X1),k3_pcomps_1(esk2_0))|v2_struct_0(k3_pcomps_1(X1))|u1_struct_0(esk2_0)=k1_xboole_0|v2_struct_0(X1)|~l1_pre_topc(k3_pcomps_1(X1))|~v2_pre_topc(k3_pcomps_1(X1))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(k3_pcomps_1(X1)),u1_struct_0(esk2_0))|esk7_3(k3_pcomps_1(X1),esk2_0,esk3_0)!=k10_xtuple_0(esk3_0)|~v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(X1))|~m1_subset_1(k8_relat_1(esk3_0,esk7_3(k3_pcomps_1(X1),esk2_0,esk3_0)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(k9_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_11283, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),X4)|~v4_relat_1(esk3_0,X4)|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_11165, negated_conjecture, (r2_hidden(X1,k8_relat_1(esk3_0,X2))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,k2_relset_1(u1_struct_0(esk1_0),X3))|X2!=k10_xtuple_0(esk3_0)|~v5_relat_1(X3,u1_struct_0(esk1_0))|~v1_relat_1(X3))).#trainneg
cnf(i_0_7124, negated_conjecture, (m1_subset_1(esk6_3(X1,esk1_0,X2),k8_relat_1(esk3_0,X3))|u1_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|X3!=k10_xtuple_0(esk3_0)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_6222, negated_conjecture, (r2_hidden(esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4))),X5)|v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),k3_pcomps_1(X1))|r1_xxreal_0(X3,k6_numbers)|v2_struct_0(X1)|u1_struct_0(esk2_0)=k1_xboole_0|X5!=k8_relat_1(esk3_0,k9_metric_1(esk2_0,X4,X3))|X4!=k1_funct_1(esk3_0,X2)|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)))),u1_struct_0(esk2_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~r2_hidden(esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4))),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7157, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|r1_xxreal_0(esk11_3(esk2_0,X4,X5),k6_numbers)|~r2_hidden(X1,k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,X4,X5),X3,X5)))|X2!=k8_relat_1(esk3_0,X4)|X5!=k1_funct_1(esk3_0,X3)|~r2_hidden(X5,X4)|~v3_pre_topc(X4,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk4_3(esk11_3(esk2_0,X4,X5),X3,X5))|~v1_xreal_0(esk11_3(esk2_0,X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11906, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(X1,k8_relat_1(esk3_0,k8_relat_1(esk3_0,X2)))|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(esk14_2(X1,k8_relat_1(esk3_0,k8_relat_1(esk3_0,X2))),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(X1,k8_relat_1(esk3_0,k8_relat_1(esk3_0,X2)))),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11863, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,X2)|X2!=k8_relat_1(X3,u1_struct_0(esk1_0))|~r2_hidden(X1,k9_xtuple_0(X3))|~v1_relat_1(X3)|~m1_subset_1(k1_funct_1(X3,X1),u1_struct_0(esk1_0))|~v1_funct_1(X3))).#trainneg
cnf(i_0_12323, negated_conjecture, (r2_hidden(esk8_3(X1,X2,X3),X4)|u1_struct_0(esk2_0)=k1_xboole_0|X4!=k8_relat_1(X1,u1_struct_0(esk1_0))|~r2_hidden(esk8_3(X1,X2,X3),k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11862, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,k8_relat_1(X2,k9_xtuple_0(esk3_0)))|~r2_hidden(X1,k9_xtuple_0(X2))|~v1_relat_1(X2)|~m1_subset_1(k1_funct_1(X2,X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_11927, negated_conjecture, (v2_struct_0(X1)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X2),k3_pcomps_1(X1))|v5_pre_topc(X4,X3,k3_pcomps_1(esk1_0))|v2_struct_0(X3)|~r2_hidden(esk12_2(X1,k8_relat_1(esk3_0,X2)),esk7_3(X3,esk1_0,X4))|~v3_pre_topc(esk7_3(X3,esk1_0,X4),k3_pcomps_1(X1))|~m1_subset_1(esk7_3(X3,esk1_0,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(esk12_2(X1,k8_relat_1(esk3_0,X2)),u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|X2!=k10_xtuple_0(esk3_0)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~v1_xreal_0(esk11_3(X1,esk7_3(X3,esk1_0,X4),esk12_2(X1,k8_relat_1(esk3_0,X2))))|~m1_subset_1(k8_relat_1(esk3_0,X2),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk1_0))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(esk1_0))|~v1_funct_1(X4))).#trainneg
cnf(i_0_11909, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(X1,k8_relat_1(esk3_0,k8_relat_1(esk3_0,X2)))|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(X1,k8_relat_1(esk3_0,k8_relat_1(esk3_0,X2)))),u1_struct_0(esk1_0))|~m1_subset_1(esk14_2(X1,k8_relat_1(esk3_0,k8_relat_1(esk3_0,X2))),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11954, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk3_0))|r2_hidden(k1_funct_1(esk3_0,X1),k10_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11928, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),k8_relat_1(X1,k9_xtuple_0(esk3_0)))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_11994, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),k8_relat_1(esk3_0,X4))|X4!=k10_xtuple_0(esk3_0)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_12224, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(k1_funct_1(X2,esk8_3(X2,X3,X4)),X1)|u1_struct_0(esk2_0)=k1_xboole_0|X3!=k10_xtuple_0(X2)|~v4_relat_1(esk3_0,X1)|~r2_hidden(X4,X3)|~v1_relat_1(X2)|~m1_subset_1(X4,u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_12440, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X1,X3),X2,X3)),X4)),X1)|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X1,X3),X2,X3)),X4)|r1_xxreal_0(esk11_3(esk2_0,X1,X3),k6_numbers)|X3!=k1_funct_1(esk3_0,X2)|~r2_hidden(X3,X1)|~v3_pre_topc(X1,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X1,X3))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X1,X3),X2,X3)),X4),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11917, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_xtuple_0(X1),k8_relat_1(X1,k8_relat_1(esk3_0,X2)))|X2!=k10_xtuple_0(esk3_0)|~v1_relat_1(X1)|~m1_subset_1(k1_funct_1(X1,esk14_2(k9_xtuple_0(X1),k8_relat_1(X1,k8_relat_1(esk3_0,X2)))),u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_11990, negated_conjecture, (m1_subset_1(k10_xtuple_0(X1),k1_zfmisc_1(k8_relat_1(esk3_0,X2)))|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(esk3_0)|~v5_relat_1(X1,u1_struct_0(esk1_0))|~v1_relat_1(X1))).#trainneg
cnf(i_0_12324, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),X4)|X4!=k8_relat_1(X1,k8_relat_1(esk3_0,X5))|X5!=k10_xtuple_0(esk3_0)|~r2_hidden(esk8_3(X1,X2,X3),k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11845, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(X1,u1_struct_0(esk1_0))|~m1_subset_1(esk14_2(X1,u1_struct_0(esk1_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11905, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),k3_pcomps_1(esk1_0))|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3))!=k8_relat_1(esk3_0,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),k1_zfmisc_1(u1_struct_0(esk1_0)))|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(X3,X2)|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12050, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(X3,k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X4,esk4_3(esk11_3(esk2_0,X5,X6),X4,X6)),X7)|r1_xxreal_0(esk11_3(esk2_0,X5,X6),k6_numbers)|X2!=k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X4,esk4_3(esk11_3(esk2_0,X5,X6),X4,X6)),X7))|X8!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,X8)|~v1_xreal_0(X3)|~m1_subset_1(esk8_3(esk3_0,X8,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X4,esk4_3(esk11_3(esk2_0,X5,X6),X4,X6)),X7),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|k9_metric_1(esk1_0,esk8_3(esk3_0,X8,X1),esk4_3(X3,esk14_2(k9_metric_1(esk1_0,X4,esk4_3(esk11_3(esk2_0,X5,X6),X4,X6)),X7),X2))!=k8_relat_1(esk3_0,X5)|X6!=k1_funct_1(esk3_0,X4)|~r2_hidden(X6,X5)|~v3_pre_topc(X5,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X5,X6))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X4,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11903, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|X1!=k8_relat_1(esk3_0,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|X2!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_11865, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|X1!=k8_relat_1(esk3_0,k8_relat_1(esk3_0,X2))|~v3_pre_topc(k8_relat_1(esk3_0,X2),k3_pcomps_1(esk2_0))|~m1_subset_1(k8_relat_1(esk3_0,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(k1_funct_1(esk3_0,esk12_2(esk1_0,X1)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11991, negated_conjecture, (r2_hidden(X1,k8_relat_1(esk3_0,X2))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,k10_xtuple_0(X3))|X2!=k10_xtuple_0(esk3_0)|~v5_relat_1(X3,u1_struct_0(esk1_0))|~v1_relat_1(X3))).#trainneg
cnf(i_0_12321, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),X4)|~v4_relat_1(esk3_0,X4)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_11866, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|X1!=k8_relat_1(esk3_0,k8_relat_1(esk3_0,X2))|~v3_pre_topc(k8_relat_1(esk3_0,X2),k3_pcomps_1(esk2_0))|~m1_subset_1(k8_relat_1(esk3_0,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(k1_funct_1(esk3_0,esk12_2(esk1_0,X1)),k9_xtuple_0(esk3_0)))).#trainneg
cnf(i_0_12055, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,k1_funct_1(esk3_0,X2),X3))|r1_xxreal_0(X3,k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X4,esk4_3(esk11_3(esk2_0,X5,X6),X4,X6)),X7)|r1_xxreal_0(esk11_3(esk2_0,X5,X6),k6_numbers)|X1!=k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X4,esk4_3(esk11_3(esk2_0,X5,X6),X4,X6)),X7))|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X4,esk4_3(esk11_3(esk2_0,X5,X6),X4,X6)),X7),u1_struct_0(esk1_0))|k9_metric_1(esk1_0,X2,esk4_3(X3,esk14_2(k9_metric_1(esk1_0,X4,esk4_3(esk11_3(esk2_0,X5,X6),X4,X6)),X7),X1))!=k8_relat_1(esk3_0,X5)|X6!=k1_funct_1(esk3_0,X4)|~r2_hidden(X6,X5)|~v3_pre_topc(X5,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X5,X6))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X4,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11920, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relat_1(X1,X2),k8_relat_1(X1,k8_relat_1(esk3_0,X3)))|X3!=k10_xtuple_0(esk3_0)|~v1_relat_1(X1)|~m1_subset_1(k1_funct_1(X1,esk14_2(k8_relat_1(X1,X2),k8_relat_1(X1,k8_relat_1(esk3_0,X3)))),u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_12082, negated_conjecture, (r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)|u1_struct_0(esk2_0)=k1_xboole_0|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|X4!=k8_relat_1(esk3_0,X2)|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11961, negated_conjecture, (r2_hidden(X1,k8_relat_1(esk3_0,X2))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,k8_relat_1(X3,X4))|X2!=k10_xtuple_0(esk3_0)|~v4_relat_1(X3,u1_struct_0(esk1_0))|~v1_relat_1(X3)|~v1_funct_1(X3))).#trainneg
cnf(i_0_7746, negated_conjecture, (m1_subset_1(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),k1_zfmisc_1(X4))|u1_struct_0(esk2_0)=k1_xboole_0|r1_xxreal_0(X2,k6_numbers)|X4!=k8_relat_1(esk3_0,k9_metric_1(esk2_0,X3,X2))|X3!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_7448, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,esk14_2(k8_relat_1(esk3_0,X1),X2)),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relat_1(esk3_0,X1),X2))).#trainneg
cnf(i_0_8944, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk14_2(k8_relat_1(X1,X2),X3),k9_xtuple_0(esk3_0))|r1_tarski(k8_relat_1(X1,X2),X3)|k8_relat_1(esk3_0,X4)!=k8_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8998, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk14_2(k8_relat_1(X1,X2),X3)),k9_xtuple_0(esk3_0))|r1_tarski(k8_relat_1(X1,X2),X3)|k8_relat_1(esk3_0,X4)!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8660, negated_conjecture, (m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,X2),X3),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,X2),X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8413, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,k1_funct_1(esk3_0,X2),X3))|r1_xxreal_0(X3,k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X4,k3_pcomps_1(esk1_0))|X1!=k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk12_2(esk1_0,X4)))|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk12_2(esk1_0,X4)),u1_struct_0(esk1_0))|k9_metric_1(esk1_0,X2,esk4_3(X3,k1_funct_1(esk3_0,esk12_2(esk1_0,X4)),X1))!=k10_xtuple_0(esk3_0)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7257, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_xtuple_0(esk3_0),X1)),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_xtuple_0(esk3_0),X1))).#trainneg
cnf(i_0_8318, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk12_2(X1,k8_relset_1(u1_struct_0(esk1_0),X2,X3,X4)),k8_relat_1(esk3_0,X5))|v3_pre_topc(k8_relset_1(u1_struct_0(esk1_0),X2,X3,X4),k3_pcomps_1(X1))|v2_struct_0(X1)|X5!=k10_xtuple_0(esk3_0)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X2)))|~m1_subset_1(k8_relset_1(u1_struct_0(esk1_0),X2,X3,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_10042, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,esk14_2(k9_metric_1(X3,X4,esk11_3(X3,k9_metric_1(esk1_0,X5,X6),X4)),X7))),k9_xtuple_0(esk3_0))|r1_tarski(k9_metric_1(X3,X4,esk11_3(X3,k9_metric_1(esk1_0,X5,X6),X4)),X7)|v2_struct_0(X3)|X2!=k10_xtuple_0(X1)|~r2_hidden(esk14_2(k9_metric_1(X3,X4,esk11_3(X3,k9_metric_1(esk1_0,X5,X6),X4)),X7),X2)|~v1_relat_1(X1)|~v1_xreal_0(X6)|~m1_subset_1(X5,u1_struct_0(esk1_0))|~v1_funct_1(X1)|~r2_hidden(X4,k9_metric_1(esk1_0,X5,X6))|~v3_pre_topc(k9_metric_1(esk1_0,X5,X6),k3_pcomps_1(X3))|~m1_subset_1(k9_metric_1(esk1_0,X5,X6),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_6534, negated_conjecture, (m1_subset_1(esk14_2(k8_relat_1(esk3_0,X1),X2),k8_relat_1(esk3_0,X3))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relat_1(esk3_0,X1),X2)|X3!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_7989, negated_conjecture, (m1_subset_1(k9_metric_1(X1,X2,esk11_3(X1,u1_struct_0(esk1_0),X2)),k1_zfmisc_1(k9_xtuple_0(esk3_0)))|u1_struct_0(esk2_0)=k1_xboole_0|v2_struct_0(X1)|~r2_hidden(X2,u1_struct_0(esk1_0))|~v3_pre_topc(u1_struct_0(esk1_0),k3_pcomps_1(X1))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_6371, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(X2,X3,X4),X5),X1,X5)),X6)),u1_struct_0(X2))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(X2,X3,X4),X5),X1,X5)),X6)|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(X2,X3,X4),X5),k6_numbers)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_metric_1(X2)|X5!=k1_funct_1(esk3_0,X1)|~r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(X2,X3,X4),X5),X1,X5)),X6),u1_struct_0(esk1_0))|~r2_hidden(X5,k9_metric_1(X2,X3,X4))|~v3_pre_topc(k9_metric_1(X2,X3,X4),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(X2,X3,X4),X5))|~m1_subset_1(k9_metric_1(X2,X3,X4),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6390, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,X1),X2)|u1_struct_0(esk2_0)=k1_xboole_0|v1_xboole_0(k10_xtuple_0(esk3_0))|~v5_relat_1(esk3_0,X2)|~r2_hidden(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8937, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(esk1_0))|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6556, negated_conjecture, (m1_subset_1(esk14_2(k9_xtuple_0(esk3_0),X1),k8_relat_1(esk3_0,X2))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_xtuple_0(esk3_0),X1)|X2!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_7449, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk14_2(k8_relat_1(esk3_0,X1),X2)),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relat_1(esk3_0,X1),X2))).#trainneg
cnf(i_0_6367, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6))),k9_metric_1(esk2_0,X4,X3))|r1_xxreal_0(X3,k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6)|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),k6_numbers)|X4!=k1_funct_1(esk3_0,X2)|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6))),u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6)),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))|X5!=k1_funct_1(esk3_0,X1)|~r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6),u1_struct_0(esk1_0))|~r2_hidden(X5,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)))|~v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12734, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(esk2_0)))|u1_struct_0(esk2_0)=k1_xboole_0|k2_pcomps_1(esk2_0)!=k8_relat_1(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_12760, negated_conjecture, (r2_hidden(esk8_3(X1,X2,X3),X4)|v2_struct_0(X5)|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X6,X4)|~v3_pre_topc(X4,k3_pcomps_1(X5))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X5))))|~m1_subset_1(X6,u1_struct_0(X5))|~l1_metric_1(X5)|~v9_metric_1(X5)|~v8_metric_1(X5)|~v7_metric_1(X5)|~v6_metric_1(X5)|k9_metric_1(X5,X6,esk11_3(X5,X4,X6))!=k8_relat_1(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_12778, negated_conjecture, (r1_tarski(X1,k8_relat_1(X2,k9_xtuple_0(esk3_0)))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(esk14_2(X1,k8_relat_1(X2,k9_xtuple_0(esk3_0))),k9_xtuple_0(X2))|~v1_relat_1(X2)|~m1_subset_1(k1_funct_1(X2,esk14_2(X1,k8_relat_1(X2,k9_xtuple_0(esk3_0)))),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_13322, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,k8_relat_1(esk3_0,X2))|X2!=k10_xtuple_0(esk3_0)|~v4_relat_1(X3,u1_struct_0(esk1_0))|~v1_relat_1(X3)|~v1_funct_1(X3)|~r2_hidden(X1,k9_xtuple_0(X3))|~m1_subset_1(k1_funct_1(X3,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13032, negated_conjecture, (v5_pre_topc(esk3_0,k3_pcomps_1(esk1_0),k3_pcomps_1(X1))|v2_struct_0(X1)|u1_struct_0(esk2_0)=k1_xboole_0|~v2_pre_topc(k3_pcomps_1(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_pcomps_1(X1)))))|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(k3_pcomps_1(X1)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v3_pre_topc(esk7_3(k3_pcomps_1(esk1_0),X1,esk3_0),k3_pcomps_1(esk2_0))|~m1_subset_1(esk7_3(k3_pcomps_1(esk1_0),X1,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_13048, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),k8_relat_1(X1,k8_relat_1(esk3_0,X4)))|X4!=k10_xtuple_0(esk3_0)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_12702, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|k8_relat_1(esk3_0,X4)!=k8_relat_1(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_12801, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),k8_relat_1(X4,X5))|u1_struct_0(esk2_0)=k1_xboole_0|X5!=k10_xtuple_0(X4)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X4)|~v1_relat_1(X1)|~v1_funct_1(X4)|~v1_funct_1(X1)|~r2_hidden(X3,k9_xtuple_0(X4))|~m1_subset_1(k1_funct_1(X4,X3),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12753, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(esk1_0))|~v1_xreal_0(X5)|~m1_subset_1(X4,u1_struct_0(esk1_0))|k9_metric_1(esk1_0,X4,X5)!=k8_relat_1(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_12686, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,k8_relat_1(X2,u1_struct_0(esk1_0)))|~r2_hidden(X1,k9_xtuple_0(X2))|~v1_relat_1(X2)|~m1_subset_1(k1_funct_1(X2,X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_12822, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,X3),k8_relat_1(X1,k9_xtuple_0(esk3_0)))|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_12832, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,esk8_3(X3,X4,X5))),k9_xtuple_0(X3))|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(X1)|~r2_hidden(esk8_3(X3,X4,X5),X2)|~v1_relat_1(X3)|~v1_relat_1(X1)|~v1_funct_1(X3)|~v1_funct_1(X1)|X4!=k10_xtuple_0(X3)|~r2_hidden(X5,X4)|~m1_subset_1(X5,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13040, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(u1_struct_0(esk1_0),k8_relat_1(esk3_0,k8_relat_1(esk3_0,X1)))|X1!=k10_xtuple_0(esk3_0)|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(u1_struct_0(esk1_0),k8_relat_1(esk3_0,k8_relat_1(esk3_0,X1)))),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13326, negated_conjecture, (k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3))=k1_xboole_0|u1_struct_0(esk2_0)=k1_xboole_0|r1_xxreal_0(X2,k6_numbers)|k8_relat_1(esk3_0,k9_metric_1(esk2_0,X3,X2))!=k1_xboole_0|X3!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_12814, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,X1),X2)|u1_struct_0(esk2_0)=k1_xboole_0|v1_xboole_0(k10_xtuple_0(esk3_0))|~v5_relat_1(esk3_0,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12839, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),k8_relat_1(X1,u1_struct_0(esk1_0)))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_12767, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,X3),X4)|u1_struct_0(esk2_0)=k1_xboole_0|X4!=k8_relat_1(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_13244, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk12_2(X1,k8_relat_1(X2,X3)),k8_relat_1(esk3_0,X4))|v3_pre_topc(k8_relat_1(X2,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|X4!=k10_xtuple_0(esk3_0)|~v4_relat_1(X2,u1_struct_0(esk1_0))|~v1_relat_1(X2)|~v1_funct_1(X2)|~m1_subset_1(k8_relat_1(X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_13382, negated_conjecture, (r2_hidden(esk12_2(X1,k8_relset_1(u1_struct_0(esk1_0),X2,X3,X4)),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k8_relset_1(u1_struct_0(esk1_0),X2,X3,X4),k3_pcomps_1(X1))|v2_struct_0(X1)|X5!=k10_xtuple_0(esk3_0)|~m1_subset_1(k8_relset_1(u1_struct_0(esk1_0),X2,X3,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X2)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_13080, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(X1,k3_pcomps_1(esk1_0))|X1!=k8_relat_1(esk3_0,k8_relat_1(esk3_0,X2))|X2!=k10_xtuple_0(esk3_0)|~v3_pre_topc(k8_relat_1(esk3_0,X2),k3_pcomps_1(esk2_0))|~m1_subset_1(k8_relat_1(esk3_0,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_funct_1(esk3_0,esk12_2(esk1_0,X1)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12805, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X1),k3_pcomps_1(X2))|v5_pre_topc(X3,X4,k3_pcomps_1(esk1_0))|v2_struct_0(X4)|v2_struct_0(X2)|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(esk12_2(X2,k8_relat_1(esk3_0,X1)),esk7_3(X4,esk1_0,X3))|~v3_pre_topc(esk7_3(X4,esk1_0,X3),k3_pcomps_1(X2))|~l1_pre_topc(X4)|~v2_pre_topc(X4)|~m1_subset_1(esk7_3(X4,esk1_0,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(k8_relat_1(esk3_0,X1),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(esk1_0))))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(esk1_0))|~v1_funct_1(X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~m1_subset_1(esk12_2(X2,k8_relat_1(esk3_0,X1)),u1_struct_0(X2)))).#trainneg
cnf(i_0_12704, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),k9_xtuple_0(esk3_0))|k8_relat_1(esk3_0,X4)!=k8_relat_1(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_12836, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,esk8_3(X3,X4,X5))),k8_relat_1(X3,X6))|u1_struct_0(esk2_0)=k1_xboole_0|X6!=k10_xtuple_0(X3)|X2!=k10_xtuple_0(X1)|~r2_hidden(esk8_3(X3,X4,X5),X2)|~v1_relat_1(X3)|~v1_relat_1(X1)|~v1_funct_1(X3)|~v1_funct_1(X1)|X4!=k10_xtuple_0(X3)|~r2_hidden(X5,X4)|~m1_subset_1(X5,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12970, negated_conjecture, (r2_hidden(k1_funct_1(X1,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,X3),X4),X2,X4)),X5))),X3)|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,X3),X4),X2,X4)),X5)|r1_xxreal_0(esk11_3(esk2_0,k8_relat_1(X1,X3),X4),k6_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|X4!=k1_funct_1(esk3_0,X2)|~r2_hidden(X4,k8_relat_1(X1,X3))|~v3_pre_topc(k8_relat_1(X1,X3),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k8_relat_1(X1,X3),X4))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,X3),X4),X2,X4)),X5),u1_struct_0(esk1_0))|~m1_subset_1(k8_relat_1(X1,X3),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12736, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(esk1_0)))|u1_struct_0(esk2_0)=k1_xboole_0|k2_pcomps_1(esk1_0)!=k8_relat_1(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_13033, negated_conjecture, (v5_pre_topc(esk3_0,k3_pcomps_1(esk1_0),k3_pcomps_1(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|~v2_pre_topc(k3_pcomps_1(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v3_pre_topc(esk7_3(k3_pcomps_1(esk1_0),esk1_0,esk3_0),k3_pcomps_1(esk2_0))|~m1_subset_1(esk7_3(k3_pcomps_1(esk1_0),esk1_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_12747, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk8_3(X1,X2,X3)),k9_metric_1(esk2_0,X4,X5))|r1_xxreal_0(X5,k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|X4!=k1_funct_1(esk3_0,X6)|~v1_xreal_0(X5)|~m1_subset_1(k1_funct_1(esk3_0,esk8_3(X1,X2,X3)),u1_struct_0(esk2_0))|~m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(esk1_0))|~m1_subset_1(X6,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))|k9_metric_1(esk1_0,X6,esk4_3(X5,X6,X4))!=k8_relat_1(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_10115, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(u1_struct_0(esk2_0),k3_pcomps_1(esk2_0))|v5_pre_topc(X1,X2,k3_pcomps_1(esk1_0))|v2_struct_0(X2)|v3_pre_topc(esk7_3(X2,esk1_0,X1),k3_pcomps_1(X3))|v2_struct_0(X3)|~v3_pre_topc(k9_metric_1(esk1_0,esk12_2(X3,esk7_3(X2,esk1_0,X1)),esk5_3(X2,esk1_0,X1)),k3_pcomps_1(esk1_0))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~m1_subset_1(k9_metric_1(esk1_0,esk12_2(X3,esk7_3(X2,esk1_0,X1)),esk5_3(X2,esk1_0,X1)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~m1_subset_1(esk7_3(X2,esk1_0,X1),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_12795, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,X2)|~v4_relat_1(esk3_0,X2)|~r2_hidden(X1,k9_xtuple_0(esk3_0))|~m1_subset_1(k1_funct_1(esk3_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9722, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,esk12_2(X3,k8_relat_1(esk3_0,X4)))),k9_xtuple_0(esk3_0))|v3_pre_topc(k8_relat_1(esk3_0,X4),k3_pcomps_1(X3))|v2_struct_0(X3)|X2!=k10_xtuple_0(X1)|~r2_hidden(esk12_2(X3,k8_relat_1(esk3_0,X4)),X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(k8_relat_1(esk3_0,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_6875, negated_conjecture, (r1_tarski(esk7_3(X1,esk1_0,X2),u1_struct_0(esk1_0))|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_9732, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,esk8_3(X3,X4,esk14_2(k9_metric_1(X5,X6,esk11_3(X5,X7,X6)),X8)))),k9_xtuple_0(esk3_0))|r1_tarski(k9_metric_1(X5,X6,esk11_3(X5,X7,X6)),X8)|v2_struct_0(X5)|X2!=k10_xtuple_0(X1)|~r2_hidden(esk8_3(X3,X4,esk14_2(k9_metric_1(X5,X6,esk11_3(X5,X7,X6)),X8)),X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|k8_relat_1(esk3_0,X9)!=k8_relat_1(X3,X7)|X4!=k10_xtuple_0(X3)|~r2_hidden(esk14_2(k9_metric_1(X5,X6,esk11_3(X5,X7,X6)),X8),X4)|~r2_hidden(X6,X7)|~v1_relat_1(X3)|~v3_pre_topc(X7,k3_pcomps_1(X5))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X5))))|~m1_subset_1(X6,u1_struct_0(X5))|~v1_funct_1(X3)|~l1_metric_1(X5)|~v9_metric_1(X5)|~v8_metric_1(X5)|~v7_metric_1(X5)|~v6_metric_1(X5))).#trainneg
cnf(i_0_9022, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk12_2(X1,k9_metric_1(esk1_0,X2,X3)),k9_xtuple_0(esk3_0))|v3_pre_topc(k9_metric_1(esk1_0,X2,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(k9_metric_1(esk1_0,X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_7156, negated_conjecture, (m1_subset_1(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),k1_zfmisc_1(X4))|u1_struct_0(esk2_0)=k1_xboole_0|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|X4!=k8_relat_1(esk3_0,X2)|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk4_3(esk11_3(esk2_0,X2,X3),X1,X3))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10350, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k1_relset_1(u1_struct_0(esk1_0),X1),k8_relat_1(esk3_0,X2))|X2!=k10_xtuple_0(esk3_0)|~v4_relat_1(X1,u1_struct_0(esk1_0))|~v1_relat_1(X1))).#trainneg
cnf(i_0_10319, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(X3,k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k1_funct_1(esk3_0,k1_funct_1(esk3_0,X4))|X5!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,X5)|~v1_xreal_0(X3)|~m1_subset_1(esk8_3(esk3_0,X5,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,X4),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|k9_metric_1(esk1_0,esk8_3(esk3_0,X5,X1),esk4_3(X3,k1_funct_1(esk3_0,X4),X2))!=k10_xtuple_0(esk3_0)|~m1_subset_1(X4,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6482, negated_conjecture, (r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(X2,X1,X3)),X4)|r1_xxreal_0(X2,k6_numbers)|k8_relat_1(esk3_0,k9_metric_1(esk2_0,X3,X2))!=k9_xtuple_0(esk3_0)|X3!=k1_funct_1(esk3_0,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_9694, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk5_3(X2,esk2_0,X3),X1,esk6_3(X2,esk2_0,X3))),X4)|r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk5_3(X2,esk2_0,X3),X1,esk6_3(X2,esk2_0,X3))),X4),X5)|v5_pre_topc(X3,X2,k3_pcomps_1(esk2_0))|r1_xxreal_0(esk5_3(X2,esk2_0,X3),k6_numbers)|v2_struct_0(X2)|X5!=k8_relat_1(esk3_0,esk7_3(X2,esk2_0,X3))|esk6_3(X2,esk2_0,X3)!=k1_funct_1(esk3_0,X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xreal_0(esk5_3(X2,esk2_0,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~m1_subset_1(esk6_3(X2,esk2_0,X3),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk5_3(X2,esk2_0,X3),X1,esk6_3(X2,esk2_0,X3))),X4),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11919, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relat_1(esk3_0,X1),k8_relat_1(esk3_0,k8_relat_1(esk3_0,X2)))|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k8_relat_1(esk3_0,X1),k8_relat_1(esk3_0,k8_relat_1(esk3_0,X2)))),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6314, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)))),X5)|v2_struct_0(X3)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),k3_pcomps_1(X1))|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),k6_numbers)|v2_struct_0(X1)|~r2_hidden(X4,X5)|~v3_pre_topc(X5,k3_pcomps_1(X3))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)|X6!=k1_funct_1(esk3_0,X2)|~r2_hidden(esk12_2(X1,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6))),u1_struct_0(esk1_0))|~r2_hidden(X6,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)))|~v3_pre_topc(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_6430, negated_conjecture, (m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4),k1_zfmisc_1(u1_struct_0(esk1_0)))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),X4)|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|k2_pcomps_1(esk1_0)!=k8_relat_1(esk3_0,X2)|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk4_3(esk11_3(esk2_0,X2,X3),X1,X3))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8820, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,esk8_3(X3,X4,X5))),X6)|u1_struct_0(esk2_0)=k1_xboole_0|~v4_relat_1(X3,X6)|~v1_relat_1(X3)|X2!=k10_xtuple_0(X1)|X4!=k10_xtuple_0(X3)|~r2_hidden(esk8_3(X3,X4,X5),X2)|~r2_hidden(X5,u1_struct_0(esk1_0))|~r2_hidden(X5,X4)|~v1_relat_1(X1)|~v1_funct_1(X3)|~v1_funct_1(X1))).#trainneg
cnf(i_0_14115, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k9_metric_1(esk1_0,X4,X5),k3_pcomps_1(X3))|v2_struct_0(X3)|X1!=k1_funct_1(esk3_0,esk12_2(X3,k9_metric_1(esk1_0,X4,X5)))|X2!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X5)|~m1_subset_1(k9_metric_1(esk1_0,X4,X5),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_14112, negated_conjecture, (m1_subset_1(esk12_2(X1,k9_metric_1(esk1_0,X2,X3)),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k9_metric_1(esk1_0,X2,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(k9_metric_1(esk1_0,X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_14124, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(esk3_0,esk12_2(X1,k9_metric_1(esk1_0,X2,X3))),X4)|v3_pre_topc(k9_metric_1(esk1_0,X2,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|X4!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X3)|~m1_subset_1(k9_metric_1(esk1_0,X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_14123, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk12_2(X1,k9_metric_1(esk1_0,X2,X3)),X4)|v3_pre_topc(k9_metric_1(esk1_0,X2,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|X4!=k8_relat_1(esk3_0,X5)|X5!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X3)|~m1_subset_1(k9_metric_1(esk1_0,X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_13489, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(X1,k8_relat_1(esk3_0,k9_xtuple_0(esk3_0)))|~r2_hidden(esk14_2(X1,k8_relat_1(esk3_0,k9_xtuple_0(esk3_0))),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(X1,k8_relat_1(esk3_0,k9_xtuple_0(esk3_0)))),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13651, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(X3,k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k1_funct_1(esk3_0,esk8_3(esk3_0,X4,X2))|X4!=k10_xtuple_0(esk3_0)|~r2_hidden(esk8_3(esk3_0,X4,X2),k1_xboole_0)|~r2_hidden(X2,X4)|~v1_xreal_0(X3)|~m1_subset_1(esk8_3(esk3_0,X4,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|k8_relat_1(esk3_0,k9_metric_1(esk2_0,X1,X3))!=k1_xboole_0)).#trainneg
cnf(i_0_13493, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(X1,k8_relat_1(esk3_0,k9_xtuple_0(esk3_0)))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(X1,k8_relat_1(esk3_0,k9_xtuple_0(esk3_0)))),u1_struct_0(esk1_0))|~m1_subset_1(esk14_2(X1,k8_relat_1(esk3_0,k9_xtuple_0(esk3_0))),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13490, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(X1,k8_relat_1(X2,u1_struct_0(esk1_0)))|~r2_hidden(esk14_2(X1,k8_relat_1(X2,u1_struct_0(esk1_0))),k9_xtuple_0(X2))|~v1_relat_1(X2)|~m1_subset_1(k1_funct_1(X2,esk14_2(X1,k8_relat_1(X2,u1_struct_0(esk1_0)))),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_13501, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),k8_relat_1(X4,k9_xtuple_0(esk3_0)))|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|~v1_relat_1(X4)|~m1_subset_1(k1_funct_1(X4,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),k8_relat_1(X4,k9_xtuple_0(esk3_0)))),u1_struct_0(esk1_0))|~v1_funct_1(X4)|k9_xtuple_0(X4)!=k8_relat_1(esk3_0,X2)|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),k8_relat_1(X4,k9_xtuple_0(esk3_0))),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13856, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X1),k3_pcomps_1(esk1_0))|v5_pre_topc(X2,X3,k3_pcomps_1(esk1_0))|v2_struct_0(X3)|X1!=k10_xtuple_0(esk3_0)|~v3_pre_topc(esk7_3(X3,esk1_0,X2),k3_pcomps_1(esk1_0))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(esk7_3(X3,esk1_0,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(esk1_0))|~v1_funct_1(X2)|~v4_relat_1(esk3_0,esk7_3(X3,esk1_0,X2)))).#trainneg
cnf(i_0_14116, negated_conjecture, (m1_subset_1(esk12_2(X1,k9_metric_1(esk1_0,X2,X3)),X4)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k9_metric_1(esk1_0,X2,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|~v4_relat_1(esk3_0,X4)|~v1_xreal_0(X3)|~m1_subset_1(k9_metric_1(esk1_0,X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_13852, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk12_2(X1,k8_relset_1(u1_struct_0(esk1_0),X2,X3,X4)),k9_xtuple_0(esk3_0))|v3_pre_topc(k8_relset_1(u1_struct_0(esk1_0),X2,X3,X4),k3_pcomps_1(X1))|v2_struct_0(X1)|~m1_subset_1(k8_relset_1(u1_struct_0(esk1_0),X2,X3,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X2)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_13518, negated_conjecture, (v1_xreal_0(esk8_3(X1,X2,X3))|u1_struct_0(esk2_0)=k1_xboole_0|k8_relat_1(X1,k8_relat_1(esk3_0,X4))!=k10_xtuple_0(X5)|~v3_valued_0(X5)|~v1_relat_1(X5)|~v1_funct_1(X5)|X4!=k10_xtuple_0(esk3_0)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_13480, negated_conjecture, (r2_hidden(esk8_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(esk2_0)))|u1_struct_0(esk2_0)=k1_xboole_0|k2_pcomps_1(esk2_0)!=k8_relat_1(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_13506, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relat_1(X1,X2),k8_relat_1(X1,k9_xtuple_0(esk3_0)))|~v1_relat_1(X1)|~m1_subset_1(k1_funct_1(X1,esk14_2(k8_relat_1(X1,X2),k8_relat_1(X1,k9_xtuple_0(esk3_0)))),u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_14117, negated_conjecture, (r2_hidden(esk12_2(X1,k9_metric_1(esk1_0,X2,X3)),X4)|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k9_metric_1(esk1_0,X2,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|~v4_relat_1(esk3_0,X4)|~v1_xreal_0(X3)|~m1_subset_1(k9_metric_1(esk1_0,X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_13496, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relat_1(esk3_0,X1),k8_relat_1(X2,k9_xtuple_0(esk3_0)))|~v1_relat_1(X2)|~m1_subset_1(k1_funct_1(X2,esk14_2(k8_relat_1(esk3_0,X1),k8_relat_1(X2,k9_xtuple_0(esk3_0)))),u1_struct_0(esk1_0))|~v1_funct_1(X2)|~v4_relat_1(esk3_0,k9_xtuple_0(X2)))).#trainneg
cnf(i_0_14164, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_xtuple_0(X1),k8_relat_1(esk3_0,X2))|X2!=k10_xtuple_0(esk3_0)|~v4_relat_1(X1,u1_struct_0(esk1_0))|~v1_relat_1(X1))).#trainneg
cnf(i_0_13857, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k8_relat_1(esk3_0,X1),k3_pcomps_1(X2))|v5_pre_topc(X3,X4,k3_pcomps_1(esk1_0))|v2_struct_0(X2)|v2_struct_0(X4)|v1_xboole_0(esk7_3(X4,esk1_0,X3))|X1!=k10_xtuple_0(esk3_0)|~v3_pre_topc(esk7_3(X4,esk1_0,X3),k3_pcomps_1(X2))|~l1_pre_topc(X4)|~v2_pre_topc(X4)|~m1_subset_1(esk7_3(X4,esk1_0,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(k8_relat_1(esk3_0,X1),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(esk1_0))))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(esk1_0))|~v1_funct_1(X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v4_relat_1(esk3_0,esk7_3(X4,esk1_0,X3)))).#trainneg
cnf(i_0_13504, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(esk7_3(X1,esk1_0,X2),k8_relat_1(esk3_0,k9_xtuple_0(esk3_0)))|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(esk7_3(X1,esk1_0,X2),k8_relat_1(esk3_0,k9_xtuple_0(esk3_0)))),u1_struct_0(esk1_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_6622, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),X7),X1,X7)),X8))),X5)|v2_struct_0(X3)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),X7),X1,X7)),X8)|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),X7),k6_numbers)|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),k6_numbers)|~r2_hidden(X4,X5)|~v3_pre_topc(X5,k3_pcomps_1(X3))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)|X6!=k1_funct_1(esk3_0,X2)|X7!=k1_funct_1(esk3_0,X1)|~r2_hidden(X7,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)))|~r2_hidden(X6,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)))|~v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),k3_pcomps_1(esk2_0))|~v3_pre_topc(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),X7))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6))|~m1_subset_1(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),X7),X1,X7)),X8))),u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),X7),X1,X7)),X8)),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),X7),X1,X7)),X8)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),X7),X1,X7)),X8),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),X6),X2,X6)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(k9_metric_1(X3,X4,esk11_3(X3,X5,X4)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13503, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_xtuple_0(X1),k8_relat_1(X1,k9_xtuple_0(esk3_0)))|~v1_relat_1(X1)|~m1_subset_1(k1_funct_1(X1,esk14_2(k9_xtuple_0(X1),k8_relat_1(X1,k9_xtuple_0(esk3_0)))),u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_9083, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|m1_subset_1(esk8_3(X1,X2,X3),k8_relat_1(X1,u1_struct_0(esk1_0)))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_13716, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,X1),k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(X3,k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k1_funct_1(esk3_0,X1)|~r2_hidden(X1,k1_xboole_0)|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,X1),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|k8_relat_1(esk3_0,k9_metric_1(esk2_0,X2,X3))!=k1_xboole_0)).#trainneg
cnf(i_0_12954, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,esk7_3(X2,X3,X4),X5),X1,X5)),X6)),u1_struct_0(X3))|v5_pre_topc(X4,X2,k3_pcomps_1(X3))|v2_struct_0(X2)|v2_struct_0(X3)|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,esk7_3(X2,X3,X4),X5),X1,X5)),X6)|r1_xxreal_0(esk11_3(esk2_0,esk7_3(X2,X3,X4),X5),k6_numbers)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X3)))))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X3)))|~v1_funct_1(X4)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)|X5!=k1_funct_1(esk3_0,X1)|~r2_hidden(X5,esk7_3(X2,X3,X4))|~v3_pre_topc(esk7_3(X2,X3,X4),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,esk7_3(X2,X3,X4),X5))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,esk7_3(X2,X3,X4),X5),X1,X5)),X6),u1_struct_0(esk1_0))|~m1_subset_1(esk7_3(X2,X3,X4),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11370, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(esk3_0,esk12_2(X1,esk7_3(X2,esk1_0,X3))),X4)|v3_pre_topc(esk7_3(X2,esk1_0,X3),k3_pcomps_1(X1))|v5_pre_topc(X3,X2,k3_pcomps_1(esk1_0))|v2_struct_0(X2)|v2_struct_0(X1)|X4!=k10_xtuple_0(esk3_0)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(esk7_3(X2,esk1_0,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X3)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_8215, negated_conjecture, (r2_hidden(X1,k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,k8_relat_1(X2,X3))|~v4_relat_1(X2,u1_struct_0(esk1_0))|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_12955, negated_conjecture, (v5_pre_topc(X1,X2,k3_pcomps_1(X3))|m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X4,esk4_3(esk11_3(esk2_0,esk7_3(X2,X3,X1),X5),X4,X5)),X6)),u1_struct_0(k3_pcomps_1(X3)))|v2_struct_0(X3)|v2_struct_0(X2)|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X4,esk4_3(esk11_3(esk2_0,esk7_3(X2,X3,X1),X5),X4,X5)),X6)|r1_xxreal_0(esk11_3(esk2_0,esk7_3(X2,X3,X1),X5),k6_numbers)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X3)))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X3)))|~v1_funct_1(X1)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)|X5!=k1_funct_1(esk3_0,X4)|~r2_hidden(X5,esk7_3(X2,X3,X1))|~v3_pre_topc(esk7_3(X2,X3,X1),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,esk7_3(X2,X3,X1),X5))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X4,esk4_3(esk11_3(esk2_0,esk7_3(X2,X3,X1),X5),X4,X5)),X6),u1_struct_0(esk1_0))|~m1_subset_1(esk7_3(X2,X3,X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X4,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13347, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|r1_xxreal_0(X4,k6_numbers)|~r2_hidden(X1,k9_metric_1(esk1_0,X3,esk4_3(X4,X3,X5)))|X2!=k8_relat_1(esk3_0,k9_metric_1(esk2_0,X5,X4))|X5!=k1_funct_1(esk3_0,X3)|~v1_xreal_0(X4)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X5,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_11546, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(esk1_0))|v5_pre_topc(X5,X4,k3_pcomps_1(esk1_0))|v2_struct_0(X4)|esk7_3(X4,esk1_0,X5)!=k8_relat_1(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_xreal_0(esk5_3(X4,esk1_0,X5))|~m1_subset_1(esk6_3(X4,esk1_0,X5),u1_struct_0(esk1_0))|~v1_funct_1(X1)|~l1_pre_topc(X4)|~v2_pre_topc(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(esk1_0))))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(esk1_0))|~v1_funct_1(X5))).#trainneg
cnf(i_0_13515, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(esk3_0,k3_pcomps_1(esk1_0),k3_pcomps_1(X1))|v2_struct_0(X1)|~v2_pre_topc(k3_pcomps_1(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k3_pcomps_1(X1)))))|~m1_subset_1(esk7_3(k3_pcomps_1(esk1_0),X1,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v1_xboole_0(esk7_3(k3_pcomps_1(esk1_0),X1,esk3_0)))).#trainneg
cnf(i_0_13975, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(esk3_0,k3_pcomps_1(esk1_0),k3_pcomps_1(esk1_0))|~v2_pre_topc(k3_pcomps_1(esk1_0))|~m1_subset_1(esk7_3(k3_pcomps_1(esk1_0),esk1_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v1_xboole_0(esk7_3(k3_pcomps_1(esk1_0),esk1_0,esk3_0)))).#trainneg
cnf(i_0_10463, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(esk7_3(X3,esk1_0,X4),X5)|v5_pre_topc(X4,X3,k3_pcomps_1(esk1_0))|v2_struct_0(X3)|X1!=k1_funct_1(esk3_0,esk14_2(esk7_3(X3,esk1_0,X4),X5))|X2!=k10_xtuple_0(esk3_0)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk1_0))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(esk1_0))|~v1_funct_1(X4))).#trainneg
cnf(i_0_12721, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k1_funct_1(X3,esk8_3(X4,X5,X6))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|k9_xtuple_0(X3)!=k8_relat_1(X4,u1_struct_0(esk1_0))|X5!=k10_xtuple_0(X4)|~r2_hidden(X6,X5)|~v1_relat_1(X4)|~m1_subset_1(X6,u1_struct_0(esk1_0))|~v1_funct_1(X4))).#trainneg
cnf(i_0_6619, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,X5),X6),X2,X6)),X7),X1,X7)),X8))),u1_struct_0(X3))|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,X5),X6),X2,X6)),X7),X1,X7)),X8)|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,X5),X6),X2,X6)),X7),k6_numbers)|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(X3,X4,X5),X6),k6_numbers)|~v1_xreal_0(X5)|~m1_subset_1(X4,u1_struct_0(X3))|~l1_metric_1(X3)|X6!=k1_funct_1(esk3_0,X2)|X7!=k1_funct_1(esk3_0,X1)|~r2_hidden(X7,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,X5),X6),X2,X6)))|~r2_hidden(X6,k9_metric_1(X3,X4,X5))|~v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,X5),X6),X2,X6)),k3_pcomps_1(esk2_0))|~v3_pre_topc(k9_metric_1(X3,X4,X5),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,X5),X6),X2,X6)),X7))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(X3,X4,X5),X6))|~m1_subset_1(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,X5),X6),X2,X6)),X7),X1,X7)),X8))),u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,X5),X6),X2,X6)),X7),X1,X7)),X8)),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,X5),X6),X2,X6)),X7),X1,X7)),X8)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,X5),X6),X2,X6)),X7),X1,X7)),X8),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k9_metric_1(X3,X4,X5),X6),X2,X6)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(k9_metric_1(X3,X4,X5),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7106, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,k1_funct_1(esk3_0,X2),X3))|r1_xxreal_0(X3,k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk8_3(esk3_0,X4,X5)))|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk8_3(esk3_0,X4,X5)),u1_struct_0(esk1_0))|k9_metric_1(esk1_0,X2,esk4_3(X3,k1_funct_1(esk3_0,esk8_3(esk3_0,X4,X5)),X1))!=k10_xtuple_0(esk3_0)|X4!=k10_xtuple_0(esk3_0)|~r2_hidden(X5,X4))).#trainneg
cnf(i_0_13560, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(esk1_0))|v5_pre_topc(X5,X4,k3_pcomps_1(esk1_0))|v2_struct_0(X4)|esk7_3(X4,esk1_0,X5)!=k8_relat_1(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_xreal_0(esk5_3(X4,esk1_0,X5))|~m1_subset_1(esk6_3(X4,esk1_0,X5),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1)|~l1_pre_topc(X4)|~v2_pre_topc(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(esk1_0))))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(esk1_0))|~v1_funct_1(X5))).#trainneg
cnf(i_0_9183, negated_conjecture, (m1_subset_1(k8_relset_1(k9_xtuple_0(esk3_0),X1,X2,X3),k1_zfmisc_1(k8_relat_1(esk3_0,X4)))|u1_struct_0(esk2_0)=k1_xboole_0|X4!=k10_xtuple_0(esk3_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k9_xtuple_0(esk3_0),X1))))).#trainneg
cnf(i_0_9812, negated_conjecture, (v1_xboole_0(k2_relset_1(X1,X2))|m1_subset_1(esk6_3(X3,esk1_0,X4),X1)|u1_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(X4,X3,k3_pcomps_1(esk1_0))|v2_struct_0(X3)|~v5_relat_1(X2,X1)|~v1_relat_1(X2)|~v4_relat_1(esk3_0,k2_relset_1(X1,X2))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk1_0))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(esk1_0))|~v1_funct_1(X4))).#trainneg
cnf(i_0_9639, negated_conjecture, (v3_pre_topc(esk7_3(X1,esk2_0,X2),k3_pcomps_1(X3))|v3_pre_topc(u1_struct_0(esk2_0),k3_pcomps_1(esk2_0))|v5_pre_topc(X2,X1,k3_pcomps_1(esk2_0))|v2_struct_0(X3)|v2_struct_0(X1)|u1_struct_0(esk2_0)=k1_xboole_0|~v3_pre_topc(k8_relat_1(esk3_0,X4),k3_pcomps_1(esk2_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk7_3(X1,esk2_0,X2),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(k8_relat_1(esk3_0,X4),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk2_0))|~v1_funct_1(X2)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)|X4!=k10_xtuple_0(esk3_0)|~m1_subset_1(esk12_2(X3,esk7_3(X1,esk2_0,X2)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13718, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,k1_funct_1(esk3_0,X2),X3))|r1_xxreal_0(X3,k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k1_funct_1(esk3_0,X2)|~r2_hidden(X2,k1_xboole_0)|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|k8_relat_1(esk3_0,k9_metric_1(esk2_0,X1,X3))!=k1_xboole_0)).#trainneg
cnf(i_0_10579, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,esk8_3(X1,X2,X3)),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_14163, negated_conjecture, (r2_hidden(X1,k8_relat_1(esk3_0,X2))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,k1_relset_1(u1_struct_0(esk1_0),X3))|X2!=k10_xtuple_0(esk3_0)|~v4_relat_1(X3,u1_struct_0(esk1_0))|~v1_relat_1(X3))).#trainneg
cnf(i_0_14912, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,k9_xtuple_0(esk3_0))|~v4_relat_1(X2,u1_struct_0(esk1_0))|~v1_relat_1(X2)|~v1_funct_1(X2)|~r2_hidden(X1,k9_xtuple_0(X2))|~m1_subset_1(k1_funct_1(X2,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14443, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(X1,k8_relat_1(esk3_0,u1_struct_0(esk1_0)))|~r2_hidden(esk14_2(X1,k8_relat_1(esk3_0,u1_struct_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(X1,k8_relat_1(esk3_0,u1_struct_0(esk1_0)))),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14453, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(X1,k8_relat_1(esk3_0,u1_struct_0(esk1_0)))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(X1,k8_relat_1(esk3_0,u1_struct_0(esk1_0)))),u1_struct_0(esk1_0))|~m1_subset_1(esk14_2(X1,k8_relat_1(esk3_0,u1_struct_0(esk1_0))),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14694, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(X1,k9_metric_1(esk2_0,X1,X2))|r1_xxreal_0(X2,k6_numbers)|k8_relat_1(esk3_0,k9_metric_1(esk2_0,X1,X2))!=k1_xboole_0|~r2_hidden(esk8_3(esk3_0,X3,X1),k1_xboole_0)|~v1_xreal_0(X2)|~m1_subset_1(esk8_3(esk3_0,X3,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|X3!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_14808, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk14_2(k8_relat_1(X1,X2),X3),k9_xtuple_0(esk3_0))|r1_tarski(k8_relat_1(X1,X2),X3)|~v4_relat_1(X1,u1_struct_0(esk1_0))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_14458, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,X3),X4),X1,X4)),k8_relat_1(esk3_0,k9_xtuple_0(esk3_0)))|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,X3),X4),k6_numbers)|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,X3),X4),X1,X4)),k8_relat_1(esk3_0,k9_xtuple_0(esk3_0))),u1_struct_0(esk1_0))|X4!=k1_funct_1(esk3_0,X1)|~r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,X3),X4),X1,X4)),k8_relat_1(esk3_0,k9_xtuple_0(esk3_0))),u1_struct_0(esk1_0))|~r2_hidden(X4,k9_metric_1(esk1_0,X2,X3))|~v3_pre_topc(k9_metric_1(esk1_0,X2,X3),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,X3),X4))|~v1_xreal_0(X3)|~m1_subset_1(k9_metric_1(esk1_0,X2,X3),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14824, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk12_2(X1,k8_relat_1(X2,X3)),k9_xtuple_0(esk3_0))|v3_pre_topc(k8_relat_1(X2,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|~v4_relat_1(X2,u1_struct_0(esk1_0))|~v1_relat_1(X2)|~v1_funct_1(X2)|~m1_subset_1(k8_relat_1(X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_14470, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),k8_relat_1(X4,u1_struct_0(esk1_0)))|r1_xxreal_0(esk11_3(esk2_0,X2,X3),k6_numbers)|~v1_relat_1(X4)|~m1_subset_1(k1_funct_1(X4,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),k8_relat_1(X4,u1_struct_0(esk1_0)))),u1_struct_0(esk1_0))|~v1_funct_1(X4)|k9_xtuple_0(X4)!=k8_relat_1(esk3_0,X2)|X3!=k1_funct_1(esk3_0,X1)|~r2_hidden(X3,X2)|~v3_pre_topc(X2,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,X2,X3))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,X2,X3),X1,X3)),k8_relat_1(X4,u1_struct_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14802, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk14_2(k8_relat_1(X1,X2),X3),k9_xtuple_0(esk3_0))|r1_tarski(k8_relat_1(X1,X2),X3)|~v4_relat_1(X4,u1_struct_0(esk1_0))|~v1_relat_1(X4)|~v1_funct_1(X4)|k8_relat_1(X4,X5)!=k8_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_14474, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relat_1(esk3_0,X1),k8_relat_1(esk3_0,u1_struct_0(esk1_0)))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k8_relat_1(esk3_0,X1),k8_relat_1(esk3_0,u1_struct_0(esk1_0)))),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14890, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,X3)),k9_xtuple_0(esk3_0))|~v4_relat_1(X4,u1_struct_0(esk1_0))|~v1_relat_1(X4)|~v1_funct_1(X4)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k8_relat_1(X4,X5))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_14920, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,esk7_3(X2,esk1_0,X3),X4),X1,X4)),X5)|v5_pre_topc(X3,X2,k3_pcomps_1(esk1_0))|r1_xxreal_0(esk11_3(esk2_0,esk7_3(X2,esk1_0,X3),X4),k6_numbers)|m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,esk7_3(X2,esk1_0,X3),X4),X1,X4)),X5)),u1_struct_0(esk1_0))|v2_struct_0(X2)|X4!=k1_funct_1(esk3_0,X1)|~r2_hidden(X4,esk7_3(X2,esk1_0,X3))|~v3_pre_topc(esk7_3(X2,esk1_0,X3),k3_pcomps_1(esk2_0))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xreal_0(esk11_3(esk2_0,esk7_3(X2,esk1_0,X3),X4))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,esk7_3(X2,esk1_0,X3),X4),X1,X4)),X5),u1_struct_0(esk1_0))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk1_0))))|~m1_subset_1(esk7_3(X2,esk1_0,X3),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk1_0))|~v1_funct_1(X3))).#trainneg
cnf(i_0_14652, negated_conjecture, (m1_subset_1(k9_xtuple_0(X1),k1_zfmisc_1(k8_relat_1(esk3_0,X2)))|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(esk3_0)|~v4_relat_1(X1,u1_struct_0(esk1_0))|~v1_relat_1(X1))).#trainneg
cnf(i_0_14653, negated_conjecture, (r2_hidden(X1,k8_relat_1(esk3_0,X2))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(X1,k9_xtuple_0(X3))|X2!=k10_xtuple_0(esk3_0)|~v4_relat_1(X3,u1_struct_0(esk1_0))|~v1_relat_1(X3))).#trainneg
cnf(i_0_14696, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,esk7_3(X2,X3,X4),X5),X1,X5)),X6)),u1_struct_0(X3))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,esk7_3(X2,X3,X4),X5),X1,X5)),X6)|v5_pre_topc(X4,X2,k3_pcomps_1(X3))|r1_xxreal_0(esk11_3(esk2_0,esk7_3(X2,X3,X4),X5),k6_numbers)|v2_struct_0(X3)|v2_struct_0(X2)|X5!=k1_funct_1(esk3_0,X1)|~r2_hidden(X5,esk7_3(X2,X3,X4))|~v3_pre_topc(esk7_3(X2,X3,X4),k3_pcomps_1(esk2_0))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xreal_0(esk11_3(esk2_0,esk7_3(X2,X3,X4),X5))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,esk7_3(X2,X3,X4),X5),X1,X5)),X6),u1_struct_0(esk1_0))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X3)))))|~m1_subset_1(esk7_3(X2,X3,X4),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(k3_pcomps_1(X3)))|~v1_funct_1(X4)|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_14475, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relat_1(X1,X2),k8_relat_1(X1,u1_struct_0(esk1_0)))|~v1_relat_1(X1)|~m1_subset_1(k1_funct_1(X1,esk14_2(k8_relat_1(X1,X2),k8_relat_1(X1,u1_struct_0(esk1_0)))),u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_14477, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,X2),k8_relat_1(esk3_0,u1_struct_0(esk1_0)))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,X2),k8_relat_1(esk3_0,u1_struct_0(esk1_0)))),u1_struct_0(esk1_0))|~v1_xreal_0(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14345, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(esk7_3(X1,esk1_0,X2),k3_pcomps_1(X3))|m1_subset_1(esk12_2(X3,esk7_3(X1,esk1_0,X2)),k9_xtuple_0(esk3_0))|v2_struct_0(X3)|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|~v1_xreal_0(esk5_3(X1,esk1_0,X2))|~m1_subset_1(esk7_3(X1,esk1_0,X2),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(esk6_3(X1,esk1_0,X2),u1_struct_0(esk1_0))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_13004, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6))),k9_metric_1(esk2_0,X4,X3))|r1_xxreal_0(X3,k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6)|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),k6_numbers)|X4!=k1_funct_1(esk3_0,X2)|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6))),u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6)),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))|X5!=k1_funct_1(esk3_0,X1)|~r2_hidden(X5,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)))|~v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),X5),X1,X5)),X6),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(X3,X2,X4)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12756, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),X4)|r1_xxreal_0(esk11_3(esk2_0,X5,X6),k6_numbers)|X4!=k8_relat_1(esk3_0,X5)|X6!=k1_funct_1(esk3_0,X7)|~r2_hidden(X6,X5)|~v3_pre_topc(X5,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk4_3(esk11_3(esk2_0,X5,X6),X7,X6))|~v1_xreal_0(esk11_3(esk2_0,X5,X6))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X7,u1_struct_0(esk1_0))|k9_metric_1(esk1_0,X7,esk4_3(esk11_3(esk2_0,X5,X6),X7,X6))!=k8_relat_1(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_13505, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k8_relat_1(esk3_0,X1),k8_relat_1(esk3_0,k9_xtuple_0(esk3_0)))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k8_relat_1(esk3_0,X1),k8_relat_1(esk3_0,k9_xtuple_0(esk3_0)))),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10766, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|m1_subset_1(esk14_2(k8_relat_1(X1,X2),X3),u1_struct_0(esk1_0))|r1_tarski(k8_relat_1(X1,X2),X3)|~v1_xreal_0(X5)|~m1_subset_1(X4,u1_struct_0(esk1_0))|k9_metric_1(esk1_0,X4,X5)!=k8_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_9024, negated_conjecture, (r2_hidden(X1,k8_relat_1(esk3_0,X2))|u1_struct_0(esk2_0)=k1_xboole_0|v2_struct_0(X3)|~r2_hidden(X1,k9_metric_1(X3,X4,esk11_3(X3,k9_xtuple_0(esk3_0),X4)))|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(X4,k9_xtuple_0(esk3_0))|~v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(X3))|~m1_subset_1(k9_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X3))))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_metric_1(X3)|~v9_metric_1(X3)|~v8_metric_1(X3)|~v7_metric_1(X3)|~v6_metric_1(X3))).#trainneg
cnf(i_0_7575, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),k8_relat_1(esk3_0,X4))|X4!=k10_xtuple_0(esk3_0)|k8_relat_1(esk3_0,X5)!=k8_relat_1(X1,u1_struct_0(esk1_0))|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,k9_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_14388, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,esk12_2(X1,k9_metric_1(esk1_0,X2,X3))),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k9_metric_1(esk1_0,X2,X3),k3_pcomps_1(X1))|v2_struct_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(k9_metric_1(esk1_0,X2,X3),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X1))))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_6341, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X2),X1,X2)),X3)),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X2),X1,X2)),X3)|r1_xxreal_0(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X2),k6_numbers)|X2!=k1_funct_1(esk3_0,X1)|~r2_hidden(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X2),X1,X2)),X3),u1_struct_0(esk1_0))|~r2_hidden(X2,k9_xtuple_0(esk3_0))|~v3_pre_topc(k9_xtuple_0(esk3_0),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_xtuple_0(esk3_0),X2))|~m1_subset_1(k9_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6631, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6)),X7)|r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6)|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),k6_numbers)|r1_xxreal_0(esk11_3(esk2_0,X3,X4),k6_numbers)|X7!=k8_relat_1(esk3_0,X3)|~r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6)),k9_xtuple_0(esk3_0))|X4!=k1_funct_1(esk3_0,X2)|X5!=k1_funct_1(esk3_0,X1)|~r2_hidden(X5,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)))|~r2_hidden(X4,X3)|~v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),k3_pcomps_1(esk2_0))|~v3_pre_topc(X3,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5))|~v1_xreal_0(esk11_3(esk2_0,X3,X4))|~m1_subset_1(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6))),u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6)),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10880, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(esk7_3(X1,esk1_0,X2),X3)|r2_hidden(k1_funct_1(X4,esk8_3(X4,X5,esk14_2(esk7_3(X1,esk1_0,X2),X3))),k9_xtuple_0(esk3_0))|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|X5!=k10_xtuple_0(X4)|~r2_hidden(esk14_2(esk7_3(X1,esk1_0,X2),X3),X5)|~v1_relat_1(X4)|~v1_xreal_0(esk5_3(X1,esk1_0,X2))|~m1_subset_1(esk6_3(X1,esk1_0,X2),u1_struct_0(esk1_0))|~v1_funct_1(X4)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_14431, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v3_pre_topc(k9_metric_1(esk1_0,X3,X4),k3_pcomps_1(X2))|v2_struct_0(X2)|~v1_xboole_0(X1)|X1!=k10_xtuple_0(esk3_0)|~v1_xreal_0(X4)|~m1_subset_1(k9_metric_1(esk1_0,X3,X4),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_7084, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,esk8_3(esk3_0,X1,X2)),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_13494, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(esk7_3(X1,esk1_0,X2),k8_relat_1(X3,k9_xtuple_0(esk3_0)))|v1_xboole_0(k9_xtuple_0(X3))|v5_pre_topc(X2,X1,k3_pcomps_1(esk1_0))|v2_struct_0(X1)|~v1_relat_1(X3)|~m1_subset_1(k1_funct_1(X3,esk14_2(esk7_3(X1,esk1_0,X2),k8_relat_1(X3,k9_xtuple_0(esk3_0)))),u1_struct_0(esk1_0))|~v1_funct_1(X3)|~v4_relat_1(esk3_0,k9_xtuple_0(X3))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk1_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk1_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_7258, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_xtuple_0(esk3_0),X1)),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_xtuple_0(esk3_0),X1))).#trainneg
cnf(i_0_8886, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(esk8_3(X1,X2,X3),k9_xtuple_0(esk3_0))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X4)))|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,k8_relset_1(u1_struct_0(esk1_0),X4,X5,X6))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8261, negated_conjecture, (m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k9_xtuple_0(esk3_0)))|u1_struct_0(esk2_0)=k1_xboole_0)).#trainneg
cnf(i_0_12812, negated_conjecture, (v1_xreal_0(k1_funct_1(esk3_0,X1))|u1_struct_0(esk2_0)=k1_xboole_0|v1_xboole_0(k10_xtuple_0(esk3_0))|k10_xtuple_0(esk3_0)!=k10_xtuple_0(X2)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9930, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,esk8_3(X3,X4,X5)),k1_zfmisc_1(u1_struct_0(esk2_0)))|u1_struct_0(esk2_0)=k1_xboole_0|k8_relat_1(X1,k9_xtuple_0(X3))!=k2_pcomps_1(esk2_0)|X2!=k10_xtuple_0(X1)|X4!=k10_xtuple_0(X3)|~r2_hidden(esk8_3(X3,X4,X5),X2)|~r2_hidden(X5,u1_struct_0(esk1_0))|~r2_hidden(X5,X4)|~v1_relat_1(X1)|~v1_relat_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_7279, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,k1_funct_1(esk3_0,X2),X3))|r1_xxreal_0(X3,k6_numbers)|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_xtuple_0(esk3_0),X4)|X1!=k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_xtuple_0(esk3_0),X4)))|~v1_xreal_0(X3)|~m1_subset_1(k1_funct_1(esk3_0,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_xtuple_0(esk3_0),X4)),u1_struct_0(esk1_0))|k9_metric_1(esk1_0,X2,esk4_3(X3,k1_funct_1(esk3_0,esk14_2(k9_xtuple_0(esk3_0),X4)),X1))!=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_4057, negated_conjecture, (r2_hidden(k1_funct_1(X1,k1_funct_1(esk3_0,esk12_2(X2,k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,X4),X5),X3,X5))))),X4)|v3_pre_topc(k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,X4),X5),X3,X5)),k3_pcomps_1(X2))|r1_xxreal_0(esk11_3(esk2_0,k8_relat_1(X1,X4),X5),k6_numbers)|v2_struct_0(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|X5!=k1_funct_1(esk3_0,X3)|~r2_hidden(X5,k8_relat_1(X1,X4))|~v3_pre_topc(k8_relat_1(X1,X4),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k8_relat_1(X1,X4),X5))|~m1_subset_1(k1_funct_1(esk3_0,esk12_2(X2,k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,X4),X5),X3,X5)))),u1_struct_0(esk2_0))|~m1_subset_1(esk12_2(X2,k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,X4),X5),X3,X5))),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X3,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,X4),X5),X3,X5)),k1_zfmisc_1(u1_struct_0(k3_pcomps_1(X2))))|~m1_subset_1(k8_relat_1(X1,X4),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_4644, negated_conjecture, (r2_hidden(X1,k9_metric_1(esk2_0,X2,X3))|r1_xxreal_0(X3,k6_numbers)|r2_hidden(X4,u1_struct_0(esk1_0))|X2!=k1_funct_1(esk3_0,X4)|X5!=k10_xtuple_0(esk3_0)|~r2_hidden(esk8_3(esk3_0,X5,X1),k1_xboole_0)|~r2_hidden(X1,X5)|~v1_xreal_0(X3)|~m1_subset_1(esk8_3(esk3_0,X5,X1),u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~v1_xreal_0(esk4_3(X3,X4,X2)))).#trainneg
cnf(i_0_6340, negated_conjecture, (r2_hidden(k1_funct_1(X1,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,X3),X4),X2,X4)),X5))),X3)|u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,X3),X4),X2,X4)),X5)|r1_xxreal_0(esk11_3(esk2_0,k8_relat_1(X1,X3),X4),k6_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|X4!=k1_funct_1(esk3_0,X2)|~r2_hidden(esk14_2(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,k8_relat_1(X1,X3),X4),X2,X4)),X5),u1_struct_0(esk1_0))|~r2_hidden(X4,k8_relat_1(X1,X3))|~v3_pre_topc(k8_relat_1(X1,X3),k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k8_relat_1(X1,X3),X4))|~m1_subset_1(k8_relat_1(X1,X3),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12484, negated_conjecture, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,u1_struct_0(esk1_0))|~v4_relat_1(esk3_0,k10_xtuple_0(esk3_0))|~r2_hidden(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6692, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk8_3(X1,X2,esk8_3(esk3_0,X3,X4))),k9_xtuple_0(esk3_0))|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(X1)|~r2_hidden(esk8_3(esk3_0,X3,X4),X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|X5!=k10_xtuple_0(esk3_0)|X3!=k10_xtuple_0(esk3_0)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_6744, negated_conjecture, (r1_tarski(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6)|r1_xxreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),k6_numbers)|r1_xxreal_0(esk11_3(esk2_0,X3,X4),k6_numbers)|m1_subset_1(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6))),X3)|u1_struct_0(esk2_0)=k1_xboole_0|X4!=k1_funct_1(esk3_0,X2)|X5!=k1_funct_1(esk3_0,X1)|~r2_hidden(X5,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)))|~r2_hidden(X4,X3)|~v3_pre_topc(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),k3_pcomps_1(esk2_0))|~v3_pre_topc(X3,k3_pcomps_1(esk2_0))|~v1_xreal_0(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5))|~v1_xreal_0(esk11_3(esk2_0,X3,X4))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6)),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6)),u1_struct_0(esk2_0))|~m1_subset_1(esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6),u1_struct_0(esk1_0))|~m1_subset_1(k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(k1_funct_1(esk3_0,esk14_2(k9_metric_1(esk1_0,X1,esk4_3(esk11_3(esk2_0,k9_metric_1(esk1_0,X2,esk4_3(esk11_3(esk2_0,X3,X4),X2,X4)),X5),X1,X5)),X6)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11850, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_tarski(k1_relset_1(u1_struct_0(esk1_0),X1),k9_xtuple_0(esk3_0))|~v4_relat_1(X1,u1_struct_0(esk1_0))|~v1_relat_1(X1))).#trainneg
cnf(i_0_7677, negated_conjecture, (r2_hidden(X1,k8_relat_1(esk3_0,X2))|u1_struct_0(esk2_0)=k1_xboole_0|v5_pre_topc(X4,X3,k3_pcomps_1(esk1_0))|v2_struct_0(X3)|~r2_hidden(X1,esk7_3(X3,esk1_0,X4))|X2!=k10_xtuple_0(esk3_0)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk1_0))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(esk1_0))|~v1_funct_1(X4))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 63
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 133
# Removed in clause preprocessing      : 10
# Initial clauses in saturation        : 123
# Processed clauses                    : 2212
# ...of these trivial                  : 20
# ...subsumed                          : 956
# ...remaining for further processing  : 1236
# Other redundant clauses eliminated   : 14
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 91
# Backward-rewritten                   : 637
# Generated clauses                    : 15875
# ...of the previous two non-trivial   : 15669
# Contextual simplify-reflections      : 212
# Paramodulations                      : 15735
# Factorizations                       : 0
# Equation resolutions                 : 138
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 506
#    Positive orientable unit clauses  : 44
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 9
#    Non-unit-clauses                  : 453
# Current number of unprocessed clauses: 13352
# ...number of literals in the above   : 186593
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1966
# Proof object given clauses           : 163
# Proof search given clauses           : 1236
# Clause-clause subsumption calls (NU) : 1480793
# Rec. Clause-clause subsumption calls : 20632
# Non-unit clause-clause subsumptions  : 1167
# Unit Clause-clause subsumption calls : 5736
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 44
# BW rewrite match successes           : 7
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 710682

# -------------------------------------------------
# User time                : 83.349 s
# System time              : 3.719 s
# Total time               : 87.068 s
# Maximum resident set size: 741088 pages
83.34user 3.93system 1:28.09elapsed 99%CPU (0avgtext+0avgdata 741088maxresident)k
0inputs+1040outputs (0major+1079020minor)pagefaults 0swaps
